Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.pred_set
% (34900)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 610
% Passive clauses: 3140
% Generated clauses: 12735
% Final active clauses: 597
% Final passive clauses: 2360
% Input formulas: 260
% Initial clauses: 545
% 
% Duplicate literals: 133
% Trivial inequalities: 30
% Fw subsumption resolutions: 1001
% Fw demodulations: 875
% Bw demodulations: 7
% 
% Simple tautologies: 256
% Equational tautologies: 21
% Forward subsumptions: 7172
% 
% Binary resolution: 3433
% Factoring: 342
% Forward superposition: 3045
% Backward superposition: 3220
% Self superposition: 19
% Equality factoring: 2
% Equality resolution: 7
% 
% Split clauses: 174
% Split components: 350
% SAT solver clauses: 276
% SAT solver unit clauses: 15
% SAT solver binary clauses: 62
% 
% Sat splits: 174
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 276
% TWLsolver calls for satisfiability: 200
% 
% Memory used [KB]: 6140
% Time elapsed: 0.601 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.pred_set
% (34920)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 515
% Passive clauses: 2124
% Generated clauses: 15917
% Final active clauses: 458
% Final passive clauses: 1189
% Input formulas: 260
% Initial clauses: 545
% 
% Duplicate literals: 99
% Trivial inequalities: 44
% Fw subsumption resolutions: 773
% Fw demodulations: 1609
% Bw demodulations: 441
% Condensations: 74
% Global subsumptions: 26
% 
% Simple tautologies: 270
% Equational tautologies: 9
% Forward subsumptions: 8364
% Fw demodulations to eq. taut.: 7
% Bw demodulations to eq. taut.: 3
% 
% Binary resolution: 1665
% Factoring: 146
% Forward superposition: 4618
% Backward superposition: 3443
% Self superposition: 57
% Equality factoring: 2
% Equality resolution: 14
% 
% Split clauses: 14
% Split components: 28
% SAT solver clauses: 2694
% SAT solver unit clauses: 658
% SAT solver binary clauses: 203
% SAT solver learnt clauses: 56
% SAT solver learnt literals: 21
% 
% Sat splits: 14
% Sat splitting refutations: 78
% 
% TWLsolver clauses: 2586
% TWLsolver calls for satisfiability: 3430
% 
% Memory used [KB]: 12665
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.pred_set
% (34945)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1376
% Passive clauses: 11911
% Generated clauses: 233446
% Final active clauses: 938
% Final passive clauses: 3240
% Input formulas: 260
% Initial clauses: 545
% Discarded non-redundant clauses: 35048
% 
% Duplicate literals: 10
% Trivial inequalities: 76
% Fw subsumption resolutions: 5417
% Fw demodulations: 77977
% Bw demodulations: 4529
% Condensations: 26
% Global subsumptions: 4
% 
% Simple tautologies: 185
% Equational tautologies: 74
% Forward subsumptions: 66282
% Fw demodulations to eq. taut.: 40016
% Bw demodulations to eq. taut.: 39
% 
% Unit resulting resolution: 96329
% Factoring: 3
% Forward superposition: 26191
% Backward superposition: 17971
% Self superposition: 124
% Equality factoring: 2
% Equality resolution: 24
% 
% Split clauses: 20
% Split components: 40
% SAT solver clauses: 16821
% SAT solver unit clauses: 8550
% SAT solver binary clauses: 5611
% 
% Sat splits: 20
% Sat splitting refutations: 4864
% 
% Memory used [KB]: 102087
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 27 next slice time: 27
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f1129,plain,(
  $false),
  inference(subsumption_resolution,[],[f972,f970])).
fof(f970,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK14,bool),sK15),s(fun(sK14,bool),sK16))))),
  inference(subsumption_resolution,[],[f969,f880])).
fof(f880,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f601])).
fof(f601,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) & (s(fun(X0,bool),X1) != s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(nnf_transformation,[],[f343])).
fof(f343,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))),
  inference(flattening,[],[f342])).
fof(f342,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X6] : ! [X9,X10] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X9),s(fun(X6,bool),X10)))) <=> s(fun(X6,bool),X9) = s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),X9),s(fun(X6,bool),X10))))),
  file('thm.pred_set.INTER_SUBSET_EQN|split|1','thm.pred_set.SUBSET_INTER_ABSORPTION')).
fof(f969,plain,(
  s(fun(sK14,bool),sK15) != s(fun(sK14,bool),'const.pred_set.INTER_2'(s(fun(sK14,bool),sK15),s(fun(sK14,bool),sK16))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK14,bool),sK15),s(fun(sK14,bool),sK16))))),
  inference(forward_demodulation,[],[f651,f892])).
fof(f892,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))) )),
  inference(cnf_transformation,[],[f358])).
fof(f358,plain,(
  ! [X0,X1,X2] : s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))),
  inference(flattening,[],[f357])).
fof(f357,plain,(
  ! [X0] : ! [X1,X2] : s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X6] : ! [X9,X10] : s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),X9),s(fun(X6,bool),X10))) = s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),X10),s(fun(X6,bool),X9)))),
  file('thm.pred_set.INTER_SUBSET_EQN|split|1','thm.pred_set.INTER_COMM')).
fof(f651,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK14,bool),sK15),s(fun(sK14,bool),sK16)))) | s(fun(sK14,bool),sK15) != s(fun(sK14,bool),'const.pred_set.INTER_2'(s(fun(sK14,bool),sK16),s(fun(sK14,bool),sK15)))),
  inference(cnf_transformation,[],[f500])).
fof(f500,plain,(
  (s(fun(sK14,bool),sK15) = s(fun(sK14,bool),'const.pred_set.INTER_2'(s(fun(sK14,bool),sK16),s(fun(sK14,bool),sK15))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK14,bool),sK15),s(fun(sK14,bool),sK16))))) & (s(fun(sK14,bool),sK15) != s(fun(sK14,bool),'const.pred_set.INTER_2'(s(fun(sK14,bool),sK16),s(fun(sK14,bool),sK15))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK14,bool),sK15),s(fun(sK14,bool),sK16)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK14,sK15,sK16])],[f499])).
fof(f499,plain,(
  ? [X0,X1,X2] : ((s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) & (s(fun(X0,bool),X1) != s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(nnf_transformation,[],[f427])).
fof(f427,plain,(
  ? [X0,X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))) <~> p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))),
  inference(ennf_transformation,[],[f263])).
fof(f263,plain,(
  ~! [X0,X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))) <=> p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))),
  inference(flattening,[],[f262])).
fof(f262,plain,(
  ~! [X0] : ! [X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))) <=> p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : (s(fun(X6,bool),X7) = s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X7))) <=> p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : (s(fun(X6,bool),X7) = s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X7))) <=> p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))))),
  file('thm.pred_set.INTER_SUBSET_EQN|split|1','thm.pred_set.INTER_SUBSET_EQN|split|1')).
fof(f972,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK14,bool),sK15),s(fun(sK14,bool),sK16))))),
  inference(subsumption_resolution,[],[f971,f881])).
fof(f881,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | s(fun(X0,bool),X1) != s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) )),
  inference(cnf_transformation,[],[f601])).
fof(f971,plain,(
  s(fun(sK14,bool),sK15) = s(fun(sK14,bool),'const.pred_set.INTER_2'(s(fun(sK14,bool),sK15),s(fun(sK14,bool),sK16))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK14,bool),sK15),s(fun(sK14,bool),sK16))))),
  inference(forward_demodulation,[],[f650,f892])).
fof(f650,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK14,bool),sK15),s(fun(sK14,bool),sK16)))) | s(fun(sK14,bool),sK15) = s(fun(sK14,bool),'const.pred_set.INTER_2'(s(fun(sK14,bool),sK16),s(fun(sK14,bool),sK15)))),
  inference(cnf_transformation,[],[f500])).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 210
% Passive clauses: 118
% Generated clauses: 160
% Final active clauses: 210
% Final passive clauses: 116
% Input formulas: 260
% Initial clauses: 310
% 
% Selected by SInE selection: 96
% SInE iterations: 1
% 
% Duplicate literals: 12
% Fw subsumption resolutions: 11
% Fw demodulations: 2
% 
% Simple tautologies: 101
% Forward subsumptions: 16
% 
% Binary resolution: 122
% Forward superposition: 16
% 
% Split clauses: 8
% Split components: 17
% SAT solver clauses: 9
% SAT solver binary clauses: 7
% 
% Sat splits: 8
% 
% Memory used [KB]: 1535
% Time elapsed: 0.041 s
% ------------------------------
% ------------------------------
% Success in time 12.3 s
