Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.relation
% (107003)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 94
% Passive clauses: 2840
% Generated clauses: 16942
% Final active clauses: 80
% Final passive clauses: 2600
% Input formulas: 6
% Initial clauses: 16
% 
% Duplicate literals: 1085
% Trivial inequalities: 162
% Fw subsumption resolutions: 1947
% Fw demodulations: 1128
% Bw demodulations: 9
% 
% Simple tautologies: 309
% Equational tautologies: 98
% Forward subsumptions: 6453
% 
% Binary resolution: 84
% Forward superposition: 6892
% Backward superposition: 4796
% Self superposition: 375
% Equality factoring: 273
% Equality resolution: 10
% 
% Split clauses: 826
% Split components: 1662
% SAT solver clauses: 1928
% SAT solver unit clauses: 329
% SAT solver binary clauses: 507
% 
% Sat splits: 826
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 1920
% TWLsolver calls for satisfiability: 93
% 
% Memory used [KB]: 9466
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.relation
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.relation
% SZS output start Proof for thm.relation
fof(f2208,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f2207,f41,f62,f1782,f34,f53,f35,f55,f74,f76,f32,f49,f36,f56,f512,f519,f517,f389,f390,f33,f47,f64,f72,f609,f1921,f1923,f39,f58,f38,f59,f37,f60,f143,f148,f31,f50,f40,f57,f2027,f2040,f30,f51,f1940,f1754,f1942,f42,f61])).
fof(f61,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP0(X0,X2,X1) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK9(X2,X1,X0)))),s(X0,sK10(X2,X1,X0)))))) ) | $spl22),
  inference(cnf_transformation,[],[f61_D])).
fof(f61_D,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP0(X0,X2,X1) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK9(X2,X1,X0)))),s(X0,sK10(X2,X1,X0)))))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f42,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | sP0(X0,X2,X1) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK9(X2,X1,X0)))),s(X0,sK10(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f28])).
fof(f28,plain,(
  ! [X0,X1,X2] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK9(X2,X1,X0)))),s(X0,sK10(X2,X1,X0))))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK9(X2,X1,X0)))),s(X0,sK10(X2,X1,X0)))))) | sP0(X0,X2,X1) | ! [X3,X4] : (~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK9,sK10])],[f27])).
fof(f27,plain,(
  ! [X0,X1,X2] : (? [X5,X6] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X6)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X5))),s(X0,X6))))) | sP0(X0,X2,X1) | ! [X3,X4] : (~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))))))),
  inference(rectify,[],[f21])).
fof(f21,plain,(
  ! [X0,X1,X2] : (? [X6,X7] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X6))),s(X0,X7)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X6))),s(X0,X7))))) | sP0(X0,X2,X1) | ! [X8,X9] : (~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X8),s(X0,X9)))) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X8))),s(X0,X9)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X8),s(X0,X9)))))))),
  inference(definition_folding,[],[f18,f20])).
fof(f20,plain,(
  ! [X0,X2,X1] : (~sP0(X0,X2,X1) | ? [X3,X4,X5] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X5)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X5))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f18,plain,(
  ! [X0,X1,X2] : (? [X6,X7] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X6))),s(X0,X7)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X6))),s(X0,X7))))) | ? [X3,X4,X5] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X5)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X5))))) | ! [X8,X9] : (~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X8),s(X0,X9)))) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X8))),s(X0,X9)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X8),s(X0,X9)))))))),
  inference(flattening,[],[f17])).
fof(f17,plain,(
  ! [X0,X1,X2] : ((? [X6,X7] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X6))),s(X0,X7)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X6))),s(X0,X7))))) | ? [X3,X4,X5] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X5))))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X5)))))) | ! [X8,X9] : (~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X8),s(X0,X9)))) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X8))),s(X0,X9)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X8),s(X0,X9)))))))),
  inference(ennf_transformation,[],[f12])).
fof(f12,plain,(
  ! [X0,X1,X2] : ((! [X6,X7] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X6))),s(X0,X7)))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X6))),s(X0,X7))))) & ! [X3,X4,X5] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X5))))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X5)))))) => ! [X8,X9] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X8),s(X0,X9)))) => (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X8))),s(X0,X9)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X8),s(X0,X9)))))))),
  inference(flattening,[],[f11])).
fof(f11,plain,(
  ! [X0] : ! [X1,X2] : ((! [X6,X7] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X6))),s(X0,X7)))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X6))),s(X0,X7))))) & ! [X3,X4,X5] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X5))))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X5)))))) => ! [X8,X9] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X8),s(X0,X9)))) => (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X8))),s(X0,X9)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X8),s(X0,X9)))))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6] : ! [X7,X8] : ((! [X9,X10] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X9))),s(X6,X10)))) => p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,X9))),s(X6,X10))))) & ! [X9,X10,X11] : ((p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,X9))),s(X6,X10)))) & p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X9),s(X6,X10)))) & p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X10))),s(X6,X11))))) => p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,X9))),s(X6,X11)))))) => ! [X12,X13] : (p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X12),s(X6,X13)))) => (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,X12))),s(X6,X13)))) & p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X12),s(X6,X13)))))))),
  file('thm.relation.TC_STRONG_INDUCT_RIGHT|49|','thm.relation.relation_5|48|_prove_ltb')).
fof(f1942,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1))))) | $spl395),
  inference(cnf_transformation,[],[f1942_D])).
fof(f1942_D,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1))))) <=> ~$spl395),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl395])])).
fof(f1754,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,X1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0),s(sK1,X1))))) ) | $spl336),
  inference(cnf_transformation,[],[f1754_D])).
fof(f1754_D,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,X1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0),s(sK1,X1))))) ) <=> ~$spl336),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl336])])).
fof(f1940,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0),s(sK1,X1)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,X1)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1)))))) ) | ($spl22 | $spl381)),
  inference(resolution,[],[f1923,f61])).
fof(f51,plain,(
  ( ! [X8,X9] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X8))),s(sK1,X9)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X8))),s(sK1,X9))))) ) | $spl6),
  inference(cnf_transformation,[],[f51_D])).
fof(f51_D,plain,(
  ( ! [X8,X9] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X8))),s(sK1,X9)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X8))),s(sK1,X9))))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f30,plain,(
  ( ! [X8,X9] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X8))),s(sK1,X9)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X8))),s(sK1,X9))))) )),
  inference(cnf_transformation,[],[f23])).
fof(f23,plain,(
  ! [X8,X9] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X8))),s(sK1,X9)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X8))),s(sK1,X9))))) & ! [X5,X6,X7] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X5))),s(sK1,X6)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X5),s(sK1,X6)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X6))),s(sK1,X7)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X5))),s(sK1,X7))))) & (p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK4),s(sK1,sK5)))) & ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK4))),s(sK1,sK5)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK1,sK2,sK3,sK4,sK5])],[f22])).
fof(f22,plain,(
  ? [X0,X1,X2] : (! [X8,X9] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X8))),s(X0,X9)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X8))),s(X0,X9))))) & ! [X5,X6,X7] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X5))),s(X0,X6)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X5),s(X0,X6)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X6))),s(X0,X7)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X5))),s(X0,X7))))) & ? [X3,X4] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4))))))),
  inference(rectify,[],[f16])).
fof(f16,plain,(
  ? [X0,X1,X2] : (! [X6,X7] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X6))),s(X0,X7)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X6))),s(X0,X7))))) & ! [X3,X4,X5] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X5)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X5))))) & ? [X8,X9] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X8),s(X0,X9)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X8))),s(X0,X9))))))),
  inference(flattening,[],[f15])).
fof(f15,plain,(
  ? [X0,X1,X2] : ((! [X6,X7] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X6))),s(X0,X7)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X6))),s(X0,X7))))) & ! [X3,X4,X5] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X5))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X5)))))) & ? [X8,X9] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X8),s(X0,X9)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X8))),s(X0,X9))))))),
  inference(ennf_transformation,[],[f9])).
fof(f9,plain,(
  ~! [X0,X1,X2] : ((! [X6,X7] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X6))),s(X0,X7)))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X6))),s(X0,X7))))) & ! [X3,X4,X5] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X5))))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X5)))))) => ! [X8,X9] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X8),s(X0,X9)))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X8))),s(X0,X9))))))),
  inference(flattening,[],[f8])).
fof(f8,plain,(
  ~! [X0] : ! [X1,X2] : ((! [X6,X7] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X6))),s(X0,X7)))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X6))),s(X0,X7))))) & ! [X3,X4,X5] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X5))))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X5)))))) => ! [X8,X9] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X8),s(X0,X9)))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X8))),s(X0,X9))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : ((! [X9,X10] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X9))),s(X6,X10)))) => p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,X9))),s(X6,X10))))) & ! [X9,X10,X11] : ((p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,X9))),s(X6,X10)))) & p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X9),s(X6,X10)))) & p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X10))),s(X6,X11))))) => p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,X9))),s(X6,X11)))))) => ! [X12,X13] : (p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X12),s(X6,X13)))) => p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,X12))),s(X6,X13))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : ((! [X9,X10] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X9))),s(X6,X10)))) => p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,X9))),s(X6,X10))))) & ! [X9,X10,X11] : ((p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,X9))),s(X6,X10)))) & p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X9),s(X6,X10)))) & p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X10))),s(X6,X11))))) => p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,X9))),s(X6,X11)))))) => ! [X12,X13] : (p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X12),s(X6,X13)))) => p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,X12))),s(X6,X13))))))),
  file('thm.relation.TC_STRONG_INDUCT_RIGHT|49|','thm.relation.TC_STRONG_INDUCT_RIGHT|49|')).
fof(f2040,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1))))) | $spl421),
  inference(cnf_transformation,[],[f2040_D])).
fof(f2040_D,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1))))) <=> ~$spl421),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl421])])).
fof(f2027,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1))))) | ($spl6 | $spl395)),
  inference(resolution,[],[f1942,f51])).
fof(f57,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK6(X2,X1,X0)))),s(X0,sK8(X2,X1,X0))))) | ~sP0(X0,X1,X2)) ) | $spl14),
  inference(cnf_transformation,[],[f57_D])).
fof(f57_D,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK6(X2,X1,X0)))),s(X0,sK8(X2,X1,X0))))) | ~sP0(X0,X1,X2)) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f40,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK6(X2,X1,X0)))),s(X0,sK8(X2,X1,X0))))) | ~sP0(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f26])).
fof(f26,plain,(
  ! [X0,X1,X2] : (~sP0(X0,X1,X2) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK6(X2,X1,X0)))),s(X0,sK7(X2,X1,X0))))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X2),s(X0,sK6(X2,X1,X0)),s(X0,sK7(X2,X1,X0))))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK7(X2,X1,X0)))),s(X0,sK8(X2,X1,X0))))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK6(X2,X1,X0)))),s(X0,sK8(X2,X1,X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK6,sK7,sK8])],[f25])).
fof(f25,plain,(
  ! [X0,X1,X2] : (~sP0(X0,X1,X2) | ? [X3,X4,X5] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X4)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X2),s(X0,X3),s(X0,X4)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X4))),s(X0,X5)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X5))))))),
  inference(rectify,[],[f24])).
fof(f24,plain,(
  ! [X0,X2,X1] : (~sP0(X0,X2,X1) | ? [X3,X4,X5] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X5)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X5))))))),
  inference(nnf_transformation,[],[f20])).
fof(f50,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X5))),s(sK1,X7)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X6))),s(sK1,X7)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X5),s(sK1,X6)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X5))),s(sK1,X6))))) ) | $spl4),
  inference(cnf_transformation,[],[f50_D])).
fof(f50_D,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X5))),s(sK1,X7)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X6))),s(sK1,X7)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X5),s(sK1,X6)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X5))),s(sK1,X6))))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f31,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X5))),s(sK1,X7)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X6))),s(sK1,X7)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X5),s(sK1,X6)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X5))),s(sK1,X6))))) )),
  inference(cnf_transformation,[],[f23])).
fof(f148,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X1))),s(sK1,sK8(X0,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(X0,sK3,sK1)))),s(sK1,X1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(X0,sK3,sK1)),s(sK1,X1)))) | ~sP0(sK1,sK3,X0)) ) | $spl58),
  inference(cnf_transformation,[],[f148_D])).
fof(f148_D,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X1))),s(sK1,sK8(X0,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(X0,sK3,sK1)))),s(sK1,X1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(X0,sK3,sK1)),s(sK1,X1)))) | ~sP0(sK1,sK3,X0)) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f143,plain,(
  ( ! [X0,X1] : (~sP0(sK1,sK3,X0) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X1))),s(sK1,sK8(X0,sK3,sK1))))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(X0,sK3,sK1)),s(sK1,X1)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(X0,sK3,sK1)))),s(sK1,X1))))) ) | ($spl4 | $spl14)),
  inference(resolution,[],[f57,f50])).
fof(f60,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK6(X2,X1,X0)))),s(X0,sK7(X2,X1,X0))))) | ~sP0(X0,X1,X2)) ) | $spl20),
  inference(cnf_transformation,[],[f60_D])).
fof(f60_D,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK6(X2,X1,X0)))),s(X0,sK7(X2,X1,X0))))) | ~sP0(X0,X1,X2)) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f37,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK6(X2,X1,X0)))),s(X0,sK7(X2,X1,X0))))) | ~sP0(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f26])).
fof(f59,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X2),s(X0,sK6(X2,X1,X0)),s(X0,sK7(X2,X1,X0))))) | ~sP0(X0,X1,X2)) ) | $spl18),
  inference(cnf_transformation,[],[f59_D])).
fof(f59_D,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X2),s(X0,sK6(X2,X1,X0)),s(X0,sK7(X2,X1,X0))))) | ~sP0(X0,X1,X2)) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f38,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X2),s(X0,sK6(X2,X1,X0)),s(X0,sK7(X2,X1,X0))))) | ~sP0(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f26])).
fof(f58,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK7(X2,X1,X0)))),s(X0,sK8(X2,X1,X0))))) | ~sP0(X0,X1,X2)) ) | $spl16),
  inference(cnf_transformation,[],[f58_D])).
fof(f58_D,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK7(X2,X1,X0)))),s(X0,sK8(X2,X1,X0))))) | ~sP0(X0,X1,X2)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f39,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK7(X2,X1,X0)))),s(X0,sK8(X2,X1,X0))))) | ~sP0(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f26])).
fof(f1923,plain,(
  ~sP0(sK1,sK3,sK2) | $spl381),
  inference(cnf_transformation,[],[f1923_D])).
fof(f1923_D,plain,(
  ~sP0(sK1,sK3,sK2) <=> ~$spl381),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl381])])).
fof(f1921,plain,(
  ~sP0(sK1,sK3,sK2) | ($spl16 | $spl18 | $spl20 | $spl58)),
  inference(subsumption_resolution,[],[f1920,f59])).
fof(f1920,plain,(
  ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK7(sK2,sK3,sK1))))) | ~sP0(sK1,sK3,sK2) | ($spl16 | $spl20 | $spl58)),
  inference(subsumption_resolution,[],[f1919,f60])).
fof(f1919,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK3,sK1)))),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK7(sK2,sK3,sK1))))) | ~sP0(sK1,sK3,sK2) | ($spl16 | $spl58)),
  inference(duplicate_literal_removal,[],[f1897])).
fof(f1897,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK3,sK1)))),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK7(sK2,sK3,sK1))))) | ~sP0(sK1,sK3,sK2) | ~sP0(sK1,sK3,sK2) | ($spl16 | $spl58)),
  inference(resolution,[],[f148,f58])).
fof(f609,plain,(
  $false | ($spl8 | $spl11 | $spl120)),
  inference(subsumption_resolution,[],[f584,f53])).
fof(f584,plain,(
  ~p(s(bool,'T_0')) | ($spl11 | $spl120)),
  inference(backward_demodulation,[],[f517,f55])).
fof(f72,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | $spl28),
  inference(cnf_transformation,[],[f72_D])).
fof(f72_D,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f64,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | $spl12),
  inference(superposition,[],[f56,f56])).
fof(f47,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK4))),s(sK1,sK5)))) | $spl1),
  inference(cnf_transformation,[],[f47_D])).
fof(f47_D,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK4))),s(sK1,sK5)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f33,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK4))),s(sK1,sK5))))),
  inference(cnf_transformation,[],[f23])).
fof(f390,plain,(
  ( ! [X9] : (s(bool,X9) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK4))),s(sK1,sK5))) | s(bool,'T_0') = s(bool,X9)) ) | $spl78),
  inference(cnf_transformation,[],[f390_D])).
fof(f390_D,plain,(
  ( ! [X9] : (s(bool,X9) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK4))),s(sK1,sK5))) | s(bool,'T_0') = s(bool,X9)) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f389,plain,(
  ( ! [X9] : (s(bool,X9) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK4))),s(sK1,sK5))) | s(bool,'T_0') = s(bool,X9)) ) | ($spl1 | $spl8 | $spl28)),
  inference(subsumption_resolution,[],[f245,f53])).
fof(f245,plain,(
  ( ! [X9] : (~p(s(bool,'T_0')) | s(bool,X9) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK4))),s(sK1,sK5))) | s(bool,'T_0') = s(bool,X9)) ) | ($spl1 | $spl28)),
  inference(superposition,[],[f47,f72])).
fof(f517,plain,(
  ( ! [X9] : (s(bool,'T_0') = s(bool,X9)) ) | $spl120),
  inference(cnf_transformation,[],[f517_D])).
fof(f517_D,plain,(
  ( ! [X9] : (s(bool,'T_0') = s(bool,X9)) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f519,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK4))),s(sK1,sK5))) | $spl122),
  inference(cnf_transformation,[],[f519_D])).
fof(f519_D,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK4))),s(sK1,sK5))) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f512,plain,(
  ( ! [X9] : (s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK4))),s(sK1,sK5))) | s(bool,'T_0') = s(bool,X9)) ) | ($spl12 | $spl78)),
  inference(duplicate_literal_removal,[],[f454])).
fof(f454,plain,(
  ( ! [X9] : (s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK4))),s(sK1,sK5))) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) ) | ($spl12 | $spl78)),
  inference(superposition,[],[f390,f56])).
fof(f56,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl12),
  inference(cnf_transformation,[],[f56_D])).
fof(f56_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f36,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.relation.TC_STRONG_INDUCT_RIGHT|49|',aHL_BOOL_CASES)).
fof(f49,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK4),s(sK1,sK5)))) | $spl2),
  inference(cnf_transformation,[],[f49_D])).
fof(f49_D,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK4),s(sK1,sK5)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f32,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK4),s(sK1,sK5))))),
  inference(cnf_transformation,[],[f23])).
fof(f76,plain,(
  s(bool,'T_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK4),s(sK1,sK5))) | $spl32),
  inference(cnf_transformation,[],[f76_D])).
fof(f76_D,plain,(
  s(bool,'T_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK4),s(sK1,sK5))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f74,plain,(
  s(bool,'T_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK4),s(sK1,sK5))) | ($spl2 | $spl11 | $spl12)),
  inference(subsumption_resolution,[],[f69,f55])).
fof(f69,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK4),s(sK1,sK5))) | ($spl2 | $spl12)),
  inference(superposition,[],[f49,f56])).
fof(f55,plain,(
  ~p(s(bool,'F_0')) | $spl11),
  inference(cnf_transformation,[],[f55_D])).
fof(f55_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f35,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f10])).
fof(f10,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.relation.TC_STRONG_INDUCT_RIGHT|49|',aHL_FALSITY)).
fof(f53,plain,(
  p(s(bool,'T_0')) | $spl8),
  inference(cnf_transformation,[],[f53_D])).
fof(f53_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f34,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.relation.TC_STRONG_INDUCT_RIGHT|49|',aHL_TRUTH)).
fof(f1782,plain,(
  $false | ($spl8 | $spl11 | $spl32 | $spl122 | $spl336)),
  inference(subsumption_resolution,[],[f1781,f53])).
fof(f1781,plain,(
  ~p(s(bool,'T_0')) | ($spl11 | $spl32 | $spl122 | $spl336)),
  inference(forward_demodulation,[],[f1780,f76])).
fof(f1780,plain,(
  ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK4),s(sK1,sK5)))) | ($spl11 | $spl122 | $spl336)),
  inference(subsumption_resolution,[],[f1765,f55])).
fof(f1765,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK4),s(sK1,sK5)))) | ($spl122 | $spl336)),
  inference(superposition,[],[f1754,f519])).
fof(f62,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP0(X0,X2,X1) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK9(X2,X1,X0)))),s(X0,sK10(X2,X1,X0)))))) ) | $spl24),
  inference(cnf_transformation,[],[f62_D])).
fof(f62_D,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP0(X0,X2,X1) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK9(X2,X1,X0)))),s(X0,sK10(X2,X1,X0)))))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f41,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | sP0(X0,X2,X1) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK9(X2,X1,X0)))),s(X0,sK10(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f28])).
fof(f2207,plain,(
  ( ! [X37,X36] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X36),s(sK1,X37)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X36))),s(sK1,X37))))) ) | ($spl24 | $spl381 | $spl421)),
  inference(subsumption_resolution,[],[f2201,f2040])).
fof(f2201,plain,(
  ( ! [X37,X36] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X36),s(sK1,X37)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X36))),s(sK1,X37)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1)))))) ) | ($spl24 | $spl381)),
  inference(resolution,[],[f62,f1923])).
% SZS output end Proof for thm.relation
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 86
% Passive clauses: 240
% Generated clauses: 2271
% Final active clauses: 53
% Final passive clauses: 100
% Input formulas: 6
% Initial clauses: 16
% 
% Duplicate literals: 111
% Trivial inequalities: 8
% Fw subsumption resolutions: 205
% Fw demodulations: 187
% Bw demodulations: 69
% Condensations: 130
% 
% Simple tautologies: 30
% Equational tautologies: 17
% Forward subsumptions: 902
% Bw demodulations to eq. taut.: 2
% 
% Binary resolution: 87
% Forward superposition: 774
% Backward superposition: 271
% Self superposition: 30
% Equality factoring: 32
% 
% Split clauses: 48
% Split components: 97
% SAT solver clauses: 427
% SAT solver unit clauses: 18
% SAT solver binary clauses: 13
% SAT solver learnt clauses: 1
% 
% Sat splits: 48
% Sat splitting refutations: 12
% 
% TWLsolver clauses: 426
% TWLsolver calls for satisfiability: 165
% 
% Memory used [KB]: 2558
% Time elapsed: 0.082 s
% ------------------------------
% ------------------------------
% Success in time 0.696 s
