Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 44
dis+11_7_41 on thm.numeral
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.numeral
% SZS output start Proof for thm.numeral
fof(f187281,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f187263,f127730,f1870,f2864,f186882,f1871,f2867,f2865,f127889,f2866])).
fof(f2866,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK17),s('type.num.num',sK16)))) | $spl3),
  inference(cnf_transformation,[],[f2866_D])).
fof(f2866_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK17),s('type.num.num',sK16)))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f127889,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK17),s('type.num.num',sK16)))) | $spl800),
  inference(resolution,[],[f127730,f2478])).
fof(f2478,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f1717])).
fof(f1717,plain,(
  ! [X0,X1] : ((p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) & (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(nnf_transformation,[],[f967])).
fof(f967,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(flattening,[],[f966])).
fof(f966,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(rectify,[],[f80])).
fof(f80,axiom,(
  ! [X8,X9] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X8),s('type.num.num',X9)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X9),s('type.num.num',X8)))))),
  file('thm.numeral.numeral_|49||48|_prove_ltb|split|3','thm.arithmetic.NOT_LESS')).
fof(f2865,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK16),s('type.num.num',sK16))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK17),s('type.num.num',sK17)))))))) | $spl1),
  inference(cnf_transformation,[],[f2865_D])).
fof(f2865_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK16),s('type.num.num',sK16))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK17),s('type.num.num',sK17)))))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f2867,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK17),s('type.num.num',sK16)))) | $spl2),
  inference(cnf_transformation,[],[f2867_D])).
fof(f2867_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK17),s('type.num.num',sK16)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f1871,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK17),s('type.num.num',sK16)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK16),s('type.num.num',sK16))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK17),s('type.num.num',sK17))))))))),
  inference(cnf_transformation,[],[f1437])).
fof(f1437,plain,(
  (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK16),s('type.num.num',sK16))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK17),s('type.num.num',sK17)))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK17),s('type.num.num',sK16))))) & (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK16),s('type.num.num',sK16))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK17),s('type.num.num',sK17)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK17),s('type.num.num',sK16)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK16,sK17])],[f1436])).
fof(f1436,plain,(
  ? [X0,X1] : ((p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X1)))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))))) & (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))))))),
  inference(nnf_transformation,[],[f1168])).
fof(f1168,plain,(
  ? [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X1)))))))) <~> ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(ennf_transformation,[],[f518])).
fof(f518,plain,(
  ~! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X1)))))))) <=> ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(flattening,[],[f517])).
fof(f517,plain,(
  ~! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X1)))))))) <=> ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X6))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num',X7)))))))) <=> ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num',X6)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X6))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num',X7)))))))) <=> ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num',X6)))))),
  file('thm.numeral.numeral_|49||48|_prove_ltb|split|3','thm.numeral.numeral_|49||48|_prove_ltb|split|3')).
fof(f186882,plain,(
  $false | ($spl1 | $spl3)),
  inference(subsumption_resolution,[],[f186881,f2866])).
fof(f186881,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK17),s('type.num.num',sK16)))) | $spl1),
  inference(forward_demodulation,[],[f186880,f2263])).
fof(f2263,plain,(
  ( ! [X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X0))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f829])).
fof(f829,plain,(
  ! [X0,X1] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X0))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X1)))))),
  inference(rectify,[],[f463])).
fof(f463,axiom,(
  ! [X9,X8] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num',X8))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X9),s('type.num.num',X9))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num',X8)))))),
  file('thm.numeral.numeral_|49||48|_prove_ltb|split|3','thm.numeral.numeral_9_prove_ltb|split|0')).
fof(f186880,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK17),s('type.num.num',sK17))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK16),s('type.num.num',sK16)))))) | $spl1),
  inference(forward_demodulation,[],[f186830,f2236])).
fof(f2236,plain,(
  ( ! [X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f802])).
fof(f802,plain,(
  ! [X0,X1] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))),
  inference(rectify,[],[f85])).
fof(f85,axiom,(
  ! [X8,X9] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X8),s('type.num.num',X9))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.num.num',X9)))),
  file('thm.numeral.numeral_|49||48|_prove_ltb|split|3','thm.arithmetic.LESS_EQ')).
fof(f186830,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK17),s('type.num.num',sK17))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK16),s('type.num.num',sK16)))))) | $spl1),
  inference(resolution,[],[f2865,f2285])).
fof(f2285,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f851])).
fof(f851,plain,(
  ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(rectify,[],[f346])).
fof(f346,axiom,(
  ! [X8,X9] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X8),s('type.num.num',X9)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X9),s('type.num.num',X8)))))),
  file('thm.numeral.numeral_|49||48|_prove_ltb|split|3','thm.arithmetic.LESS_EQ_CASES')).
fof(f2864,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK16),s('type.num.num',sK16))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK17),s('type.num.num',sK17)))))))) | $spl0),
  inference(cnf_transformation,[],[f2864_D])).
fof(f2864_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK16),s('type.num.num',sK16))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK17),s('type.num.num',sK17)))))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f1870,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK17),s('type.num.num',sK16)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK16),s('type.num.num',sK16))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK17),s('type.num.num',sK17))))))))),
  inference(cnf_transformation,[],[f1437])).
fof(f127730,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK17)))) | $spl800),
  inference(cnf_transformation,[],[f127730_D])).
fof(f127730_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK17)))) <=> ~$spl800),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl800])])).
fof(f187263,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK17)))) | $spl0),
  inference(forward_demodulation,[],[f187262,f2264])).
fof(f2264,plain,(
  ( ! [X0,X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X0))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f830])).
fof(f830,plain,(
  ! [X0,X1] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X0))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X1)))))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X9,X8] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X9),s('type.num.num',X8))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X9),s('type.num.num',X9))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num',X8)))))),
  file('thm.numeral.numeral_|49||48|_prove_ltb|split|3','thm.numeral.numeral_9_prove_ltb|split|1')).
fof(f187262,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK16),s('type.num.num',sK16))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK17),s('type.num.num',sK17)))))) | $spl0),
  inference(subsumption_resolution,[],[f187149,f2499])).
fof(f2499,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X1))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f987])).
fof(f987,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X1))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X0)))))),
  inference(flattening,[],[f986])).
fof(f986,plain,(
  ! [X0,X1] : ~s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X1))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X0)))))),
  inference(rectify,[],[f95])).
fof(f95,axiom,(
  ! [X9,X8] : ~s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num',X8))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X9),s('type.num.num',X9)))))),
  file('thm.numeral.numeral_|49||48|_prove_ltb|split|3','thm.arithmetic.NOT_ODD_EQ_EVEN')).
fof(f187149,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK16),s('type.num.num',sK16))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK17),s('type.num.num',sK17))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK16),s('type.num.num',sK16))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK17),s('type.num.num',sK17)))))) | $spl0),
  inference(resolution,[],[f2864,f2400])).
fof(f2400,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f1675])).
fof(f1675,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))) & ((s('type.num.num',X0) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) & ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))))))),
  inference(flattening,[],[f1674])).
fof(f1674,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))) | (s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))))) & ((s('type.num.num',X0) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) & ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))))))),
  inference(nnf_transformation,[],[f941])).
fof(f941,plain,(
  ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))) <=> (s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X8,X9] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X8),s('type.num.num','const.num.SUC_1'(s('type.num.num',X9)))))) <=> (s('type.num.num',X8) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X8),s('type.num.num',X9))))))),
  file('thm.numeral.numeral_|49||48|_prove_ltb|split|3','thm.arithmetic.arithmetic_2|49|7_prove_ltb')).
% SZS output end Proof for thm.numeral
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 2826
% Passive clauses: 42717
% Generated clauses: 185263
% Final active clauses: 2578
% Final passive clauses: 37099
% Input formulas: 515
% Initial clauses: 981
% 
% Duplicate literals: 289
% Trivial inequalities: 625
% Fw subsumption resolutions: 6191
% Fw demodulations: 38867
% Bw demodulations: 217
% 
% Simple tautologies: 603
% Equational tautologies: 56
% Forward subsumptions: 94852
% Fw demodulations to eq. taut.: 73
% 
% Binary resolution: 29894
% Factoring: 749
% Forward superposition: 52168
% Backward superposition: 54494
% Self superposition: 113
% Equality factoring: 2
% Equality resolution: 50
% 
% Split clauses: 1621
% Split components: 3244
% SAT solver clauses: 2170
% SAT solver unit clauses: 55
% SAT solver binary clauses: 288
% SAT solver learnt clauses: 2
% SAT solver learnt literals: 1
% 
% Sat splits: 1621
% Sat splitting refutations: 24
% 
% TWLsolver clauses: 2166
% TWLsolver calls for satisfiability: 872
% 
% Memory used [KB]: 105158
% Time elapsed: 3.824 s
% ------------------------------
% ------------------------------
% Success in time 3.87 s
