Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.pred_set
% (99766)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1161
% Passive clauses: 7086
% Generated clauses: 39048
% Final active clauses: 1103
% Final passive clauses: 5096
% Input formulas: 515
% Initial clauses: 912
% 
% Duplicate literals: 117
% Trivial inequalities: 48
% Fw subsumption resolutions: 1110
% Fw demodulations: 2795
% Bw demodulations: 57
% 
% Simple tautologies: 333
% Equational tautologies: 62
% Forward subsumptions: 26424
% Fw demodulations to eq. taut.: 4
% 
% Binary resolution: 14017
% Factoring: 348
% Forward superposition: 9277
% Backward superposition: 9876
% Self superposition: 45
% Equality factoring: 2
% Equality resolution: 5
% 
% Split clauses: 1038
% Split components: 2078
% SAT solver clauses: 1844
% SAT solver unit clauses: 17
% SAT solver binary clauses: 240
% 
% Sat splits: 1038
% Sat splitting refutations: 8
% 
% TWLsolver clauses: 1844
% TWLsolver calls for satisfiability: 620
% 
% Memory used [KB]: 19829
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.pred_set
% (99778)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 602
% Passive clauses: 4217
% Generated clauses: 27894
% Final active clauses: 583
% Final passive clauses: 3460
% Input formulas: 515
% Initial clauses: 912
% 
% Duplicate literals: 90
% Trivial inequalities: 30
% Fw subsumption resolutions: 874
% Fw demodulations: 3992
% Bw demodulations: 172
% Condensations: 71
% Global subsumptions: 40
% 
% Simple tautologies: 336
% Equational tautologies: 15
% Forward subsumptions: 13245
% Fw demodulations to eq. taut.: 9
% Bw demodulations to eq. taut.: 2
% 
% Binary resolution: 3085
% Factoring: 80
% Forward superposition: 8495
% Backward superposition: 5660
% Self superposition: 69
% Equality factoring: 2
% Equality resolution: 11
% 
% Split clauses: 24
% Split components: 48
% SAT solver clauses: 5970
% SAT solver unit clauses: 1233
% SAT solver binary clauses: 264
% SAT solver learnt clauses: 80
% SAT solver learnt literals: 21
% 
% Sat splits: 24
% 
% TWLsolver clauses: 5811
% TWLsolver calls for satisfiability: 5241
% 
% Memory used [KB]: 26225
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f65125,plain,(
  $false),
  inference(subsumption_resolution,[],[f65119,f53001])).
fof(f53001,plain,(
  ( ! [X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),'=_1'(s(X0,X1))))))) )),
  inference(forward_demodulation,[],[f52623,f1976])).
fof(f1976,plain,(
  ( ! [X0,X1] : (s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) )),
  inference(cnf_transformation,[],[f734])).
fof(f734,plain,(
  ! [X0,X1] : s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f733])).
fof(f733,plain,(
  ! [X0] : ! [X1] : s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f441])).
fof(f441,axiom,(
  ! [X6] : ! [X5] : s(bool,'|dollar|exists_1'(s(fun(X6,bool),X5))) = s(bool,happ(s(fun(X6,bool),X5),s(X6,'const.min.@_1'(s(fun(X6,bool),X5)))))),
  file('thm.pred_set.chooser_def_compute|split|2','thm.bool.EXISTS_DEF')).
fof(f52623,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),'=_1'(s(X0,X1))),s(X0,'const.min.@_1'(s(fun(X0,bool),'=_1'(s(X0,X1))))))))) )),
  inference(unit_resulting_resolution,[],[f1968,f2434])).
fof(f2434,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),'=_1'(s(X0,X1))),s(X0,X2)))) | s(X0,X1) != s(X0,X2)) )),
  inference(cnf_transformation,[],[f1638])).
fof(f1638,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X0,bool),'=_1'(s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),'=_1'(s(X0,X1))),s(X0,X2))))))),
  inference(nnf_transformation,[],[f1119])).
fof(f1119,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),'=_1'(s(X0,X1))),s(X0,X2)))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f510])).
fof(f510,axiom,(
  ! [X55,X56,X57] : (p(s(bool,happ(s(fun(X55,bool),'=_1'(s(X55,X56))),s(X55,X57)))) <=> s(X55,X56) = s(X55,X57))),
  file('thm.pred_set.chooser_def_compute|split|2',arity_1_HL_EQ)).
fof(f1968,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,'const.min.@_1'(s(fun(X0,bool),'=_1'(s(X0,X1)))))) )),
  inference(cnf_transformation,[],[f718])).
fof(f718,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,'const.min.@_1'(s(fun(X0,bool),'=_1'(s(X0,X1)))))),
  inference(flattening,[],[f717])).
fof(f717,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,'const.min.@_1'(s(fun(X0,bool),'=_1'(s(X0,X1)))))),
  inference(rectify,[],[f512])).
fof(f512,axiom,(
  ! [X6] : ! [X15] : s(X6,X15) = s(X6,'const.min.@_1'(s(fun(X6,bool),'=_1'(s(X6,X15)))))),
  file('thm.pred_set.chooser_def_compute|split|2','thm.bool.SELECT_REFL_2')).
fof(f65119,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK17,bool),'=_1'(s(sK17,'const.pred_set.chooser_2'(s(fun(sK17,bool),sK18),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK19))))))))))),
  inference(backward_demodulation,[],[f64987,f52389])).
fof(f52389,plain,(
  ~p(s(bool,happ(s(fun(sK17,bool),'=_1'(s(sK17,'const.pred_set.chooser_2'(s(fun(sK17,bool),sK18),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK19))))))),s(sK17,'const.pred_set.chooser_2'(s(fun(sK17,bool),sK18),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK19))))))))),
  inference(forward_demodulation,[],[f52388,f1778])).
fof(f1778,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f566])).
fof(f566,plain,(
  ! [X0] : s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f327])).
fof(f327,axiom,(
  ! [X8] : s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X8))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X8)))))),
  file('thm.pred_set.chooser_def_compute|split|2','thm.numeral.numeral_suc|split|1')).
fof(f52388,plain,(
  ~p(s(bool,happ(s(fun(sK17,bool),'=_1'(s(sK17,'const.pred_set.chooser_2'(s(fun(sK17,bool),sK18),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK19))))))))),s(sK17,'const.pred_set.chooser_2'(s(fun(sK17,bool),sK18),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK19))))))))),
  inference(forward_demodulation,[],[f52214,f2028])).
fof(f2028,plain,(
  ( ! [X2,X0,X1] : (s(X0,'const.pred_set.chooser_2'(s(fun(X0,bool),X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))))) = s(X0,'const.pred_set.chooser_2'(s(fun(X0,bool),'const.pred_set.REST_1'(s(fun(X0,bool),X1))),s('type.num.num',X2)))) )),
  inference(cnf_transformation,[],[f802])).
fof(f802,plain,(
  ! [X0,X1,X2] : s(X0,'const.pred_set.chooser_2'(s(fun(X0,bool),X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))))) = s(X0,'const.pred_set.chooser_2'(s(fun(X0,bool),'const.pred_set.REST_1'(s(fun(X0,bool),X1))),s('type.num.num',X2)))),
  inference(flattening,[],[f801])).
fof(f801,plain,(
  ! [X0] : ! [X1,X2] : s(X0,'const.pred_set.chooser_2'(s(fun(X0,bool),X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))))) = s(X0,'const.pred_set.chooser_2'(s(fun(X0,bool),'const.pred_set.REST_1'(s(fun(X0,bool),X1))),s('type.num.num',X2)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : ! [X7,X8] : s(X6,'const.pred_set.chooser_2'(s(fun(X6,bool),X7),s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))))) = s(X6,'const.pred_set.chooser_2'(s(fun(X6,bool),'const.pred_set.REST_1'(s(fun(X6,bool),X7))),s('type.num.num',X8)))),
  file('thm.pred_set.chooser_def_compute|split|2','thm.pred_set.chooser_def|split|1')).
fof(f52214,plain,(
  ~p(s(bool,happ(s(fun(sK17,bool),'=_1'(s(sK17,'const.pred_set.chooser_2'(s(fun(sK17,bool),'const.pred_set.REST_1'(s(fun(sK17,bool),sK18))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK19))))))),s(sK17,'const.pred_set.chooser_2'(s(fun(sK17,bool),sK18),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK19))))))))),
  inference(unit_resulting_resolution,[],[f2680,f2433])).
fof(f2433,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),'=_1'(s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) )),
  inference(cnf_transformation,[],[f1638])).
fof(f2680,plain,(
  s(sK17,'const.pred_set.chooser_2'(s(fun(sK17,bool),sK18),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK19))))) != s(sK17,'const.pred_set.chooser_2'(s(fun(sK17,bool),'const.pred_set.REST_1'(s(fun(sK17,bool),sK18))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK19)))))),
  inference(forward_demodulation,[],[f2676,f1739])).
fof(f1739,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f527])).
fof(f527,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X15] : s('type.num.num',X15) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X15)))),
  file('thm.pred_set.chooser_def_compute|split|2','thm.arithmetic.NUMERAL_DEF')).
fof(f2676,plain,(
  s(sK17,'const.pred_set.chooser_2'(s(fun(sK17,bool),sK18),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK19))))))) != s(sK17,'const.pred_set.chooser_2'(s(fun(sK17,bool),'const.pred_set.REST_1'(s(fun(sK17,bool),sK18))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK19)))))),
  inference(backward_demodulation,[],[f1739,f1710])).
fof(f1710,plain,(
  s(sK17,'const.pred_set.chooser_2'(s(fun(sK17,bool),sK18),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK19))))))) != s(sK17,'const.pred_set.chooser_2'(s(fun(sK17,bool),'const.pred_set.REST_1'(s(fun(sK17,bool),sK18))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK19)))))))),
  inference(cnf_transformation,[],[f1376])).
fof(f1376,plain,(
  s(sK17,'const.pred_set.chooser_2'(s(fun(sK17,bool),sK18),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK19))))))) != s(sK17,'const.pred_set.chooser_2'(s(fun(sK17,bool),'const.pred_set.REST_1'(s(fun(sK17,bool),sK18))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK19)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK17,sK18,sK19])],[f1224])).
fof(f1224,plain,(
  ? [X0,X1,X2] : s(X0,'const.pred_set.chooser_2'(s(fun(X0,bool),X1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X2))))))) != s(X0,'const.pred_set.chooser_2'(s(fun(X0,bool),'const.pred_set.REST_1'(s(fun(X0,bool),X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X2)))))))),
  inference(ennf_transformation,[],[f518])).
fof(f518,plain,(
  ~! [X0,X1,X2] : s(X0,'const.pred_set.chooser_2'(s(fun(X0,bool),X1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X2))))))) = s(X0,'const.pred_set.chooser_2'(s(fun(X0,bool),'const.pred_set.REST_1'(s(fun(X0,bool),X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X2)))))))),
  inference(flattening,[],[f517])).
fof(f517,plain,(
  ~! [X0] : ! [X1,X2] : s(X0,'const.pred_set.chooser_2'(s(fun(X0,bool),X1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X2))))))) = s(X0,'const.pred_set.chooser_2'(s(fun(X0,bool),'const.pred_set.REST_1'(s(fun(X0,bool),X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X2)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : s(X6,'const.pred_set.chooser_2'(s(fun(X6,bool),X7),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X8))))))) = s(X6,'const.pred_set.chooser_2'(s(fun(X6,bool),'const.pred_set.REST_1'(s(fun(X6,bool),X7))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X8)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : s(X6,'const.pred_set.chooser_2'(s(fun(X6,bool),X7),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X8))))))) = s(X6,'const.pred_set.chooser_2'(s(fun(X6,bool),'const.pred_set.REST_1'(s(fun(X6,bool),X7))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X8)))))))),
  file('thm.pred_set.chooser_def_compute|split|2','thm.pred_set.chooser_def_compute|split|2')).
fof(f64987,plain,(
  ( ! [X0,X1] : (s(bool,'|dollar|exists_1'(s(fun(X0,bool),'=_1'(s(X0,X1))))) = s(bool,happ(s(fun(X0,bool),'=_1'(s(X0,X1))),s(X0,X1)))) )),
  inference(superposition,[],[f1976,f1968])).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 950
% Passive clauses: 8460
% Generated clauses: 64049
% Final active clauses: 882
% Final passive clauses: 5373
% Input formulas: 515
% Initial clauses: 912
% 
% Duplicate literals: 13
% Trivial inequalities: 26
% Fw subsumption resolutions: 2050
% Fw demodulations: 14318
% Bw demodulations: 932
% Condensations: 55
% 
% Simple tautologies: 159
% Equational tautologies: 34
% Forward subsumptions: 36708
% Fw demodulations to eq. taut.: 545
% Bw demodulations to eq. taut.: 19
% 
% Unit resulting resolution: 19789
% Factoring: 1
% Forward superposition: 14133
% Backward superposition: 11072
% Self superposition: 85
% Equality factoring: 2
% Equality resolution: 9
% 
% Split clauses: 4
% Split components: 8
% SAT solver clauses: 10019
% SAT solver unit clauses: 6826
% SAT solver binary clauses: 2724
% 
% Sat splits: 4
% Sat splitting refutations: 1549
% 
% Memory used [KB]: 32366
% Time elapsed: 1.302 s
% ------------------------------
% ------------------------------
% Success in time 2.63 s
