Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f2143,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f1299,f918,f26,f39,f1290,f561,f27,f41,f37,f1097,f28,f42,f40,f38,f385,f327,f332,f330,f2138,f2007,f50,f54,f52,f1162])).
fof(f1162,plain,(
  $false | $spl6),
  inference(resolution,[],[f1136,f29])).
fof(f29,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pred_set.DISJOINT_INSERT|39|',aHL_TRUTH)).
fof(f1136,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl6),
  inference(superposition,[],[f30,f52])).
fof(f30,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f11])).
fof(f11,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.pred_set.DISJOINT_INSERT|39|',aHL_FALSITY)).
fof(f52,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl6),
  inference(cnf_transformation,[],[f52_D])).
fof(f52_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f54,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl9),
  inference(cnf_transformation,[],[f54_D])).
fof(f54_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f50,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f31])).
fof(f31,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.pred_set.DISJOINT_INSERT|39|',aHL_BOOL_CASES)).
fof(f2007,plain,(
  $false | ($spl3 | $spl9 | $spl28)),
  inference(subsumption_resolution,[],[f2006,f30])).
fof(f2006,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl9 | $spl28)),
  inference(forward_demodulation,[],[f2005,f1283])).
fof(f1283,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))) | ($spl3 | $spl9)),
  inference(subsumption_resolution,[],[f1269,f54])).
fof(f1269,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))) | $spl3),
  inference(superposition,[],[f1196,f245])).
fof(f245,plain,(
  ( ! [X24,X23,X22] : (s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X22,bool),X24),s(fun(X22,bool),X23))) | s(bool,'T_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X22,bool),X23),s(fun(X22,bool),X24)))) )),
  inference(superposition,[],[f32,f31])).
fof(f32,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))) )),
  inference(cnf_transformation,[],[f13])).
fof(f13,plain,(
  ! [X0,X1,X2] : s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))),
  inference(flattening,[],[f12])).
fof(f12,plain,(
  ! [X0] : ! [X1,X2] : s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6] : ! [X8,X9] : s(bool,'const.pred_set.DISJOINT_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X9))) = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X6,bool),X9),s(fun(X6,bool),X8)))),
  file('thm.pred_set.DISJOINT_INSERT|39|','thm.pred_set.DISJOINT_SYM')).
fof(f1196,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),sK2))) | $spl3),
  inference(subsumption_resolution,[],[f1188,f29])).
fof(f1188,plain,(
  ~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),sK2))) | $spl3),
  inference(superposition,[],[f1112,f245])).
fof(f1112,plain,(
  ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))) | $spl3),
  inference(forward_demodulation,[],[f40,f32])).
fof(f2005,plain,(
  p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))) | $spl28),
  inference(subsumption_resolution,[],[f1989,f29])).
fof(f1989,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))) | $spl28),
  inference(superposition,[],[f33,f918])).
fof(f33,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X3)))) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f24])).
fof(f24,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X3)))) | (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3)))))) & (~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3)))) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X3))))))),
  inference(flattening,[],[f23])).
fof(f23,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X3)))) | (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3)))))) & ((~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X3))))))),
  inference(nnf_transformation,[],[f15])).
fof(f15,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X3)))) <=> (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))))),
  inference(flattening,[],[f14])).
fof(f14,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X3)))) <=> (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : ! [X7,X8,X9] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X6,bool),'const.pred_set.INSERT_2'(s(X6,X7),s(fun(X6,bool),X8))),s(fun(X6,bool),X9)))) <=> (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X9)))) & ~p(s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),X9))))))),
  file('thm.pred_set.DISJOINT_INSERT|39|','thm.pred_set.DISJOINT_INSERT')).
fof(f2138,plain,(
  $false | ($spl4 | $spl22)),
  inference(subsumption_resolution,[],[f2129,f29])).
fof(f2129,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl22)),
  inference(superposition,[],[f2024,f561])).
fof(f2024,plain,(
  ( ! [X0] : (~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),X0))))))) ) | $spl4),
  inference(forward_demodulation,[],[f2017,f32])).
fof(f2017,plain,(
  ( ! [X0] : (~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),X0))),s(fun(sK0,bool),sK3))))) ) | $spl4),
  inference(resolution,[],[f42,f34])).
fof(f34,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3)))) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f24])).
fof(f330,plain,(
  ( ! [X9] : (s(bool,'T_0') = s(bool,X9)) ) | $spl10),
  inference(cnf_transformation,[],[f330_D])).
fof(f330_D,plain,(
  ( ! [X9] : (s(bool,'T_0') = s(bool,X9)) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f332,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))) | $spl12),
  inference(cnf_transformation,[],[f332_D])).
fof(f332_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f327,plain,(
  ( ! [X9] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))) | s(bool,'T_0') = s(bool,X9)) ) | $spl5),
  inference(duplicate_literal_removal,[],[f284])).
fof(f284,plain,(
  ( ! [X9] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) ) | $spl5),
  inference(superposition,[],[f216,f31])).
fof(f216,plain,(
  ( ! [X9] : (s(bool,X9) = s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))) | s(bool,'T_0') = s(bool,X9)) ) | $spl5),
  inference(subsumption_resolution,[],[f142,f29])).
fof(f142,plain,(
  ( ! [X9] : (~p(s(bool,'T_0')) | s(bool,X9) = s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))) | s(bool,'T_0') = s(bool,X9)) ) | $spl5),
  inference(superposition,[],[f41,f43])).
fof(f43,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f31,f31])).
fof(f385,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f336,f29])).
fof(f336,plain,(
  ~p(s(bool,'T_0')) | $spl10),
  inference(backward_demodulation,[],[f330,f30])).
fof(f38,plain,(
  ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)))))) | $spl1),
  inference(cnf_transformation,[],[f38_D])).
fof(f38_D,plain,(
  ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f40,plain,(
  ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),sK2)))) | $spl3),
  inference(cnf_transformation,[],[f40_D])).
fof(f40_D,plain,(
  ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),sK2)))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f42,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)))) | $spl4),
  inference(cnf_transformation,[],[f42_D])).
fof(f42_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f28,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)))) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),sK2)))) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))))))),
  inference(cnf_transformation,[],[f22])).
fof(f22,plain,(
  (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)))))) | (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),sK2)))) & ~p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)))))) & (~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),sK2)))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f21])).
fof(f21,plain,(
  ? [X0,X1,X2,X3] : ((p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X3),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))))) | (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X3),s(fun(X0,bool),X2)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3)))))) & (~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X3),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))))) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X3),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))))),
  inference(flattening,[],[f20])).
fof(f20,plain,(
  ? [X0,X1,X2,X3] : ((p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X3),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))))) | (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X3),s(fun(X0,bool),X2)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3)))))) & (~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X3),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))))) | (~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X3),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3)))))))),
  inference(nnf_transformation,[],[f18])).
fof(f18,plain,(
  ? [X0,X1,X2,X3] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X3),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))))) <~> (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X3),s(fun(X0,bool),X2)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))))),
  inference(ennf_transformation,[],[f10])).
fof(f10,plain,(
  ~! [X0,X1,X2,X3] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X3),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))))) <=> (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X3),s(fun(X0,bool),X2)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))))),
  inference(flattening,[],[f9])).
fof(f9,plain,(
  ~! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X3),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))))) <=> (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X3),s(fun(X0,bool),X2)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X6,bool),X9),s(fun(X6,bool),'const.pred_set.INSERT_2'(s(X6,X7),s(fun(X6,bool),X8)))))) <=> (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X6,bool),X9),s(fun(X6,bool),X8)))) & ~p(s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),X9))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X6,bool),X9),s(fun(X6,bool),'const.pred_set.INSERT_2'(s(X6,X7),s(fun(X6,bool),X8)))))) <=> (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X6,bool),X9),s(fun(X6,bool),X8)))) & ~p(s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),X9))))))),
  file('thm.pred_set.DISJOINT_INSERT|39|','thm.pred_set.DISJOINT_INSERT|39|')).
fof(f1097,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl1 | $spl2 | $spl12)),
  inference(subsumption_resolution,[],[f1096,f30])).
fof(f1096,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,X0)) ) | ($spl1 | $spl2 | $spl12)),
  inference(forward_demodulation,[],[f1095,f332])).
fof(f1095,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f1080,f45])).
fof(f45,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f30,f31])).
fof(f1080,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f1043,f238])).
fof(f238,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl1),
  inference(subsumption_resolution,[],[f233,f29])).
fof(f233,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl1),
  inference(superposition,[],[f38,f43])).
fof(f1043,plain,(
  ( ! [X4] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X4),s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.bool.IN_2'(s(sK0,X4),s(fun(sK0,bool),sK3))))) ) | $spl2),
  inference(forward_demodulation,[],[f1042,f32])).
fof(f1042,plain,(
  ( ! [X4] : (p(s(bool,'const.bool.IN_2'(s(sK0,X4),s(fun(sK0,bool),sK3)))) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X4),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3))))) ) | $spl2),
  inference(subsumption_resolution,[],[f1025,f29])).
fof(f1025,plain,(
  ( ! [X4] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,X4),s(fun(sK0,bool),sK3)))) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X4),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3))))) ) | $spl2),
  inference(superposition,[],[f35,f240])).
fof(f240,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))) | $spl2),
  inference(superposition,[],[f32,f51])).
fof(f51,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),sK2))) | $spl2),
  inference(subsumption_resolution,[],[f47,f30])).
fof(f47,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),sK2))) | $spl2),
  inference(superposition,[],[f39,f31])).
fof(f35,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3)))) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f24])).
fof(f37,plain,(
  p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)))))) | $spl0),
  inference(cnf_transformation,[],[f37_D])).
fof(f37_D,plain,(
  p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f41,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)))) | $spl5),
  inference(cnf_transformation,[],[f41_D])).
fof(f41_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f27,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)))) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))))))),
  inference(cnf_transformation,[],[f22])).
fof(f561,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))))) | $spl22),
  inference(cnf_transformation,[],[f561_D])).
fof(f561_D,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f1290,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))))) | $spl0),
  inference(resolution,[],[f37,f45])).
fof(f39,plain,(
  p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),sK2)))) | $spl2),
  inference(cnf_transformation,[],[f39_D])).
fof(f39_D,plain,(
  p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),sK2)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f26,plain,(
  p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),sK2)))) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))))))),
  inference(cnf_transformation,[],[f22])).
fof(f918,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3))) | $spl28),
  inference(cnf_transformation,[],[f918_D])).
fof(f918_D,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3))) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f1299,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3))) | $spl0),
  inference(subsumption_resolution,[],[f1292,f30])).
fof(f1292,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3))) | $spl0),
  inference(superposition,[],[f37,f245])).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 73
% Passive clauses: 352
% Generated clauses: 2165
% Final active clauses: 20
% Final passive clauses: 121
% Input formulas: 7
% Initial clauses: 11
% 
% Duplicate literals: 176
% Trivial inequalities: 29
% Fw subsumption resolutions: 166
% Fw demodulations: 162
% Bw demodulations: 23
% 
% Simple tautologies: 47
% Equational tautologies: 20
% Forward subsumptions: 1121
% 
% Binary resolution: 31
% Forward superposition: 797
% Backward superposition: 579
% Self superposition: 52
% Equality factoring: 39
% 
% Split clauses: 44
% Split components: 89
% SAT solver clauses: 150
% SAT solver unit clauses: 19
% SAT solver binary clauses: 46
% SAT solver learnt clauses: 3
% SAT solver learnt literals: 3
% 
% Sat splits: 44
% Sat splitting refutations: 17
% 
% TWLsolver clauses: 147
% TWLsolver calls for satisfiability: 35
% 
% Memory used [KB]: 1407
% Time elapsed: 0.058 s
% ------------------------------
% ------------------------------
% Success in time 0.06 s
