Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.util_prob
% (57727)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 627
% Passive clauses: 9469
% Generated clauses: 18412
% Final active clauses: 620
% Final passive clauses: 8467
% Input formulas: 37
% Initial clauses: 159
% 
% Duplicate literals: 337
% Trivial inequalities: 72
% Fw subsumption resolutions: 1640
% Fw demodulations: 149
% Bw demodulations: 4
% 
% Simple tautologies: 191
% Equational tautologies: 55
% Forward subsumptions: 5908
% 
% Binary resolution: 11023
% Factoring: 1146
% Forward superposition: 2051
% Backward superposition: 1621
% Self superposition: 24
% Equality factoring: 28
% Equality resolution: 24
% 
% Split clauses: 606
% Split components: 1222
% SAT solver clauses: 974
% SAT solver unit clauses: 25
% SAT solver binary clauses: 218
% 
% Sat splits: 606
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 973
% TWLsolver calls for satisfiability: 382
% 
% Memory used [KB]: 10746
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.util_prob
% (57757)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 410
% Passive clauses: 2971
% Generated clauses: 10723
% Final active clauses: 405
% Final passive clauses: 2556
% Input formulas: 37
% Initial clauses: 159
% 
% Duplicate literals: 170
% Trivial inequalities: 39
% Fw subsumption resolutions: 795
% Fw demodulations: 301
% Bw demodulations: 10
% Condensations: 197
% Global subsumptions: 13
% 
% Simple tautologies: 211
% Equational tautologies: 16
% Forward subsumptions: 2627
% 
% Binary resolution: 3834
% Factoring: 198
% Forward superposition: 1198
% Backward superposition: 773
% Self superposition: 22
% Equality factoring: 26
% Equality resolution: 6
% 
% Split clauses: 58
% Split components: 121
% SAT solver clauses: 3564
% SAT solver unit clauses: 165
% SAT solver binary clauses: 73
% SAT solver learnt clauses: 26
% SAT solver learnt literals: 10
% 
% Sat splits: 58
% 
% TWLsolver clauses: 3512
% TWLsolver calls for satisfiability: 1516
% 
% Memory used [KB]: 14200
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.util_prob
% (57771)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1574
% Passive clauses: 6748
% Generated clauses: 111510
% Final active clauses: 205
% Final passive clauses: 786
% Input formulas: 37
% Initial clauses: 159
% 
% Duplicate literals: 7642
% Trivial inequalities: 274
% Fw subsumption resolutions: 5868
% Fw demodulations: 15152
% Bw demodulations: 1934
% Condensations: 10349
% Global subsumptions: 1956
% 
% Simple tautologies: 401
% Equational tautologies: 1554
% Forward subsumptions: 58111
% Bw demodulations to eq. taut.: 213
% 
% Unit resulting resolution: 17797
% Factoring: 5
% Forward superposition: 15928
% Backward superposition: 26598
% Self superposition: 420
% Equality factoring: 364
% Equality resolution: 10
% 
% Split clauses: 1080
% Split components: 2181
% SAT solver clauses: 13166
% SAT solver unit clauses: 1181
% SAT solver binary clauses: 2593
% 
% Sat splits: 1080
% Sat splitting refutations: 2300
% 
% Memory used [KB]: 102343
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.util_prob
% (58033)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 2398
% Passive clauses: 32333
% Generated clauses: 216412
% Final active clauses: 634
% Final passive clauses: 6769
% Input formulas: 37
% Initial clauses: 154
% 
% Selected by SInE selection: 35
% SInE iterations: 1
% 
% Duplicate literals: 8225
% Trivial inequalities: 1036
% Fw subsumption resolutions: 31120
% Fw demodulations: 15895
% Bw demodulations: 126
% 
% Simple tautologies: 10871
% Equational tautologies: 967
% Forward subsumptions: 101397
% Bw demodulations to eq. taut.: 25
% 
% Binary resolution: 136170
% Factoring: 226
% Forward superposition: 10729
% Backward superposition: 7648
% Self superposition: 8
% Equality factoring: 2
% Equality resolution: 47
% 
% Split clauses: 16740
% Split components: 38809
% SAT solver clauses: 18429
% SAT solver unit clauses: 48
% SAT solver binary clauses: 956
% 
% Sat splits: 16740
% Sat splitting refutations: 122
% 
% Memory used [KB]: 55777
% Time elapsed: 4.0000 s
% ------------------------------
% ------------------------------
% remaining time: 142 next slice time: 6
lrs+1002_2:3_br=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:sas=minisat:stl=30:sac=on:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none:sp=reverse_arity:urr=on_4 on thm.util_prob
% (58145)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 326
% Passive clauses: 2482
% Generated clauses: 13491
% Final active clauses: 320
% Final passive clauses: 2078
% Input formulas: 37
% Initial clauses: 159
% 
% Duplicate literals: 538
% Trivial inequalities: 225
% Fw subsumption resolutions: 2157
% Fw demodulations: 3
% Bw demodulations: 26
% Global subsumptions: 16
% 
% Simple tautologies: 565
% Equational tautologies: 38
% Forward subsumptions: 7308
% 
% Unit resulting resolution: 1077
% Factoring: 18
% Forward superposition: 4416
% Backward superposition: 4688
% Self superposition: 88
% Equality factoring: 6
% Equality resolution: 22
% 
% Split clauses: 23
% Split components: 47
% SAT solver clauses: 2579
% SAT solver unit clauses: 167
% SAT solver binary clauses: 118
% 
% Sat splits: 23
% Sat splitting refutations: 56
% 
% Memory used [KB]: 7547
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 135 next slice time: 15
lrs+1011_12_bs=on:bsr=unit_only:cond=on:gs=on:gsaa=from_current:gsssp=full:nwc=1.1:sas=minisat:stl=60:sos=all:sac=on:sdd=off:sser=off:sfr=on:ssfp=100000:ssfq=1.2:ssnc=none:sp=reverse_arity:updr=off_14 on thm.util_prob
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.util_prob
% SZS output start Proof for thm.util_prob
fof(f5140,plain,(
  $false),
  inference(subsumption_resolution,[],[f5037,f213])).
fof(f213,plain,(
  s(fun(fun(sK10,bool),bool),'const.pred_set.EMPTY_0') != s(fun(fun(sK10,bool),bool),sK12)),
  inference(cnf_transformation,[],[f139])).
fof(f139,plain,(
  ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X3),s(fun(fun(sK10,bool),bool),sK12)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),X3),s(fun(sK10,bool),sK11))))) & s(fun(fun(sK10,bool),bool),'const.pred_set.EMPTY_0') != s(fun(fun(sK10,bool),bool),sK12) & ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK10,sK11,sK12])],[f95])).
fof(f95,plain,(
  ? [X0,X1,X2] : (! [X3] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),X2)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X3),s(fun(X0,bool),X1))))) & s(fun(fun(X0,bool),bool),'const.pred_set.EMPTY_0') != s(fun(fun(X0,bool),bool),X2) & ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(X0,bool),bool),X2))),s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f94])).
fof(f94,plain,(
  ? [X0,X1,X2] : ((! [X3] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),X2)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X3),s(fun(X0,bool),X1))))) & s(fun(fun(X0,bool),bool),'const.pred_set.EMPTY_0') != s(fun(fun(X0,bool),bool),X2)) & ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(X0,bool),bool),X2))),s(fun(X0,bool),X1)))))),
  inference(ennf_transformation,[],[f40])).
fof(f40,plain,(
  ~! [X0,X1,X2] : ((! [X3] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),X2)))) => p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X3),s(fun(X0,bool),X1))))) & s(fun(fun(X0,bool),bool),'const.pred_set.EMPTY_0') != s(fun(fun(X0,bool),bool),X2)) => p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(X0,bool),bool),X2))),s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f39])).
fof(f39,plain,(
  ~! [X0] : ! [X1,X2] : ((! [X3] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),X2)))) => p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X3),s(fun(X0,bool),X1))))) & ~s(fun(fun(X0,bool),bool),'const.pred_set.EMPTY_0') = s(fun(fun(X0,bool),bool),X2)) => p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(X0,bool),bool),X2))),s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : ((! [X9] : (p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),X9),s(fun(fun(X6,bool),bool),X8)))) => p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X9),s(fun(X6,bool),X7))))) & ~s(fun(fun(X6,bool),bool),'const.pred_set.EMPTY_0') = s(fun(fun(X6,bool),bool),X8)) => p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(X6,bool),bool),X8))),s(fun(X6,bool),X7)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : ((! [X9] : (p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),X9),s(fun(fun(X6,bool),bool),X8)))) => p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X9),s(fun(X6,bool),X7))))) & ~s(fun(fun(X6,bool),bool),'const.pred_set.EMPTY_0') = s(fun(fun(X6,bool),bool),X8)) => p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(X6,bool),bool),X8))),s(fun(X6,bool),X7)))))),
  file('thm.util_prob.BIGINTER_SUBSET','thm.util_prob.BIGINTER_SUBSET')).
fof(f5037,plain,(
  s(fun(fun(sK10,bool),bool),'const.pred_set.EMPTY_0') = s(fun(fun(sK10,bool),bool),sK12)),
  inference(resolution,[],[f5014,f255])).
fof(f255,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,'const.pred_set.CHOICE_1'(s(fun(X0,bool),X1))),s(fun(X0,bool),X1)))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1)) )),
  inference(cnf_transformation,[],[f105])).
fof(f105,plain,(
  ! [X0,X1] : (s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1) | p(s(bool,'const.bool.IN_2'(s(X0,'const.pred_set.CHOICE_1'(s(fun(X0,bool),X1))),s(fun(X0,bool),X1)))))),
  inference(ennf_transformation,[],[f65])).
fof(f65,plain,(
  ! [X0,X1] : (s(fun(X0,bool),'const.pred_set.EMPTY_0') != s(fun(X0,bool),X1) => p(s(bool,'const.bool.IN_2'(s(X0,'const.pred_set.CHOICE_1'(s(fun(X0,bool),X1))),s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f64])).
fof(f64,plain,(
  ! [X0] : ! [X1] : (~s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1) => p(s(bool,'const.bool.IN_2'(s(X0,'const.pred_set.CHOICE_1'(s(fun(X0,bool),X1))),s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X6] : ! [X8] : (~s(fun(X6,bool),'const.pred_set.EMPTY_0') = s(fun(X6,bool),X8) => p(s(bool,'const.bool.IN_2'(s(X6,'const.pred_set.CHOICE_1'(s(fun(X6,bool),X8))),s(fun(X6,bool),X8)))))),
  file('thm.util_prob.BIGINTER_SUBSET','thm.pred_set.CHOICE_DEF')).
fof(f5014,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X0),s(fun(fun(sK10,bool),bool),sK12))))) )),
  inference(global_subsumption,[],[f214,f213,f212,f393,f391,f392,f394,f395,f396,f398,f401,f402,f498,f499,f409,f410,f411,f412,f417,f418,f419,f420,f421,f422,f500,f423,f424,f425,f501,f426,f427,f428,f429,f430,f431,f432,f435,f436,f437,f438,f439,f440,f441,f442,f443,f444,f447,f448,f449,f451,f452,f453,f454,f455,f456,f457,f458,f469,f470,f471,f472,f473,f474,f475,f476,f487,f488,f490,f532,f533,f519,f520,f521,f522,f523,f524,f525,f526,f539,f540,f541,f542,f543,f544,f545,f546,f554,f555,f556,f557,f558,f559,f560,f561,f572,f573,f574,f575,f576,f577,f578,f579,f588,f589,f590,f591,f592,f593,f594,f595,f609,f610,f611,f612,f613,f614,f615,f616,f627,f628,f629,f630,f631,f632,f633,f634,f642,f643,f644,f646,f649,f650,f743,f744,f657,f658,f659,f660,f665,f666,f667,f668,f669,f670,f745,f671,f672,f673,f746,f674,f675,f676,f677,f678,f679,f680,f683,f684,f685,f686,f687,f688,f689,f690,f691,f692,f695,f696,f697,f699,f700,f701,f702,f703,f704,f705,f706,f717,f718,f719,f720,f721,f732,f733,f735,f774,f775,f764,f765,f766,f767,f768,f781,f782,f783,f784,f785,f793,f794,f795,f796,f797,f808,f809,f810,f811,f812,f815,f816,f818,f822,f824,f825,f828,f829,f830,f831,f836,f837,f838,f839,f840,f841,f842,f843,f844,f845,f846,f847,f848,f849,f850,f851,f852,f853,f854,f855,f856,f857,f858,f859,f860,f861,f864,f865,f867,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f879,f880,f881,f887,f888,f889,f890,f891,f892,f893,f904,f905,f906,f918,f919,f920,f921,f922,f936,f937,f938,f939,f940,f951,f952,f953,f954,f955,f963,f964,f965,f1061,f973,f974,f977,f978,f979,f980,f985,f986,f987,f988,f989,f990,f991,f992,f993,f994,f995,f996,f1065,f997,f998,f999,f1000,f1001,f1002,f1003,f1004,f1005,f1006,f1007,f1008,f1009,f1010,f1013,f1014,f1016,f1017,f1018,f1019,f1020,f1021,f1022,f1023,f1024,f1025,f1026,f1028,f1029,f1036,f1037,f1038,f1040,f1041,f1052,f1053,f1054,f1092,f1093,f1094,f1095,f1104,f1105,f1106,f1127,f1128,f1135,f1137,f1148,f1150,f1152,f1153,f1154,f1155,f1156,f1157,f1158,f1159,f1170,f1171,f1172,f1173,f1191,f1195,f1216,f1217,f1244,f1247,f1249,f1260,f1261,f1262,f1263,f1277,f1278,f1279,f1280,f1301,f1309,f1310,f1340,f1342,f1350,f1351,f1352,f1353,f1380,f1367,f1368,f1369,f1370,f1388,f1389,f1390,f1391,f1392,f1393,f1394,f1395,f1403,f1404,f1405,f1406,f1615,f1491,f1492,f1497,f1498,f1592,f1662,f1674,f1675,f1676,f1677,f1695,f1696,f1697,f1698,f1722,f1723,f1725,f1733,f1734,f1735,f1736,f1760,f1747,f1748,f1749,f1750,f1780,f1781,f1782,f1783,f1801,f1802,f1803,f1804,f1837,f1838,f1839,f1840,f1858,f1859,f1860,f1861,f1885,f1886,f1887,f1889,f1890,f1900,f1901,f1902,f1904,f1905,f1922,f1923,f1924,f1946,f1947,f1955,f1969,f1970,f1971,f1972,f1973,f1974,f1975,f1976,f1994,f1996,f1998,f2006,f2007,f2008,f2009,f2040,f2039,f2020,f2021,f2022,f2023,f2058,f2059,f2060,f2061,f2072,f2073,f2074,f2075,f2108,f2109,f2110,f2111,f2122,f2123,f2124,f2125,f2131,f2153,f2168,f2169,f2178,f2180,f2189,f2193,f2194,f2199,f2203,f2204,f2219,f2225,f2226,f2232,f2233,f2234,f2235,f2236,f2237,f2238,f2239,f2246,f2247,f2251,f2252,f2254,f2265,f2322,f2331,f2332,f2333,f2335,f2339,f2340,f2371,f2395,f2403,f2404,f2405,f2406,f2407,f2408,f2409,f2410,f2451,f2452,f2530,f2531,f2532,f2533,f2551,f2552,f2553,f2554,f2598,f2599,f2649,f2667,f2655,f2656,f2657,f2659,f2660,f2673,f2674,f2675,f2677,f2678,f2695,f2696,f2697,f2698,f2699,f2700,f2701,f2702,f2712,f2716,f2717,f2718,f2720,f2721,f2737,f2738,f2739,f2741,f2742,f2750,f2756,f2757,f2758,f2760,f2761,f2773,f2774,f2775,f2777,f2778,f2790,f2791,f2792,f2794,f2795,f2807,f2808,f2809,f2811,f2812,f2818,f2819,f2820,f2822,f2825,f2826,f2918,f2919,f2833,f2834,f2835,f2836,f2841,f2842,f2843,f2844,f2845,f2920,f2846,f2847,f2848,f2849,f2850,f2851,f2852,f2853,f2854,f2855,f2858,f2859,f2860,f2861,f2862,f2863,f2864,f2865,f2866,f2867,f2870,f2871,f2872,f2874,f2875,f2876,f2877,f2878,f2879,f2880,f2881,f2891,f2892,f2893,f2895,f2896,f2907,f2908,f2910,f2935,f2958,f2946,f2947,f2948,f2950,f2951,f2963,f2964,f2965,f2966,f2970,f2973,f2975,f2976,f2977,f2979,f2980,f2994,f2995,f2996,f2998,f2999,f3019,f3009,f3010,f3011,f3013,f3014,f3026,f3027,f3028,f3030,f3031,f3048,f3049,f3050,f3051,f3052,f3053,f3054,f3055,f3065,f3066,f3067,f3069,f3070,f3081,f3082,f3083,f3085,f3086,f3103,f3104,f3105,f3106,f3107,f3108,f3109,f3110,f3125,f3126,f3127,f3128,f3129,f3130,f3131,f3132,f3154,f3155,f3156,f3157,f3158,f3159,f3160,f3161,f3183,f3184,f3185,f3186,f3187,f3188,f3189,f3190,f3209,f3210,f3211,f3212,f3213,f3214,f3215,f3216,f3232,f3233,f3234,f3235,f3236,f3237,f3238,f3239,f3263,f3264,f3265,f3266,f3267,f3268,f3269,f3270,f3278,f3279,f3280,f3281,f3282,f3293,f3294,f3295,f3296,f3308,f3309,f3310,f3311,f3312,f3320,f3321,f3322,f3323,f3406,f3407,f3412,f3413,f3524,f3519,f3594,f3602,f3603,f3604,f3605,f3606,f3612,f3613,f3626,f3627,f3628,f3629,f3630,f3631,f3632,f3633,f3662,f3663,f3664,f3665,f3666,f3680,f3681,f3682,f3683,f3684,f3701,f3702,f3703,f3704,f3705,f3722,f3723,f3724,f3725,f3726,f3739,f3740,f3741,f3742,f3743,f3760,f3761,f3762,f3763,f3782,f3783,f3784,f3785,f3824,f3825,f3826,f3827,f3846,f3847,f3848,f3849,f3878,f3880,f3891,f3892,f3893,f3894,f3905,f3906,f3907,f3908,f3945,f3946,f3947,f3948,f3949,f3971,f3972,f3973,f3974,f3975,f3998,f3999,f4000,f4001,f4002,f4021,f4022,f4023,f4024,f4025,f4064,f4068,f4053,f4054,f4055,f4056,f4057,f4058,f4059,f4060,f4079,f4080,f4081,f4082,f4083,f4111,f4112,f4113,f4114,f4115,f4125,f4126,f4127,f4129,f4130,f4138,f4139,f4140,f4141,f4224,f4225,f4230,f4231,f4348,f4347,f4342,f4425,f4433,f4434,f4435,f4437,f4438,f4452,f4453,f4454,f4456,f4457,f4473,f4474,f4475,f4477,f4478,f4500,f4501,f4502,f4504,f4505,f4519,f4520,f4521,f4523,f4524,f4535,f4536,f4537,f4538,f4549,f4550,f4551,f4553,f4554,f4562,f4563,f4564,f4565,f4648,f4649,f4654,f4655,f4771,f4766,f4845,f4854,f4855,f4856,f4858,f4859,f4876,f4877,f4878,f4880,f4881,f4900,f4901,f4902,f4904,f4905,f4918])).
fof(f4918,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),X0)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X0),s(fun(fun(sK10,bool),bool),sK12))))) )),
  inference(resolution,[],[f394,f261])).
fof(f261,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X4)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X4),s(fun(fun(X0,bool),bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(X0,bool),bool),X2))))))) )),
  inference(cnf_transformation,[],[f164])).
fof(f164,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(X0,bool),bool),X2)))))) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X4),s(fun(fun(X0,bool),bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X4)))))) & ((p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),sK14(X2,X1,X0)),s(fun(fun(X0,bool),bool),X2)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),sK14(X2,X1,X0)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(X0,bool),bool),X2))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK14])],[f163])).
fof(f163,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(X0,bool),bool),X2)))))) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X4),s(fun(fun(X0,bool),bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X4)))))) & (? [X3] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),X2)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(X0,bool),bool),X2))))))))),
  inference(rectify,[],[f162])).
fof(f162,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(X0,bool),bool),X2)))))) | ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3)))))) & (? [X3] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),X2)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(X0,bool),bool),X2))))))))),
  inference(nnf_transformation,[],[f107])).
fof(f107,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(X0,bool),bool),X2)))))) <=> ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(X0,bool),bool),X2)))))) <=> ! [X3] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),X2)))) => p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(X0,bool),bool),X2)))))) <=> ! [X3] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),X2)))) => p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X6] : ! [X12,X19] : (p(s(bool,'const.bool.IN_2'(s(X6,X12),s(fun(X6,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(X6,bool),bool),X19)))))) <=> ! [X15] : (p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),X15),s(fun(fun(X6,bool),bool),X19)))) => p(s(bool,'const.bool.IN_2'(s(X6,X12),s(fun(X6,bool),X15))))))),
  file('thm.util_prob.BIGINTER_SUBSET','thm.pred_set.IN_BIGINTER')).
fof(f4905,plain,(
  ( ! [X4,X3] : (sP1(X4,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X3))),s(fun(fun(sK10,bool),bool),X3))) | ~p(s(bool,X4)) | s(bool,f(s(fun(sK10,bool),sK17(X3,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X3,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f2835,f299])).
fof(f299,plain,(
  ( ! [X2,X0,X3,X1] : (s(fun(X0,X1),X2) = s(fun(X0,X1),X3) | s(X1,f(s(X0,sK17(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK17(X3,X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f179])).
fof(f179,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK17(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK17(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK17])],[f118])).
fof(f118,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f85])).
fof(f85,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f84])).
fof(f84,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.util_prob.BIGINTER_SUBSET',aHL_EXT)).
fof(f4904,plain,(
  ( ! [X2,X1] : (sP1(X2,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X1))),s(fun(fun(sK10,bool),bool),X1))) | ~p(s(bool,X2)) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X1,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X1,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f2835,f299])).
fof(f4902,plain,(
  ( ! [X4,X3] : (sP1(X4,'const.bool.IN_2'(s(fun(sK10,bool),X3),s(fun(fun(sK10,bool),bool),sK12))) | ~p(s(bool,X4)) | s(bool,f(s(sK10,sK17(X3,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10)))) != s(bool,g(s(sK10,sK17(X3,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10))))) )),
  inference(superposition,[],[f2835,f299])).
fof(f4901,plain,(
  ( ! [X2,X1] : (sP1(X2,'const.bool.IN_2'(s(fun(sK10,bool),X1),s(fun(fun(sK10,bool),bool),sK12))) | ~p(s(bool,X2)) | s(bool,f(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X1,bool,sK10)))) != s(bool,g(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X1,bool,sK10))))) )),
  inference(superposition,[],[f2835,f299])).
fof(f4900,plain,(
  ( ! [X0] : (sP1(X0,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK10,bool),bool),sK12))) | ~p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))),s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))))) )),
  inference(superposition,[],[f2835,f255])).
fof(f4881,plain,(
  ( ! [X4,X3] : (~sP1(X4,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X3))),s(fun(fun(sK10,bool),bool),X3))) | p(s(bool,X4)) | s(bool,f(s(fun(sK10,bool),sK17(X3,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X3,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f2834,f299])).
fof(f4880,plain,(
  ( ! [X2,X1] : (~sP1(X2,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X1))),s(fun(fun(sK10,bool),bool),X1))) | p(s(bool,X2)) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X1,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X1,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f2834,f299])).
fof(f4878,plain,(
  ( ! [X4,X3] : (~sP1(X4,'const.bool.IN_2'(s(fun(sK10,bool),X3),s(fun(fun(sK10,bool),bool),sK12))) | p(s(bool,X4)) | s(bool,f(s(sK10,sK17(X3,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10)))) != s(bool,g(s(sK10,sK17(X3,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10))))) )),
  inference(superposition,[],[f2834,f299])).
fof(f4877,plain,(
  ( ! [X2,X1] : (~sP1(X2,'const.bool.IN_2'(s(fun(sK10,bool),X1),s(fun(fun(sK10,bool),bool),sK12))) | p(s(bool,X2)) | s(bool,f(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X1,bool,sK10)))) != s(bool,g(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X1,bool,sK10))))) )),
  inference(superposition,[],[f2834,f299])).
fof(f4876,plain,(
  ( ! [X0] : (~sP1(X0,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK10,bool),bool),sK12))) | p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))),s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))))) )),
  inference(superposition,[],[f2834,f255])).
fof(f4859,plain,(
  ( ! [X4,X3] : (~sP1('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X3))),s(fun(fun(sK10,bool),bool),X3)),X4) | p(s(bool,X4)) | s(bool,f(s(fun(sK10,bool),sK17(X3,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X3,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f2833,f299])).
fof(f4858,plain,(
  ( ! [X2,X1] : (~sP1('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X1))),s(fun(fun(sK10,bool),bool),X1)),X2) | p(s(bool,X2)) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X1,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X1,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f2833,f299])).
fof(f4856,plain,(
  ( ! [X4,X3] : (~sP1('const.bool.IN_2'(s(fun(sK10,bool),X3),s(fun(fun(sK10,bool),bool),sK12)),X4) | p(s(bool,X4)) | s(bool,f(s(sK10,sK17(X3,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10)))) != s(bool,g(s(sK10,sK17(X3,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10))))) )),
  inference(superposition,[],[f2833,f299])).
fof(f4855,plain,(
  ( ! [X2,X1] : (~sP1('const.bool.IN_2'(s(fun(sK10,bool),X1),s(fun(fun(sK10,bool),bool),sK12)),X2) | p(s(bool,X2)) | s(bool,f(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X1,bool,sK10)))) != s(bool,g(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X1,bool,sK10))))) )),
  inference(superposition,[],[f2833,f299])).
fof(f4854,plain,(
  ( ! [X0] : (~sP1('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK10,bool),bool),sK12)),X0) | p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))),s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))))) )),
  inference(superposition,[],[f2833,f255])).
fof(f4845,plain,(
  ( ! [X292] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),sK14(X292,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),bool))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(bool,bool),bool),X292))))))) )),
  inference(forward_demodulation,[],[f4767,f2822])).
fof(f4767,plain,(
  ( ! [X292] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(bool,bool),bool),X292)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12))),s(fun(bool,bool),sK14(X292,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),bool)))))) )),
  inference(superposition,[],[f263,f2822])).
fof(f263,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(X0,bool),bool),X2)))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),sK14(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f164])).
fof(f4766,plain,(
  ( ! [X291] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(bool,bool),bool),X291)))))) | p(s(bool,'const.bool.IN_2'(s(fun(bool,bool),sK14(X291,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),bool)),s(fun(fun(bool,bool),bool),X291))))) )),
  inference(superposition,[],[f262,f2822])).
fof(f262,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(X0,bool),bool),X2)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),sK14(X2,X1,X0)),s(fun(fun(X0,bool),bool),X2))))) )),
  inference(cnf_transformation,[],[f164])).
fof(f4771,plain,(
  ( ! [X282] : (s(bool,X282) != s(bool,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12))) | s(bool,'F_0') = s(bool,X282)) )),
  inference(trivial_inequality_removal,[],[f4761])).
fof(f4761,plain,(
  ( ! [X282] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,X282) != s(bool,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12))) | s(bool,'F_0') = s(bool,X282)) )),
  inference(superposition,[],[f2040,f2822])).
fof(f4655,plain,(
  ( ! [X107,X106] : (s(bool,'F_0') != s(bool,X106) | p(s(bool,X107)) | ~sP5(X107,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X106)) )),
  inference(superposition,[],[f328,f2822])).
fof(f328,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | s(bool,X0) != s(bool,X1) | ~sP5(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f193])).
fof(f193,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP5(X2,X1,X0)) & (~sP5(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f132])).
fof(f132,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP5(X2,X1,X0))),
  inference(definition_folding,[],[f88,f131,f130,f129,f128])).
fof(f128,plain,(
  ! [X0,X1,X2] : (sP2(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f129,plain,(
  ! [X1,X0,X2] : (sP3(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f130,plain,(
  ! [X2,X0,X1] : (sP4(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f131,plain,(
  ! [X2,X1,X0] : (sP5(X2,X1,X0) <=> (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f88,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f87])).
fof(f87,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X20,X21,X22] : ((p(s(bool,X22)) <=> s(bool,X20) = s(bool,X21)) <=> ((p(s(bool,X22)) | p(s(bool,X21)) | p(s(bool,X20))) & (p(s(bool,X22)) | ~p(s(bool,X20)) | ~p(s(bool,X21))) & (p(s(bool,X21)) | ~p(s(bool,X20)) | ~p(s(bool,X22))) & (p(s(bool,X20)) | ~p(s(bool,X21)) | ~p(s(bool,X22)))))),
  file('thm.util_prob.BIGINTER_SUBSET','thm.sat.sat_2|49|_prove_ltb')).
fof(f4654,plain,(
  ( ! [X105,X104] : (s(bool,'F_0') != s(bool,X104) | p(s(bool,X105)) | ~sP5(X105,X104,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)))) )),
  inference(superposition,[],[f328,f2822])).
fof(f4649,plain,(
  ( ! [X94,X95] : (s(bool,'F_0') != s(bool,X94) | sP5(X95,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X94) | ~p(s(bool,X95))) )),
  inference(superposition,[],[f326,f2822])).
fof(f326,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f193])).
fof(f4648,plain,(
  ( ! [X92,X93] : (s(bool,'F_0') != s(bool,X92) | sP5(X93,X92,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12))) | ~p(s(bool,X93))) )),
  inference(superposition,[],[f326,f2822])).
fof(f4565,plain,(
  ( ! [X10,X11] : (s(bool,'F_0') = s(bool,X10) | ~p(s(bool,X11)) | ~sP5(X11,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X10)) )),
  inference(superposition,[],[f2822,f327])).
fof(f327,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X2)) | ~sP5(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f193])).
fof(f4564,plain,(
  ( ! [X8,X9] : (s(bool,'F_0') = s(bool,X8) | ~p(s(bool,X9)) | ~sP5(X9,X8,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)))) )),
  inference(superposition,[],[f2822,f327])).
fof(f4563,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') = s(bool,X6) | sP5(X7,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X6) | p(s(bool,X7))) )),
  inference(superposition,[],[f2822,f325])).
fof(f325,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f193])).
fof(f4562,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') = s(bool,X4) | sP5(X5,X4,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12))) | p(s(bool,X5))) )),
  inference(superposition,[],[f2822,f325])).
fof(f4554,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X1))),s(fun(fun(sK10,bool),bool),X1))) | s(bool,f(s(fun(sK10,bool),sK17(X1,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X1,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f2822,f299])).
fof(f4553,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X0))),s(fun(fun(sK10,bool),bool),X0))) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X0,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X0,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f2822,f299])).
fof(f4551,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X1),s(fun(fun(sK10,bool),bool),sK12))) | s(bool,f(s(sK10,sK17(X1,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10)))) != s(bool,g(s(sK10,sK17(X1,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10))))) )),
  inference(superposition,[],[f2822,f299])).
fof(f4550,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X0),s(fun(fun(sK10,bool),bool),sK12))) | s(bool,f(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X0,bool,sK10)))) != s(bool,g(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X0,bool,sK10))))) )),
  inference(superposition,[],[f2822,f299])).
fof(f4549,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK10,bool),bool),sK12))) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))),s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))))),
  inference(superposition,[],[f2822,f255])).
fof(f4538,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') != s(bool,X10) | ~p(s(bool,X11)) | ~sP5(X11,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X10)) )),
  inference(superposition,[],[f2820,f327])).
fof(f4537,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') != s(bool,X8) | ~p(s(bool,X9)) | ~sP5(X9,X8,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)))) )),
  inference(superposition,[],[f2820,f327])).
fof(f4536,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X6) | sP5(X7,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X6) | p(s(bool,X7))) )),
  inference(superposition,[],[f2820,f325])).
fof(f4535,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | sP5(X5,X4,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12))) | p(s(bool,X5))) )),
  inference(superposition,[],[f2820,f325])).
fof(f4524,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X1))),s(fun(fun(sK10,bool),bool),X1))) | s(bool,f(s(fun(sK10,bool),sK17(X1,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X1,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f2820,f299])).
fof(f4523,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X0))),s(fun(fun(sK10,bool),bool),X0))) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X0,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X0,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f2820,f299])).
fof(f4521,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X1),s(fun(fun(sK10,bool),bool),sK12))) | s(bool,f(s(sK10,sK17(X1,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10)))) != s(bool,g(s(sK10,sK17(X1,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10))))) )),
  inference(superposition,[],[f2820,f299])).
fof(f4520,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X0),s(fun(fun(sK10,bool),bool),sK12))) | s(bool,f(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X0,bool,sK10)))) != s(bool,g(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X0,bool,sK10))))) )),
  inference(superposition,[],[f2820,f299])).
fof(f4519,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK10,bool),bool),sK12))) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))),s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))))),
  inference(superposition,[],[f2820,f255])).
fof(f4505,plain,(
  ( ! [X4,X3] : (sP1(X4,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),X3))),s(fun(fun(sK10,bool),bool),X3))) | p(s(bool,X4)) | s(bool,f(s(fun(sK10,bool),sK17(X3,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X3,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f980,f299])).
fof(f4504,plain,(
  ( ! [X2,X1] : (sP1(X2,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),X1))),s(fun(fun(sK10,bool),bool),X1))) | p(s(bool,X2)) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X1,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X1,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f980,f299])).
fof(f4502,plain,(
  ( ! [X4,X3] : (sP1(X4,'const.bool.IN_2'(s(fun(sK10,bool),X3),s(fun(fun(sK10,bool),bool),sK12))) | p(s(bool,X4)) | s(bool,f(s(sK10,sK17(X3,'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10)))) != s(bool,g(s(sK10,sK17(X3,'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10))))) )),
  inference(superposition,[],[f980,f299])).
fof(f4501,plain,(
  ( ! [X2,X1] : (sP1(X2,'const.bool.IN_2'(s(fun(sK10,bool),X1),s(fun(fun(sK10,bool),bool),sK12))) | p(s(bool,X2)) | s(bool,f(s(sK10,sK17('const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12)),X1,bool,sK10)))) != s(bool,g(s(sK10,sK17('const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12)),X1,bool,sK10))))) )),
  inference(superposition,[],[f980,f299])).
fof(f4500,plain,(
  ( ! [X0] : (sP1(X0,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK10,bool),bool),sK12))) | p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))))),s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))))))) )),
  inference(superposition,[],[f980,f255])).
fof(f4478,plain,(
  ( ! [X4,X3] : (sP1('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),X3))),s(fun(fun(sK10,bool),bool),X3)),X4) | p(s(bool,X4)) | s(bool,f(s(fun(sK10,bool),sK17(X3,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X3,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f979,f299])).
fof(f4477,plain,(
  ( ! [X2,X1] : (sP1('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),X1))),s(fun(fun(sK10,bool),bool),X1)),X2) | p(s(bool,X2)) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X1,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X1,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f979,f299])).
fof(f4475,plain,(
  ( ! [X4,X3] : (sP1('const.bool.IN_2'(s(fun(sK10,bool),X3),s(fun(fun(sK10,bool),bool),sK12)),X4) | p(s(bool,X4)) | s(bool,f(s(sK10,sK17(X3,'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10)))) != s(bool,g(s(sK10,sK17(X3,'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10))))) )),
  inference(superposition,[],[f979,f299])).
fof(f4474,plain,(
  ( ! [X2,X1] : (sP1('const.bool.IN_2'(s(fun(sK10,bool),X1),s(fun(fun(sK10,bool),bool),sK12)),X2) | p(s(bool,X2)) | s(bool,f(s(sK10,sK17('const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12)),X1,bool,sK10)))) != s(bool,g(s(sK10,sK17('const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12)),X1,bool,sK10))))) )),
  inference(superposition,[],[f979,f299])).
fof(f4473,plain,(
  ( ! [X0] : (sP1('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK10,bool),bool),sK12)),X0) | p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))))),s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))))))) )),
  inference(superposition,[],[f979,f255])).
fof(f4457,plain,(
  ( ! [X4,X3] : (~sP1(X4,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),X3))),s(fun(fun(sK10,bool),bool),X3))) | ~p(s(bool,X4)) | s(bool,f(s(fun(sK10,bool),sK17(X3,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X3,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f978,f299])).
fof(f4456,plain,(
  ( ! [X2,X1] : (~sP1(X2,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),X1))),s(fun(fun(sK10,bool),bool),X1))) | ~p(s(bool,X2)) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X1,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X1,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f978,f299])).
fof(f4454,plain,(
  ( ! [X4,X3] : (~sP1(X4,'const.bool.IN_2'(s(fun(sK10,bool),X3),s(fun(fun(sK10,bool),bool),sK12))) | ~p(s(bool,X4)) | s(bool,f(s(sK10,sK17(X3,'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10)))) != s(bool,g(s(sK10,sK17(X3,'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10))))) )),
  inference(superposition,[],[f978,f299])).
fof(f4453,plain,(
  ( ! [X2,X1] : (~sP1(X2,'const.bool.IN_2'(s(fun(sK10,bool),X1),s(fun(fun(sK10,bool),bool),sK12))) | ~p(s(bool,X2)) | s(bool,f(s(sK10,sK17('const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12)),X1,bool,sK10)))) != s(bool,g(s(sK10,sK17('const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12)),X1,bool,sK10))))) )),
  inference(superposition,[],[f978,f299])).
fof(f4452,plain,(
  ( ! [X0] : (~sP1(X0,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK10,bool),bool),sK12))) | ~p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))))),s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))))))) )),
  inference(superposition,[],[f978,f255])).
fof(f4438,plain,(
  ( ! [X4,X3] : (~sP1('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),X3))),s(fun(fun(sK10,bool),bool),X3)),X4) | ~p(s(bool,X4)) | s(bool,f(s(fun(sK10,bool),sK17(X3,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X3,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f977,f299])).
fof(f4437,plain,(
  ( ! [X2,X1] : (~sP1('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),X1))),s(fun(fun(sK10,bool),bool),X1)),X2) | ~p(s(bool,X2)) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X1,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X1,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f977,f299])).
fof(f4435,plain,(
  ( ! [X4,X3] : (~sP1('const.bool.IN_2'(s(fun(sK10,bool),X3),s(fun(fun(sK10,bool),bool),sK12)),X4) | ~p(s(bool,X4)) | s(bool,f(s(sK10,sK17(X3,'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10)))) != s(bool,g(s(sK10,sK17(X3,'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10))))) )),
  inference(superposition,[],[f977,f299])).
fof(f4434,plain,(
  ( ! [X2,X1] : (~sP1('const.bool.IN_2'(s(fun(sK10,bool),X1),s(fun(fun(sK10,bool),bool),sK12)),X2) | ~p(s(bool,X2)) | s(bool,f(s(sK10,sK17('const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12)),X1,bool,sK10)))) != s(bool,g(s(sK10,sK17('const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12)),X1,bool,sK10))))) )),
  inference(superposition,[],[f977,f299])).
fof(f4433,plain,(
  ( ! [X0] : (~sP1('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK10,bool),bool),sK12)),X0) | ~p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))))),s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))))))) )),
  inference(superposition,[],[f977,f255])).
fof(f4425,plain,(
  ( ! [X292] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),sK14(X292,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),bool))))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(bool,bool),bool),X292))))))) )),
  inference(forward_demodulation,[],[f4343,f965])).
fof(f4343,plain,(
  ( ! [X292] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(bool,bool),bool),X292)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12))),s(fun(bool,bool),sK14(X292,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),bool)))))) )),
  inference(superposition,[],[f263,f965])).
fof(f4342,plain,(
  ( ! [X291] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(bool,bool),bool),X291)))))) | p(s(bool,'const.bool.IN_2'(s(fun(bool,bool),sK14(X291,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),bool)),s(fun(fun(bool,bool),bool),X291))))) )),
  inference(superposition,[],[f262,f965])).
fof(f4347,plain,(
  ( ! [X278] : (s(bool,X278) != s(bool,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12))) | s(bool,'F_0') != s(bool,X278)) )),
  inference(trivial_inequality_removal,[],[f4334])).
fof(f4334,plain,(
  ( ! [X278] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X278) != s(bool,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12))) | s(bool,'F_0') != s(bool,X278)) )),
  inference(superposition,[],[f1994,f965])).
fof(f4348,plain,(
  ( ! [X259] : (s(bool,X259) != s(bool,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12))) | s(bool,'T_0') = s(bool,X259)) )),
  inference(trivial_inequality_removal,[],[f4321])).
fof(f4321,plain,(
  ( ! [X259] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X259) != s(bool,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12))) | s(bool,'T_0') = s(bool,X259)) )),
  inference(superposition,[],[f1249,f965])).
fof(f4231,plain,(
  ( ! [X107,X106] : (s(bool,'T_0') != s(bool,X106) | p(s(bool,X107)) | ~sP5(X107,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X106)) )),
  inference(superposition,[],[f328,f965])).
fof(f4230,plain,(
  ( ! [X105,X104] : (s(bool,'T_0') != s(bool,X104) | p(s(bool,X105)) | ~sP5(X105,X104,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)))) )),
  inference(superposition,[],[f328,f965])).
fof(f4225,plain,(
  ( ! [X94,X95] : (s(bool,'T_0') != s(bool,X94) | sP5(X95,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X94) | ~p(s(bool,X95))) )),
  inference(superposition,[],[f326,f965])).
fof(f4224,plain,(
  ( ! [X92,X93] : (s(bool,'T_0') != s(bool,X92) | sP5(X93,X92,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12))) | ~p(s(bool,X93))) )),
  inference(superposition,[],[f326,f965])).
fof(f4141,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') = s(bool,X10) | ~p(s(bool,X11)) | ~sP5(X11,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X10)) )),
  inference(superposition,[],[f965,f327])).
fof(f4140,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') = s(bool,X8) | ~p(s(bool,X9)) | ~sP5(X9,X8,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)))) )),
  inference(superposition,[],[f965,f327])).
fof(f4139,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') = s(bool,X6) | sP5(X7,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X6) | p(s(bool,X7))) )),
  inference(superposition,[],[f965,f325])).
fof(f4138,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') = s(bool,X4) | sP5(X5,X4,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12))) | p(s(bool,X5))) )),
  inference(superposition,[],[f965,f325])).
fof(f4130,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),X1))),s(fun(fun(sK10,bool),bool),X1))) | s(bool,f(s(fun(sK10,bool),sK17(X1,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X1,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f965,f299])).
fof(f4129,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),X0))),s(fun(fun(sK10,bool),bool),X0))) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X0,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X0,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f965,f299])).
fof(f4127,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X1),s(fun(fun(sK10,bool),bool),sK12))) | s(bool,f(s(sK10,sK17(X1,'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10)))) != s(bool,g(s(sK10,sK17(X1,'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10))))) )),
  inference(superposition,[],[f965,f299])).
fof(f4126,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X0),s(fun(fun(sK10,bool),bool),sK12))) | s(bool,f(s(sK10,sK17('const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12)),X0,bool,sK10)))) != s(bool,g(s(sK10,sK17('const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12)),X0,bool,sK10))))) )),
  inference(superposition,[],[f965,f299])).
fof(f4125,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK10,bool),bool),sK12))) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))))),s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))))))),
  inference(superposition,[],[f965,f255])).
fof(f4115,plain,(
  ( ! [X6,X7,X5] : (sP9(X6,X7,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),X5))) | ~p(s(bool,X6)) | s(bool,f(s(sK10,sK17(X5,sK11,bool,sK10)))) != s(bool,g(s(sK10,sK17(X5,sK11,bool,sK10))))) )),
  inference(superposition,[],[f705,f299])).
fof(f4114,plain,(
  ( ! [X4,X2,X3] : (sP9(X3,X4,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),X2))) | ~p(s(bool,X3)) | s(bool,f(s(sK10,sK17(sK11,X2,bool,sK10)))) != s(bool,g(s(sK10,sK17(sK11,X2,bool,sK10))))) )),
  inference(superposition,[],[f705,f299])).
fof(f4113,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),sK11))),s(fun(sK10,bool),sK11))))) )),
  inference(superposition,[],[f705,f255])).
fof(f4112,plain,(
  ( ! [X6,X7,X5] : (sP9(X6,X7,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X5)),sK10)),s(fun(sK10,bool),sK11))) | ~p(s(bool,X6)) | s(bool,f(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f705,f299])).
fof(f4111,plain,(
  ( ! [X4,X2,X3] : (sP9(X3,X4,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X2)),sK10)),s(fun(sK10,bool),sK11))) | ~p(s(bool,X3)) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f705,f299])).
fof(f4083,plain,(
  ( ! [X6,X7,X5] : (~sP9('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),X5)),X6,X7) | ~p(s(bool,X6)) | s(bool,f(s(sK10,sK17(X5,sK11,bool,sK10)))) != s(bool,g(s(sK10,sK17(X5,sK11,bool,sK10))))) )),
  inference(superposition,[],[f697,f299])).
fof(f4082,plain,(
  ( ! [X4,X2,X3] : (~sP9('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),X2)),X3,X4) | ~p(s(bool,X3)) | s(bool,f(s(sK10,sK17(sK11,X2,bool,sK10)))) != s(bool,g(s(sK10,sK17(sK11,X2,bool,sK10))))) )),
  inference(superposition,[],[f697,f299])).
fof(f4081,plain,(
  ( ! [X0,X1] : (~sP9('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),'const.pred_set.EMPTY_0')),X0,X1) | ~p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),sK11))),s(fun(sK10,bool),sK11))))) )),
  inference(superposition,[],[f697,f255])).
fof(f4080,plain,(
  ( ! [X6,X7,X5] : (~sP9('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X5)),sK10)),s(fun(sK10,bool),sK11)),X6,X7) | ~p(s(bool,X6)) | s(bool,f(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f697,f299])).
fof(f4079,plain,(
  ( ! [X4,X2,X3] : (~sP9('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X2)),sK10)),s(fun(sK10,bool),sK11)),X3,X4) | ~p(s(bool,X3)) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f697,f299])).
fof(f4060,plain,(
  ( ! [X6,X7,X5] : (~sP5('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),X5)),X6,X7) | s(bool,X6) != s(bool,X7) | s(bool,f(s(sK10,sK17(X5,sK11,bool,sK10)))) != s(bool,g(s(sK10,sK17(X5,sK11,bool,sK10))))) )),
  inference(superposition,[],[f430,f299])).
fof(f4059,plain,(
  ( ! [X4,X2,X3] : (~sP5('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),X2)),X3,X4) | s(bool,X3) != s(bool,X4) | s(bool,f(s(sK10,sK17(sK11,X2,bool,sK10)))) != s(bool,g(s(sK10,sK17(sK11,X2,bool,sK10))))) )),
  inference(superposition,[],[f430,f299])).
fof(f4058,plain,(
  ( ! [X0,X1] : (~sP5('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),'const.pred_set.EMPTY_0')),X0,X1) | s(bool,X0) != s(bool,X1) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),sK11))),s(fun(sK10,bool),sK11))))) )),
  inference(superposition,[],[f430,f255])).
fof(f4057,plain,(
  ( ! [X6,X7,X5] : (~sP5('const.pred_set.SUBSET_2'(s(fun(sK10,bool),X5),s(fun(sK10,bool),sK11)),X6,X7) | s(bool,X6) != s(bool,X7) | s(bool,f(s(sK10,sK17(X5,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10)))) != s(bool,g(s(sK10,sK17(X5,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10))))) )),
  inference(superposition,[],[f430,f299])).
fof(f4056,plain,(
  ( ! [X4,X2,X3] : (~sP5('const.pred_set.SUBSET_2'(s(fun(sK10,bool),X2),s(fun(sK10,bool),sK11)),X3,X4) | s(bool,X3) != s(bool,X4) | s(bool,f(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X2,bool,sK10)))) != s(bool,g(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X2,bool,sK10))))) )),
  inference(superposition,[],[f430,f299])).
fof(f4055,plain,(
  ( ! [X0,X1] : (~sP5('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.EMPTY_0'),s(fun(sK10,bool),sK11)),X0,X1) | s(bool,X0) != s(bool,X1) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))),s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))))) )),
  inference(superposition,[],[f430,f255])).
fof(f4054,plain,(
  ( ! [X6,X7,X5] : (~sP5('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X5))),s(fun(sK10,bool),sK11)),X6,X7) | s(bool,X6) != s(bool,X7) | s(bool,f(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f430,f299])).
fof(f4053,plain,(
  ( ! [X4,X2,X3] : (~sP5('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X2))),s(fun(sK10,bool),sK11)),X3,X4) | s(bool,X3) != s(bool,X4) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f430,f299])).
fof(f4068,plain,(
  ( ! [X31] : (s(bool,'F_0') != s(bool,X31) | ~sP2(X31,'F_0','const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)))) )),
  inference(subsumption_resolution,[],[f4049,f422])).
fof(f4049,plain,(
  ( ! [X31] : (s(bool,'F_0') != s(bool,X31) | ~sP4('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)),X31,'F_0') | ~sP2(X31,'F_0','const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)))) )),
  inference(resolution,[],[f430,f1191])).
fof(f4064,plain,(
  ( ! [X2,X3] : (s(bool,X2) != s(bool,X3) | ~sP3(X3,X2,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11))) | ~sP2(X2,X3,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)))) )),
  inference(subsumption_resolution,[],[f4063,f221])).
fof(f221,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f43])).
fof(f43,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.util_prob.BIGINTER_SUBSET',aHL_FALSITY)).
fof(f4063,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | s(bool,X2) != s(bool,X3) | ~sP3(X3,X2,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11))) | ~sP2(X2,X3,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)))) )),
  inference(forward_demodulation,[],[f4062,f398])).
fof(f4062,plain,(
  ( ! [X2,X3] : (s(bool,X2) != s(bool,X3) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)))) | ~sP3(X3,X2,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11))) | ~sP2(X2,X3,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)))) )),
  inference(subsumption_resolution,[],[f4034,f316])).
fof(f316,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f187])).
fof(f187,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f186])).
fof(f186,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(flattening,[],[f185])).
fof(f185,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(nnf_transformation,[],[f130])).
fof(f4034,plain,(
  ( ! [X2,X3] : (s(bool,X2) != s(bool,X3) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)))) | ~sP4('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)),X2,X3) | ~sP3(X3,X2,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11))) | ~sP2(X2,X3,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)))) )),
  inference(resolution,[],[f430,f312])).
fof(f312,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | p(s(bool,X0)) | ~sP4(X0,X2,X1) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f184])).
fof(f184,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | (sP2(X2,X1,X0) & sP3(X1,X2,X0) & sP4(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP2(X2,X1,X0) | ~sP3(X1,X2,X0) | ~sP4(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f183])).
fof(f183,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP5(X2,X1,X0)))),
  inference(flattening,[],[f182])).
fof(f182,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP5(X2,X1,X0)))),
  inference(nnf_transformation,[],[f131])).
fof(f4025,plain,(
  ( ! [X6,X7,X5] : (~sP9(X6,X7,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),X5))) | p(s(bool,X6)) | s(bool,f(s(sK10,sK17(X5,sK11,bool,sK10)))) != s(bool,g(s(sK10,sK17(X5,sK11,bool,sK10))))) )),
  inference(superposition,[],[f696,f299])).
fof(f4024,plain,(
  ( ! [X4,X2,X3] : (~sP9(X3,X4,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),X2))) | p(s(bool,X3)) | s(bool,f(s(sK10,sK17(sK11,X2,bool,sK10)))) != s(bool,g(s(sK10,sK17(sK11,X2,bool,sK10))))) )),
  inference(superposition,[],[f696,f299])).
fof(f4023,plain,(
  ( ! [X0,X1] : (~sP9(X0,X1,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),sK11))),s(fun(sK10,bool),sK11))))) )),
  inference(superposition,[],[f696,f255])).
fof(f4022,plain,(
  ( ! [X6,X7,X5] : (~sP9(X6,X7,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X5)),sK10)),s(fun(sK10,bool),sK11))) | p(s(bool,X6)) | s(bool,f(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f696,f299])).
fof(f4021,plain,(
  ( ! [X4,X2,X3] : (~sP9(X3,X4,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X2)),sK10)),s(fun(sK10,bool),sK11))) | p(s(bool,X3)) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f696,f299])).
fof(f4002,plain,(
  ( ! [X6,X7,X5] : (~sP9('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),X5)),X6,X7) | p(s(bool,X7)) | s(bool,f(s(sK10,sK17(X5,sK11,bool,sK10)))) != s(bool,g(s(sK10,sK17(X5,sK11,bool,sK10))))) )),
  inference(superposition,[],[f695,f299])).
fof(f4001,plain,(
  ( ! [X4,X2,X3] : (~sP9('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),X2)),X3,X4) | p(s(bool,X4)) | s(bool,f(s(sK10,sK17(sK11,X2,bool,sK10)))) != s(bool,g(s(sK10,sK17(sK11,X2,bool,sK10))))) )),
  inference(superposition,[],[f695,f299])).
fof(f4000,plain,(
  ( ! [X0,X1] : (~sP9('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),'const.pred_set.EMPTY_0')),X0,X1) | p(s(bool,X1)) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),sK11))),s(fun(sK10,bool),sK11))))) )),
  inference(superposition,[],[f695,f255])).
fof(f3999,plain,(
  ( ! [X6,X7,X5] : (~sP9('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X5)),sK10)),s(fun(sK10,bool),sK11)),X6,X7) | p(s(bool,X7)) | s(bool,f(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f695,f299])).
fof(f3998,plain,(
  ( ! [X4,X2,X3] : (~sP9('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X2)),sK10)),s(fun(sK10,bool),sK11)),X3,X4) | p(s(bool,X4)) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f695,f299])).
fof(f3975,plain,(
  ( ! [X6,X7,X5] : (~sP7('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),X5)),X6,X7) | ~p(s(bool,X6)) | s(bool,f(s(sK10,sK17(X5,sK11,bool,sK10)))) != s(bool,g(s(sK10,sK17(X5,sK11,bool,sK10))))) )),
  inference(superposition,[],[f680,f299])).
fof(f3974,plain,(
  ( ! [X4,X2,X3] : (~sP7('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),X2)),X3,X4) | ~p(s(bool,X3)) | s(bool,f(s(sK10,sK17(sK11,X2,bool,sK10)))) != s(bool,g(s(sK10,sK17(sK11,X2,bool,sK10))))) )),
  inference(superposition,[],[f680,f299])).
fof(f3973,plain,(
  ( ! [X0,X1] : (~sP7('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),'const.pred_set.EMPTY_0')),X0,X1) | ~p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),sK11))),s(fun(sK10,bool),sK11))))) )),
  inference(superposition,[],[f680,f255])).
fof(f3972,plain,(
  ( ! [X6,X7,X5] : (~sP7('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X5)),sK10)),s(fun(sK10,bool),sK11)),X6,X7) | ~p(s(bool,X6)) | s(bool,f(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f680,f299])).
fof(f3971,plain,(
  ( ! [X4,X2,X3] : (~sP7('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X2)),sK10)),s(fun(sK10,bool),sK11)),X3,X4) | ~p(s(bool,X3)) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f680,f299])).
fof(f3949,plain,(
  ( ! [X6,X7,X5] : (~sP7('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),X5)),X6,X7) | ~p(s(bool,X7)) | s(bool,f(s(sK10,sK17(X5,sK11,bool,sK10)))) != s(bool,g(s(sK10,sK17(X5,sK11,bool,sK10))))) )),
  inference(superposition,[],[f679,f299])).
fof(f3948,plain,(
  ( ! [X4,X2,X3] : (~sP7('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),X2)),X3,X4) | ~p(s(bool,X4)) | s(bool,f(s(sK10,sK17(sK11,X2,bool,sK10)))) != s(bool,g(s(sK10,sK17(sK11,X2,bool,sK10))))) )),
  inference(superposition,[],[f679,f299])).
fof(f3947,plain,(
  ( ! [X0,X1] : (~sP7('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),'const.pred_set.EMPTY_0')),X0,X1) | ~p(s(bool,X1)) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),sK11))),s(fun(sK10,bool),sK11))))) )),
  inference(superposition,[],[f679,f255])).
fof(f3946,plain,(
  ( ! [X6,X7,X5] : (~sP7('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X5)),sK10)),s(fun(sK10,bool),sK11)),X6,X7) | ~p(s(bool,X7)) | s(bool,f(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f679,f299])).
fof(f3945,plain,(
  ( ! [X4,X2,X3] : (~sP7('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X2)),sK10)),s(fun(sK10,bool),sK11)),X3,X4) | ~p(s(bool,X4)) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f679,f299])).
fof(f3908,plain,(
  ( ! [X52,X50,X53,X51,X49] : (s(bool,'F_0') != s(bool,X50) | s(bool,X49) != s(bool,X51) | s(bool,X51) != s(bool,X52) | ~p(s(bool,X52)) | ~p(s(bool,X53)) | ~sP5(X53,X49,X50)) )),
  inference(superposition,[],[f1996,f327])).
fof(f3907,plain,(
  ( ! [X47,X45,X48,X46,X44] : (s(bool,'F_0') != s(bool,X45) | s(bool,X44) != s(bool,X46) | s(bool,X46) != s(bool,X47) | ~p(s(bool,X47)) | ~p(s(bool,X48)) | ~sP5(X48,X45,X44)) )),
  inference(superposition,[],[f1996,f327])).
fof(f3906,plain,(
  ( ! [X39,X43,X41,X42,X40] : (s(bool,'F_0') != s(bool,X40) | s(bool,X39) != s(bool,X41) | s(bool,X41) != s(bool,X42) | ~p(s(bool,X42)) | sP5(X43,X39,X40) | p(s(bool,X43))) )),
  inference(superposition,[],[f1996,f325])).
fof(f3905,plain,(
  ( ! [X37,X35,X38,X36,X34] : (s(bool,'F_0') != s(bool,X35) | s(bool,X34) != s(bool,X36) | s(bool,X36) != s(bool,X37) | ~p(s(bool,X37)) | sP5(X38,X35,X34) | p(s(bool,X38))) )),
  inference(superposition,[],[f1996,f325])).
fof(f3894,plain,(
  ( ! [X61,X64,X62,X65,X63] : (s(bool,X61) != s(bool,X62) | s(bool,X62) != s(bool,X63) | s(bool,X63) != s(bool,X64) | ~p(s(bool,X64)) | ~p(s(bool,X65)) | ~sP5(X65,'F_0',X61)) )),
  inference(superposition,[],[f1996,f327])).
fof(f3893,plain,(
  ( ! [X59,X57,X60,X58,X56] : (s(bool,X56) != s(bool,X57) | s(bool,X57) != s(bool,X58) | s(bool,X58) != s(bool,X59) | ~p(s(bool,X59)) | ~p(s(bool,X60)) | ~sP5(X60,X56,'F_0')) )),
  inference(superposition,[],[f1996,f327])).
fof(f3892,plain,(
  ( ! [X54,X52,X55,X53,X51] : (s(bool,X51) != s(bool,X52) | s(bool,X52) != s(bool,X53) | s(bool,X53) != s(bool,X54) | ~p(s(bool,X54)) | sP5(X55,'F_0',X51) | p(s(bool,X55))) )),
  inference(superposition,[],[f1996,f325])).
fof(f3891,plain,(
  ( ! [X47,X50,X48,X46,X49] : (s(bool,X46) != s(bool,X47) | s(bool,X47) != s(bool,X48) | s(bool,X48) != s(bool,X49) | ~p(s(bool,X49)) | sP5(X50,X46,'F_0') | p(s(bool,X50))) )),
  inference(superposition,[],[f1996,f325])).
fof(f3880,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) != s(bool,X9) | s(bool,X9) != s(bool,X10) | s(bool,X10) != s(bool,X11) | ~p(s(bool,X11)) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f1996,f222])).
fof(f222,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.util_prob.BIGINTER_SUBSET',aHL_BOOL_CASES)).
fof(f3878,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) != s(bool,X2) | s(bool,X2) != s(bool,X3) | ~p(s(bool,X3)) | p(s(bool,X0))) )),
  inference(superposition,[],[f1996,f235])).
fof(f235,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f52])).
fof(f52,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X9] : (s(bool,'F_0') = s(bool,X9) <=> ~p(s(bool,X9)))),
  file('thm.util_prob.BIGINTER_SUBSET','thm.bool.EQ_CLAUSES|split|2')).
fof(f3849,plain,(
  ( ! [X52,X50,X53,X51,X49] : (s(bool,X50) != s(bool,X51) | s(bool,X49) != s(bool,X52) | s(bool,'T_0') = s(bool,X52) | ~p(s(bool,X51)) | ~p(s(bool,X53)) | ~sP5(X53,X49,X50)) )),
  inference(superposition,[],[f1723,f327])).
fof(f3848,plain,(
  ( ! [X47,X45,X48,X46,X44] : (s(bool,X45) != s(bool,X46) | s(bool,X44) != s(bool,X47) | s(bool,'T_0') = s(bool,X47) | ~p(s(bool,X46)) | ~p(s(bool,X48)) | ~sP5(X48,X45,X44)) )),
  inference(superposition,[],[f1723,f327])).
fof(f3847,plain,(
  ( ! [X39,X43,X41,X42,X40] : (s(bool,X40) != s(bool,X41) | s(bool,X39) != s(bool,X42) | s(bool,'T_0') = s(bool,X42) | ~p(s(bool,X41)) | sP5(X43,X39,X40) | p(s(bool,X43))) )),
  inference(superposition,[],[f1723,f325])).
fof(f3846,plain,(
  ( ! [X37,X35,X38,X36,X34] : (s(bool,X35) != s(bool,X36) | s(bool,X34) != s(bool,X37) | s(bool,'T_0') = s(bool,X37) | ~p(s(bool,X36)) | sP5(X38,X35,X34) | p(s(bool,X38))) )),
  inference(superposition,[],[f1723,f325])).
fof(f3827,plain,(
  ( ! [X52,X50,X53,X51,X49] : (s(bool,X50) != s(bool,X51) | s(bool,X51) != s(bool,X52) | s(bool,'T_0') = s(bool,X52) | ~p(s(bool,X49)) | ~p(s(bool,X53)) | ~sP5(X53,X49,X50)) )),
  inference(superposition,[],[f1723,f327])).
fof(f3826,plain,(
  ( ! [X47,X45,X48,X46,X44] : (s(bool,X45) != s(bool,X46) | s(bool,X46) != s(bool,X47) | s(bool,'T_0') = s(bool,X47) | ~p(s(bool,X44)) | ~p(s(bool,X48)) | ~sP5(X48,X45,X44)) )),
  inference(superposition,[],[f1723,f327])).
fof(f3825,plain,(
  ( ! [X39,X43,X41,X42,X40] : (s(bool,X40) != s(bool,X41) | s(bool,X41) != s(bool,X42) | s(bool,'T_0') = s(bool,X42) | ~p(s(bool,X39)) | sP5(X43,X39,X40) | p(s(bool,X43))) )),
  inference(superposition,[],[f1723,f325])).
fof(f3824,plain,(
  ( ! [X37,X35,X38,X36,X34] : (s(bool,X35) != s(bool,X36) | s(bool,X36) != s(bool,X37) | s(bool,'T_0') = s(bool,X37) | ~p(s(bool,X34)) | sP5(X38,X35,X34) | p(s(bool,X38))) )),
  inference(superposition,[],[f1723,f325])).
fof(f3785,plain,(
  ( ! [X52,X50,X53,X51,X49] : (s(bool,X50) != s(bool,X51) | s(bool,X49) != s(bool,X52) | p(s(bool,X52)) | s(bool,'F_0') = s(bool,X51) | ~p(s(bool,X53)) | ~sP5(X53,X49,X50)) )),
  inference(superposition,[],[f1342,f327])).
fof(f3784,plain,(
  ( ! [X47,X45,X48,X46,X44] : (s(bool,X45) != s(bool,X46) | s(bool,X44) != s(bool,X47) | p(s(bool,X47)) | s(bool,'F_0') = s(bool,X46) | ~p(s(bool,X48)) | ~sP5(X48,X45,X44)) )),
  inference(superposition,[],[f1342,f327])).
fof(f3783,plain,(
  ( ! [X39,X43,X41,X42,X40] : (s(bool,X40) != s(bool,X41) | s(bool,X39) != s(bool,X42) | p(s(bool,X42)) | s(bool,'F_0') = s(bool,X41) | sP5(X43,X39,X40) | p(s(bool,X43))) )),
  inference(superposition,[],[f1342,f325])).
fof(f3782,plain,(
  ( ! [X37,X35,X38,X36,X34] : (s(bool,X35) != s(bool,X36) | s(bool,X34) != s(bool,X37) | p(s(bool,X37)) | s(bool,'F_0') = s(bool,X36) | sP5(X38,X35,X34) | p(s(bool,X38))) )),
  inference(superposition,[],[f1342,f325])).
fof(f3763,plain,(
  ( ! [X52,X50,X53,X51,X49] : (s(bool,X50) != s(bool,X51) | s(bool,X51) != s(bool,X52) | p(s(bool,X52)) | s(bool,'F_0') = s(bool,X49) | ~p(s(bool,X53)) | ~sP5(X53,X49,X50)) )),
  inference(superposition,[],[f1342,f327])).
fof(f3762,plain,(
  ( ! [X47,X45,X48,X46,X44] : (s(bool,X45) != s(bool,X46) | s(bool,X46) != s(bool,X47) | p(s(bool,X47)) | s(bool,'F_0') = s(bool,X44) | ~p(s(bool,X48)) | ~sP5(X48,X45,X44)) )),
  inference(superposition,[],[f1342,f327])).
fof(f3761,plain,(
  ( ! [X39,X43,X41,X42,X40] : (s(bool,X40) != s(bool,X41) | s(bool,X41) != s(bool,X42) | p(s(bool,X42)) | s(bool,'F_0') = s(bool,X39) | sP5(X43,X39,X40) | p(s(bool,X43))) )),
  inference(superposition,[],[f1342,f325])).
fof(f3760,plain,(
  ( ! [X37,X35,X38,X36,X34] : (s(bool,X35) != s(bool,X36) | s(bool,X36) != s(bool,X37) | p(s(bool,X37)) | s(bool,'F_0') = s(bool,X34) | sP5(X38,X35,X34) | p(s(bool,X38))) )),
  inference(superposition,[],[f1342,f325])).
fof(f3743,plain,(
  ( ! [X4,X3] : (~sP0(X4,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),X3))) | p(s(bool,X4)) | s(bool,f(s(sK10,sK17(X3,sK11,bool,sK10)))) != s(bool,g(s(sK10,sK17(X3,sK11,bool,sK10))))) )),
  inference(superposition,[],[f744,f299])).
fof(f3742,plain,(
  ( ! [X2,X1] : (~sP0(X2,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),X1))) | p(s(bool,X2)) | s(bool,f(s(sK10,sK17(sK11,X1,bool,sK10)))) != s(bool,g(s(sK10,sK17(sK11,X1,bool,sK10))))) )),
  inference(superposition,[],[f744,f299])).
fof(f3741,plain,(
  ( ! [X0] : (~sP0(X0,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),sK11))),s(fun(sK10,bool),sK11))))) )),
  inference(superposition,[],[f744,f255])).
fof(f3740,plain,(
  ( ! [X4,X3] : (~sP0(X4,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X3)),sK10)),s(fun(sK10,bool),sK11))) | p(s(bool,X4)) | s(bool,f(s(fun(sK10,bool),sK17(X3,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X3,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f744,f299])).
fof(f3739,plain,(
  ( ! [X2,X1] : (~sP0(X2,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X1)),sK10)),s(fun(sK10,bool),sK11))) | p(s(bool,X2)) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X1,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X1,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f744,f299])).
fof(f3726,plain,(
  ( ! [X4,X3] : (~sP0('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),X3)),X4) | p(s(bool,X4)) | s(bool,f(s(sK10,sK17(X3,sK11,bool,sK10)))) != s(bool,g(s(sK10,sK17(X3,sK11,bool,sK10))))) )),
  inference(superposition,[],[f743,f299])).
fof(f3725,plain,(
  ( ! [X2,X1] : (~sP0('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),X1)),X2) | p(s(bool,X2)) | s(bool,f(s(sK10,sK17(sK11,X1,bool,sK10)))) != s(bool,g(s(sK10,sK17(sK11,X1,bool,sK10))))) )),
  inference(superposition,[],[f743,f299])).
fof(f3724,plain,(
  ( ! [X0] : (~sP0('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),'const.pred_set.EMPTY_0')),X0) | p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),sK11))),s(fun(sK10,bool),sK11))))) )),
  inference(superposition,[],[f743,f255])).
fof(f3723,plain,(
  ( ! [X4,X3] : (~sP0('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X3)),sK10)),s(fun(sK10,bool),sK11)),X4) | p(s(bool,X4)) | s(bool,f(s(fun(sK10,bool),sK17(X3,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X3,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f743,f299])).
fof(f3722,plain,(
  ( ! [X2,X1] : (~sP0('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X1)),sK10)),s(fun(sK10,bool),sK11)),X2) | p(s(bool,X2)) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X1,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X1,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f743,f299])).
fof(f3705,plain,(
  ( ! [X4,X3] : (sP1('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),X3)),X4) | ~p(s(bool,X4)) | s(bool,f(s(sK10,sK17(X3,sK11,bool,sK10)))) != s(bool,g(s(sK10,sK17(X3,sK11,bool,sK10))))) )),
  inference(superposition,[],[f660,f299])).
fof(f3704,plain,(
  ( ! [X2,X1] : (sP1('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),X1)),X2) | ~p(s(bool,X2)) | s(bool,f(s(sK10,sK17(sK11,X1,bool,sK10)))) != s(bool,g(s(sK10,sK17(sK11,X1,bool,sK10))))) )),
  inference(superposition,[],[f660,f299])).
fof(f3703,plain,(
  ( ! [X0] : (sP1('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),'const.pred_set.EMPTY_0')),X0) | ~p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),sK11))),s(fun(sK10,bool),sK11))))) )),
  inference(superposition,[],[f660,f255])).
fof(f3702,plain,(
  ( ! [X4,X3] : (sP1('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X3)),sK10)),s(fun(sK10,bool),sK11)),X4) | ~p(s(bool,X4)) | s(bool,f(s(fun(sK10,bool),sK17(X3,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X3,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f660,f299])).
fof(f3701,plain,(
  ( ! [X2,X1] : (sP1('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X1)),sK10)),s(fun(sK10,bool),sK11)),X2) | ~p(s(bool,X2)) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X1,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X1,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f660,f299])).
fof(f3684,plain,(
  ( ! [X4,X3] : (sP1(X4,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),X3))) | ~p(s(bool,X4)) | s(bool,f(s(sK10,sK17(X3,sK11,bool,sK10)))) != s(bool,g(s(sK10,sK17(X3,sK11,bool,sK10))))) )),
  inference(superposition,[],[f659,f299])).
fof(f3683,plain,(
  ( ! [X2,X1] : (sP1(X2,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),X1))) | ~p(s(bool,X2)) | s(bool,f(s(sK10,sK17(sK11,X1,bool,sK10)))) != s(bool,g(s(sK10,sK17(sK11,X1,bool,sK10))))) )),
  inference(superposition,[],[f659,f299])).
fof(f3682,plain,(
  ( ! [X0] : (sP1(X0,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),sK11))),s(fun(sK10,bool),sK11))))) )),
  inference(superposition,[],[f659,f255])).
fof(f3681,plain,(
  ( ! [X4,X3] : (sP1(X4,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X3)),sK10)),s(fun(sK10,bool),sK11))) | ~p(s(bool,X4)) | s(bool,f(s(fun(sK10,bool),sK17(X3,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X3,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f659,f299])).
fof(f3680,plain,(
  ( ! [X2,X1] : (sP1(X2,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X1)),sK10)),s(fun(sK10,bool),sK11))) | ~p(s(bool,X2)) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X1,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X1,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f659,f299])).
fof(f3666,plain,(
  ( ! [X4,X3] : (~sP1(X4,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),X3))) | p(s(bool,X4)) | s(bool,f(s(sK10,sK17(X3,sK11,bool,sK10)))) != s(bool,g(s(sK10,sK17(X3,sK11,bool,sK10))))) )),
  inference(superposition,[],[f658,f299])).
fof(f3665,plain,(
  ( ! [X2,X1] : (~sP1(X2,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),X1))) | p(s(bool,X2)) | s(bool,f(s(sK10,sK17(sK11,X1,bool,sK10)))) != s(bool,g(s(sK10,sK17(sK11,X1,bool,sK10))))) )),
  inference(superposition,[],[f658,f299])).
fof(f3664,plain,(
  ( ! [X0] : (~sP1(X0,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),sK11))),s(fun(sK10,bool),sK11))))) )),
  inference(superposition,[],[f658,f255])).
fof(f3663,plain,(
  ( ! [X4,X3] : (~sP1(X4,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X3)),sK10)),s(fun(sK10,bool),sK11))) | p(s(bool,X4)) | s(bool,f(s(fun(sK10,bool),sK17(X3,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X3,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f658,f299])).
fof(f3662,plain,(
  ( ! [X2,X1] : (~sP1(X2,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X1)),sK10)),s(fun(sK10,bool),sK11))) | p(s(bool,X2)) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X1,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X1,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f658,f299])).
fof(f3633,plain,(
  ( ! [X6,X7,X5] : (sP5('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),X5)),X6,X7) | s(bool,X6) = s(bool,X7) | s(bool,f(s(sK10,sK17(X5,sK11,bool,sK10)))) != s(bool,g(s(sK10,sK17(X5,sK11,bool,sK10))))) )),
  inference(superposition,[],[f429,f299])).
fof(f3632,plain,(
  ( ! [X4,X2,X3] : (sP5('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),X2)),X3,X4) | s(bool,X3) = s(bool,X4) | s(bool,f(s(sK10,sK17(sK11,X2,bool,sK10)))) != s(bool,g(s(sK10,sK17(sK11,X2,bool,sK10))))) )),
  inference(superposition,[],[f429,f299])).
fof(f3631,plain,(
  ( ! [X0,X1] : (sP5('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),'const.pred_set.EMPTY_0')),X0,X1) | s(bool,X0) = s(bool,X1) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),sK11))),s(fun(sK10,bool),sK11))))) )),
  inference(superposition,[],[f429,f255])).
fof(f3630,plain,(
  ( ! [X6,X7,X5] : (sP5('const.pred_set.SUBSET_2'(s(fun(sK10,bool),X5),s(fun(sK10,bool),sK11)),X6,X7) | s(bool,X6) = s(bool,X7) | s(bool,f(s(sK10,sK17(X5,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10)))) != s(bool,g(s(sK10,sK17(X5,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10))))) )),
  inference(superposition,[],[f429,f299])).
fof(f3629,plain,(
  ( ! [X4,X2,X3] : (sP5('const.pred_set.SUBSET_2'(s(fun(sK10,bool),X2),s(fun(sK10,bool),sK11)),X3,X4) | s(bool,X3) = s(bool,X4) | s(bool,f(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X2,bool,sK10)))) != s(bool,g(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X2,bool,sK10))))) )),
  inference(superposition,[],[f429,f299])).
fof(f3628,plain,(
  ( ! [X0,X1] : (sP5('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.EMPTY_0'),s(fun(sK10,bool),sK11)),X0,X1) | s(bool,X0) = s(bool,X1) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))),s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))))) )),
  inference(superposition,[],[f429,f255])).
fof(f3627,plain,(
  ( ! [X6,X7,X5] : (sP5('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X5))),s(fun(sK10,bool),sK11)),X6,X7) | s(bool,X6) = s(bool,X7) | s(bool,f(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f429,f299])).
fof(f3626,plain,(
  ( ! [X4,X2,X3] : (sP5('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X2))),s(fun(sK10,bool),sK11)),X3,X4) | s(bool,X3) = s(bool,X4) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f429,f299])).
fof(f3613,plain,(
  ( ! [X2,X3] : (s(bool,X2) = s(bool,X3) | sP3(X2,X3,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)))) )),
  inference(resolution,[],[f429,f307])).
fof(f307,plain,(
  ( ! [X2,X0,X1] : (sP3(X1,X2,X0) | ~sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f184])).
fof(f3612,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | sP2(X1,X0,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)))) )),
  inference(resolution,[],[f429,f306])).
fof(f306,plain,(
  ( ! [X2,X0,X1] : (sP2(X2,X1,X0) | ~sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f184])).
fof(f3606,plain,(
  ( ! [X4,X3] : (~sP1('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),X3)),X4) | p(s(bool,X4)) | s(bool,f(s(sK10,sK17(X3,sK11,bool,sK10)))) != s(bool,g(s(sK10,sK17(X3,sK11,bool,sK10))))) )),
  inference(superposition,[],[f657,f299])).
fof(f3605,plain,(
  ( ! [X2,X1] : (~sP1('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),X1)),X2) | p(s(bool,X2)) | s(bool,f(s(sK10,sK17(sK11,X1,bool,sK10)))) != s(bool,g(s(sK10,sK17(sK11,X1,bool,sK10))))) )),
  inference(superposition,[],[f657,f299])).
fof(f3604,plain,(
  ( ! [X0] : (~sP1('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),'const.pred_set.EMPTY_0')),X0) | p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),sK11))),s(fun(sK10,bool),sK11))))) )),
  inference(superposition,[],[f657,f255])).
fof(f3603,plain,(
  ( ! [X4,X3] : (~sP1('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X3)),sK10)),s(fun(sK10,bool),sK11)),X4) | p(s(bool,X4)) | s(bool,f(s(fun(sK10,bool),sK17(X3,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X3,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f657,f299])).
fof(f3602,plain,(
  ( ! [X2,X1] : (~sP1('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X1)),sK10)),s(fun(sK10,bool),sK11)),X2) | p(s(bool,X2)) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X1,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X1,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f657,f299])).
fof(f3594,plain,(
  ( ! [X281] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),sK14(X281,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)),bool))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(bool,bool),bool),X281))))))) )),
  inference(forward_demodulation,[],[f3520,f646])).
fof(f3520,plain,(
  ( ! [X281] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(bool,bool),bool),X281)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11))),s(fun(bool,bool),sK14(X281,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)),bool)))))) )),
  inference(superposition,[],[f263,f646])).
fof(f3519,plain,(
  ( ! [X280] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(bool,bool),bool),X280)))))) | p(s(bool,'const.bool.IN_2'(s(fun(bool,bool),sK14(X280,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)),bool)),s(fun(fun(bool,bool),bool),X280))))) )),
  inference(superposition,[],[f262,f646])).
fof(f3524,plain,(
  ( ! [X271] : (s(bool,X271) != s(bool,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11))) | s(bool,'F_0') = s(bool,X271)) )),
  inference(trivial_inequality_removal,[],[f3514])).
fof(f3514,plain,(
  ( ! [X271] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,X271) != s(bool,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11))) | s(bool,'F_0') = s(bool,X271)) )),
  inference(superposition,[],[f2040,f646])).
fof(f3413,plain,(
  ( ! [X105,X106] : (s(bool,'F_0') != s(bool,X105) | p(s(bool,X106)) | ~sP5(X106,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)),X105)) )),
  inference(superposition,[],[f328,f646])).
fof(f3412,plain,(
  ( ! [X103,X104] : (s(bool,'F_0') != s(bool,X103) | p(s(bool,X104)) | ~sP5(X104,X103,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)))) )),
  inference(superposition,[],[f328,f646])).
fof(f3407,plain,(
  ( ! [X94,X93] : (s(bool,'F_0') != s(bool,X93) | sP5(X94,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)),X93) | ~p(s(bool,X94))) )),
  inference(superposition,[],[f326,f646])).
fof(f3406,plain,(
  ( ! [X92,X91] : (s(bool,'F_0') != s(bool,X91) | sP5(X92,X91,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11))) | ~p(s(bool,X92))) )),
  inference(superposition,[],[f326,f646])).
fof(f3323,plain,(
  ( ! [X10,X11] : (s(bool,'F_0') = s(bool,X10) | ~p(s(bool,X11)) | ~sP5(X11,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)),X10)) )),
  inference(superposition,[],[f646,f327])).
fof(f3322,plain,(
  ( ! [X8,X9] : (s(bool,'F_0') = s(bool,X8) | ~p(s(bool,X9)) | ~sP5(X9,X8,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)))) )),
  inference(superposition,[],[f646,f327])).
fof(f3321,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') = s(bool,X6) | sP5(X7,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)),X6) | p(s(bool,X7))) )),
  inference(superposition,[],[f646,f325])).
fof(f3320,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') = s(bool,X4) | sP5(X5,X4,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11))) | p(s(bool,X5))) )),
  inference(superposition,[],[f646,f325])).
fof(f3312,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),X1))) | s(bool,f(s(sK10,sK17(X1,sK11,bool,sK10)))) != s(bool,g(s(sK10,sK17(X1,sK11,bool,sK10))))) )),
  inference(superposition,[],[f646,f299])).
fof(f3311,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),X0))) | s(bool,f(s(sK10,sK17(sK11,X0,bool,sK10)))) != s(bool,g(s(sK10,sK17(sK11,X0,bool,sK10))))) )),
  inference(superposition,[],[f646,f299])).
fof(f3310,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),sK11))),s(fun(sK10,bool),sK11))))),
  inference(superposition,[],[f646,f255])).
fof(f3309,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X1)),sK10)),s(fun(sK10,bool),sK11))) | s(bool,f(s(fun(sK10,bool),sK17(X1,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X1,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f646,f299])).
fof(f3308,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X0)),sK10)),s(fun(sK10,bool),sK11))) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X0,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X0,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f646,f299])).
fof(f3296,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') != s(bool,X10) | ~p(s(bool,X11)) | ~sP5(X11,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)),X10)) )),
  inference(superposition,[],[f644,f327])).
fof(f3295,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') != s(bool,X8) | ~p(s(bool,X9)) | ~sP5(X9,X8,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)))) )),
  inference(superposition,[],[f644,f327])).
fof(f3294,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X6) | sP5(X7,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)),X6) | p(s(bool,X7))) )),
  inference(superposition,[],[f644,f325])).
fof(f3293,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | sP5(X5,X4,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11))) | p(s(bool,X5))) )),
  inference(superposition,[],[f644,f325])).
fof(f3282,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),X1))) | s(bool,f(s(sK10,sK17(X1,sK11,bool,sK10)))) != s(bool,g(s(sK10,sK17(X1,sK11,bool,sK10))))) )),
  inference(superposition,[],[f644,f299])).
fof(f3281,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),X0))) | s(bool,f(s(sK10,sK17(sK11,X0,bool,sK10)))) != s(bool,g(s(sK10,sK17(sK11,X0,bool,sK10))))) )),
  inference(superposition,[],[f644,f299])).
fof(f3280,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),sK11))),s(fun(sK10,bool),sK11))))),
  inference(superposition,[],[f644,f255])).
fof(f3279,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X1)),sK10)),s(fun(sK10,bool),sK11))) | s(bool,f(s(fun(sK10,bool),sK17(X1,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X1,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f644,f299])).
fof(f3278,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X0)),sK10)),s(fun(sK10,bool),sK11))) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X0,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X0,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f644,f299])).
fof(f3270,plain,(
  ( ! [X6,X7,X5] : (sP9(X6,X7,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),X5))) | ~p(s(bool,X6)) | s(bool,f(s(sK10,sK17(X5,sK11,bool,sK10)))) != s(bool,g(s(sK10,sK17(X5,sK11,bool,sK10))))) )),
  inference(superposition,[],[f457,f299])).
fof(f3269,plain,(
  ( ! [X4,X2,X3] : (sP9(X3,X4,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),X2))) | ~p(s(bool,X3)) | s(bool,f(s(sK10,sK17(sK11,X2,bool,sK10)))) != s(bool,g(s(sK10,sK17(sK11,X2,bool,sK10))))) )),
  inference(superposition,[],[f457,f299])).
fof(f3268,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),sK11))),s(fun(sK10,bool),sK11))))) )),
  inference(superposition,[],[f457,f255])).
fof(f3267,plain,(
  ( ! [X6,X7,X5] : (sP9(X6,X7,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),X5),s(fun(sK10,bool),sK11))) | ~p(s(bool,X6)) | s(bool,f(s(sK10,sK17(X5,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10)))) != s(bool,g(s(sK10,sK17(X5,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10))))) )),
  inference(superposition,[],[f457,f299])).
fof(f3266,plain,(
  ( ! [X4,X2,X3] : (sP9(X3,X4,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),X2),s(fun(sK10,bool),sK11))) | ~p(s(bool,X3)) | s(bool,f(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X2,bool,sK10)))) != s(bool,g(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X2,bool,sK10))))) )),
  inference(superposition,[],[f457,f299])).
fof(f3265,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.EMPTY_0'),s(fun(sK10,bool),sK11))) | ~p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))),s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))))) )),
  inference(superposition,[],[f457,f255])).
fof(f3264,plain,(
  ( ! [X6,X7,X5] : (sP9(X6,X7,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X5))),s(fun(sK10,bool),sK11))) | ~p(s(bool,X6)) | s(bool,f(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f457,f299])).
fof(f3263,plain,(
  ( ! [X4,X2,X3] : (sP9(X3,X4,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X2))),s(fun(sK10,bool),sK11))) | ~p(s(bool,X3)) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f457,f299])).
fof(f3239,plain,(
  ( ! [X6,X7,X5] : (~sP9('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),X5)),X6,X7) | ~p(s(bool,X6)) | s(bool,f(s(sK10,sK17(X5,sK11,bool,sK10)))) != s(bool,g(s(sK10,sK17(X5,sK11,bool,sK10))))) )),
  inference(superposition,[],[f449,f299])).
fof(f3238,plain,(
  ( ! [X4,X2,X3] : (~sP9('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),X2)),X3,X4) | ~p(s(bool,X3)) | s(bool,f(s(sK10,sK17(sK11,X2,bool,sK10)))) != s(bool,g(s(sK10,sK17(sK11,X2,bool,sK10))))) )),
  inference(superposition,[],[f449,f299])).
fof(f3237,plain,(
  ( ! [X0,X1] : (~sP9('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),'const.pred_set.EMPTY_0')),X0,X1) | ~p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),sK11))),s(fun(sK10,bool),sK11))))) )),
  inference(superposition,[],[f449,f255])).
fof(f3236,plain,(
  ( ! [X6,X7,X5] : (~sP9('const.pred_set.SUBSET_2'(s(fun(sK10,bool),X5),s(fun(sK10,bool),sK11)),X6,X7) | ~p(s(bool,X6)) | s(bool,f(s(sK10,sK17(X5,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10)))) != s(bool,g(s(sK10,sK17(X5,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10))))) )),
  inference(superposition,[],[f449,f299])).
fof(f3235,plain,(
  ( ! [X4,X2,X3] : (~sP9('const.pred_set.SUBSET_2'(s(fun(sK10,bool),X2),s(fun(sK10,bool),sK11)),X3,X4) | ~p(s(bool,X3)) | s(bool,f(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X2,bool,sK10)))) != s(bool,g(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X2,bool,sK10))))) )),
  inference(superposition,[],[f449,f299])).
fof(f3234,plain,(
  ( ! [X0,X1] : (~sP9('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.EMPTY_0'),s(fun(sK10,bool),sK11)),X0,X1) | ~p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))),s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))))) )),
  inference(superposition,[],[f449,f255])).
fof(f3233,plain,(
  ( ! [X6,X7,X5] : (~sP9('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X5))),s(fun(sK10,bool),sK11)),X6,X7) | ~p(s(bool,X6)) | s(bool,f(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f449,f299])).
fof(f3232,plain,(
  ( ! [X4,X2,X3] : (~sP9('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X2))),s(fun(sK10,bool),sK11)),X3,X4) | ~p(s(bool,X3)) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f449,f299])).
fof(f3216,plain,(
  ( ! [X6,X7,X5] : (~sP9(X6,X7,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),X5))) | p(s(bool,X6)) | s(bool,f(s(sK10,sK17(X5,sK11,bool,sK10)))) != s(bool,g(s(sK10,sK17(X5,sK11,bool,sK10))))) )),
  inference(superposition,[],[f448,f299])).
fof(f3215,plain,(
  ( ! [X4,X2,X3] : (~sP9(X3,X4,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),X2))) | p(s(bool,X3)) | s(bool,f(s(sK10,sK17(sK11,X2,bool,sK10)))) != s(bool,g(s(sK10,sK17(sK11,X2,bool,sK10))))) )),
  inference(superposition,[],[f448,f299])).
fof(f3214,plain,(
  ( ! [X0,X1] : (~sP9(X0,X1,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),sK11))),s(fun(sK10,bool),sK11))))) )),
  inference(superposition,[],[f448,f255])).
fof(f3213,plain,(
  ( ! [X6,X7,X5] : (~sP9(X6,X7,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),X5),s(fun(sK10,bool),sK11))) | p(s(bool,X6)) | s(bool,f(s(sK10,sK17(X5,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10)))) != s(bool,g(s(sK10,sK17(X5,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10))))) )),
  inference(superposition,[],[f448,f299])).
fof(f3212,plain,(
  ( ! [X4,X2,X3] : (~sP9(X3,X4,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),X2),s(fun(sK10,bool),sK11))) | p(s(bool,X3)) | s(bool,f(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X2,bool,sK10)))) != s(bool,g(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X2,bool,sK10))))) )),
  inference(superposition,[],[f448,f299])).
fof(f3211,plain,(
  ( ! [X0,X1] : (~sP9(X0,X1,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.EMPTY_0'),s(fun(sK10,bool),sK11))) | p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))),s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))))) )),
  inference(superposition,[],[f448,f255])).
fof(f3210,plain,(
  ( ! [X6,X7,X5] : (~sP9(X6,X7,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X5))),s(fun(sK10,bool),sK11))) | p(s(bool,X6)) | s(bool,f(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f448,f299])).
fof(f3209,plain,(
  ( ! [X4,X2,X3] : (~sP9(X3,X4,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X2))),s(fun(sK10,bool),sK11))) | p(s(bool,X3)) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f448,f299])).
fof(f3190,plain,(
  ( ! [X6,X7,X5] : (~sP9('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),X5)),X6,X7) | p(s(bool,X7)) | s(bool,f(s(sK10,sK17(X5,sK11,bool,sK10)))) != s(bool,g(s(sK10,sK17(X5,sK11,bool,sK10))))) )),
  inference(superposition,[],[f447,f299])).
fof(f3189,plain,(
  ( ! [X4,X2,X3] : (~sP9('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),X2)),X3,X4) | p(s(bool,X4)) | s(bool,f(s(sK10,sK17(sK11,X2,bool,sK10)))) != s(bool,g(s(sK10,sK17(sK11,X2,bool,sK10))))) )),
  inference(superposition,[],[f447,f299])).
fof(f3188,plain,(
  ( ! [X0,X1] : (~sP9('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),'const.pred_set.EMPTY_0')),X0,X1) | p(s(bool,X1)) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),sK11))),s(fun(sK10,bool),sK11))))) )),
  inference(superposition,[],[f447,f255])).
fof(f3187,plain,(
  ( ! [X6,X7,X5] : (~sP9('const.pred_set.SUBSET_2'(s(fun(sK10,bool),X5),s(fun(sK10,bool),sK11)),X6,X7) | p(s(bool,X7)) | s(bool,f(s(sK10,sK17(X5,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10)))) != s(bool,g(s(sK10,sK17(X5,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10))))) )),
  inference(superposition,[],[f447,f299])).
fof(f3186,plain,(
  ( ! [X4,X2,X3] : (~sP9('const.pred_set.SUBSET_2'(s(fun(sK10,bool),X2),s(fun(sK10,bool),sK11)),X3,X4) | p(s(bool,X4)) | s(bool,f(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X2,bool,sK10)))) != s(bool,g(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X2,bool,sK10))))) )),
  inference(superposition,[],[f447,f299])).
fof(f3185,plain,(
  ( ! [X0,X1] : (~sP9('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.EMPTY_0'),s(fun(sK10,bool),sK11)),X0,X1) | p(s(bool,X1)) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))),s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))))) )),
  inference(superposition,[],[f447,f255])).
fof(f3184,plain,(
  ( ! [X6,X7,X5] : (~sP9('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X5))),s(fun(sK10,bool),sK11)),X6,X7) | p(s(bool,X7)) | s(bool,f(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f447,f299])).
fof(f3183,plain,(
  ( ! [X4,X2,X3] : (~sP9('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X2))),s(fun(sK10,bool),sK11)),X3,X4) | p(s(bool,X4)) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f447,f299])).
fof(f3161,plain,(
  ( ! [X6,X7,X5] : (~sP7('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),X5)),X6,X7) | ~p(s(bool,X6)) | s(bool,f(s(sK10,sK17(X5,sK11,bool,sK10)))) != s(bool,g(s(sK10,sK17(X5,sK11,bool,sK10))))) )),
  inference(superposition,[],[f432,f299])).
fof(f3160,plain,(
  ( ! [X4,X2,X3] : (~sP7('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),X2)),X3,X4) | ~p(s(bool,X3)) | s(bool,f(s(sK10,sK17(sK11,X2,bool,sK10)))) != s(bool,g(s(sK10,sK17(sK11,X2,bool,sK10))))) )),
  inference(superposition,[],[f432,f299])).
fof(f3159,plain,(
  ( ! [X0,X1] : (~sP7('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),'const.pred_set.EMPTY_0')),X0,X1) | ~p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),sK11))),s(fun(sK10,bool),sK11))))) )),
  inference(superposition,[],[f432,f255])).
fof(f3158,plain,(
  ( ! [X6,X7,X5] : (~sP7('const.pred_set.SUBSET_2'(s(fun(sK10,bool),X5),s(fun(sK10,bool),sK11)),X6,X7) | ~p(s(bool,X6)) | s(bool,f(s(sK10,sK17(X5,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10)))) != s(bool,g(s(sK10,sK17(X5,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10))))) )),
  inference(superposition,[],[f432,f299])).
fof(f3157,plain,(
  ( ! [X4,X2,X3] : (~sP7('const.pred_set.SUBSET_2'(s(fun(sK10,bool),X2),s(fun(sK10,bool),sK11)),X3,X4) | ~p(s(bool,X3)) | s(bool,f(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X2,bool,sK10)))) != s(bool,g(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X2,bool,sK10))))) )),
  inference(superposition,[],[f432,f299])).
fof(f3156,plain,(
  ( ! [X0,X1] : (~sP7('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.EMPTY_0'),s(fun(sK10,bool),sK11)),X0,X1) | ~p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))),s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))))) )),
  inference(superposition,[],[f432,f255])).
fof(f3155,plain,(
  ( ! [X6,X7,X5] : (~sP7('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X5))),s(fun(sK10,bool),sK11)),X6,X7) | ~p(s(bool,X6)) | s(bool,f(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f432,f299])).
fof(f3154,plain,(
  ( ! [X4,X2,X3] : (~sP7('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X2))),s(fun(sK10,bool),sK11)),X3,X4) | ~p(s(bool,X3)) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f432,f299])).
fof(f3132,plain,(
  ( ! [X6,X7,X5] : (~sP7('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),X5)),X6,X7) | ~p(s(bool,X7)) | s(bool,f(s(sK10,sK17(X5,sK11,bool,sK10)))) != s(bool,g(s(sK10,sK17(X5,sK11,bool,sK10))))) )),
  inference(superposition,[],[f431,f299])).
fof(f3131,plain,(
  ( ! [X4,X2,X3] : (~sP7('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),X2)),X3,X4) | ~p(s(bool,X4)) | s(bool,f(s(sK10,sK17(sK11,X2,bool,sK10)))) != s(bool,g(s(sK10,sK17(sK11,X2,bool,sK10))))) )),
  inference(superposition,[],[f431,f299])).
fof(f3130,plain,(
  ( ! [X0,X1] : (~sP7('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),'const.pred_set.EMPTY_0')),X0,X1) | ~p(s(bool,X1)) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),sK11))),s(fun(sK10,bool),sK11))))) )),
  inference(superposition,[],[f431,f255])).
fof(f3129,plain,(
  ( ! [X6,X7,X5] : (~sP7('const.pred_set.SUBSET_2'(s(fun(sK10,bool),X5),s(fun(sK10,bool),sK11)),X6,X7) | ~p(s(bool,X7)) | s(bool,f(s(sK10,sK17(X5,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10)))) != s(bool,g(s(sK10,sK17(X5,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10))))) )),
  inference(superposition,[],[f431,f299])).
fof(f3128,plain,(
  ( ! [X4,X2,X3] : (~sP7('const.pred_set.SUBSET_2'(s(fun(sK10,bool),X2),s(fun(sK10,bool),sK11)),X3,X4) | ~p(s(bool,X4)) | s(bool,f(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X2,bool,sK10)))) != s(bool,g(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X2,bool,sK10))))) )),
  inference(superposition,[],[f431,f299])).
fof(f3127,plain,(
  ( ! [X0,X1] : (~sP7('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.EMPTY_0'),s(fun(sK10,bool),sK11)),X0,X1) | ~p(s(bool,X1)) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))),s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))))) )),
  inference(superposition,[],[f431,f255])).
fof(f3126,plain,(
  ( ! [X6,X7,X5] : (~sP7('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X5))),s(fun(sK10,bool),sK11)),X6,X7) | ~p(s(bool,X7)) | s(bool,f(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f431,f299])).
fof(f3125,plain,(
  ( ! [X4,X2,X3] : (~sP7('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X2))),s(fun(sK10,bool),sK11)),X3,X4) | ~p(s(bool,X4)) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f431,f299])).
fof(f3110,plain,(
  ( ! [X4,X3] : (~sP0(X4,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),X3))) | p(s(bool,X4)) | s(bool,f(s(sK10,sK17(X3,sK11,bool,sK10)))) != s(bool,g(s(sK10,sK17(X3,sK11,bool,sK10))))) )),
  inference(superposition,[],[f499,f299])).
fof(f3109,plain,(
  ( ! [X2,X1] : (~sP0(X2,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),X1))) | p(s(bool,X2)) | s(bool,f(s(sK10,sK17(sK11,X1,bool,sK10)))) != s(bool,g(s(sK10,sK17(sK11,X1,bool,sK10))))) )),
  inference(superposition,[],[f499,f299])).
fof(f3108,plain,(
  ( ! [X0] : (~sP0(X0,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),sK11))),s(fun(sK10,bool),sK11))))) )),
  inference(superposition,[],[f499,f255])).
fof(f3107,plain,(
  ( ! [X4,X3] : (~sP0(X4,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),X3),s(fun(sK10,bool),sK11))) | p(s(bool,X4)) | s(bool,f(s(sK10,sK17(X3,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10)))) != s(bool,g(s(sK10,sK17(X3,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10))))) )),
  inference(superposition,[],[f499,f299])).
fof(f3106,plain,(
  ( ! [X2,X1] : (~sP0(X2,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),X1),s(fun(sK10,bool),sK11))) | p(s(bool,X2)) | s(bool,f(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X1,bool,sK10)))) != s(bool,g(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X1,bool,sK10))))) )),
  inference(superposition,[],[f499,f299])).
fof(f3105,plain,(
  ( ! [X0] : (~sP0(X0,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.EMPTY_0'),s(fun(sK10,bool),sK11))) | p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))),s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))))) )),
  inference(superposition,[],[f499,f255])).
fof(f3104,plain,(
  ( ! [X4,X3] : (~sP0(X4,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X3))),s(fun(sK10,bool),sK11))) | p(s(bool,X4)) | s(bool,f(s(fun(sK10,bool),sK17(X3,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X3,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f499,f299])).
fof(f3103,plain,(
  ( ! [X2,X1] : (~sP0(X2,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X1))),s(fun(sK10,bool),sK11))) | p(s(bool,X2)) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X1,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X1,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f499,f299])).
fof(f3086,plain,(
  ( ! [X6,X7,X5] : (sP8('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X5))),s(fun(fun(sK10,bool),bool),X5)),X6,X7) | s(bool,f(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f2877,f299])).
fof(f3085,plain,(
  ( ! [X4,X2,X3] : (sP8('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X2))),s(fun(fun(sK10,bool),bool),X2)),X3,X4) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f2877,f299])).
fof(f3083,plain,(
  ( ! [X6,X7,X5] : (sP8('const.bool.IN_2'(s(fun(sK10,bool),X5),s(fun(fun(sK10,bool),bool),sK12)),X6,X7) | s(bool,f(s(sK10,sK17(X5,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10)))) != s(bool,g(s(sK10,sK17(X5,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10))))) )),
  inference(superposition,[],[f2877,f299])).
fof(f3082,plain,(
  ( ! [X4,X2,X3] : (sP8('const.bool.IN_2'(s(fun(sK10,bool),X2),s(fun(fun(sK10,bool),bool),sK12)),X3,X4) | s(bool,f(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X2,bool,sK10)))) != s(bool,g(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X2,bool,sK10))))) )),
  inference(superposition,[],[f2877,f299])).
fof(f3081,plain,(
  ( ! [X0,X1] : (sP8('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK10,bool),bool),sK12)),X0,X1) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))),s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))))) )),
  inference(superposition,[],[f2877,f255])).
fof(f3070,plain,(
  ( ! [X6,X7,X5] : (sP8(X6,X7,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X5))),s(fun(fun(sK10,bool),bool),X5))) | s(bool,f(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f2876,f299])).
fof(f3069,plain,(
  ( ! [X4,X2,X3] : (sP8(X3,X4,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X2))),s(fun(fun(sK10,bool),bool),X2))) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f2876,f299])).
fof(f3067,plain,(
  ( ! [X6,X7,X5] : (sP8(X6,X7,'const.bool.IN_2'(s(fun(sK10,bool),X5),s(fun(fun(sK10,bool),bool),sK12))) | s(bool,f(s(sK10,sK17(X5,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10)))) != s(bool,g(s(sK10,sK17(X5,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10))))) )),
  inference(superposition,[],[f2876,f299])).
fof(f3066,plain,(
  ( ! [X4,X2,X3] : (sP8(X3,X4,'const.bool.IN_2'(s(fun(sK10,bool),X2),s(fun(fun(sK10,bool),bool),sK12))) | s(bool,f(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X2,bool,sK10)))) != s(bool,g(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X2,bool,sK10))))) )),
  inference(superposition,[],[f2876,f299])).
fof(f3065,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK10,bool),bool),sK12))) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))),s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))))) )),
  inference(superposition,[],[f2876,f255])).
fof(f3055,plain,(
  ( ! [X4,X3] : (~sP0('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),X3)),X4) | p(s(bool,X4)) | s(bool,f(s(sK10,sK17(X3,sK11,bool,sK10)))) != s(bool,g(s(sK10,sK17(X3,sK11,bool,sK10))))) )),
  inference(superposition,[],[f498,f299])).
fof(f3054,plain,(
  ( ! [X2,X1] : (~sP0('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),X1)),X2) | p(s(bool,X2)) | s(bool,f(s(sK10,sK17(sK11,X1,bool,sK10)))) != s(bool,g(s(sK10,sK17(sK11,X1,bool,sK10))))) )),
  inference(superposition,[],[f498,f299])).
fof(f3053,plain,(
  ( ! [X0] : (~sP0('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),'const.pred_set.EMPTY_0')),X0) | p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),sK11))),s(fun(sK10,bool),sK11))))) )),
  inference(superposition,[],[f498,f255])).
fof(f3052,plain,(
  ( ! [X4,X3] : (~sP0('const.pred_set.SUBSET_2'(s(fun(sK10,bool),X3),s(fun(sK10,bool),sK11)),X4) | p(s(bool,X4)) | s(bool,f(s(sK10,sK17(X3,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10)))) != s(bool,g(s(sK10,sK17(X3,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10))))) )),
  inference(superposition,[],[f498,f299])).
fof(f3051,plain,(
  ( ! [X2,X1] : (~sP0('const.pred_set.SUBSET_2'(s(fun(sK10,bool),X1),s(fun(sK10,bool),sK11)),X2) | p(s(bool,X2)) | s(bool,f(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X1,bool,sK10)))) != s(bool,g(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X1,bool,sK10))))) )),
  inference(superposition,[],[f498,f299])).
fof(f3050,plain,(
  ( ! [X0] : (~sP0('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.EMPTY_0'),s(fun(sK10,bool),sK11)),X0) | p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))),s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))))) )),
  inference(superposition,[],[f498,f255])).
fof(f3049,plain,(
  ( ! [X4,X3] : (~sP0('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X3))),s(fun(sK10,bool),sK11)),X4) | p(s(bool,X4)) | s(bool,f(s(fun(sK10,bool),sK17(X3,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X3,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f498,f299])).
fof(f3048,plain,(
  ( ! [X2,X1] : (~sP0('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X1))),s(fun(sK10,bool),sK11)),X2) | p(s(bool,X2)) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X1,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X1,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f498,f299])).
fof(f3031,plain,(
  ( ! [X6,X7,X5] : (sP6('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X5))),s(fun(fun(sK10,bool),bool),X5)),X6,X7) | s(bool,f(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f2862,f299])).
fof(f3030,plain,(
  ( ! [X4,X2,X3] : (sP6('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X2))),s(fun(fun(sK10,bool),bool),X2)),X3,X4) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f2862,f299])).
fof(f3028,plain,(
  ( ! [X6,X7,X5] : (sP6('const.bool.IN_2'(s(fun(sK10,bool),X5),s(fun(fun(sK10,bool),bool),sK12)),X6,X7) | s(bool,f(s(sK10,sK17(X5,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10)))) != s(bool,g(s(sK10,sK17(X5,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10))))) )),
  inference(superposition,[],[f2862,f299])).
fof(f3027,plain,(
  ( ! [X4,X2,X3] : (sP6('const.bool.IN_2'(s(fun(sK10,bool),X2),s(fun(fun(sK10,bool),bool),sK12)),X3,X4) | s(bool,f(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X2,bool,sK10)))) != s(bool,g(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X2,bool,sK10))))) )),
  inference(superposition,[],[f2862,f299])).
fof(f3026,plain,(
  ( ! [X0,X1] : (sP6('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK10,bool),bool),sK12)),X0,X1) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))),s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))))) )),
  inference(superposition,[],[f2862,f255])).
fof(f3014,plain,(
  ( ! [X6,X7,X5] : (sP3('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X5))),s(fun(fun(sK10,bool),bool),X5)),X6,X7) | s(bool,f(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f2848,f299])).
fof(f3013,plain,(
  ( ! [X4,X2,X3] : (sP3('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X2))),s(fun(fun(sK10,bool),bool),X2)),X3,X4) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f2848,f299])).
fof(f3011,plain,(
  ( ! [X6,X7,X5] : (sP3('const.bool.IN_2'(s(fun(sK10,bool),X5),s(fun(fun(sK10,bool),bool),sK12)),X6,X7) | s(bool,f(s(sK10,sK17(X5,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10)))) != s(bool,g(s(sK10,sK17(X5,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10))))) )),
  inference(superposition,[],[f2848,f299])).
fof(f3010,plain,(
  ( ! [X4,X2,X3] : (sP3('const.bool.IN_2'(s(fun(sK10,bool),X2),s(fun(fun(sK10,bool),bool),sK12)),X3,X4) | s(bool,f(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X2,bool,sK10)))) != s(bool,g(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X2,bool,sK10))))) )),
  inference(superposition,[],[f2848,f299])).
fof(f3009,plain,(
  ( ! [X0,X1] : (sP3('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK10,bool),bool),sK12)),X0,X1) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))),s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))))) )),
  inference(superposition,[],[f2848,f255])).
fof(f3019,plain,(
  ( ! [X6,X7] : (sP5(X6,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X7) | p(s(bool,X6)) | ~sP2(X7,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X6)) )),
  inference(subsumption_resolution,[],[f3008,f316])).
fof(f3008,plain,(
  ( ! [X6,X7] : (sP5(X6,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X7) | p(s(bool,X6)) | ~sP4(X6,X7,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12))) | ~sP2(X7,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X6)) )),
  inference(resolution,[],[f2848,f312])).
fof(f2999,plain,(
  ( ! [X6,X7,X5] : (sP3(X6,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X5))),s(fun(fun(sK10,bool),bool),X5)),X7) | s(bool,f(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f2847,f299])).
fof(f2998,plain,(
  ( ! [X4,X2,X3] : (sP3(X3,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X2))),s(fun(fun(sK10,bool),bool),X2)),X4) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f2847,f299])).
fof(f2996,plain,(
  ( ! [X6,X7,X5] : (sP3(X6,'const.bool.IN_2'(s(fun(sK10,bool),X5),s(fun(fun(sK10,bool),bool),sK12)),X7) | s(bool,f(s(sK10,sK17(X5,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10)))) != s(bool,g(s(sK10,sK17(X5,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10))))) )),
  inference(superposition,[],[f2847,f299])).
fof(f2995,plain,(
  ( ! [X4,X2,X3] : (sP3(X3,'const.bool.IN_2'(s(fun(sK10,bool),X2),s(fun(fun(sK10,bool),bool),sK12)),X4) | s(bool,f(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X2,bool,sK10)))) != s(bool,g(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X2,bool,sK10))))) )),
  inference(superposition,[],[f2847,f299])).
fof(f2994,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK10,bool),bool),sK12)),X1) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))),s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))))) )),
  inference(superposition,[],[f2847,f255])).
fof(f2980,plain,(
  ( ! [X6,X7,X5] : (sP4('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X5))),s(fun(fun(sK10,bool),bool),X5)),X6,X7) | s(bool,f(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f2845,f299])).
fof(f2979,plain,(
  ( ! [X4,X2,X3] : (sP4('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X2))),s(fun(fun(sK10,bool),bool),X2)),X3,X4) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f2845,f299])).
fof(f2977,plain,(
  ( ! [X6,X7,X5] : (sP4('const.bool.IN_2'(s(fun(sK10,bool),X5),s(fun(fun(sK10,bool),bool),sK12)),X6,X7) | s(bool,f(s(sK10,sK17(X5,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10)))) != s(bool,g(s(sK10,sK17(X5,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10))))) )),
  inference(superposition,[],[f2845,f299])).
fof(f2976,plain,(
  ( ! [X4,X2,X3] : (sP4('const.bool.IN_2'(s(fun(sK10,bool),X2),s(fun(fun(sK10,bool),bool),sK12)),X3,X4) | s(bool,f(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X2,bool,sK10)))) != s(bool,g(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X2,bool,sK10))))) )),
  inference(superposition,[],[f2845,f299])).
fof(f2975,plain,(
  ( ! [X0,X1] : (sP4('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK10,bool),bool),sK12)),X0,X1) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))),s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))))) )),
  inference(superposition,[],[f2845,f255])).
fof(f2973,plain,(
  ( ! [X18] : (sP5('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),'F_0',X18) | s(bool,'T_0') != s(bool,X18)) )),
  inference(resolution,[],[f2845,f1137])).
fof(f2970,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') != s(bool,X14) | sP5('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),'F_0',X15) | s(bool,X14) != s(bool,X15)) )),
  inference(resolution,[],[f2845,f1955])).
fof(f2966,plain,(
  ( ! [X6,X7] : (sP5('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X6,X7) | ~p(s(bool,X7)) | s(bool,'F_0') != s(bool,X6)) )),
  inference(resolution,[],[f2845,f2178])).
fof(f2965,plain,(
  ( ! [X4,X5] : (sP5('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X4,X5) | p(s(bool,X4)) | s(bool,'T_0') != s(bool,X5)) )),
  inference(resolution,[],[f2845,f2180])).
fof(f2964,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | sP5('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X2,X3) | ~sP2(X3,X2,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)))) )),
  inference(resolution,[],[f2845,f2225])).
fof(f2963,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP5('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X1,X0) | ~sP3(X1,X0,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)))) )),
  inference(resolution,[],[f2845,f2331])).
fof(f2951,plain,(
  ( ! [X6,X7,X5] : (sP4(X6,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X5))),s(fun(fun(sK10,bool),bool),X5)),X7) | s(bool,f(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f2844,f299])).
fof(f2950,plain,(
  ( ! [X4,X2,X3] : (sP4(X3,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X2))),s(fun(fun(sK10,bool),bool),X2)),X4) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f2844,f299])).
fof(f2948,plain,(
  ( ! [X6,X7,X5] : (sP4(X6,'const.bool.IN_2'(s(fun(sK10,bool),X5),s(fun(fun(sK10,bool),bool),sK12)),X7) | s(bool,f(s(sK10,sK17(X5,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10)))) != s(bool,g(s(sK10,sK17(X5,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10))))) )),
  inference(superposition,[],[f2844,f299])).
fof(f2947,plain,(
  ( ! [X4,X2,X3] : (sP4(X3,'const.bool.IN_2'(s(fun(sK10,bool),X2),s(fun(fun(sK10,bool),bool),sK12)),X4) | s(bool,f(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X2,bool,sK10)))) != s(bool,g(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X2,bool,sK10))))) )),
  inference(superposition,[],[f2844,f299])).
fof(f2946,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK10,bool),bool),sK12)),X1) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))),s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))))) )),
  inference(superposition,[],[f2844,f255])).
fof(f2958,plain,(
  ( ! [X10,X11] : (sP5(X10,X11,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12))) | p(s(bool,X10)) | ~sP2('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X11,X10)) )),
  inference(subsumption_resolution,[],[f2939,f2847])).
fof(f2939,plain,(
  ( ! [X10,X11] : (sP5(X10,X11,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12))) | p(s(bool,X10)) | ~sP3(X11,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X10) | ~sP2('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X11,X10)) )),
  inference(resolution,[],[f2844,f312])).
fof(f2935,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | sP5(X3,X2,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12))) | ~sP2('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X2,X3)) )),
  inference(resolution,[],[f2844,f2225])).
fof(f2910,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | ~p(s(bool,X11)) | ~sP5(X11,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X10)) )),
  inference(superposition,[],[f1615,f327])).
fof(f2908,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | sP5(X7,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X6) | p(s(bool,X7))) )),
  inference(superposition,[],[f1615,f325])).
fof(f2907,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | sP5(X5,X4,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12))) | p(s(bool,X5))) )),
  inference(superposition,[],[f1615,f325])).
fof(f2896,plain,(
  ( ! [X1] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X1))),s(fun(fun(sK10,bool),bool),X1)))) | s(bool,f(s(fun(sK10,bool),sK17(X1,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X1,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f1615,f299])).
fof(f2895,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X0))),s(fun(fun(sK10,bool),bool),X0)))) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X0,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X0,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f1615,f299])).
fof(f2893,plain,(
  ( ! [X1] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X1),s(fun(fun(sK10,bool),bool),sK12)))) | s(bool,f(s(sK10,sK17(X1,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10)))) != s(bool,g(s(sK10,sK17(X1,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10))))) )),
  inference(superposition,[],[f1615,f299])).
fof(f2892,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X0),s(fun(fun(sK10,bool),bool),sK12)))) | s(bool,f(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X0,bool,sK10)))) != s(bool,g(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X0,bool,sK10))))) )),
  inference(superposition,[],[f1615,f299])).
fof(f2891,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK10,bool),bool),sK12)))) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))),s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))))),
  inference(superposition,[],[f1615,f255])).
fof(f2881,plain,(
  ( ! [X103,X102] : (~p(s(bool,X102)) | ~p(s(bool,X103)) | ~sP9(X103,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X102)) )),
  inference(resolution,[],[f1615,f360])).
fof(f360,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f209])).
fof(f209,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f208])).
fof(f208,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f138])).
fof(f138,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP9(X2,X0,X1))),
  inference(definition_folding,[],[f121,f137,f136])).
fof(f136,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f137,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f121,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f92])).
fof(f92,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f91])).
fof(f91,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X20,X21,X22] : ((p(s(bool,X22)) <=> (p(s(bool,X21)) => p(s(bool,X20)))) <=> ((p(s(bool,X22)) | p(s(bool,X21))) & (p(s(bool,X22)) | ~p(s(bool,X20))) & (~p(s(bool,X21)) | p(s(bool,X20)) | ~p(s(bool,X22)))))),
  file('thm.util_prob.BIGINTER_SUBSET','thm.sat.sat_27_prove_ltb')).
fof(f2880,plain,(
  ( ! [X101,X100] : (sP9(X100,X101,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12))) | ~p(s(bool,X100))) )),
  inference(resolution,[],[f1615,f358])).
fof(f358,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f209])).
fof(f2879,plain,(
  ( ! [X99,X98] : (sP9('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X98,X99) | p(s(bool,X98)) | ~p(s(bool,X99))) )),
  inference(resolution,[],[f1615,f357])).
fof(f357,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f209])).
fof(f2878,plain,(
  ( ! [X97,X96] : (sP9(X96,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X97) | ~p(s(bool,X97)) | p(s(bool,X96))) )),
  inference(resolution,[],[f1615,f357])).
fof(f2877,plain,(
  ( ! [X94,X95] : (sP8('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X94,X95)) )),
  inference(resolution,[],[f1615,f356])).
fof(f356,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f207])).
fof(f207,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f206])).
fof(f206,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f205])).
fof(f205,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f136])).
fof(f2876,plain,(
  ( ! [X92,X93] : (sP8(X92,X93,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)))) )),
  inference(resolution,[],[f1615,f354])).
fof(f354,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f207])).
fof(f2875,plain,(
  ( ! [X90,X91] : (~p(s(bool,X90)) | ~p(s(bool,X91)) | ~sP8(X90,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X91)) )),
  inference(resolution,[],[f1615,f353])).
fof(f353,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f207])).
fof(f2874,plain,(
  ( ! [X88,X89] : (sP9(X88,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X89) | ~sP8(X88,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X89) | ~p(s(bool,X89))) )),
  inference(resolution,[],[f1615,f352])).
fof(f352,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f204])).
fof(f204,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP8(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f203])).
fof(f203,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1) | sP9(X2,X0,X1)))),
  inference(flattening,[],[f202])).
fof(f202,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1)) | sP9(X2,X0,X1)))),
  inference(nnf_transformation,[],[f137])).
fof(f2872,plain,(
  ( ! [X85,X84] : (~p(s(bool,X84)) | ~sP9('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X84,X85)) )),
  inference(resolution,[],[f1615,f347])).
fof(f347,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | ~sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f204])).
fof(f2871,plain,(
  ( ! [X83,X82] : (p(s(bool,X82)) | ~sP9(X82,X83,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)))) )),
  inference(resolution,[],[f1615,f346])).
fof(f346,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | p(s(bool,X0)) | ~sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f204])).
fof(f2870,plain,(
  ( ! [X80,X81] : (p(s(bool,X80)) | ~sP9('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X81,X80)) )),
  inference(resolution,[],[f1615,f346])).
fof(f2867,plain,(
  ( ! [X74,X75] : (p(s(bool,X74)) | ~p(s(bool,X75)) | ~sP7(X75,X74,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)))) )),
  inference(resolution,[],[f1615,f343])).
fof(f343,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP7(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f201])).
fof(f201,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f200])).
fof(f200,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f135])).
fof(f135,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP7(X2,X0,X1))),
  inference(definition_folding,[],[f90,f134,f133])).
fof(f133,plain,(
  ! [X2,X0,X1] : (sP6(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f134,plain,(
  ! [X2,X0,X1] : (sP7(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP6(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f90,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f89])).
fof(f89,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X20,X21,X22] : ((p(s(bool,X22)) <=> (p(s(bool,X21)) | p(s(bool,X20)))) <=> ((p(s(bool,X22)) | ~p(s(bool,X21))) & (p(s(bool,X22)) | ~p(s(bool,X20))) & (p(s(bool,X21)) | p(s(bool,X20)) | ~p(s(bool,X22)))))),
  file('thm.util_prob.BIGINTER_SUBSET','thm.sat.sat_25_prove_ltb')).
fof(f2866,plain,(
  ( ! [X72,X73] : (p(s(bool,X72)) | ~p(s(bool,X73)) | ~sP7(X73,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X72)) )),
  inference(resolution,[],[f1615,f343])).
fof(f2865,plain,(
  ( ! [X70,X71] : (sP7('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X70,X71) | p(s(bool,X70)) | p(s(bool,X71))) )),
  inference(resolution,[],[f1615,f340])).
fof(f340,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f201])).
fof(f2864,plain,(
  ( ! [X68,X69] : (sP7(X68,X69,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12))) | p(s(bool,X69)) | p(s(bool,X68))) )),
  inference(resolution,[],[f1615,f340])).
fof(f2863,plain,(
  ( ! [X66,X67] : (sP7(X66,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X67) | p(s(bool,X67)) | p(s(bool,X66))) )),
  inference(resolution,[],[f1615,f340])).
fof(f2862,plain,(
  ( ! [X64,X65] : (sP6('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X64,X65)) )),
  inference(resolution,[],[f1615,f339])).
fof(f339,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f199])).
fof(f199,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(rectify,[],[f198])).
fof(f198,plain,(
  ! [X2,X0,X1] : ((~sP6(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP6(X2,X0,X1)))),
  inference(flattening,[],[f197])).
fof(f197,plain,(
  ! [X2,X0,X1] : ((~sP6(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP6(X2,X0,X1)))),
  inference(nnf_transformation,[],[f133])).
fof(f2861,plain,(
  ( ! [X62,X63] : (~p(s(bool,X62)) | p(s(bool,X63)) | ~sP6(X62,X63,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)))) )),
  inference(resolution,[],[f1615,f336])).
fof(f336,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f199])).
fof(f2860,plain,(
  ( ! [X61,X60] : (~p(s(bool,X60)) | p(s(bool,X61)) | ~sP6(X60,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X61)) )),
  inference(resolution,[],[f1615,f336])).
fof(f2859,plain,(
  ( ! [X59,X58] : (sP7(X58,X59,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12))) | ~sP6(X58,X59,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12))) | p(s(bool,X59))) )),
  inference(resolution,[],[f1615,f335])).
fof(f335,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~sP6(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f196])).
fof(f196,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP6(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP6(X0,X1,X2) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f195])).
fof(f195,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP6(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP6(X2,X0,X1) | sP7(X2,X0,X1)))),
  inference(flattening,[],[f194])).
fof(f194,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP6(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP6(X2,X0,X1)) | sP7(X2,X0,X1)))),
  inference(nnf_transformation,[],[f134])).
fof(f2858,plain,(
  ( ! [X57,X56] : (sP7(X56,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X57) | ~sP6(X56,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X57) | p(s(bool,X57))) )),
  inference(resolution,[],[f1615,f335])).
fof(f2855,plain,(
  ( ! [X50,X51] : (~p(s(bool,X50)) | ~sP7('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X50,X51)) )),
  inference(resolution,[],[f1615,f330])).
fof(f330,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | ~sP7(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f196])).
fof(f2854,plain,(
  ( ! [X48,X49] : (~p(s(bool,X48)) | ~sP7('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X49,X48)) )),
  inference(resolution,[],[f1615,f329])).
fof(f329,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X0)) | ~sP7(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f196])).
fof(f2853,plain,(
  ( ! [X47,X46] : (s(bool,X46) != s(bool,X47) | ~sP5('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X47,X46)) )),
  inference(resolution,[],[f1615,f328])).
fof(f2852,plain,(
  ( ! [X45,X44] : (sP5('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X44,X45) | s(bool,X44) = s(bool,X45)) )),
  inference(resolution,[],[f1615,f325])).
fof(f2851,plain,(
  ( ! [X43,X42] : (p(s(bool,X42)) | p(s(bool,X43)) | ~sP2(X42,X43,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)))) )),
  inference(resolution,[],[f1615,f321])).
fof(f321,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(flattening,[],[f191])).
fof(f191,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(nnf_transformation,[],[f128])).
fof(f2850,plain,(
  ( ! [X41,X40] : (p(s(bool,X40)) | p(s(bool,X41)) | ~sP2(X40,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X41)) )),
  inference(resolution,[],[f1615,f321])).
fof(f2849,plain,(
  ( ! [X39,X38] : (p(s(bool,X38)) | p(s(bool,X39)) | ~sP2('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X38,X39)) )),
  inference(resolution,[],[f1615,f321])).
fof(f2848,plain,(
  ( ! [X37,X36] : (sP3('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X36,X37)) )),
  inference(resolution,[],[f1615,f320])).
fof(f320,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f190])).
fof(f190,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f189])).
fof(f189,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(flattening,[],[f188])).
fof(f188,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(nnf_transformation,[],[f129])).
fof(f2847,plain,(
  ( ! [X35,X34] : (sP3(X34,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X35)) )),
  inference(resolution,[],[f1615,f319])).
fof(f319,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f190])).
fof(f2846,plain,(
  ( ! [X33,X32] : (~p(s(bool,X32)) | ~p(s(bool,X33)) | ~sP3(X32,X33,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)))) )),
  inference(resolution,[],[f1615,f317])).
fof(f317,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f190])).
fof(f2920,plain,(
  ( ! [X24,X25] : (sP5('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X24,X25) | ~sP3(X24,X25,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12))) | ~sP2(X25,X24,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)))) )),
  inference(subsumption_resolution,[],[f2842,f2845])).
fof(f2845,plain,(
  ( ! [X30,X31] : (sP4('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X30,X31)) )),
  inference(resolution,[],[f1615,f316])).
fof(f2844,plain,(
  ( ! [X28,X29] : (sP4(X28,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X29)) )),
  inference(resolution,[],[f1615,f315])).
fof(f315,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f187])).
fof(f2843,plain,(
  ( ! [X26,X27] : (~p(s(bool,X26)) | ~p(s(bool,X27)) | ~sP4(X26,X27,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)))) )),
  inference(resolution,[],[f1615,f313])).
fof(f313,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f187])).
fof(f2842,plain,(
  ( ! [X24,X25] : (sP5('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X24,X25) | ~sP4('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X25,X24) | ~sP3(X24,X25,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12))) | ~sP2(X25,X24,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)))) )),
  inference(resolution,[],[f1615,f312])).
fof(f2841,plain,(
  ( ! [X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | ~sP5(X22,X23,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)))) )),
  inference(resolution,[],[f1615,f309])).
fof(f309,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f184])).
fof(f2836,plain,(
  ( ! [X17] : (sP1('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X17) | ~p(s(bool,X17))) )),
  inference(resolution,[],[f1615,f294])).
fof(f294,plain,(
  ( ! [X0,X1] : (sP1(X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f177])).
fof(f177,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1)))),
  inference(rectify,[],[f176])).
fof(f176,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0)))),
  inference(flattening,[],[f175])).
fof(f175,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP1(X1,X0)))),
  inference(nnf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X1,X0] : (sP1(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f2835,plain,(
  ( ! [X16] : (sP1(X16,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12))) | ~p(s(bool,X16))) )),
  inference(resolution,[],[f1615,f291])).
fof(f291,plain,(
  ( ! [X0,X1] : (sP1(X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f177])).
fof(f2834,plain,(
  ( ! [X15] : (~sP1(X15,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12))) | p(s(bool,X15))) )),
  inference(resolution,[],[f1615,f289])).
fof(f289,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | ~sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f177])).
fof(f2833,plain,(
  ( ! [X14] : (~sP1('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X14) | p(s(bool,X14))) )),
  inference(resolution,[],[f1615,f289])).
fof(f2919,plain,(
  ( ! [X11] : (p(s(bool,X11)) | ~sP0(X11,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)))) )),
  inference(subsumption_resolution,[],[f2830,f221])).
fof(f2830,plain,(
  ( ! [X11] : (p(s(bool,'F_0')) | p(s(bool,X11)) | ~sP0(X11,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)))) )),
  inference(resolution,[],[f1615,f274])).
fof(f274,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f170])).
fof(f170,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f169])).
fof(f169,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f124])).
fof(f124,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f2918,plain,(
  ( ! [X10] : (p(s(bool,X10)) | ~sP0('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X10)) )),
  inference(subsumption_resolution,[],[f2829,f221])).
fof(f2829,plain,(
  ( ! [X10] : (p(s(bool,'F_0')) | p(s(bool,X10)) | ~sP0('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X10)) )),
  inference(resolution,[],[f1615,f274])).
fof(f2826,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,happ(s(fun(X5,bool),X6),s(X5,X7)))) | ~p(s(bool,happ(s(fun(X5,bool),X6),s(X5,sK16(X6,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X5)))))) )),
  inference(resolution,[],[f1615,f266])).
fof(f266,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK16(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f168])).
fof(f168,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK15(X2,X1,X0))))) & ~p(s(bool,X1))) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) | p(s(bool,X1))) & ((~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK16(X2,X1,X0))))) & ~p(s(bool,X1))) | ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | p(s(bool,X1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK15,sK16])],[f167])).
fof(f167,plain,(
  ! [X0,X1,X2] : ((? [X6] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))) & ~p(s(bool,X1))) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) | p(s(bool,X1))) & ((? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) & ~p(s(bool,X1))) | ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | p(s(bool,X1)))))),
  inference(rectify,[],[f166])).
fof(f166,plain,(
  ! [X0,X1,X2] : ((? [X3] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) & ~p(s(bool,X1))) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,X1))) & ((? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) & ~p(s(bool,X1))) | ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | p(s(bool,X1)))))),
  inference(flattening,[],[f165])).
fof(f165,plain,(
  ! [X0,X1,X2] : ((? [X3] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) & ~p(s(bool,X1))) | (! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,X1)))) & ((? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) & ~p(s(bool,X1))) | ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | p(s(bool,X1)))))),
  inference(nnf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0,X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | p(s(bool,X1))) <=> (! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,X1))))),
  inference(flattening,[],[f72])).
fof(f72,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | p(s(bool,X1))) <=> (! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,X1))))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X6] : ! [X14,X15] : (! [X12] : (p(s(bool,happ(s(fun(X6,bool),X15),s(X6,X12)))) | p(s(bool,X14))) <=> (! [X12] : p(s(bool,happ(s(fun(X6,bool),X15),s(X6,X12)))) | p(s(bool,X14))))),
  file('thm.util_prob.BIGINTER_SUBSET','thm.bool.LEFT_FORALL_OR_THM')).
fof(f2825,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,happ(s(fun(X2,bool),X3),s(X2,X4)))) | ~p(s(bool,happ(s(fun(X2,bool),X3),s(X2,sK15(X3,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X2)))))) )),
  inference(resolution,[],[f1615,f264])).
fof(f264,plain,(
  ( ! [X2,X0,X5,X1] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK15(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f168])).
fof(f2822,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)))),
  inference(resolution,[],[f1615,f233])).
fof(f233,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f144])).
fof(f144,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f50])).
fof(f50,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X9] : (s(bool,'F_0') = s(bool,X9) <=> ~p(s(bool,X9)))),
  file('thm.util_prob.BIGINTER_SUBSET','thm.bool.EQ_CLAUSES|split|3')).
fof(f2820,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)))),
  inference(resolution,[],[f1615,f228])).
fof(f228,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f142])).
fof(f142,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f48])).
fof(f48,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X9] : (s(bool,'T_0') = s(bool,X9) <=> p(s(bool,X9)))),
  file('thm.util_prob.BIGINTER_SUBSET','thm.bool.EQ_CLAUSES|split|1')).
fof(f2819,plain,(
  ( ! [X1] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),X1)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK10,bool),bool),X1),s(fun(fun(sK10,bool),bool),sK12))))) )),
  inference(resolution,[],[f1615,f258])).
fof(f258,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f161])).
fof(f161,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))))) & ((p(s(bool,'const.bool.IN_2'(s(X0,sK13(X2,X1,X0)),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,sK13(X2,X1,X0)),s(fun(X0,bool),X2))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13])],[f160])).
fof(f160,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))))) & (? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f159])).
fof(f159,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))) & (? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(nnf_transformation,[],[f106])).
fof(f106,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(ennf_transformation,[],[f69])).
fof(f69,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) => p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(flattening,[],[f68])).
fof(f68,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) => p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X6] : ! [X8,X9] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X9)))) <=> ! [X12] : (p(s(bool,'const.bool.IN_2'(s(X6,X12),s(fun(X6,bool),X8)))) => p(s(bool,'const.bool.IN_2'(s(X6,X12),s(fun(X6,bool),X9))))))),
  file('thm.util_prob.BIGINTER_SUBSET','thm.pred_set.SUBSET_DEF')).
fof(f2818,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(fun(fun(sK10,bool),bool),sK12),s(fun(fun(fun(sK10,bool),bool),bool),X0)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),'const.pred_set.BIGINTER_1'(s(fun(fun(fun(sK10,bool),bool),bool),X0))))))) )),
  inference(resolution,[],[f1615,f261])).
fof(f2812,plain,(
  ( ! [X6,X7,X5] : (sP8(X6,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),X5))),s(fun(fun(sK10,bool),bool),X5)),X7) | s(bool,f(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f1020,f299])).
fof(f2811,plain,(
  ( ! [X4,X2,X3] : (sP8(X3,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),X2))),s(fun(fun(sK10,bool),bool),X2)),X4) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f1020,f299])).
fof(f2809,plain,(
  ( ! [X6,X7,X5] : (sP8(X6,'const.bool.IN_2'(s(fun(sK10,bool),X5),s(fun(fun(sK10,bool),bool),sK12)),X7) | s(bool,f(s(sK10,sK17(X5,'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10)))) != s(bool,g(s(sK10,sK17(X5,'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10))))) )),
  inference(superposition,[],[f1020,f299])).
fof(f2808,plain,(
  ( ! [X4,X2,X3] : (sP8(X3,'const.bool.IN_2'(s(fun(sK10,bool),X2),s(fun(fun(sK10,bool),bool),sK12)),X4) | s(bool,f(s(sK10,sK17('const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12)),X2,bool,sK10)))) != s(bool,g(s(sK10,sK17('const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12)),X2,bool,sK10))))) )),
  inference(superposition,[],[f1020,f299])).
fof(f2807,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK10,bool),bool),sK12)),X1) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))))),s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))))))) )),
  inference(superposition,[],[f1020,f255])).
fof(f2795,plain,(
  ( ! [X6,X7,X5] : (sP6(X6,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),X5))),s(fun(fun(sK10,bool),bool),X5)),X7) | s(bool,f(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f1005,f299])).
fof(f2794,plain,(
  ( ! [X4,X2,X3] : (sP6(X3,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),X2))),s(fun(fun(sK10,bool),bool),X2)),X4) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f1005,f299])).
fof(f2792,plain,(
  ( ! [X6,X7,X5] : (sP6(X6,'const.bool.IN_2'(s(fun(sK10,bool),X5),s(fun(fun(sK10,bool),bool),sK12)),X7) | s(bool,f(s(sK10,sK17(X5,'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10)))) != s(bool,g(s(sK10,sK17(X5,'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10))))) )),
  inference(superposition,[],[f1005,f299])).
fof(f2791,plain,(
  ( ! [X4,X2,X3] : (sP6(X3,'const.bool.IN_2'(s(fun(sK10,bool),X2),s(fun(fun(sK10,bool),bool),sK12)),X4) | s(bool,f(s(sK10,sK17('const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12)),X2,bool,sK10)))) != s(bool,g(s(sK10,sK17('const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12)),X2,bool,sK10))))) )),
  inference(superposition,[],[f1005,f299])).
fof(f2790,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK10,bool),bool),sK12)),X1) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))))),s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))))))) )),
  inference(superposition,[],[f1005,f255])).
fof(f2778,plain,(
  ( ! [X6,X7,X5] : (sP6(X6,X7,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),X5))),s(fun(fun(sK10,bool),bool),X5))) | s(bool,f(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f1004,f299])).
fof(f2777,plain,(
  ( ! [X4,X2,X3] : (sP6(X3,X4,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),X2))),s(fun(fun(sK10,bool),bool),X2))) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f1004,f299])).
fof(f2775,plain,(
  ( ! [X6,X7,X5] : (sP6(X6,X7,'const.bool.IN_2'(s(fun(sK10,bool),X5),s(fun(fun(sK10,bool),bool),sK12))) | s(bool,f(s(sK10,sK17(X5,'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10)))) != s(bool,g(s(sK10,sK17(X5,'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10))))) )),
  inference(superposition,[],[f1004,f299])).
fof(f2774,plain,(
  ( ! [X4,X2,X3] : (sP6(X3,X4,'const.bool.IN_2'(s(fun(sK10,bool),X2),s(fun(fun(sK10,bool),bool),sK12))) | s(bool,f(s(sK10,sK17('const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12)),X2,bool,sK10)))) != s(bool,g(s(sK10,sK17('const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12)),X2,bool,sK10))))) )),
  inference(superposition,[],[f1004,f299])).
fof(f2773,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK10,bool),bool),sK12))) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))))),s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))))))) )),
  inference(superposition,[],[f1004,f255])).
fof(f2761,plain,(
  ( ! [X6,X7,X5] : (sP2('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),X5))),s(fun(fun(sK10,bool),bool),X5)),X6,X7) | s(bool,f(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f996,f299])).
fof(f2760,plain,(
  ( ! [X4,X2,X3] : (sP2('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),X2))),s(fun(fun(sK10,bool),bool),X2)),X3,X4) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f996,f299])).
fof(f2758,plain,(
  ( ! [X6,X7,X5] : (sP2('const.bool.IN_2'(s(fun(sK10,bool),X5),s(fun(fun(sK10,bool),bool),sK12)),X6,X7) | s(bool,f(s(sK10,sK17(X5,'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10)))) != s(bool,g(s(sK10,sK17(X5,'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10))))) )),
  inference(superposition,[],[f996,f299])).
fof(f2757,plain,(
  ( ! [X4,X2,X3] : (sP2('const.bool.IN_2'(s(fun(sK10,bool),X2),s(fun(fun(sK10,bool),bool),sK12)),X3,X4) | s(bool,f(s(sK10,sK17('const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12)),X2,bool,sK10)))) != s(bool,g(s(sK10,sK17('const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12)),X2,bool,sK10))))) )),
  inference(superposition,[],[f996,f299])).
fof(f2756,plain,(
  ( ! [X0,X1] : (sP2('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK10,bool),bool),sK12)),X0,X1) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))))),s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))))))) )),
  inference(superposition,[],[f996,f255])).
fof(f2750,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12))) | p(s(bool,X0)) | s(bool,'F_0') != s(bool,X1)) )),
  inference(resolution,[],[f996,f2226])).
fof(f2742,plain,(
  ( ! [X6,X7,X5] : (sP2(X6,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),X5))),s(fun(fun(sK10,bool),bool),X5)),X7) | s(bool,f(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f995,f299])).
fof(f2741,plain,(
  ( ! [X4,X2,X3] : (sP2(X3,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),X2))),s(fun(fun(sK10,bool),bool),X2)),X4) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f995,f299])).
fof(f2739,plain,(
  ( ! [X6,X7,X5] : (sP2(X6,'const.bool.IN_2'(s(fun(sK10,bool),X5),s(fun(fun(sK10,bool),bool),sK12)),X7) | s(bool,f(s(sK10,sK17(X5,'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10)))) != s(bool,g(s(sK10,sK17(X5,'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10))))) )),
  inference(superposition,[],[f995,f299])).
fof(f2738,plain,(
  ( ! [X4,X2,X3] : (sP2(X3,'const.bool.IN_2'(s(fun(sK10,bool),X2),s(fun(fun(sK10,bool),bool),sK12)),X4) | s(bool,f(s(sK10,sK17('const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12)),X2,bool,sK10)))) != s(bool,g(s(sK10,sK17('const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12)),X2,bool,sK10))))) )),
  inference(superposition,[],[f995,f299])).
fof(f2737,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK10,bool),bool),sK12)),X1) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))))),s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))))))) )),
  inference(superposition,[],[f995,f255])).
fof(f2721,plain,(
  ( ! [X6,X7,X5] : (sP2(X6,X7,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),X5))),s(fun(fun(sK10,bool),bool),X5))) | s(bool,f(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f994,f299])).
fof(f2720,plain,(
  ( ! [X4,X2,X3] : (sP2(X3,X4,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),X2))),s(fun(fun(sK10,bool),bool),X2))) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f994,f299])).
fof(f2718,plain,(
  ( ! [X6,X7,X5] : (sP2(X6,X7,'const.bool.IN_2'(s(fun(sK10,bool),X5),s(fun(fun(sK10,bool),bool),sK12))) | s(bool,f(s(sK10,sK17(X5,'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10)))) != s(bool,g(s(sK10,sK17(X5,'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10))))) )),
  inference(superposition,[],[f994,f299])).
fof(f2717,plain,(
  ( ! [X4,X2,X3] : (sP2(X3,X4,'const.bool.IN_2'(s(fun(sK10,bool),X2),s(fun(fun(sK10,bool),bool),sK12))) | s(bool,f(s(sK10,sK17('const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12)),X2,bool,sK10)))) != s(bool,g(s(sK10,sK17('const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12)),X2,bool,sK10))))) )),
  inference(superposition,[],[f994,f299])).
fof(f2716,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK10,bool),bool),sK12))) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))))),s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))))))) )),
  inference(superposition,[],[f994,f255])).
fof(f2712,plain,(
  ( ! [X4,X5] : (sP5('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X4,X5) | p(s(bool,X4)) | p(s(bool,X5))) )),
  inference(resolution,[],[f994,f2246])).
fof(f2702,plain,(
  ( ! [X4,X3] : (sP1('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),X3)),X4) | ~p(s(bool,X4)) | s(bool,f(s(sK10,sK17(X3,sK11,bool,sK10)))) != s(bool,g(s(sK10,sK17(X3,sK11,bool,sK10))))) )),
  inference(superposition,[],[f412,f299])).
fof(f2701,plain,(
  ( ! [X2,X1] : (sP1('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),X1)),X2) | ~p(s(bool,X2)) | s(bool,f(s(sK10,sK17(sK11,X1,bool,sK10)))) != s(bool,g(s(sK10,sK17(sK11,X1,bool,sK10))))) )),
  inference(superposition,[],[f412,f299])).
fof(f2700,plain,(
  ( ! [X0] : (sP1('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),'const.pred_set.EMPTY_0')),X0) | ~p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),sK11))),s(fun(sK10,bool),sK11))))) )),
  inference(superposition,[],[f412,f255])).
fof(f2699,plain,(
  ( ! [X4,X3] : (sP1('const.pred_set.SUBSET_2'(s(fun(sK10,bool),X3),s(fun(sK10,bool),sK11)),X4) | ~p(s(bool,X4)) | s(bool,f(s(sK10,sK17(X3,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10)))) != s(bool,g(s(sK10,sK17(X3,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10))))) )),
  inference(superposition,[],[f412,f299])).
fof(f2698,plain,(
  ( ! [X2,X1] : (sP1('const.pred_set.SUBSET_2'(s(fun(sK10,bool),X1),s(fun(sK10,bool),sK11)),X2) | ~p(s(bool,X2)) | s(bool,f(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X1,bool,sK10)))) != s(bool,g(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X1,bool,sK10))))) )),
  inference(superposition,[],[f412,f299])).
fof(f2697,plain,(
  ( ! [X0] : (sP1('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.EMPTY_0'),s(fun(sK10,bool),sK11)),X0) | ~p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))),s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))))) )),
  inference(superposition,[],[f412,f255])).
fof(f2696,plain,(
  ( ! [X4,X3] : (sP1('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X3))),s(fun(sK10,bool),sK11)),X4) | ~p(s(bool,X4)) | s(bool,f(s(fun(sK10,bool),sK17(X3,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X3,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f412,f299])).
fof(f2695,plain,(
  ( ! [X2,X1] : (sP1('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X1))),s(fun(sK10,bool),sK11)),X2) | ~p(s(bool,X2)) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X1,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X1,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f412,f299])).
fof(f2678,plain,(
  ( ! [X6,X7,X5] : (sP3(X6,X7,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),X5))),s(fun(fun(sK10,bool),bool),X5))) | s(bool,f(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f993,f299])).
fof(f2677,plain,(
  ( ! [X4,X2,X3] : (sP3(X3,X4,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),X2))),s(fun(fun(sK10,bool),bool),X2))) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f993,f299])).
fof(f2675,plain,(
  ( ! [X6,X7,X5] : (sP3(X6,X7,'const.bool.IN_2'(s(fun(sK10,bool),X5),s(fun(fun(sK10,bool),bool),sK12))) | s(bool,f(s(sK10,sK17(X5,'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10)))) != s(bool,g(s(sK10,sK17(X5,'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10))))) )),
  inference(superposition,[],[f993,f299])).
fof(f2674,plain,(
  ( ! [X4,X2,X3] : (sP3(X3,X4,'const.bool.IN_2'(s(fun(sK10,bool),X2),s(fun(fun(sK10,bool),bool),sK12))) | s(bool,f(s(sK10,sK17('const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12)),X2,bool,sK10)))) != s(bool,g(s(sK10,sK17('const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12)),X2,bool,sK10))))) )),
  inference(superposition,[],[f993,f299])).
fof(f2673,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK10,bool),bool),sK12))) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))))),s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))))))) )),
  inference(superposition,[],[f993,f255])).
fof(f2660,plain,(
  ( ! [X6,X7,X5] : (sP4(X6,X7,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),X5))),s(fun(fun(sK10,bool),bool),X5))) | s(bool,f(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f990,f299])).
fof(f2659,plain,(
  ( ! [X4,X2,X3] : (sP4(X3,X4,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),X2))),s(fun(fun(sK10,bool),bool),X2))) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f990,f299])).
fof(f2657,plain,(
  ( ! [X6,X7,X5] : (sP4(X6,X7,'const.bool.IN_2'(s(fun(sK10,bool),X5),s(fun(fun(sK10,bool),bool),sK12))) | s(bool,f(s(sK10,sK17(X5,'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10)))) != s(bool,g(s(sK10,sK17(X5,'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10))))) )),
  inference(superposition,[],[f990,f299])).
fof(f2656,plain,(
  ( ! [X4,X2,X3] : (sP4(X3,X4,'const.bool.IN_2'(s(fun(sK10,bool),X2),s(fun(fun(sK10,bool),bool),sK12))) | s(bool,f(s(sK10,sK17('const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12)),X2,bool,sK10)))) != s(bool,g(s(sK10,sK17('const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12)),X2,bool,sK10))))) )),
  inference(superposition,[],[f990,f299])).
fof(f2655,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK10,bool),bool),sK12))) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))))),s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))))))) )),
  inference(superposition,[],[f990,f255])).
fof(f2667,plain,(
  ( ! [X10,X11] : (sP5(X10,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X11) | p(s(bool,X10)) | ~sP3('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X11,X10)) )),
  inference(subsumption_resolution,[],[f2654,f995])).
fof(f2654,plain,(
  ( ! [X10,X11] : (sP5(X10,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X11) | p(s(bool,X10)) | ~sP3('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X11,X10) | ~sP2(X11,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X10)) )),
  inference(resolution,[],[f990,f312])).
fof(f2649,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP5(X1,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X0) | ~sP3('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X0,X1)) )),
  inference(resolution,[],[f990,f2331])).
fof(f2599,plain,(
  ( ! [X19,X18] : (s(bool,'T_0') != s(bool,X18) | sP5('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)),'F_0',X19) | s(bool,X18) != s(bool,X19)) )),
  inference(resolution,[],[f1955,f670])).
fof(f2598,plain,(
  ( ! [X17,X16] : (s(bool,'T_0') != s(bool,X16) | sP5('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)),'F_0',X17) | s(bool,X16) != s(bool,X17)) )),
  inference(resolution,[],[f1955,f422])).
fof(f2554,plain,(
  ( ! [X52,X50,X53,X51,X49] : (s(bool,X50) != s(bool,X51) | s(bool,X49) != s(bool,X52) | p(s(bool,X52)) | ~p(s(bool,X51)) | ~p(s(bool,X53)) | ~sP5(X53,X49,X50)) )),
  inference(superposition,[],[f1340,f327])).
fof(f2553,plain,(
  ( ! [X47,X45,X48,X46,X44] : (s(bool,X45) != s(bool,X46) | s(bool,X44) != s(bool,X47) | p(s(bool,X47)) | ~p(s(bool,X46)) | ~p(s(bool,X48)) | ~sP5(X48,X45,X44)) )),
  inference(superposition,[],[f1340,f327])).
fof(f2552,plain,(
  ( ! [X39,X43,X41,X42,X40] : (s(bool,X40) != s(bool,X41) | s(bool,X39) != s(bool,X42) | p(s(bool,X42)) | ~p(s(bool,X41)) | sP5(X43,X39,X40) | p(s(bool,X43))) )),
  inference(superposition,[],[f1340,f325])).
fof(f2551,plain,(
  ( ! [X37,X35,X38,X36,X34] : (s(bool,X35) != s(bool,X36) | s(bool,X34) != s(bool,X37) | p(s(bool,X37)) | ~p(s(bool,X36)) | sP5(X38,X35,X34) | p(s(bool,X38))) )),
  inference(superposition,[],[f1340,f325])).
fof(f2533,plain,(
  ( ! [X52,X50,X53,X51,X49] : (s(bool,X50) != s(bool,X51) | s(bool,X51) != s(bool,X52) | p(s(bool,X52)) | ~p(s(bool,X49)) | ~p(s(bool,X53)) | ~sP5(X53,X49,X50)) )),
  inference(superposition,[],[f1340,f327])).
fof(f2532,plain,(
  ( ! [X47,X45,X48,X46,X44] : (s(bool,X45) != s(bool,X46) | s(bool,X46) != s(bool,X47) | p(s(bool,X47)) | ~p(s(bool,X44)) | ~p(s(bool,X48)) | ~sP5(X48,X45,X44)) )),
  inference(superposition,[],[f1340,f327])).
fof(f2531,plain,(
  ( ! [X39,X43,X41,X42,X40] : (s(bool,X40) != s(bool,X41) | s(bool,X41) != s(bool,X42) | p(s(bool,X42)) | ~p(s(bool,X39)) | sP5(X43,X39,X40) | p(s(bool,X43))) )),
  inference(superposition,[],[f1340,f325])).
fof(f2530,plain,(
  ( ! [X37,X35,X38,X36,X34] : (s(bool,X35) != s(bool,X36) | s(bool,X36) != s(bool,X37) | p(s(bool,X37)) | ~p(s(bool,X34)) | sP5(X38,X35,X34) | p(s(bool,X38))) )),
  inference(superposition,[],[f1340,f325])).
fof(f2452,plain,(
  ( ! [X10,X11,X9] : (sP3(X9,X10,X11) | s(bool,'F_0') != s(bool,X10) | s(bool,'T_0') != s(bool,X9)) )),
  inference(subsumption_resolution,[],[f2441,f318])).
fof(f318,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f190])).
fof(f2441,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') != s(bool,X9) | s(bool,'F_0') != s(bool,X10) | p(s(bool,X11)) | sP3(X9,X10,X11)) )),
  inference(resolution,[],[f1278,f307])).
fof(f2451,plain,(
  ( ! [X6,X8,X7] : (sP2(X7,X6,X8) | s(bool,'F_0') != s(bool,X7) | s(bool,'T_0') != s(bool,X6)) )),
  inference(subsumption_resolution,[],[f2440,f322])).
fof(f322,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f192])).
fof(f2440,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X6) | s(bool,'F_0') != s(bool,X7) | p(s(bool,X8)) | sP2(X7,X6,X8)) )),
  inference(resolution,[],[f1278,f306])).
fof(f2410,plain,(
  ( ! [X4,X3] : (sP1(X4,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),X3))) | ~p(s(bool,X4)) | s(bool,f(s(sK10,sK17(X3,sK11,bool,sK10)))) != s(bool,g(s(sK10,sK17(X3,sK11,bool,sK10))))) )),
  inference(superposition,[],[f411,f299])).
fof(f2409,plain,(
  ( ! [X2,X1] : (sP1(X2,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),X1))) | ~p(s(bool,X2)) | s(bool,f(s(sK10,sK17(sK11,X1,bool,sK10)))) != s(bool,g(s(sK10,sK17(sK11,X1,bool,sK10))))) )),
  inference(superposition,[],[f411,f299])).
fof(f2408,plain,(
  ( ! [X0] : (sP1(X0,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),sK11))),s(fun(sK10,bool),sK11))))) )),
  inference(superposition,[],[f411,f255])).
fof(f2407,plain,(
  ( ! [X4,X3] : (sP1(X4,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),X3),s(fun(sK10,bool),sK11))) | ~p(s(bool,X4)) | s(bool,f(s(sK10,sK17(X3,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10)))) != s(bool,g(s(sK10,sK17(X3,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10))))) )),
  inference(superposition,[],[f411,f299])).
fof(f2406,plain,(
  ( ! [X2,X1] : (sP1(X2,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),X1),s(fun(sK10,bool),sK11))) | ~p(s(bool,X2)) | s(bool,f(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X1,bool,sK10)))) != s(bool,g(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X1,bool,sK10))))) )),
  inference(superposition,[],[f411,f299])).
fof(f2405,plain,(
  ( ! [X0] : (sP1(X0,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.EMPTY_0'),s(fun(sK10,bool),sK11))) | ~p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))),s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))))) )),
  inference(superposition,[],[f411,f255])).
fof(f2404,plain,(
  ( ! [X4,X3] : (sP1(X4,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X3))),s(fun(sK10,bool),sK11))) | ~p(s(bool,X4)) | s(bool,f(s(fun(sK10,bool),sK17(X3,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X3,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f411,f299])).
fof(f2403,plain,(
  ( ! [X2,X1] : (sP1(X2,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X1))),s(fun(sK10,bool),sK11))) | ~p(s(bool,X2)) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X1,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X1,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f411,f299])).
fof(f2395,plain,(
  ( ! [X4,X5,X3] : (~sP3(X5,X3,X4) | p(s(bool,X4)) | s(bool,'T_0') != s(bool,X3) | s(bool,X3) != s(bool,X5)) )),
  inference(duplicate_literal_removal,[],[f2384])).
fof(f2384,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,X3) | p(s(bool,X4)) | ~sP3(X5,X3,X4) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5)) )),
  inference(resolution,[],[f2332,f328])).
fof(f2371,plain,(
  ( ! [X4,X5,X3] : (~sP3(X4,X3,X5) | ~p(s(bool,X3)) | p(s(bool,X5)) | s(bool,X3) != s(bool,X4)) )),
  inference(duplicate_literal_removal,[],[f2360])).
fof(f2360,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | ~sP3(X4,X3,X5) | p(s(bool,X5)) | p(s(bool,X5)) | s(bool,X3) != s(bool,X4)) )),
  inference(resolution,[],[f2335,f328])).
fof(f2340,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | sP5('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)),X19,X18) | ~sP3(X19,X18,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)))) )),
  inference(resolution,[],[f2331,f670])).
fof(f2339,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | sP5('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)),X17,X16) | ~sP3(X17,X16,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)))) )),
  inference(resolution,[],[f2331,f422])).
fof(f2335,plain,(
  ( ! [X6,X8,X7] : (sP5(X7,X8,X6) | ~p(s(bool,X6)) | ~sP3(X8,X6,X7) | p(s(bool,X7))) )),
  inference(resolution,[],[f2331,f316])).
fof(f2333,plain,(
  ( ! [X2,X0,X1] : (sP5(X1,X2,X0) | ~p(s(bool,X0)) | ~sP3(X2,X0,X1) | ~p(s(bool,X2))) )),
  inference(resolution,[],[f2331,f314])).
fof(f314,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f187])).
fof(f2332,plain,(
  ( ! [X14,X12,X13] : (sP5(X13,X14,X12) | s(bool,'T_0') != s(bool,X12) | p(s(bool,X13)) | ~sP3(X14,X12,X13)) )),
  inference(subsumption_resolution,[],[f2327,f316])).
fof(f2327,plain,(
  ( ! [X14,X12,X13] : (s(bool,'T_0') != s(bool,X12) | sP5(X13,X14,X12) | p(s(bool,X13)) | ~sP4(X13,X12,X14) | ~sP3(X14,X12,X13)) )),
  inference(resolution,[],[f2322,f312])).
fof(f2331,plain,(
  ( ! [X10,X11,X9] : (~sP4(X10,X9,X11) | ~p(s(bool,X9)) | sP5(X10,X11,X9) | ~sP3(X11,X9,X10)) )),
  inference(subsumption_resolution,[],[f2326,f229])).
fof(f229,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f142])).
fof(f2326,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') != s(bool,X9) | sP5(X10,X11,X9) | ~p(s(bool,X9)) | ~sP4(X10,X9,X11) | ~sP3(X11,X9,X10)) )),
  inference(resolution,[],[f2322,f310])).
fof(f310,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | ~p(s(bool,X2)) | ~sP4(X0,X2,X1) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f184])).
fof(f2322,plain,(
  ( ! [X6,X8,X7] : (sP2(X7,X6,X8) | s(bool,'T_0') != s(bool,X7)) )),
  inference(subsumption_resolution,[],[f2321,f322])).
fof(f2321,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X7) | p(s(bool,X8)) | sP2(X7,X6,X8)) )),
  inference(subsumption_resolution,[],[f2309,f323])).
fof(f323,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f192])).
fof(f2309,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | s(bool,'T_0') != s(bool,X7) | p(s(bool,X8)) | sP2(X7,X6,X8)) )),
  inference(resolution,[],[f2199,f306])).
fof(f2265,plain,(
  ( ! [X10,X11,X9] : (sP5(X9,X10,X11) | p(s(bool,X10)) | p(s(bool,X9)) | ~p(s(bool,X11))) )),
  inference(resolution,[],[f2247,f324])).
fof(f324,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f192])).
fof(f2254,plain,(
  ( ! [X4,X5,X3] : (sP5(X3,X4,X5) | p(s(bool,X4)) | p(s(bool,X5)) | ~p(s(bool,X3))) )),
  inference(resolution,[],[f2246,f322])).
fof(f2252,plain,(
  ( ! [X19,X18] : (p(s(bool,X18)) | sP5('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)),X18,X19) | ~sP2(X19,X18,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)))) )),
  inference(resolution,[],[f2225,f670])).
fof(f2251,plain,(
  ( ! [X17,X16] : (p(s(bool,X16)) | sP5('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)),X16,X17) | ~sP2(X17,X16,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)))) )),
  inference(resolution,[],[f2225,f422])).
fof(f2247,plain,(
  ( ! [X6,X8,X7] : (~sP2(X8,X6,X7) | sP5(X7,X6,X8) | p(s(bool,X6)) | p(s(bool,X7))) )),
  inference(resolution,[],[f2225,f316])).
fof(f2246,plain,(
  ( ! [X4,X5,X3] : (~sP2(X5,X3,X4) | sP5(X4,X3,X5) | p(s(bool,X3)) | p(s(bool,X5))) )),
  inference(resolution,[],[f2225,f315])).
fof(f2239,plain,(
  ( ! [X4,X3] : (~sP1(X4,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),X3))) | p(s(bool,X4)) | s(bool,f(s(sK10,sK17(X3,sK11,bool,sK10)))) != s(bool,g(s(sK10,sK17(X3,sK11,bool,sK10))))) )),
  inference(superposition,[],[f410,f299])).
fof(f2238,plain,(
  ( ! [X2,X1] : (~sP1(X2,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),X1))) | p(s(bool,X2)) | s(bool,f(s(sK10,sK17(sK11,X1,bool,sK10)))) != s(bool,g(s(sK10,sK17(sK11,X1,bool,sK10))))) )),
  inference(superposition,[],[f410,f299])).
fof(f2237,plain,(
  ( ! [X0] : (~sP1(X0,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),sK11))),s(fun(sK10,bool),sK11))))) )),
  inference(superposition,[],[f410,f255])).
fof(f2236,plain,(
  ( ! [X4,X3] : (~sP1(X4,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),X3),s(fun(sK10,bool),sK11))) | p(s(bool,X4)) | s(bool,f(s(sK10,sK17(X3,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10)))) != s(bool,g(s(sK10,sK17(X3,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10))))) )),
  inference(superposition,[],[f410,f299])).
fof(f2235,plain,(
  ( ! [X2,X1] : (~sP1(X2,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),X1),s(fun(sK10,bool),sK11))) | p(s(bool,X2)) | s(bool,f(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X1,bool,sK10)))) != s(bool,g(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X1,bool,sK10))))) )),
  inference(superposition,[],[f410,f299])).
fof(f2234,plain,(
  ( ! [X0] : (~sP1(X0,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.EMPTY_0'),s(fun(sK10,bool),sK11))) | p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))),s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))))) )),
  inference(superposition,[],[f410,f255])).
fof(f2233,plain,(
  ( ! [X4,X3] : (~sP1(X4,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X3))),s(fun(sK10,bool),sK11))) | p(s(bool,X4)) | s(bool,f(s(fun(sK10,bool),sK17(X3,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X3,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f410,f299])).
fof(f2232,plain,(
  ( ! [X2,X1] : (~sP1(X2,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X1))),s(fun(sK10,bool),sK11))) | p(s(bool,X2)) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X1,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X1,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f410,f299])).
fof(f2226,plain,(
  ( ! [X10,X11,X9] : (~sP2(X11,X9,X10) | sP5(X10,X9,X11) | p(s(bool,X10)) | s(bool,'F_0') != s(bool,X9)) )),
  inference(subsumption_resolution,[],[f2223,f316])).
fof(f2223,plain,(
  ( ! [X10,X11,X9] : (s(bool,'F_0') != s(bool,X9) | sP5(X10,X9,X11) | p(s(bool,X10)) | ~sP4(X10,X11,X9) | ~sP2(X11,X9,X10)) )),
  inference(resolution,[],[f2219,f312])).
fof(f2225,plain,(
  ( ! [X6,X8,X7] : (~sP4(X7,X8,X6) | p(s(bool,X6)) | sP5(X7,X6,X8) | ~sP2(X8,X6,X7)) )),
  inference(subsumption_resolution,[],[f2222,f233])).
fof(f2222,plain,(
  ( ! [X6,X8,X7] : (s(bool,'F_0') != s(bool,X6) | sP5(X7,X6,X8) | p(s(bool,X6)) | ~sP4(X7,X8,X6) | ~sP2(X8,X6,X7)) )),
  inference(resolution,[],[f2219,f311])).
fof(f311,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | p(s(bool,X1)) | ~sP4(X0,X2,X1) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f184])).
fof(f2219,plain,(
  ( ! [X10,X11,X9] : (sP3(X10,X9,X11) | s(bool,'F_0') != s(bool,X10)) )),
  inference(subsumption_resolution,[],[f2218,f319])).
fof(f2218,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | s(bool,'F_0') != s(bool,X10) | sP3(X10,X9,X11)) )),
  inference(subsumption_resolution,[],[f2208,f318])).
fof(f2208,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | s(bool,'F_0') != s(bool,X10) | p(s(bool,X11)) | sP3(X10,X9,X11)) )),
  inference(resolution,[],[f2189,f307])).
fof(f2204,plain,(
  ( ! [X19,X18] : (sP5('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)),X18,X19) | p(s(bool,X18)) | s(bool,'T_0') != s(bool,X19)) )),
  inference(resolution,[],[f2180,f670])).
fof(f2203,plain,(
  ( ! [X17,X16] : (sP5('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)),X16,X17) | p(s(bool,X16)) | s(bool,'T_0') != s(bool,X17)) )),
  inference(resolution,[],[f2180,f422])).
fof(f2199,plain,(
  ( ! [X6,X8,X7] : (sP5(X6,X7,X8) | p(s(bool,X7)) | s(bool,'T_0') != s(bool,X8) | p(s(bool,X6))) )),
  inference(resolution,[],[f2180,f316])).
fof(f2194,plain,(
  ( ! [X19,X18] : (sP5('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)),X18,X19) | ~p(s(bool,X19)) | s(bool,'F_0') != s(bool,X18)) )),
  inference(resolution,[],[f2178,f670])).
fof(f2193,plain,(
  ( ! [X17,X16] : (sP5('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)),X16,X17) | ~p(s(bool,X17)) | s(bool,'F_0') != s(bool,X16)) )),
  inference(resolution,[],[f2178,f422])).
fof(f2189,plain,(
  ( ! [X6,X8,X7] : (sP5(X6,X7,X8) | ~p(s(bool,X8)) | s(bool,'F_0') != s(bool,X7) | p(s(bool,X6))) )),
  inference(resolution,[],[f2178,f316])).
fof(f2180,plain,(
  ( ! [X6,X8,X7] : (~sP4(X8,X7,X6) | sP5(X8,X6,X7) | p(s(bool,X6)) | s(bool,'T_0') != s(bool,X7)) )),
  inference(subsumption_resolution,[],[f2179,f233])).
fof(f2179,plain,(
  ( ! [X6,X8,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,'T_0') != s(bool,X7) | sP5(X8,X6,X7) | p(s(bool,X6)) | ~sP4(X8,X7,X6)) )),
  inference(subsumption_resolution,[],[f2172,f2169])).
fof(f2172,plain,(
  ( ! [X6,X8,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,'T_0') != s(bool,X7) | sP5(X8,X6,X7) | p(s(bool,X6)) | ~sP4(X8,X7,X6) | ~sP3(X6,X7,X8)) )),
  inference(resolution,[],[f2168,f311])).
fof(f2178,plain,(
  ( ! [X4,X5,X3] : (~sP4(X5,X4,X3) | sP5(X5,X3,X4) | ~p(s(bool,X4)) | s(bool,'F_0') != s(bool,X3)) )),
  inference(subsumption_resolution,[],[f2177,f229])).
fof(f2177,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X3) | s(bool,'T_0') != s(bool,X4) | sP5(X5,X3,X4) | ~p(s(bool,X4)) | ~sP4(X5,X4,X3)) )),
  inference(subsumption_resolution,[],[f2171,f2169])).
fof(f2171,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X3) | s(bool,'T_0') != s(bool,X4) | sP5(X5,X3,X4) | ~p(s(bool,X4)) | ~sP4(X5,X4,X3) | ~sP3(X3,X4,X5)) )),
  inference(resolution,[],[f2168,f310])).
fof(f2169,plain,(
  ( ! [X10,X11,X9] : (sP3(X10,X9,X11) | s(bool,'F_0') != s(bool,X10) | s(bool,'T_0') != s(bool,X9)) )),
  inference(subsumption_resolution,[],[f2158,f318])).
fof(f2158,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') != s(bool,X9) | s(bool,'F_0') != s(bool,X10) | p(s(bool,X11)) | sP3(X10,X9,X11)) )),
  inference(resolution,[],[f1277,f307])).
fof(f2168,plain,(
  ( ! [X6,X8,X7] : (sP2(X6,X7,X8) | s(bool,'F_0') != s(bool,X7) | s(bool,'T_0') != s(bool,X6)) )),
  inference(subsumption_resolution,[],[f2157,f322])).
fof(f2157,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X6) | s(bool,'F_0') != s(bool,X7) | p(s(bool,X8)) | sP2(X6,X7,X8)) )),
  inference(resolution,[],[f1277,f306])).
fof(f2153,plain,(
  ( ! [X6,X4,X5] : (~sP4(X5,X6,'F_0') | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,'T_0') != s(bool,X4)) )),
  inference(subsumption_resolution,[],[f2142,f1946])).
fof(f2142,plain,(
  ( ! [X6,X4,X5] : (s(bool,'T_0') != s(bool,X4) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | ~sP4(X5,X6,'F_0') | ~sP2(X6,'F_0',X5)) )),
  inference(resolution,[],[f1263,f1191])).
fof(f2131,plain,(
  ( ! [X49] : (s(bool,X49) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11))) | s(bool,'F_0') = s(bool,X49)) )),
  inference(trivial_inequality_removal,[],[f2129])).
fof(f2129,plain,(
  ( ! [X49] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,X49) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11))) | s(bool,'F_0') = s(bool,X49)) )),
  inference(superposition,[],[f2040,f398])).
fof(f2125,plain,(
  ( ! [X39,X37,X38,X36] : (s(bool,'F_0') != s(bool,X37) | s(bool,X36) != s(bool,X38) | s(bool,'F_0') = s(bool,X38) | ~p(s(bool,X39)) | ~sP5(X39,X36,X37)) )),
  inference(superposition,[],[f2040,f327])).
fof(f2124,plain,(
  ( ! [X35,X33,X34,X32] : (s(bool,'F_0') != s(bool,X33) | s(bool,X32) != s(bool,X34) | s(bool,'F_0') = s(bool,X34) | ~p(s(bool,X35)) | ~sP5(X35,X33,X32)) )),
  inference(superposition,[],[f2040,f327])).
fof(f2123,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'F_0') != s(bool,X29) | s(bool,X28) != s(bool,X30) | s(bool,'F_0') = s(bool,X30) | sP5(X31,X28,X29) | p(s(bool,X31))) )),
  inference(superposition,[],[f2040,f325])).
fof(f2122,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,'F_0') != s(bool,X25) | s(bool,X24) != s(bool,X26) | s(bool,'F_0') = s(bool,X26) | sP5(X27,X25,X24) | p(s(bool,X27))) )),
  inference(superposition,[],[f2040,f325])).
fof(f2111,plain,(
  ( ! [X47,X45,X48,X46] : (s(bool,X45) != s(bool,X46) | s(bool,X46) != s(bool,X47) | s(bool,'F_0') = s(bool,X47) | ~p(s(bool,X48)) | ~sP5(X48,'F_0',X45)) )),
  inference(superposition,[],[f2040,f327])).
fof(f2110,plain,(
  ( ! [X43,X41,X44,X42] : (s(bool,X41) != s(bool,X42) | s(bool,X42) != s(bool,X43) | s(bool,'F_0') = s(bool,X43) | ~p(s(bool,X44)) | ~sP5(X44,X41,'F_0')) )),
  inference(superposition,[],[f2040,f327])).
fof(f2109,plain,(
  ( ! [X39,X37,X38,X40] : (s(bool,X37) != s(bool,X38) | s(bool,X38) != s(bool,X39) | s(bool,'F_0') = s(bool,X39) | sP5(X40,'F_0',X37) | p(s(bool,X40))) )),
  inference(superposition,[],[f2040,f325])).
fof(f2108,plain,(
  ( ! [X35,X33,X36,X34] : (s(bool,X33) != s(bool,X34) | s(bool,X34) != s(bool,X35) | s(bool,'F_0') = s(bool,X35) | sP5(X36,X33,'F_0') | p(s(bool,X36))) )),
  inference(superposition,[],[f2040,f325])).
fof(f2075,plain,(
  ( ! [X39,X37,X38,X36] : (s(bool,'F_0') != s(bool,X37) | s(bool,X36) != s(bool,X38) | ~p(s(bool,X38)) | ~p(s(bool,X39)) | ~sP5(X39,X36,X37)) )),
  inference(superposition,[],[f2039,f327])).
fof(f2074,plain,(
  ( ! [X35,X33,X34,X32] : (s(bool,'F_0') != s(bool,X33) | s(bool,X32) != s(bool,X34) | ~p(s(bool,X34)) | ~p(s(bool,X35)) | ~sP5(X35,X33,X32)) )),
  inference(superposition,[],[f2039,f327])).
fof(f2073,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'F_0') != s(bool,X29) | s(bool,X28) != s(bool,X30) | ~p(s(bool,X30)) | sP5(X31,X28,X29) | p(s(bool,X31))) )),
  inference(superposition,[],[f2039,f325])).
fof(f2072,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,'F_0') != s(bool,X25) | s(bool,X24) != s(bool,X26) | ~p(s(bool,X26)) | sP5(X27,X25,X24) | p(s(bool,X27))) )),
  inference(superposition,[],[f2039,f325])).
fof(f2061,plain,(
  ( ! [X47,X45,X48,X46] : (s(bool,X45) != s(bool,X46) | s(bool,X46) != s(bool,X47) | ~p(s(bool,X47)) | ~p(s(bool,X48)) | ~sP5(X48,'F_0',X45)) )),
  inference(superposition,[],[f2039,f327])).
fof(f2060,plain,(
  ( ! [X43,X41,X44,X42] : (s(bool,X41) != s(bool,X42) | s(bool,X42) != s(bool,X43) | ~p(s(bool,X43)) | ~p(s(bool,X44)) | ~sP5(X44,X41,'F_0')) )),
  inference(superposition,[],[f2039,f327])).
fof(f2059,plain,(
  ( ! [X39,X37,X38,X40] : (s(bool,X37) != s(bool,X38) | s(bool,X38) != s(bool,X39) | ~p(s(bool,X39)) | sP5(X40,'F_0',X37) | p(s(bool,X40))) )),
  inference(superposition,[],[f2039,f325])).
fof(f2058,plain,(
  ( ! [X35,X33,X36,X34] : (s(bool,X33) != s(bool,X34) | s(bool,X34) != s(bool,X35) | ~p(s(bool,X35)) | sP5(X36,X33,'F_0') | p(s(bool,X36))) )),
  inference(superposition,[],[f2039,f325])).
fof(f2023,plain,(
  ( ! [X39,X37,X38,X36] : (s(bool,'T_0') != s(bool,X37) | s(bool,X36) != s(bool,X38) | s(bool,'F_0') != s(bool,X38) | ~p(s(bool,X39)) | ~sP5(X39,X36,X37)) )),
  inference(superposition,[],[f1994,f327])).
fof(f2022,plain,(
  ( ! [X35,X33,X34,X32] : (s(bool,'T_0') != s(bool,X33) | s(bool,X32) != s(bool,X34) | s(bool,'F_0') != s(bool,X34) | ~p(s(bool,X35)) | ~sP5(X35,X33,X32)) )),
  inference(superposition,[],[f1994,f327])).
fof(f2021,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'T_0') != s(bool,X29) | s(bool,X28) != s(bool,X30) | s(bool,'F_0') != s(bool,X30) | sP5(X31,X28,X29) | p(s(bool,X31))) )),
  inference(superposition,[],[f1994,f325])).
fof(f2020,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,'T_0') != s(bool,X25) | s(bool,X24) != s(bool,X26) | s(bool,'F_0') != s(bool,X26) | sP5(X27,X25,X24) | p(s(bool,X27))) )),
  inference(superposition,[],[f1994,f325])).
fof(f2039,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,X5) | s(bool,X4) != s(bool,X5) | ~p(s(bool,X4))) )),
  inference(trivial_inequality_removal,[],[f2012])).
fof(f2012,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X4) != s(bool,X5) | s(bool,'F_0') != s(bool,X5) | ~p(s(bool,X4))) )),
  inference(superposition,[],[f1994,f229])).
fof(f2040,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X0)) )),
  inference(trivial_inequality_removal,[],[f2010])).
fof(f2010,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) != s(bool,X1) | s(bool,'F_0') != s(bool,X1) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f1994,f222])).
fof(f2009,plain,(
  ( ! [X39,X41,X42,X40] : (s(bool,X39) != s(bool,X40) | s(bool,X40) != s(bool,X41) | s(bool,'F_0') != s(bool,X41) | ~p(s(bool,X42)) | ~sP5(X42,'T_0',X39)) )),
  inference(superposition,[],[f1994,f327])).
fof(f2008,plain,(
  ( ! [X37,X35,X38,X36] : (s(bool,X35) != s(bool,X36) | s(bool,X36) != s(bool,X37) | s(bool,'F_0') != s(bool,X37) | ~p(s(bool,X38)) | ~sP5(X38,X35,'T_0')) )),
  inference(superposition,[],[f1994,f327])).
fof(f2007,plain,(
  ( ! [X33,X31,X34,X32] : (s(bool,X31) != s(bool,X32) | s(bool,X32) != s(bool,X33) | s(bool,'F_0') != s(bool,X33) | sP5(X34,'T_0',X31) | p(s(bool,X34))) )),
  inference(superposition,[],[f1994,f325])).
fof(f2006,plain,(
  ( ! [X30,X28,X29,X27] : (s(bool,X27) != s(bool,X28) | s(bool,X28) != s(bool,X29) | s(bool,'F_0') != s(bool,X29) | sP5(X30,X27,'T_0') | p(s(bool,X30))) )),
  inference(superposition,[],[f1994,f325])).
fof(f1998,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,X7) | s(bool,X7) != s(bool,X8) | s(bool,'F_0') != s(bool,X8) | s(bool,'F_0') = s(bool,X6)) )),
  inference(superposition,[],[f1994,f222])).
fof(f1996,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,X2) | s(bool,X1) != s(bool,X2) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X0))) )),
  inference(superposition,[],[f1994,f231])).
fof(f231,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f143])).
fof(f143,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X9] : (s(bool,'T_0') = s(bool,X9) <=> p(s(bool,X9)))),
  file('thm.util_prob.BIGINTER_SUBSET','thm.bool.EQ_CLAUSES|split|0')).
fof(f1994,plain,(
  ( ! [X12,X10] : (s(bool,'T_0') != s(bool,X10) | s(bool,X10) != s(bool,X12) | s(bool,'F_0') != s(bool,X12)) )),
  inference(subsumption_resolution,[],[f1993,f1922])).
fof(f1993,plain,(
  ( ! [X12,X10,X11] : (s(bool,'T_0') != s(bool,X10) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,'F_0') != s(bool,X12)) )),
  inference(duplicate_literal_removal,[],[f1985])).
fof(f1985,plain,(
  ( ! [X12,X10,X11] : (s(bool,'T_0') != s(bool,X10) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,'F_0') != s(bool,X12) | ~p(s(bool,X11))) )),
  inference(resolution,[],[f1262,f326])).
fof(f1976,plain,(
  ( ! [X4,X3] : (~sP1('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),X3)),X4) | p(s(bool,X4)) | s(bool,f(s(sK10,sK17(X3,sK11,bool,sK10)))) != s(bool,g(s(sK10,sK17(X3,sK11,bool,sK10))))) )),
  inference(superposition,[],[f409,f299])).
fof(f1975,plain,(
  ( ! [X2,X1] : (~sP1('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),X1)),X2) | p(s(bool,X2)) | s(bool,f(s(sK10,sK17(sK11,X1,bool,sK10)))) != s(bool,g(s(sK10,sK17(sK11,X1,bool,sK10))))) )),
  inference(superposition,[],[f409,f299])).
fof(f1974,plain,(
  ( ! [X0] : (~sP1('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),'const.pred_set.EMPTY_0')),X0) | p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),sK11))),s(fun(sK10,bool),sK11))))) )),
  inference(superposition,[],[f409,f255])).
fof(f1973,plain,(
  ( ! [X4,X3] : (~sP1('const.pred_set.SUBSET_2'(s(fun(sK10,bool),X3),s(fun(sK10,bool),sK11)),X4) | p(s(bool,X4)) | s(bool,f(s(sK10,sK17(X3,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10)))) != s(bool,g(s(sK10,sK17(X3,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10))))) )),
  inference(superposition,[],[f409,f299])).
fof(f1972,plain,(
  ( ! [X2,X1] : (~sP1('const.pred_set.SUBSET_2'(s(fun(sK10,bool),X1),s(fun(sK10,bool),sK11)),X2) | p(s(bool,X2)) | s(bool,f(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X1,bool,sK10)))) != s(bool,g(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X1,bool,sK10))))) )),
  inference(superposition,[],[f409,f299])).
fof(f1971,plain,(
  ( ! [X0] : (~sP1('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.EMPTY_0'),s(fun(sK10,bool),sK11)),X0) | p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))),s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))))) )),
  inference(superposition,[],[f409,f255])).
fof(f1970,plain,(
  ( ! [X4,X3] : (~sP1('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X3))),s(fun(sK10,bool),sK11)),X4) | p(s(bool,X4)) | s(bool,f(s(fun(sK10,bool),sK17(X3,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X3,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f409,f299])).
fof(f1969,plain,(
  ( ! [X2,X1] : (~sP1('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X1))),s(fun(sK10,bool),sK11)),X2) | p(s(bool,X2)) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X1,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X1,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f409,f299])).
fof(f1955,plain,(
  ( ! [X6,X8,X7] : (~sP4(X8,X7,'F_0') | s(bool,'T_0') != s(bool,X6) | sP5(X8,'F_0',X7) | s(bool,X6) != s(bool,X7)) )),
  inference(subsumption_resolution,[],[f1954,f1947])).
fof(f1954,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,X7) | s(bool,'T_0') != s(bool,X6) | sP5(X8,'F_0',X7) | ~sP4(X8,X7,'F_0') | ~sP3('F_0',X7,X8)) )),
  inference(subsumption_resolution,[],[f1950,f1247])).
fof(f1950,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,X7) | s(bool,'T_0') != s(bool,X6) | sP5(X8,'F_0',X7) | ~p(s(bool,X7)) | ~sP4(X8,X7,'F_0') | ~sP3('F_0',X7,X8)) )),
  inference(resolution,[],[f1946,f310])).
fof(f1947,plain,(
  ( ! [X17,X15,X16] : (sP3('F_0',X16,X17) | s(bool,X15) != s(bool,X16) | s(bool,'T_0') != s(bool,X15)) )),
  inference(subsumption_resolution,[],[f1941,f318])).
fof(f1941,plain,(
  ( ! [X17,X15,X16] : (s(bool,'T_0') != s(bool,X15) | s(bool,X15) != s(bool,X16) | p(s(bool,X17)) | sP3('F_0',X16,X17)) )),
  inference(resolution,[],[f1261,f307])).
fof(f1946,plain,(
  ( ! [X14,X12,X13] : (sP2(X13,'F_0',X14) | s(bool,X12) != s(bool,X13) | s(bool,'T_0') != s(bool,X12)) )),
  inference(subsumption_resolution,[],[f1940,f322])).
fof(f1940,plain,(
  ( ! [X14,X12,X13] : (s(bool,'T_0') != s(bool,X12) | s(bool,X12) != s(bool,X13) | p(s(bool,X14)) | sP2(X13,'F_0',X14)) )),
  inference(resolution,[],[f1261,f306])).
fof(f1924,plain,(
  ( ! [X14,X12,X13] : (sP3(X13,'F_0',X14) | s(bool,X12) != s(bool,X13) | s(bool,'T_0') != s(bool,X12)) )),
  inference(subsumption_resolution,[],[f1917,f318])).
fof(f1917,plain,(
  ( ! [X14,X12,X13] : (s(bool,'T_0') != s(bool,X12) | s(bool,X12) != s(bool,X13) | p(s(bool,X14)) | sP3(X13,'F_0',X14)) )),
  inference(resolution,[],[f1260,f307])).
fof(f1923,plain,(
  ( ! [X10,X11,X9] : (sP2('F_0',X10,X11) | s(bool,X9) != s(bool,X10) | s(bool,'T_0') != s(bool,X9)) )),
  inference(subsumption_resolution,[],[f1916,f322])).
fof(f1916,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') != s(bool,X9) | s(bool,X9) != s(bool,X10) | p(s(bool,X11)) | sP2('F_0',X10,X11)) )),
  inference(resolution,[],[f1260,f306])).
fof(f1922,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X6) | s(bool,X6) != s(bool,X7) | p(s(bool,X8)) | s(bool,'F_0') != s(bool,X7)) )),
  inference(duplicate_literal_removal,[],[f1915])).
fof(f1915,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X6) | s(bool,X6) != s(bool,X7) | p(s(bool,X8)) | p(s(bool,X8)) | s(bool,'F_0') != s(bool,X7)) )),
  inference(resolution,[],[f1260,f328])).
fof(f1905,plain,(
  ( ! [X4,X3] : (sP0('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),X3))),s(fun(fun(sK10,bool),bool),X3)),X4) | s(bool,f(s(fun(sK10,bool),sK17(X3,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X3,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f974,f299])).
fof(f1904,plain,(
  ( ! [X2,X1] : (sP0('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),X1))),s(fun(fun(sK10,bool),bool),X1)),X2) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X1,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X1,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f974,f299])).
fof(f1902,plain,(
  ( ! [X4,X3] : (sP0('const.bool.IN_2'(s(fun(sK10,bool),X3),s(fun(fun(sK10,bool),bool),sK12)),X4) | s(bool,f(s(sK10,sK17(X3,'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10)))) != s(bool,g(s(sK10,sK17(X3,'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10))))) )),
  inference(superposition,[],[f974,f299])).
fof(f1901,plain,(
  ( ! [X2,X1] : (sP0('const.bool.IN_2'(s(fun(sK10,bool),X1),s(fun(fun(sK10,bool),bool),sK12)),X2) | s(bool,f(s(sK10,sK17('const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12)),X1,bool,sK10)))) != s(bool,g(s(sK10,sK17('const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12)),X1,bool,sK10))))) )),
  inference(superposition,[],[f974,f299])).
fof(f1900,plain,(
  ( ! [X0] : (sP0('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK10,bool),bool),sK12)),X0) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))))),s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))))))) )),
  inference(superposition,[],[f974,f255])).
fof(f1890,plain,(
  ( ! [X4,X3] : (sP0(X4,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),X3))),s(fun(fun(sK10,bool),bool),X3))) | s(bool,f(s(fun(sK10,bool),sK17(X3,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X3,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f973,f299])).
fof(f1889,plain,(
  ( ! [X2,X1] : (sP0(X2,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),X1))),s(fun(fun(sK10,bool),bool),X1))) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X1,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X1,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f973,f299])).
fof(f1887,plain,(
  ( ! [X4,X3] : (sP0(X4,'const.bool.IN_2'(s(fun(sK10,bool),X3),s(fun(fun(sK10,bool),bool),sK12))) | s(bool,f(s(sK10,sK17(X3,'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10)))) != s(bool,g(s(sK10,sK17(X3,'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10))))) )),
  inference(superposition,[],[f973,f299])).
fof(f1886,plain,(
  ( ! [X2,X1] : (sP0(X2,'const.bool.IN_2'(s(fun(sK10,bool),X1),s(fun(fun(sK10,bool),bool),sK12))) | s(bool,f(s(sK10,sK17('const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12)),X1,bool,sK10)))) != s(bool,g(s(sK10,sK17('const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12)),X1,bool,sK10))))) )),
  inference(superposition,[],[f973,f299])).
fof(f1885,plain,(
  ( ! [X0] : (sP0(X0,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK10,bool),bool),sK12))) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))))),s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))))))) )),
  inference(superposition,[],[f973,f255])).
fof(f1861,plain,(
  ( ! [X39,X37,X38,X36] : (s(bool,X37) != s(bool,X38) | s(bool,'F_0') = s(bool,X36) | s(bool,'T_0') = s(bool,X38) | ~p(s(bool,X39)) | ~sP5(X39,X36,X37)) )),
  inference(superposition,[],[f1722,f327])).
fof(f1860,plain,(
  ( ! [X35,X33,X34,X32] : (s(bool,X33) != s(bool,X34) | s(bool,'F_0') = s(bool,X32) | s(bool,'T_0') = s(bool,X34) | ~p(s(bool,X35)) | ~sP5(X35,X33,X32)) )),
  inference(superposition,[],[f1722,f327])).
fof(f1859,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X29) != s(bool,X30) | s(bool,'F_0') = s(bool,X28) | s(bool,'T_0') = s(bool,X30) | sP5(X31,X28,X29) | p(s(bool,X31))) )),
  inference(superposition,[],[f1722,f325])).
fof(f1858,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,X25) != s(bool,X26) | s(bool,'F_0') = s(bool,X24) | s(bool,'T_0') = s(bool,X26) | sP5(X27,X25,X24) | p(s(bool,X27))) )),
  inference(superposition,[],[f1722,f325])).
fof(f1840,plain,(
  ( ! [X39,X37,X38,X36] : (s(bool,X37) != s(bool,X38) | s(bool,'F_0') = s(bool,X38) | s(bool,'T_0') = s(bool,X36) | ~p(s(bool,X39)) | ~sP5(X39,X36,X37)) )),
  inference(superposition,[],[f1722,f327])).
fof(f1839,plain,(
  ( ! [X35,X33,X34,X32] : (s(bool,X33) != s(bool,X34) | s(bool,'F_0') = s(bool,X34) | s(bool,'T_0') = s(bool,X32) | ~p(s(bool,X35)) | ~sP5(X35,X33,X32)) )),
  inference(superposition,[],[f1722,f327])).
fof(f1838,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X29) != s(bool,X30) | s(bool,'F_0') = s(bool,X30) | s(bool,'T_0') = s(bool,X28) | sP5(X31,X28,X29) | p(s(bool,X31))) )),
  inference(superposition,[],[f1722,f325])).
fof(f1837,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,X25) != s(bool,X26) | s(bool,'F_0') = s(bool,X26) | s(bool,'T_0') = s(bool,X24) | sP5(X27,X25,X24) | p(s(bool,X27))) )),
  inference(superposition,[],[f1722,f325])).
fof(f1804,plain,(
  ( ! [X39,X37,X38,X36] : (s(bool,X37) != s(bool,X38) | s(bool,'T_0') = s(bool,X36) | ~p(s(bool,X38)) | ~p(s(bool,X39)) | ~sP5(X39,X36,X37)) )),
  inference(superposition,[],[f1760,f327])).
fof(f1803,plain,(
  ( ! [X35,X33,X34,X32] : (s(bool,X33) != s(bool,X34) | s(bool,'T_0') = s(bool,X32) | ~p(s(bool,X34)) | ~p(s(bool,X35)) | ~sP5(X35,X33,X32)) )),
  inference(superposition,[],[f1760,f327])).
fof(f1802,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X29) != s(bool,X30) | s(bool,'T_0') = s(bool,X28) | ~p(s(bool,X30)) | sP5(X31,X28,X29) | p(s(bool,X31))) )),
  inference(superposition,[],[f1760,f325])).
fof(f1801,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,X25) != s(bool,X26) | s(bool,'T_0') = s(bool,X24) | ~p(s(bool,X26)) | sP5(X27,X25,X24) | p(s(bool,X27))) )),
  inference(superposition,[],[f1760,f325])).
fof(f1783,plain,(
  ( ! [X39,X37,X38,X36] : (s(bool,X37) != s(bool,X38) | s(bool,'T_0') = s(bool,X38) | ~p(s(bool,X36)) | ~p(s(bool,X39)) | ~sP5(X39,X36,X37)) )),
  inference(superposition,[],[f1760,f327])).
fof(f1782,plain,(
  ( ! [X35,X33,X34,X32] : (s(bool,X33) != s(bool,X34) | s(bool,'T_0') = s(bool,X34) | ~p(s(bool,X32)) | ~p(s(bool,X35)) | ~sP5(X35,X33,X32)) )),
  inference(superposition,[],[f1760,f327])).
fof(f1781,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X29) != s(bool,X30) | s(bool,'T_0') = s(bool,X30) | ~p(s(bool,X28)) | sP5(X31,X28,X29) | p(s(bool,X31))) )),
  inference(superposition,[],[f1760,f325])).
fof(f1780,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,X25) != s(bool,X26) | s(bool,'T_0') = s(bool,X26) | ~p(s(bool,X24)) | sP5(X27,X25,X24) | p(s(bool,X27))) )),
  inference(superposition,[],[f1760,f325])).
fof(f1750,plain,(
  ( ! [X39,X37,X38,X36] : (s(bool,'T_0') != s(bool,X37) | s(bool,X36) != s(bool,X38) | s(bool,'T_0') = s(bool,X38) | ~p(s(bool,X39)) | ~sP5(X39,X36,X37)) )),
  inference(superposition,[],[f1249,f327])).
fof(f1749,plain,(
  ( ! [X35,X33,X34,X32] : (s(bool,'T_0') != s(bool,X33) | s(bool,X32) != s(bool,X34) | s(bool,'T_0') = s(bool,X34) | ~p(s(bool,X35)) | ~sP5(X35,X33,X32)) )),
  inference(superposition,[],[f1249,f327])).
fof(f1748,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'T_0') != s(bool,X29) | s(bool,X28) != s(bool,X30) | s(bool,'T_0') = s(bool,X30) | sP5(X31,X28,X29) | p(s(bool,X31))) )),
  inference(superposition,[],[f1249,f325])).
fof(f1747,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,'T_0') != s(bool,X25) | s(bool,X24) != s(bool,X26) | s(bool,'T_0') = s(bool,X26) | sP5(X27,X25,X24) | p(s(bool,X27))) )),
  inference(superposition,[],[f1249,f325])).
fof(f1760,plain,(
  ( ! [X4,X5] : (s(bool,X4) != s(bool,X5) | s(bool,'T_0') = s(bool,X5) | ~p(s(bool,X4))) )),
  inference(trivial_inequality_removal,[],[f1739])).
fof(f1739,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X4) != s(bool,X5) | s(bool,'T_0') = s(bool,X5) | ~p(s(bool,X4))) )),
  inference(superposition,[],[f1249,f229])).
fof(f1736,plain,(
  ( ! [X39,X41,X42,X40] : (s(bool,X39) != s(bool,X40) | s(bool,X40) != s(bool,X41) | s(bool,'T_0') = s(bool,X41) | ~p(s(bool,X42)) | ~sP5(X42,'T_0',X39)) )),
  inference(superposition,[],[f1249,f327])).
fof(f1735,plain,(
  ( ! [X37,X35,X38,X36] : (s(bool,X35) != s(bool,X36) | s(bool,X36) != s(bool,X37) | s(bool,'T_0') = s(bool,X37) | ~p(s(bool,X38)) | ~sP5(X38,X35,'T_0')) )),
  inference(superposition,[],[f1249,f327])).
fof(f1734,plain,(
  ( ! [X33,X31,X34,X32] : (s(bool,X31) != s(bool,X32) | s(bool,X32) != s(bool,X33) | s(bool,'T_0') = s(bool,X33) | sP5(X34,'T_0',X31) | p(s(bool,X34))) )),
  inference(superposition,[],[f1249,f325])).
fof(f1733,plain,(
  ( ! [X30,X28,X29,X27] : (s(bool,X27) != s(bool,X28) | s(bool,X28) != s(bool,X29) | s(bool,'T_0') = s(bool,X29) | sP5(X30,X27,'T_0') | p(s(bool,X30))) )),
  inference(superposition,[],[f1249,f325])).
fof(f1725,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,X7) | s(bool,X7) != s(bool,X8) | s(bool,'T_0') = s(bool,X8) | s(bool,'F_0') = s(bool,X6)) )),
  inference(superposition,[],[f1249,f222])).
fof(f1723,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) != s(bool,X2) | s(bool,'T_0') = s(bool,X2) | ~p(s(bool,X0))) )),
  inference(superposition,[],[f1249,f231])).
fof(f1722,plain,(
  ( ! [X50,X51] : (s(bool,X50) != s(bool,X51) | s(bool,'F_0') = s(bool,X51) | s(bool,'T_0') = s(bool,X50)) )),
  inference(subsumption_resolution,[],[f1704,f221])).
fof(f1704,plain,(
  ( ! [X50,X51] : (s(bool,X50) != s(bool,X51) | p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,X51) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f1380,f222])).
fof(f1698,plain,(
  ( ! [X39,X37,X38,X36] : (s(bool,X37) != s(bool,X38) | p(s(bool,X36)) | s(bool,'F_0') = s(bool,X38) | ~p(s(bool,X39)) | ~sP5(X39,X36,X37)) )),
  inference(superposition,[],[f1380,f327])).
fof(f1697,plain,(
  ( ! [X35,X33,X34,X32] : (s(bool,X33) != s(bool,X34) | p(s(bool,X32)) | s(bool,'F_0') = s(bool,X34) | ~p(s(bool,X35)) | ~sP5(X35,X33,X32)) )),
  inference(superposition,[],[f1380,f327])).
fof(f1696,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X29) != s(bool,X30) | p(s(bool,X28)) | s(bool,'F_0') = s(bool,X30) | sP5(X31,X28,X29) | p(s(bool,X31))) )),
  inference(superposition,[],[f1380,f325])).
fof(f1695,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,X25) != s(bool,X26) | p(s(bool,X24)) | s(bool,'F_0') = s(bool,X26) | sP5(X27,X25,X24) | p(s(bool,X27))) )),
  inference(superposition,[],[f1380,f325])).
fof(f1677,plain,(
  ( ! [X39,X37,X38,X36] : (s(bool,X37) != s(bool,X38) | p(s(bool,X38)) | s(bool,'F_0') = s(bool,X36) | ~p(s(bool,X39)) | ~sP5(X39,X36,X37)) )),
  inference(superposition,[],[f1380,f327])).
fof(f1676,plain,(
  ( ! [X35,X33,X34,X32] : (s(bool,X33) != s(bool,X34) | p(s(bool,X34)) | s(bool,'F_0') = s(bool,X32) | ~p(s(bool,X35)) | ~sP5(X35,X33,X32)) )),
  inference(superposition,[],[f1380,f327])).
fof(f1675,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X29) != s(bool,X30) | p(s(bool,X30)) | s(bool,'F_0') = s(bool,X28) | sP5(X31,X28,X29) | p(s(bool,X31))) )),
  inference(superposition,[],[f1380,f325])).
fof(f1674,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,X25) != s(bool,X26) | p(s(bool,X26)) | s(bool,'F_0') = s(bool,X24) | sP5(X27,X25,X24) | p(s(bool,X27))) )),
  inference(superposition,[],[f1380,f325])).
fof(f1662,plain,(
  ( ! [X269] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),sK14(X269,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)),bool))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(bool,bool),bool),X269))))))) )),
  inference(forward_demodulation,[],[f1593,f398])).
fof(f1593,plain,(
  ( ! [X269] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(bool,bool),bool),X269)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11))),s(fun(bool,bool),sK14(X269,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)),bool)))))) )),
  inference(superposition,[],[f263,f398])).
fof(f1592,plain,(
  ( ! [X268] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(bool,bool),bool),X268)))))) | p(s(bool,'const.bool.IN_2'(s(fun(bool,bool),sK14(X268,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)),bool)),s(fun(fun(bool,bool),bool),X268))))) )),
  inference(superposition,[],[f262,f398])).
fof(f1498,plain,(
  ( ! [X107,X108] : (s(bool,'F_0') != s(bool,X107) | p(s(bool,X108)) | ~sP5(X108,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)),X107)) )),
  inference(superposition,[],[f328,f398])).
fof(f1497,plain,(
  ( ! [X105,X106] : (s(bool,'F_0') != s(bool,X105) | p(s(bool,X106)) | ~sP5(X106,X105,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)))) )),
  inference(superposition,[],[f328,f398])).
fof(f1492,plain,(
  ( ! [X95,X96] : (s(bool,'F_0') != s(bool,X95) | sP5(X96,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)),X95) | ~p(s(bool,X96))) )),
  inference(superposition,[],[f326,f398])).
fof(f1491,plain,(
  ( ! [X94,X93] : (s(bool,'F_0') != s(bool,X93) | sP5(X94,X93,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11))) | ~p(s(bool,X94))) )),
  inference(superposition,[],[f326,f398])).
fof(f1615,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12))))),
  inference(subsumption_resolution,[],[f1408,f221])).
fof(f1408,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12))))),
  inference(superposition,[],[f212,f398])).
fof(f1406,plain,(
  ( ! [X10,X11] : (s(bool,'F_0') = s(bool,X10) | ~p(s(bool,X11)) | ~sP5(X11,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)),X10)) )),
  inference(superposition,[],[f398,f327])).
fof(f1405,plain,(
  ( ! [X8,X9] : (s(bool,'F_0') = s(bool,X8) | ~p(s(bool,X9)) | ~sP5(X9,X8,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)))) )),
  inference(superposition,[],[f398,f327])).
fof(f1404,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') = s(bool,X6) | sP5(X7,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)),X6) | p(s(bool,X7))) )),
  inference(superposition,[],[f398,f325])).
fof(f1403,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') = s(bool,X4) | sP5(X5,X4,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11))) | p(s(bool,X5))) )),
  inference(superposition,[],[f398,f325])).
fof(f1395,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),X1))) | s(bool,f(s(sK10,sK17(X1,sK11,bool,sK10)))) != s(bool,g(s(sK10,sK17(X1,sK11,bool,sK10))))) )),
  inference(superposition,[],[f398,f299])).
fof(f1394,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),X0))) | s(bool,f(s(sK10,sK17(sK11,X0,bool,sK10)))) != s(bool,g(s(sK10,sK17(sK11,X0,bool,sK10))))) )),
  inference(superposition,[],[f398,f299])).
fof(f1393,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),sK11))),s(fun(sK10,bool),sK11))))),
  inference(superposition,[],[f398,f255])).
fof(f1392,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),X1),s(fun(sK10,bool),sK11))) | s(bool,f(s(sK10,sK17(X1,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10)))) != s(bool,g(s(sK10,sK17(X1,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10))))) )),
  inference(superposition,[],[f398,f299])).
fof(f1391,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),X0),s(fun(sK10,bool),sK11))) | s(bool,f(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X0,bool,sK10)))) != s(bool,g(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X0,bool,sK10))))) )),
  inference(superposition,[],[f398,f299])).
fof(f1390,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.EMPTY_0'),s(fun(sK10,bool),sK11))) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))),s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))))),
  inference(superposition,[],[f398,f255])).
fof(f1389,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X1))),s(fun(sK10,bool),sK11))) | s(bool,f(s(fun(sK10,bool),sK17(X1,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X1,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f398,f299])).
fof(f1388,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X0))),s(fun(sK10,bool),sK11))) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X0,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X0,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f398,f299])).
fof(f1370,plain,(
  ( ! [X45,X43,X44,X42] : (s(bool,'T_0') != s(bool,X43) | s(bool,X42) != s(bool,X44) | p(s(bool,X44)) | ~p(s(bool,X45)) | ~sP5(X45,X42,X43)) )),
  inference(superposition,[],[f1247,f327])).
fof(f1369,plain,(
  ( ! [X39,X41,X38,X40] : (s(bool,'T_0') != s(bool,X39) | s(bool,X38) != s(bool,X40) | p(s(bool,X40)) | ~p(s(bool,X41)) | ~sP5(X41,X39,X38)) )),
  inference(superposition,[],[f1247,f327])).
fof(f1368,plain,(
  ( ! [X37,X35,X36,X34] : (s(bool,'T_0') != s(bool,X35) | s(bool,X34) != s(bool,X36) | p(s(bool,X36)) | sP5(X37,X34,X35) | p(s(bool,X37))) )),
  inference(superposition,[],[f1247,f325])).
fof(f1367,plain,(
  ( ! [X30,X33,X31,X32] : (s(bool,'T_0') != s(bool,X31) | s(bool,X30) != s(bool,X32) | p(s(bool,X32)) | sP5(X33,X31,X30) | p(s(bool,X33))) )),
  inference(superposition,[],[f1247,f325])).
fof(f1380,plain,(
  ( ! [X6,X7] : (s(bool,X6) != s(bool,X7) | p(s(bool,X7)) | s(bool,'F_0') = s(bool,X6)) )),
  inference(trivial_inequality_removal,[],[f1357])).
fof(f1357,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X6) != s(bool,X7) | p(s(bool,X7)) | s(bool,'F_0') = s(bool,X6)) )),
  inference(superposition,[],[f1247,f222])).
fof(f1353,plain,(
  ( ! [X39,X41,X42,X40] : (s(bool,X39) != s(bool,X40) | s(bool,X40) != s(bool,X41) | p(s(bool,X41)) | ~p(s(bool,X42)) | ~sP5(X42,'T_0',X39)) )),
  inference(superposition,[],[f1247,f327])).
fof(f1352,plain,(
  ( ! [X37,X35,X38,X36] : (s(bool,X35) != s(bool,X36) | s(bool,X36) != s(bool,X37) | p(s(bool,X37)) | ~p(s(bool,X38)) | ~sP5(X38,X35,'T_0')) )),
  inference(superposition,[],[f1247,f327])).
fof(f1351,plain,(
  ( ! [X33,X31,X34,X32] : (s(bool,X31) != s(bool,X32) | s(bool,X32) != s(bool,X33) | p(s(bool,X33)) | sP5(X34,'T_0',X31) | p(s(bool,X34))) )),
  inference(superposition,[],[f1247,f325])).
fof(f1350,plain,(
  ( ! [X30,X28,X29,X27] : (s(bool,X27) != s(bool,X28) | s(bool,X28) != s(bool,X29) | p(s(bool,X29)) | sP5(X30,X27,'T_0') | p(s(bool,X30))) )),
  inference(superposition,[],[f1247,f325])).
fof(f1342,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,X7) | s(bool,X7) != s(bool,X8) | p(s(bool,X8)) | s(bool,'F_0') = s(bool,X6)) )),
  inference(superposition,[],[f1247,f222])).
fof(f1340,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) != s(bool,X2) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(superposition,[],[f1247,f231])).
fof(f1310,plain,(
  ( ! [X11] : (sP5('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)),'F_0',X11) | s(bool,'T_0') != s(bool,X11)) )),
  inference(resolution,[],[f1137,f670])).
fof(f1309,plain,(
  ( ! [X10] : (sP5('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)),'F_0',X10) | s(bool,'T_0') != s(bool,X10)) )),
  inference(resolution,[],[f1137,f422])).
fof(f1301,plain,(
  ( ! [X0,X1] : (~sP4(X0,X1,'F_0') | s(bool,'T_0') != s(bool,X1) | ~p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f1289,f1127])).
fof(f1289,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s(bool,'T_0') != s(bool,X1) | ~sP4(X0,X1,'F_0') | ~sP2(X1,'F_0',X0)) )),
  inference(resolution,[],[f1095,f1191])).
fof(f1280,plain,(
  ( ! [X28,X26,X27] : (~sP5(X28,X26,X27) | s(bool,'T_0') != s(bool,X26) | ~p(s(bool,X28)) | s(bool,'F_0') != s(bool,X27)) )),
  inference(superposition,[],[f1244,f327])).
fof(f1279,plain,(
  ( ! [X24,X23,X25] : (~sP5(X25,X24,X23) | s(bool,'T_0') != s(bool,X23) | ~p(s(bool,X25)) | s(bool,'F_0') != s(bool,X24)) )),
  inference(superposition,[],[f1244,f327])).
fof(f1278,plain,(
  ( ! [X21,X22,X20] : (sP5(X22,X20,X21) | s(bool,'T_0') != s(bool,X20) | s(bool,'F_0') != s(bool,X21) | p(s(bool,X22))) )),
  inference(superposition,[],[f1244,f325])).
fof(f1277,plain,(
  ( ! [X19,X17,X18] : (sP5(X19,X18,X17) | s(bool,'T_0') != s(bool,X17) | s(bool,'F_0') != s(bool,X18) | p(s(bool,X19))) )),
  inference(superposition,[],[f1244,f325])).
fof(f1263,plain,(
  ( ! [X30,X31,X29] : (~sP5(X31,'F_0',X29) | s(bool,'T_0') != s(bool,X30) | ~p(s(bool,X31)) | s(bool,X29) != s(bool,X30)) )),
  inference(superposition,[],[f1244,f327])).
fof(f1262,plain,(
  ( ! [X28,X26,X27] : (~sP5(X28,X26,'F_0') | s(bool,'T_0') != s(bool,X27) | ~p(s(bool,X28)) | s(bool,X26) != s(bool,X27)) )),
  inference(superposition,[],[f1244,f327])).
fof(f1261,plain,(
  ( ! [X24,X23,X25] : (sP5(X25,'F_0',X23) | s(bool,'T_0') != s(bool,X24) | s(bool,X23) != s(bool,X24) | p(s(bool,X25))) )),
  inference(superposition,[],[f1244,f325])).
fof(f1260,plain,(
  ( ! [X21,X22,X20] : (sP5(X22,X20,'F_0') | s(bool,'T_0') != s(bool,X21) | s(bool,X20) != s(bool,X21) | p(s(bool,X22))) )),
  inference(superposition,[],[f1244,f325])).
fof(f1249,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X5) | s(bool,X4) != s(bool,X5) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f1244,f222])).
fof(f1247,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,X0) != s(bool,X1) | p(s(bool,X0))) )),
  inference(superposition,[],[f1244,f235])).
fof(f1244,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,X5) | s(bool,'T_0') != s(bool,X5)) )),
  inference(subsumption_resolution,[],[f1242,f1104])).
fof(f1242,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | s(bool,'T_0') != s(bool,X5) | s(bool,'F_0') != s(bool,X5)) )),
  inference(duplicate_literal_removal,[],[f1234])).
fof(f1234,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | s(bool,'T_0') != s(bool,X5) | s(bool,'F_0') != s(bool,X5) | ~p(s(bool,X4))) )),
  inference(resolution,[],[f1094,f326])).
fof(f1217,plain,(
  ( ! [X2,X3] : (~sP2(X3,'F_0',X2) | p(s(bool,X2)) | s(bool,'F_0') != s(bool,X3)) )),
  inference(subsumption_resolution,[],[f1211,f316])).
fof(f1211,plain,(
  ( ! [X2,X3] : (~sP4(X2,X3,'F_0') | ~sP2(X3,'F_0',X2) | p(s(bool,X2)) | s(bool,'F_0') != s(bool,X3)) )),
  inference(resolution,[],[f1191,f328])).
fof(f1216,plain,(
  ( ! [X0,X1] : (~sP4(X0,X1,'F_0') | s(bool,'F_0') = s(bool,X1) | ~p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f1210,f322])).
fof(f1210,plain,(
  ( ! [X0,X1] : (~sP4(X0,X1,'F_0') | ~sP2(X1,'F_0',X0) | s(bool,'F_0') = s(bool,X1) | ~p(s(bool,X0))) )),
  inference(resolution,[],[f1191,f327])).
fof(f1195,plain,(
  ( ! [X6,X7] : (~sP5(X7,'F_0',X6) | ~p(s(bool,X7)) | ~p(s(bool,X6))) )),
  inference(resolution,[],[f1148,f308])).
fof(f308,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X2,X1) | ~sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f184])).
fof(f1191,plain,(
  ( ! [X4,X5] : (sP5(X4,'F_0',X5) | ~sP4(X4,X5,'F_0') | ~sP2(X5,'F_0',X4)) )),
  inference(subsumption_resolution,[],[f1189,f221])).
fof(f1189,plain,(
  ( ! [X4,X5] : (~sP4(X4,X5,'F_0') | sP5(X4,'F_0',X5) | p(s(bool,'F_0')) | ~sP2(X5,'F_0',X4)) )),
  inference(duplicate_literal_removal,[],[f1186])).
fof(f1186,plain,(
  ( ! [X4,X5] : (~sP4(X4,X5,'F_0') | sP5(X4,'F_0',X5) | p(s(bool,'F_0')) | ~sP4(X4,X5,'F_0') | ~sP2(X5,'F_0',X4)) )),
  inference(resolution,[],[f1150,f311])).
fof(f1173,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') != s(bool,X10) | ~p(s(bool,X11)) | ~sP5(X11,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)),X10)) )),
  inference(superposition,[],[f396,f327])).
fof(f1172,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') != s(bool,X8) | ~p(s(bool,X9)) | ~sP5(X9,X8,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)))) )),
  inference(superposition,[],[f396,f327])).
fof(f1171,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X6) | sP5(X7,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)),X6) | p(s(bool,X7))) )),
  inference(superposition,[],[f396,f325])).
fof(f1170,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | sP5(X5,X4,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11))) | p(s(bool,X5))) )),
  inference(superposition,[],[f396,f325])).
fof(f1159,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),X1))) | s(bool,f(s(sK10,sK17(X1,sK11,bool,sK10)))) != s(bool,g(s(sK10,sK17(X1,sK11,bool,sK10))))) )),
  inference(superposition,[],[f396,f299])).
fof(f1158,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),X0))) | s(bool,f(s(sK10,sK17(sK11,X0,bool,sK10)))) != s(bool,g(s(sK10,sK17(sK11,X0,bool,sK10))))) )),
  inference(superposition,[],[f396,f299])).
fof(f1157,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),sK11))),s(fun(sK10,bool),sK11))))),
  inference(superposition,[],[f396,f255])).
fof(f1156,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),X1),s(fun(sK10,bool),sK11))) | s(bool,f(s(sK10,sK17(X1,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10)))) != s(bool,g(s(sK10,sK17(X1,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10))))) )),
  inference(superposition,[],[f396,f299])).
fof(f1155,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),X0),s(fun(sK10,bool),sK11))) | s(bool,f(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X0,bool,sK10)))) != s(bool,g(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X0,bool,sK10))))) )),
  inference(superposition,[],[f396,f299])).
fof(f1154,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.EMPTY_0'),s(fun(sK10,bool),sK11))) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))),s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))))),
  inference(superposition,[],[f396,f255])).
fof(f1153,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X1))),s(fun(sK10,bool),sK11))) | s(bool,f(s(fun(sK10,bool),sK17(X1,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X1,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f396,f299])).
fof(f1152,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X0))),s(fun(sK10,bool),sK11))) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X0,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X0,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f396,f299])).
fof(f1150,plain,(
  ( ! [X6,X7] : (sP3('F_0',X6,X7) | ~sP4(X7,X6,'F_0')) )),
  inference(subsumption_resolution,[],[f1145,f319])).
fof(f1145,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | ~sP4(X7,X6,'F_0') | sP3('F_0',X6,X7)) )),
  inference(resolution,[],[f1135,f307])).
fof(f1148,plain,(
  ( ! [X0,X1] : (~sP4(X1,X0,'F_0') | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(subsumption_resolution,[],[f1142,f232])).
fof(f232,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f144])).
fof(f1142,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~sP4(X1,X0,'F_0') | s(bool,'F_0') = s(bool,X0) | ~p(s(bool,X1))) )),
  inference(resolution,[],[f1135,f327])).
fof(f1137,plain,(
  ( ! [X4,X5] : (~sP4(X5,X4,'F_0') | sP5(X5,'F_0',X4) | s(bool,'T_0') != s(bool,X4)) )),
  inference(subsumption_resolution,[],[f1136,f221])).
fof(f1136,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | sP5(X5,'F_0',X4) | p(s(bool,'F_0')) | ~sP4(X5,X4,'F_0')) )),
  inference(subsumption_resolution,[],[f1131,f320])).
fof(f1131,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | sP5(X5,'F_0',X4) | p(s(bool,'F_0')) | ~sP4(X5,X4,'F_0') | ~sP3('F_0',X4,X5)) )),
  inference(resolution,[],[f1127,f311])).
fof(f1135,plain,(
  ( ! [X2,X3] : (sP5(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP4(X3,X2,'F_0')) )),
  inference(subsumption_resolution,[],[f1134,f229])).
fof(f1134,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | sP5(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP4(X3,X2,'F_0')) )),
  inference(subsumption_resolution,[],[f1130,f1128])).
fof(f1130,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | sP5(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP4(X3,X2,'F_0') | ~sP3('F_0',X2,X3)) )),
  inference(resolution,[],[f1127,f310])).
fof(f1128,plain,(
  ( ! [X6,X7] : (sP3('F_0',X6,X7) | s(bool,'T_0') != s(bool,X6)) )),
  inference(subsumption_resolution,[],[f1123,f318])).
fof(f1123,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X6) | p(s(bool,X7)) | sP3('F_0',X6,X7)) )),
  inference(resolution,[],[f1093,f307])).
fof(f1127,plain,(
  ( ! [X4,X5] : (sP2(X4,'F_0',X5) | s(bool,'T_0') != s(bool,X4)) )),
  inference(subsumption_resolution,[],[f1122,f322])).
fof(f1122,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | p(s(bool,X5)) | sP2(X4,'F_0',X5)) )),
  inference(resolution,[],[f1093,f306])).
fof(f1106,plain,(
  ( ! [X6,X7] : (sP3(X6,'F_0',X7) | s(bool,'T_0') != s(bool,X6)) )),
  inference(subsumption_resolution,[],[f1101,f318])).
fof(f1101,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X6) | p(s(bool,X7)) | sP3(X6,'F_0',X7)) )),
  inference(resolution,[],[f1092,f307])).
fof(f1105,plain,(
  ( ! [X4,X5] : (sP2('F_0',X4,X5) | s(bool,'T_0') != s(bool,X4)) )),
  inference(subsumption_resolution,[],[f1100,f322])).
fof(f1100,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | p(s(bool,X5)) | sP2('F_0',X4,X5)) )),
  inference(resolution,[],[f1092,f306])).
fof(f1104,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) )),
  inference(duplicate_literal_removal,[],[f1099])).
fof(f1099,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) )),
  inference(resolution,[],[f1092,f328])).
fof(f1095,plain,(
  ( ! [X14,X13] : (~sP5(X14,'F_0',X13) | ~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X13)) )),
  inference(superposition,[],[f1061,f327])).
fof(f1094,plain,(
  ( ! [X12,X11] : (~sP5(X12,X11,'F_0') | ~p(s(bool,X12)) | s(bool,'T_0') != s(bool,X11)) )),
  inference(superposition,[],[f1061,f327])).
fof(f1093,plain,(
  ( ! [X10,X9] : (sP5(X10,'F_0',X9) | s(bool,'T_0') != s(bool,X9) | p(s(bool,X10))) )),
  inference(superposition,[],[f1061,f325])).
fof(f1092,plain,(
  ( ! [X8,X7] : (sP5(X8,X7,'F_0') | s(bool,'T_0') != s(bool,X7) | p(s(bool,X8))) )),
  inference(superposition,[],[f1061,f325])).
fof(f1054,plain,(
  ( ! [X8,X9] : (p(s(bool,X8)) | ~p(s(bool,X9)) | ~sP5(X9,X8,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)))) )),
  inference(superposition,[],[f393,f327])).
fof(f1053,plain,(
  ( ! [X6,X7] : (p(s(bool,X6)) | sP5(X7,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X6) | p(s(bool,X7))) )),
  inference(superposition,[],[f393,f325])).
fof(f1052,plain,(
  ( ! [X4,X5] : (p(s(bool,X4)) | sP5(X5,X4,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12))) | p(s(bool,X5))) )),
  inference(superposition,[],[f393,f325])).
fof(f1041,plain,(
  ( ! [X1] : (p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),X1))),s(fun(fun(sK10,bool),bool),X1)))) | s(bool,f(s(fun(sK10,bool),sK17(X1,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X1,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f393,f299])).
fof(f1040,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),X0))),s(fun(fun(sK10,bool),bool),X0)))) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X0,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X0,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f393,f299])).
fof(f1038,plain,(
  ( ! [X1] : (p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X1),s(fun(fun(sK10,bool),bool),sK12)))) | s(bool,f(s(sK10,sK17(X1,'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10)))) != s(bool,g(s(sK10,sK17(X1,'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10))))) )),
  inference(superposition,[],[f393,f299])).
fof(f1037,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X0),s(fun(fun(sK10,bool),bool),sK12)))) | s(bool,f(s(sK10,sK17('const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12)),X0,bool,sK10)))) != s(bool,g(s(sK10,sK17('const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12)),X0,bool,sK10))))) )),
  inference(superposition,[],[f393,f299])).
fof(f1036,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK10,bool),bool),sK12)))) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))))),s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))))))),
  inference(superposition,[],[f393,f255])).
fof(f1029,plain,(
  ( ! [X105,X104] : (sP9('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X104,X105) | ~sP8('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X104,X105)) )),
  inference(resolution,[],[f393,f377])).
fof(f377,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(duplicate_literal_removal,[],[f349])).
fof(f349,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f204])).
fof(f1028,plain,(
  ( ! [X103,X102] : (sP7('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X102,X103) | ~sP6('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X102,X103)) )),
  inference(resolution,[],[f393,f376])).
fof(f376,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~sP6(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(duplicate_literal_removal,[],[f332])).
fof(f332,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~sP6(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f196])).
fof(f1026,plain,(
  ( ! [X99,X98] : (p(s(bool,X98)) | ~p(s(bool,X99)) | ~sP9('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X98,X99)) )),
  inference(resolution,[],[f393,f360])).
fof(f1025,plain,(
  ( ! [X97,X96] : (p(s(bool,X96)) | ~p(s(bool,X97)) | ~sP9(X97,X96,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)))) )),
  inference(resolution,[],[f393,f360])).
fof(f1024,plain,(
  ( ! [X94,X95] : (sP9('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X94,X95) | ~p(s(bool,X94))) )),
  inference(resolution,[],[f393,f359])).
fof(f359,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f209])).
fof(f1023,plain,(
  ( ! [X92,X93] : (sP9(X92,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X93) | ~p(s(bool,X92))) )),
  inference(resolution,[],[f393,f359])).
fof(f1022,plain,(
  ( ! [X90,X91] : (sP9('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X90,X91) | p(s(bool,X91))) )),
  inference(resolution,[],[f393,f358])).
fof(f1021,plain,(
  ( ! [X88,X89] : (sP9(X88,X89,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12))) | p(s(bool,X89)) | p(s(bool,X88))) )),
  inference(resolution,[],[f393,f357])).
fof(f1020,plain,(
  ( ! [X87,X86] : (sP8(X86,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X87)) )),
  inference(resolution,[],[f393,f355])).
fof(f355,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f207])).
fof(f1019,plain,(
  ( ! [X85,X84] : (~p(s(bool,X84)) | p(s(bool,X85)) | ~sP8(X84,X85,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)))) )),
  inference(resolution,[],[f393,f353])).
fof(f1018,plain,(
  ( ! [X83,X82] : (p(s(bool,X82)) | ~p(s(bool,X83)) | ~sP8('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X82,X83)) )),
  inference(resolution,[],[f393,f353])).
fof(f1017,plain,(
  ( ! [X80,X81] : (sP9(X80,X81,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12))) | ~sP8(X80,X81,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12))) | p(s(bool,X81))) )),
  inference(resolution,[],[f393,f352])).
fof(f1016,plain,(
  ( ! [X78,X79] : (sP9('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X78,X79) | ~sP8('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X78,X79) | p(s(bool,X78))) )),
  inference(resolution,[],[f393,f351])).
fof(f351,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f204])).
fof(f1014,plain,(
  ( ! [X74,X75] : (sP9('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X74,X75) | ~sP8('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X74,X75) | ~p(s(bool,X75))) )),
  inference(resolution,[],[f393,f350])).
fof(f350,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f204])).
fof(f1013,plain,(
  ( ! [X72,X73] : (p(s(bool,X72)) | ~sP9(X72,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X73)) )),
  inference(resolution,[],[f393,f347])).
fof(f1010,plain,(
  ( ! [X66,X67] : (p(s(bool,X66)) | p(s(bool,X67)) | ~sP7('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X66,X67)) )),
  inference(resolution,[],[f393,f343])).
fof(f1009,plain,(
  ( ! [X64,X65] : (sP7('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X64,X65) | ~p(s(bool,X64))) )),
  inference(resolution,[],[f393,f342])).
fof(f342,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f201])).
fof(f1008,plain,(
  ( ! [X62,X63] : (sP7(X62,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X63) | ~p(s(bool,X62))) )),
  inference(resolution,[],[f393,f342])).
fof(f1007,plain,(
  ( ! [X61,X60] : (sP7('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X60,X61) | ~p(s(bool,X61))) )),
  inference(resolution,[],[f393,f341])).
fof(f341,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f201])).
fof(f1006,plain,(
  ( ! [X59,X58] : (sP7(X58,X59,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12))) | ~p(s(bool,X58))) )),
  inference(resolution,[],[f393,f341])).
fof(f1005,plain,(
  ( ! [X57,X56] : (sP6(X56,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X57)) )),
  inference(resolution,[],[f393,f338])).
fof(f338,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f199])).
fof(f1004,plain,(
  ( ! [X54,X55] : (sP6(X54,X55,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)))) )),
  inference(resolution,[],[f393,f337])).
fof(f337,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f199])).
fof(f1003,plain,(
  ( ! [X52,X53] : (p(s(bool,X52)) | p(s(bool,X53)) | ~sP6('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X52,X53)) )),
  inference(resolution,[],[f393,f336])).
fof(f1002,plain,(
  ( ! [X50,X51] : (sP7('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X50,X51) | ~sP6('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X50,X51) | p(s(bool,X50))) )),
  inference(resolution,[],[f393,f334])).
fof(f334,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~sP6(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f196])).
fof(f1001,plain,(
  ( ! [X48,X49] : (sP7('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X48,X49) | ~sP6('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X48,X49) | p(s(bool,X49))) )),
  inference(resolution,[],[f393,f333])).
fof(f333,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~sP6(X0,X1,X2) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f196])).
fof(f1000,plain,(
  ( ! [X47,X46] : (p(s(bool,X46)) | ~sP7(X46,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X47)) )),
  inference(resolution,[],[f393,f330])).
fof(f999,plain,(
  ( ! [X45,X44] : (p(s(bool,X44)) | ~sP7(X44,X45,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)))) )),
  inference(resolution,[],[f393,f329])).
fof(f998,plain,(
  ( ! [X43,X42] : (s(bool,X42) = s(bool,X43) | ~sP5('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X43,X42)) )),
  inference(resolution,[],[f393,f327])).
fof(f997,plain,(
  ( ! [X41,X40] : (sP5('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X40,X41) | s(bool,X40) != s(bool,X41)) )),
  inference(resolution,[],[f393,f326])).
fof(f1065,plain,(
  ( ! [X21,X20] : (sP5(X20,X21,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12))) | ~sP4(X20,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X21) | ~sP3(X21,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X20)) )),
  inference(subsumption_resolution,[],[f987,f996])).
fof(f996,plain,(
  ( ! [X39,X38] : (sP2('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X38,X39)) )),
  inference(resolution,[],[f393,f324])).
fof(f995,plain,(
  ( ! [X37,X36] : (sP2(X36,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X37)) )),
  inference(resolution,[],[f393,f323])).
fof(f994,plain,(
  ( ! [X35,X34] : (sP2(X34,X35,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)))) )),
  inference(resolution,[],[f393,f322])).
fof(f993,plain,(
  ( ! [X33,X32] : (sP3(X32,X33,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)))) )),
  inference(resolution,[],[f393,f318])).
fof(f992,plain,(
  ( ! [X30,X31] : (~p(s(bool,X30)) | p(s(bool,X31)) | ~sP3(X30,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X31)) )),
  inference(resolution,[],[f393,f317])).
fof(f991,plain,(
  ( ! [X28,X29] : (~p(s(bool,X28)) | p(s(bool,X29)) | ~sP3('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X28,X29)) )),
  inference(resolution,[],[f393,f317])).
fof(f990,plain,(
  ( ! [X26,X27] : (sP4(X26,X27,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)))) )),
  inference(resolution,[],[f393,f314])).
fof(f989,plain,(
  ( ! [X24,X25] : (~p(s(bool,X24)) | p(s(bool,X25)) | ~sP4(X24,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X25)) )),
  inference(resolution,[],[f393,f313])).
fof(f988,plain,(
  ( ! [X23,X22] : (~p(s(bool,X22)) | p(s(bool,X23)) | ~sP4('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X22,X23)) )),
  inference(resolution,[],[f393,f313])).
fof(f987,plain,(
  ( ! [X21,X20] : (sP5(X20,X21,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12))) | ~sP4(X20,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X21) | ~sP3(X21,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X20) | ~sP2('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X21,X20)) )),
  inference(resolution,[],[f393,f310])).
fof(f986,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | p(s(bool,X19)) | ~sP5(X18,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X19)) )),
  inference(resolution,[],[f393,f309])).
fof(f985,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | p(s(bool,X17)) | ~sP5('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X16,X17)) )),
  inference(resolution,[],[f393,f309])).
fof(f980,plain,(
  ( ! [X11] : (sP1(X11,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12))) | p(s(bool,X11))) )),
  inference(resolution,[],[f393,f294])).
fof(f979,plain,(
  ( ! [X10] : (sP1('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X10) | p(s(bool,X10))) )),
  inference(resolution,[],[f393,f291])).
fof(f978,plain,(
  ( ! [X9] : (~sP1(X9,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12))) | ~p(s(bool,X9))) )),
  inference(resolution,[],[f393,f290])).
fof(f290,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | ~sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f177])).
fof(f977,plain,(
  ( ! [X8] : (~sP1('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X8) | ~p(s(bool,X8))) )),
  inference(resolution,[],[f393,f290])).
fof(f974,plain,(
  ( ! [X5] : (sP0('const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)),X5)) )),
  inference(resolution,[],[f393,f276])).
fof(f276,plain,(
  ( ! [X0,X1] : (sP0(X0,X1) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f170])).
fof(f973,plain,(
  ( ! [X4] : (sP0(X4,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)))) )),
  inference(resolution,[],[f393,f275])).
fof(f275,plain,(
  ( ! [X0,X1] : (sP0(X0,X1) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f170])).
fof(f1061,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(forward_demodulation,[],[f967,f965])).
fof(f967,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)))),
  inference(resolution,[],[f393,f232])).
fof(f965,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12)))),
  inference(resolution,[],[f393,f229])).
fof(f964,plain,(
  ( ! [X1] : (p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),X1)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK10,bool),bool),sK12),s(fun(fun(sK10,bool),bool),X1))))) )),
  inference(resolution,[],[f393,f258])).
fof(f963,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK10,X0),s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12)))))) | ~p(s(bool,'const.bool.IN_2'(s(sK10,X0),s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))))) )),
  inference(resolution,[],[f393,f261])).
fof(f955,plain,(
  ( ! [X6,X7,X5] : (sP8('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),X5)),X6,X7) | s(bool,f(s(sK10,sK17(X5,sK11,bool,sK10)))) != s(bool,g(s(sK10,sK17(X5,sK11,bool,sK10))))) )),
  inference(superposition,[],[f702,f299])).
fof(f954,plain,(
  ( ! [X4,X2,X3] : (sP8('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),X2)),X3,X4) | s(bool,f(s(sK10,sK17(sK11,X2,bool,sK10)))) != s(bool,g(s(sK10,sK17(sK11,X2,bool,sK10))))) )),
  inference(superposition,[],[f702,f299])).
fof(f953,plain,(
  ( ! [X0,X1] : (sP8('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),'const.pred_set.EMPTY_0')),X0,X1) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),sK11))),s(fun(sK10,bool),sK11))))) )),
  inference(superposition,[],[f702,f255])).
fof(f952,plain,(
  ( ! [X6,X7,X5] : (sP8('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X5)),sK10)),s(fun(sK10,bool),sK11)),X6,X7) | s(bool,f(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f702,f299])).
fof(f951,plain,(
  ( ! [X4,X2,X3] : (sP8('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X2)),sK10)),s(fun(sK10,bool),sK11)),X3,X4) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f702,f299])).
fof(f940,plain,(
  ( ! [X6,X7,X5] : (sP8(X6,X7,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),X5))) | s(bool,f(s(sK10,sK17(X5,sK11,bool,sK10)))) != s(bool,g(s(sK10,sK17(X5,sK11,bool,sK10))))) )),
  inference(superposition,[],[f701,f299])).
fof(f939,plain,(
  ( ! [X4,X2,X3] : (sP8(X3,X4,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),X2))) | s(bool,f(s(sK10,sK17(sK11,X2,bool,sK10)))) != s(bool,g(s(sK10,sK17(sK11,X2,bool,sK10))))) )),
  inference(superposition,[],[f701,f299])).
fof(f938,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),sK11))),s(fun(sK10,bool),sK11))))) )),
  inference(superposition,[],[f701,f255])).
fof(f937,plain,(
  ( ! [X6,X7,X5] : (sP8(X6,X7,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X5)),sK10)),s(fun(sK10,bool),sK11))) | s(bool,f(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f701,f299])).
fof(f936,plain,(
  ( ! [X4,X2,X3] : (sP8(X3,X4,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X2)),sK10)),s(fun(sK10,bool),sK11))) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f701,f299])).
fof(f922,plain,(
  ( ! [X6,X7,X5] : (sP6('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),X5)),X6,X7) | s(bool,f(s(sK10,sK17(X5,sK11,bool,sK10)))) != s(bool,g(s(sK10,sK17(X5,sK11,bool,sK10))))) )),
  inference(superposition,[],[f687,f299])).
fof(f921,plain,(
  ( ! [X4,X2,X3] : (sP6('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),X2)),X3,X4) | s(bool,f(s(sK10,sK17(sK11,X2,bool,sK10)))) != s(bool,g(s(sK10,sK17(sK11,X2,bool,sK10))))) )),
  inference(superposition,[],[f687,f299])).
fof(f920,plain,(
  ( ! [X0,X1] : (sP6('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),'const.pred_set.EMPTY_0')),X0,X1) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),sK11))),s(fun(sK10,bool),sK11))))) )),
  inference(superposition,[],[f687,f255])).
fof(f919,plain,(
  ( ! [X6,X7,X5] : (sP6('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X5)),sK10)),s(fun(sK10,bool),sK11)),X6,X7) | s(bool,f(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f687,f299])).
fof(f918,plain,(
  ( ! [X4,X2,X3] : (sP6('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X2)),sK10)),s(fun(sK10,bool),sK11)),X3,X4) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f687,f299])).
fof(f906,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X21)) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X20),s(fun(fun(sK10,bool),bool),sK12)))) | ~p(s(bool,X22)) | ~sP5(X22,X21,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),X20),s(fun(sK10,bool),sK11)))) )),
  inference(superposition,[],[f212,f327])).
fof(f905,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X18)) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X17),s(fun(fun(sK10,bool),bool),sK12)))) | sP5(X19,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),X17),s(fun(sK10,bool),sK11)),X18) | p(s(bool,X19))) )),
  inference(superposition,[],[f212,f325])).
fof(f904,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X14),s(fun(fun(sK10,bool),bool),sK12)))) | sP5(X16,X15,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),X14),s(fun(sK10,bool),sK11))) | p(s(bool,X16))) )),
  inference(superposition,[],[f212,f325])).
fof(f893,plain,(
  ( ! [X4,X3] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),X4),s(fun(sK10,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X4),s(fun(fun(sK10,bool),bool),sK12)))) | s(bool,f(s(sK10,sK17(X3,sK11,bool,sK10)))) != s(bool,g(s(sK10,sK17(X3,sK11,bool,sK10))))) )),
  inference(superposition,[],[f212,f299])).
fof(f892,plain,(
  ( ! [X2,X1] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),X2),s(fun(sK10,bool),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X2),s(fun(fun(sK10,bool),bool),sK12)))) | s(bool,f(s(sK10,sK17(sK11,X1,bool,sK10)))) != s(bool,g(s(sK10,sK17(sK11,X1,bool,sK10))))) )),
  inference(superposition,[],[f212,f299])).
fof(f891,plain,(
  ( ! [X0] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),X0),s(fun(sK10,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X0),s(fun(fun(sK10,bool),bool),sK12)))) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),sK11))),s(fun(sK10,bool),sK11))))) )),
  inference(superposition,[],[f212,f255])).
fof(f890,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),X5),s(fun(sK10,bool),sK11)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X4),s(fun(fun(sK10,bool),bool),sK12)))) | s(bool,f(s(sK10,sK17(X5,X4,bool,sK10)))) != s(bool,g(s(sK10,sK17(X5,X4,bool,sK10))))) )),
  inference(superposition,[],[f212,f299])).
fof(f889,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),X3),s(fun(sK10,bool),sK11)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X2),s(fun(fun(sK10,bool),bool),sK12)))) | s(bool,f(s(sK10,sK17(X2,X3,bool,sK10)))) != s(bool,g(s(sK10,sK17(X2,X3,bool,sK10))))) )),
  inference(superposition,[],[f212,f299])).
fof(f888,plain,(
  ( ! [X1] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.EMPTY_0'),s(fun(sK10,bool),sK11)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X1),s(fun(fun(sK10,bool),bool),sK12)))) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),X1))),s(fun(sK10,bool),X1))))) )),
  inference(superposition,[],[f212,f255])).
fof(f887,plain,(
  ( ! [X0] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),X0),s(fun(sK10,bool),sK11)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK10,bool),bool),sK12)))) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),X0))),s(fun(sK10,bool),X0))))) )),
  inference(superposition,[],[f212,f255])).
fof(f881,plain,(
  ( ! [X171,X172] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X171),s(fun(fun(sK10,bool),bool),sK12)))) | s(bool,X172) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),X171),s(fun(sK10,bool),sK11))) | p(s(bool,X172))) )),
  inference(resolution,[],[f212,f379])).
fof(f379,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | ~p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(condensation,[],[f378])).
fof(f378,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2)) )),
  inference(condensation,[],[f370])).
fof(f370,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f211])).
fof(f211,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f210])).
fof(f210,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f122])).
fof(f122,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f93])).
fof(f93,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X12,X17,X13,X18] : ((s(bool,X12) = s(bool,X17) & (p(s(bool,X17)) => s(bool,X13) = s(bool,X18))) => ((p(s(bool,X12)) => p(s(bool,X13))) <=> (p(s(bool,X17)) => p(s(bool,X18)))))),
  file('thm.util_prob.BIGINTER_SUBSET','thm.bool.IMP_CONG')).
fof(f880,plain,(
  ( ! [X169,X170,X168] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X168),s(fun(fun(sK10,bool),bool),sK12)))) | sP9('const.pred_set.SUBSET_2'(s(fun(sK10,bool),X168),s(fun(sK10,bool),sK11)),X169,X170) | ~sP8('const.pred_set.SUBSET_2'(s(fun(sK10,bool),X168),s(fun(sK10,bool),sK11)),X169,X170)) )),
  inference(resolution,[],[f212,f377])).
fof(f879,plain,(
  ( ! [X167,X165,X166] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X165),s(fun(fun(sK10,bool),bool),sK12)))) | sP7('const.pred_set.SUBSET_2'(s(fun(sK10,bool),X165),s(fun(sK10,bool),sK11)),X166,X167) | ~sP6('const.pred_set.SUBSET_2'(s(fun(sK10,bool),X165),s(fun(sK10,bool),sK11)),X166,X167)) )),
  inference(resolution,[],[f212,f376])).
fof(f877,plain,(
  ( ! [X161,X159,X160] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X159),s(fun(fun(sK10,bool),bool),sK12)))) | p(s(bool,X160)) | ~p(s(bool,X161)) | ~sP9('const.pred_set.SUBSET_2'(s(fun(sK10,bool),X159),s(fun(sK10,bool),sK11)),X160,X161)) )),
  inference(resolution,[],[f212,f360])).
fof(f876,plain,(
  ( ! [X158,X156,X157] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X156),s(fun(fun(sK10,bool),bool),sK12)))) | p(s(bool,X157)) | ~p(s(bool,X158)) | ~sP9(X158,X157,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),X156),s(fun(sK10,bool),sK11)))) )),
  inference(resolution,[],[f212,f360])).
fof(f875,plain,(
  ( ! [X154,X155,X153] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X153),s(fun(fun(sK10,bool),bool),sK12)))) | sP9('const.pred_set.SUBSET_2'(s(fun(sK10,bool),X153),s(fun(sK10,bool),sK11)),X154,X155) | ~p(s(bool,X154))) )),
  inference(resolution,[],[f212,f359])).
fof(f874,plain,(
  ( ! [X152,X151,X150] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X150),s(fun(fun(sK10,bool),bool),sK12)))) | sP9(X151,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),X150),s(fun(sK10,bool),sK11)),X152) | ~p(s(bool,X151))) )),
  inference(resolution,[],[f212,f359])).
fof(f873,plain,(
  ( ! [X149,X147,X148] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X147),s(fun(fun(sK10,bool),bool),sK12)))) | sP9('const.pred_set.SUBSET_2'(s(fun(sK10,bool),X147),s(fun(sK10,bool),sK11)),X148,X149) | p(s(bool,X149))) )),
  inference(resolution,[],[f212,f358])).
fof(f872,plain,(
  ( ! [X146,X144,X145] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X144),s(fun(fun(sK10,bool),bool),sK12)))) | sP9(X145,X146,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),X144),s(fun(sK10,bool),sK11))) | p(s(bool,X146)) | p(s(bool,X145))) )),
  inference(resolution,[],[f212,f357])).
fof(f871,plain,(
  ( ! [X142,X143,X141] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X141),s(fun(fun(sK10,bool),bool),sK12)))) | sP8(X142,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),X141),s(fun(sK10,bool),sK11)),X143)) )),
  inference(resolution,[],[f212,f355])).
fof(f870,plain,(
  ( ! [X140,X138,X139] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X138),s(fun(fun(sK10,bool),bool),sK12)))) | ~p(s(bool,X139)) | p(s(bool,X140)) | ~sP8(X139,X140,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),X138),s(fun(sK10,bool),sK11)))) )),
  inference(resolution,[],[f212,f353])).
fof(f869,plain,(
  ( ! [X136,X135,X137] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X135),s(fun(fun(sK10,bool),bool),sK12)))) | p(s(bool,X136)) | ~p(s(bool,X137)) | ~sP8('const.pred_set.SUBSET_2'(s(fun(sK10,bool),X135),s(fun(sK10,bool),sK11)),X136,X137)) )),
  inference(resolution,[],[f212,f353])).
fof(f868,plain,(
  ( ! [X134,X132,X133] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X132),s(fun(fun(sK10,bool),bool),sK12)))) | sP9(X133,X134,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),X132),s(fun(sK10,bool),sK11))) | ~sP8(X133,X134,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),X132),s(fun(sK10,bool),sK11))) | p(s(bool,X134))) )),
  inference(resolution,[],[f212,f352])).
fof(f867,plain,(
  ( ! [X130,X131,X129] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X129),s(fun(fun(sK10,bool),bool),sK12)))) | sP9('const.pred_set.SUBSET_2'(s(fun(sK10,bool),X129),s(fun(sK10,bool),sK11)),X130,X131) | ~sP8('const.pred_set.SUBSET_2'(s(fun(sK10,bool),X129),s(fun(sK10,bool),sK11)),X130,X131) | p(s(bool,X130))) )),
  inference(resolution,[],[f212,f351])).
fof(f865,plain,(
  ( ! [X125,X123,X124] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X123),s(fun(fun(sK10,bool),bool),sK12)))) | sP9('const.pred_set.SUBSET_2'(s(fun(sK10,bool),X123),s(fun(sK10,bool),sK11)),X124,X125) | ~sP8('const.pred_set.SUBSET_2'(s(fun(sK10,bool),X123),s(fun(sK10,bool),sK11)),X124,X125) | ~p(s(bool,X125))) )),
  inference(resolution,[],[f212,f350])).
fof(f864,plain,(
  ( ! [X121,X122,X120] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X120),s(fun(fun(sK10,bool),bool),sK12)))) | p(s(bool,X121)) | ~sP9(X121,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),X120),s(fun(sK10,bool),sK11)),X122)) )),
  inference(resolution,[],[f212,f347])).
fof(f861,plain,(
  ( ! [X111,X112,X113] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X111),s(fun(fun(sK10,bool),bool),sK12)))) | p(s(bool,X112)) | p(s(bool,X113)) | ~sP7('const.pred_set.SUBSET_2'(s(fun(sK10,bool),X111),s(fun(sK10,bool),sK11)),X112,X113)) )),
  inference(resolution,[],[f212,f343])).
fof(f860,plain,(
  ( ! [X109,X110,X108] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X108),s(fun(fun(sK10,bool),bool),sK12)))) | sP7('const.pred_set.SUBSET_2'(s(fun(sK10,bool),X108),s(fun(sK10,bool),sK11)),X109,X110) | ~p(s(bool,X109))) )),
  inference(resolution,[],[f212,f342])).
fof(f859,plain,(
  ( ! [X107,X105,X106] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X105),s(fun(fun(sK10,bool),bool),sK12)))) | sP7(X106,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),X105),s(fun(sK10,bool),sK11)),X107) | ~p(s(bool,X106))) )),
  inference(resolution,[],[f212,f342])).
fof(f858,plain,(
  ( ! [X103,X102,X104] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X102),s(fun(fun(sK10,bool),bool),sK12)))) | sP7('const.pred_set.SUBSET_2'(s(fun(sK10,bool),X102),s(fun(sK10,bool),sK11)),X103,X104) | ~p(s(bool,X104))) )),
  inference(resolution,[],[f212,f341])).
fof(f857,plain,(
  ( ! [X101,X99,X100] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X99),s(fun(fun(sK10,bool),bool),sK12)))) | sP7(X100,X101,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),X99),s(fun(sK10,bool),sK11))) | ~p(s(bool,X100))) )),
  inference(resolution,[],[f212,f341])).
fof(f856,plain,(
  ( ! [X97,X98,X96] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X96),s(fun(fun(sK10,bool),bool),sK12)))) | sP6(X97,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),X96),s(fun(sK10,bool),sK11)),X98)) )),
  inference(resolution,[],[f212,f338])).
fof(f855,plain,(
  ( ! [X94,X95,X93] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X93),s(fun(fun(sK10,bool),bool),sK12)))) | sP6(X94,X95,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),X93),s(fun(sK10,bool),sK11)))) )),
  inference(resolution,[],[f212,f337])).
fof(f854,plain,(
  ( ! [X92,X90,X91] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X90),s(fun(fun(sK10,bool),bool),sK12)))) | p(s(bool,X91)) | p(s(bool,X92)) | ~sP6('const.pred_set.SUBSET_2'(s(fun(sK10,bool),X90),s(fun(sK10,bool),sK11)),X91,X92)) )),
  inference(resolution,[],[f212,f336])).
fof(f853,plain,(
  ( ! [X88,X87,X89] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X87),s(fun(fun(sK10,bool),bool),sK12)))) | sP7('const.pred_set.SUBSET_2'(s(fun(sK10,bool),X87),s(fun(sK10,bool),sK11)),X88,X89) | ~sP6('const.pred_set.SUBSET_2'(s(fun(sK10,bool),X87),s(fun(sK10,bool),sK11)),X88,X89) | p(s(bool,X88))) )),
  inference(resolution,[],[f212,f334])).
fof(f852,plain,(
  ( ! [X85,X86,X84] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X84),s(fun(fun(sK10,bool),bool),sK12)))) | sP7('const.pred_set.SUBSET_2'(s(fun(sK10,bool),X84),s(fun(sK10,bool),sK11)),X85,X86) | ~sP6('const.pred_set.SUBSET_2'(s(fun(sK10,bool),X84),s(fun(sK10,bool),sK11)),X85,X86) | p(s(bool,X86))) )),
  inference(resolution,[],[f212,f333])).
fof(f851,plain,(
  ( ! [X83,X81,X82] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X81),s(fun(fun(sK10,bool),bool),sK12)))) | p(s(bool,X82)) | ~sP7(X82,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),X81),s(fun(sK10,bool),sK11)),X83)) )),
  inference(resolution,[],[f212,f330])).
fof(f850,plain,(
  ( ! [X80,X78,X79] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X78),s(fun(fun(sK10,bool),bool),sK12)))) | p(s(bool,X79)) | ~sP7(X79,X80,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),X78),s(fun(sK10,bool),sK11)))) )),
  inference(resolution,[],[f212,f329])).
fof(f849,plain,(
  ( ! [X76,X77,X75] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X75),s(fun(fun(sK10,bool),bool),sK12)))) | s(bool,X76) = s(bool,X77) | ~sP5('const.pred_set.SUBSET_2'(s(fun(sK10,bool),X75),s(fun(sK10,bool),sK11)),X77,X76)) )),
  inference(resolution,[],[f212,f327])).
fof(f848,plain,(
  ( ! [X74,X72,X73] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X72),s(fun(fun(sK10,bool),bool),sK12)))) | sP5('const.pred_set.SUBSET_2'(s(fun(sK10,bool),X72),s(fun(sK10,bool),sK11)),X73,X74) | s(bool,X73) != s(bool,X74)) )),
  inference(resolution,[],[f212,f326])).
fof(f847,plain,(
  ( ! [X70,X71,X69] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X69),s(fun(fun(sK10,bool),bool),sK12)))) | sP2('const.pred_set.SUBSET_2'(s(fun(sK10,bool),X69),s(fun(sK10,bool),sK11)),X70,X71)) )),
  inference(resolution,[],[f212,f324])).
fof(f846,plain,(
  ( ! [X68,X66,X67] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X66),s(fun(fun(sK10,bool),bool),sK12)))) | sP2(X67,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),X66),s(fun(sK10,bool),sK11)),X68)) )),
  inference(resolution,[],[f212,f323])).
fof(f845,plain,(
  ( ! [X64,X65,X63] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X63),s(fun(fun(sK10,bool),bool),sK12)))) | sP2(X64,X65,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),X63),s(fun(sK10,bool),sK11)))) )),
  inference(resolution,[],[f212,f322])).
fof(f844,plain,(
  ( ! [X61,X62,X60] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X60),s(fun(fun(sK10,bool),bool),sK12)))) | sP3(X61,X62,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),X60),s(fun(sK10,bool),sK11)))) )),
  inference(resolution,[],[f212,f318])).
fof(f843,plain,(
  ( ! [X59,X57,X58] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X57),s(fun(fun(sK10,bool),bool),sK12)))) | ~p(s(bool,X58)) | p(s(bool,X59)) | ~sP3(X58,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),X57),s(fun(sK10,bool),sK11)),X59)) )),
  inference(resolution,[],[f212,f317])).
fof(f842,plain,(
  ( ! [X54,X56,X55] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X54),s(fun(fun(sK10,bool),bool),sK12)))) | ~p(s(bool,X55)) | p(s(bool,X56)) | ~sP3('const.pred_set.SUBSET_2'(s(fun(sK10,bool),X54),s(fun(sK10,bool),sK11)),X55,X56)) )),
  inference(resolution,[],[f212,f317])).
fof(f841,plain,(
  ( ! [X52,X53,X51] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X51),s(fun(fun(sK10,bool),bool),sK12)))) | sP4(X52,X53,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),X51),s(fun(sK10,bool),sK11)))) )),
  inference(resolution,[],[f212,f314])).
fof(f840,plain,(
  ( ! [X50,X48,X49] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X48),s(fun(fun(sK10,bool),bool),sK12)))) | ~p(s(bool,X49)) | p(s(bool,X50)) | ~sP4(X49,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),X48),s(fun(sK10,bool),sK11)),X50)) )),
  inference(resolution,[],[f212,f313])).
fof(f839,plain,(
  ( ! [X47,X45,X46] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X45),s(fun(fun(sK10,bool),bool),sK12)))) | ~p(s(bool,X46)) | p(s(bool,X47)) | ~sP4('const.pred_set.SUBSET_2'(s(fun(sK10,bool),X45),s(fun(sK10,bool),sK11)),X46,X47)) )),
  inference(resolution,[],[f212,f313])).
fof(f838,plain,(
  ( ! [X43,X44,X42] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X42),s(fun(fun(sK10,bool),bool),sK12)))) | sP5(X43,X44,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),X42),s(fun(sK10,bool),sK11))) | ~sP4(X43,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),X42),s(fun(sK10,bool),sK11)),X44) | ~sP3(X44,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),X42),s(fun(sK10,bool),sK11)),X43) | ~sP2('const.pred_set.SUBSET_2'(s(fun(sK10,bool),X42),s(fun(sK10,bool),sK11)),X44,X43)) )),
  inference(resolution,[],[f212,f310])).
fof(f837,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X39),s(fun(fun(sK10,bool),bool),sK12)))) | ~p(s(bool,X40)) | p(s(bool,X41)) | ~sP5(X40,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),X39),s(fun(sK10,bool),sK11)),X41)) )),
  inference(resolution,[],[f212,f309])).
fof(f836,plain,(
  ( ! [X37,X38,X36] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X36),s(fun(fun(sK10,bool),bool),sK12)))) | ~p(s(bool,X37)) | p(s(bool,X38)) | ~sP5('const.pred_set.SUBSET_2'(s(fun(sK10,bool),X36),s(fun(sK10,bool),sK11)),X37,X38)) )),
  inference(resolution,[],[f212,f309])).
fof(f831,plain,(
  ( ! [X26,X27] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X26),s(fun(fun(sK10,bool),bool),sK12)))) | sP1(X27,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),X26),s(fun(sK10,bool),sK11))) | p(s(bool,X27))) )),
  inference(resolution,[],[f212,f294])).
fof(f830,plain,(
  ( ! [X24,X25] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X24),s(fun(fun(sK10,bool),bool),sK12)))) | sP1('const.pred_set.SUBSET_2'(s(fun(sK10,bool),X24),s(fun(sK10,bool),sK11)),X25) | p(s(bool,X25))) )),
  inference(resolution,[],[f212,f291])).
fof(f829,plain,(
  ( ! [X23,X22] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X22),s(fun(fun(sK10,bool),bool),sK12)))) | ~p(s(bool,X23)) | ~sP1(X23,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),X22),s(fun(sK10,bool),sK11)))) )),
  inference(resolution,[],[f212,f290])).
fof(f828,plain,(
  ( ! [X21,X20] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X20),s(fun(fun(sK10,bool),bool),sK12)))) | ~p(s(bool,X21)) | ~sP1('const.pred_set.SUBSET_2'(s(fun(sK10,bool),X20),s(fun(sK10,bool),sK11)),X21)) )),
  inference(resolution,[],[f212,f290])).
fof(f825,plain,(
  ( ! [X14,X15] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X14),s(fun(fun(sK10,bool),bool),sK12)))) | sP0('const.pred_set.SUBSET_2'(s(fun(sK10,bool),X14),s(fun(sK10,bool),sK11)),X15)) )),
  inference(resolution,[],[f212,f276])).
fof(f824,plain,(
  ( ! [X12,X13] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X12),s(fun(fun(sK10,bool),bool),sK12)))) | sP0(X13,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),X12),s(fun(sK10,bool),sK11)))) )),
  inference(resolution,[],[f212,f275])).
fof(f822,plain,(
  ( ! [X8,X9] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X8),s(fun(fun(sK10,bool),bool),sK12)))) | s(bool,X9) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),X8),s(fun(sK10,bool),sK11))) | ~p(s(bool,X9))) )),
  inference(resolution,[],[f212,f271])).
fof(f271,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f108])).
fof(f108,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X10,X11] : ((p(s(bool,X10)) => p(s(bool,X11))) => ((p(s(bool,X11)) => p(s(bool,X10))) => s(bool,X10) = s(bool,X11)))),
  file('thm.util_prob.BIGINTER_SUBSET','thm.bool.IMP_ANTISYM_AX')).
fof(f818,plain,(
  ( ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X4),s(fun(fun(sK10,bool),bool),sK12)))) | s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),X4),s(fun(sK10,bool),sK11)))) )),
  inference(resolution,[],[f212,f232])).
fof(f816,plain,(
  ( ! [X2] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X2),s(fun(fun(sK10,bool),bool),sK12)))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),X2),s(fun(sK10,bool),sK11)))) )),
  inference(resolution,[],[f212,f229])).
fof(f815,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X0),s(fun(fun(sK10,bool),bool),sK12)))) | p(s(bool,'const.bool.IN_2'(s(sK10,X1),s(fun(sK10,bool),sK11)))) | ~p(s(bool,'const.bool.IN_2'(s(sK10,X1),s(fun(sK10,bool),X0))))) )),
  inference(resolution,[],[f212,f258])).
fof(f812,plain,(
  ( ! [X6,X7,X5] : (sP3('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),X5)),X6,X7) | s(bool,f(s(sK10,sK17(X5,sK11,bool,sK10)))) != s(bool,g(s(sK10,sK17(X5,sK11,bool,sK10))))) )),
  inference(superposition,[],[f673,f299])).
fof(f811,plain,(
  ( ! [X4,X2,X3] : (sP3('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),X2)),X3,X4) | s(bool,f(s(sK10,sK17(sK11,X2,bool,sK10)))) != s(bool,g(s(sK10,sK17(sK11,X2,bool,sK10))))) )),
  inference(superposition,[],[f673,f299])).
fof(f810,plain,(
  ( ! [X0,X1] : (sP3('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),'const.pred_set.EMPTY_0')),X0,X1) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),sK11))),s(fun(sK10,bool),sK11))))) )),
  inference(superposition,[],[f673,f255])).
fof(f809,plain,(
  ( ! [X6,X7,X5] : (sP3('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X5)),sK10)),s(fun(sK10,bool),sK11)),X6,X7) | s(bool,f(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f673,f299])).
fof(f808,plain,(
  ( ! [X4,X2,X3] : (sP3('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X2)),sK10)),s(fun(sK10,bool),sK11)),X3,X4) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f673,f299])).
fof(f797,plain,(
  ( ! [X6,X7,X5] : (sP3(X6,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),X5)),X7) | s(bool,f(s(sK10,sK17(X5,sK11,bool,sK10)))) != s(bool,g(s(sK10,sK17(X5,sK11,bool,sK10))))) )),
  inference(superposition,[],[f672,f299])).
fof(f796,plain,(
  ( ! [X4,X2,X3] : (sP3(X3,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),X2)),X4) | s(bool,f(s(sK10,sK17(sK11,X2,bool,sK10)))) != s(bool,g(s(sK10,sK17(sK11,X2,bool,sK10))))) )),
  inference(superposition,[],[f672,f299])).
fof(f795,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),'const.pred_set.EMPTY_0')),X1) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),sK11))),s(fun(sK10,bool),sK11))))) )),
  inference(superposition,[],[f672,f255])).
fof(f794,plain,(
  ( ! [X6,X7,X5] : (sP3(X6,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X5)),sK10)),s(fun(sK10,bool),sK11)),X7) | s(bool,f(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f672,f299])).
fof(f793,plain,(
  ( ! [X4,X2,X3] : (sP3(X3,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X2)),sK10)),s(fun(sK10,bool),sK11)),X4) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f672,f299])).
fof(f785,plain,(
  ( ! [X6,X7,X5] : (sP4('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),X5)),X6,X7) | s(bool,f(s(sK10,sK17(X5,sK11,bool,sK10)))) != s(bool,g(s(sK10,sK17(X5,sK11,bool,sK10))))) )),
  inference(superposition,[],[f670,f299])).
fof(f784,plain,(
  ( ! [X4,X2,X3] : (sP4('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),X2)),X3,X4) | s(bool,f(s(sK10,sK17(sK11,X2,bool,sK10)))) != s(bool,g(s(sK10,sK17(sK11,X2,bool,sK10))))) )),
  inference(superposition,[],[f670,f299])).
fof(f783,plain,(
  ( ! [X0,X1] : (sP4('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),'const.pred_set.EMPTY_0')),X0,X1) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),sK11))),s(fun(sK10,bool),sK11))))) )),
  inference(superposition,[],[f670,f255])).
fof(f782,plain,(
  ( ! [X6,X7,X5] : (sP4('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X5)),sK10)),s(fun(sK10,bool),sK11)),X6,X7) | s(bool,f(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f670,f299])).
fof(f781,plain,(
  ( ! [X4,X2,X3] : (sP4('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X2)),sK10)),s(fun(sK10,bool),sK11)),X3,X4) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f670,f299])).
fof(f768,plain,(
  ( ! [X6,X7,X5] : (sP4(X6,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),X5)),X7) | s(bool,f(s(sK10,sK17(X5,sK11,bool,sK10)))) != s(bool,g(s(sK10,sK17(X5,sK11,bool,sK10))))) )),
  inference(superposition,[],[f669,f299])).
fof(f767,plain,(
  ( ! [X4,X2,X3] : (sP4(X3,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),X2)),X4) | s(bool,f(s(sK10,sK17(sK11,X2,bool,sK10)))) != s(bool,g(s(sK10,sK17(sK11,X2,bool,sK10))))) )),
  inference(superposition,[],[f669,f299])).
fof(f766,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),'const.pred_set.EMPTY_0')),X1) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),sK11))),s(fun(sK10,bool),sK11))))) )),
  inference(superposition,[],[f669,f255])).
fof(f765,plain,(
  ( ! [X6,X7,X5] : (sP4(X6,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X5)),sK10)),s(fun(sK10,bool),sK11)),X7) | s(bool,f(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f669,f299])).
fof(f764,plain,(
  ( ! [X4,X2,X3] : (sP4(X3,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X2)),sK10)),s(fun(sK10,bool),sK11)),X4) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f669,f299])).
fof(f775,plain,(
  ( ! [X6,X7] : (sP5(X6,X7,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11))) | p(s(bool,X6)) | ~sP2('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)),X7,X6)) )),
  inference(subsumption_resolution,[],[f762,f672])).
fof(f762,plain,(
  ( ! [X6,X7] : (sP5(X6,X7,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11))) | p(s(bool,X6)) | ~sP3(X7,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)),X6) | ~sP2('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)),X7,X6)) )),
  inference(resolution,[],[f669,f312])).
fof(f774,plain,(
  ( ! [X4,X5] : (sP5(X4,X5,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11))) | p(s(bool,X5)) | ~sP2('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)),X5,X4)) )),
  inference(subsumption_resolution,[],[f761,f320])).
fof(f761,plain,(
  ( ! [X4,X5] : (sP5(X4,X5,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11))) | p(s(bool,X5)) | ~sP3(X5,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)),X4) | ~sP2('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)),X5,X4)) )),
  inference(resolution,[],[f669,f311])).
fof(f735,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | ~p(s(bool,X11)) | ~sP5(X11,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)),X10)) )),
  inference(superposition,[],[f395,f327])).
fof(f733,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | sP5(X7,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)),X6) | p(s(bool,X7))) )),
  inference(superposition,[],[f395,f325])).
fof(f732,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | sP5(X5,X4,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11))) | p(s(bool,X5))) )),
  inference(superposition,[],[f395,f325])).
fof(f721,plain,(
  ( ! [X1] : (~p(s(bool,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),X1)))) | s(bool,f(s(sK10,sK17(X1,sK11,bool,sK10)))) != s(bool,g(s(sK10,sK17(X1,sK11,bool,sK10))))) )),
  inference(superposition,[],[f395,f299])).
fof(f720,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),X0)))) | s(bool,f(s(sK10,sK17(sK11,X0,bool,sK10)))) != s(bool,g(s(sK10,sK17(sK11,X0,bool,sK10))))) )),
  inference(superposition,[],[f395,f299])).
fof(f719,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),sK11))),s(fun(sK10,bool),sK11))))),
  inference(superposition,[],[f395,f255])).
fof(f718,plain,(
  ( ! [X1] : (~p(s(bool,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X1)),sK10)),s(fun(sK10,bool),sK11)))) | s(bool,f(s(fun(sK10,bool),sK17(X1,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X1,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f395,f299])).
fof(f717,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X0)),sK10)),s(fun(sK10,bool),sK11)))) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X0,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X0,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f395,f299])).
fof(f706,plain,(
  ( ! [X105,X104] : (~p(s(bool,X104)) | ~p(s(bool,X105)) | ~sP9(X105,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)),X104)) )),
  inference(resolution,[],[f395,f360])).
fof(f705,plain,(
  ( ! [X103,X102] : (sP9(X102,X103,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11))) | ~p(s(bool,X102))) )),
  inference(resolution,[],[f395,f358])).
fof(f704,plain,(
  ( ! [X101,X100] : (sP9('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)),X100,X101) | p(s(bool,X100)) | ~p(s(bool,X101))) )),
  inference(resolution,[],[f395,f357])).
fof(f703,plain,(
  ( ! [X99,X98] : (sP9(X98,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)),X99) | ~p(s(bool,X99)) | p(s(bool,X98))) )),
  inference(resolution,[],[f395,f357])).
fof(f702,plain,(
  ( ! [X97,X96] : (sP8('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)),X96,X97)) )),
  inference(resolution,[],[f395,f356])).
fof(f701,plain,(
  ( ! [X94,X95] : (sP8(X94,X95,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)))) )),
  inference(resolution,[],[f395,f354])).
fof(f700,plain,(
  ( ! [X92,X93] : (~p(s(bool,X92)) | ~p(s(bool,X93)) | ~sP8(X92,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)),X93)) )),
  inference(resolution,[],[f395,f353])).
fof(f699,plain,(
  ( ! [X90,X91] : (sP9(X90,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)),X91) | ~sP8(X90,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)),X91) | ~p(s(bool,X91))) )),
  inference(resolution,[],[f395,f352])).
fof(f697,plain,(
  ( ! [X87,X86] : (~sP9('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)),X86,X87) | ~p(s(bool,X86))) )),
  inference(resolution,[],[f395,f347])).
fof(f696,plain,(
  ( ! [X85,X84] : (~sP9(X84,X85,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11))) | p(s(bool,X84))) )),
  inference(resolution,[],[f395,f346])).
fof(f695,plain,(
  ( ! [X83,X82] : (~sP9('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)),X83,X82) | p(s(bool,X82))) )),
  inference(resolution,[],[f395,f346])).
fof(f692,plain,(
  ( ! [X76,X77] : (p(s(bool,X76)) | ~p(s(bool,X77)) | ~sP7(X77,X76,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)))) )),
  inference(resolution,[],[f395,f343])).
fof(f691,plain,(
  ( ! [X74,X75] : (p(s(bool,X74)) | ~p(s(bool,X75)) | ~sP7(X75,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)),X74)) )),
  inference(resolution,[],[f395,f343])).
fof(f690,plain,(
  ( ! [X72,X73] : (sP7('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)),X72,X73) | p(s(bool,X72)) | p(s(bool,X73))) )),
  inference(resolution,[],[f395,f340])).
fof(f689,plain,(
  ( ! [X70,X71] : (sP7(X70,X71,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11))) | p(s(bool,X71)) | p(s(bool,X70))) )),
  inference(resolution,[],[f395,f340])).
fof(f688,plain,(
  ( ! [X68,X69] : (sP7(X68,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)),X69) | p(s(bool,X69)) | p(s(bool,X68))) )),
  inference(resolution,[],[f395,f340])).
fof(f687,plain,(
  ( ! [X66,X67] : (sP6('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)),X66,X67)) )),
  inference(resolution,[],[f395,f339])).
fof(f686,plain,(
  ( ! [X64,X65] : (~p(s(bool,X64)) | p(s(bool,X65)) | ~sP6(X64,X65,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)))) )),
  inference(resolution,[],[f395,f336])).
fof(f685,plain,(
  ( ! [X62,X63] : (~p(s(bool,X62)) | p(s(bool,X63)) | ~sP6(X62,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)),X63)) )),
  inference(resolution,[],[f395,f336])).
fof(f684,plain,(
  ( ! [X61,X60] : (sP7(X60,X61,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11))) | ~sP6(X60,X61,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11))) | p(s(bool,X61))) )),
  inference(resolution,[],[f395,f335])).
fof(f683,plain,(
  ( ! [X59,X58] : (sP7(X58,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)),X59) | ~sP6(X58,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)),X59) | p(s(bool,X59))) )),
  inference(resolution,[],[f395,f335])).
fof(f680,plain,(
  ( ! [X52,X53] : (~sP7('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)),X52,X53) | ~p(s(bool,X52))) )),
  inference(resolution,[],[f395,f330])).
fof(f679,plain,(
  ( ! [X50,X51] : (~sP7('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)),X51,X50) | ~p(s(bool,X50))) )),
  inference(resolution,[],[f395,f329])).
fof(f678,plain,(
  ( ! [X48,X49] : (s(bool,X48) != s(bool,X49) | ~sP5('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)),X49,X48)) )),
  inference(resolution,[],[f395,f328])).
fof(f677,plain,(
  ( ! [X47,X46] : (sP5('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)),X46,X47) | s(bool,X46) = s(bool,X47)) )),
  inference(resolution,[],[f395,f325])).
fof(f676,plain,(
  ( ! [X45,X44] : (p(s(bool,X44)) | p(s(bool,X45)) | ~sP2(X44,X45,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)))) )),
  inference(resolution,[],[f395,f321])).
fof(f675,plain,(
  ( ! [X43,X42] : (p(s(bool,X42)) | p(s(bool,X43)) | ~sP2(X42,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)),X43)) )),
  inference(resolution,[],[f395,f321])).
fof(f674,plain,(
  ( ! [X41,X40] : (p(s(bool,X40)) | p(s(bool,X41)) | ~sP2('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)),X40,X41)) )),
  inference(resolution,[],[f395,f321])).
fof(f746,plain,(
  ( ! [X24,X25] : (sP5(X24,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)),X25) | ~sP4(X24,X25,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11))) | ~sP2(X25,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)),X24)) )),
  inference(subsumption_resolution,[],[f666,f673])).
fof(f673,plain,(
  ( ! [X39,X38] : (sP3('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)),X38,X39)) )),
  inference(resolution,[],[f395,f320])).
fof(f672,plain,(
  ( ! [X37,X36] : (sP3(X36,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)),X37)) )),
  inference(resolution,[],[f395,f319])).
fof(f671,plain,(
  ( ! [X35,X34] : (~p(s(bool,X34)) | ~p(s(bool,X35)) | ~sP3(X34,X35,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)))) )),
  inference(resolution,[],[f395,f317])).
fof(f745,plain,(
  ( ! [X26,X27] : (sP5('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)),X26,X27) | ~sP3(X26,X27,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11))) | ~sP2(X27,X26,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)))) )),
  inference(subsumption_resolution,[],[f667,f670])).
fof(f670,plain,(
  ( ! [X33,X32] : (sP4('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)),X32,X33)) )),
  inference(resolution,[],[f395,f316])).
fof(f669,plain,(
  ( ! [X30,X31] : (sP4(X30,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)),X31)) )),
  inference(resolution,[],[f395,f315])).
fof(f668,plain,(
  ( ! [X28,X29] : (~p(s(bool,X28)) | ~p(s(bool,X29)) | ~sP4(X28,X29,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)))) )),
  inference(resolution,[],[f395,f313])).
fof(f667,plain,(
  ( ! [X26,X27] : (sP5('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)),X26,X27) | ~sP4('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)),X27,X26) | ~sP3(X26,X27,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11))) | ~sP2(X27,X26,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)))) )),
  inference(resolution,[],[f395,f312])).
fof(f666,plain,(
  ( ! [X24,X25] : (sP5(X24,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)),X25) | ~sP4(X24,X25,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11))) | ~sP3('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)),X25,X24) | ~sP2(X25,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)),X24)) )),
  inference(resolution,[],[f395,f311])).
fof(f665,plain,(
  ( ! [X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | ~sP5(X22,X23,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)))) )),
  inference(resolution,[],[f395,f309])).
fof(f660,plain,(
  ( ! [X17] : (sP1('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)),X17) | ~p(s(bool,X17))) )),
  inference(resolution,[],[f395,f294])).
fof(f659,plain,(
  ( ! [X16] : (sP1(X16,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11))) | ~p(s(bool,X16))) )),
  inference(resolution,[],[f395,f291])).
fof(f658,plain,(
  ( ! [X15] : (~sP1(X15,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11))) | p(s(bool,X15))) )),
  inference(resolution,[],[f395,f289])).
fof(f657,plain,(
  ( ! [X14] : (~sP1('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)),X14) | p(s(bool,X14))) )),
  inference(resolution,[],[f395,f289])).
fof(f744,plain,(
  ( ! [X11] : (~sP0(X11,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11))) | p(s(bool,X11))) )),
  inference(subsumption_resolution,[],[f654,f221])).
fof(f654,plain,(
  ( ! [X11] : (p(s(bool,'F_0')) | p(s(bool,X11)) | ~sP0(X11,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)))) )),
  inference(resolution,[],[f395,f274])).
fof(f743,plain,(
  ( ! [X10] : (~sP0('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)),X10) | p(s(bool,X10))) )),
  inference(subsumption_resolution,[],[f653,f221])).
fof(f653,plain,(
  ( ! [X10] : (p(s(bool,'F_0')) | p(s(bool,X10)) | ~sP0('const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)),X10)) )),
  inference(resolution,[],[f395,f274])).
fof(f650,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,happ(s(fun(X5,bool),X6),s(X5,X7)))) | ~p(s(bool,happ(s(fun(X5,bool),X6),s(X5,sK16(X6,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)),X5)))))) )),
  inference(resolution,[],[f395,f266])).
fof(f649,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,happ(s(fun(X2,bool),X3),s(X2,X4)))) | ~p(s(bool,happ(s(fun(X2,bool),X3),s(X2,sK15(X3,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)),X2)))))) )),
  inference(resolution,[],[f395,f264])).
fof(f646,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)))),
  inference(resolution,[],[f395,f233])).
fof(f644,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11)))),
  inference(resolution,[],[f395,f228])).
fof(f643,plain,(
  ( ! [X1] : (~p(s(bool,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),X1)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),X1),s(fun(sK10,bool),sK11))))) )),
  inference(resolution,[],[f395,f258])).
fof(f642,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),X0)))) | ~p(s(bool,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X0))))))) )),
  inference(resolution,[],[f395,f261])).
fof(f634,plain,(
  ( ! [X6,X7,X5] : (sP8('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),X5)),X6,X7) | s(bool,f(s(sK10,sK17(X5,sK11,bool,sK10)))) != s(bool,g(s(sK10,sK17(X5,sK11,bool,sK10))))) )),
  inference(superposition,[],[f454,f299])).
fof(f633,plain,(
  ( ! [X4,X2,X3] : (sP8('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),X2)),X3,X4) | s(bool,f(s(sK10,sK17(sK11,X2,bool,sK10)))) != s(bool,g(s(sK10,sK17(sK11,X2,bool,sK10))))) )),
  inference(superposition,[],[f454,f299])).
fof(f632,plain,(
  ( ! [X0,X1] : (sP8('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),'const.pred_set.EMPTY_0')),X0,X1) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),sK11))),s(fun(sK10,bool),sK11))))) )),
  inference(superposition,[],[f454,f255])).
fof(f631,plain,(
  ( ! [X6,X7,X5] : (sP8('const.pred_set.SUBSET_2'(s(fun(sK10,bool),X5),s(fun(sK10,bool),sK11)),X6,X7) | s(bool,f(s(sK10,sK17(X5,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10)))) != s(bool,g(s(sK10,sK17(X5,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10))))) )),
  inference(superposition,[],[f454,f299])).
fof(f630,plain,(
  ( ! [X4,X2,X3] : (sP8('const.pred_set.SUBSET_2'(s(fun(sK10,bool),X2),s(fun(sK10,bool),sK11)),X3,X4) | s(bool,f(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X2,bool,sK10)))) != s(bool,g(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X2,bool,sK10))))) )),
  inference(superposition,[],[f454,f299])).
fof(f629,plain,(
  ( ! [X0,X1] : (sP8('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.EMPTY_0'),s(fun(sK10,bool),sK11)),X0,X1) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))),s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))))) )),
  inference(superposition,[],[f454,f255])).
fof(f628,plain,(
  ( ! [X6,X7,X5] : (sP8('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X5))),s(fun(sK10,bool),sK11)),X6,X7) | s(bool,f(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f454,f299])).
fof(f627,plain,(
  ( ! [X4,X2,X3] : (sP8('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X2))),s(fun(sK10,bool),sK11)),X3,X4) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f454,f299])).
fof(f616,plain,(
  ( ! [X6,X7,X5] : (sP8(X6,X7,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),X5))) | s(bool,f(s(sK10,sK17(X5,sK11,bool,sK10)))) != s(bool,g(s(sK10,sK17(X5,sK11,bool,sK10))))) )),
  inference(superposition,[],[f453,f299])).
fof(f615,plain,(
  ( ! [X4,X2,X3] : (sP8(X3,X4,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),X2))) | s(bool,f(s(sK10,sK17(sK11,X2,bool,sK10)))) != s(bool,g(s(sK10,sK17(sK11,X2,bool,sK10))))) )),
  inference(superposition,[],[f453,f299])).
fof(f614,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),sK11))),s(fun(sK10,bool),sK11))))) )),
  inference(superposition,[],[f453,f255])).
fof(f613,plain,(
  ( ! [X6,X7,X5] : (sP8(X6,X7,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),X5),s(fun(sK10,bool),sK11))) | s(bool,f(s(sK10,sK17(X5,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10)))) != s(bool,g(s(sK10,sK17(X5,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10))))) )),
  inference(superposition,[],[f453,f299])).
fof(f612,plain,(
  ( ! [X4,X2,X3] : (sP8(X3,X4,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),X2),s(fun(sK10,bool),sK11))) | s(bool,f(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X2,bool,sK10)))) != s(bool,g(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X2,bool,sK10))))) )),
  inference(superposition,[],[f453,f299])).
fof(f611,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.EMPTY_0'),s(fun(sK10,bool),sK11))) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))),s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))))) )),
  inference(superposition,[],[f453,f255])).
fof(f610,plain,(
  ( ! [X6,X7,X5] : (sP8(X6,X7,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X5))),s(fun(sK10,bool),sK11))) | s(bool,f(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f453,f299])).
fof(f609,plain,(
  ( ! [X4,X2,X3] : (sP8(X3,X4,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X2))),s(fun(sK10,bool),sK11))) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f453,f299])).
fof(f595,plain,(
  ( ! [X6,X7,X5] : (sP6('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),X5)),X6,X7) | s(bool,f(s(sK10,sK17(X5,sK11,bool,sK10)))) != s(bool,g(s(sK10,sK17(X5,sK11,bool,sK10))))) )),
  inference(superposition,[],[f439,f299])).
fof(f594,plain,(
  ( ! [X4,X2,X3] : (sP6('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),X2)),X3,X4) | s(bool,f(s(sK10,sK17(sK11,X2,bool,sK10)))) != s(bool,g(s(sK10,sK17(sK11,X2,bool,sK10))))) )),
  inference(superposition,[],[f439,f299])).
fof(f593,plain,(
  ( ! [X0,X1] : (sP6('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),'const.pred_set.EMPTY_0')),X0,X1) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),sK11))),s(fun(sK10,bool),sK11))))) )),
  inference(superposition,[],[f439,f255])).
fof(f592,plain,(
  ( ! [X6,X7,X5] : (sP6('const.pred_set.SUBSET_2'(s(fun(sK10,bool),X5),s(fun(sK10,bool),sK11)),X6,X7) | s(bool,f(s(sK10,sK17(X5,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10)))) != s(bool,g(s(sK10,sK17(X5,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10))))) )),
  inference(superposition,[],[f439,f299])).
fof(f591,plain,(
  ( ! [X4,X2,X3] : (sP6('const.pred_set.SUBSET_2'(s(fun(sK10,bool),X2),s(fun(sK10,bool),sK11)),X3,X4) | s(bool,f(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X2,bool,sK10)))) != s(bool,g(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X2,bool,sK10))))) )),
  inference(superposition,[],[f439,f299])).
fof(f590,plain,(
  ( ! [X0,X1] : (sP6('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.EMPTY_0'),s(fun(sK10,bool),sK11)),X0,X1) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))),s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))))) )),
  inference(superposition,[],[f439,f255])).
fof(f589,plain,(
  ( ! [X6,X7,X5] : (sP6('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X5))),s(fun(sK10,bool),sK11)),X6,X7) | s(bool,f(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f439,f299])).
fof(f588,plain,(
  ( ! [X4,X2,X3] : (sP6('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X2))),s(fun(sK10,bool),sK11)),X3,X4) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f439,f299])).
fof(f579,plain,(
  ( ! [X6,X7,X5] : (sP3('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),X5)),X6,X7) | s(bool,f(s(sK10,sK17(X5,sK11,bool,sK10)))) != s(bool,g(s(sK10,sK17(X5,sK11,bool,sK10))))) )),
  inference(superposition,[],[f425,f299])).
fof(f578,plain,(
  ( ! [X4,X2,X3] : (sP3('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),X2)),X3,X4) | s(bool,f(s(sK10,sK17(sK11,X2,bool,sK10)))) != s(bool,g(s(sK10,sK17(sK11,X2,bool,sK10))))) )),
  inference(superposition,[],[f425,f299])).
fof(f577,plain,(
  ( ! [X0,X1] : (sP3('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),'const.pred_set.EMPTY_0')),X0,X1) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),sK11))),s(fun(sK10,bool),sK11))))) )),
  inference(superposition,[],[f425,f255])).
fof(f576,plain,(
  ( ! [X6,X7,X5] : (sP3('const.pred_set.SUBSET_2'(s(fun(sK10,bool),X5),s(fun(sK10,bool),sK11)),X6,X7) | s(bool,f(s(sK10,sK17(X5,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10)))) != s(bool,g(s(sK10,sK17(X5,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10))))) )),
  inference(superposition,[],[f425,f299])).
fof(f575,plain,(
  ( ! [X4,X2,X3] : (sP3('const.pred_set.SUBSET_2'(s(fun(sK10,bool),X2),s(fun(sK10,bool),sK11)),X3,X4) | s(bool,f(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X2,bool,sK10)))) != s(bool,g(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X2,bool,sK10))))) )),
  inference(superposition,[],[f425,f299])).
fof(f574,plain,(
  ( ! [X0,X1] : (sP3('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.EMPTY_0'),s(fun(sK10,bool),sK11)),X0,X1) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))),s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))))) )),
  inference(superposition,[],[f425,f255])).
fof(f573,plain,(
  ( ! [X6,X7,X5] : (sP3('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X5))),s(fun(sK10,bool),sK11)),X6,X7) | s(bool,f(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f425,f299])).
fof(f572,plain,(
  ( ! [X4,X2,X3] : (sP3('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X2))),s(fun(sK10,bool),sK11)),X3,X4) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f425,f299])).
fof(f561,plain,(
  ( ! [X6,X7,X5] : (sP3(X6,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),X5)),X7) | s(bool,f(s(sK10,sK17(X5,sK11,bool,sK10)))) != s(bool,g(s(sK10,sK17(X5,sK11,bool,sK10))))) )),
  inference(superposition,[],[f424,f299])).
fof(f560,plain,(
  ( ! [X4,X2,X3] : (sP3(X3,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),X2)),X4) | s(bool,f(s(sK10,sK17(sK11,X2,bool,sK10)))) != s(bool,g(s(sK10,sK17(sK11,X2,bool,sK10))))) )),
  inference(superposition,[],[f424,f299])).
fof(f559,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),'const.pred_set.EMPTY_0')),X1) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),sK11))),s(fun(sK10,bool),sK11))))) )),
  inference(superposition,[],[f424,f255])).
fof(f558,plain,(
  ( ! [X6,X7,X5] : (sP3(X6,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),X5),s(fun(sK10,bool),sK11)),X7) | s(bool,f(s(sK10,sK17(X5,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10)))) != s(bool,g(s(sK10,sK17(X5,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10))))) )),
  inference(superposition,[],[f424,f299])).
fof(f557,plain,(
  ( ! [X4,X2,X3] : (sP3(X3,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),X2),s(fun(sK10,bool),sK11)),X4) | s(bool,f(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X2,bool,sK10)))) != s(bool,g(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X2,bool,sK10))))) )),
  inference(superposition,[],[f424,f299])).
fof(f556,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.EMPTY_0'),s(fun(sK10,bool),sK11)),X1) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))),s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))))) )),
  inference(superposition,[],[f424,f255])).
fof(f555,plain,(
  ( ! [X6,X7,X5] : (sP3(X6,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X5))),s(fun(sK10,bool),sK11)),X7) | s(bool,f(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f424,f299])).
fof(f554,plain,(
  ( ! [X4,X2,X3] : (sP3(X3,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X2))),s(fun(sK10,bool),sK11)),X4) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f424,f299])).
fof(f546,plain,(
  ( ! [X6,X7,X5] : (sP4('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),X5)),X6,X7) | s(bool,f(s(sK10,sK17(X5,sK11,bool,sK10)))) != s(bool,g(s(sK10,sK17(X5,sK11,bool,sK10))))) )),
  inference(superposition,[],[f422,f299])).
fof(f545,plain,(
  ( ! [X4,X2,X3] : (sP4('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),X2)),X3,X4) | s(bool,f(s(sK10,sK17(sK11,X2,bool,sK10)))) != s(bool,g(s(sK10,sK17(sK11,X2,bool,sK10))))) )),
  inference(superposition,[],[f422,f299])).
fof(f544,plain,(
  ( ! [X0,X1] : (sP4('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),'const.pred_set.EMPTY_0')),X0,X1) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),sK11))),s(fun(sK10,bool),sK11))))) )),
  inference(superposition,[],[f422,f255])).
fof(f543,plain,(
  ( ! [X6,X7,X5] : (sP4('const.pred_set.SUBSET_2'(s(fun(sK10,bool),X5),s(fun(sK10,bool),sK11)),X6,X7) | s(bool,f(s(sK10,sK17(X5,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10)))) != s(bool,g(s(sK10,sK17(X5,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10))))) )),
  inference(superposition,[],[f422,f299])).
fof(f542,plain,(
  ( ! [X4,X2,X3] : (sP4('const.pred_set.SUBSET_2'(s(fun(sK10,bool),X2),s(fun(sK10,bool),sK11)),X3,X4) | s(bool,f(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X2,bool,sK10)))) != s(bool,g(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X2,bool,sK10))))) )),
  inference(superposition,[],[f422,f299])).
fof(f541,plain,(
  ( ! [X0,X1] : (sP4('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.EMPTY_0'),s(fun(sK10,bool),sK11)),X0,X1) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))),s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))))) )),
  inference(superposition,[],[f422,f255])).
fof(f540,plain,(
  ( ! [X6,X7,X5] : (sP4('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X5))),s(fun(sK10,bool),sK11)),X6,X7) | s(bool,f(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f422,f299])).
fof(f539,plain,(
  ( ! [X4,X2,X3] : (sP4('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X2))),s(fun(sK10,bool),sK11)),X3,X4) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f422,f299])).
fof(f526,plain,(
  ( ! [X6,X7,X5] : (sP4(X6,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),X5)),X7) | s(bool,f(s(sK10,sK17(X5,sK11,bool,sK10)))) != s(bool,g(s(sK10,sK17(X5,sK11,bool,sK10))))) )),
  inference(superposition,[],[f421,f299])).
fof(f525,plain,(
  ( ! [X4,X2,X3] : (sP4(X3,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),X2)),X4) | s(bool,f(s(sK10,sK17(sK11,X2,bool,sK10)))) != s(bool,g(s(sK10,sK17(sK11,X2,bool,sK10))))) )),
  inference(superposition,[],[f421,f299])).
fof(f524,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),'const.pred_set.EMPTY_0')),X1) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),sK11))),s(fun(sK10,bool),sK11))))) )),
  inference(superposition,[],[f421,f255])).
fof(f523,plain,(
  ( ! [X6,X7,X5] : (sP4(X6,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),X5),s(fun(sK10,bool),sK11)),X7) | s(bool,f(s(sK10,sK17(X5,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10)))) != s(bool,g(s(sK10,sK17(X5,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10))))) )),
  inference(superposition,[],[f421,f299])).
fof(f522,plain,(
  ( ! [X4,X2,X3] : (sP4(X3,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),X2),s(fun(sK10,bool),sK11)),X4) | s(bool,f(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X2,bool,sK10)))) != s(bool,g(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X2,bool,sK10))))) )),
  inference(superposition,[],[f421,f299])).
fof(f521,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.EMPTY_0'),s(fun(sK10,bool),sK11)),X1) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))),s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))))) )),
  inference(superposition,[],[f421,f255])).
fof(f520,plain,(
  ( ! [X6,X7,X5] : (sP4(X6,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X5))),s(fun(sK10,bool),sK11)),X7) | s(bool,f(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X5,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f421,f299])).
fof(f519,plain,(
  ( ! [X4,X2,X3] : (sP4(X3,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X2))),s(fun(sK10,bool),sK11)),X4) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X2,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f421,f299])).
fof(f533,plain,(
  ( ! [X6,X7] : (sP5(X6,X7,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11))) | p(s(bool,X6)) | ~sP2('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)),X7,X6)) )),
  inference(subsumption_resolution,[],[f517,f424])).
fof(f517,plain,(
  ( ! [X6,X7] : (sP5(X6,X7,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11))) | p(s(bool,X6)) | ~sP3(X7,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)),X6) | ~sP2('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)),X7,X6)) )),
  inference(resolution,[],[f421,f312])).
fof(f532,plain,(
  ( ! [X4,X5] : (sP5(X4,X5,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11))) | p(s(bool,X5)) | ~sP2('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)),X5,X4)) )),
  inference(subsumption_resolution,[],[f516,f320])).
fof(f516,plain,(
  ( ! [X4,X5] : (sP5(X4,X5,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11))) | p(s(bool,X5)) | ~sP3(X5,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)),X4) | ~sP2('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)),X5,X4)) )),
  inference(resolution,[],[f421,f311])).
fof(f490,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | ~p(s(bool,X11)) | ~sP5(X11,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)),X10)) )),
  inference(superposition,[],[f214,f327])).
fof(f488,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | sP5(X7,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)),X6) | p(s(bool,X7))) )),
  inference(superposition,[],[f214,f325])).
fof(f487,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | sP5(X5,X4,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11))) | p(s(bool,X5))) )),
  inference(superposition,[],[f214,f325])).
fof(f476,plain,(
  ( ! [X1] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),X1)))) | s(bool,f(s(sK10,sK17(X1,sK11,bool,sK10)))) != s(bool,g(s(sK10,sK17(X1,sK11,bool,sK10))))) )),
  inference(superposition,[],[f214,f299])).
fof(f475,plain,(
  ( ! [X0] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),X0)))) | s(bool,f(s(sK10,sK17(sK11,X0,bool,sK10)))) != s(bool,g(s(sK10,sK17(sK11,X0,bool,sK10))))) )),
  inference(superposition,[],[f214,f299])).
fof(f474,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),sK11))),s(fun(sK10,bool),sK11))))),
  inference(superposition,[],[f214,f255])).
fof(f473,plain,(
  ( ! [X1] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),X1),s(fun(sK10,bool),sK11)))) | s(bool,f(s(sK10,sK17(X1,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10)))) != s(bool,g(s(sK10,sK17(X1,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),bool,sK10))))) )),
  inference(superposition,[],[f214,f299])).
fof(f472,plain,(
  ( ! [X0] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),X0),s(fun(sK10,bool),sK11)))) | s(bool,f(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X0,bool,sK10)))) != s(bool,g(s(sK10,sK17('const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),X0,bool,sK10))))) )),
  inference(superposition,[],[f214,f299])).
fof(f471,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.EMPTY_0'),s(fun(sK10,bool),sK11)))) | p(s(bool,'const.bool.IN_2'(s(sK10,'const.pred_set.CHOICE_1'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))),s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))))),
  inference(superposition,[],[f214,f255])).
fof(f470,plain,(
  ( ! [X1] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X1))),s(fun(sK10,bool),sK11)))) | s(bool,f(s(fun(sK10,bool),sK17(X1,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X1,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f214,f299])).
fof(f469,plain,(
  ( ! [X0] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),X0))),s(fun(sK10,bool),sK11)))) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X0,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X0,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f214,f299])).
fof(f458,plain,(
  ( ! [X103,X102] : (~p(s(bool,X102)) | ~p(s(bool,X103)) | ~sP9(X103,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)),X102)) )),
  inference(resolution,[],[f214,f360])).
fof(f457,plain,(
  ( ! [X101,X100] : (sP9(X100,X101,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11))) | ~p(s(bool,X100))) )),
  inference(resolution,[],[f214,f358])).
fof(f456,plain,(
  ( ! [X99,X98] : (sP9('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)),X98,X99) | p(s(bool,X98)) | ~p(s(bool,X99))) )),
  inference(resolution,[],[f214,f357])).
fof(f455,plain,(
  ( ! [X97,X96] : (sP9(X96,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)),X97) | ~p(s(bool,X97)) | p(s(bool,X96))) )),
  inference(resolution,[],[f214,f357])).
fof(f454,plain,(
  ( ! [X94,X95] : (sP8('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)),X94,X95)) )),
  inference(resolution,[],[f214,f356])).
fof(f453,plain,(
  ( ! [X92,X93] : (sP8(X92,X93,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)))) )),
  inference(resolution,[],[f214,f354])).
fof(f452,plain,(
  ( ! [X90,X91] : (~p(s(bool,X90)) | ~p(s(bool,X91)) | ~sP8(X90,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)),X91)) )),
  inference(resolution,[],[f214,f353])).
fof(f451,plain,(
  ( ! [X88,X89] : (sP9(X88,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)),X89) | ~sP8(X88,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)),X89) | ~p(s(bool,X89))) )),
  inference(resolution,[],[f214,f352])).
fof(f449,plain,(
  ( ! [X85,X84] : (~sP9('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)),X84,X85) | ~p(s(bool,X84))) )),
  inference(resolution,[],[f214,f347])).
fof(f448,plain,(
  ( ! [X83,X82] : (~sP9(X82,X83,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11))) | p(s(bool,X82))) )),
  inference(resolution,[],[f214,f346])).
fof(f447,plain,(
  ( ! [X80,X81] : (~sP9('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)),X81,X80) | p(s(bool,X80))) )),
  inference(resolution,[],[f214,f346])).
fof(f444,plain,(
  ( ! [X74,X75] : (p(s(bool,X74)) | ~p(s(bool,X75)) | ~sP7(X75,X74,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)))) )),
  inference(resolution,[],[f214,f343])).
fof(f443,plain,(
  ( ! [X72,X73] : (p(s(bool,X72)) | ~p(s(bool,X73)) | ~sP7(X73,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)),X72)) )),
  inference(resolution,[],[f214,f343])).
fof(f442,plain,(
  ( ! [X70,X71] : (sP7('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)),X70,X71) | p(s(bool,X70)) | p(s(bool,X71))) )),
  inference(resolution,[],[f214,f340])).
fof(f441,plain,(
  ( ! [X68,X69] : (sP7(X68,X69,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11))) | p(s(bool,X69)) | p(s(bool,X68))) )),
  inference(resolution,[],[f214,f340])).
fof(f440,plain,(
  ( ! [X66,X67] : (sP7(X66,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)),X67) | p(s(bool,X67)) | p(s(bool,X66))) )),
  inference(resolution,[],[f214,f340])).
fof(f439,plain,(
  ( ! [X64,X65] : (sP6('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)),X64,X65)) )),
  inference(resolution,[],[f214,f339])).
fof(f438,plain,(
  ( ! [X62,X63] : (~p(s(bool,X62)) | p(s(bool,X63)) | ~sP6(X62,X63,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)))) )),
  inference(resolution,[],[f214,f336])).
fof(f437,plain,(
  ( ! [X61,X60] : (~p(s(bool,X60)) | p(s(bool,X61)) | ~sP6(X60,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)),X61)) )),
  inference(resolution,[],[f214,f336])).
fof(f436,plain,(
  ( ! [X59,X58] : (sP7(X58,X59,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11))) | ~sP6(X58,X59,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11))) | p(s(bool,X59))) )),
  inference(resolution,[],[f214,f335])).
fof(f435,plain,(
  ( ! [X57,X56] : (sP7(X56,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)),X57) | ~sP6(X56,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)),X57) | p(s(bool,X57))) )),
  inference(resolution,[],[f214,f335])).
fof(f432,plain,(
  ( ! [X50,X51] : (~sP7('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)),X50,X51) | ~p(s(bool,X50))) )),
  inference(resolution,[],[f214,f330])).
fof(f431,plain,(
  ( ! [X48,X49] : (~sP7('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)),X49,X48) | ~p(s(bool,X48))) )),
  inference(resolution,[],[f214,f329])).
fof(f430,plain,(
  ( ! [X47,X46] : (~sP5('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)),X47,X46) | s(bool,X46) != s(bool,X47)) )),
  inference(resolution,[],[f214,f328])).
fof(f429,plain,(
  ( ! [X45,X44] : (sP5('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)),X44,X45) | s(bool,X44) = s(bool,X45)) )),
  inference(resolution,[],[f214,f325])).
fof(f428,plain,(
  ( ! [X43,X42] : (p(s(bool,X42)) | p(s(bool,X43)) | ~sP2(X42,X43,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)))) )),
  inference(resolution,[],[f214,f321])).
fof(f427,plain,(
  ( ! [X41,X40] : (p(s(bool,X40)) | p(s(bool,X41)) | ~sP2(X40,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)),X41)) )),
  inference(resolution,[],[f214,f321])).
fof(f426,plain,(
  ( ! [X39,X38] : (p(s(bool,X38)) | p(s(bool,X39)) | ~sP2('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)),X38,X39)) )),
  inference(resolution,[],[f214,f321])).
fof(f501,plain,(
  ( ! [X23,X22] : (sP5(X22,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)),X23) | ~sP4(X22,X23,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11))) | ~sP2(X23,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)),X22)) )),
  inference(subsumption_resolution,[],[f418,f425])).
fof(f425,plain,(
  ( ! [X37,X36] : (sP3('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)),X36,X37)) )),
  inference(resolution,[],[f214,f320])).
fof(f424,plain,(
  ( ! [X35,X34] : (sP3(X34,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)),X35)) )),
  inference(resolution,[],[f214,f319])).
fof(f423,plain,(
  ( ! [X33,X32] : (~p(s(bool,X32)) | ~p(s(bool,X33)) | ~sP3(X32,X33,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)))) )),
  inference(resolution,[],[f214,f317])).
fof(f500,plain,(
  ( ! [X24,X25] : (sP5('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)),X24,X25) | ~sP3(X24,X25,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11))) | ~sP2(X25,X24,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)))) )),
  inference(subsumption_resolution,[],[f419,f422])).
fof(f422,plain,(
  ( ! [X30,X31] : (sP4('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)),X30,X31)) )),
  inference(resolution,[],[f214,f316])).
fof(f421,plain,(
  ( ! [X28,X29] : (sP4(X28,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)),X29)) )),
  inference(resolution,[],[f214,f315])).
fof(f420,plain,(
  ( ! [X26,X27] : (~p(s(bool,X26)) | ~p(s(bool,X27)) | ~sP4(X26,X27,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)))) )),
  inference(resolution,[],[f214,f313])).
fof(f419,plain,(
  ( ! [X24,X25] : (sP5('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)),X24,X25) | ~sP4('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)),X25,X24) | ~sP3(X24,X25,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11))) | ~sP2(X25,X24,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)))) )),
  inference(resolution,[],[f214,f312])).
fof(f418,plain,(
  ( ! [X23,X22] : (sP5(X22,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)),X23) | ~sP4(X22,X23,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11))) | ~sP3('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)),X23,X22) | ~sP2(X23,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)),X22)) )),
  inference(resolution,[],[f214,f311])).
fof(f417,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | ~p(s(bool,X21)) | ~sP5(X20,X21,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)))) )),
  inference(resolution,[],[f214,f309])).
fof(f412,plain,(
  ( ! [X15] : (sP1('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)),X15) | ~p(s(bool,X15))) )),
  inference(resolution,[],[f214,f294])).
fof(f411,plain,(
  ( ! [X14] : (sP1(X14,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11))) | ~p(s(bool,X14))) )),
  inference(resolution,[],[f214,f291])).
fof(f410,plain,(
  ( ! [X13] : (~sP1(X13,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11))) | p(s(bool,X13))) )),
  inference(resolution,[],[f214,f289])).
fof(f409,plain,(
  ( ! [X12] : (~sP1('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)),X12) | p(s(bool,X12))) )),
  inference(resolution,[],[f214,f289])).
fof(f499,plain,(
  ( ! [X9] : (~sP0(X9,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11))) | p(s(bool,X9))) )),
  inference(subsumption_resolution,[],[f406,f221])).
fof(f406,plain,(
  ( ! [X9] : (p(s(bool,'F_0')) | p(s(bool,X9)) | ~sP0(X9,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)))) )),
  inference(resolution,[],[f214,f274])).
fof(f498,plain,(
  ( ! [X8] : (~sP0('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)),X8) | p(s(bool,X8))) )),
  inference(subsumption_resolution,[],[f405,f221])).
fof(f405,plain,(
  ( ! [X8] : (p(s(bool,'F_0')) | p(s(bool,X8)) | ~sP0('const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)),X8)) )),
  inference(resolution,[],[f214,f274])).
fof(f402,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,happ(s(fun(X3,bool),X4),s(X3,X5)))) | ~p(s(bool,happ(s(fun(X3,bool),X4),s(X3,sK16(X4,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)),X3)))))) )),
  inference(resolution,[],[f214,f266])).
fof(f401,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK15(X1,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)),X0)))))) )),
  inference(resolution,[],[f214,f264])).
fof(f398,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)))),
  inference(resolution,[],[f214,f233])).
fof(f396,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11)))),
  inference(resolution,[],[f214,f228])).
fof(f395,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),sK11))))),
  inference(resolution,[],[f214,f260])).
fof(f260,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,sK13(X2,X1,X0)),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f161])).
fof(f394,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK10,sK13(sK11,'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12)),sK10)),s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))))))),
  inference(resolution,[],[f214,f259])).
fof(f259,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK13(X2,X1,X0)),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f161])).
fof(f392,plain,(
  ( ! [X1] : (s(fun(fun(sK10,bool),bool),'const.pred_set.EMPTY_0') != s(fun(fun(sK10,bool),bool),X1) | s(bool,f(s(fun(sK10,bool),sK17(X1,sK12,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(X1,sK12,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f213,f299])).
fof(f391,plain,(
  ( ! [X0] : (s(fun(fun(sK10,bool),bool),'const.pred_set.EMPTY_0') != s(fun(fun(sK10,bool),bool),X0) | s(bool,f(s(fun(sK10,bool),sK17(sK12,X0,bool,fun(sK10,bool))))) != s(bool,g(s(fun(sK10,bool),sK17(sK12,X0,bool,fun(sK10,bool)))))) )),
  inference(superposition,[],[f213,f299])).
fof(f393,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12))))),
  inference(trivial_inequality_removal,[],[f390])).
fof(f390,plain,(
  s(fun(fun(sK10,bool),bool),'const.pred_set.EMPTY_0') != s(fun(fun(sK10,bool),bool),'const.pred_set.EMPTY_0') | p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),'const.pred_set.CHOICE_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(fun(sK10,bool),bool),sK12))))),
  inference(superposition,[],[f213,f255])).
fof(f212,plain,(
  ( ! [X3] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),X3),s(fun(sK10,bool),sK11)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),X3),s(fun(fun(sK10,bool),bool),sK12))))) )),
  inference(cnf_transformation,[],[f139])).
fof(f214,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),'const.pred_set.BIGINTER_1'(s(fun(fun(sK10,bool),bool),sK12))),s(fun(sK10,bool),sK11))))),
  inference(cnf_transformation,[],[f139])).
% SZS output end Proof for thm.util_prob
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 261
% Passive clauses: 1052
% Generated clauses: 4734
% Final active clauses: 250
% Final passive clauses: 738
% Input formulas: 37
% Initial clauses: 159
% 
% Duplicate literals: 179
% Trivial inequalities: 76
% Fw subsumption resolutions: 364
% Bw subsumption resolutions: 10
% Fw demodulations: 510
% Bw demodulations: 5
% Condensations: 218
% Global subsumptions: 1
% 
% Simple tautologies: 376
% Forward subsumptions: 1925
% Backward subsumptions: 164
% 
% Binary resolution: 1349
% Forward superposition: 1228
% Backward superposition: 800
% Equality resolution: 13
% 
% SAT solver clauses: 1055
% SAT solver unit clauses: 63
% SAT solver binary clauses: 379
% 
% Memory used [KB]: 3965
% Time elapsed: 0.195 s
% ------------------------------
% ------------------------------
% Success in time 16.623 s
