Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.pred_set
% (85546)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 170
% Passive clauses: 1715
% Generated clauses: 30399
% Final active clauses: 141
% Final passive clauses: 1405
% Input formulas: 13
% Initial clauses: 20
% 
% Duplicate literals: 1794
% Trivial inequalities: 197
% Fw subsumption resolutions: 3672
% Fw demodulations: 5909
% Bw demodulations: 26
% 
% Simple tautologies: 270
% Equational tautologies: 66
% Forward subsumptions: 14554
% 
% Binary resolution: 331
% Factoring: 1
% Forward superposition: 8872
% Backward superposition: 8944
% Self superposition: 255
% Equality factoring: 174
% Equality resolution: 7
% 
% Split clauses: 1543
% Split components: 3087
% SAT solver clauses: 2348
% SAT solver unit clauses: 17
% SAT solver binary clauses: 226
% 
% Sat splits: 1543
% Sat splitting refutations: 6
% 
% TWLsolver clauses: 2348
% TWLsolver calls for satisfiability: 123
% 
% Memory used [KB]: 14711
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 294 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.pred_set
% (85558)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 143
% Passive clauses: 453
% Generated clauses: 12144
% Final active clauses: 111
% Final passive clauses: 251
% Input formulas: 13
% Initial clauses: 20
% 
% Duplicate literals: 1438
% Trivial inequalities: 111
% Fw subsumption resolutions: 756
% Fw demodulations: 910
% Bw demodulations: 86
% Condensations: 1925
% Global subsumptions: 1
% 
% Simple tautologies: 147
% Equational tautologies: 52
% Forward subsumptions: 4657
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 164
% Factoring: 5
% Forward superposition: 2852
% Backward superposition: 3148
% Self superposition: 167
% Equality factoring: 152
% Equality resolution: 5
% 
% Split clauses: 263
% Split components: 531
% SAT solver clauses: 1361
% SAT solver unit clauses: 29
% SAT solver binary clauses: 71
% SAT solver learnt clauses: 2
% 
% Sat splits: 263
% Sat splitting refutations: 13
% 
% TWLsolver clauses: 1357
% TWLsolver calls for satisfiability: 166
% 
% Memory used [KB]: 8571
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.pred_set
% (85585)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 497
% Passive clauses: 3794
% Generated clauses: 173440
% Final active clauses: 68
% Final passive clauses: 201
% Input formulas: 13
% Initial clauses: 20
% 
% Duplicate literals: 8771
% Trivial inequalities: 1167
% Fw subsumption resolutions: 15321
% Fw demodulations: 22357
% Bw demodulations: 2195
% Condensations: 16357
% Global subsumptions: 2163
% 
% Simple tautologies: 1356
% Equational tautologies: 2349
% Forward subsumptions: 93713
% Fw demodulations to eq. taut.: 10
% Bw demodulations to eq. taut.: 107
% 
% Unit resulting resolution: 9562
% Forward superposition: 40188
% Backward superposition: 31239
% Self superposition: 727
% Equality factoring: 526
% Equality resolution: 55
% 
% Split clauses: 748
% Split components: 1510
% SAT solver clauses: 9811
% SAT solver unit clauses: 100
% SAT solver binary clauses: 963
% 
% Sat splits: 748
% Sat splitting refutations: 2356
% 
% Memory used [KB]: 67418
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f17692,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f121,f138,f137,f749,f750,f766,f1121,f1139,f1138,f1424,f1438,f1444,f1450,f1426,f1476,f1480,f1430,f1504,f1506,f1440,f1510,f1512,f1446,f1528,f1530,f1452,f1482,f1534,f1478,f1428,f1536,f1539,f1540,f1833,f1945,f1834,f1946,f1944,f2516,f2530,f2536,f2542,f2568,f2572,f2522,f2596,f2598,f2532,f2602,f2604,f2538,f2620,f2622,f2544,f2626,f2628,f2912,f2926,f2932,f2938,f2914,f2964,f2968,f2918,f2992,f2994,f2928,f2998,f3000,f3016,f3018,f2940,f2970,f3022,f2916,f3024,f3282,f3376,f3283,f3415,f3431,f3380,f3432,f3433,f3444,f770,f3445,f3520,f772,f3536,f3545,f3563,f3564,f1558,f3812,f1554,f1556,f3925,f3944,f3941,f764,f4353,f4374,f4372,f4354,f4378,f4376,f4387,f4393,f4399,f4403,f4407,f1788,f4409,f4410,f4411,f4414,f4415,f4416,f1113,f1114,f1115,f669,f3534,f635,f440,f4383,f2570,f4348,f4423,f4270,f4426,f4425,f4173,f4429,f4064,f2574,f4430,f4420,f2520,f2518,f4431,f4434,f4437,f4440,f4443,f4446,f4449,f4452,f4455,f4458,f4461,f4464,f4467,f151,f4469,f4473,f136,f4477,f3535,f4479,f4413,f4389,f135,f5126,f5249,f5127,f5250,f5248,f5344,f5358,f5388,f5498,f5389,f5499,f5497,f5881,f5882,f5883,f5901,f5884,f5915,f5917,f3456,f5921,f768,f3455,f5932,f5933,f5968,f5972,f3457,f5982,f6232,f6248,f6233,f6250,f6234,f6252,f6235,f6439,f6627,f6453,f6658,f6628,f6661,f6455,f6254,f8115,f8306,f8819,f8854,f8853,f8852,f11756,f11779,f11778,f12494,f12499,f12493,f12500,f12609,f12614,f12608,f12615,f12832,f2917,f2934,f12980,f12983,f12985,f4412,f4388,f13621,f13637,f13649,f13652,f13655,f13657,f13660,f13662,f13667,f13669,f13672,f12591,f12342,f13681,f13690,f10600,f13694,f4421,f13697,f13710,f13720,f13721,f9086,f13728,f13729,f13755,f13760,f8522,f4424,f13783,f13843,f13870,f13873,f13875,f13837,f13832,f13878,f13881,f13815,f13807,f13802,f13782,f13886,f13773,f13768,f13762,f13888,f13754,f13750,f13891,f13739,f13734,f5357,f13723,f13894,f13896,f13899,f13696,f4428,f13699,f13901,f13903,f4422,f12476,f12477,f12207,f12208,f13952,f13955,f13960,f9085,f13961,f9052,f13966,f13973,f13984,f8520,f13985,f8488,f8378,f8375,f8374,f8358,f8343,f8342,f8311,f8310,f13988,f13991,f13994,f14299,f14300,f14301,f14302,f14643,f14646,f12427,f14644,f12426,f14647,f14661,f14662,f14663,f14664,f14665,f14666,f14667,f14668,f9898,f14669,f14671,f14674,f14676,f14681,f14684,f14686,f9566,f14690,f14691,f14692,f14693,f14694,f7533,f14695,f14696,f14697,f7447,f14698,f14699,f7317,f14700,f7275,f5983,f5958,f5956,f774,f14614,f14607,f14592,f14584,f13968,f14536,f14533,f14528,f13883,f5897,f14525,f8117,f14881,f776,f14882,f14893,f3443,f15414,f15415,f15416,f5899,f15417,f5895,f15420,f13670,f15421,f13673,f15422,f15423,f8119,f15440,f15441,f5900,f15452,f5898,f15486,f15506,f15489,f15507,f15490,f15508,f15491,f5903,f13664,f15539,f15937,f15986,f15987,f16024,f16005,f16025,f16001,f16121,f16143,f16103,f15967,f15965,f16145,f16147,f16203,f16204,f16225,f16241,f16239,f16253,f16224,f16222,f16254,f16220,f16218,f16308,f16309,f16310,f16313,f16003,f16365,f16332,f16368,f16333,f16369,f16334,f16371,f16007,f16141,f16662,f16664,f16665,f16661,f16663,f16794,f16796,f16797,f16793,f16795,f17068,f17080,f17067,f17081,f17156,f17158,f17159,f17155,f17157,f17211,f17213,f3378,f17214,f17210,f17212,f4395,f17447,f17449,f17452,f17466,f17467,f17468,f17496,f17469,f17470,f17498,f17471,f17495,f17472,f17499,f17473,f17501,f17474,f17502,f17475,f17503,f17478,f17479,f17505,f17480,f17507,f17481,f2966,f17482,f17465,f17483,f17451,f17510,f17521,f17537,f17522,f17494,f17500,f17525,f17509,f17539,f17691])).
fof(f17691,plain,(
  $false | $spl118),
  inference(subsumption_resolution,[],[f17682,f4412])).
fof(f17682,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | $spl118),
  inference(resolution,[],[f17245,f52])).
fof(f52,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f40])).
fof(f40,plain,(
  ! [X0,X1] : ((p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) & (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(nnf_transformation,[],[f21])).
fof(f21,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(flattening,[],[f20])).
fof(f20,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X9,X10] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num',X10)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X10),s('type.num.num',X9)))))),
  file('thm.pred_set.LESS_CARD_DIFF','thm.arithmetic.NOT_LESS')).
fof(f17245,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | $spl118),
  inference(resolution,[],[f4412,f54])).
fof(f54,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f42])).
fof(f42,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num',X0) = s('type.num.num',X1)) & ((~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) & s('type.num.num',X0) != s('type.num.num',X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(flattening,[],[f41])).
fof(f41,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num',X0) = s('type.num.num',X1))) & ((~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) & s('type.num.num',X0) != s('type.num.num',X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(nnf_transformation,[],[f22])).
fof(f22,plain,(
  ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num',X0) = s('type.num.num',X1)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X9,X10] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X9),s('type.num.num',X10)))) <=> (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num',X10)))) | s('type.num.num',X9) = s('type.num.num',X10)))),
  file('thm.pred_set.LESS_CARD_DIFF','thm.arithmetic.LESS_OR_EQ')).
fof(f17539,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | $spl303),
  inference(cnf_transformation,[],[f17539_D])).
fof(f17539_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) <=> ~$spl303),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl303])])).
fof(f17509,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))) | $spl298),
  inference(cnf_transformation,[],[f17509_D])).
fof(f17509_D,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))) <=> ~$spl298),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl298])])).
fof(f17525,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | $spl274),
  inference(resolution,[],[f17484,f53])).
fof(f53,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) = s('type.num.num',X1) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f42])).
fof(f17484,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | $spl274),
  inference(resolution,[],[f17448,f52])).
fof(f17448,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | $spl274),
  inference(cnf_transformation,[],[f17448_D])).
fof(f17448_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) <=> ~$spl274),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl274])])).
fof(f17500,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl286),
  inference(cnf_transformation,[],[f17500_D])).
fof(f17500_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) <=> ~$spl286),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl286])])).
fof(f17494,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | $spl280),
  inference(cnf_transformation,[],[f17494_D])).
fof(f17494_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) <=> ~$spl280),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl280])])).
fof(f17522,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl274),
  inference(resolution,[],[f17484,f320])).
fof(f320,plain,(
  ( ! [X2] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X2)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))))) )),
  inference(resolution,[],[f78,f51])).
fof(f51,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f40])).
fof(f78,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))))))) )),
  inference(superposition,[],[f47,f53])).
fof(f47,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))))))),
  inference(cnf_transformation,[],[f39])).
fof(f39,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))) & (p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK2)))) & p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) & ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f34])).
fof(f34,plain,(
  ? [X0,X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) & ? [X2] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))) & p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X2)))))) & ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))))))))))),
  inference(flattening,[],[f33])).
fof(f33,plain,(
  ? [X0,X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) & ? [X2] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))) & (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X2)))))) & ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))))))))))),
  inference(ennf_transformation,[],[f16])).
fof(f16,plain,(
  ~! [X0,X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) => ! [X2] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))) => (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X2)))))) => p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))))))))))),
  inference(flattening,[],[f15])).
fof(f15,plain,(
  ~! [X0] : ! [X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) => ! [X2] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))) => (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X2)))))) => p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))))))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X7)))) => ! [X8] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X8)))) => (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X6,bool),X7))),s('type.num.num','const.pred_set.CARD_1'(s(fun(X6,bool),X8)))))) => p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(X6,bool),'const.pred_set.DIFF_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X7)))))))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X7)))) => ! [X8] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X8)))) => (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X6,bool),X7))),s('type.num.num','const.pred_set.CARD_1'(s(fun(X6,bool),X8)))))) => p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(X6,bool),'const.pred_set.DIFF_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X7)))))))))))),
  file('thm.pred_set.LESS_CARD_DIFF','thm.pred_set.LESS_CARD_DIFF')).
fof(f17537,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))) | $spl301),
  inference(cnf_transformation,[],[f17537_D])).
fof(f17537_D,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))) <=> ~$spl301),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl301])])).
fof(f17521,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))) | $spl274),
  inference(resolution,[],[f17484,f319])).
fof(f319,plain,(
  ( ! [X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | s('type.num.num',X1) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))) )),
  inference(resolution,[],[f78,f55])).
fof(f55,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num',X0) != s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f42])).
fof(f17510,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | $spl276),
  inference(resolution,[],[f17451,f890])).
fof(f890,plain,(
  ( ! [X2] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num',X2)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X2))))) )),
  inference(resolution,[],[f119,f51])).
fof(f119,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num',X0))))) )),
  inference(superposition,[],[f85,f53])).
fof(f85,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))))),
  inference(resolution,[],[f74,f52])).
fof(f74,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))))))),
  inference(resolution,[],[f47,f54])).
fof(f17451,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | $spl276),
  inference(cnf_transformation,[],[f17451_D])).
fof(f17451_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) <=> ~$spl276),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl276])])).
fof(f17483,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | $spl274),
  inference(resolution,[],[f17448,f53])).
fof(f17465,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | $spl279),
  inference(cnf_transformation,[],[f17465_D])).
fof(f17465_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) <=> ~$spl279),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl279])])).
fof(f17482,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | ($spl17 | $spl274)),
  inference(resolution,[],[f17448,f16160])).
fof(f16160,plain,(
  ( ! [X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num',X1)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1))))) ) | $spl17),
  inference(superposition,[],[f773,f53])).
fof(f773,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | $spl17),
  inference(cnf_transformation,[],[f773_D])).
fof(f773_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f2966,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.num.|48|_0')))) | $spl84),
  inference(cnf_transformation,[],[f2966_D])).
fof(f2966_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.num.|48|_0')))) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f17481,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.num.|48|_0')))) | ($spl18 | $spl274)),
  inference(resolution,[],[f17448,f14902])).
fof(f14902,plain,(
  ( ! [X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.num.|48|_0'))))) ) | $spl18),
  inference(superposition,[],[f776,f53])).
fof(f17507,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | $spl297),
  inference(cnf_transformation,[],[f17507_D])).
fof(f17507_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) <=> ~$spl297),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl297])])).
fof(f17480,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | ($spl172 | $spl274)),
  inference(resolution,[],[f17448,f8899])).
fof(f8899,plain,(
  ( ! [X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num',X1)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X1))))) ) | $spl172),
  inference(superposition,[],[f8881,f53])).
fof(f8881,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | $spl172),
  inference(resolution,[],[f8852,f52])).
fof(f17505,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | $spl294),
  inference(cnf_transformation,[],[f17505_D])).
fof(f17505_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) <=> ~$spl294),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl294])])).
fof(f17479,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | ($spl172 | $spl274)),
  inference(resolution,[],[f17448,f8885])).
fof(f8885,plain,(
  ( ! [X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X1))))) ) | $spl172),
  inference(superposition,[],[f8852,f53])).
fof(f17478,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | $spl274),
  inference(resolution,[],[f17448,f701])).
fof(f701,plain,(
  ( ! [X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X1))))) )),
  inference(resolution,[],[f320,f52])).
fof(f17503,plain,(
  ( ! [X2] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num',X2)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num',X2)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X2))))) ) | $spl292),
  inference(cnf_transformation,[],[f17503_D])).
fof(f17503_D,plain,(
  ( ! [X2] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num',X2)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num',X2)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X2))))) ) <=> ~$spl292),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl292])])).
fof(f17475,plain,(
  ( ! [X2] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num',X2)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X2)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num',X2))))) ) | $spl274),
  inference(resolution,[],[f17448,f425])).
fof(f425,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num',X2)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num',X3)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X3)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num',X2)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X3))))) )),
  inference(superposition,[],[f83,f53])).
fof(f83,plain,(
  ( ! [X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X1))))) )),
  inference(superposition,[],[f47,f53])).
fof(f17502,plain,(
  ( ! [X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X1))))) ) | $spl290),
  inference(cnf_transformation,[],[f17502_D])).
fof(f17502_D,plain,(
  ( ! [X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X1))))) ) <=> ~$spl290),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl290])])).
fof(f17474,plain,(
  ( ! [X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))))) ) | $spl274),
  inference(resolution,[],[f17448,f384])).
fof(f384,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num',X2)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X3)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num',X2)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X3),s('type.num.num',X2))))) )),
  inference(superposition,[],[f82,f53])).
fof(f82,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X0))))) )),
  inference(superposition,[],[f47,f53])).
fof(f17501,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X0))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))))) ) | $spl288),
  inference(cnf_transformation,[],[f17501_D])).
fof(f17501_D,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X0))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))))) ) <=> ~$spl288),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl288])])).
fof(f17473,plain,(
  ( ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X0))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))))) ) | $spl274),
  inference(resolution,[],[f17448,f217])).
fof(f217,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num',X2)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num',X2)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X3))))),s('type.num.num',X2))))) )),
  inference(superposition,[],[f69,f53])).
fof(f69,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X0))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))))))) )),
  inference(resolution,[],[f46,f56])).
fof(f56,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1)))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f36])).
fof(f36,plain,(
  ! [X0,X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) | ! [X2] : p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1)))))))),
  inference(ennf_transformation,[],[f24])).
fof(f24,plain,(
  ! [X0,X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) => ! [X2] : p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1)))))))),
  inference(flattening,[],[f23])).
fof(f23,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) => ! [X2] : p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1)))))))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X6] : ! [X8] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X8)))) => ! [X7] : p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X7))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(X6,bool),X8)))))))),
  file('thm.pred_set.LESS_CARD_DIFF','thm.pred_set.CARD_INTER_LESS_EQ')).
fof(f46,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK2))))),
  inference(cnf_transformation,[],[f39])).
fof(f17499,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl287),
  inference(cnf_transformation,[],[f17499_D])).
fof(f17499_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) <=> ~$spl287),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl287])])).
fof(f17472,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl274),
  inference(resolution,[],[f17448,f116])).
fof(f116,plain,(
  ( ! [X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num',X1)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))))) )),
  inference(superposition,[],[f85,f53])).
fof(f17495,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | $spl283),
  inference(cnf_transformation,[],[f17495_D])).
fof(f17495_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) <=> ~$spl283),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl283])])).
fof(f17471,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | $spl274),
  inference(resolution,[],[f17448,f105])).
fof(f105,plain,(
  ( ! [X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X1))))) )),
  inference(superposition,[],[f75,f53])).
fof(f75,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))))),
  inference(resolution,[],[f47,f52])).
fof(f17498,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl285),
  inference(cnf_transformation,[],[f17498_D])).
fof(f17498_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) <=> ~$spl285),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl285])])).
fof(f17470,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl274),
  inference(resolution,[],[f17448,f104])).
fof(f104,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num',X0)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))))) )),
  inference(superposition,[],[f75,f53])).
fof(f17469,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl274),
  inference(resolution,[],[f17448,f101])).
fof(f101,plain,(
  ( ! [X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))))) )),
  inference(superposition,[],[f75,f53])).
fof(f17496,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | $spl282),
  inference(cnf_transformation,[],[f17496_D])).
fof(f17496_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) <=> ~$spl282),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl282])])).
fof(f17468,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | $spl274),
  inference(resolution,[],[f17448,f93])).
fof(f93,plain,(
  ( ! [X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X1))))) )),
  inference(superposition,[],[f74,f53])).
fof(f17467,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | $spl274),
  inference(resolution,[],[f17448,f83])).
fof(f17466,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | $spl278),
  inference(cnf_transformation,[],[f17466_D])).
fof(f17466_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) <=> ~$spl278),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl278])])).
fof(f17452,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.num.|48|_0')))) | ($spl18 | $spl275)),
  inference(resolution,[],[f17449,f16591])).
fof(f16591,plain,(
  ( ! [X3] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num',X3)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X3)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.num.|48|_0'))))) ) | $spl18),
  inference(resolution,[],[f15650,f54])).
fof(f15650,plain,(
  ( ! [X2] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num',X2)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X2))))) ) | $spl18),
  inference(resolution,[],[f14903,f51])).
fof(f14903,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num',X0))))) ) | $spl18),
  inference(superposition,[],[f776,f53])).
fof(f17449,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | $spl275),
  inference(cnf_transformation,[],[f17449_D])).
fof(f17449_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) <=> ~$spl275),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl275])])).
fof(f17447,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | ($spl17 | $spl74)),
  inference(subsumption_resolution,[],[f17446,f2916])).
fof(f17446,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | $spl17),
  inference(forward_demodulation,[],[f17445,f50])).
fof(f50,plain,(
  ( ! [X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f19])).
fof(f19,plain,(
  ! [X0,X1] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1)))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X10,X9] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num',X10))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X10),s('type.num.num',X9)))))),
  file('thm.pred_set.LESS_CARD_DIFF','thm.arithmetic.SUB_LESS_|48|')).
fof(f17445,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))))) | $spl17),
  inference(forward_demodulation,[],[f17435,f957])).
fof(f957,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))),
  inference(resolution,[],[f70,f46])).
fof(f70,plain,(
  ( ! [X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X1)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),X1),s(fun(sK0,bool),sK2))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),X1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X1),s(fun(sK0,bool),sK2)))))))) )),
  inference(resolution,[],[f46,f49])).
fof(f49,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f35])).
fof(f35,plain,(
  ! [X0,X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) | ! [X2] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))))))))),
  inference(ennf_transformation,[],[f18])).
fof(f18,plain,(
  ! [X0,X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) => ! [X2] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))) => s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))))))))),
  inference(flattening,[],[f17])).
fof(f17,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) => ! [X2] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))) => s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))))))))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X6] : ! [X7] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X7)))) => ! [X8] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X8)))) => s('type.num.num','const.pred_set.CARD_1'(s(fun(X6,bool),'const.pred_set.DIFF_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X7))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X6,bool),X8))),s('type.num.num','const.pred_set.CARD_1'(s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X7)))))))))),
  file('thm.pred_set.LESS_CARD_DIFF','thm.pred_set.CARD_DIFF')).
fof(f17435,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))))) | $spl17),
  inference(superposition,[],[f16160,f957])).
fof(f4395,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | $spl116),
  inference(cnf_transformation,[],[f4395_D])).
fof(f4395_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f17212,plain,(
  ( ! [X10] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num',X10)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X10),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X10),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num',X10))))) ) | $spl272),
  inference(cnf_transformation,[],[f17212_D])).
fof(f17212_D,plain,(
  ( ! [X10] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num',X10)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X10),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X10),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num',X10))))) ) <=> ~$spl272),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl272])])).
fof(f17210,plain,(
  ( ! [X13] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num',X13)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num',X13)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X13),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X13),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) ) | $spl10),
  inference(duplicate_literal_removal,[],[f17195])).
fof(f17195,plain,(
  ( ! [X13] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num',X13)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num',X13)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X13),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X13),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) ) | $spl10),
  inference(resolution,[],[f1151,f6717])).
fof(f6717,plain,(
  ( ! [X2] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X2))))) ) | $spl10),
  inference(resolution,[],[f5936,f51])).
fof(f5936,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))))) ) | $spl10),
  inference(superposition,[],[f768,f53])).
fof(f1151,plain,(
  ( ! [X6,X5] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X6)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num',X5)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num',X5)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num',X6)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X5),s('type.num.num',X6))))) )),
  inference(superposition,[],[f133,f53])).
fof(f133,plain,(
  ( ! [X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num',X1))))) )),
  inference(superposition,[],[f48,f53])).
fof(f48,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))))))),
  inference(cnf_transformation,[],[f39])).
fof(f17214,plain,(
  ( ! [X12] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num',X12)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num',X12)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X12),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X12),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) ) | ($spl10 | $spl13)),
  inference(subsumption_resolution,[],[f17194,f769])).
fof(f769,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.num.|48|_0')))) | $spl13),
  inference(cnf_transformation,[],[f769_D])).
fof(f769_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.num.|48|_0')))) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f17194,plain,(
  ( ! [X12] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num',X12)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num',X12)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X12),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X12),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) ) | $spl10),
  inference(resolution,[],[f1151,f8084])).
fof(f8084,plain,(
  ( ! [X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X3)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X3))))) ) | $spl10),
  inference(resolution,[],[f5940,f51])).
fof(f5940,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))) ) | $spl10),
  inference(superposition,[],[f768,f53])).
fof(f3378,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | $spl92),
  inference(cnf_transformation,[],[f3378_D])).
fof(f3378_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f17213,plain,(
  ( ! [X11] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num',X11)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num',X11)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X11),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) ) | ($spl13 | $spl27)),
  inference(subsumption_resolution,[],[f17193,f769])).
fof(f17193,plain,(
  ( ! [X11] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num',X11)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num',X11)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X11),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.num.|48|_0'))))) ) | $spl27),
  inference(resolution,[],[f1151,f8593])).
fof(f8593,plain,(
  ( ! [X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X3),s('type.num.num',X3)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X3))))) ) | $spl27),
  inference(resolution,[],[f4501,f51])).
fof(f4501,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X0))))) ) | $spl27),
  inference(superposition,[],[f2311,f53])).
fof(f2311,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl27),
  inference(resolution,[],[f1817,f51])).
fof(f1817,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl27),
  inference(backward_demodulation,[],[f1816,f1427])).
fof(f1427,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl27),
  inference(cnf_transformation,[],[f1427_D])).
fof(f1427_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) <=> ~$spl27),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl27])])).
fof(f1816,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl27),
  inference(subsumption_resolution,[],[f1803,f64])).
fof(f64,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X0))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))))) )),
  inference(resolution,[],[f45,f56])).
fof(f45,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1))))),
  inference(cnf_transformation,[],[f39])).
fof(f1803,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl27),
  inference(resolution,[],[f1427,f53])).
fof(f17211,plain,(
  ( ! [X10] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num',X10)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num',X10)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X10),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X10),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) ) | ($spl10 | $spl13)),
  inference(subsumption_resolution,[],[f17192,f769])).
fof(f17192,plain,(
  ( ! [X10] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num',X10)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num',X10)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X10),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X10),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) ) | $spl10),
  inference(resolution,[],[f1151,f8735])).
fof(f8735,plain,(
  ( ! [X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X3)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X3))))) ) | $spl10),
  inference(resolution,[],[f5966,f51])).
fof(f5966,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))) ) | $spl10),
  inference(superposition,[],[f3434,f53])).
fof(f3434,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl10),
  inference(resolution,[],[f768,f54])).
fof(f17157,plain,(
  ( ! [X10] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num',X10)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X10)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X10)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num',X10))))) ) | $spl270),
  inference(cnf_transformation,[],[f17157_D])).
fof(f17157_D,plain,(
  ( ! [X10] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num',X10)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X10)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X10)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num',X10))))) ) <=> ~$spl270),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl270])])).
fof(f17155,plain,(
  ( ! [X13] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num',X13)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num',X13)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X13)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X13)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) ) | $spl10),
  inference(duplicate_literal_removal,[],[f17140])).
fof(f17140,plain,(
  ( ! [X13] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num',X13)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num',X13)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X13)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X13)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) ) | $spl10),
  inference(resolution,[],[f1150,f6717])).
fof(f1150,plain,(
  ( ! [X4,X3] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X4)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num',X3)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num',X3)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num',X3)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X4),s('type.num.num',X3))))) )),
  inference(superposition,[],[f133,f53])).
fof(f17159,plain,(
  ( ! [X12] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num',X12)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num',X12)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X12)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X12)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) ) | ($spl10 | $spl13)),
  inference(subsumption_resolution,[],[f17139,f769])).
fof(f17139,plain,(
  ( ! [X12] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num',X12)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num',X12)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X12)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X12)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) ) | $spl10),
  inference(resolution,[],[f1150,f8084])).
fof(f17158,plain,(
  ( ! [X11] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num',X11)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num',X11)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X11)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X11)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) ) | ($spl13 | $spl27)),
  inference(subsumption_resolution,[],[f17138,f769])).
fof(f17138,plain,(
  ( ! [X11] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num',X11)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num',X11)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X11)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X11)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.num.|48|_0'))))) ) | $spl27),
  inference(resolution,[],[f1150,f8593])).
fof(f17156,plain,(
  ( ! [X10] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num',X10)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num',X10)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X10)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X10)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) ) | ($spl10 | $spl13)),
  inference(subsumption_resolution,[],[f17137,f769])).
fof(f17137,plain,(
  ( ! [X10] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num',X10)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num',X10)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X10)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X10)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) ) | $spl10),
  inference(resolution,[],[f1150,f8735])).
fof(f17081,plain,(
  ( ! [X34] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X34),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X34),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X34),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X34),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))))))) ) | $spl268),
  inference(cnf_transformation,[],[f17081_D])).
fof(f17081_D,plain,(
  ( ! [X34] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X34),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X34),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X34),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X34),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))))))) ) <=> ~$spl268),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl268])])).
fof(f17067,plain,(
  ( ! [X34] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X34),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X34),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X34),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X34),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))))))) ) | $spl18),
  inference(duplicate_literal_removal,[],[f17013])).
fof(f17013,plain,(
  ( ! [X34] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X34),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X34),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X34),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X34),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))))))) ) | $spl18),
  inference(resolution,[],[f16591,f444])).
fof(f444,plain,(
  ( ! [X6,X5] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num',X6)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X5),s('type.num.num',X6))))) )),
  inference(superposition,[],[f124,f53])).
fof(f124,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))))))) )),
  inference(superposition,[],[f48,f53])).
fof(f17080,plain,(
  ( ! [X33] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X33),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num',X33)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num',X33)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X33),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))))))) ) | $spl266),
  inference(cnf_transformation,[],[f17080_D])).
fof(f17080_D,plain,(
  ( ! [X33] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X33),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num',X33)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num',X33)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X33),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))))))) ) <=> ~$spl266),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl266])])).
fof(f17068,plain,(
  ( ! [X33] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X33),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X33),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num',X33)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num',X33))))) ) | $spl18),
  inference(duplicate_literal_removal,[],[f17012])).
fof(f17012,plain,(
  ( ! [X33] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X33),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X33),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num',X33)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num',X33))))) ) | $spl18),
  inference(resolution,[],[f16591,f443])).
fof(f443,plain,(
  ( ! [X4,X3] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X4),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num',X3)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X4),s('type.num.num',X3))))) )),
  inference(superposition,[],[f124,f53])).
fof(f16795,plain,(
  ( ! [X10] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X10),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X10),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X10),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X10),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))))))) ) | $spl264),
  inference(cnf_transformation,[],[f16795_D])).
fof(f16795_D,plain,(
  ( ! [X10] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X10),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X10),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X10),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X10),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))))))) ) <=> ~$spl264),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl264])])).
fof(f16793,plain,(
  ( ! [X13] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X13),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X13),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X13),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X13),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) ) | $spl10),
  inference(duplicate_literal_removal,[],[f16778])).
fof(f16778,plain,(
  ( ! [X13] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X13),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X13),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X13),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X13),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) ) | $spl10),
  inference(resolution,[],[f1088,f6717])).
fof(f1088,plain,(
  ( ! [X6,X5] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X6)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X5),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num',X6)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X5),s('type.num.num',X6))))) )),
  inference(superposition,[],[f132,f53])).
fof(f132,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))))))) )),
  inference(superposition,[],[f48,f53])).
fof(f16797,plain,(
  ( ! [X12] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X12),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X12),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X12),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X12),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) ) | ($spl10 | $spl13)),
  inference(subsumption_resolution,[],[f16777,f769])).
fof(f16777,plain,(
  ( ! [X12] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X12),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X12),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X12),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X12),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) ) | $spl10),
  inference(resolution,[],[f1088,f8084])).
fof(f16796,plain,(
  ( ! [X11] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X11),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X11),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) ) | ($spl13 | $spl27)),
  inference(subsumption_resolution,[],[f16776,f769])).
fof(f16776,plain,(
  ( ! [X11] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X11),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X11),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.num.|48|_0'))))) ) | $spl27),
  inference(resolution,[],[f1088,f8593])).
fof(f16794,plain,(
  ( ! [X10] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X10),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X10),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X10),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X10),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) ) | ($spl10 | $spl13)),
  inference(subsumption_resolution,[],[f16775,f769])).
fof(f16775,plain,(
  ( ! [X10] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X10),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X10),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X10),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X10),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) ) | $spl10),
  inference(resolution,[],[f1088,f8735])).
fof(f16663,plain,(
  ( ! [X10] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X10),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X10)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X10)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X10),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))))))) ) | $spl262),
  inference(cnf_transformation,[],[f16663_D])).
fof(f16663_D,plain,(
  ( ! [X10] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X10),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X10)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X10)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X10),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))))))) ) <=> ~$spl262),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl262])])).
fof(f16661,plain,(
  ( ! [X13] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X13),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X13),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X13)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X13)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) ) | $spl10),
  inference(duplicate_literal_removal,[],[f16646])).
fof(f16646,plain,(
  ( ! [X13] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X13),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X13),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X13)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X13)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) ) | $spl10),
  inference(resolution,[],[f1087,f6717])).
fof(f1087,plain,(
  ( ! [X4,X3] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X4)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X3),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num',X3)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X4),s('type.num.num',X3))))) )),
  inference(superposition,[],[f132,f53])).
fof(f16665,plain,(
  ( ! [X12] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X12),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X12),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X12)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X12)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) ) | ($spl10 | $spl13)),
  inference(subsumption_resolution,[],[f16645,f769])).
fof(f16645,plain,(
  ( ! [X12] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X12),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X12),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X12)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X12)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) ) | $spl10),
  inference(resolution,[],[f1087,f8084])).
fof(f16664,plain,(
  ( ! [X11] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X11),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X11)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X11)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) ) | ($spl13 | $spl27)),
  inference(subsumption_resolution,[],[f16644,f769])).
fof(f16644,plain,(
  ( ! [X11] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X11),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X11)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X11)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.num.|48|_0'))))) ) | $spl27),
  inference(resolution,[],[f1087,f8593])).
fof(f16662,plain,(
  ( ! [X10] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X10),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X10),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X10)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X10)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) ) | ($spl10 | $spl13)),
  inference(subsumption_resolution,[],[f16643,f769])).
fof(f16643,plain,(
  ( ! [X10] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X10),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X10),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X10)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X10)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) ) | $spl10),
  inference(resolution,[],[f1087,f8735])).
fof(f16141,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | $spl242),
  inference(cnf_transformation,[],[f16141_D])).
fof(f16141_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) <=> ~$spl242),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl242])])).
fof(f16007,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl240),
  inference(cnf_transformation,[],[f16007_D])).
fof(f16007_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) <=> ~$spl240),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl240])])).
fof(f16371,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | ($spl27 | $spl238)),
  inference(resolution,[],[f16004,f4501])).
fof(f16004,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl238),
  inference(cnf_transformation,[],[f16004_D])).
fof(f16004_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) <=> ~$spl238),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl238])])).
fof(f16334,plain,(
  ( ! [X2] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X2)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num',X2)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))))))) ) | $spl260),
  inference(cnf_transformation,[],[f16334_D])).
fof(f16334_D,plain,(
  ( ! [X2] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X2)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num',X2)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))))))) ) <=> ~$spl260),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl260])])).
fof(f16369,plain,(
  ( ! [X4] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X4)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num',X4))))) ) | $spl238),
  inference(resolution,[],[f16004,f340])).
fof(f340,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(superposition,[],[f79,f53])).
fof(f79,plain,(
  ( ! [X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))))))) )),
  inference(superposition,[],[f47,f53])).
fof(f16333,plain,(
  ( ! [X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))))))) ) | $spl258),
  inference(cnf_transformation,[],[f16333_D])).
fof(f16333_D,plain,(
  ( ! [X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))))))) ) <=> ~$spl258),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl258])])).
fof(f16368,plain,(
  ( ! [X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X3),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))))) ) | $spl238),
  inference(resolution,[],[f16004,f330])).
fof(f330,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) )),
  inference(superposition,[],[f78,f53])).
fof(f16332,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X0))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))))) ) | $spl256),
  inference(cnf_transformation,[],[f16332_D])).
fof(f16332_D,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X0))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))))) ) <=> ~$spl256),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl256])])).
fof(f16365,plain,(
  ( ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X0))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))))) ) | $spl238),
  inference(resolution,[],[f16004,f200])).
fof(f200,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X1))))),s('type.num.num',X0))))) )),
  inference(superposition,[],[f64,f53])).
fof(f16003,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | $spl236),
  inference(cnf_transformation,[],[f16003_D])).
fof(f16003_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) <=> ~$spl236),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl236])])).
fof(f16313,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | ($spl27 | $spl234)),
  inference(resolution,[],[f16000,f4502])).
fof(f4502,plain,(
  ( ! [X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X1))))) ) | $spl27),
  inference(superposition,[],[f2311,f53])).
fof(f16000,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | $spl234),
  inference(cnf_transformation,[],[f16000_D])).
fof(f16000_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) <=> ~$spl234),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl234])])).
fof(f16310,plain,(
  ( ! [X2] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X2)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num',X2))))) ) | $spl234),
  inference(resolution,[],[f16000,f341])).
fof(f341,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X2)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X3)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X2)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X3))))) )),
  inference(superposition,[],[f79,f53])).
fof(f16309,plain,(
  ( ! [X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))))) ) | $spl234),
  inference(resolution,[],[f16000,f331])).
fof(f331,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X2)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X2)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X3),s('type.num.num',X2))))) )),
  inference(superposition,[],[f78,f53])).
fof(f16308,plain,(
  ( ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X0))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))))) ) | $spl234),
  inference(resolution,[],[f16000,f201])).
fof(f201,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X2)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X2)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X3))))),s('type.num.num',X2))))) )),
  inference(superposition,[],[f64,f53])).
fof(f16218,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | $spl245),
  inference(cnf_transformation,[],[f16218_D])).
fof(f16218_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) <=> ~$spl245),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl245])])).
fof(f16220,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | $spl246),
  inference(cnf_transformation,[],[f16220_D])).
fof(f16220_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) <=> ~$spl246),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl246])])).
fof(f16254,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | ($spl1 | $spl230)),
  inference(resolution,[],[f16205,f4532])).
fof(f4532,plain,(
  ( ! [X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num',X1))))) ) | $spl1),
  inference(superposition,[],[f135,f53])).
fof(f16205,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | $spl230),
  inference(resolution,[],[f15965,f54])).
fof(f16222,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | $spl249),
  inference(cnf_transformation,[],[f16222_D])).
fof(f16222_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) <=> ~$spl249),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl249])])).
fof(f16224,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | $spl250),
  inference(cnf_transformation,[],[f16224_D])).
fof(f16224_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) <=> ~$spl250),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl250])])).
fof(f16253,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | ($spl1 | $spl230)),
  inference(resolution,[],[f16205,f4531])).
fof(f4531,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))))))) ) | $spl1),
  inference(superposition,[],[f135,f53])).
fof(f16239,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.num.|48|_0')))) | $spl253),
  inference(cnf_transformation,[],[f16239_D])).
fof(f16239_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.num.|48|_0')))) <=> ~$spl253),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl253])])).
fof(f16241,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))) | $spl254),
  inference(cnf_transformation,[],[f16241_D])).
fof(f16241_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))) <=> ~$spl254),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl254])])).
fof(f16225,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.num.|48|_0')))) | $spl233),
  inference(resolution,[],[f15966,f53])).
fof(f15966,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.num.|48|_0')))) | $spl233),
  inference(cnf_transformation,[],[f15966_D])).
fof(f15966_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.num.|48|_0')))) <=> ~$spl233),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl233])])).
fof(f16204,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | $spl230),
  inference(resolution,[],[f15965,f132])).
fof(f16203,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | $spl230),
  inference(resolution,[],[f15965,f133])).
fof(f16147,plain,(
  $false | ($spl16 | $spl18)),
  inference(subsumption_resolution,[],[f6025,f776])).
fof(f6025,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | $spl16),
  inference(resolution,[],[f5944,f52])).
fof(f5944,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | $spl16),
  inference(resolution,[],[f774,f54])).
fof(f16145,plain,(
  $false | ($spl16 | $spl97)),
  inference(subsumption_resolution,[],[f5944,f3443])).
fof(f15965,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | $spl230),
  inference(cnf_transformation,[],[f15965_D])).
fof(f15965_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) <=> ~$spl230),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl230])])).
fof(f15967,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.num.|48|_0')))) | $spl232),
  inference(cnf_transformation,[],[f15967_D])).
fof(f15967_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.num.|48|_0')))) <=> ~$spl232),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl232])])).
fof(f16103,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | ($spl10 | $spl239)),
  inference(resolution,[],[f16005,f6924])).
fof(f6924,plain,(
  ( ! [X2] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X2))))) ) | $spl10),
  inference(resolution,[],[f5962,f51])).
fof(f5962,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))))) ) | $spl10),
  inference(superposition,[],[f3434,f53])).
fof(f16143,plain,(
  $false | ($spl235 | $spl239)),
  inference(subsumption_resolution,[],[f16125,f16005])).
fof(f16125,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl235),
  inference(resolution,[],[f16079,f51])).
fof(f16079,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | $spl235),
  inference(resolution,[],[f16001,f54])).
fof(f16121,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | ($spl27 | $spl235)),
  inference(resolution,[],[f16079,f8593])).
fof(f16001,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | $spl235),
  inference(cnf_transformation,[],[f16001_D])).
fof(f16001_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) <=> ~$spl235),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl235])])).
fof(f16025,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | ($spl10 | $spl232)),
  inference(resolution,[],[f15988,f5991])).
fof(f5991,plain,(
  ( ! [X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X1))))) ) | $spl10),
  inference(superposition,[],[f3435,f53])).
fof(f3435,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.num.|48|_0')))) | $spl10),
  inference(resolution,[],[f768,f52])).
fof(f15988,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.num.|48|_0')))) | $spl232),
  inference(resolution,[],[f15967,f54])).
fof(f16005,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl239),
  inference(cnf_transformation,[],[f16005_D])).
fof(f16005_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) <=> ~$spl239),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl239])])).
fof(f16024,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl10 | $spl232)),
  inference(resolution,[],[f15988,f5990])).
fof(f5990,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))))) ) | $spl10),
  inference(superposition,[],[f3435,f53])).
fof(f15987,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl13 | $spl232)),
  inference(resolution,[],[f15967,f5977])).
fof(f5977,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))))) ) | $spl13),
  inference(superposition,[],[f769,f53])).
fof(f15986,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | ($spl13 | $spl232)),
  inference(resolution,[],[f15967,f5978])).
fof(f5978,plain,(
  ( ! [X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X1))))) ) | $spl13),
  inference(superposition,[],[f769,f53])).
fof(f15937,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | ($spl17 | $spl74)),
  inference(resolution,[],[f15560,f2916])).
fof(f15560,plain,(
  ( ! [X2] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X2))))) ) | $spl17),
  inference(resolution,[],[f14885,f51])).
fof(f14885,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))))))) ) | $spl17),
  inference(superposition,[],[f773,f53])).
fof(f15539,plain,(
  $false | ($spl27 | $spl57 | $spl154 | $spl203)),
  inference(subsumption_resolution,[],[f15538,f13883])).
fof(f15538,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | ($spl27 | $spl57 | $spl154)),
  inference(forward_demodulation,[],[f10611,f13341])).
fof(f13341,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl27 | $spl57)),
  inference(forward_demodulation,[],[f13026,f4602])).
fof(f4602,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))) | $spl27),
  inference(backward_demodulation,[],[f1816,f844])).
fof(f844,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))),
  inference(resolution,[],[f65,f45])).
fof(f65,plain,(
  ( ! [X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X1)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),X1),s(fun(sK0,bool),sK1))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),X1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X1),s(fun(sK0,bool),sK1)))))))) )),
  inference(resolution,[],[f45,f49])).
fof(f13026,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))) | $spl57),
  inference(backward_demodulation,[],[f4856,f904])).
fof(f904,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))),
  inference(resolution,[],[f66,f46])).
fof(f66,plain,(
  ( ! [X2] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X2)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X2))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X2)))))))) )),
  inference(resolution,[],[f45,f49])).
fof(f4856,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl57),
  inference(subsumption_resolution,[],[f4843,f64])).
fof(f4843,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl57),
  inference(resolution,[],[f2519,f53])).
fof(f2519,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl57),
  inference(cnf_transformation,[],[f2519_D])).
fof(f2519_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) <=> ~$spl57),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl57])])).
fof(f10611,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ($spl27 | $spl154)),
  inference(superposition,[],[f6248,f1816])).
fof(f13664,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | $spl188),
  inference(cnf_transformation,[],[f13664_D])).
fof(f13664_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f5903,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | $spl150),
  inference(cnf_transformation,[],[f5903_D])).
fof(f5903_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f15491,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | ($spl148 | $spl172)),
  inference(resolution,[],[f5901,f8884])).
fof(f8884,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X0))))) ) | $spl172),
  inference(superposition,[],[f8852,f53])).
fof(f15508,plain,(
  ( ! [X4] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num',X4)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num',X4)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X4))))) ) | $spl228),
  inference(cnf_transformation,[],[f15508_D])).
fof(f15508_D,plain,(
  ( ! [X4] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num',X4)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num',X4)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X4))))) ) <=> ~$spl228),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl228])])).
fof(f15490,plain,(
  ( ! [X4] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num',X4)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X4)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num',X4))))) ) | $spl148),
  inference(resolution,[],[f5901,f424])).
fof(f424,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(superposition,[],[f83,f53])).
fof(f15507,plain,(
  ( ! [X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X3),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X3))))) ) | $spl226),
  inference(cnf_transformation,[],[f15507_D])).
fof(f15507_D,plain,(
  ( ! [X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X3),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X3))))) ) <=> ~$spl226),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl226])])).
fof(f15489,plain,(
  ( ! [X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X3)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X3),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))))))) ) | $spl148),
  inference(resolution,[],[f5901,f383])).
fof(f383,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) )),
  inference(superposition,[],[f82,f53])).
fof(f15506,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X0))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))))))) ) | $spl224),
  inference(cnf_transformation,[],[f15506_D])).
fof(f15506_D,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X0))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))))))) ) <=> ~$spl224),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl224])])).
fof(f15486,plain,(
  ( ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X0))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))))))) ) | $spl148),
  inference(resolution,[],[f5901,f216])).
fof(f216,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X1))))),s('type.num.num',X0))))) )),
  inference(superposition,[],[f69,f53])).
fof(f5898,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | $spl146),
  inference(cnf_transformation,[],[f5898_D])).
fof(f5898_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f15452,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | $spl203),
  inference(resolution,[],[f13883,f511])).
fof(f511,plain,(
  ( ! [X3] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X3)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num',X3))))) )),
  inference(resolution,[],[f92,f51])).
fof(f92,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X0))))) )),
  inference(superposition,[],[f74,f53])).
fof(f5900,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | $spl149),
  inference(cnf_transformation,[],[f5900_D])).
fof(f5900_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) <=> ~$spl149),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl149])])).
fof(f15441,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | $spl170),
  inference(resolution,[],[f8119,f115])).
fof(f115,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))))))) )),
  inference(superposition,[],[f85,f53])).
fof(f15440,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | $spl170),
  inference(resolution,[],[f8119,f870])).
fof(f870,plain,(
  ( ! [X2] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num',X2)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))))))) )),
  inference(resolution,[],[f116,f51])).
fof(f8119,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl170),
  inference(cnf_transformation,[],[f8119_D])).
fof(f8119_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f15423,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | ($spl27 | $spl168)),
  inference(resolution,[],[f8116,f4501])).
fof(f8116,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl168),
  inference(cnf_transformation,[],[f8116_D])).
fof(f8116_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f15422,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | $spl168),
  inference(resolution,[],[f8116,f570])).
fof(f570,plain,(
  ( ! [X2] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num',X2)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))))))) )),
  inference(resolution,[],[f101,f51])).
fof(f13673,plain,(
  ( ! [X6] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num',X6)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X6))))) ) | $spl192),
  inference(cnf_transformation,[],[f13673_D])).
fof(f13673_D,plain,(
  ( ! [X6] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num',X6)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X6))))) ) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f15421,plain,(
  ( ! [X4] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X4)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num',X4))))) ) | $spl168),
  inference(resolution,[],[f8116,f340])).
fof(f13670,plain,(
  ( ! [X6] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))))) ) | $spl190),
  inference(cnf_transformation,[],[f13670_D])).
fof(f13670_D,plain,(
  ( ! [X6] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))))) ) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f15420,plain,(
  ( ! [X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X3),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))))))) ) | $spl168),
  inference(resolution,[],[f8116,f330])).
fof(f5895,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X0))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))))))) ) | $spl142),
  inference(cnf_transformation,[],[f5895_D])).
fof(f5895_D,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X0))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))))))) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f15417,plain,(
  ( ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X0))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))))))) ) | $spl168),
  inference(resolution,[],[f8116,f200])).
fof(f5899,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | $spl147),
  inference(cnf_transformation,[],[f5899_D])).
fof(f5899_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) <=> ~$spl147),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl147])])).
fof(f15416,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | $spl168),
  inference(resolution,[],[f8116,f119])).
fof(f15415,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | $spl168),
  inference(resolution,[],[f8116,f100])).
fof(f100,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))))))) )),
  inference(superposition,[],[f75,f53])).
fof(f15414,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | $spl168),
  inference(resolution,[],[f8116,f78])).
fof(f3443,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | $spl97),
  inference(cnf_transformation,[],[f3443_D])).
fof(f3443_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) <=> ~$spl97),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl97])])).
fof(f14893,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))) | $spl222),
  inference(cnf_transformation,[],[f14893_D])).
fof(f14893_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))) <=> ~$spl222),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl222])])).
fof(f14882,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | $spl17),
  inference(resolution,[],[f773,f53])).
fof(f776,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | $spl18),
  inference(cnf_transformation,[],[f776_D])).
fof(f776_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f14881,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | ($spl17 | $spl172)),
  inference(resolution,[],[f773,f8923])).
fof(f8923,plain,(
  ( ! [X5] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X5),s('type.num.num',X5)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num',X5))))) ) | $spl172),
  inference(resolution,[],[f8884,f51])).
fof(f8117,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl169),
  inference(cnf_transformation,[],[f8117_D])).
fof(f8117_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) <=> ~$spl169),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl169])])).
fof(f14525,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl13 | $spl25 | $spl27 | $spl57)),
  inference(forward_demodulation,[],[f14524,f13341])).
fof(f14524,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl13 | $spl25 | $spl27 | $spl57)),
  inference(forward_demodulation,[],[f14523,f13341])).
fof(f14523,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl13 | $spl25 | $spl27 | $spl57)),
  inference(subsumption_resolution,[],[f14197,f1567])).
fof(f1567,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) | $spl25),
  inference(resolution,[],[f1426,f51])).
fof(f14197,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl13 | $spl27 | $spl57)),
  inference(superposition,[],[f5977,f13341])).
fof(f5897,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | $spl144),
  inference(cnf_transformation,[],[f5897_D])).
fof(f5897_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f13883,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | $spl203),
  inference(cnf_transformation,[],[f13883_D])).
fof(f13883_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) <=> ~$spl203),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl203])])).
fof(f14528,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | ($spl13 | $spl25 | $spl27 | $spl57)),
  inference(forward_demodulation,[],[f14527,f13341])).
fof(f14527,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ($spl13 | $spl25 | $spl27 | $spl57)),
  inference(forward_demodulation,[],[f14526,f13341])).
fof(f14526,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ($spl13 | $spl25 | $spl27 | $spl57)),
  inference(subsumption_resolution,[],[f14198,f1567])).
fof(f14198,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ($spl13 | $spl27 | $spl57)),
  inference(superposition,[],[f5978,f13341])).
fof(f14533,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl10 | $spl25 | $spl27 | $spl57)),
  inference(forward_demodulation,[],[f14532,f13341])).
fof(f14532,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl10 | $spl25 | $spl27 | $spl57)),
  inference(forward_demodulation,[],[f14531,f13341])).
fof(f14531,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl10 | $spl25 | $spl27 | $spl57)),
  inference(subsumption_resolution,[],[f14201,f1578])).
fof(f1578,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) | $spl25),
  inference(resolution,[],[f1565,f51])).
fof(f1565,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | $spl25),
  inference(resolution,[],[f1426,f54])).
fof(f14201,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl10 | $spl27 | $spl57)),
  inference(superposition,[],[f5990,f13341])).
fof(f14536,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | ($spl10 | $spl25 | $spl27 | $spl57)),
  inference(forward_demodulation,[],[f14535,f13341])).
fof(f14535,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ($spl10 | $spl25 | $spl27 | $spl57)),
  inference(forward_demodulation,[],[f14534,f13341])).
fof(f14534,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ($spl10 | $spl25 | $spl27 | $spl57)),
  inference(subsumption_resolution,[],[f14202,f1578])).
fof(f14202,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ($spl10 | $spl27 | $spl57)),
  inference(superposition,[],[f5991,f13341])).
fof(f13968,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl205),
  inference(cnf_transformation,[],[f13968_D])).
fof(f13968_D,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) <=> ~$spl205),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl205])])).
fof(f14584,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | ($spl10 | $spl27 | $spl57)),
  inference(forward_demodulation,[],[f14583,f13341])).
fof(f14583,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl10 | $spl27 | $spl57)),
  inference(subsumption_resolution,[],[f14582,f1817])).
fof(f14582,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl10 | $spl27 | $spl57)),
  inference(forward_demodulation,[],[f14581,f4863])).
fof(f4863,plain,(
  s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))))) | $spl27),
  inference(superposition,[],[f50,f4602])).
fof(f14581,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl10 | $spl27 | $spl57)),
  inference(forward_demodulation,[],[f14248,f13341])).
fof(f14248,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl10 | $spl27 | $spl57)),
  inference(superposition,[],[f8083,f13341])).
fof(f8083,plain,(
  ( ! [X2] : (s('type.num.num',X2) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X2)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))))) ) | $spl10),
  inference(resolution,[],[f5940,f55])).
fof(f14592,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | ($spl10 | $spl27 | $spl57)),
  inference(forward_demodulation,[],[f14591,f13341])).
fof(f14591,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ($spl10 | $spl27 | $spl57)),
  inference(subsumption_resolution,[],[f14590,f1817])).
fof(f14590,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ($spl10 | $spl27 | $spl57)),
  inference(forward_demodulation,[],[f14589,f4863])).
fof(f14589,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ($spl10 | $spl27 | $spl57)),
  inference(forward_demodulation,[],[f14251,f13341])).
fof(f14251,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ($spl10 | $spl27 | $spl57)),
  inference(superposition,[],[f8216,f13341])).
fof(f8216,plain,(
  ( ! [X2] : (s('type.num.num',X2) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X2)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X2))))) ) | $spl10),
  inference(resolution,[],[f5941,f55])).
fof(f5941,plain,(
  ( ! [X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1))))) ) | $spl10),
  inference(superposition,[],[f768,f53])).
fof(f14607,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | ($spl10 | $spl25 | $spl27 | $spl57)),
  inference(forward_demodulation,[],[f14606,f13341])).
fof(f14606,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl10 | $spl25 | $spl27 | $spl57)),
  inference(subsumption_resolution,[],[f14605,f1426])).
fof(f14605,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl10 | $spl27 | $spl57)),
  inference(forward_demodulation,[],[f14261,f13341])).
fof(f14261,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl10 | $spl27 | $spl57)),
  inference(superposition,[],[f8734,f13341])).
fof(f8734,plain,(
  ( ! [X2] : (s('type.num.num',X2) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X2)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))))) ) | $spl10),
  inference(resolution,[],[f5966,f55])).
fof(f14614,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | ($spl10 | $spl25 | $spl27 | $spl57)),
  inference(forward_demodulation,[],[f14613,f13341])).
fof(f14613,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ($spl10 | $spl25 | $spl27 | $spl57)),
  inference(subsumption_resolution,[],[f14612,f1426])).
fof(f14612,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ($spl10 | $spl27 | $spl57)),
  inference(forward_demodulation,[],[f14264,f13341])).
fof(f14264,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ($spl10 | $spl27 | $spl57)),
  inference(superposition,[],[f8789,f13341])).
fof(f8789,plain,(
  ( ! [X2] : (s('type.num.num',X2) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X2)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X2))))) ) | $spl10),
  inference(resolution,[],[f5967,f55])).
fof(f5967,plain,(
  ( ! [X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1))))) ) | $spl10),
  inference(superposition,[],[f3434,f53])).
fof(f774,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | $spl16),
  inference(cnf_transformation,[],[f774_D])).
fof(f774_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f5956,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | $spl10),
  inference(resolution,[],[f3434,f570])).
fof(f5958,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | $spl10),
  inference(resolution,[],[f3434,f100])).
fof(f5983,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | $spl10),
  inference(resolution,[],[f3435,f511])).
fof(f7275,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | $spl10),
  inference(resolution,[],[f6717,f85])).
fof(f14700,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ($spl10 | $spl27)),
  inference(subsumption_resolution,[],[f7293,f1817])).
fof(f7293,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl10),
  inference(duplicate_literal_removal,[],[f7276])).
fof(f7276,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl10),
  inference(resolution,[],[f6717,f890])).
fof(f7317,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | $spl13),
  inference(resolution,[],[f6844,f47])).
fof(f6844,plain,(
  ( ! [X2] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X2)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X2))))) ) | $spl13),
  inference(resolution,[],[f5979,f51])).
fof(f5979,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X0))))) ) | $spl13),
  inference(superposition,[],[f769,f53])).
fof(f14699,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ($spl13 | $spl27)),
  inference(subsumption_resolution,[],[f7343,f1817])).
fof(f7343,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl13),
  inference(duplicate_literal_removal,[],[f7325])).
fof(f7325,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl13),
  inference(resolution,[],[f6844,f320])).
fof(f14698,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ($spl13 | $spl27)),
  inference(subsumption_resolution,[],[f7342,f1817])).
fof(f7342,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl13),
  inference(trivial_inequality_removal,[],[f7326])).
fof(f7326,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | $spl13),
  inference(resolution,[],[f6844,f319])).
fof(f7447,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | $spl10),
  inference(resolution,[],[f6924,f75])).
fof(f14697,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ($spl10 | $spl27)),
  inference(subsumption_resolution,[],[f7495,f1817])).
fof(f7495,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl10),
  inference(subsumption_resolution,[],[f7456,f51])).
fof(f7456,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl10),
  inference(resolution,[],[f6924,f104])).
fof(f14696,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ($spl10 | $spl27)),
  inference(subsumption_resolution,[],[f7496,f1817])).
fof(f7496,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl10),
  inference(subsumption_resolution,[],[f7457,f51])).
fof(f7457,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl10),
  inference(resolution,[],[f6924,f105])).
fof(f14695,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ($spl10 | $spl27)),
  inference(subsumption_resolution,[],[f7494,f1817])).
fof(f7494,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl10),
  inference(duplicate_literal_removal,[],[f7460])).
fof(f7460,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl10),
  inference(resolution,[],[f6924,f701])).
fof(f7533,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | $spl10),
  inference(resolution,[],[f7003,f74])).
fof(f7003,plain,(
  ( ! [X2] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X2)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X2))))) ) | $spl10),
  inference(resolution,[],[f5992,f51])).
fof(f5992,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X0))))) ) | $spl10),
  inference(superposition,[],[f3435,f53])).
fof(f14694,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ($spl10 | $spl27)),
  inference(subsumption_resolution,[],[f7557,f1817])).
fof(f7557,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl10),
  inference(duplicate_literal_removal,[],[f7540])).
fof(f7540,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl10),
  inference(resolution,[],[f7003,f489])).
fof(f489,plain,(
  ( ! [X2] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X2)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))))) )),
  inference(resolution,[],[f89,f51])).
fof(f89,plain,(
  ( ! [X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))))))) )),
  inference(superposition,[],[f74,f53])).
fof(f14693,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ($spl10 | $spl27)),
  inference(subsumption_resolution,[],[f7558,f1817])).
fof(f7558,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl10),
  inference(subsumption_resolution,[],[f7541,f51])).
fof(f7541,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl10),
  inference(resolution,[],[f7003,f88])).
fof(f88,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))))) )),
  inference(superposition,[],[f74,f53])).
fof(f14692,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | ($spl10 | $spl27)),
  inference(subsumption_resolution,[],[f9068,f1817])).
fof(f9068,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | $spl10),
  inference(subsumption_resolution,[],[f9054,f51])).
fof(f9054,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | $spl10),
  inference(resolution,[],[f5990,f511])).
fof(f14691,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | ($spl10 | $spl27)),
  inference(subsumption_resolution,[],[f9101,f1817])).
fof(f9101,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | $spl10),
  inference(subsumption_resolution,[],[f9087,f51])).
fof(f9087,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | $spl10),
  inference(resolution,[],[f5991,f511])).
fof(f14690,plain,(
  ( ! [X23] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X23),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num',X23)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num',X23)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X23))))) ) | $spl220),
  inference(cnf_transformation,[],[f14690_D])).
fof(f14690_D,plain,(
  ( ! [X23] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X23),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num',X23)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num',X23)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X23))))) ) <=> ~$spl220),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl220])])).
fof(f9566,plain,(
  ( ! [X23] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X23),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X23)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num',X23)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num',X23)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))))))) )),
  inference(resolution,[],[f379,f634])).
fof(f634,plain,(
  ( ! [X2] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num',X2)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X2))))) )),
  inference(resolution,[],[f148,f51])).
fof(f148,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num',X0))))) )),
  inference(superposition,[],[f122,f53])).
fof(f122,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0'))))),
  inference(resolution,[],[f48,f51])).
fof(f379,plain,(
  ( ! [X4,X3] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X4),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X3)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num',X3)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X4),s('type.num.num',X3))))) )),
  inference(superposition,[],[f82,f53])).
fof(f14686,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | ($spl10 | $spl172)),
  inference(subsumption_resolution,[],[f14685,f870])).
fof(f14685,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl10 | $spl172)),
  inference(subsumption_resolution,[],[f9720,f9646])).
fof(f9646,plain,(
  ( ! [X2] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num',X2)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num',X2)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))))))) ) | $spl172),
  inference(resolution,[],[f8899,f51])).
fof(f9720,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl10 | $spl172)),
  inference(resolution,[],[f8923,f5936])).
fof(f14684,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | ($spl10 | $spl172)),
  inference(subsumption_resolution,[],[f14683,f570])).
fof(f14683,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl10 | $spl172)),
  inference(subsumption_resolution,[],[f9721,f9646])).
fof(f9721,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl10 | $spl172)),
  inference(resolution,[],[f8923,f5962])).
fof(f14681,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | ($spl10 | $spl172)),
  inference(subsumption_resolution,[],[f14680,f9646])).
fof(f14680,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | ($spl10 | $spl172)),
  inference(subsumption_resolution,[],[f9725,f511])).
fof(f9725,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.num.|48|_0')))) | ($spl10 | $spl172)),
  inference(resolution,[],[f8923,f5992])).
fof(f14676,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | ($spl10 | $spl172)),
  inference(subsumption_resolution,[],[f14675,f870])).
fof(f14675,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl10 | $spl172)),
  inference(subsumption_resolution,[],[f9762,f9646])).
fof(f9762,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl10 | $spl172)),
  inference(resolution,[],[f8923,f5937])).
fof(f5937,plain,(
  ( ! [X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))))) ) | $spl10),
  inference(superposition,[],[f768,f53])).
fof(f14674,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | ($spl10 | $spl172)),
  inference(subsumption_resolution,[],[f14673,f570])).
fof(f14673,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl10 | $spl172)),
  inference(subsumption_resolution,[],[f9763,f9646])).
fof(f9763,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl10 | $spl172)),
  inference(resolution,[],[f8923,f5963])).
fof(f5963,plain,(
  ( ! [X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))))) ) | $spl10),
  inference(superposition,[],[f3434,f53])).
fof(f14671,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | ($spl10 | $spl172)),
  inference(subsumption_resolution,[],[f14670,f9646])).
fof(f14670,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | ($spl10 | $spl172)),
  inference(subsumption_resolution,[],[f9765,f511])).
fof(f9765,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.num.|48|_0')))) | ($spl10 | $spl172)),
  inference(resolution,[],[f8923,f5993])).
fof(f5993,plain,(
  ( ! [X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X1))))) ) | $spl10),
  inference(superposition,[],[f3435,f53])).
fof(f14669,plain,(
  ( ! [X24] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X24),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X24),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X24),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X24))))) ) | $spl218),
  inference(cnf_transformation,[],[f14669_D])).
fof(f14669_D,plain,(
  ( ! [X24] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X24),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X24),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X24),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X24))))) ) <=> ~$spl218),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl218])])).
fof(f9898,plain,(
  ( ! [X24] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X24),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X24)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X24),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X24),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))))))) )),
  inference(resolution,[],[f380,f634])).
fof(f380,plain,(
  ( ! [X6,X5] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X5)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num',X6)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X5),s('type.num.num',X6))))) )),
  inference(superposition,[],[f82,f53])).
fof(f14668,plain,(
  ( ! [X11] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X11)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X11))))) ) | $spl216),
  inference(cnf_transformation,[],[f14668_D])).
fof(f14668_D,plain,(
  ( ! [X11] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X11)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X11))))) ) <=> ~$spl216),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl216])])).
fof(f14667,plain,(
  ( ! [X11] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X11)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X11)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))))))) ) | $spl10),
  inference(subsumption_resolution,[],[f9987,f120])).
fof(f120,plain,(
  ( ! [X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X1)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num',X1))))) )),
  inference(superposition,[],[f85,f53])).
fof(f9987,plain,(
  ( ! [X11] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num',X11)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X11)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X11)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))))))) ) | $spl10),
  inference(duplicate_literal_removal,[],[f9971])).
fof(f9971,plain,(
  ( ! [X11] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num',X11)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X11)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X11)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X11)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))))))) ) | $spl10),
  inference(resolution,[],[f427,f6924])).
fof(f427,plain,(
  ( ! [X4,X3] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num',X4)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num',X3)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X3)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num',X3)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X4),s('type.num.num',X3))))) )),
  inference(superposition,[],[f83,f53])).
fof(f14666,plain,(
  ( ! [X11] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X11)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))))) ) | $spl214),
  inference(cnf_transformation,[],[f14666_D])).
fof(f14666_D,plain,(
  ( ! [X11] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X11)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))))) ) <=> ~$spl214),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl214])])).
fof(f14665,plain,(
  ( ! [X11] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X11)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))))))) ) | $spl10),
  inference(subsumption_resolution,[],[f10368,f890])).
fof(f10368,plain,(
  ( ! [X11] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num',X11)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X11)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))))))) ) | $spl10),
  inference(subsumption_resolution,[],[f10347,f51])).
fof(f10347,plain,(
  ( ! [X11] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num',X11)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X11)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X11),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))))))) ) | $spl10),
  inference(resolution,[],[f428,f6924])).
fof(f428,plain,(
  ( ! [X6,X5] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num',X6)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num',X5)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X5)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num',X6)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X5),s('type.num.num',X6))))) )),
  inference(superposition,[],[f83,f53])).
fof(f14664,plain,(
  ( ! [X11] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X11)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X11))))) ) | $spl212),
  inference(cnf_transformation,[],[f14664_D])).
fof(f14664_D,plain,(
  ( ! [X11] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X11)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X11))))) ) <=> ~$spl212),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl212])])).
fof(f14663,plain,(
  ( ! [X11] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X11)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X11)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0'))))) ) | $spl13),
  inference(subsumption_resolution,[],[f10522,f5978])).
fof(f10522,plain,(
  ( ! [X11] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X11)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X11)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0'))))) )),
  inference(resolution,[],[f443,f511])).
fof(f14662,plain,(
  ( ! [X11] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X11),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))))) ) | $spl210),
  inference(cnf_transformation,[],[f14662_D])).
fof(f14662_D,plain,(
  ( ! [X11] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X11),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))))) ) <=> ~$spl210),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl210])])).
fof(f14661,plain,(
  ( ! [X11] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X11),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0'))))) ) | $spl13),
  inference(subsumption_resolution,[],[f10678,f5977])).
fof(f10678,plain,(
  ( ! [X11] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X11),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0'))))) )),
  inference(resolution,[],[f444,f511])).
fof(f14647,plain,(
  ( ! [X12] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num',X12)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X12)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X12),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num',X12))))) ) | $spl208),
  inference(cnf_transformation,[],[f14647_D])).
fof(f14647_D,plain,(
  ( ! [X12] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num',X12)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X12)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X12),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num',X12))))) ) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f12426,plain,(
  ( ! [X10] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X10),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X10)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num',X10)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num',X10))))) ) | $spl25),
  inference(resolution,[],[f11884,f379])).
fof(f11884,plain,(
  ( ! [X2] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num',X2)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X2))))) ) | $spl25),
  inference(resolution,[],[f1801,f51])).
fof(f1801,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num',X0))))) ) | $spl25),
  inference(superposition,[],[f1578,f53])).
fof(f14644,plain,(
  ( ! [X13] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X13),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X13)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X13),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X13),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))))))) ) | $spl206),
  inference(cnf_transformation,[],[f14644_D])).
fof(f14644_D,plain,(
  ( ! [X13] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X13),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X13)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X13),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X13),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))))))) ) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f12427,plain,(
  ( ! [X11] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X11)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X11),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))))))) ) | $spl25),
  inference(resolution,[],[f11884,f380])).
fof(f14646,plain,(
  ( ! [X12] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num',X12)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num',X12)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X12),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X12))))) ) | ($spl27 | $spl55 | $spl57)),
  inference(forward_demodulation,[],[f14645,f13341])).
fof(f14645,plain,(
  ( ! [X12] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num',X12)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X12),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X12)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num',X12))))) ) | ($spl27 | $spl55 | $spl57)),
  inference(forward_demodulation,[],[f12539,f13341])).
fof(f12539,plain,(
  ( ! [X12] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X12),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X12)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num',X12)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num',X12))))) ) | $spl55),
  inference(resolution,[],[f11969,f379])).
fof(f11969,plain,(
  ( ! [X2] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num',X2)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X2))))) ) | $spl55),
  inference(resolution,[],[f4600,f51])).
fof(f4600,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num',X0))))) ) | $spl55),
  inference(superposition,[],[f4560,f53])).
fof(f4560,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.num.|48|_0')))) | $spl55),
  inference(resolution,[],[f4545,f51])).
fof(f4545,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | $spl55),
  inference(resolution,[],[f2518,f54])).
fof(f14643,plain,(
  ( ! [X13] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X13),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X13),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X13),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X13))))) ) | ($spl27 | $spl55 | $spl57)),
  inference(forward_demodulation,[],[f14642,f13341])).
fof(f14642,plain,(
  ( ! [X13] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X13),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X13),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X13)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X13),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))) ) | ($spl27 | $spl55 | $spl57)),
  inference(forward_demodulation,[],[f12540,f13341])).
fof(f12540,plain,(
  ( ! [X13] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X13),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X13)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X13),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X13),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))) ) | $spl55),
  inference(resolution,[],[f11969,f380])).
fof(f14302,plain,(
  $false | ($spl27 | $spl57 | $spl149 | $spl163)),
  inference(subsumption_resolution,[],[f14017,f8345])).
fof(f8345,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | $spl149),
  inference(resolution,[],[f8312,f51])).
fof(f8312,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | $spl149),
  inference(resolution,[],[f5900,f54])).
fof(f14017,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | ($spl27 | $spl57 | $spl163)),
  inference(backward_demodulation,[],[f13341,f6453])).
fof(f14301,plain,(
  $false | ($spl27 | $spl57 | $spl149 | $spl160)),
  inference(subsumption_resolution,[],[f14016,f8312])).
fof(f14016,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ($spl27 | $spl57 | $spl160)),
  inference(backward_demodulation,[],[f13341,f6254])).
fof(f14300,plain,(
  $false | ($spl27 | $spl57 | $spl149 | $spl158)),
  inference(subsumption_resolution,[],[f14015,f5900])).
fof(f14015,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ($spl27 | $spl57 | $spl158)),
  inference(backward_demodulation,[],[f13341,f6252])).
fof(f14299,plain,(
  $false | ($spl27 | $spl57 | $spl149 | $spl157)),
  inference(subsumption_resolution,[],[f14014,f8314])).
fof(f8314,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | $spl149),
  inference(resolution,[],[f5900,f51])).
fof(f14014,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | ($spl27 | $spl57 | $spl157)),
  inference(backward_demodulation,[],[f13341,f6250])).
fof(f13994,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | ($spl10 | $spl27)),
  inference(forward_demodulation,[],[f13993,f4602])).
fof(f13993,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl10 | $spl27)),
  inference(subsumption_resolution,[],[f13992,f1817])).
fof(f13992,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl10 | $spl27)),
  inference(forward_demodulation,[],[f8191,f50])).
fof(f8191,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl10 | $spl27)),
  inference(superposition,[],[f8083,f4602])).
fof(f13991,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | ($spl10 | $spl27)),
  inference(subsumption_resolution,[],[f13990,f1817])).
fof(f13990,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | ($spl10 | $spl27)),
  inference(forward_demodulation,[],[f13989,f50])).
fof(f13989,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | ($spl10 | $spl27)),
  inference(forward_demodulation,[],[f8231,f4602])).
fof(f8231,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | ($spl10 | $spl27)),
  inference(superposition,[],[f5941,f4602])).
fof(f13988,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | ($spl10 | $spl27)),
  inference(forward_demodulation,[],[f13987,f4602])).
fof(f13987,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | ($spl10 | $spl27)),
  inference(subsumption_resolution,[],[f13986,f1817])).
fof(f13986,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | ($spl10 | $spl27)),
  inference(forward_demodulation,[],[f8252,f50])).
fof(f8252,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | ($spl10 | $spl27)),
  inference(superposition,[],[f8216,f4602])).
fof(f8310,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl149),
  inference(resolution,[],[f5900,f489])).
fof(f8311,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl149),
  inference(resolution,[],[f5900,f88])).
fof(f8342,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl149),
  inference(resolution,[],[f8312,f320])).
fof(f8343,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl149),
  inference(resolution,[],[f8312,f319])).
fof(f8358,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | $spl149),
  inference(resolution,[],[f8314,f890])).
fof(f8374,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl149),
  inference(resolution,[],[f8345,f104])).
fof(f8375,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | $spl149),
  inference(resolution,[],[f8345,f105])).
fof(f8378,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | $spl149),
  inference(resolution,[],[f8345,f701])).
fof(f8488,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl13 | $spl25)),
  inference(resolution,[],[f5977,f1567])).
fof(f13985,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl13 | $spl25 | $spl27)),
  inference(forward_demodulation,[],[f8515,f4602])).
fof(f8515,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl13 | $spl25 | $spl27)),
  inference(forward_demodulation,[],[f8514,f4602])).
fof(f8514,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl13 | $spl25 | $spl27)),
  inference(subsumption_resolution,[],[f8498,f1567])).
fof(f8498,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl13 | $spl27)),
  inference(superposition,[],[f5977,f4602])).
fof(f8520,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | ($spl13 | $spl25)),
  inference(resolution,[],[f5978,f1567])).
fof(f13984,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | ($spl13 | $spl25 | $spl27)),
  inference(forward_demodulation,[],[f8547,f4602])).
fof(f8547,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | ($spl13 | $spl25 | $spl27)),
  inference(forward_demodulation,[],[f8546,f4602])).
fof(f8546,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | ($spl13 | $spl25 | $spl27)),
  inference(subsumption_resolution,[],[f8530,f1567])).
fof(f8530,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | ($spl13 | $spl27)),
  inference(superposition,[],[f5978,f4602])).
fof(f13973,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | ($spl10 | $spl25 | $spl27)),
  inference(forward_demodulation,[],[f13972,f4602])).
fof(f13972,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl10 | $spl25 | $spl27)),
  inference(subsumption_resolution,[],[f13971,f1426])).
fof(f13971,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl10 | $spl27)),
  inference(forward_demodulation,[],[f8770,f4602])).
fof(f8770,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl10 | $spl27)),
  inference(superposition,[],[f8734,f4602])).
fof(f13966,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | ($spl10 | $spl25 | $spl27)),
  inference(forward_demodulation,[],[f13965,f4602])).
fof(f13965,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | ($spl10 | $spl25 | $spl27)),
  inference(subsumption_resolution,[],[f13964,f1426])).
fof(f13964,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | ($spl10 | $spl27)),
  inference(forward_demodulation,[],[f8904,f4602])).
fof(f8904,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | ($spl10 | $spl27)),
  inference(superposition,[],[f8789,f4602])).
fof(f9052,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl10 | $spl25)),
  inference(resolution,[],[f5990,f1578])).
fof(f13961,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl10 | $spl25 | $spl27)),
  inference(forward_demodulation,[],[f9078,f4602])).
fof(f9078,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl10 | $spl25 | $spl27)),
  inference(forward_demodulation,[],[f9077,f4602])).
fof(f9077,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl10 | $spl25 | $spl27)),
  inference(subsumption_resolution,[],[f9061,f1578])).
fof(f9061,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl10 | $spl27)),
  inference(superposition,[],[f5990,f4602])).
fof(f9085,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | ($spl10 | $spl25)),
  inference(resolution,[],[f5991,f1578])).
fof(f13960,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | ($spl10 | $spl25 | $spl27)),
  inference(forward_demodulation,[],[f9111,f4602])).
fof(f9111,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | ($spl10 | $spl25 | $spl27)),
  inference(forward_demodulation,[],[f9110,f4602])).
fof(f9110,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | ($spl10 | $spl25 | $spl27)),
  inference(subsumption_resolution,[],[f9094,f1578])).
fof(f9094,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | ($spl10 | $spl27)),
  inference(superposition,[],[f5991,f4602])).
fof(f13955,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl10 | $spl27)),
  inference(subsumption_resolution,[],[f13954,f1817])).
fof(f13954,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl10 | $spl27)),
  inference(forward_demodulation,[],[f13953,f50])).
fof(f13953,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | ($spl10 | $spl27)),
  inference(forward_demodulation,[],[f9386,f4602])).
fof(f9386,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | ($spl10 | $spl27)),
  inference(superposition,[],[f8084,f4602])).
fof(f13952,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | ($spl10 | $spl27)),
  inference(subsumption_resolution,[],[f13951,f1817])).
fof(f13951,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | ($spl10 | $spl27)),
  inference(forward_demodulation,[],[f13950,f50])).
fof(f13950,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | ($spl10 | $spl27)),
  inference(forward_demodulation,[],[f9402,f4602])).
fof(f9402,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | ($spl10 | $spl27)),
  inference(superposition,[],[f8084,f4602])).
fof(f12208,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | ($spl13 | $spl25)),
  inference(duplicate_literal_removal,[],[f12191])).
fof(f12191,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | ($spl13 | $spl25)),
  inference(resolution,[],[f11496,f5978])).
fof(f11496,plain,(
  ( ! [X2] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num',X2)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X2))))) ) | $spl25),
  inference(resolution,[],[f1596,f51])).
fof(f1596,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num',X0))))) ) | $spl25),
  inference(superposition,[],[f1567,f53])).
fof(f12207,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl13 | $spl25)),
  inference(duplicate_literal_removal,[],[f12192])).
fof(f12192,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl13 | $spl25)),
  inference(resolution,[],[f11496,f5977])).
fof(f12477,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl10 | $spl25)),
  inference(duplicate_literal_removal,[],[f12450])).
fof(f12450,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl10 | $spl25)),
  inference(resolution,[],[f11884,f5990])).
fof(f12476,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | ($spl10 | $spl25)),
  inference(duplicate_literal_removal,[],[f12451])).
fof(f12451,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | ($spl10 | $spl25)),
  inference(resolution,[],[f11884,f5991])).
fof(f4422,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl123),
  inference(cnf_transformation,[],[f4422_D])).
fof(f4422_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) <=> ~$spl123),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl123])])).
fof(f13903,plain,(
  ( ! [X6] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num',X6)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X6)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))))))) ) | ($spl27 | $spl57)),
  inference(forward_demodulation,[],[f13902,f4602])).
fof(f13902,plain,(
  ( ! [X6] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.num.num',X6)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X6)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))))))) ) | ($spl27 | $spl57)),
  inference(forward_demodulation,[],[f13678,f4856])).
fof(f13678,plain,(
  ( ! [X6] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X6)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num',X6))))) ) | ($spl27 | $spl57)),
  inference(forward_demodulation,[],[f13677,f4602])).
fof(f13677,plain,(
  ( ! [X6] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X6)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num',X6))))) ) | $spl57),
  inference(forward_demodulation,[],[f12286,f4856])).
fof(f12286,plain,(
  ( ! [X6] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X6)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num',X6))))) )),
  inference(superposition,[],[f340,f904])).
fof(f13901,plain,(
  ( ! [X6] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))))))) ) | ($spl27 | $spl57)),
  inference(forward_demodulation,[],[f13900,f4602])).
fof(f13900,plain,(
  ( ! [X6] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))))))) ) | ($spl27 | $spl57)),
  inference(forward_demodulation,[],[f13687,f4856])).
fof(f13687,plain,(
  ( ! [X6] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))))) ) | ($spl27 | $spl57)),
  inference(forward_demodulation,[],[f13686,f4602])).
fof(f13686,plain,(
  ( ! [X6] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))))) ) | $spl57),
  inference(forward_demodulation,[],[f12015,f4856])).
fof(f12015,plain,(
  ( ! [X6] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))))) )),
  inference(superposition,[],[f330,f904])).
fof(f13699,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | $spl197),
  inference(cnf_transformation,[],[f13699_D])).
fof(f13699_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) <=> ~$spl197),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl197])])).
fof(f4428,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | $spl126),
  inference(cnf_transformation,[],[f4428_D])).
fof(f4428_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f13696,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ($spl27 | $spl57)),
  inference(forward_demodulation,[],[f13695,f4602])).
fof(f13695,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ($spl27 | $spl57)),
  inference(forward_demodulation,[],[f10598,f4856])).
fof(f10598,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl27),
  inference(forward_demodulation,[],[f10576,f904])).
fof(f10576,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl27),
  inference(superposition,[],[f10050,f904])).
fof(f10050,plain,(
  ( ! [X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))))) ) | $spl27),
  inference(resolution,[],[f8593,f52])).
fof(f13899,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ($spl10 | $spl25 | $spl27 | $spl57)),
  inference(subsumption_resolution,[],[f13898,f1426])).
fof(f13898,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ($spl10 | $spl27 | $spl57)),
  inference(forward_demodulation,[],[f13897,f4602])).
fof(f13897,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ($spl10 | $spl27 | $spl57)),
  inference(forward_demodulation,[],[f13701,f4856])).
fof(f13701,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | ($spl10 | $spl27 | $spl57)),
  inference(forward_demodulation,[],[f13700,f4602])).
fof(f13700,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | ($spl10 | $spl57)),
  inference(forward_demodulation,[],[f10306,f4856])).
fof(f10306,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | $spl10),
  inference(superposition,[],[f8735,f904])).
fof(f13896,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ($spl27 | $spl57)),
  inference(forward_demodulation,[],[f13895,f4602])).
fof(f13895,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ($spl27 | $spl57)),
  inference(forward_demodulation,[],[f13707,f4856])).
fof(f13707,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | ($spl27 | $spl57)),
  inference(forward_demodulation,[],[f13706,f4602])).
fof(f13706,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | ($spl27 | $spl57)),
  inference(forward_demodulation,[],[f10171,f4856])).
fof(f10171,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | $spl27),
  inference(superposition,[],[f8593,f904])).
fof(f13894,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ($spl10 | $spl27 | $spl57)),
  inference(subsumption_resolution,[],[f13893,f1817])).
fof(f13893,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ($spl10 | $spl27 | $spl57)),
  inference(forward_demodulation,[],[f13892,f4856])).
fof(f13892,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ($spl10 | $spl27 | $spl57)),
  inference(forward_demodulation,[],[f13712,f50])).
fof(f13712,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | ($spl10 | $spl27 | $spl57)),
  inference(forward_demodulation,[],[f13711,f4602])).
fof(f13711,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | ($spl10 | $spl57)),
  inference(forward_demodulation,[],[f9401,f4856])).
fof(f9401,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | $spl10),
  inference(superposition,[],[f8084,f904])).
fof(f13723,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ($spl10 | $spl27 | $spl55 | $spl57)),
  inference(forward_demodulation,[],[f13722,f4602])).
fof(f13722,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ($spl10 | $spl55 | $spl57)),
  inference(forward_demodulation,[],[f9109,f4856])).
fof(f9109,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | ($spl10 | $spl55)),
  inference(forward_demodulation,[],[f9108,f904])).
fof(f9108,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | ($spl10 | $spl55)),
  inference(subsumption_resolution,[],[f9093,f4560])).
fof(f9093,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | $spl10),
  inference(superposition,[],[f5991,f904])).
fof(f5357,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl135),
  inference(cnf_transformation,[],[f5357_D])).
fof(f5357_D,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl135),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl135])])).
fof(f13734,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl10 | $spl25 | $spl27 | $spl57)),
  inference(forward_demodulation,[],[f13733,f4602])).
fof(f13733,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl10 | $spl25 | $spl27 | $spl57)),
  inference(forward_demodulation,[],[f13732,f4856])).
fof(f13732,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | ($spl10 | $spl25 | $spl27 | $spl57)),
  inference(subsumption_resolution,[],[f13731,f1426])).
fof(f13731,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | ($spl10 | $spl27 | $spl57)),
  inference(forward_demodulation,[],[f13730,f4602])).
fof(f13730,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | ($spl10 | $spl57)),
  inference(forward_demodulation,[],[f8903,f4856])).
fof(f8903,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | $spl10),
  inference(superposition,[],[f8789,f904])).
fof(f13739,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl10 | $spl25 | $spl27 | $spl57)),
  inference(forward_demodulation,[],[f13738,f4602])).
fof(f13738,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl10 | $spl25 | $spl27 | $spl57)),
  inference(forward_demodulation,[],[f13737,f4856])).
fof(f13737,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl10 | $spl25 | $spl27 | $spl57)),
  inference(subsumption_resolution,[],[f13736,f1426])).
fof(f13736,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl10 | $spl27 | $spl57)),
  inference(forward_demodulation,[],[f13735,f4602])).
fof(f13735,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl10 | $spl57)),
  inference(forward_demodulation,[],[f8769,f4856])).
fof(f8769,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl10),
  inference(superposition,[],[f8734,f904])).
fof(f13891,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl10 | $spl25 | $spl27 | $spl57)),
  inference(subsumption_resolution,[],[f13890,f1426])).
fof(f13890,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl10 | $spl27 | $spl57)),
  inference(forward_demodulation,[],[f13889,f4602])).
fof(f13889,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl10 | $spl27 | $spl57)),
  inference(forward_demodulation,[],[f13742,f4856])).
fof(f13742,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | ($spl10 | $spl27 | $spl57)),
  inference(forward_demodulation,[],[f13741,f4602])).
fof(f13741,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | ($spl10 | $spl57)),
  inference(forward_demodulation,[],[f8748,f4856])).
fof(f8748,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | $spl10),
  inference(superposition,[],[f5966,f904])).
fof(f13750,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl27 | $spl57)),
  inference(forward_demodulation,[],[f13749,f4602])).
fof(f13749,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl27 | $spl57)),
  inference(forward_demodulation,[],[f13748,f4856])).
fof(f13748,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | ($spl27 | $spl57)),
  inference(forward_demodulation,[],[f13747,f4602])).
fof(f13747,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | ($spl27 | $spl57)),
  inference(forward_demodulation,[],[f8717,f4856])).
fof(f8717,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | $spl27),
  inference(superposition,[],[f8684,f904])).
fof(f8684,plain,(
  ( ! [X2] : (s('type.num.num',X2) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X2)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X2))))) ) | $spl27),
  inference(resolution,[],[f4502,f55])).
fof(f13754,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl27 | $spl57)),
  inference(forward_demodulation,[],[f13753,f4602])).
fof(f13753,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl27 | $spl57)),
  inference(forward_demodulation,[],[f13752,f4856])).
fof(f13752,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl27 | $spl57)),
  inference(forward_demodulation,[],[f13751,f4602])).
fof(f13751,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl27 | $spl57)),
  inference(forward_demodulation,[],[f8627,f4856])).
fof(f8627,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl27),
  inference(superposition,[],[f8592,f904])).
fof(f8592,plain,(
  ( ! [X2] : (s('type.num.num',X2) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X2)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))))) ) | $spl27),
  inference(resolution,[],[f4501,f55])).
fof(f13888,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl27 | $spl57)),
  inference(forward_demodulation,[],[f13887,f4602])).
fof(f13887,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl27 | $spl57)),
  inference(forward_demodulation,[],[f13757,f4856])).
fof(f13757,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | ($spl27 | $spl57)),
  inference(forward_demodulation,[],[f13756,f4602])).
fof(f13756,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | ($spl27 | $spl57)),
  inference(forward_demodulation,[],[f8606,f4856])).
fof(f8606,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | $spl27),
  inference(superposition,[],[f4501,f904])).
fof(f13762,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ($spl13 | $spl27 | $spl55 | $spl57)),
  inference(forward_demodulation,[],[f13761,f4602])).
fof(f13761,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ($spl13 | $spl55 | $spl57)),
  inference(forward_demodulation,[],[f8545,f4856])).
fof(f8545,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | ($spl13 | $spl55)),
  inference(forward_demodulation,[],[f8544,f904])).
fof(f8544,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | ($spl13 | $spl55)),
  inference(subsumption_resolution,[],[f8529,f4547])).
fof(f4547,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.num.|48|_0')))) | $spl55),
  inference(resolution,[],[f2518,f51])).
fof(f8529,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | $spl13),
  inference(superposition,[],[f5978,f904])).
fof(f13768,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl10 | $spl27 | $spl57)),
  inference(forward_demodulation,[],[f13767,f4602])).
fof(f13767,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl10 | $spl27 | $spl57)),
  inference(forward_demodulation,[],[f13766,f4856])).
fof(f13766,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | ($spl10 | $spl27 | $spl57)),
  inference(subsumption_resolution,[],[f13765,f1817])).
fof(f13765,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | ($spl10 | $spl57)),
  inference(forward_demodulation,[],[f13764,f4856])).
fof(f13764,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | $spl10),
  inference(forward_demodulation,[],[f8251,f50])).
fof(f8251,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | $spl10),
  inference(superposition,[],[f8216,f904])).
fof(f13773,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl10 | $spl27 | $spl57)),
  inference(forward_demodulation,[],[f13772,f4602])).
fof(f13772,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl10 | $spl27 | $spl57)),
  inference(forward_demodulation,[],[f13771,f4856])).
fof(f13771,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl10 | $spl27 | $spl57)),
  inference(subsumption_resolution,[],[f13770,f1817])).
fof(f13770,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl10 | $spl57)),
  inference(forward_demodulation,[],[f13769,f4856])).
fof(f13769,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl10),
  inference(forward_demodulation,[],[f8190,f50])).
fof(f8190,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl10),
  inference(superposition,[],[f8083,f904])).
fof(f13886,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl10 | $spl27 | $spl57)),
  inference(subsumption_resolution,[],[f13885,f1817])).
fof(f13885,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl10 | $spl27 | $spl57)),
  inference(forward_demodulation,[],[f13884,f4856])).
fof(f13884,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl10 | $spl27 | $spl57)),
  inference(forward_demodulation,[],[f13776,f50])).
fof(f13776,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | ($spl10 | $spl27 | $spl57)),
  inference(forward_demodulation,[],[f13775,f4602])).
fof(f13775,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | ($spl10 | $spl57)),
  inference(forward_demodulation,[],[f8097,f4856])).
fof(f8097,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | $spl10),
  inference(superposition,[],[f5940,f904])).
fof(f13782,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ($spl27 | $spl57)),
  inference(forward_demodulation,[],[f13781,f4602])).
fof(f13781,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ($spl27 | $spl57)),
  inference(forward_demodulation,[],[f8072,f4856])).
fof(f8072,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | $spl27),
  inference(forward_demodulation,[],[f8053,f904])).
fof(f8053,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | $spl27),
  inference(superposition,[],[f4487,f904])).
fof(f4487,plain,(
  ( ! [X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X1))))) ) | $spl27),
  inference(superposition,[],[f1817,f53])).
fof(f13802,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl27 | $spl57 | $spl149)),
  inference(forward_demodulation,[],[f13801,f4602])).
fof(f13801,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl27 | $spl57 | $spl149)),
  inference(forward_demodulation,[],[f13800,f4856])).
fof(f13800,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | ($spl27 | $spl57 | $spl149)),
  inference(subsumption_resolution,[],[f13799,f8314])).
fof(f13799,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | ($spl27 | $spl57)),
  inference(forward_demodulation,[],[f13798,f4602])).
fof(f13798,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | $spl57),
  inference(forward_demodulation,[],[f6551,f4856])).
fof(f6551,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))))),
  inference(superposition,[],[f920,f904])).
fof(f920,plain,(
  ( ! [X1] : (s('type.num.num',X1) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X1))))) )),
  inference(resolution,[],[f120,f55])).
fof(f13807,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl27 | $spl57 | $spl149)),
  inference(forward_demodulation,[],[f13806,f4602])).
fof(f13806,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl27 | $spl57 | $spl149)),
  inference(forward_demodulation,[],[f13805,f4856])).
fof(f13805,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl27 | $spl57 | $spl149)),
  inference(subsumption_resolution,[],[f13804,f8314])).
fof(f13804,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl27 | $spl57)),
  inference(forward_demodulation,[],[f13803,f4602])).
fof(f13803,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl57),
  inference(forward_demodulation,[],[f6549,f4856])).
fof(f6549,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))))),
  inference(superposition,[],[f889,f904])).
fof(f889,plain,(
  ( ! [X1] : (s('type.num.num',X1) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))))) )),
  inference(resolution,[],[f119,f55])).
fof(f13815,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl27 | $spl57 | $spl149)),
  inference(forward_demodulation,[],[f13814,f4602])).
fof(f13814,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl27 | $spl57 | $spl149)),
  inference(forward_demodulation,[],[f13813,f4856])).
fof(f13813,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | ($spl27 | $spl57 | $spl149)),
  inference(subsumption_resolution,[],[f13812,f5900])).
fof(f13812,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | ($spl27 | $spl57)),
  inference(forward_demodulation,[],[f13811,f4602])).
fof(f13811,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | $spl57),
  inference(forward_demodulation,[],[f6534,f4856])).
fof(f6534,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))))),
  inference(superposition,[],[f488,f904])).
fof(f488,plain,(
  ( ! [X1] : (s('type.num.num',X1) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X1))))) )),
  inference(resolution,[],[f89,f55])).
fof(f13881,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl27 | $spl57 | $spl149)),
  inference(subsumption_resolution,[],[f13880,f8345])).
fof(f13880,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl27 | $spl57)),
  inference(forward_demodulation,[],[f13879,f4602])).
fof(f13879,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl27 | $spl57)),
  inference(forward_demodulation,[],[f13817,f4856])).
fof(f13817,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | ($spl27 | $spl57)),
  inference(forward_demodulation,[],[f13816,f4602])).
fof(f13816,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | $spl57),
  inference(forward_demodulation,[],[f6526,f4856])).
fof(f6526,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))))),
  inference(superposition,[],[f346,f904])).
fof(f346,plain,(
  ( ! [X1] : (s('type.num.num',X1) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num',X1))))) )),
  inference(resolution,[],[f319,f52])).
fof(f13878,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl27 | $spl57 | $spl149)),
  inference(subsumption_resolution,[],[f13877,f5900])).
fof(f13877,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl27 | $spl57)),
  inference(forward_demodulation,[],[f13876,f4602])).
fof(f13876,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl27 | $spl57)),
  inference(forward_demodulation,[],[f13823,f4856])).
fof(f13823,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ($spl27 | $spl57)),
  inference(forward_demodulation,[],[f13822,f4602])).
fof(f13822,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | $spl57),
  inference(forward_demodulation,[],[f6525,f4856])).
fof(f6525,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))))))),
  inference(superposition,[],[f345,f904])).
fof(f345,plain,(
  ( ! [X0] : (s('type.num.num',X0) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))))))) )),
  inference(resolution,[],[f319,f54])).
fof(f13832,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl27 | $spl57 | $spl149)),
  inference(forward_demodulation,[],[f13831,f4602])).
fof(f13831,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl27 | $spl57 | $spl149)),
  inference(forward_demodulation,[],[f13830,f4856])).
fof(f13830,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | ($spl27 | $spl57 | $spl149)),
  inference(subsumption_resolution,[],[f13829,f8312])).
fof(f13829,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | ($spl27 | $spl57)),
  inference(forward_demodulation,[],[f13828,f4602])).
fof(f13828,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | $spl57),
  inference(forward_demodulation,[],[f6524,f4856])).
fof(f6524,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))))),
  inference(superposition,[],[f335,f904])).
fof(f335,plain,(
  ( ! [X1] : (s('type.num.num',X1) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X1))))) )),
  inference(resolution,[],[f79,f55])).
fof(f13837,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ($spl27 | $spl57 | $spl149)),
  inference(forward_demodulation,[],[f13836,f4602])).
fof(f13836,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ($spl27 | $spl57 | $spl149)),
  inference(forward_demodulation,[],[f13835,f4856])).
fof(f13835,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl27 | $spl57 | $spl149)),
  inference(subsumption_resolution,[],[f13834,f8312])).
fof(f13834,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl27 | $spl57)),
  inference(forward_demodulation,[],[f13833,f4602])).
fof(f13833,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl57),
  inference(forward_demodulation,[],[f6522,f4856])).
fof(f6522,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))))),
  inference(superposition,[],[f320,f904])).
fof(f13875,plain,(
  ( ! [X24] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X24))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))))) ) | ($spl27 | $spl57)),
  inference(forward_demodulation,[],[f13874,f4602])).
fof(f13874,plain,(
  ( ! [X24] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X24))))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))))) ) | ($spl27 | $spl57)),
  inference(forward_demodulation,[],[f13840,f4856])).
fof(f13840,plain,(
  ( ! [X24] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X24))))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))))) ) | ($spl27 | $spl57)),
  inference(forward_demodulation,[],[f13839,f4602])).
fof(f13839,plain,(
  ( ! [X24] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X24))))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))))) ) | $spl57),
  inference(forward_demodulation,[],[f6517,f4856])).
fof(f6517,plain,(
  ( ! [X24] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X24))))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))))) )),
  inference(superposition,[],[f200,f904])).
fof(f13873,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl27 | $spl57 | $spl149)),
  inference(subsumption_resolution,[],[f13872,f8314])).
fof(f13872,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl27 | $spl57)),
  inference(forward_demodulation,[],[f13871,f4602])).
fof(f13871,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl27 | $spl57)),
  inference(forward_demodulation,[],[f13845,f4856])).
fof(f13845,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | ($spl27 | $spl57)),
  inference(forward_demodulation,[],[f13844,f4602])).
fof(f13844,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | $spl57),
  inference(forward_demodulation,[],[f6509,f4856])).
fof(f6509,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))))),
  inference(superposition,[],[f119,f904])).
fof(f13870,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl27 | $spl57 | $spl149)),
  inference(subsumption_resolution,[],[f13869,f8312])).
fof(f13869,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl27 | $spl57)),
  inference(forward_demodulation,[],[f13868,f4602])).
fof(f13868,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl27 | $spl57)),
  inference(forward_demodulation,[],[f13863,f4856])).
fof(f13863,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ($spl27 | $spl57)),
  inference(forward_demodulation,[],[f13862,f4602])).
fof(f13862,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | $spl57),
  inference(forward_demodulation,[],[f6493,f4856])).
fof(f6493,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))))))),
  inference(superposition,[],[f78,f904])).
fof(f13843,plain,(
  ( ! [X24] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X24))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))))) ) | ($spl27 | $spl57 | $spl169)),
  inference(forward_demodulation,[],[f13842,f4602])).
fof(f13842,plain,(
  ( ! [X24] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X24))))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))))) ) | ($spl27 | $spl57 | $spl169)),
  inference(forward_demodulation,[],[f13841,f4856])).
fof(f13841,plain,(
  ( ! [X24] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X24))))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))))) ) | ($spl27 | $spl57 | $spl169)),
  inference(subsumption_resolution,[],[f13840,f8125])).
fof(f8125,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl169),
  inference(resolution,[],[f8117,f54])).
fof(f13783,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ($spl27 | $spl57 | $spl169)),
  inference(subsumption_resolution,[],[f13782,f8151])).
fof(f8151,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | $spl169),
  inference(resolution,[],[f8125,f51])).
fof(f4424,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | $spl125),
  inference(cnf_transformation,[],[f4424_D])).
fof(f4424_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) <=> ~$spl125),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl125])])).
fof(f8522,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ($spl13 | $spl55)),
  inference(resolution,[],[f5978,f4547])).
fof(f13760,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl27 | $spl57 | $spl169)),
  inference(forward_demodulation,[],[f13759,f4602])).
fof(f13759,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl27 | $spl57 | $spl169)),
  inference(forward_demodulation,[],[f13758,f4856])).
fof(f13758,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | ($spl27 | $spl57 | $spl169)),
  inference(subsumption_resolution,[],[f13757,f8125])).
fof(f13755,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl27 | $spl57 | $spl169)),
  inference(subsumption_resolution,[],[f13754,f8125])).
fof(f13729,plain,(
  ( ! [X6] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num',X6)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num',X6))))) ) | $spl200),
  inference(cnf_transformation,[],[f13729_D])).
fof(f13729_D,plain,(
  ( ! [X6] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num',X6)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num',X6))))) ) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f13728,plain,(
  ( ! [X6] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num',X6)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num',X6)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))))))) ) | ($spl27 | $spl57)),
  inference(forward_demodulation,[],[f13727,f4602])).
fof(f13727,plain,(
  ( ! [X6] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.num.num',X6)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num',X6)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))))))) ) | ($spl27 | $spl57)),
  inference(forward_demodulation,[],[f13726,f4856])).
fof(f13726,plain,(
  ( ! [X6] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num',X6)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num',X6))))) ) | ($spl27 | $spl57)),
  inference(forward_demodulation,[],[f13725,f4602])).
fof(f13725,plain,(
  ( ! [X6] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.num.num',X6)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num',X6))))) ) | $spl57),
  inference(forward_demodulation,[],[f8988,f4856])).
fof(f8988,plain,(
  ( ! [X6] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num',X6)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num',X6))))) )),
  inference(superposition,[],[f326,f904])).
fof(f326,plain,(
  ( ! [X4,X3] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X4),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num',X3)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X4),s('type.num.num',X3))))) )),
  inference(superposition,[],[f78,f53])).
fof(f9086,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ($spl10 | $spl55)),
  inference(resolution,[],[f5991,f4560])).
fof(f13721,plain,(
  ( ! [X6] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))))))) ) | $spl198),
  inference(cnf_transformation,[],[f13721_D])).
fof(f13721_D,plain,(
  ( ! [X6] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))))))) ) <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f13720,plain,(
  ( ! [X6] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))))))) ) | ($spl27 | $spl57)),
  inference(forward_demodulation,[],[f13719,f4602])).
fof(f13719,plain,(
  ( ! [X6] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))))))) ) | ($spl27 | $spl57)),
  inference(forward_demodulation,[],[f13718,f4856])).
fof(f13718,plain,(
  ( ! [X6] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))))) ) | ($spl27 | $spl57)),
  inference(forward_demodulation,[],[f13717,f4602])).
fof(f13717,plain,(
  ( ! [X6] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))))) ) | $spl57),
  inference(forward_demodulation,[],[f9167,f4856])).
fof(f9167,plain,(
  ( ! [X6] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))))) )),
  inference(superposition,[],[f327,f904])).
fof(f327,plain,(
  ( ! [X6,X5] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num',X6)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X5),s('type.num.num',X6))))) )),
  inference(superposition,[],[f78,f53])).
fof(f13710,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ($spl27 | $spl57 | $spl169)),
  inference(forward_demodulation,[],[f13709,f4602])).
fof(f13709,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ($spl27 | $spl57 | $spl169)),
  inference(forward_demodulation,[],[f13708,f4856])).
fof(f13708,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | ($spl27 | $spl57 | $spl169)),
  inference(subsumption_resolution,[],[f13707,f8125])).
fof(f13697,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ($spl27 | $spl57 | $spl169)),
  inference(subsumption_resolution,[],[f13696,f8125])).
fof(f4421,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl122),
  inference(cnf_transformation,[],[f4421_D])).
fof(f4421_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f13694,plain,(
  ( ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X0))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))) ) | $spl194),
  inference(cnf_transformation,[],[f13694_D])).
fof(f13694_D,plain,(
  ( ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X0))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))) ) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f10600,plain,(
  ( ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X0))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))))) ) | $spl154),
  inference(resolution,[],[f6248,f197])).
fof(f197,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X2))))),s('type.num.num',X3)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X2))))),s('type.num.num',X3)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X3),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))))) )),
  inference(superposition,[],[f64,f53])).
fof(f13690,plain,(
  ( ! [X6] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))))))) ) | ($spl27 | $spl57 | $spl169)),
  inference(forward_demodulation,[],[f13689,f4602])).
fof(f13689,plain,(
  ( ! [X6] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))))))) ) | ($spl27 | $spl57 | $spl169)),
  inference(forward_demodulation,[],[f13688,f4856])).
fof(f13688,plain,(
  ( ! [X6] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))))) ) | ($spl27 | $spl57 | $spl169)),
  inference(subsumption_resolution,[],[f13687,f8125])).
fof(f13681,plain,(
  ( ! [X6] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num',X6)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X6)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))))))) ) | ($spl27 | $spl57 | $spl169)),
  inference(forward_demodulation,[],[f13680,f4602])).
fof(f13680,plain,(
  ( ! [X6] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.num.num',X6)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X6)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))))))) ) | ($spl27 | $spl57 | $spl169)),
  inference(forward_demodulation,[],[f13679,f4856])).
fof(f13679,plain,(
  ( ! [X6] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X6)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num',X6))))) ) | ($spl27 | $spl57 | $spl169)),
  inference(subsumption_resolution,[],[f13678,f8125])).
fof(f12342,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ($spl13 | $spl55)),
  inference(duplicate_literal_removal,[],[f12323])).
fof(f12323,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ($spl13 | $spl55)),
  inference(resolution,[],[f11667,f5978])).
fof(f11667,plain,(
  ( ! [X2] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num',X2)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X2))))) ) | $spl55),
  inference(resolution,[],[f4582,f51])).
fof(f4582,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num',X0))))) ) | $spl55),
  inference(superposition,[],[f4547,f53])).
fof(f12591,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ($spl10 | $spl55)),
  inference(duplicate_literal_removal,[],[f12564])).
fof(f12564,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ($spl10 | $spl55)),
  inference(resolution,[],[f11969,f5991])).
fof(f13672,plain,(
  ( ! [X6] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num',X6)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X6)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))))))) ) | ($spl27 | $spl57)),
  inference(forward_demodulation,[],[f13671,f4602])).
fof(f13671,plain,(
  ( ! [X6] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.num.num',X6)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X6)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))))))) ) | $spl57),
  inference(forward_demodulation,[],[f12990,f4856])).
fof(f12990,plain,(
  ( ! [X6] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X6)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num',X6))))) )),
  inference(forward_demodulation,[],[f12379,f904])).
fof(f12379,plain,(
  ( ! [X6] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X6)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num',X6))))) )),
  inference(superposition,[],[f341,f904])).
fof(f13669,plain,(
  ( ! [X6] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))))))) ) | ($spl27 | $spl57)),
  inference(forward_demodulation,[],[f13668,f4602])).
fof(f13668,plain,(
  ( ! [X6] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))))))) ) | $spl57),
  inference(forward_demodulation,[],[f12994,f4856])).
fof(f12994,plain,(
  ( ! [X6] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))))) )),
  inference(forward_demodulation,[],[f12112,f904])).
fof(f12112,plain,(
  ( ! [X6] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))))) )),
  inference(superposition,[],[f331,f904])).
fof(f13667,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ($spl10 | $spl25 | $spl27 | $spl57)),
  inference(subsumption_resolution,[],[f13666,f1426])).
fof(f13666,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ($spl10 | $spl27 | $spl57)),
  inference(forward_demodulation,[],[f13665,f4602])).
fof(f13665,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ($spl10 | $spl57)),
  inference(forward_demodulation,[],[f13008,f4856])).
fof(f13008,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | $spl10),
  inference(forward_demodulation,[],[f8803,f904])).
fof(f8803,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | $spl10),
  inference(superposition,[],[f5967,f904])).
fof(f13662,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ($spl27 | $spl57)),
  inference(forward_demodulation,[],[f13661,f4602])).
fof(f13661,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ($spl27 | $spl57)),
  inference(forward_demodulation,[],[f13009,f4856])).
fof(f13009,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | $spl27),
  inference(forward_demodulation,[],[f8698,f904])).
fof(f8698,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | $spl27),
  inference(superposition,[],[f4502,f904])).
fof(f13660,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ($spl10 | $spl27 | $spl57)),
  inference(subsumption_resolution,[],[f13659,f1817])).
fof(f13659,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ($spl10 | $spl57)),
  inference(forward_demodulation,[],[f13658,f4856])).
fof(f13658,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | $spl10),
  inference(forward_demodulation,[],[f13011,f50])).
fof(f13011,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | $spl10),
  inference(forward_demodulation,[],[f8230,f904])).
fof(f8230,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | $spl10),
  inference(superposition,[],[f5941,f904])).
fof(f13657,plain,(
  ( ! [X25] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X25))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))) ) | ($spl27 | $spl57)),
  inference(forward_demodulation,[],[f13656,f4602])).
fof(f13656,plain,(
  ( ! [X25] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X25))))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))) ) | $spl57),
  inference(forward_demodulation,[],[f13019,f4856])).
fof(f13019,plain,(
  ( ! [X25] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X25))))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))))) )),
  inference(forward_demodulation,[],[f6518,f904])).
fof(f6518,plain,(
  ( ! [X25] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X25))))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))))) )),
  inference(superposition,[],[f201,f904])).
fof(f13655,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ($spl27 | $spl57 | $spl149)),
  inference(subsumption_resolution,[],[f13654,f8314])).
fof(f13654,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ($spl27 | $spl57)),
  inference(forward_demodulation,[],[f13653,f4602])).
fof(f13653,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | $spl57),
  inference(forward_demodulation,[],[f13020,f4856])).
fof(f13020,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))))),
  inference(forward_demodulation,[],[f6510,f904])).
fof(f6510,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))))),
  inference(superposition,[],[f120,f904])).
fof(f13652,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ($spl27 | $spl57 | $spl149)),
  inference(subsumption_resolution,[],[f13651,f5900])).
fof(f13651,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ($spl27 | $spl57)),
  inference(forward_demodulation,[],[f13650,f4602])).
fof(f13650,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | $spl57),
  inference(forward_demodulation,[],[f13022,f4856])).
fof(f13022,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))))))),
  inference(forward_demodulation,[],[f6498,f904])).
fof(f6498,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))))))),
  inference(superposition,[],[f89,f904])).
fof(f13649,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ($spl27 | $spl57 | $spl149)),
  inference(subsumption_resolution,[],[f13648,f8312])).
fof(f13648,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ($spl27 | $spl57)),
  inference(forward_demodulation,[],[f13647,f4602])).
fof(f13647,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | $spl57),
  inference(forward_demodulation,[],[f13023,f4856])).
fof(f13023,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))))))),
  inference(forward_demodulation,[],[f6494,f904])).
fof(f6494,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))))))),
  inference(superposition,[],[f79,f904])).
fof(f13637,plain,(
  $false | ($spl120 | $spl123)),
  inference(subsumption_resolution,[],[f13625,f4422])).
fof(f13625,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl120),
  inference(resolution,[],[f4419,f54])).
fof(f4419,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl120),
  inference(cnf_transformation,[],[f4419_D])).
fof(f4419_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f13621,plain,(
  $false | ($spl57 | $spl115)),
  inference(subsumption_resolution,[],[f13620,f47])).
fof(f13620,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ($spl57 | $spl115)),
  inference(forward_demodulation,[],[f4388,f4856])).
fof(f4388,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | $spl115),
  inference(cnf_transformation,[],[f4388_D])).
fof(f4388_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) <=> ~$spl115),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl115])])).
fof(f4412,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | $spl118),
  inference(cnf_transformation,[],[f4412_D])).
fof(f4412_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f12985,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))))))),
  inference(subsumption_resolution,[],[f12984,f51])).
fof(f12984,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))))))),
  inference(forward_demodulation,[],[f12798,f851])).
fof(f851,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))),
  inference(forward_demodulation,[],[f845,f62])).
fof(f62,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))) )),
  inference(cnf_transformation,[],[f30])).
fof(f30,plain,(
  ! [X0,X1,X2] : s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))),
  inference(flattening,[],[f29])).
fof(f29,plain,(
  ! [X0] : ! [X1,X2] : s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6] : ! [X8,X7] : s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8))) = s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X7)))),
  file('thm.pred_set.LESS_CARD_DIFF','thm.pred_set.INTER_COMM')).
fof(f845,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))),
  inference(resolution,[],[f65,f46])).
fof(f12798,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))))))),
  inference(superposition,[],[f1089,f851])).
fof(f1089,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))))) )),
  inference(superposition,[],[f132,f50])).
fof(f12983,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))))))),
  inference(subsumption_resolution,[],[f12982,f51])).
fof(f12982,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))))))),
  inference(forward_demodulation,[],[f12934,f851])).
fof(f12934,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))))))),
  inference(superposition,[],[f1152,f851])).
fof(f1152,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))))) )),
  inference(superposition,[],[f133,f50])).
fof(f12980,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f12971,f12835])).
fof(f12835,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | $spl56),
  inference(subsumption_resolution,[],[f12834,f51])).
fof(f12834,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | $spl56),
  inference(forward_demodulation,[],[f12833,f851])).
fof(f12833,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | $spl56),
  inference(subsumption_resolution,[],[f12798,f2646])).
fof(f2646,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | $spl56),
  inference(subsumption_resolution,[],[f2632,f208])).
fof(f208,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X0),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))))))) )),
  inference(superposition,[],[f69,f62])).
fof(f2632,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | $spl56),
  inference(resolution,[],[f2520,f115])).
fof(f12971,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | $spl56),
  inference(resolution,[],[f12864,f52])).
fof(f12864,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | $spl56),
  inference(resolution,[],[f12835,f54])).
fof(f2934,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | $spl81),
  inference(cnf_transformation,[],[f2934_D])).
fof(f2934_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) <=> ~$spl81),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl81])])).
fof(f2917,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | $spl76),
  inference(cnf_transformation,[],[f2917_D])).
fof(f2917_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f12832,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | $spl74),
  inference(forward_demodulation,[],[f12831,f957])).
fof(f12831,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | $spl74),
  inference(subsumption_resolution,[],[f12797,f2916])).
fof(f12797,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))))))),
  inference(superposition,[],[f1089,f957])).
fof(f12615,plain,(
  ( ! [X15] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X15),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X15),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X15),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X15),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))))))) ) | $spl186),
  inference(cnf_transformation,[],[f12615_D])).
fof(f12615_D,plain,(
  ( ! [X15] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X15),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X15),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X15),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X15),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))))))) ) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f12608,plain,(
  ( ! [X15] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X15),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X15),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X15),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X15),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))) ) | $spl55),
  inference(duplicate_literal_removal,[],[f12547])).
fof(f12547,plain,(
  ( ! [X15] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X15),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X15),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X15),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X15),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))) ) | $spl55),
  inference(resolution,[],[f11969,f444])).
fof(f12614,plain,(
  ( ! [X14] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X14),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num',X14)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num',X14)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X14),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))))))) ) | $spl184),
  inference(cnf_transformation,[],[f12614_D])).
fof(f12614_D,plain,(
  ( ! [X14] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X14),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num',X14)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num',X14)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X14),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))))))) ) <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f12609,plain,(
  ( ! [X14] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X14),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X14),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num',X14)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num',X14))))) ) | $spl55),
  inference(duplicate_literal_removal,[],[f12546])).
fof(f12546,plain,(
  ( ! [X14] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X14),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X14),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num',X14)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num',X14))))) ) | $spl55),
  inference(resolution,[],[f11969,f443])).
fof(f12500,plain,(
  ( ! [X13] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X13),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X13),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X13),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X13),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))))))) ) | $spl182),
  inference(cnf_transformation,[],[f12500_D])).
fof(f12500_D,plain,(
  ( ! [X13] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X13),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X13),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X13),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X13),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))))))) ) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f12493,plain,(
  ( ! [X13] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X13),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X13),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X13),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X13),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))))))) ) | $spl25),
  inference(duplicate_literal_removal,[],[f12434])).
fof(f12434,plain,(
  ( ! [X13] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X13),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X13),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X13),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X13),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))))))) ) | $spl25),
  inference(resolution,[],[f11884,f444])).
fof(f12499,plain,(
  ( ! [X12] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X12),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num',X12)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num',X12)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X12),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))))))) ) | $spl180),
  inference(cnf_transformation,[],[f12499_D])).
fof(f12499_D,plain,(
  ( ! [X12] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X12),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num',X12)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num',X12)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X12),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))))))) ) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f12494,plain,(
  ( ! [X12] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X12),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X12),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num',X12)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num',X12))))) ) | $spl25),
  inference(duplicate_literal_removal,[],[f12433])).
fof(f12433,plain,(
  ( ! [X12] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X12),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X12),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num',X12)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num',X12))))) ) | $spl25),
  inference(resolution,[],[f11884,f443])).
fof(f11778,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))) | $spl177),
  inference(cnf_transformation,[],[f11778_D])).
fof(f11778_D,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))) <=> ~$spl177),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl177])])).
fof(f11779,plain,(
  ( ! [X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),X3))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),X3))))))))) ) | $spl178),
  inference(cnf_transformation,[],[f11779_D])).
fof(f11779_D,plain,(
  ( ! [X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),X3))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),X3))))))))) ) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f11756,plain,(
  ( ! [X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),X3))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),X3)))))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))))) ) | $spl25),
  inference(resolution,[],[f11199,f56])).
fof(f11199,plain,(
  ( ! [X2] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X2))))) ) | $spl25),
  inference(resolution,[],[f1569,f51])).
fof(f1569,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))))))) ) | $spl25),
  inference(superposition,[],[f1426,f53])).
fof(f8852,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | $spl172),
  inference(cnf_transformation,[],[f8852_D])).
fof(f8852_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f8853,plain,(
  ( ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X0))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))))))) ) | $spl174),
  inference(cnf_transformation,[],[f8853_D])).
fof(f8853_D,plain,(
  ( ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X0))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))))))) ) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f8854,plain,(
  ( ! [X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))))))) )),
  inference(subsumption_resolution,[],[f8822,f46])).
fof(f8822,plain,(
  ( ! [X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK2))))) )),
  inference(resolution,[],[f213,f56])).
fof(f213,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X2))))),s('type.num.num',X3)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X2))))),s('type.num.num',X3)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X3),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))))))) )),
  inference(superposition,[],[f69,f53])).
fof(f8819,plain,(
  ( ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X0))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))))))) )),
  inference(resolution,[],[f213,f69])).
fof(f8306,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ($spl16 | $spl27)),
  inference(subsumption_resolution,[],[f8305,f1817])).
fof(f8305,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ($spl16 | $spl27)),
  inference(forward_demodulation,[],[f8304,f50])).
fof(f8304,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | ($spl16 | $spl27)),
  inference(forward_demodulation,[],[f8286,f4602])).
fof(f8286,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | ($spl16 | $spl27)),
  inference(superposition,[],[f5950,f4602])).
fof(f5950,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))) ) | $spl16),
  inference(superposition,[],[f774,f53])).
fof(f8115,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl10 | $spl27)),
  inference(subsumption_resolution,[],[f8114,f1817])).
fof(f8114,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl10 | $spl27)),
  inference(forward_demodulation,[],[f8113,f50])).
fof(f8113,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | ($spl10 | $spl27)),
  inference(forward_demodulation,[],[f8098,f4602])).
fof(f8098,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) | ($spl10 | $spl27)),
  inference(superposition,[],[f5940,f4602])).
fof(f6254,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | $spl160),
  inference(cnf_transformation,[],[f6254_D])).
fof(f6254_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f6455,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl164),
  inference(cnf_transformation,[],[f6455_D])).
fof(f6455_D,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f6661,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | $spl158),
  inference(resolution,[],[f6252,f53])).
fof(f6628,plain,(
  ( ! [X27] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X27))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))) ) | $spl166),
  inference(cnf_transformation,[],[f6628_D])).
fof(f6628_D,plain,(
  ( ! [X27] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X27))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))) ) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f6658,plain,(
  ( ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X0))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))) ) | $spl158),
  inference(resolution,[],[f6252,f216])).
fof(f6453,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | $spl163),
  inference(cnf_transformation,[],[f6453_D])).
fof(f6453_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) <=> ~$spl163),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl163])])).
fof(f6627,plain,(
  ( ! [X27] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X27))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))) ) | $spl157),
  inference(forward_demodulation,[],[f6626,f904])).
fof(f6626,plain,(
  ( ! [X27] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X27))))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))))) ) | $spl157),
  inference(subsumption_resolution,[],[f6625,f6250])).
fof(f6625,plain,(
  ( ! [X27] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X27))))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))))) )),
  inference(forward_demodulation,[],[f6520,f904])).
fof(f6520,plain,(
  ( ! [X27] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X27))))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))))) )),
  inference(superposition,[],[f217,f904])).
fof(f6439,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | $spl157),
  inference(resolution,[],[f6250,f53])).
fof(f6235,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | $spl121),
  inference(resolution,[],[f5345,f79])).
fof(f5345,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | $spl121),
  inference(resolution,[],[f4420,f51])).
fof(f6252,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | $spl158),
  inference(cnf_transformation,[],[f6252_D])).
fof(f6252_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f6234,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | $spl121),
  inference(resolution,[],[f5345,f89])).
fof(f6250,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | $spl157),
  inference(cnf_transformation,[],[f6250_D])).
fof(f6250_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) <=> ~$spl157),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl157])])).
fof(f6233,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | $spl121),
  inference(resolution,[],[f5345,f120])).
fof(f6248,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X0))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))) ) | $spl154),
  inference(cnf_transformation,[],[f6248_D])).
fof(f6248_D,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X0))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f6232,plain,(
  ( ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X0))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))) ) | $spl121),
  inference(resolution,[],[f5345,f201])).
fof(f5982,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.num.|48|_0')))) | $spl153),
  inference(cnf_transformation,[],[f5982_D])).
fof(f5982_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.num.|48|_0')))) <=> ~$spl153),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl153])])).
fof(f3457,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | $spl100),
  inference(cnf_transformation,[],[f3457_D])).
fof(f3457_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f5972,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.num.|48|_0')))) | $spl13),
  inference(resolution,[],[f769,f53])).
fof(f5968,plain,(
  $false | ($spl10 | $spl12)),
  inference(subsumption_resolution,[],[f5961,f770])).
fof(f5961,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.num.|48|_0')))) | $spl10),
  inference(resolution,[],[f3434,f52])).
fof(f5933,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | $spl10),
  inference(resolution,[],[f768,f115])).
fof(f5932,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | $spl10),
  inference(resolution,[],[f768,f870])).
fof(f3455,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl99),
  inference(cnf_transformation,[],[f3455_D])).
fof(f3455_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) <=> ~$spl99),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl99])])).
fof(f768,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl10),
  inference(cnf_transformation,[],[f768_D])).
fof(f768_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f5921,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl12 | $spl27)),
  inference(forward_demodulation,[],[f5920,f1816])).
fof(f5920,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | ($spl12 | $spl27)),
  inference(subsumption_resolution,[],[f4712,f187])).
fof(f187,plain,(
  ( ! [X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X1))))))))) )),
  inference(resolution,[],[f64,f52])).
fof(f4712,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | ($spl12 | $spl27)),
  inference(superposition,[],[f3465,f1816])).
fof(f3465,plain,(
  ( ! [X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X1))))) ) | $spl12),
  inference(superposition,[],[f770,f53])).
fof(f3456,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | $spl101),
  inference(cnf_transformation,[],[f3456_D])).
fof(f3456_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) <=> ~$spl101),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl101])])).
fof(f5917,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | ($spl12 | $spl27)),
  inference(forward_demodulation,[],[f5916,f1816])).
fof(f5916,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | ($spl12 | $spl27)),
  inference(subsumption_resolution,[],[f4721,f187])).
fof(f4721,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | ($spl12 | $spl27)),
  inference(superposition,[],[f4144,f1816])).
fof(f4144,plain,(
  ( ! [X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X1) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1))))) ) | $spl12),
  inference(resolution,[],[f3465,f55])).
fof(f5915,plain,(
  $false | ($spl10 | $spl99)),
  inference(subsumption_resolution,[],[f3434,f3455])).
fof(f5884,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ($spl27 | $spl57 | $spl121)),
  inference(resolution,[],[f5666,f79])).
fof(f5666,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | ($spl27 | $spl57 | $spl121)),
  inference(resolution,[],[f5374,f51])).
fof(f5374,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ($spl27 | $spl57 | $spl121)),
  inference(backward_demodulation,[],[f5239,f4420])).
fof(f5239,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl27 | $spl57)),
  inference(forward_demodulation,[],[f5110,f4602])).
fof(f5110,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))) | $spl57),
  inference(backward_demodulation,[],[f4856,f904])).
fof(f5901,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | $spl148),
  inference(cnf_transformation,[],[f5901_D])).
fof(f5901_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f5883,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ($spl27 | $spl57 | $spl121)),
  inference(resolution,[],[f5666,f89])).
fof(f5882,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | ($spl27 | $spl57 | $spl121)),
  inference(resolution,[],[f5666,f120])).
fof(f5881,plain,(
  ( ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X0))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))))))) ) | ($spl27 | $spl57 | $spl121)),
  inference(resolution,[],[f5666,f201])).
fof(f5497,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) | $spl137),
  inference(cnf_transformation,[],[f5497_D])).
fof(f5497_D,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) <=> ~$spl137),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl137])])).
fof(f5499,plain,(
  ( ! [X1] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),X1))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),X1))))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X1))))) ) | $spl140),
  inference(cnf_transformation,[],[f5499_D])).
fof(f5499_D,plain,(
  ( ! [X1] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),X1))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),X1))))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X1))))) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f5389,plain,(
  ( ! [X1] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),X1))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),X1))))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X1))))) ) | ($spl27 | $spl57)),
  inference(superposition,[],[f49,f5239])).
fof(f5498,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),X0))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))))))) ) | $spl138),
  inference(cnf_transformation,[],[f5498_D])).
fof(f5498_D,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),X0))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))))))) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f5388,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),X0))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))) ) | ($spl27 | $spl57)),
  inference(superposition,[],[f56,f5239])).
fof(f5358,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl134),
  inference(cnf_transformation,[],[f5358_D])).
fof(f5358_D,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f5344,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl121),
  inference(resolution,[],[f4420,f53])).
fof(f5248,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) | $spl129),
  inference(cnf_transformation,[],[f5248_D])).
fof(f5248_D,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) <=> ~$spl129),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl129])])).
fof(f5250,plain,(
  ( ! [X1] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),X1))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),X1))))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X1))))) ) | $spl132),
  inference(cnf_transformation,[],[f5250_D])).
fof(f5250_D,plain,(
  ( ! [X1] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),X1))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),X1))))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X1))))) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f5127,plain,(
  ( ! [X1] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),X1))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),X1))))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X1))))) ) | $spl57),
  inference(superposition,[],[f49,f4856])).
fof(f5249,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),X0))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))))) ) | $spl130),
  inference(cnf_transformation,[],[f5249_D])).
fof(f5249_D,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),X0))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))))) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f5126,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),X0))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))) ) | $spl57),
  inference(superposition,[],[f56,f4856])).
fof(f135,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | $spl1),
  inference(cnf_transformation,[],[f135_D])).
fof(f135_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f4389,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | $spl114),
  inference(cnf_transformation,[],[f4389_D])).
fof(f4389_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f4413,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | $spl119),
  inference(cnf_transformation,[],[f4413_D])).
fof(f4413_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) <=> ~$spl119),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl119])])).
fof(f4479,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))))))),
  inference(forward_demodulation,[],[f4478,f851])).
fof(f4478,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))))))),
  inference(forward_demodulation,[],[f1654,f50])).
fof(f1654,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))))))),
  inference(superposition,[],[f125,f851])).
fof(f125,plain,(
  ( ! [X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))))))) )),
  inference(superposition,[],[f48,f53])).
fof(f3535,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | $spl102),
  inference(cnf_transformation,[],[f3535_D])).
fof(f3535_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f4477,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))))))),
  inference(forward_demodulation,[],[f4476,f851])).
fof(f4476,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))))),
  inference(forward_demodulation,[],[f1658,f50])).
fof(f1658,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))))),
  inference(superposition,[],[f149,f851])).
fof(f149,plain,(
  ( ! [X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num',X1))))) )),
  inference(superposition,[],[f122,f53])).
fof(f136,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) | $spl3),
  inference(cnf_transformation,[],[f136_D])).
fof(f136_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f4473,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))),
  inference(forward_demodulation,[],[f4472,f50])).
fof(f4472,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))))),
  inference(forward_demodulation,[],[f1675,f851])).
fof(f1675,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))))),
  inference(superposition,[],[f448,f851])).
fof(f448,plain,(
  ( ! [X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X1) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1))))) )),
  inference(resolution,[],[f125,f55])).
fof(f4469,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))),
  inference(forward_demodulation,[],[f4468,f50])).
fof(f4468,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))))),
  inference(forward_demodulation,[],[f1687,f851])).
fof(f1687,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))))),
  inference(superposition,[],[f649,f851])).
fof(f649,plain,(
  ( ! [X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X1) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1))))) )),
  inference(resolution,[],[f149,f55])).
fof(f151,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) | $spl4),
  inference(cnf_transformation,[],[f151_D])).
fof(f151_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f4467,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) | ($spl6 | $spl11)),
  inference(forward_demodulation,[],[f4466,f851])).
fof(f4466,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.num.|48|_0')))) | ($spl6 | $spl11)),
  inference(subsumption_resolution,[],[f4465,f764])).
fof(f4465,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.num.|48|_0')))) | $spl11),
  inference(forward_demodulation,[],[f4023,f851])).
fof(f4023,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.num.|48|_0')))) | $spl11),
  inference(superposition,[],[f3994,f851])).
fof(f3994,plain,(
  ( ! [X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X1) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.num.|48|_0'))))) ) | $spl11),
  inference(resolution,[],[f3448,f55])).
fof(f3448,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))))) ) | $spl11),
  inference(superposition,[],[f767,f53])).
fof(f767,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl11),
  inference(cnf_transformation,[],[f767_D])).
fof(f767_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f4464,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | ($spl6 | $spl11)),
  inference(subsumption_resolution,[],[f4463,f764])).
fof(f4463,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | $spl11),
  inference(forward_demodulation,[],[f4462,f851])).
fof(f4462,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl11),
  inference(forward_demodulation,[],[f4043,f50])).
fof(f4043,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl11),
  inference(superposition,[],[f3449,f851])).
fof(f3449,plain,(
  ( ! [X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))))) ) | $spl11),
  inference(superposition,[],[f767,f53])).
fof(f4461,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) | ($spl6 | $spl11)),
  inference(forward_demodulation,[],[f4460,f50])).
fof(f4460,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | ($spl6 | $spl11)),
  inference(subsumption_resolution,[],[f4459,f764])).
fof(f4459,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | $spl11),
  inference(forward_demodulation,[],[f4055,f851])).
fof(f4055,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | $spl11),
  inference(superposition,[],[f4035,f851])).
fof(f4035,plain,(
  ( ! [X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X1) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1))))) ) | $spl11),
  inference(resolution,[],[f3449,f55])).
fof(f4458,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) | ($spl12 | $spl95)),
  inference(forward_demodulation,[],[f4457,f851])).
fof(f4457,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.num.|48|_0')))) | ($spl12 | $spl95)),
  inference(subsumption_resolution,[],[f4456,f3505])).
fof(f3505,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | $spl95),
  inference(resolution,[],[f3380,f54])).
fof(f4456,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.num.|48|_0')))) | $spl12),
  inference(forward_demodulation,[],[f4132,f851])).
fof(f4132,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.num.|48|_0')))) | $spl12),
  inference(superposition,[],[f4103,f851])).
fof(f4103,plain,(
  ( ! [X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X1) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.num.|48|_0'))))) ) | $spl12),
  inference(resolution,[],[f3464,f55])).
fof(f3464,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X0))))) ) | $spl12),
  inference(superposition,[],[f770,f53])).
fof(f4455,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | ($spl12 | $spl95)),
  inference(subsumption_resolution,[],[f4454,f3505])).
fof(f4454,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | $spl12),
  inference(forward_demodulation,[],[f4453,f851])).
fof(f4453,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | $spl12),
  inference(forward_demodulation,[],[f4152,f50])).
fof(f4152,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | $spl12),
  inference(superposition,[],[f3465,f851])).
fof(f4452,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) | ($spl12 | $spl95)),
  inference(forward_demodulation,[],[f4451,f50])).
fof(f4451,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | ($spl12 | $spl95)),
  inference(subsumption_resolution,[],[f4450,f3505])).
fof(f4450,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | $spl12),
  inference(forward_demodulation,[],[f4164,f851])).
fof(f4164,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | $spl12),
  inference(superposition,[],[f4144,f851])).
fof(f4449,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) | ($spl11 | $spl95)),
  inference(forward_demodulation,[],[f4448,f851])).
fof(f4448,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.num.|48|_0')))) | ($spl11 | $spl95)),
  inference(subsumption_resolution,[],[f4447,f3380])).
fof(f4447,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.num.|48|_0')))) | $spl11),
  inference(forward_demodulation,[],[f4229,f851])).
fof(f4229,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.num.|48|_0')))) | $spl11),
  inference(superposition,[],[f4177,f851])).
fof(f4177,plain,(
  ( ! [X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X1) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.num.|48|_0'))))) ) | $spl11),
  inference(resolution,[],[f3480,f55])).
fof(f3480,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X0))))) ) | $spl11),
  inference(superposition,[],[f3446,f53])).
fof(f3446,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.num.|48|_0')))) | $spl11),
  inference(resolution,[],[f767,f51])).
fof(f4446,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | ($spl11 | $spl95)),
  inference(subsumption_resolution,[],[f4445,f3380])).
fof(f4445,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | $spl11),
  inference(forward_demodulation,[],[f4444,f851])).
fof(f4444,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | $spl11),
  inference(forward_demodulation,[],[f4249,f50])).
fof(f4249,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | $spl11),
  inference(superposition,[],[f3481,f851])).
fof(f3481,plain,(
  ( ! [X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X1))))) ) | $spl11),
  inference(superposition,[],[f3446,f53])).
fof(f4443,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) | ($spl11 | $spl95)),
  inference(forward_demodulation,[],[f4442,f50])).
fof(f4442,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | ($spl11 | $spl95)),
  inference(subsumption_resolution,[],[f4441,f3380])).
fof(f4441,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | $spl11),
  inference(forward_demodulation,[],[f4261,f851])).
fof(f4261,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | $spl11),
  inference(superposition,[],[f4241,f851])).
fof(f4241,plain,(
  ( ! [X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X1) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1))))) ) | $spl11),
  inference(resolution,[],[f3481,f55])).
fof(f4440,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) | ($spl90 | $spl99)),
  inference(forward_demodulation,[],[f4439,f851])).
fof(f4439,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.num.|48|_0')))) | ($spl90 | $spl99)),
  inference(subsumption_resolution,[],[f4438,f3375])).
fof(f3375,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl90),
  inference(cnf_transformation,[],[f3375_D])).
fof(f3375_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f4438,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.num.|48|_0')))) | $spl99),
  inference(forward_demodulation,[],[f4303,f851])).
fof(f4303,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.num.|48|_0')))) | $spl99),
  inference(superposition,[],[f4274,f851])).
fof(f4274,plain,(
  ( ! [X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X1) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.num.|48|_0'))))) ) | $spl99),
  inference(resolution,[],[f3486,f55])).
fof(f3486,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))))) ) | $spl99),
  inference(superposition,[],[f3455,f53])).
fof(f4437,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | ($spl90 | $spl99)),
  inference(subsumption_resolution,[],[f4436,f3375])).
fof(f4436,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | $spl99),
  inference(forward_demodulation,[],[f4435,f851])).
fof(f4435,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl99),
  inference(forward_demodulation,[],[f4327,f50])).
fof(f4327,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl99),
  inference(superposition,[],[f3487,f851])).
fof(f3487,plain,(
  ( ! [X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))))) ) | $spl99),
  inference(superposition,[],[f3455,f53])).
fof(f4434,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) | ($spl90 | $spl99)),
  inference(forward_demodulation,[],[f4433,f50])).
fof(f4433,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | ($spl90 | $spl99)),
  inference(subsumption_resolution,[],[f4432,f3375])).
fof(f4432,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | $spl99),
  inference(forward_demodulation,[],[f4339,f851])).
fof(f4339,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | $spl99),
  inference(superposition,[],[f4319,f851])).
fof(f4319,plain,(
  ( ! [X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X1) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1))))) ) | $spl99),
  inference(resolution,[],[f3487,f55])).
fof(f4431,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f137,f153])).
fof(f153,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) | $spl1),
  inference(resolution,[],[f135,f55])).
fof(f2518,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | $spl55),
  inference(cnf_transformation,[],[f2518_D])).
fof(f2518_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) <=> ~$spl55),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl55])])).
fof(f2520,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl56),
  inference(cnf_transformation,[],[f2520_D])).
fof(f2520_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f4420,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl121),
  inference(cnf_transformation,[],[f4420_D])).
fof(f4420_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) <=> ~$spl121),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl121])])).
fof(f4430,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | $spl11),
  inference(forward_demodulation,[],[f4050,f904])).
fof(f4050,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl11),
  inference(forward_demodulation,[],[f4044,f50])).
fof(f4044,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl11),
  inference(superposition,[],[f3449,f904])).
fof(f2574,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl69),
  inference(cnf_transformation,[],[f2574_D])).
fof(f2574_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl69),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl69])])).
fof(f4064,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl11),
  inference(forward_demodulation,[],[f4063,f50])).
fof(f4063,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | $spl11),
  inference(forward_demodulation,[],[f4056,f904])).
fof(f4056,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | $spl11),
  inference(superposition,[],[f4035,f904])).
fof(f4429,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | $spl12),
  inference(forward_demodulation,[],[f4159,f904])).
fof(f4159,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | $spl12),
  inference(forward_demodulation,[],[f4153,f50])).
fof(f4153,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | $spl12),
  inference(superposition,[],[f3465,f904])).
fof(f4173,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl12),
  inference(forward_demodulation,[],[f4172,f50])).
fof(f4172,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | $spl12),
  inference(forward_demodulation,[],[f4165,f904])).
fof(f4165,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | $spl12),
  inference(superposition,[],[f4144,f904])).
fof(f4425,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | $spl124),
  inference(cnf_transformation,[],[f4425_D])).
fof(f4425_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f4426,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | $spl11),
  inference(forward_demodulation,[],[f4256,f904])).
fof(f4256,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | $spl11),
  inference(forward_demodulation,[],[f4250,f50])).
fof(f4250,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | $spl11),
  inference(superposition,[],[f3481,f904])).
fof(f4270,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl11),
  inference(forward_demodulation,[],[f4269,f50])).
fof(f4269,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | $spl11),
  inference(forward_demodulation,[],[f4262,f904])).
fof(f4262,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | $spl11),
  inference(superposition,[],[f4241,f904])).
fof(f4423,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | $spl99),
  inference(forward_demodulation,[],[f4334,f904])).
fof(f4334,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl99),
  inference(forward_demodulation,[],[f4328,f50])).
fof(f4328,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl99),
  inference(superposition,[],[f3487,f904])).
fof(f4348,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl99),
  inference(forward_demodulation,[],[f4347,f50])).
fof(f4347,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | $spl99),
  inference(forward_demodulation,[],[f4340,f904])).
fof(f4340,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | $spl99),
  inference(superposition,[],[f4319,f904])).
fof(f2570,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.num.|48|_0')))) | $spl66),
  inference(cnf_transformation,[],[f2570_D])).
fof(f2570_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.num.|48|_0')))) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f4383,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl11),
  inference(forward_demodulation,[],[f4382,f50])).
fof(f4382,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | $spl11),
  inference(forward_demodulation,[],[f4362,f904])).
fof(f4362,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | $spl11),
  inference(superposition,[],[f3995,f904])).
fof(f3995,plain,(
  ( ! [X2] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X2))))) ) | $spl11),
  inference(resolution,[],[f3448,f51])).
fof(f440,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))))))),
  inference(resolution,[],[f124,f122])).
fof(f635,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))))))),
  inference(resolution,[],[f148,f122])).
fof(f3534,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | $spl103),
  inference(cnf_transformation,[],[f3534_D])).
fof(f3534_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) <=> ~$spl103),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl103])])).
fof(f669,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | $spl1),
  inference(resolution,[],[f156,f122])).
fof(f156,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))))))) ) | $spl1),
  inference(superposition,[],[f135,f53])).
fof(f1115,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | $spl1),
  inference(duplicate_literal_removal,[],[f1099])).
fof(f1099,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | $spl1),
  inference(resolution,[],[f634,f156])).
fof(f1114,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))))))),
  inference(duplicate_literal_removal,[],[f1100])).
fof(f1100,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))))))),
  inference(resolution,[],[f634,f148])).
fof(f1113,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))))))),
  inference(duplicate_literal_removal,[],[f1101])).
fof(f1101,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))))))),
  inference(resolution,[],[f634,f124])).
fof(f4416,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0'))))),
  inference(forward_demodulation,[],[f1757,f851])).
fof(f1757,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))))))),
  inference(forward_demodulation,[],[f1756,f851])).
fof(f1756,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))))))),
  inference(subsumption_resolution,[],[f1653,f122])).
fof(f1653,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))))))),
  inference(superposition,[],[f124,f851])).
fof(f4415,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0'))))),
  inference(forward_demodulation,[],[f1759,f851])).
fof(f1759,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))))),
  inference(forward_demodulation,[],[f1758,f851])).
fof(f1758,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))))),
  inference(subsumption_resolution,[],[f1657,f122])).
fof(f1657,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))))),
  inference(superposition,[],[f148,f851])).
fof(f4414,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) | $spl1),
  inference(forward_demodulation,[],[f1761,f851])).
fof(f1761,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | $spl1),
  inference(forward_demodulation,[],[f1760,f851])).
fof(f1760,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | $spl1),
  inference(subsumption_resolution,[],[f1659,f122])).
fof(f1659,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | $spl1),
  inference(superposition,[],[f156,f851])).
fof(f4411,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))))))),
  inference(forward_demodulation,[],[f1772,f50])).
fof(f1772,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))))),
  inference(forward_demodulation,[],[f1674,f851])).
fof(f1674,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))))),
  inference(superposition,[],[f439,f851])).
fof(f439,plain,(
  ( ! [X2] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X2))))) )),
  inference(resolution,[],[f124,f51])).
fof(f4410,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))))))),
  inference(forward_demodulation,[],[f1778,f50])).
fof(f1778,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))))),
  inference(forward_demodulation,[],[f1686,f851])).
fof(f1686,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))))),
  inference(superposition,[],[f634,f851])).
fof(f4409,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | $spl1),
  inference(forward_demodulation,[],[f1779,f50])).
fof(f1779,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | $spl1),
  inference(forward_demodulation,[],[f1689,f851])).
fof(f1689,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | $spl1),
  inference(superposition,[],[f668,f851])).
fof(f668,plain,(
  ( ! [X2] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X2))))) ) | $spl1),
  inference(resolution,[],[f156,f51])).
fof(f1788,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | $spl1),
  inference(forward_demodulation,[],[f1787,f851])).
fof(f1787,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.num.|48|_0')))) | $spl1),
  inference(forward_demodulation,[],[f1703,f50])).
fof(f1703,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.num.|48|_0')))) | $spl1),
  inference(trivial_inequality_removal,[],[f1701])).
fof(f1701,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.num.|48|_0')))) | $spl1),
  inference(superposition,[],[f1119,f851])).
fof(f1119,plain,(
  ( ! [X1] : (s('type.num.num',X1) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.num.|48|_0'))))) ) | $spl1),
  inference(resolution,[],[f668,f55])).
fof(f4407,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | ($spl6 | $spl11)),
  inference(subsumption_resolution,[],[f4011,f764])).
fof(f4011,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | $spl11),
  inference(duplicate_literal_removal,[],[f4000])).
fof(f4000,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | $spl11),
  inference(resolution,[],[f3448,f634])).
fof(f4403,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | ($spl12 | $spl95)),
  inference(subsumption_resolution,[],[f4120,f3505])).
fof(f4120,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | $spl12),
  inference(duplicate_literal_removal,[],[f4109])).
fof(f4109,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | $spl12),
  inference(resolution,[],[f3464,f634])).
fof(f4399,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | ($spl11 | $spl95)),
  inference(subsumption_resolution,[],[f4194,f3380])).
fof(f4194,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | $spl11),
  inference(duplicate_literal_removal,[],[f4183])).
fof(f4183,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | $spl11),
  inference(resolution,[],[f3480,f634])).
fof(f4393,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | ($spl90 | $spl99)),
  inference(subsumption_resolution,[],[f4291,f3375])).
fof(f4291,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | $spl99),
  inference(duplicate_literal_removal,[],[f4280])).
fof(f4280,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | $spl99),
  inference(resolution,[],[f3486,f634])).
fof(f4387,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) | ($spl6 | $spl11)),
  inference(forward_demodulation,[],[f4381,f50])).
fof(f4381,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | ($spl6 | $spl11)),
  inference(forward_demodulation,[],[f4380,f851])).
fof(f4380,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | ($spl6 | $spl11)),
  inference(subsumption_resolution,[],[f4361,f764])).
fof(f4361,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | $spl11),
  inference(superposition,[],[f3995,f851])).
fof(f4376,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | $spl110),
  inference(cnf_transformation,[],[f4376_D])).
fof(f4376_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f4378,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) | $spl112),
  inference(cnf_transformation,[],[f4378_D])).
fof(f4378_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f4354,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | ($spl11 | $spl26)),
  inference(resolution,[],[f3995,f1428])).
fof(f4372,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | $spl106),
  inference(cnf_transformation,[],[f4372_D])).
fof(f4372_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f4374,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.num.|48|_0')))) | $spl108),
  inference(cnf_transformation,[],[f4374_D])).
fof(f4374_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.num.|48|_0')))) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f4353,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ($spl11 | $spl56)),
  inference(resolution,[],[f3995,f2520])).
fof(f764,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl6),
  inference(cnf_transformation,[],[f764_D])).
fof(f764_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f3941,plain,(
  ( ! [X12] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X12))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))))))) ) | $spl104),
  inference(cnf_transformation,[],[f3941_D])).
fof(f3941_D,plain,(
  ( ! [X12] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X12))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))))))) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f3944,plain,(
  ( ! [X4] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X4))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))))) ) | $spl90),
  inference(forward_demodulation,[],[f3943,f851])).
fof(f3943,plain,(
  ( ! [X4] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X4))))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))))) ) | $spl90),
  inference(forward_demodulation,[],[f3942,f851])).
fof(f3942,plain,(
  ( ! [X4] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X4))))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))))) ) | $spl90),
  inference(subsumption_resolution,[],[f3930,f3375])).
fof(f3930,plain,(
  ( ! [X4] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X4))))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))))) )),
  inference(superposition,[],[f200,f851])).
fof(f3925,plain,(
  ( ! [X12] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X12))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))))))) ) | $spl90),
  inference(resolution,[],[f200,f3375])).
fof(f1556,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | $spl44),
  inference(cnf_transformation,[],[f1556_D])).
fof(f1556_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f1554,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | $spl42),
  inference(cnf_transformation,[],[f1554_D])).
fof(f1554_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f3812,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | $spl26),
  inference(resolution,[],[f1542,f511])).
fof(f1542,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | $spl26),
  inference(resolution,[],[f1428,f52])).
fof(f1558,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | $spl47),
  inference(cnf_transformation,[],[f1558_D])).
fof(f1558_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) <=> ~$spl47),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl47])])).
fof(f3564,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | $spl26),
  inference(resolution,[],[f1428,f115])).
fof(f3563,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | $spl26),
  inference(resolution,[],[f1428,f870])).
fof(f3545,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | ($spl27 | $spl90)),
  inference(resolution,[],[f3375,f2328])).
fof(f2328,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X0))))) ) | $spl27),
  inference(superposition,[],[f2311,f53])).
fof(f3536,plain,(
  $false | ($spl91 | $spl95)),
  inference(subsumption_resolution,[],[f3522,f3380])).
fof(f3522,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | $spl91),
  inference(resolution,[],[f3416,f51])).
fof(f3416,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl91),
  inference(resolution,[],[f3376,f54])).
fof(f772,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | $spl14),
  inference(cnf_transformation,[],[f772_D])).
fof(f772_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f3520,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | ($spl27 | $spl91)),
  inference(resolution,[],[f3416,f3159])).
fof(f3159,plain,(
  ( ! [X2] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X2)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X2))))) ) | $spl27),
  inference(resolution,[],[f2328,f51])).
fof(f3445,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl11),
  inference(resolution,[],[f767,f53])).
fof(f770,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.num.|48|_0')))) | $spl12),
  inference(cnf_transformation,[],[f770_D])).
fof(f770_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.num.|48|_0')))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f3444,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.num.|48|_0')))) | ($spl11 | $spl27)),
  inference(resolution,[],[f767,f3159])).
fof(f3433,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | $spl10),
  inference(resolution,[],[f768,f115])).
fof(f3432,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | $spl10),
  inference(resolution,[],[f768,f870])).
fof(f3380,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | $spl95),
  inference(cnf_transformation,[],[f3380_D])).
fof(f3380_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) <=> ~$spl95),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl95])])).
fof(f3431,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | $spl10),
  inference(resolution,[],[f768,f132])).
fof(f3415,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl91),
  inference(resolution,[],[f3376,f634])).
fof(f3283,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | $spl27),
  inference(resolution,[],[f3159,f132])).
fof(f3376,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl91),
  inference(cnf_transformation,[],[f3376_D])).
fof(f3376_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) <=> ~$spl91),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl91])])).
fof(f3282,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl27),
  inference(resolution,[],[f3159,f133])).
fof(f3024,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))) | $spl89),
  inference(cnf_transformation,[],[f3024_D])).
fof(f3024_D,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))) <=> ~$spl89),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl89])])).
fof(f2916,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | $spl74),
  inference(cnf_transformation,[],[f2916_D])).
fof(f2916_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f3022,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))) | $spl1),
  inference(forward_demodulation,[],[f3021,f957])).
fof(f3021,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.num.|48|_0')))) | $spl1),
  inference(forward_demodulation,[],[f2839,f50])).
fof(f2839,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.num.|48|_0')))) | $spl1),
  inference(superposition,[],[f1119,f957])).
fof(f2970,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))) | $spl87),
  inference(cnf_transformation,[],[f2970_D])).
fof(f2970_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))) <=> ~$spl87),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl87])])).
fof(f2940,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | $spl82),
  inference(cnf_transformation,[],[f2940_D])).
fof(f2940_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f3018,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))) | $spl4),
  inference(forward_demodulation,[],[f3017,f50])).
fof(f3017,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))))) | $spl4),
  inference(forward_demodulation,[],[f2837,f957])).
fof(f2837,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))))) | $spl4),
  inference(superposition,[],[f966,f957])).
fof(f966,plain,(
  ( ! [X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X1) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1))))) ) | $spl4),
  inference(resolution,[],[f185,f55])).
fof(f185,plain,(
  ( ! [X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num',X1))))) ) | $spl4),
  inference(superposition,[],[f151,f53])).
fof(f3016,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | $spl4),
  inference(forward_demodulation,[],[f3015,f50])).
fof(f3015,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))))) | $spl4),
  inference(forward_demodulation,[],[f2836,f957])).
fof(f2836,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))))) | $spl4),
  inference(superposition,[],[f942,f957])).
fof(f942,plain,(
  ( ! [X2] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num',X2)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X2))))) ) | $spl4),
  inference(resolution,[],[f184,f51])).
fof(f184,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num',X0))))) ) | $spl4),
  inference(superposition,[],[f151,f53])).
fof(f3000,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))) | $spl1),
  inference(forward_demodulation,[],[f2999,f50])).
fof(f2999,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))))) | $spl1),
  inference(forward_demodulation,[],[f2828,f957])).
fof(f2828,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))))) | $spl1),
  inference(superposition,[],[f683,f957])).
fof(f683,plain,(
  ( ! [X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X1) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1))))) ) | $spl1),
  inference(resolution,[],[f157,f55])).
fof(f157,plain,(
  ( ! [X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))))))) ) | $spl1),
  inference(superposition,[],[f135,f53])).
fof(f2998,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | $spl1),
  inference(forward_demodulation,[],[f2997,f50])).
fof(f2997,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))))) | $spl1),
  inference(forward_demodulation,[],[f2827,f957])).
fof(f2827,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))))) | $spl1),
  inference(superposition,[],[f668,f957])).
fof(f2928,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | $spl78),
  inference(cnf_transformation,[],[f2928_D])).
fof(f2928_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f2994,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))),
  inference(forward_demodulation,[],[f2993,f50])).
fof(f2993,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))))))),
  inference(forward_demodulation,[],[f2825,f957])).
fof(f2825,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))))))),
  inference(superposition,[],[f649,f957])).
fof(f2992,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))))),
  inference(forward_demodulation,[],[f2991,f50])).
fof(f2991,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))))))),
  inference(forward_demodulation,[],[f2824,f957])).
fof(f2824,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))))))),
  inference(superposition,[],[f634,f957])).
fof(f2918,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | $spl77),
  inference(cnf_transformation,[],[f2918_D])).
fof(f2918_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) <=> ~$spl77),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl77])])).
fof(f2968,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))),
  inference(forward_demodulation,[],[f2967,f50])).
fof(f2967,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))))))),
  inference(forward_demodulation,[],[f2813,f957])).
fof(f2813,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))))))),
  inference(superposition,[],[f448,f957])).
fof(f2964,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))))))),
  inference(forward_demodulation,[],[f2963,f50])).
fof(f2963,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))))))),
  inference(forward_demodulation,[],[f2812,f957])).
fof(f2812,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))))))),
  inference(superposition,[],[f439,f957])).
fof(f2914,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | $spl73),
  inference(cnf_transformation,[],[f2914_D])).
fof(f2914_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) <=> ~$spl73),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl73])])).
fof(f2938,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | $spl4),
  inference(forward_demodulation,[],[f2937,f957])).
fof(f2937,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))))) | $spl4),
  inference(forward_demodulation,[],[f2800,f50])).
fof(f2800,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))))) | $spl4),
  inference(superposition,[],[f185,f957])).
fof(f2932,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | $spl1),
  inference(forward_demodulation,[],[f2931,f957])).
fof(f2931,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | $spl1),
  inference(forward_demodulation,[],[f2798,f50])).
fof(f2798,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | $spl1),
  inference(superposition,[],[f157,f957])).
fof(f2926,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))))),
  inference(forward_demodulation,[],[f2925,f957])).
fof(f2925,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))))))),
  inference(forward_demodulation,[],[f2796,f50])).
fof(f2796,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))))))),
  inference(superposition,[],[f149,f957])).
fof(f2912,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))))),
  inference(forward_demodulation,[],[f2911,f957])).
fof(f2911,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))))))),
  inference(forward_demodulation,[],[f2792,f50])).
fof(f2792,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))))))),
  inference(superposition,[],[f125,f957])).
fof(f2628,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) | $spl71),
  inference(cnf_transformation,[],[f2628_D])).
fof(f2628_D,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) <=> ~$spl71),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl71])])).
fof(f2626,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) | $spl1),
  inference(forward_demodulation,[],[f2625,f904])).
fof(f2625,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.num.|48|_0')))) | $spl1),
  inference(forward_demodulation,[],[f2443,f50])).
fof(f2443,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.num.|48|_0')))) | $spl1),
  inference(superposition,[],[f1119,f904])).
fof(f2544,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | $spl64),
  inference(cnf_transformation,[],[f2544_D])).
fof(f2544_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f2622,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl4),
  inference(forward_demodulation,[],[f2621,f50])).
fof(f2621,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | $spl4),
  inference(forward_demodulation,[],[f2441,f904])).
fof(f2441,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | $spl4),
  inference(superposition,[],[f966,f904])).
fof(f2620,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | $spl4),
  inference(forward_demodulation,[],[f2619,f50])).
fof(f2619,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | $spl4),
  inference(forward_demodulation,[],[f2440,f904])).
fof(f2440,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | $spl4),
  inference(superposition,[],[f942,f904])).
fof(f2538,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | $spl63),
  inference(cnf_transformation,[],[f2538_D])).
fof(f2538_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) <=> ~$spl63),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl63])])).
fof(f2604,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl1),
  inference(forward_demodulation,[],[f2603,f50])).
fof(f2603,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | $spl1),
  inference(forward_demodulation,[],[f2432,f904])).
fof(f2432,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | $spl1),
  inference(superposition,[],[f683,f904])).
fof(f2602,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | $spl1),
  inference(forward_demodulation,[],[f2601,f50])).
fof(f2601,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | $spl1),
  inference(forward_demodulation,[],[f2431,f904])).
fof(f2431,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | $spl1),
  inference(superposition,[],[f668,f904])).
fof(f2532,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | $spl60),
  inference(cnf_transformation,[],[f2532_D])).
fof(f2532_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f2598,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))),
  inference(forward_demodulation,[],[f2597,f50])).
fof(f2597,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))))),
  inference(forward_demodulation,[],[f2429,f904])).
fof(f2429,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))))),
  inference(superposition,[],[f649,f904])).
fof(f2596,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))),
  inference(forward_demodulation,[],[f2595,f50])).
fof(f2595,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))))),
  inference(forward_demodulation,[],[f2428,f904])).
fof(f2428,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))))),
  inference(superposition,[],[f634,f904])).
fof(f2522,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | $spl59),
  inference(cnf_transformation,[],[f2522_D])).
fof(f2522_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) <=> ~$spl59),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl59])])).
fof(f2572,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))),
  inference(forward_demodulation,[],[f2571,f50])).
fof(f2571,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))))),
  inference(forward_demodulation,[],[f2417,f904])).
fof(f2417,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))))),
  inference(superposition,[],[f448,f904])).
fof(f2568,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))))))),
  inference(forward_demodulation,[],[f2567,f50])).
fof(f2567,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))))),
  inference(forward_demodulation,[],[f2416,f904])).
fof(f2416,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))))),
  inference(superposition,[],[f439,f904])).
fof(f2542,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | $spl4),
  inference(forward_demodulation,[],[f2541,f904])).
fof(f2541,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | $spl4),
  inference(forward_demodulation,[],[f2404,f50])).
fof(f2404,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | $spl4),
  inference(superposition,[],[f185,f904])).
fof(f2536,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | $spl1),
  inference(forward_demodulation,[],[f2535,f904])).
fof(f2535,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | $spl1),
  inference(forward_demodulation,[],[f2402,f50])).
fof(f2402,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | $spl1),
  inference(superposition,[],[f157,f904])).
fof(f2530,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))),
  inference(forward_demodulation,[],[f2529,f904])).
fof(f2529,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))))),
  inference(forward_demodulation,[],[f2400,f50])).
fof(f2400,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))))),
  inference(superposition,[],[f149,f904])).
fof(f2516,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))),
  inference(forward_demodulation,[],[f2515,f904])).
fof(f2515,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))))))),
  inference(forward_demodulation,[],[f2396,f50])).
fof(f2396,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))))))),
  inference(superposition,[],[f125,f904])).
fof(f1944,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))) | $spl49),
  inference(cnf_transformation,[],[f1944_D])).
fof(f1944_D,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))) <=> ~$spl49),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl49])])).
fof(f1946,plain,(
  ( ! [X1] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),X1))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),X1))))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X1))))) ) | $spl52),
  inference(cnf_transformation,[],[f1946_D])).
fof(f1946_D,plain,(
  ( ! [X1] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),X1))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),X1))))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X1))))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f1834,plain,(
  ( ! [X1] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),X1))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),X1))))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X1))))) ) | $spl27),
  inference(superposition,[],[f49,f1816])).
fof(f1945,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),X0))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))))) ) | $spl50),
  inference(cnf_transformation,[],[f1945_D])).
fof(f1945_D,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),X0))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f1833,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),X0))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))))) ) | $spl27),
  inference(superposition,[],[f56,f1816])).
fof(f1540,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | $spl26),
  inference(resolution,[],[f1428,f115])).
fof(f1539,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | $spl26),
  inference(resolution,[],[f1428,f870])).
fof(f1536,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) | $spl41),
  inference(cnf_transformation,[],[f1536_D])).
fof(f1536_D,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) <=> ~$spl41),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl41])])).
fof(f1428,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | $spl26),
  inference(cnf_transformation,[],[f1428_D])).
fof(f1428_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f1478,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) | $spl36),
  inference(cnf_transformation,[],[f1478_D])).
fof(f1478_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f1534,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) | $spl1),
  inference(forward_demodulation,[],[f1533,f844])).
fof(f1533,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))))),s('type.num.num','const.num.|48|_0')))) | $spl1),
  inference(forward_demodulation,[],[f1351,f50])).
fof(f1351,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))))),s('type.num.num','const.num.|48|_0')))) | $spl1),
  inference(superposition,[],[f1119,f844])).
fof(f1482,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl39),
  inference(cnf_transformation,[],[f1482_D])).
fof(f1482_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) <=> ~$spl39),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl39])])).
fof(f1452,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | $spl34),
  inference(cnf_transformation,[],[f1452_D])).
fof(f1452_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f1530,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl4),
  inference(forward_demodulation,[],[f1529,f50])).
fof(f1529,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))))) | $spl4),
  inference(forward_demodulation,[],[f1349,f844])).
fof(f1349,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))))) | $spl4),
  inference(superposition,[],[f966,f844])).
fof(f1528,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | $spl4),
  inference(forward_demodulation,[],[f1527,f50])).
fof(f1527,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))))) | $spl4),
  inference(forward_demodulation,[],[f1348,f844])).
fof(f1348,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))))) | $spl4),
  inference(superposition,[],[f942,f844])).
fof(f1446,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | $spl33),
  inference(cnf_transformation,[],[f1446_D])).
fof(f1446_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f1512,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl1),
  inference(forward_demodulation,[],[f1511,f50])).
fof(f1511,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))))) | $spl1),
  inference(forward_demodulation,[],[f1340,f844])).
fof(f1340,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))))) | $spl1),
  inference(superposition,[],[f683,f844])).
fof(f1510,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | $spl1),
  inference(forward_demodulation,[],[f1509,f50])).
fof(f1509,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))))) | $spl1),
  inference(forward_demodulation,[],[f1339,f844])).
fof(f1339,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))))) | $spl1),
  inference(superposition,[],[f668,f844])).
fof(f1440,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | $spl30),
  inference(cnf_transformation,[],[f1440_D])).
fof(f1440_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f1506,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))),
  inference(forward_demodulation,[],[f1505,f50])).
fof(f1505,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))))))))),
  inference(forward_demodulation,[],[f1337,f844])).
fof(f1337,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))))))))),
  inference(superposition,[],[f649,f844])).
fof(f1504,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))))))),
  inference(forward_demodulation,[],[f1503,f50])).
fof(f1503,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))))))))),
  inference(forward_demodulation,[],[f1336,f844])).
fof(f1336,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))))))))),
  inference(superposition,[],[f634,f844])).
fof(f1430,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | $spl29),
  inference(cnf_transformation,[],[f1430_D])).
fof(f1430_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f1480,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))),
  inference(forward_demodulation,[],[f1479,f50])).
fof(f1479,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))))))))),
  inference(forward_demodulation,[],[f1325,f844])).
fof(f1325,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))))))))),
  inference(superposition,[],[f448,f844])).
fof(f1476,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))))))),
  inference(forward_demodulation,[],[f1475,f50])).
fof(f1475,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))))))))),
  inference(forward_demodulation,[],[f1324,f844])).
fof(f1324,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))))))))),
  inference(superposition,[],[f439,f844])).
fof(f1426,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | $spl25),
  inference(cnf_transformation,[],[f1426_D])).
fof(f1426_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f1450,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | $spl4),
  inference(forward_demodulation,[],[f1449,f844])).
fof(f1449,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))))) | $spl4),
  inference(forward_demodulation,[],[f1312,f50])).
fof(f1312,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))))) | $spl4),
  inference(superposition,[],[f185,f844])).
fof(f1444,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | $spl1),
  inference(forward_demodulation,[],[f1443,f844])).
fof(f1443,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | $spl1),
  inference(forward_demodulation,[],[f1310,f50])).
fof(f1310,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | $spl1),
  inference(superposition,[],[f157,f844])).
fof(f1438,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))))))),
  inference(forward_demodulation,[],[f1437,f844])).
fof(f1437,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))))))))),
  inference(forward_demodulation,[],[f1308,f50])).
fof(f1308,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))))))))),
  inference(superposition,[],[f149,f844])).
fof(f1424,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))))))),
  inference(forward_demodulation,[],[f1423,f844])).
fof(f1423,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))))))),
  inference(forward_demodulation,[],[f1304,f50])).
fof(f1304,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))))))),
  inference(superposition,[],[f125,f844])).
fof(f1138,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))) | $spl21),
  inference(cnf_transformation,[],[f1138_D])).
fof(f1138_D,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f1139,plain,(
  ( ! [X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),X3))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),X3))))))))) ) | $spl22),
  inference(cnf_transformation,[],[f1139_D])).
fof(f1139_D,plain,(
  ( ! [X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),X3))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),X3))))))))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f1121,plain,(
  ( ! [X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),X3))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),X3)))))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))))) ) | $spl1),
  inference(resolution,[],[f668,f56])).
fof(f766,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | $spl8),
  inference(cnf_transformation,[],[f766_D])).
fof(f766_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f750,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))))))),
  inference(resolution,[],[f439,f372])).
fof(f372,plain,(
  ( ! [X2] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2))),s('type.num.num',X2)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num',X2))))) )),
  inference(resolution,[],[f82,f51])).
fof(f749,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))))),
  inference(resolution,[],[f439,f320])).
fof(f137,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) | $spl2),
  inference(cnf_transformation,[],[f137_D])).
fof(f137_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f138,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),s('type.num.num','const.num.|48|_0'))))),
  inference(resolution,[],[f122,f53])).
fof(f121,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))))))),
  inference(resolution,[],[f48,f53])).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 728
% Passive clauses: 5094
% Generated clauses: 18098
% Final active clauses: 489
% Final passive clauses: 2883
% Input formulas: 13
% Initial clauses: 19
% 
% Selected by SInE selection: 12
% SInE iterations: 1
% 
% Duplicate literals: 563
% Trivial inequalities: 59
% Fw subsumption resolutions: 1289
% Fw demodulations: 3947
% Bw demodulations: 59
% 
% Simple tautologies: 162
% Equational tautologies: 18
% Forward subsumptions: 6851
% 
% Binary resolution: 3573
% Factoring: 9
% Forward superposition: 5711
% Backward superposition: 2135
% Equality resolution: 113
% 
% Split clauses: 338
% Split components: 771
% SAT solver clauses: 355
% SAT solver unit clauses: 2
% SAT solver binary clauses: 19
% 
% Sat splits: 338
% Sat splitting refutations: 16
% 
% Memory used [KB]: 6652
% Time elapsed: 0.641 s
% ------------------------------
% ------------------------------
% Success in time 12.335 s
