Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.relation
% (50915)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 899
% Passive clauses: 7769
% Generated clauses: 25530
% Final active clauses: 876
% Final passive clauses: 6427
% Input formulas: 65
% Initial clauses: 231
% 
% Duplicate literals: 570
% Trivial inequalities: 125
% Fw subsumption resolutions: 2382
% Fw demodulations: 2234
% Bw demodulations: 21
% 
% Simple tautologies: 325
% Equational tautologies: 60
% Forward subsumptions: 11518
% 
% Binary resolution: 8834
% Factoring: 1486
% Forward superposition: 3690
% Backward superposition: 5641
% Self superposition: 37
% Equality factoring: 30
% Equality resolution: 21
% 
% Split clauses: 650
% Split components: 1313
% SAT solver clauses: 954
% SAT solver unit clauses: 14
% SAT solver binary clauses: 155
% 
% Sat splits: 650
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 950
% TWLsolver calls for satisfiability: 378
% 
% Memory used [KB]: 11257
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.relation
% (50943)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 449
% Passive clauses: 2712
% Generated clauses: 12560
% Final active clauses: 438
% Final passive clauses: 2250
% Input formulas: 65
% Initial clauses: 231
% 
% Duplicate literals: 214
% Trivial inequalities: 106
% Fw subsumption resolutions: 835
% Fw demodulations: 867
% Bw demodulations: 20
% Condensations: 332
% Global subsumptions: 22
% 
% Simple tautologies: 323
% Equational tautologies: 17
% Forward subsumptions: 3851
% 
% Binary resolution: 3543
% Factoring: 213
% Forward superposition: 2096
% Backward superposition: 1232
% Self superposition: 29
% Equality factoring: 26
% Equality resolution: 14
% 
% Split clauses: 83
% Split components: 171
% SAT solver clauses: 3468
% SAT solver unit clauses: 225
% SAT solver binary clauses: 86
% SAT solver learnt clauses: 44
% SAT solver learnt literals: 10
% 
% Sat splits: 83
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 3380
% TWLsolver calls for satisfiability: 1806
% 
% Memory used [KB]: 13432
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.relation
% (50959)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 2801
% Passive clauses: 14132
% Generated clauses: 212482
% Final active clauses: 2136
% Final passive clauses: 5654
% Input formulas: 65
% Initial clauses: 231
% 
% Duplicate literals: 4870
% Trivial inequalities: 365
% Fw subsumption resolutions: 9703
% Fw demodulations: 31630
% Bw demodulations: 3233
% Condensations: 15538
% Global subsumptions: 267
% 
% Simple tautologies: 377
% Equational tautologies: 1030
% Forward subsumptions: 133131
% Bw demodulations to eq. taut.: 286
% 
% Unit resulting resolution: 27455
% Factoring: 1
% Forward superposition: 65894
% Backward superposition: 48408
% Self superposition: 178
% Equality factoring: 114
% Equality resolution: 22
% 
% Split clauses: 527
% Split components: 1067
% SAT solver clauses: 16377
% SAT solver unit clauses: 1840
% SAT solver binary clauses: 4020
% 
% Sat splits: 527
% Sat splitting refutations: 922
% 
% Memory used [KB]: 101704
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.relation
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.relation
% SZS output start Proof for thm.relation
fof(f21796,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f553,f1158,f1160,f1268,f1263,f1334,f1335,f1336,f1337,f1671,f1929,f1938,f2026,f2029,f2027,f4027,f4058,f4059,f4338,f4561,f4562,f4807,f4960,f4963,f4974,f4977,f4986,f4987,f4988,f4992,f5002,f5006,f5016,f5020,f5030,f5041,f5043,f5034,f5059,f5047,f5048,f5049,f5051,f5079,f5068,f5069,f5071,f5083,f5084,f5099,f5088,f5089,f5102,f5103,f5115,f5118,f5119,f5123,f5124,f5134,f5138,f5150,f5151,f5162,f5163,f5164,f5165,f5176,f5177,f5178,f5179,f5192,f5193,f5225,f5209,f5210,f5211,f5212,f5214,f5215,f5217,f5232,f5233,f5239,f5241,f5266,f5255,f5256,f5257,f5258,f5274,f5276,f7975,f8054,f8059,f8238,f8485,f8558,f8562,f5862,f5287,f8728,f8729,f8736,f8737,f8738,f8742,f8749,f8750,f8754,f8761,f8762,f8766,f8773,f8774,f8783,f8785,f8778,f8784,f8801,f8789,f8790,f8803,f8791,f8793,f8802,f8794,f8795,f8819,f8820,f8810,f8811,f8821,f8813,f8822,f8823,f8824,f8825,f8838,f8839,f8829,f8830,f8840,f8841,f8842,f8852,f8853,f8855,f8856,f8857,f8861,f8862,f4962,f8881,f8882,f8891,f8892,f8893,f8897,f8906,f8910,f8919,f8923,f8932,f8943,f8944,f8936,f8962,f8948,f8949,f8950,f8952,f8953,f8954,f8977,f8966,f8967,f8969,f8978,f8979,f8980,f8994,f8984,f8985,f8995,f8996,f8997,f9008,f9010,f9011,f9012,f9016,f9017,f9027,f9031,f9040,f9041,f9051,f9052,f9053,f9054,f9064,f9065,f9066,f9067,f9079,f9080,f9115,f9095,f9096,f9097,f9098,f9100,f9101,f9102,f9103,f9106,f9107,f9108,f9229,f9219,f9218,f9217,f9304,f9305,f9306,f9307,f9308,f9309,f9316,f9317,f9322,f9323,f9325,f9328,f9330,f9354,f9343,f9344,f9345,f9346,f9359,f9361,f9515,f9514,f9562,f9565,f9566,f9584,f9583,f9612,f9624,f9625,f9630,f9638,f9639,f9640,f9641,f9646,f9660,f9647,f9659,f9658,f9657,f9674,f9673,f9688,f9687,f10098,f10097,f10112,f10111,f10128,f10127,f10145,f10144,f13161,f13164,f13165,f13244,f13260,f13261,f13266,f13410,f13412,f13411,f13415,f13502,f13522,f9969,f670,f668,f13693,f13707,f13708,f13717,f13718,f13719,f13723,f5001,f13732,f13736,f5015,f13745,f13749,f5029,f13758,f13769,f13770,f13762,f5042,f13789,f13774,f13775,f5061,f13776,f13778,f5060,f13779,f13780,f13781,f5064,f13806,f5080,f13795,f13796,f5081,f13798,f5082,f13824,f5100,f13810,f13811,f13812,f8835,f13813,f13825,f13814,f13815,f13840,f5116,f13829,f13830,f13831,f5113,f13832,f13842,f13791,f13846,f13847,f5131,f13848,f9024,f13850,f5063,f13851,f13860,f13864,f5149,f13873,f13874,f13884,f13885,f5161,f13886,f13887,f13897,f13898,f5175,f13899,f13900,f5189,f13912,f13913,f5203,f13949,f13928,f13929,f5228,f13930,f5229,f13931,f5226,f13933,f13934,f13935,f13951,f13936,f13937,f13940,f13941,f1940,f13942,f9117,f13954,f5251,f5145,f13955,f13960,f13961,f13964,f13967,f13969,f9339,f13992,f5267,f13981,f13982,f1933,f13983,f5270,f13984,f5269,f13997,f13999,f5286,f14021,f14020,f4976,f14606,f14605,f14622,f14621,f14640,f14639,f5146,f15182,f13416,f16668,f16905,f16908,f16909,f16910,f4342,f16981,f8570,f16982,f13531,f16997,f16998,f8572,f17000,f17186,f17188,f17189,f17190,f17187,f17191,f17193,f17192,f17194,f17197,f17196,f17198,f17199,f5096,f17195,f17200,f17201,f17183,f17182,f17204,f17206,f17202,f17550,f13506,f17574,f17554,f4042,f17953,f17955,f17954,f17956,f17959,f17960,f17961,f17957,f17964,f18033,f18034,f18035,f18302,f18440,f18443,f18444,f18445,f8806,f18503,f8850,f18504,f18515,f18538,f8868,f18539,f8867,f18541,f18455,f18542,f18552,f18543,f8805,f19293,f17965,f19991,f20158,f20160,f1159,f20329,f20331,f20330,f20332,f20334,f20335,f20336,f17958,f20333,f20340,f20341,f20423,f20422,f20426,f20427,f20428,f2028,f20429,f5101,f21041,f1931,f4815,f21789])).
fof(f21789,plain,(
  $false | $spl2),
  inference(subsumption_resolution,[],[f21674,f670])).
fof(f21674,plain,(
  ~p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))))),
  inference(resolution,[],[f676,f321])).
fof(f321,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2))))) )),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ! [X0,X1,X2] : p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2))))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ! [X0] : ! [X1,X2] : p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2))))),
  inference(rectify,[],[f48])).
fof(f48,axiom,(
  ! [X6] : ! [X7,X8] : p(s(bool,'const.relation.RTC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X8),s(X6,X8))))),
  file('thm.relation.RC_RTC','thm.relation.RTC_RULES|split|0')).
fof(f676,plain,(
  ( ! [X0] : (~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,X0),s(sK12,sK15)))) | ~p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,X0))))) )),
  inference(resolution,[],[f319,f320])).
fof(f320,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f160])).
fof(f160,plain,(
  ! [X0,X1,X2,X3,X4] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4)))))),
  inference(flattening,[],[f159])).
fof(f159,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4))))) | p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4)))))),
  inference(ennf_transformation,[],[f70])).
fof(f70,plain,(
  ! [X0,X1,X2,X3,X4] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4))))) => p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4)))))),
  inference(flattening,[],[f69])).
fof(f69,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4))))) => p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4)))))),
  inference(rectify,[],[f49])).
fof(f49,axiom,(
  ! [X6] : ! [X7,X8,X9,X18] : ((p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X8))),s(X6,X9)))) & p(s(bool,'const.relation.RTC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X9),s(X6,X18))))) => p(s(bool,'const.relation.RTC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X8),s(X6,X18)))))),
  file('thm.relation.RC_RTC','thm.relation.RTC_RULES|split|1')).
fof(f319,plain,(
  ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))),
  inference(cnf_transformation,[],[f211])).
fof(f211,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) & ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13,sK14,sK15])],[f158])).
fof(f158,plain,(
  ? [X0,X1,X2,X3] : (p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) & ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))))),
  inference(ennf_transformation,[],[f68])).
fof(f68,plain,(
  ~! [X0,X1,X2,X3] : (p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) => p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))))),
  inference(flattening,[],[f67])).
fof(f67,plain,(
  ~! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) => p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : (p(s(bool,'const.relation.RC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X8),s(X6,X9)))) => p(s(bool,'const.relation.RTC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X8),s(X6,X9)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : (p(s(bool,'const.relation.RC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X8),s(X6,X9)))) => p(s(bool,'const.relation.RTC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X8),s(X6,X9)))))),
  file('thm.relation.RC_RTC','thm.relation.RC_RTC')).
fof(f4815,plain,(
  ( ! [X38] : (p(s(bool,X38)) | s(bool,X38) != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) ) | $spl26),
  inference(cnf_transformation,[],[f4815_D])).
fof(f4815_D,plain,(
  ( ! [X38] : (p(s(bool,X38)) | s(bool,X38) != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f1931,plain,(
  ( ! [X8] : (~p(s(bool,X8))) ) | $spl8),
  inference(cnf_transformation,[],[f1931_D])).
fof(f1931_D,plain,(
  ( ! [X8] : (~p(s(bool,X8))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f21041,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X24) != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))) | p(s(bool,X24))) ) | ($spl2 | $spl6)),
  inference(resolution,[],[f17804,f3366])).
fof(f3366,plain,(
  ( ! [X101,X102] : (sP8(X101,X102,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) ) | $spl2),
  inference(resolution,[],[f670,f381])).
fof(f381,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f240])).
fof(f240,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f239])).
fof(f239,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f238])).
fof(f238,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f204])).
fof(f204,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f17804,plain,(
  ( ! [X2,X0,X3,X1] : (~sP8(X2,X1,X3) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X0) != s(bool,X3) | p(s(bool,X0))) ) | $spl6),
  inference(subsumption_resolution,[],[f17803,f391])).
fof(f391,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f245])).
fof(f245,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(flattening,[],[f244])).
fof(f244,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(nnf_transformation,[],[f202])).
fof(f202,plain,(
  ! [X0,X1,X2] : (sP6(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f17803,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X0) != s(bool,X3) | ~sP8(X2,X1,X3) | ~sP6(X1,X3,X2)) ) | $spl6),
  inference(subsumption_resolution,[],[f17802,f385])).
fof(f385,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f243])).
fof(f243,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f242])).
fof(f242,plain,(
  ! [X1,X0,X2] : ((~sP7(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP7(X1,X0,X2)))),
  inference(flattening,[],[f241])).
fof(f241,plain,(
  ! [X1,X0,X2] : ((~sP7(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP7(X1,X0,X2)))),
  inference(nnf_transformation,[],[f203])).
fof(f203,plain,(
  ! [X1,X0,X2] : (sP7(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f17802,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X0) != s(bool,X3) | ~sP8(X2,X1,X3) | ~sP7(X3,X1,X2) | ~sP6(X1,X3,X2)) ) | $spl6),
  inference(duplicate_literal_removal,[],[f17759])).
fof(f17759,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X0) != s(bool,X3) | ~p(s(bool,X1)) | ~sP8(X2,X1,X3) | ~sP7(X3,X1,X2) | ~sP6(X1,X3,X2)) ) | $spl6),
  inference(resolution,[],[f2034,f377])).
fof(f377,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~p(s(bool,X2)) | ~sP8(X0,X2,X1) | ~sP7(X1,X2,X0) | ~sP6(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f237])).
fof(f237,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | (sP6(X2,X1,X0) & sP7(X1,X2,X0) & sP8(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP6(X2,X1,X0) | ~sP7(X1,X2,X0) | ~sP8(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f236])).
fof(f236,plain,(
  ! [X2,X1,X0] : ((~sP9(X2,X1,X0) | (sP6(X0,X1,X2) & sP7(X1,X0,X2) & sP8(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP6(X0,X1,X2) | ~sP7(X1,X0,X2) | ~sP8(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP9(X2,X1,X0)))),
  inference(flattening,[],[f235])).
fof(f235,plain,(
  ! [X2,X1,X0] : ((~sP9(X2,X1,X0) | (sP6(X0,X1,X2) & sP7(X1,X0,X2) & sP8(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP6(X0,X1,X2) | ~sP7(X1,X0,X2) | ~sP8(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP9(X2,X1,X0)))),
  inference(nnf_transformation,[],[f205])).
fof(f205,plain,(
  ! [X2,X1,X0] : (sP9(X2,X1,X0) <=> (sP6(X0,X1,X2) & sP7(X1,X0,X2) & sP8(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f2034,plain,(
  ( ! [X10,X8,X11,X9] : (~sP9(X11,X9,X8) | p(s(bool,X10)) | ~p(s(bool,X8)) | ~p(s(bool,X11)) | s(bool,X9) != s(bool,X10)) ) | $spl6),
  inference(superposition,[],[f2022,f394])).
fof(f394,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X2)) | ~sP9(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f246])).
fof(f246,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP9(X2,X1,X0)) & (~sP9(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f206])).
fof(f206,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP9(X2,X1,X0))),
  inference(definition_folding,[],[f80,f205,f204,f203,f202])).
fof(f80,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f79])).
fof(f79,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f55])).
fof(f55,axiom,(
  ! [X19,X20,X21] : ((p(s(bool,X21)) <=> s(bool,X19) = s(bool,X20)) <=> ((p(s(bool,X21)) | p(s(bool,X20)) | p(s(bool,X19))) & (p(s(bool,X21)) | ~p(s(bool,X19)) | ~p(s(bool,X20))) & (p(s(bool,X20)) | ~p(s(bool,X19)) | ~p(s(bool,X21))) & (p(s(bool,X19)) | ~p(s(bool,X20)) | ~p(s(bool,X21)))))),
  file('thm.relation.RC_RTC','thm.sat.sat_2|49|_prove_ltb')).
fof(f2022,plain,(
  ( ! [X35,X34] : (s(bool,X34) != s(bool,X35) | p(s(bool,X35)) | ~p(s(bool,X34))) ) | $spl6),
  inference(trivial_inequality_removal,[],[f1997])).
fof(f1997,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X34) != s(bool,X35) | p(s(bool,X35)) | ~p(s(bool,X34))) ) | $spl6),
  inference(superposition,[],[f1351,f513])).
fof(f513,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f305])).
fof(f305,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f130])).
fof(f130,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X10] : (s(bool,'T_0') = s(bool,X10) <=> p(s(bool,X10)))),
  file('thm.relation.RC_RTC','thm.bool.EQ_CLAUSES|split|0')).
fof(f1351,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X5) | s(bool,X4) != s(bool,X5) | p(s(bool,X4))) ) | $spl6),
  inference(superposition,[],[f1160,f509])).
fof(f509,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f303])).
fof(f303,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f128])).
fof(f128,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f127])).
fof(f127,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X10] : (s(bool,'F_0') = s(bool,X10) <=> ~p(s(bool,X10)))),
  file('thm.relation.RC_RTC','thm.bool.EQ_CLAUSES|split|3')).
fof(f5101,plain,(
  ( ! [X10] : (s(bool,'F_0') != s(bool,X10)) ) | $spl58),
  inference(cnf_transformation,[],[f5101_D])).
fof(f5101_D,plain,(
  ( ! [X10] : (s(bool,'F_0') != s(bool,X10)) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f20429,plain,(
  ( ! [X43,X42] : (s(bool,'F_0') != s(bool,X43) | ~p(s(bool,X42))) ) | $spl174),
  inference(subsumption_resolution,[],[f20406,f17958])).
fof(f20406,plain,(
  ( ! [X43,X42] : (s(bool,'F_0') != s(bool,X43) | ~p(s(bool,X42)) | s(bool,'T_0') = s(bool,X42)) ) | $spl174),
  inference(superposition,[],[f17958,f539])).
fof(f539,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.relation.RC_RTC',aHL_BOOL_CASES)).
fof(f2028,plain,(
  ( ! [X23] : (s(bool,'T_0') != s(bool,X23)) ) | $spl16),
  inference(cnf_transformation,[],[f2028_D])).
fof(f2028_D,plain,(
  ( ! [X23] : (s(bool,'T_0') != s(bool,X23)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f20428,plain,(
  ( ! [X41,X40] : (s(bool,'T_0') != s(bool,X41) | ~p(s(bool,X40))) ) | $spl174),
  inference(subsumption_resolution,[],[f20405,f17958])).
fof(f20405,plain,(
  ( ! [X41,X40] : (s(bool,'T_0') != s(bool,X41) | ~p(s(bool,X40)) | s(bool,'F_0') = s(bool,X40)) ) | $spl174),
  inference(superposition,[],[f17958,f539])).
fof(f20427,plain,(
  ( ! [X39,X38] : (s(bool,'F_0') != s(bool,X39) | ~p(s(bool,X38))) ) | $spl174),
  inference(subsumption_resolution,[],[f20404,f17958])).
fof(f20404,plain,(
  ( ! [X39,X38] : (s(bool,'F_0') != s(bool,X39) | ~p(s(bool,X38)) | s(bool,'T_0') = s(bool,X38)) ) | $spl174),
  inference(superposition,[],[f17958,f538])).
fof(f538,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f153])).
fof(f153,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X10] : (s(bool,'T_0') = s(bool,X10) | s(bool,'F_0') = s(bool,X10))),
  file('thm.relation.RC_RTC','thm.bool.BOOL_CASES_AX')).
fof(f20426,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') != s(bool,X37) | ~p(s(bool,X36))) ) | $spl174),
  inference(subsumption_resolution,[],[f20403,f17958])).
fof(f20403,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') != s(bool,X37) | ~p(s(bool,X36)) | s(bool,'F_0') = s(bool,X36)) ) | $spl174),
  inference(superposition,[],[f17958,f538])).
fof(f20422,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') != s(bool,X35) | ~p(s(bool,X34))) ) | $spl174),
  inference(duplicate_literal_removal,[],[f20402])).
fof(f20402,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') != s(bool,X35) | ~p(s(bool,X34)) | ~p(s(bool,X34))) ) | $spl174),
  inference(superposition,[],[f17958,f513])).
fof(f20423,plain,(
  ( ! [X33,X32] : (s(bool,'T_0') != s(bool,X33) | ~p(s(bool,X32))) ) | $spl174),
  inference(duplicate_literal_removal,[],[f20401])).
fof(f20401,plain,(
  ( ! [X33,X32] : (s(bool,'T_0') != s(bool,X33) | ~p(s(bool,X32)) | ~p(s(bool,X32))) ) | $spl174),
  inference(superposition,[],[f17958,f511])).
fof(f511,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f304])).
fof(f304,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f129])).
fof(f129,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X10] : (s(bool,'T_0') = s(bool,X10) <=> p(s(bool,X10)))),
  file('thm.relation.RC_RTC','thm.bool.EQ_CLAUSES|split|1')).
fof(f20341,plain,(
  ( ! [X112,X113] : (p(s(bool,X112)) | p(s(bool,X113)) | sP9(X113,X112,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | $spl202),
  inference(cnf_transformation,[],[f20341_D])).
fof(f20341_D,plain,(
  ( ! [X112,X113] : (p(s(bool,X112)) | p(s(bool,X113)) | sP9(X113,X112,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f20340,plain,(
  ( ! [X111,X112,X113] : (s(bool,'T_0') != s(bool,X111) | p(s(bool,X112)) | sP9(X113,X112,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | p(s(bool,X113))) ) | ($spl6 | $spl28)),
  inference(subsumption_resolution,[],[f20283,f510])).
fof(f510,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f304])).
fof(f20283,plain,(
  ( ! [X111,X112,X113] : (s(bool,'T_0') != s(bool,X111) | p(s(bool,X112)) | ~p(s(bool,X111)) | sP9(X113,X112,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | p(s(bool,X113))) ) | ($spl6 | $spl28)),
  inference(superposition,[],[f2061,f18636])).
fof(f18636,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | $spl28),
  inference(resolution,[],[f18351,f511])).
fof(f18351,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | $spl28),
  inference(resolution,[],[f9439,f791])).
fof(f791,plain,(
  ( ! [X128] : (~sP11('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X128) | p(s(bool,X128))) )),
  inference(subsumption_resolution,[],[f745,f480])).
fof(f480,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f285])).
fof(f285,plain,(
  ! [X0] : (((~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | p(s(bool,X0))) & (~p(s(bool,X0)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f284])).
fof(f284,plain,(
  ! [X0] : (((~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f116])).
fof(f116,plain,(
  ! [X0] : ((p(s(bool,X0)) | p(s(bool,'F_0'))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X10] : ((p(s(bool,X10)) | p(s(bool,'F_0'))) <=> p(s(bool,X10)))),
  file('thm.relation.RC_RTC','thm.bool.OR_CLAUSES|split|3')).
fof(f745,plain,(
  ( ! [X128] : (p(s(bool,'F_0')) | p(s(bool,X128)) | ~sP11('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X128)) )),
  inference(resolution,[],[f319,f432])).
fof(f432,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP11(X0,X1)) )),
  inference(cnf_transformation,[],[f262])).
fof(f262,plain,(
  ! [X0,X1] : ((~sP11(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP11(X0,X1)))),
  inference(flattening,[],[f261])).
fof(f261,plain,(
  ! [X0,X1] : ((~sP11(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP11(X0,X1)))),
  inference(nnf_transformation,[],[f209])).
fof(f209,plain,(
  ! [X0,X1] : (sP11(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f9439,plain,(
  ( ! [X126] : (sP11(X126,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | $spl28),
  inference(resolution,[],[f4962,f433])).
fof(f433,plain,(
  ( ! [X0,X1] : (sP11(X0,X1) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f262])).
fof(f2061,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X5) != s(bool,X6) | p(s(bool,X4)) | ~p(s(bool,X6)) | sP9(X7,X4,X5) | p(s(bool,X7))) ) | $spl6),
  inference(superposition,[],[f2022,f392])).
fof(f392,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f246])).
fof(f20333,plain,(
  ( ! [X78,X79] : (sP9(X79,X78,'T_0') | p(s(bool,X79)) | p(s(bool,X78))) ) | $spl200),
  inference(cnf_transformation,[],[f20333_D])).
fof(f20333_D,plain,(
  ( ! [X78,X79] : (sP9(X79,X78,'T_0') | p(s(bool,X79)) | p(s(bool,X78))) ) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f17958,plain,(
  ( ! [X76,X77] : (s(bool,X76) != s(bool,X77) | ~p(s(bool,X77))) ) | $spl174),
  inference(cnf_transformation,[],[f17958_D])).
fof(f17958_D,plain,(
  ( ! [X76,X77] : (s(bool,X76) != s(bool,X77) | ~p(s(bool,X77))) ) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f20336,plain,(
  ( ! [X90,X88,X91,X89] : (s(bool,X88) != s(bool,X89) | p(s(bool,X90)) | sP9(X91,X90,'T_0') | p(s(bool,X91)) | ~p(s(bool,X88))) ) | $spl6),
  inference(subsumption_resolution,[],[f20277,f2022])).
fof(f20277,plain,(
  ( ! [X90,X88,X91,X89] : (s(bool,X88) != s(bool,X89) | p(s(bool,X90)) | ~p(s(bool,X89)) | sP9(X91,X90,'T_0') | p(s(bool,X91)) | ~p(s(bool,X88))) ) | $spl6),
  inference(superposition,[],[f2061,f511])).
fof(f20335,plain,(
  ( ! [X87,X85,X86,X84] : (s(bool,X84) != s(bool,X85) | p(s(bool,X86)) | sP9(X87,X86,'T_0') | p(s(bool,X87)) | ~p(s(bool,X84))) ) | $spl6),
  inference(subsumption_resolution,[],[f20276,f2022])).
fof(f20276,plain,(
  ( ! [X87,X85,X86,X84] : (s(bool,X84) != s(bool,X85) | p(s(bool,X86)) | ~p(s(bool,X85)) | sP9(X87,X86,'T_0') | p(s(bool,X87)) | ~p(s(bool,X84))) ) | $spl6),
  inference(superposition,[],[f2061,f513])).
fof(f20334,plain,(
  ( ! [X80,X83,X81,X82] : (s(bool,X80) != s(bool,X81) | p(s(bool,X82)) | ~p(s(bool,X81)) | sP9(X83,X82,'T_0') | p(s(bool,X83))) ) | ($spl6 | $spl24)),
  inference(subsumption_resolution,[],[f20275,f4562])).
fof(f20275,plain,(
  ( ! [X80,X83,X81,X82] : (s(bool,X80) != s(bool,X81) | p(s(bool,X82)) | ~p(s(bool,X81)) | sP9(X83,X82,'T_0') | p(s(bool,X83)) | s(bool,'F_0') = s(bool,X80)) ) | $spl6),
  inference(superposition,[],[f2061,f538])).
fof(f20332,plain,(
  ( ! [X78,X76,X79,X77] : (s(bool,X76) != s(bool,X77) | p(s(bool,X78)) | ~p(s(bool,X77)) | sP9(X79,X78,'T_0') | p(s(bool,X79))) ) | ($spl6 | $spl24)),
  inference(subsumption_resolution,[],[f20274,f4562])).
fof(f20274,plain,(
  ( ! [X78,X76,X79,X77] : (s(bool,X76) != s(bool,X77) | p(s(bool,X78)) | ~p(s(bool,X77)) | sP9(X79,X78,'T_0') | p(s(bool,X79)) | s(bool,'F_0') = s(bool,X76)) ) | $spl6),
  inference(superposition,[],[f2061,f539])).
fof(f20330,plain,(
  ( ! [X62,X60,X63] : (sP9(X63,X62,X60) | s(bool,'F_0') = s(bool,X60) | p(s(bool,X63)) | p(s(bool,X62))) ) | $spl198),
  inference(cnf_transformation,[],[f20330_D])).
fof(f20330_D,plain,(
  ( ! [X62,X60,X63] : (sP9(X63,X62,X60) | s(bool,'F_0') = s(bool,X60) | p(s(bool,X63)) | p(s(bool,X62))) ) <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f20331,plain,(
  ( ! [X70,X68,X71,X69] : (s(bool,'T_0') != s(bool,X69) | p(s(bool,X70)) | sP9(X71,X70,X68) | p(s(bool,X71)) | s(bool,'F_0') = s(bool,X68)) ) | $spl6),
  inference(subsumption_resolution,[],[f20272,f510])).
fof(f20272,plain,(
  ( ! [X70,X68,X71,X69] : (s(bool,'T_0') != s(bool,X69) | p(s(bool,X70)) | ~p(s(bool,X69)) | sP9(X71,X70,X68) | p(s(bool,X71)) | s(bool,'F_0') = s(bool,X68)) ) | $spl6),
  inference(superposition,[],[f2061,f539])).
fof(f20329,plain,(
  ( ! [X61,X62,X60,X63] : (s(bool,'T_0') != s(bool,X61) | p(s(bool,X62)) | sP9(X63,X62,X60) | p(s(bool,X63)) | s(bool,'F_0') = s(bool,X60)) ) | $spl6),
  inference(subsumption_resolution,[],[f20270,f510])).
fof(f20270,plain,(
  ( ! [X61,X62,X60,X63] : (s(bool,'T_0') != s(bool,X61) | p(s(bool,X62)) | ~p(s(bool,X61)) | sP9(X63,X62,X60) | p(s(bool,X63)) | s(bool,'F_0') = s(bool,X60)) ) | $spl6),
  inference(superposition,[],[f2061,f538])).
fof(f1159,plain,(
  ( ! [X3] : (p(s(bool,X3))) ) | $spl4),
  inference(cnf_transformation,[],[f1159_D])).
fof(f1159_D,plain,(
  ( ! [X3] : (p(s(bool,X3))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f20160,plain,(
  ( ! [X17,X15] : (s(bool,'F_0') != s(bool,X17) | s(bool,X15) != s(bool,X17) | s(bool,'F_0') = s(bool,X15)) ) | $spl196),
  inference(cnf_transformation,[],[f20160_D])).
fof(f20160_D,plain,(
  ( ! [X17,X15] : (s(bool,'F_0') != s(bool,X17) | s(bool,X15) != s(bool,X17) | s(bool,'F_0') = s(bool,X15)) ) <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f20158,plain,(
  ( ! [X17,X15,X16] : (s(bool,'F_0') = s(bool,X15) | p(s(bool,X16)) | s(bool,'F_0') != s(bool,X17) | s(bool,X15) != s(bool,X17)) ) | ($spl166 | $spl170)),
  inference(duplicate_literal_removal,[],[f20136])).
fof(f20136,plain,(
  ( ! [X17,X15,X16] : (s(bool,'F_0') = s(bool,X15) | p(s(bool,X16)) | s(bool,'F_0') != s(bool,X17) | p(s(bool,X16)) | s(bool,X15) != s(bool,X17)) ) | ($spl166 | $spl170)),
  inference(resolution,[],[f18736,f395])).
fof(f395,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | s(bool,X0) != s(bool,X1) | ~sP9(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f246])).
fof(f18736,plain,(
  ( ! [X26,X24,X25] : (sP9(X25,X24,X26) | s(bool,'F_0') = s(bool,X24) | p(s(bool,X25)) | s(bool,'F_0') != s(bool,X26)) ) | ($spl166 | $spl170)),
  inference(resolution,[],[f18701,f17475])).
fof(f17475,plain,(
  ( ! [X6,X8,X7] : (~sP6(X6,X8,X7) | sP9(X7,X8,X6) | p(s(bool,X7)) | s(bool,'F_0') != s(bool,X6)) ) | $spl166),
  inference(subsumption_resolution,[],[f17468,f383])).
fof(f383,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f240])).
fof(f17468,plain,(
  ( ! [X6,X8,X7] : (s(bool,'F_0') != s(bool,X6) | sP9(X7,X8,X6) | p(s(bool,X7)) | ~sP8(X7,X6,X8) | ~sP6(X6,X8,X7)) ) | $spl166),
  inference(resolution,[],[f17459,f379])).
fof(f379,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | p(s(bool,X0)) | ~sP8(X0,X2,X1) | ~sP7(X1,X2,X0) | ~sP6(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f237])).
fof(f17459,plain,(
  ( ! [X4,X5,X3] : (sP7(X5,X4,X3) | s(bool,'F_0') != s(bool,X4)) ) | $spl166),
  inference(subsumption_resolution,[],[f17458,f385])).
fof(f17458,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | s(bool,'F_0') != s(bool,X4) | sP7(X5,X4,X3)) ) | $spl166),
  inference(subsumption_resolution,[],[f17433,f387])).
fof(f387,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f243])).
fof(f17433,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | s(bool,'F_0') != s(bool,X4) | ~p(s(bool,X5)) | sP7(X5,X4,X3)) ) | $spl166),
  inference(resolution,[],[f17202,f374])).
fof(f374,plain,(
  ( ! [X2,X0,X1] : (sP7(X1,X2,X0) | ~sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f237])).
fof(f18701,plain,(
  ( ! [X2,X0,X1] : (sP6(X2,X0,X1) | s(bool,'F_0') = s(bool,X0)) ) | $spl170),
  inference(subsumption_resolution,[],[f18700,f389])).
fof(f389,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f245])).
fof(f18700,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X1)) | sP6(X2,X0,X1)) ) | $spl170),
  inference(subsumption_resolution,[],[f18670,f391])).
fof(f18670,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | sP6(X2,X0,X1)) ) | $spl170),
  inference(resolution,[],[f17954,f373])).
fof(f373,plain,(
  ( ! [X2,X0,X1] : (sP6(X2,X1,X0) | ~sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f237])).
fof(f19991,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | p(s(bool,X11))) ) | ($spl28 | $spl112)),
  inference(subsumption_resolution,[],[f19990,f511])).
fof(f19990,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') != s(bool,X10) | ~p(s(bool,X10)) | p(s(bool,X11))) ) | ($spl28 | $spl112)),
  inference(forward_demodulation,[],[f19969,f18636])).
fof(f19969,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | $spl112),
  inference(resolution,[],[f8806,f395])).
fof(f17965,plain,(
  ( ! [X109,X110] : (p(s(bool,X109)) | p(s(bool,X110)) | sP9(X110,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X109)) ) | $spl176),
  inference(cnf_transformation,[],[f17965_D])).
fof(f17965_D,plain,(
  ( ! [X109,X110] : (p(s(bool,X109)) | p(s(bool,X110)) | sP9(X110,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X109)) ) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f19293,plain,(
  ( ! [X315,X316,X317] : (s(bool,'T_0') != s(bool,X315) | p(s(bool,X316)) | sP9(X317,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X316) | p(s(bool,X317))) ) | ($spl6 | $spl28)),
  inference(subsumption_resolution,[],[f19145,f510])).
fof(f19145,plain,(
  ( ! [X315,X316,X317] : (s(bool,'T_0') != s(bool,X315) | p(s(bool,X316)) | ~p(s(bool,X315)) | sP9(X317,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X316) | p(s(bool,X317))) ) | ($spl6 | $spl28)),
  inference(superposition,[],[f2060,f18636])).
fof(f2060,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X1) != s(bool,X2) | p(s(bool,X0)) | ~p(s(bool,X2)) | sP9(X3,X1,X0) | p(s(bool,X3))) ) | $spl6),
  inference(superposition,[],[f2022,f392])).
fof(f8805,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | $spl111),
  inference(cnf_transformation,[],[f8805_D])).
fof(f8805_D,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) <=> ~$spl111),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl111])])).
fof(f18543,plain,(
  ( ! [X18] : (p(s(bool,X18)) | s(bool,'F_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | $spl28),
  inference(resolution,[],[f9428,f1855])).
fof(f1855,plain,(
  ( ! [X4,X5] : (~sP6(X5,'F_0',X4) | p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5)) )),
  inference(duplicate_literal_removal,[],[f1850])).
fof(f1850,plain,(
  ( ! [X4,X5] : (p(s(bool,X4)) | ~sP6(X5,'F_0',X4) | p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5)) )),
  inference(resolution,[],[f1823,f395])).
fof(f1823,plain,(
  ( ! [X6,X7] : (sP9(X6,'F_0',X7) | p(s(bool,X6)) | ~sP6(X7,'F_0',X6)) )),
  inference(subsumption_resolution,[],[f1821,f383])).
fof(f1821,plain,(
  ( ! [X6,X7] : (sP9(X6,'F_0',X7) | p(s(bool,X6)) | ~sP8(X6,X7,'F_0') | ~sP6(X7,'F_0',X6)) )),
  inference(resolution,[],[f1817,f379])).
fof(f1817,plain,(
  ( ! [X4,X5] : (sP7('F_0',X4,X5)) )),
  inference(subsumption_resolution,[],[f1813,f385])).
fof(f1813,plain,(
  ( ! [X4,X5] : (sP7('F_0',X4,X5) | p(s(bool,X5))) )),
  inference(resolution,[],[f1810,f383])).
fof(f1810,plain,(
  ( ! [X6,X7] : (~sP8(X7,X6,'F_0') | sP7('F_0',X6,X7)) )),
  inference(subsumption_resolution,[],[f1806,f386])).
fof(f386,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f243])).
fof(f1806,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | ~sP8(X7,X6,'F_0') | sP7('F_0',X6,X7)) )),
  inference(resolution,[],[f1800,f374])).
fof(f1800,plain,(
  ( ! [X2,X3] : (sP9(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP8(X3,X2,'F_0')) )),
  inference(subsumption_resolution,[],[f1799,f511])).
fof(f1799,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | sP9(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP8(X3,X2,'F_0')) )),
  inference(subsumption_resolution,[],[f1796,f1786])).
fof(f1786,plain,(
  ( ! [X4,X5] : (sP6(X4,'F_0',X5) | s(bool,'T_0') != s(bool,X4)) )),
  inference(subsumption_resolution,[],[f1781,f389])).
fof(f1781,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | p(s(bool,X5)) | sP6(X4,'F_0',X5)) )),
  inference(resolution,[],[f1133,f373])).
fof(f1133,plain,(
  ( ! [X6,X7] : (sP9(X7,'F_0',X6) | s(bool,'T_0') != s(bool,X6) | p(s(bool,X7))) )),
  inference(superposition,[],[f1124,f392])).
fof(f1124,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(subsumption_resolution,[],[f1110,f318])).
fof(f318,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))),
  inference(cnf_transformation,[],[f211])).
fof(f1110,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))),
  inference(superposition,[],[f634,f511])).
fof(f634,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))),
  inference(resolution,[],[f318,f506])).
fof(f506,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f302])).
fof(f302,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f125])).
fof(f125,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X10] : (s(bool,'F_0') = s(bool,X10) <=> ~p(s(bool,X10)))),
  file('thm.relation.RC_RTC','thm.bool.EQ_CLAUSES|split|2')).
fof(f1796,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | sP9(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP8(X3,X2,'F_0') | ~sP6(X2,'F_0',X3)) )),
  inference(resolution,[],[f1787,f377])).
fof(f1787,plain,(
  ( ! [X6,X7] : (sP7('F_0',X6,X7) | s(bool,'T_0') != s(bool,X6)) )),
  inference(subsumption_resolution,[],[f1782,f385])).
fof(f1782,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X6) | p(s(bool,X7)) | sP7('F_0',X6,X7)) )),
  inference(resolution,[],[f1133,f374])).
fof(f9428,plain,(
  ( ! [X112,X113] : (sP6('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X112,X113)) ) | $spl28),
  inference(resolution,[],[f4962,f391])).
fof(f18552,plain,(
  ( ! [X14,X15] : (sP9(X14,X15,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | p(s(bool,X14))) ) | $spl194),
  inference(cnf_transformation,[],[f18552_D])).
fof(f18552_D,plain,(
  ( ! [X14,X15] : (sP9(X14,X15,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | p(s(bool,X14))) ) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f18542,plain,(
  ( ! [X17,X16] : (sP9(X16,X17,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | p(s(bool,X16)) | s(bool,'F_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | ($spl28 | $spl166)),
  inference(resolution,[],[f9428,f17475])).
fof(f18455,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | $spl178),
  inference(cnf_transformation,[],[f18455_D])).
fof(f18455_D,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f18541,plain,(
  ( ! [X14,X15] : (sP9(X14,X15,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | p(s(bool,X14)) | s(bool,'T_0') = s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | ($spl28 | $spl158)),
  inference(resolution,[],[f9428,f17431])).
fof(f17431,plain,(
  ( ! [X6,X8,X7] : (~sP6(X6,X8,X7) | sP9(X7,X8,X6) | p(s(bool,X7)) | s(bool,'T_0') = s(bool,X6)) ) | $spl158),
  inference(subsumption_resolution,[],[f17424,f383])).
fof(f17424,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') = s(bool,X6) | sP9(X7,X8,X6) | p(s(bool,X7)) | ~sP8(X7,X6,X8) | ~sP6(X6,X8,X7)) ) | $spl158),
  inference(resolution,[],[f17415,f379])).
fof(f17415,plain,(
  ( ! [X4,X5,X3] : (sP7(X5,X3,X4) | s(bool,'T_0') = s(bool,X3)) ) | $spl158),
  inference(subsumption_resolution,[],[f17414,f385])).
fof(f17414,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') = s(bool,X3) | p(s(bool,X4)) | sP7(X5,X3,X4)) ) | $spl158),
  inference(subsumption_resolution,[],[f17390,f387])).
fof(f17390,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') = s(bool,X3) | p(s(bool,X4)) | ~p(s(bool,X5)) | sP7(X5,X3,X4)) ) | $spl158),
  inference(resolution,[],[f17192,f374])).
fof(f8867,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | $spl128),
  inference(cnf_transformation,[],[f8867_D])).
fof(f8867_D,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f18539,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | s(bool,'F_0') != s(bool,X11) | s(bool,X11) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | ($spl6 | $spl28)),
  inference(resolution,[],[f9428,f4784])).
fof(f4784,plain,(
  ( ! [X4,X5,X3] : (~sP6(X5,X3,X4) | p(s(bool,X4)) | s(bool,'F_0') != s(bool,X3) | s(bool,X3) != s(bool,X5)) ) | $spl6),
  inference(duplicate_literal_removal,[],[f4765])).
fof(f4765,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X3) | p(s(bool,X4)) | ~sP6(X5,X3,X4) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5)) ) | $spl6),
  inference(resolution,[],[f4576,f395])).
fof(f4576,plain,(
  ( ! [X17,X15,X16] : (sP9(X16,X15,X17) | s(bool,'F_0') != s(bool,X15) | p(s(bool,X16)) | ~sP6(X17,X15,X16)) ) | $spl6),
  inference(subsumption_resolution,[],[f4574,f383])).
fof(f4574,plain,(
  ( ! [X17,X15,X16] : (s(bool,'F_0') != s(bool,X15) | sP9(X16,X15,X17) | p(s(bool,X16)) | ~sP8(X16,X17,X15) | ~sP6(X17,X15,X16)) ) | $spl6),
  inference(resolution,[],[f4564,f379])).
fof(f4564,plain,(
  ( ! [X10,X11,X9] : (sP7(X10,X9,X11) | s(bool,'F_0') != s(bool,X10)) ) | $spl6),
  inference(subsumption_resolution,[],[f4563,f386])).
fof(f4563,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | s(bool,'F_0') != s(bool,X10) | sP7(X10,X9,X11)) ) | $spl6),
  inference(subsumption_resolution,[],[f4544,f385])).
fof(f4544,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | s(bool,'F_0') != s(bool,X10) | p(s(bool,X11)) | sP7(X10,X9,X11)) ) | $spl6),
  inference(resolution,[],[f4335,f374])).
fof(f4335,plain,(
  ( ! [X6,X8,X7] : (sP9(X6,X7,X8) | ~p(s(bool,X8)) | s(bool,'F_0') != s(bool,X7) | p(s(bool,X6))) ) | $spl6),
  inference(resolution,[],[f4304,f383])).
fof(f4304,plain,(
  ( ! [X4,X5,X3] : (~sP8(X5,X4,X3) | sP9(X5,X3,X4) | ~p(s(bool,X4)) | s(bool,'F_0') != s(bool,X3)) ) | $spl6),
  inference(subsumption_resolution,[],[f4303,f511])).
fof(f4303,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X3) | s(bool,'T_0') != s(bool,X4) | sP9(X5,X3,X4) | ~p(s(bool,X4)) | ~sP8(X5,X4,X3)) ) | $spl6),
  inference(subsumption_resolution,[],[f4300,f4280])).
fof(f4280,plain,(
  ( ! [X6,X8,X7] : (sP6(X6,X7,X8) | s(bool,'F_0') != s(bool,X7) | s(bool,'T_0') != s(bool,X6)) ) | $spl6),
  inference(subsumption_resolution,[],[f4260,f389])).
fof(f4260,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X6) | s(bool,'F_0') != s(bool,X7) | p(s(bool,X8)) | sP6(X6,X7,X8)) ) | $spl6),
  inference(resolution,[],[f1371,f373])).
fof(f1371,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X1,X0) | s(bool,'T_0') != s(bool,X0) | s(bool,'F_0') != s(bool,X1) | p(s(bool,X2))) ) | $spl6),
  inference(superposition,[],[f1160,f392])).
fof(f4300,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X3) | s(bool,'T_0') != s(bool,X4) | sP9(X5,X3,X4) | ~p(s(bool,X4)) | ~sP8(X5,X4,X3) | ~sP6(X4,X3,X5)) ) | $spl6),
  inference(resolution,[],[f4281,f377])).
fof(f4281,plain,(
  ( ! [X10,X11,X9] : (sP7(X10,X9,X11) | s(bool,'F_0') != s(bool,X10) | s(bool,'T_0') != s(bool,X9)) ) | $spl6),
  inference(subsumption_resolution,[],[f4261,f385])).
fof(f4261,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') != s(bool,X9) | s(bool,'F_0') != s(bool,X10) | p(s(bool,X11)) | sP7(X10,X9,X11)) ) | $spl6),
  inference(resolution,[],[f1371,f374])).
fof(f8868,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | $spl130),
  inference(cnf_transformation,[],[f8868_D])).
fof(f8868_D,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f18538,plain,(
  ( ! [X8,X9] : (p(s(bool,X8)) | p(s(bool,X9)) | s(bool,X8) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | ($spl6 | $spl28)),
  inference(resolution,[],[f9428,f4655])).
fof(f4655,plain,(
  ( ! [X4,X5,X3] : (~sP6(X4,X3,X5) | p(s(bool,X3)) | p(s(bool,X5)) | s(bool,X3) != s(bool,X4)) ) | $spl6),
  inference(duplicate_literal_removal,[],[f4639])).
fof(f4639,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~sP6(X4,X3,X5) | p(s(bool,X5)) | p(s(bool,X5)) | s(bool,X3) != s(bool,X4)) ) | $spl6),
  inference(resolution,[],[f4579,f395])).
fof(f4579,plain,(
  ( ! [X6,X8,X7] : (sP9(X7,X6,X8) | p(s(bool,X6)) | ~sP6(X8,X6,X7) | p(s(bool,X7))) ) | $spl6),
  inference(resolution,[],[f4575,f383])).
fof(f4575,plain,(
  ( ! [X14,X12,X13] : (~sP8(X13,X14,X12) | p(s(bool,X12)) | sP9(X13,X12,X14) | ~sP6(X14,X12,X13)) ) | $spl6),
  inference(subsumption_resolution,[],[f4573,f507])).
fof(f507,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f302])).
fof(f4573,plain,(
  ( ! [X14,X12,X13] : (s(bool,'F_0') != s(bool,X12) | sP9(X13,X12,X14) | p(s(bool,X12)) | ~sP8(X13,X14,X12) | ~sP6(X14,X12,X13)) ) | $spl6),
  inference(resolution,[],[f4564,f378])).
fof(f378,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | p(s(bool,X1)) | ~sP8(X0,X2,X1) | ~sP7(X1,X2,X0) | ~sP6(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f237])).
fof(f18515,plain,(
  ( ! [X12,X13] : (sP9(X12,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X13) | p(s(bool,X12))) ) | $spl192),
  inference(cnf_transformation,[],[f18515_D])).
fof(f18515_D,plain,(
  ( ! [X12,X13] : (sP9(X12,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X13) | p(s(bool,X12))) ) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f18504,plain,(
  ( ! [X12,X13] : (sP9(X12,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X13) | p(s(bool,X12)) | s(bool,'T_0') = s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | ($spl28 | $spl140)),
  inference(resolution,[],[f9427,f13495])).
fof(f13495,plain,(
  ( ! [X6,X8,X7] : (~sP6(X8,X6,X7) | sP9(X7,X6,X8) | p(s(bool,X7)) | s(bool,'T_0') = s(bool,X6)) ) | $spl140),
  inference(subsumption_resolution,[],[f13487,f383])).
fof(f13487,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') = s(bool,X6) | sP9(X7,X6,X8) | p(s(bool,X7)) | ~sP8(X7,X8,X6) | ~sP6(X8,X6,X7)) ) | $spl140),
  inference(resolution,[],[f13480,f379])).
fof(f13480,plain,(
  ( ! [X4,X5,X3] : (sP7(X3,X5,X4) | s(bool,'T_0') = s(bool,X3)) ) | $spl140),
  inference(subsumption_resolution,[],[f13479,f385])).
fof(f13479,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') = s(bool,X3) | p(s(bool,X4)) | sP7(X3,X5,X4)) ) | $spl140),
  inference(subsumption_resolution,[],[f13453,f386])).
fof(f13453,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') = s(bool,X3) | p(s(bool,X4)) | ~p(s(bool,X5)) | sP7(X3,X5,X4)) ) | $spl140),
  inference(resolution,[],[f13411,f374])).
fof(f9427,plain,(
  ( ! [X111,X110] : (sP6(X110,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X111)) ) | $spl28),
  inference(resolution,[],[f4962,f390])).
fof(f390,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f245])).
fof(f8850,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | $spl124),
  inference(cnf_transformation,[],[f8850_D])).
fof(f8850_D,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f18503,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | s(bool,'F_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | s(bool,X11) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | ($spl6 | $spl28)),
  inference(resolution,[],[f9427,f4784])).
fof(f8806,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X11) | ~p(s(bool,X11))) ) | $spl112),
  inference(cnf_transformation,[],[f8806_D])).
fof(f8806_D,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X11) | ~p(s(bool,X11))) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f18445,plain,(
  ( ! [X21,X20] : (sP9(X20,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X21) | ~p(s(bool,X21)) | s(bool,'T_0') = s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | ($spl28 | $spl140)),
  inference(resolution,[],[f9422,f13494])).
fof(f13494,plain,(
  ( ! [X2,X0,X1] : (~sP8(X1,X2,X0) | sP9(X1,X0,X2) | ~p(s(bool,X2)) | s(bool,'T_0') = s(bool,X0)) ) | $spl140),
  inference(subsumption_resolution,[],[f13485,f391])).
fof(f13485,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP9(X1,X0,X2) | ~p(s(bool,X2)) | ~sP8(X1,X2,X0) | ~sP6(X2,X0,X1)) ) | $spl140),
  inference(resolution,[],[f13480,f377])).
fof(f9422,plain,(
  ( ! [X101,X100] : (sP8(X100,X101,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | $spl28),
  inference(resolution,[],[f4962,f381])).
fof(f18444,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'F_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | s(bool,'T_0') != s(bool,X19)) ) | ($spl6 | $spl28)),
  inference(resolution,[],[f9422,f8227])).
fof(f8227,plain,(
  ( ! [X4,X5,X3] : (~sP8(X4,X3,X5) | ~p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5) | s(bool,'T_0') != s(bool,X3)) ) | $spl6),
  inference(subsumption_resolution,[],[f8226,f506])).
fof(f8226,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5) | p(s(bool,X5)) | ~sP8(X4,X3,X5)) ) | $spl6),
  inference(subsumption_resolution,[],[f8225,f4280])).
fof(f8225,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5) | p(s(bool,X5)) | ~sP8(X4,X3,X5) | ~sP6(X3,X5,X4)) ) | $spl6),
  inference(subsumption_resolution,[],[f8190,f4281])).
fof(f8190,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5) | p(s(bool,X5)) | ~sP8(X4,X3,X5) | ~sP7(X5,X3,X4) | ~sP6(X3,X5,X4)) ) | $spl6),
  inference(resolution,[],[f1373,f378])).
fof(f1373,plain,(
  ( ! [X6,X8,X7] : (~sP9(X8,X7,X6) | s(bool,'T_0') != s(bool,X6) | ~p(s(bool,X8)) | s(bool,'F_0') != s(bool,X7)) ) | $spl6),
  inference(superposition,[],[f1160,f394])).
fof(f18443,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | ~p(s(bool,X16)) | ~p(s(bool,X17))) ) | ($spl6 | $spl28)),
  inference(resolution,[],[f9422,f8224])).
fof(f8224,plain,(
  ( ! [X2,X0,X1] : (~sP8(X1,X0,X2) | s(bool,'F_0') != s(bool,X2) | ~p(s(bool,X0)) | ~p(s(bool,X1))) ) | $spl6),
  inference(subsumption_resolution,[],[f8223,f511])).
fof(f8223,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X1)) | s(bool,'F_0') != s(bool,X2) | ~p(s(bool,X0)) | ~sP8(X1,X0,X2)) ) | $spl6),
  inference(subsumption_resolution,[],[f8222,f4280])).
fof(f8222,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X1)) | s(bool,'F_0') != s(bool,X2) | ~p(s(bool,X0)) | ~sP8(X1,X0,X2) | ~sP6(X0,X2,X1)) ) | $spl6),
  inference(subsumption_resolution,[],[f8189,f4281])).
fof(f8189,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X1)) | s(bool,'F_0') != s(bool,X2) | ~p(s(bool,X0)) | ~sP8(X1,X0,X2) | ~sP7(X2,X0,X1) | ~sP6(X0,X2,X1)) ) | $spl6),
  inference(resolution,[],[f1373,f377])).
fof(f18440,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X11) | ~p(s(bool,X11)) | s(bool,'F_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | ($spl6 | $spl28)),
  inference(resolution,[],[f9422,f4304])).
fof(f18302,plain,(
  $false | $spl29),
  inference(trivial_inequality_removal,[],[f18189])).
fof(f18189,plain,(
  s(sK12,sK14) != s(sK12,sK14) | $spl29),
  inference(resolution,[],[f4961,f533])).
fof(f533,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | s(X0,X2) != s(X0,X3)) )),
  inference(cnf_transformation,[],[f309])).
fof(f309,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | s(X0,X2) = s(X0,X3) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))))) & ((s(X0,X2) != s(X0,X3) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))))) | p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))))),
  inference(flattening,[],[f308])).
fof(f308,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | (s(X0,X2) = s(X0,X3) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))))) & ((s(X0,X2) != s(X0,X3) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))))) | p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))))),
  inference(nnf_transformation,[],[f148])).
fof(f148,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) <=> (s(X0,X2) = s(X0,X3) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))))))),
  inference(flattening,[],[f147])).
fof(f147,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) <=> (s(X0,X2) = s(X0,X3) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))))))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X6] : ! [X7,X8,X9] : (p(s(bool,'const.relation.RC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X8),s(X6,X9)))) <=> (s(X6,X8) = s(X6,X9) | p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X8))),s(X6,X9))))))),
  file('thm.relation.RC_RTC','thm.relation.RC_DEF')).
fof(f4961,plain,(
  ~p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | $spl29),
  inference(cnf_transformation,[],[f4961_D])).
fof(f4961_D,plain,(
  ~p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f18035,plain,(
  $false | ($spl16 | $spl28)),
  inference(trivial_inequality_removal,[],[f18027])).
fof(f18027,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl16 | $spl28)),
  inference(superposition,[],[f2028,f9454])).
fof(f9454,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | $spl28),
  inference(resolution,[],[f4962,f511])).
fof(f18034,plain,(
  $false | $spl16),
  inference(trivial_inequality_removal,[],[f18028])).
fof(f18028,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl16),
  inference(superposition,[],[f2028,f636])).
fof(f636,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))),
  inference(resolution,[],[f318,f511])).
fof(f18033,plain,(
  $false | $spl16),
  inference(equality_resolution,[],[f2028])).
fof(f17964,plain,(
  ( ! [X109,X110,X108] : (s(bool,'T_0') != s(bool,X108) | p(s(bool,X109)) | sP9(X110,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X109) | p(s(bool,X110))) ) | ($spl6 | $spl28)),
  inference(subsumption_resolution,[],[f17906,f510])).
fof(f17906,plain,(
  ( ! [X109,X110,X108] : (s(bool,'T_0') != s(bool,X108) | p(s(bool,X109)) | ~p(s(bool,X108)) | sP9(X110,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X109) | p(s(bool,X110))) ) | ($spl6 | $spl28)),
  inference(superposition,[],[f2060,f9454])).
fof(f17957,plain,(
  ( ! [X78,X79] : (sP9(X79,'T_0',X78) | p(s(bool,X79)) | p(s(bool,X78))) ) | $spl172),
  inference(cnf_transformation,[],[f17957_D])).
fof(f17957_D,plain,(
  ( ! [X78,X79] : (sP9(X79,'T_0',X78) | p(s(bool,X79)) | p(s(bool,X78))) ) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f17961,plain,(
  ( ! [X90,X88,X91,X89] : (s(bool,X88) != s(bool,X89) | p(s(bool,X90)) | sP9(X91,'T_0',X90) | p(s(bool,X91)) | ~p(s(bool,X88))) ) | $spl6),
  inference(subsumption_resolution,[],[f17901,f2022])).
fof(f17901,plain,(
  ( ! [X90,X88,X91,X89] : (s(bool,X88) != s(bool,X89) | p(s(bool,X90)) | ~p(s(bool,X89)) | sP9(X91,'T_0',X90) | p(s(bool,X91)) | ~p(s(bool,X88))) ) | $spl6),
  inference(superposition,[],[f2060,f511])).
fof(f17960,plain,(
  ( ! [X87,X85,X86,X84] : (s(bool,X84) != s(bool,X85) | p(s(bool,X86)) | sP9(X87,'T_0',X86) | p(s(bool,X87)) | ~p(s(bool,X84))) ) | $spl6),
  inference(subsumption_resolution,[],[f17900,f2022])).
fof(f17900,plain,(
  ( ! [X87,X85,X86,X84] : (s(bool,X84) != s(bool,X85) | p(s(bool,X86)) | ~p(s(bool,X85)) | sP9(X87,'T_0',X86) | p(s(bool,X87)) | ~p(s(bool,X84))) ) | $spl6),
  inference(superposition,[],[f2060,f513])).
fof(f17959,plain,(
  ( ! [X80,X83,X81,X82] : (s(bool,X80) != s(bool,X81) | p(s(bool,X82)) | ~p(s(bool,X81)) | sP9(X83,'T_0',X82) | p(s(bool,X83))) ) | ($spl6 | $spl24)),
  inference(subsumption_resolution,[],[f17899,f4562])).
fof(f17899,plain,(
  ( ! [X80,X83,X81,X82] : (s(bool,X80) != s(bool,X81) | p(s(bool,X82)) | ~p(s(bool,X81)) | sP9(X83,'T_0',X82) | p(s(bool,X83)) | s(bool,'F_0') = s(bool,X80)) ) | $spl6),
  inference(superposition,[],[f2060,f538])).
fof(f17956,plain,(
  ( ! [X78,X76,X79,X77] : (s(bool,X76) != s(bool,X77) | p(s(bool,X78)) | ~p(s(bool,X77)) | sP9(X79,'T_0',X78) | p(s(bool,X79))) ) | ($spl6 | $spl24)),
  inference(subsumption_resolution,[],[f17898,f4562])).
fof(f17898,plain,(
  ( ! [X78,X76,X79,X77] : (s(bool,X76) != s(bool,X77) | p(s(bool,X78)) | ~p(s(bool,X77)) | sP9(X79,'T_0',X78) | p(s(bool,X79)) | s(bool,'F_0') = s(bool,X76)) ) | $spl6),
  inference(superposition,[],[f2060,f539])).
fof(f17954,plain,(
  ( ! [X62,X60,X63] : (sP9(X63,X60,X62) | s(bool,'F_0') = s(bool,X60) | p(s(bool,X63)) | p(s(bool,X62))) ) | $spl170),
  inference(cnf_transformation,[],[f17954_D])).
fof(f17954_D,plain,(
  ( ! [X62,X60,X63] : (sP9(X63,X60,X62) | s(bool,'F_0') = s(bool,X60) | p(s(bool,X63)) | p(s(bool,X62))) ) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f17955,plain,(
  ( ! [X70,X68,X71,X69] : (s(bool,'T_0') != s(bool,X69) | p(s(bool,X70)) | sP9(X71,X68,X70) | p(s(bool,X71)) | s(bool,'F_0') = s(bool,X68)) ) | $spl6),
  inference(subsumption_resolution,[],[f17896,f510])).
fof(f17896,plain,(
  ( ! [X70,X68,X71,X69] : (s(bool,'T_0') != s(bool,X69) | p(s(bool,X70)) | ~p(s(bool,X69)) | sP9(X71,X68,X70) | p(s(bool,X71)) | s(bool,'F_0') = s(bool,X68)) ) | $spl6),
  inference(superposition,[],[f2060,f539])).
fof(f17953,plain,(
  ( ! [X61,X62,X60,X63] : (s(bool,'T_0') != s(bool,X61) | p(s(bool,X62)) | sP9(X63,X60,X62) | p(s(bool,X63)) | s(bool,'F_0') = s(bool,X60)) ) | $spl6),
  inference(subsumption_resolution,[],[f17894,f510])).
fof(f17894,plain,(
  ( ! [X61,X62,X60,X63] : (s(bool,'T_0') != s(bool,X61) | p(s(bool,X62)) | ~p(s(bool,X61)) | sP9(X63,X60,X62) | p(s(bool,X63)) | s(bool,'F_0') = s(bool,X60)) ) | $spl6),
  inference(superposition,[],[f2060,f538])).
fof(f4042,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))) | $spl19),
  inference(cnf_transformation,[],[f4042_D])).
fof(f4042_D,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f17554,plain,(
  ( ! [X47,X48] : (sP9(X47,X48,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))) | p(s(bool,X47))) ) | $spl168),
  inference(cnf_transformation,[],[f17554_D])).
fof(f17554_D,plain,(
  ( ! [X47,X48] : (sP9(X47,X48,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))) | p(s(bool,X47))) ) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f17574,plain,(
  ( ! [X47,X48] : (sP9(X47,X48,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))) | p(s(bool,X47)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) ) | ($spl2 | $spl166)),
  inference(resolution,[],[f17475,f3372])).
fof(f3372,plain,(
  ( ! [X114,X113] : (sP6(happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X113,X114)) ) | $spl2),
  inference(resolution,[],[f670,f391])).
fof(f13506,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))) | $spl144),
  inference(cnf_transformation,[],[f13506_D])).
fof(f13506_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f17550,plain,(
  ( ! [X47,X48] : (sP9(X47,X48,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))) | p(s(bool,X47)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) ) | ($spl2 | $spl158)),
  inference(resolution,[],[f17431,f3372])).
fof(f17202,plain,(
  ( ! [X47,X45,X46] : (sP9(X47,X46,X45) | p(s(bool,X47)) | s(bool,'F_0') != s(bool,X45) | ~p(s(bool,X46))) ) | $spl166),
  inference(cnf_transformation,[],[f17202_D])).
fof(f17202_D,plain,(
  ( ! [X47,X45,X46] : (sP9(X47,X46,X45) | p(s(bool,X47)) | s(bool,'F_0') != s(bool,X45) | ~p(s(bool,X46))) ) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f17206,plain,(
  ( ! [X74,X72,X75,X73] : (s(bool,'F_0') != s(bool,X73) | p(s(bool,X72)) | ~p(s(bool,X74)) | sP9(X75,X74,X73) | p(s(bool,X75))) ) | $spl6),
  inference(subsumption_resolution,[],[f17166,f510])).
fof(f17166,plain,(
  ( ! [X74,X72,X75,X73] : (s(bool,'F_0') != s(bool,X73) | p(s(bool,X72)) | ~p(s(bool,X74)) | sP9(X75,X74,X73) | p(s(bool,X75)) | s(bool,'T_0') = s(bool,X72)) ) | $spl6),
  inference(superposition,[],[f2033,f539])).
fof(f2033,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X5) != s(bool,X6) | p(s(bool,X6)) | ~p(s(bool,X4)) | sP9(X7,X4,X5) | p(s(bool,X7))) ) | $spl6),
  inference(superposition,[],[f2022,f392])).
fof(f17204,plain,(
  ( ! [X66,X64,X67,X65] : (s(bool,'F_0') != s(bool,X65) | p(s(bool,X64)) | ~p(s(bool,X66)) | sP9(X67,X66,X65) | p(s(bool,X67))) ) | $spl6),
  inference(subsumption_resolution,[],[f17164,f510])).
fof(f17164,plain,(
  ( ! [X66,X64,X67,X65] : (s(bool,'F_0') != s(bool,X65) | p(s(bool,X64)) | ~p(s(bool,X66)) | sP9(X67,X66,X65) | p(s(bool,X67)) | s(bool,'T_0') = s(bool,X64)) ) | $spl6),
  inference(superposition,[],[f2033,f538])).
fof(f17182,plain,(
  ( ! [X50,X48,X51,X49] : (s(bool,'F_0') != s(bool,X49) | p(s(bool,X48)) | ~p(s(bool,X50)) | sP9(X51,X50,X49) | p(s(bool,X51))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f17160])).
fof(f17160,plain,(
  ( ! [X50,X48,X51,X49] : (s(bool,'F_0') != s(bool,X49) | p(s(bool,X48)) | ~p(s(bool,X50)) | sP9(X51,X50,X49) | p(s(bool,X51)) | p(s(bool,X48))) ) | $spl6),
  inference(superposition,[],[f2033,f509])).
fof(f17183,plain,(
  ( ! [X47,X45,X46,X44] : (s(bool,'F_0') != s(bool,X45) | p(s(bool,X44)) | ~p(s(bool,X46)) | sP9(X47,X46,X45) | p(s(bool,X47))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f17159])).
fof(f17159,plain,(
  ( ! [X47,X45,X46,X44] : (s(bool,'F_0') != s(bool,X45) | p(s(bool,X44)) | ~p(s(bool,X46)) | sP9(X47,X46,X45) | p(s(bool,X47)) | p(s(bool,X44))) ) | $spl6),
  inference(superposition,[],[f2033,f507])).
fof(f17201,plain,(
  ( ! [X116,X115] : (~p(s(bool,X115)) | p(s(bool,X116)) | sP9(X116,X115,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) | $spl164),
  inference(cnf_transformation,[],[f17201_D])).
fof(f17201_D,plain,(
  ( ! [X116,X115] : (~p(s(bool,X115)) | p(s(bool,X116)) | sP9(X116,X115,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f17200,plain,(
  ( ! [X116,X114,X115] : (p(s(bool,X114)) | ~p(s(bool,X115)) | sP9(X116,X115,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | p(s(bool,X116))) ) | $spl6),
  inference(subsumption_resolution,[],[f17147,f507])).
fof(f17147,plain,(
  ( ! [X116,X114,X115] : (s(bool,'F_0') != s(bool,X114) | p(s(bool,X114)) | ~p(s(bool,X115)) | sP9(X116,X115,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | p(s(bool,X116))) ) | $spl6),
  inference(superposition,[],[f2033,f758])).
fof(f758,plain,(
  s(bool,'F_0') = s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))),
  inference(resolution,[],[f319,f507])).
fof(f17195,plain,(
  ( ! [X94,X95] : (sP9(X95,X94,'F_0') | p(s(bool,X95)) | ~p(s(bool,X94))) ) | $spl160),
  inference(cnf_transformation,[],[f17195_D])).
fof(f17195_D,plain,(
  ( ! [X94,X95] : (sP9(X95,X94,'F_0') | p(s(bool,X95)) | ~p(s(bool,X94))) ) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f5096,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,X0) != s(bool,X1)) ) | $spl54),
  inference(cnf_transformation,[],[f5096_D])).
fof(f5096_D,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,X0) != s(bool,X1)) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f17199,plain,(
  ( ! [X107,X105,X106,X104] : (s(bool,X104) != s(bool,X105) | p(s(bool,X105)) | ~p(s(bool,X106)) | sP9(X107,X106,'F_0') | p(s(bool,X107))) ) | $spl6),
  inference(subsumption_resolution,[],[f17144,f2022])).
fof(f17144,plain,(
  ( ! [X107,X105,X106,X104] : (s(bool,X104) != s(bool,X105) | p(s(bool,X105)) | ~p(s(bool,X106)) | sP9(X107,X106,'F_0') | p(s(bool,X107)) | p(s(bool,X104))) ) | $spl6),
  inference(superposition,[],[f2033,f507])).
fof(f17198,plain,(
  ( ! [X103,X101,X102,X100] : (s(bool,X100) != s(bool,X101) | p(s(bool,X101)) | ~p(s(bool,X102)) | sP9(X103,X102,'F_0') | p(s(bool,X103))) ) | $spl6),
  inference(subsumption_resolution,[],[f17143,f2022])).
fof(f17143,plain,(
  ( ! [X103,X101,X102,X100] : (s(bool,X100) != s(bool,X101) | p(s(bool,X101)) | ~p(s(bool,X102)) | sP9(X103,X102,'F_0') | p(s(bool,X103)) | p(s(bool,X100))) ) | $spl6),
  inference(superposition,[],[f2033,f509])).
fof(f17196,plain,(
  ( ! [X92,X93] : (s(bool,X92) != s(bool,X93) | s(bool,'T_0') = s(bool,X92)) ) | $spl162),
  inference(cnf_transformation,[],[f17196_D])).
fof(f17196_D,plain,(
  ( ! [X92,X93] : (s(bool,X92) != s(bool,X93) | s(bool,'T_0') = s(bool,X92)) ) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f17197,plain,(
  ( ! [X99,X97,X98,X96] : (s(bool,X96) != s(bool,X97) | ~p(s(bool,X98)) | sP9(X99,X98,'F_0') | p(s(bool,X99)) | s(bool,'T_0') = s(bool,X96)) ) | $spl6),
  inference(subsumption_resolution,[],[f17142,f2114])).
fof(f2114,plain,(
  ( ! [X52,X53] : (s(bool,X52) != s(bool,X53) | ~p(s(bool,X53)) | s(bool,'T_0') = s(bool,X52)) ) | $spl6),
  inference(subsumption_resolution,[],[f2080,f462])).
fof(f462,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f273])).
fof(f273,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f272])).
fof(f272,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f178])).
fof(f178,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f110])).
fof(f110,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f109])).
fof(f109,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X10] : ((p(s(bool,X10)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X10)))),
  file('thm.relation.RC_RTC','thm.bool.IMP_CLAUSES|split|4')).
fof(f2080,plain,(
  ( ! [X52,X53] : (s(bool,X52) != s(bool,X53) | p(s(bool,'F_0')) | ~p(s(bool,X53)) | s(bool,'T_0') = s(bool,X52)) ) | $spl6),
  inference(superposition,[],[f2022,f539])).
fof(f17142,plain,(
  ( ! [X99,X97,X98,X96] : (s(bool,X96) != s(bool,X97) | p(s(bool,X97)) | ~p(s(bool,X98)) | sP9(X99,X98,'F_0') | p(s(bool,X99)) | s(bool,'T_0') = s(bool,X96)) ) | $spl6),
  inference(superposition,[],[f2033,f538])).
fof(f17194,plain,(
  ( ! [X94,X92,X95,X93] : (s(bool,X92) != s(bool,X93) | ~p(s(bool,X94)) | sP9(X95,X94,'F_0') | p(s(bool,X95)) | s(bool,'T_0') = s(bool,X92)) ) | $spl6),
  inference(subsumption_resolution,[],[f17141,f2114])).
fof(f17141,plain,(
  ( ! [X94,X92,X95,X93] : (s(bool,X92) != s(bool,X93) | p(s(bool,X93)) | ~p(s(bool,X94)) | sP9(X95,X94,'F_0') | p(s(bool,X95)) | s(bool,'T_0') = s(bool,X92)) ) | $spl6),
  inference(superposition,[],[f2033,f539])).
fof(f17192,plain,(
  ( ! [X66,X64,X67] : (sP9(X67,X66,X64) | s(bool,'T_0') = s(bool,X64) | p(s(bool,X67)) | ~p(s(bool,X66))) ) | $spl158),
  inference(cnf_transformation,[],[f17192_D])).
fof(f17192_D,plain,(
  ( ! [X66,X64,X67] : (sP9(X67,X66,X64) | s(bool,'T_0') = s(bool,X64) | p(s(bool,X67)) | ~p(s(bool,X66))) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f17193,plain,(
  ( ! [X74,X72,X75,X73] : (p(s(bool,X73)) | ~p(s(bool,X74)) | sP9(X75,X74,X72) | p(s(bool,X75)) | s(bool,'T_0') = s(bool,X72)) ) | $spl6),
  inference(subsumption_resolution,[],[f17136,f507])).
fof(f17136,plain,(
  ( ! [X74,X72,X75,X73] : (s(bool,'F_0') != s(bool,X73) | p(s(bool,X73)) | ~p(s(bool,X74)) | sP9(X75,X74,X72) | p(s(bool,X75)) | s(bool,'T_0') = s(bool,X72)) ) | $spl6),
  inference(superposition,[],[f2033,f539])).
fof(f17191,plain,(
  ( ! [X66,X64,X67,X65] : (p(s(bool,X65)) | ~p(s(bool,X66)) | sP9(X67,X66,X64) | p(s(bool,X67)) | s(bool,'T_0') = s(bool,X64)) ) | $spl6),
  inference(subsumption_resolution,[],[f17134,f507])).
fof(f17134,plain,(
  ( ! [X66,X64,X67,X65] : (s(bool,'F_0') != s(bool,X65) | p(s(bool,X65)) | ~p(s(bool,X66)) | sP9(X67,X66,X64) | p(s(bool,X67)) | s(bool,'T_0') = s(bool,X64)) ) | $spl6),
  inference(superposition,[],[f2033,f538])).
fof(f17187,plain,(
  ( ! [X28,X24,X27] : (~p(s(bool,X27)) | p(s(bool,X24)) | p(s(bool,X28)) | sP9(X28,X27,X24)) ) | $spl156),
  inference(cnf_transformation,[],[f17187_D])).
fof(f17187_D,plain,(
  ( ! [X28,X24,X27] : (~p(s(bool,X27)) | p(s(bool,X24)) | p(s(bool,X28)) | sP9(X28,X27,X24)) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f17190,plain,(
  ( ! [X50,X48,X51,X49] : (p(s(bool,X49)) | ~p(s(bool,X50)) | sP9(X51,X50,X48) | p(s(bool,X51)) | p(s(bool,X48))) ) | $spl6),
  inference(subsumption_resolution,[],[f17130,f507])).
fof(f17130,plain,(
  ( ! [X50,X48,X51,X49] : (s(bool,'F_0') != s(bool,X49) | p(s(bool,X49)) | ~p(s(bool,X50)) | sP9(X51,X50,X48) | p(s(bool,X51)) | p(s(bool,X48))) ) | $spl6),
  inference(superposition,[],[f2033,f509])).
fof(f17189,plain,(
  ( ! [X47,X45,X46,X44] : (p(s(bool,X45)) | ~p(s(bool,X46)) | sP9(X47,X46,X44) | p(s(bool,X47)) | p(s(bool,X44))) ) | $spl6),
  inference(subsumption_resolution,[],[f17129,f507])).
fof(f17129,plain,(
  ( ! [X47,X45,X46,X44] : (s(bool,'F_0') != s(bool,X45) | p(s(bool,X45)) | ~p(s(bool,X46)) | sP9(X47,X46,X44) | p(s(bool,X47)) | p(s(bool,X44))) ) | $spl6),
  inference(superposition,[],[f2033,f507])).
fof(f17188,plain,(
  ( ! [X30,X33,X31,X29,X32] : (s(bool,X30) != s(bool,X31) | p(s(bool,X31)) | ~p(s(bool,X32)) | sP9(X33,X32,X29) | p(s(bool,X33)) | p(s(bool,X29))) ) | $spl6),
  inference(subsumption_resolution,[],[f17126,f2022])).
fof(f17126,plain,(
  ( ! [X30,X33,X31,X29,X32] : (s(bool,X30) != s(bool,X31) | p(s(bool,X31)) | ~p(s(bool,X32)) | sP9(X33,X32,X29) | p(s(bool,X33)) | p(s(bool,X29)) | p(s(bool,X30))) ) | $spl6),
  inference(superposition,[],[f2033,f449])).
fof(f449,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f177])).
fof(f177,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f176])).
fof(f176,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f101])).
fof(f101,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X11,X12] : ((p(s(bool,X11)) => p(s(bool,X12))) => ((p(s(bool,X12)) => p(s(bool,X11))) => s(bool,X11) = s(bool,X12)))),
  file('thm.relation.RC_RTC','thm.bool.IMP_ANTISYM_AX')).
fof(f17186,plain,(
  ( ! [X28,X26,X24,X27,X25] : (s(bool,X25) != s(bool,X26) | p(s(bool,X26)) | ~p(s(bool,X27)) | sP9(X28,X27,X24) | p(s(bool,X28)) | p(s(bool,X24))) ) | $spl6),
  inference(subsumption_resolution,[],[f17125,f2022])).
fof(f17125,plain,(
  ( ! [X28,X26,X24,X27,X25] : (s(bool,X25) != s(bool,X26) | p(s(bool,X26)) | ~p(s(bool,X27)) | sP9(X28,X27,X24) | p(s(bool,X28)) | p(s(bool,X25)) | p(s(bool,X24))) ) | $spl6),
  inference(superposition,[],[f2033,f449])).
fof(f17000,plain,(
  ( ! [X14] : (p(s(bool,X14)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) ) | $spl2),
  inference(resolution,[],[f3372,f1855])).
fof(f8572,plain,(
  ( ! [X42] : (s(bool,'F_0') != s(bool,X42) | s(bool,X42) != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) ) | $spl96),
  inference(cnf_transformation,[],[f8572_D])).
fof(f8572_D,plain,(
  ( ! [X42] : (s(bool,'F_0') != s(bool,X42) | s(bool,X42) != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f16998,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | s(bool,'F_0') != s(bool,X11) | s(bool,X11) != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) ) | ($spl2 | $spl6)),
  inference(resolution,[],[f3372,f4784])).
fof(f16997,plain,(
  ( ! [X8,X9] : (p(s(bool,X8)) | p(s(bool,X9)) | s(bool,X8) != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) ) | ($spl2 | $spl6)),
  inference(resolution,[],[f3372,f4655])).
fof(f13531,plain,(
  ( ! [X37,X36] : (sP9(X36,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X37) | p(s(bool,X36))) ) | $spl146),
  inference(cnf_transformation,[],[f13531_D])).
fof(f13531_D,plain,(
  ( ! [X37,X36] : (sP9(X36,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X37) | p(s(bool,X36))) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f16982,plain,(
  ( ! [X12,X13] : (sP9(X12,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X13) | p(s(bool,X12)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) ) | ($spl2 | $spl140)),
  inference(resolution,[],[f3371,f13495])).
fof(f3371,plain,(
  ( ! [X111,X112] : (sP6(X111,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X112)) ) | $spl2),
  inference(resolution,[],[f670,f390])).
fof(f8570,plain,(
  ( ! [X33] : (s(bool,X33) != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) ) | $spl94),
  inference(cnf_transformation,[],[f8570_D])).
fof(f8570_D,plain,(
  ( ! [X33] : (s(bool,X33) != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f16981,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))) | s(bool,X11) != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) ) | ($spl2 | $spl6)),
  inference(resolution,[],[f3371,f4784])).
fof(f4342,plain,(
  ( ! [X14,X15] : (sP9(X14,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X15) | ~p(s(bool,X15))) ) | $spl22),
  inference(cnf_transformation,[],[f4342_D])).
fof(f4342_D,plain,(
  ( ! [X14,X15] : (sP9(X14,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X15) | ~p(s(bool,X15))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f16910,plain,(
  ( ! [X21,X20] : (sP9(X20,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X21) | ~p(s(bool,X21)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) ) | ($spl2 | $spl140)),
  inference(resolution,[],[f3366,f13494])).
fof(f16909,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))) | s(bool,'T_0') != s(bool,X19)) ) | ($spl2 | $spl6)),
  inference(resolution,[],[f3366,f8227])).
fof(f16908,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))) | ~p(s(bool,X16)) | ~p(s(bool,X17))) ) | ($spl2 | $spl6)),
  inference(resolution,[],[f3366,f8224])).
fof(f16905,plain,(
  ( ! [X10,X11] : (sP9(X10,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X11) | ~p(s(bool,X11)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) ) | ($spl2 | $spl6)),
  inference(resolution,[],[f3366,f4304])).
fof(f16668,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | p(s(bool,X11))) ) | $spl46),
  inference(subsumption_resolution,[],[f16667,f511])).
fof(f16667,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') != s(bool,X10) | ~p(s(bool,X10)) | p(s(bool,X11))) ) | $spl46),
  inference(forward_demodulation,[],[f16650,f636])).
fof(f16650,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) | $spl46),
  inference(resolution,[],[f5064,f395])).
fof(f13416,plain,(
  ( ! [X116,X115] : (~p(s(bool,X115)) | p(s(bool,X116)) | sP9(X116,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X115)) ) | $spl142),
  inference(cnf_transformation,[],[f13416_D])).
fof(f13416_D,plain,(
  ( ! [X116,X115] : (~p(s(bool,X115)) | p(s(bool,X116)) | sP9(X116,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X115)) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f15182,plain,(
  ( ! [X306,X305,X307] : (p(s(bool,X305)) | ~p(s(bool,X306)) | sP9(X307,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X306) | p(s(bool,X307))) ) | $spl6),
  inference(subsumption_resolution,[],[f15012,f507])).
fof(f15012,plain,(
  ( ! [X306,X305,X307] : (s(bool,'F_0') != s(bool,X305) | p(s(bool,X305)) | ~p(s(bool,X306)) | sP9(X307,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X306) | p(s(bool,X307))) ) | $spl6),
  inference(superposition,[],[f2032,f758])).
fof(f2032,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X1) != s(bool,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | sP9(X3,X1,X0) | p(s(bool,X3))) ) | $spl6),
  inference(superposition,[],[f2022,f392])).
fof(f5146,plain,(
  ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | $spl67),
  inference(cnf_transformation,[],[f5146_D])).
fof(f5146_D,plain,(
  ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) <=> ~$spl67),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl67])])).
fof(f14639,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(duplicate_literal_removal,[],[f14629])).
fof(f14629,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | ~p(s(bool,X4)) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f738,f417])).
fof(f417,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | ~sP10(X0,X1)) )),
  inference(cnf_transformation,[],[f257])).
fof(f257,plain,(
  ! [X0,X1] : ((~sP10(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP10(X0,X1)))),
  inference(rectify,[],[f256])).
fof(f256,plain,(
  ! [X1,X0] : ((~sP10(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP10(X1,X0)))),
  inference(flattening,[],[f255])).
fof(f255,plain,(
  ! [X1,X0] : ((~sP10(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP10(X1,X0)))),
  inference(nnf_transformation,[],[f207])).
fof(f207,plain,(
  ! [X1,X0] : (sP10(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f738,plain,(
  ( ! [X121] : (sP10(X121,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | ~p(s(bool,X121))) )),
  inference(resolution,[],[f319,f418])).
fof(f418,plain,(
  ( ! [X0,X1] : (sP10(X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f257])).
fof(f14640,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(duplicate_literal_removal,[],[f14626])).
fof(f14626,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~p(s(bool,X1))) )),
  inference(resolution,[],[f738,f424])).
fof(f424,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | ~sP10(X1,X0)) )),
  inference(cnf_transformation,[],[f258])).
fof(f258,plain,(
  ! [X0,X1] : ((((p(s(bool,X1)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)))) | sP10(X1,X0)) & (~sP10(X1,X0) | ((~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X0)) | p(s(bool,X1))))))),
  inference(nnf_transformation,[],[f208])).
fof(f208,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> sP10(X1,X0))),
  inference(definition_folding,[],[f88,f207])).
fof(f88,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  inference(flattening,[],[f87])).
fof(f87,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  inference(rectify,[],[f59])).
fof(f59,axiom,(
  ! [X20,X21] : ((p(s(bool,X21)) <=> ~p(s(bool,X20))) <=> ((p(s(bool,X21)) | p(s(bool,X20))) & (~p(s(bool,X20)) | ~p(s(bool,X21)))))),
  file('thm.relation.RC_RTC','thm.sat.sat_29_prove_ltb')).
fof(f14621,plain,(
  ( ! [X2] : (p(s(bool,X2)) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(duplicate_literal_removal,[],[f14611])).
fof(f14611,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,X2)) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f737,f418])).
fof(f737,plain,(
  ( ! [X120] : (~sP10('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X120) | p(s(bool,X120))) )),
  inference(resolution,[],[f319,f416])).
fof(f416,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | ~sP10(X0,X1)) )),
  inference(cnf_transformation,[],[f257])).
fof(f14622,plain,(
  ( ! [X1] : (p(s(bool,X1)) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(duplicate_literal_removal,[],[f14610])).
fof(f14610,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,X1)) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f737,f423])).
fof(f423,plain,(
  ( ! [X0,X1] : (sP10(X1,X0) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f258])).
fof(f14605,plain,(
  ( ! [X3] : (p(s(bool,X3)) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(duplicate_literal_removal,[],[f14596])).
fof(f14596,plain,(
  ( ! [X3] : (p(s(bool,X3)) | p(s(bool,X3)) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f736,f421])).
fof(f421,plain,(
  ( ! [X0,X1] : (sP10(X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f257])).
fof(f736,plain,(
  ( ! [X119] : (~sP10(X119,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | p(s(bool,X119))) )),
  inference(resolution,[],[f319,f416])).
fof(f14606,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(duplicate_literal_removal,[],[f14593])).
fof(f14593,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X0))) )),
  inference(resolution,[],[f736,f422])).
fof(f422,plain,(
  ( ! [X0,X1] : (sP10(X1,X0) | ~p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f258])).
fof(f4976,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | $spl30),
  inference(cnf_transformation,[],[f4976_D])).
fof(f4976_D,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f14020,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(duplicate_literal_removal,[],[f14012])).
fof(f14012,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~p(s(bool,X3))) )),
  inference(resolution,[],[f613,f421])).
fof(f613,plain,(
  ( ! [X118] : (~sP10('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X118) | ~p(s(bool,X118))) )),
  inference(resolution,[],[f318,f417])).
fof(f14021,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(duplicate_literal_removal,[],[f14009])).
fof(f14009,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f613,f422])).
fof(f5286,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X11) | ~p(s(bool,X11))) ) | $spl92),
  inference(cnf_transformation,[],[f5286_D])).
fof(f5286_D,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X11) | ~p(s(bool,X11))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f13999,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X11) | ~p(s(bool,X11)) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) ) | $spl6),
  inference(resolution,[],[f730,f4396])).
fof(f4396,plain,(
  ( ! [X2,X0,X1] : (~sP7(X2,X0,X1) | sP9(X1,X2,X0) | ~p(s(bool,X0)) | ~p(s(bool,X2))) ) | $spl6),
  inference(resolution,[],[f4377,f381])).
fof(f4377,plain,(
  ( ! [X4,X5,X3] : (~sP8(X4,X3,X5) | ~p(s(bool,X3)) | sP9(X4,X5,X3) | ~sP7(X5,X3,X4)) ) | $spl6),
  inference(subsumption_resolution,[],[f4371,f511])).
fof(f4371,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,X3) | sP9(X4,X5,X3) | ~p(s(bool,X3)) | ~sP8(X4,X3,X5) | ~sP7(X5,X3,X4)) ) | $spl6),
  inference(resolution,[],[f4364,f377])).
fof(f4364,plain,(
  ( ! [X6,X8,X7] : (sP6(X7,X6,X8) | s(bool,'T_0') != s(bool,X7)) ) | $spl6),
  inference(subsumption_resolution,[],[f4363,f389])).
fof(f4363,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X7) | p(s(bool,X8)) | sP6(X7,X6,X8)) ) | $spl6),
  inference(subsumption_resolution,[],[f4346,f390])).
fof(f4346,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | s(bool,'T_0') != s(bool,X7) | p(s(bool,X8)) | sP6(X7,X6,X8)) ) | $spl6),
  inference(resolution,[],[f4326,f373])).
fof(f4326,plain,(
  ( ! [X6,X8,X7] : (sP9(X6,X7,X8) | p(s(bool,X7)) | s(bool,'T_0') != s(bool,X8) | p(s(bool,X6))) ) | $spl6),
  inference(resolution,[],[f4298,f383])).
fof(f4298,plain,(
  ( ! [X6,X8,X7] : (~sP8(X8,X7,X6) | sP9(X8,X6,X7) | p(s(bool,X6)) | s(bool,'T_0') != s(bool,X7)) ) | $spl6),
  inference(subsumption_resolution,[],[f4297,f387])).
fof(f4297,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X7) | sP9(X8,X6,X7) | p(s(bool,X6)) | ~sP8(X8,X7,X6) | ~sP7(X6,X7,X8)) ) | $spl6),
  inference(subsumption_resolution,[],[f4288,f507])).
fof(f4288,plain,(
  ( ! [X6,X8,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,'T_0') != s(bool,X7) | sP9(X8,X6,X7) | p(s(bool,X6)) | ~sP8(X8,X7,X6) | ~sP7(X6,X7,X8)) ) | $spl6),
  inference(resolution,[],[f4280,f378])).
fof(f730,plain,(
  ( ! [X107,X108] : (sP7('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X107,X108)) )),
  inference(resolution,[],[f319,f387])).
fof(f13997,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~p(s(bool,X6)) | p(s(bool,X7))) )),
  inference(resolution,[],[f730,f384])).
fof(f384,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f243])).
fof(f5269,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | p(s(bool,X8))) ) | $spl88),
  inference(cnf_transformation,[],[f5269_D])).
fof(f5269_D,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | p(s(bool,X8))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f13984,plain,(
  ( ! [X12,X13] : (sP9(X12,X13,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X12))) ) | $spl6),
  inference(resolution,[],[f729,f4398])).
fof(f4398,plain,(
  ( ! [X6,X8,X7] : (~sP7(X8,X6,X7) | sP9(X7,X8,X6) | ~p(s(bool,X6)) | p(s(bool,X7))) ) | $spl6),
  inference(resolution,[],[f4377,f383])).
fof(f729,plain,(
  ( ! [X105,X106] : (sP7(X105,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X106)) )),
  inference(resolution,[],[f319,f386])).
fof(f5270,plain,(
  ( ! [X10,X11] : (sP9(X10,X11,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | ~p(s(bool,X11))) ) | $spl90),
  inference(cnf_transformation,[],[f5270_D])).
fof(f5270_D,plain,(
  ( ! [X10,X11] : (sP9(X10,X11,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | ~p(s(bool,X11))) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f13983,plain,(
  ( ! [X10,X11] : (sP9(X10,X11,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~p(s(bool,X11))) ) | $spl6),
  inference(resolution,[],[f729,f4396])).
fof(f1933,plain,(
  s(bool,'T_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | $spl11),
  inference(cnf_transformation,[],[f1933_D])).
fof(f1933_D,plain,(
  s(bool,'T_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f13982,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | p(s(bool,X8)) | s(bool,'T_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) | $spl6),
  inference(resolution,[],[f729,f4378])).
fof(f4378,plain,(
  ( ! [X10,X11,X9] : (~sP7(X11,X9,X10) | sP9(X10,X11,X9) | p(s(bool,X10)) | s(bool,'T_0') != s(bool,X9)) ) | $spl6),
  inference(subsumption_resolution,[],[f4373,f383])).
fof(f4373,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') != s(bool,X9) | sP9(X10,X11,X9) | p(s(bool,X10)) | ~sP8(X10,X9,X11) | ~sP7(X11,X9,X10)) ) | $spl6),
  inference(resolution,[],[f4364,f379])).
fof(f13981,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X7))) )),
  inference(resolution,[],[f729,f384])).
fof(f5267,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) | $spl86),
  inference(cnf_transformation,[],[f5267_D])).
fof(f5267_D,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f13992,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(subsumption_resolution,[],[f13991,f726])).
fof(f726,plain,(
  ( ! [X99,X100] : (sP8(X99,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X100)) )),
  inference(resolution,[],[f319,f382])).
fof(f382,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f240])).
fof(f13991,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~sP8(X0,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X1)) )),
  inference(subsumption_resolution,[],[f13978,f391])).
fof(f13978,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~sP8(X0,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X1) | ~sP6('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X1,X0)) )),
  inference(resolution,[],[f729,f377])).
fof(f9339,plain,(
  ( ! [X28,X27] : (s(bool,X27) != s(bool,X28) | s(bool,'T_0') != s(bool,X27)) ) | $spl136),
  inference(cnf_transformation,[],[f9339_D])).
fof(f9339_D,plain,(
  ( ! [X28,X27] : (s(bool,X27) != s(bool,X28) | s(bool,'T_0') != s(bool,X27)) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f13969,plain,(
  ( ! [X30,X29] : (~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | s(bool,X29) != s(bool,X30) | s(bool,'T_0') != s(bool,X29)) ) | $spl6),
  inference(resolution,[],[f727,f4249])).
fof(f4249,plain,(
  ( ! [X12,X10,X11] : (~sP8(X11,X12,'F_0') | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,'T_0') != s(bool,X10)) ) | $spl6),
  inference(subsumption_resolution,[],[f4237,f1351])).
fof(f4237,plain,(
  ( ! [X12,X10,X11] : (s(bool,'T_0') != s(bool,X10) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | ~p(s(bool,X12)) | ~sP8(X11,X12,'F_0')) ) | $spl6),
  inference(resolution,[],[f1356,f1800])).
fof(f1356,plain,(
  ( ! [X19,X17,X18] : (~sP9(X19,'F_0',X17) | s(bool,'T_0') != s(bool,X18) | ~p(s(bool,X19)) | s(bool,X17) != s(bool,X18)) ) | $spl6),
  inference(superposition,[],[f1160,f394])).
fof(f727,plain,(
  ( ! [X101,X102] : (sP8('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X101,X102)) )),
  inference(resolution,[],[f319,f383])).
fof(f13967,plain,(
  ( ! [X26] : (s(bool,'T_0') != s(bool,X26) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f727,f1924])).
fof(f1924,plain,(
  ( ! [X14,X15] : (~sP8(X14,X15,'F_0') | s(bool,'T_0') != s(bool,X15) | ~p(s(bool,X14))) )),
  inference(subsumption_resolution,[],[f1923,f389])).
fof(f1923,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | ~sP8(X14,X15,'F_0') | ~sP6(X15,'F_0',X14)) )),
  inference(subsumption_resolution,[],[f1922,f462])).
fof(f1922,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | p(s(bool,'F_0')) | ~sP8(X14,X15,'F_0') | ~sP6(X15,'F_0',X14)) )),
  inference(subsumption_resolution,[],[f1915,f387])).
fof(f1915,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | p(s(bool,'F_0')) | ~sP8(X14,X15,'F_0') | ~sP7('F_0',X15,X14) | ~sP6(X15,'F_0',X14)) )),
  inference(resolution,[],[f1135,f378])).
fof(f1135,plain,(
  ( ! [X10,X11] : (~sP9(X11,'F_0',X10) | ~p(s(bool,X11)) | s(bool,'T_0') != s(bool,X10)) )),
  inference(superposition,[],[f1124,f394])).
fof(f13964,plain,(
  ( ! [X23] : (~p(s(bool,X23)) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f727,f1809])).
fof(f1809,plain,(
  ( ! [X0,X1] : (~sP8(X1,X0,'F_0') | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(subsumption_resolution,[],[f1803,f506])).
fof(f1803,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~sP8(X1,X0,'F_0') | s(bool,'F_0') = s(bool,X0) | ~p(s(bool,X1))) )),
  inference(resolution,[],[f1800,f394])).
fof(f13961,plain,(
  ( ! [X19,X18] : (~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | s(bool,'F_0') != s(bool,X18) | s(bool,'T_0') != s(bool,X19)) ) | $spl6),
  inference(resolution,[],[f727,f8227])).
fof(f13960,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,X16) | ~p(s(bool,X17)) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) ) | $spl6),
  inference(resolution,[],[f727,f8224])).
fof(f13955,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~p(s(bool,X6)) | p(s(bool,X7))) )),
  inference(resolution,[],[f727,f380])).
fof(f380,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f240])).
fof(f5145,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | $spl66),
  inference(cnf_transformation,[],[f5145_D])).
fof(f5145_D,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f5251,plain,(
  ( ! [X4,X5] : (sP9('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X4,X5) | ~sP6(X5,X4,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | ~sP7(X4,X5,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) | $spl84),
  inference(cnf_transformation,[],[f5251_D])).
fof(f5251_D,plain,(
  ( ! [X4,X5] : (sP9('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X4,X5) | ~sP6(X5,X4,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | ~sP7(X4,X5,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f13954,plain,(
  ( ! [X4,X5] : (sP9('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X4,X5) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~sP7(X4,X5,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | ~sP6(X5,X4,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(resolution,[],[f727,f379])).
fof(f9117,plain,(
  ( ! [X25] : (s(bool,'T_0') != s(bool,X25) | s(bool,X25) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) | $spl134),
  inference(cnf_transformation,[],[f9117_D])).
fof(f9117_D,plain,(
  ( ! [X25] : (s(bool,'T_0') != s(bool,X25) | s(bool,X25) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f13942,plain,(
  ( ! [X30,X29] : (~p(s(bool,X29)) | s(bool,X30) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | s(bool,'T_0') != s(bool,X30)) ) | $spl6),
  inference(resolution,[],[f726,f4249])).
fof(f1940,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) | $spl12),
  inference(cnf_transformation,[],[f1940_D])).
fof(f1940_D,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f13941,plain,(
  ( ! [X28,X27] : (s(bool,'T_0') != s(bool,X27) | sP9(X28,'F_0','const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | s(bool,X27) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) | $spl6),
  inference(resolution,[],[f726,f4191])).
fof(f4191,plain,(
  ( ! [X12,X13,X11] : (~sP8(X13,X12,'F_0') | s(bool,'T_0') != s(bool,X11) | sP9(X13,'F_0',X12) | s(bool,X11) != s(bool,X12)) ) | $spl6),
  inference(subsumption_resolution,[],[f4190,f1817])).
fof(f4190,plain,(
  ( ! [X12,X13,X11] : (s(bool,X11) != s(bool,X12) | s(bool,'T_0') != s(bool,X11) | sP9(X13,'F_0',X12) | ~sP8(X13,X12,'F_0') | ~sP7('F_0',X12,X13)) ) | $spl6),
  inference(subsumption_resolution,[],[f4186,f1351])).
fof(f4186,plain,(
  ( ! [X12,X13,X11] : (s(bool,X11) != s(bool,X12) | s(bool,'T_0') != s(bool,X11) | sP9(X13,'F_0',X12) | ~p(s(bool,X12)) | ~sP8(X13,X12,'F_0') | ~sP7('F_0',X12,X13)) ) | $spl6),
  inference(resolution,[],[f4177,f377])).
fof(f4177,plain,(
  ( ! [X17,X15,X16] : (sP6(X16,'F_0',X17) | s(bool,X15) != s(bool,X16) | s(bool,'T_0') != s(bool,X15)) ) | $spl6),
  inference(subsumption_resolution,[],[f4166,f389])).
fof(f4166,plain,(
  ( ! [X17,X15,X16] : (s(bool,'T_0') != s(bool,X15) | s(bool,X15) != s(bool,X16) | p(s(bool,X17)) | sP6(X16,'F_0',X17)) ) | $spl6),
  inference(resolution,[],[f1354,f373])).
fof(f1354,plain,(
  ( ! [X12,X13,X11] : (sP9(X13,'F_0',X11) | s(bool,'T_0') != s(bool,X12) | s(bool,X11) != s(bool,X12) | p(s(bool,X13))) ) | $spl6),
  inference(superposition,[],[f1160,f392])).
fof(f13940,plain,(
  ( ! [X26] : (s(bool,'T_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | ~p(s(bool,X26))) )),
  inference(resolution,[],[f726,f1924])).
fof(f13937,plain,(
  ( ! [X23] : (~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~p(s(bool,X23))) )),
  inference(resolution,[],[f726,f1809])).
fof(f13936,plain,(
  ( ! [X22] : (sP9(X22,'F_0','const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | s(bool,'T_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(resolution,[],[f726,f1794])).
fof(f1794,plain,(
  ( ! [X4,X5] : (~sP8(X5,X4,'F_0') | sP9(X5,'F_0',X4) | s(bool,'T_0') != s(bool,X4)) )),
  inference(subsumption_resolution,[],[f1793,f525])).
fof(f525,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f144])).
fof(f144,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.relation.RC_RTC',aHL_FALSITY)).
fof(f1793,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | sP9(X5,'F_0',X4) | p(s(bool,'F_0')) | ~sP8(X5,X4,'F_0')) )),
  inference(subsumption_resolution,[],[f1790,f387])).
fof(f1790,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | sP9(X5,'F_0',X4) | p(s(bool,'F_0')) | ~sP8(X5,X4,'F_0') | ~sP7('F_0',X4,X5)) )),
  inference(resolution,[],[f1786,f378])).
fof(f13951,plain,(
  ( ! [X21,X20] : (sP9(X20,X21,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | s(bool,'T_0') = s(bool,X21)) ) | $spl154),
  inference(cnf_transformation,[],[f13951_D])).
fof(f13951_D,plain,(
  ( ! [X21,X20] : (sP9(X20,X21,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | s(bool,'T_0') = s(bool,X21)) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f13935,plain,(
  ( ! [X21,X20] : (sP9(X20,X21,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | s(bool,'T_0') = s(bool,X21)) ) | $spl140),
  inference(resolution,[],[f726,f13494])).
fof(f13934,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'F_0') != s(bool,X19) | s(bool,'T_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) | $spl6),
  inference(resolution,[],[f726,f8227])).
fof(f13933,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,X16) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~p(s(bool,X17))) ) | $spl6),
  inference(resolution,[],[f726,f8224])).
fof(f5226,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | ~sP7(X1,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X0)) ) | $spl78),
  inference(cnf_transformation,[],[f5226_D])).
fof(f5226_D,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | ~sP7(X1,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X0)) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f13931,plain,(
  ( ! [X12,X13] : (~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | sP9(X12,X13,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | ~sP7(X13,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X12)) ) | $spl6),
  inference(resolution,[],[f726,f4377])).
fof(f5229,plain,(
  ( ! [X10,X11] : (sP9(X10,X11,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | s(bool,'F_0') != s(bool,X11)) ) | $spl82),
  inference(cnf_transformation,[],[f5229_D])).
fof(f5229_D,plain,(
  ( ! [X10,X11] : (sP9(X10,X11,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | s(bool,'F_0') != s(bool,X11)) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f13930,plain,(
  ( ! [X10,X11] : (sP9(X10,X11,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | s(bool,'F_0') != s(bool,X11)) ) | $spl6),
  inference(resolution,[],[f726,f4304])).
fof(f5228,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | p(s(bool,X9))) ) | $spl80),
  inference(cnf_transformation,[],[f5228_D])).
fof(f5228_D,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | p(s(bool,X9))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f13929,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | p(s(bool,X9)) | s(bool,'T_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) | $spl6),
  inference(resolution,[],[f726,f4298])).
fof(f13928,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X7))) )),
  inference(resolution,[],[f726,f380])).
fof(f13949,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~sP7(X1,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X0)) )),
  inference(subsumption_resolution,[],[f13925,f391])).
fof(f13925,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~sP7(X1,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X0) | ~sP6('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X1,X0)) )),
  inference(resolution,[],[f726,f377])).
fof(f5203,plain,(
  ( ! [X2,X3] : (sP5(X2,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X3) | ~p(s(bool,X3))) ) | $spl76),
  inference(cnf_transformation,[],[f5203_D])).
fof(f5203_D,plain,(
  ( ! [X2,X3] : (sP5(X2,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X3) | ~p(s(bool,X3))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f13913,plain,(
  ( ! [X2,X3] : (sP5(X2,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X3) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~p(s(bool,X3))) )),
  inference(resolution,[],[f714,f359])).
fof(f359,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP4(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f230])).
fof(f230,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | (sP4(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP4(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f229])).
fof(f229,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | (sP4(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP4(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP5(X2,X0,X1)))),
  inference(flattening,[],[f228])).
fof(f228,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | (sP4(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP4(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP5(X2,X0,X1)))),
  inference(nnf_transformation,[],[f200])).
fof(f200,plain,(
  ! [X2,X0,X1] : (sP5(X2,X0,X1) <=> (sP4(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f714,plain,(
  ( ! [X76,X75] : (sP4('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X75,X76)) )),
  inference(resolution,[],[f319,f366])).
fof(f366,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f232])).
fof(f232,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(flattening,[],[f231])).
fof(f231,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(nnf_transformation,[],[f199])).
fof(f199,plain,(
  ! [X0,X1,X2] : (sP4(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f13912,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(resolution,[],[f714,f363])).
fof(f363,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f232])).
fof(f5189,plain,(
  ( ! [X2,X3] : (sP5(X2,X3,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | ~p(s(bool,X3))) ) | $spl74),
  inference(cnf_transformation,[],[f5189_D])).
fof(f5189_D,plain,(
  ( ! [X2,X3] : (sP5(X2,X3,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | ~p(s(bool,X3))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f13900,plain,(
  ( ! [X2,X3] : (sP5(X2,X3,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | ~p(s(bool,X3)) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f713,f359])).
fof(f713,plain,(
  ( ! [X74,X73] : (sP4(X73,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X74)) )),
  inference(resolution,[],[f319,f365])).
fof(f365,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f232])).
fof(f13899,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X1))) )),
  inference(resolution,[],[f713,f363])).
fof(f5175,plain,(
  ( ! [X2,X3] : (sP3('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X2,X3)) ) | $spl72),
  inference(cnf_transformation,[],[f5175_D])).
fof(f5175_D,plain,(
  ( ! [X2,X3] : (sP3('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X2,X3)) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f13898,plain,(
  ( ! [X6,X7] : (sP3('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X6,X7) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(subsumption_resolution,[],[f13889,f352])).
fof(f352,plain,(
  ( ! [X2,X0,X1] : (sP3(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f227])).
fof(f227,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP3(X2,X0,X1)) & (~sP3(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f226])).
fof(f226,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP3(X2,X0,X1)) & (~sP3(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f198])).
fof(f198,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP3(X2,X0,X1))),
  inference(definition_folding,[],[f76,f197,f196])).
fof(f196,plain,(
  ! [X2,X0,X1] : (sP2(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f197,plain,(
  ! [X2,X0,X1] : (sP3(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP2(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f76,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f75])).
fof(f75,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X19,X20,X21] : ((p(s(bool,X21)) <=> (p(s(bool,X20)) | p(s(bool,X19)))) <=> ((p(s(bool,X21)) | ~p(s(bool,X20))) & (p(s(bool,X21)) | ~p(s(bool,X19))) & (p(s(bool,X20)) | p(s(bool,X19)) | ~p(s(bool,X21)))))),
  file('thm.relation.RC_RTC','thm.sat.sat_25_prove_ltb')).
fof(f13889,plain,(
  ( ! [X6,X7] : (sP3('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X6,X7) | p(s(bool,X6)) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f700,f344])).
fof(f344,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~sP2(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f222])).
fof(f222,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP2(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP2(X0,X1,X2) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f221])).
fof(f221,plain,(
  ! [X2,X0,X1] : ((~sP3(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP2(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP2(X2,X0,X1) | sP3(X2,X0,X1)))),
  inference(flattening,[],[f220])).
fof(f220,plain,(
  ! [X2,X0,X1] : ((~sP3(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP2(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP2(X2,X0,X1)) | sP3(X2,X0,X1)))),
  inference(nnf_transformation,[],[f197])).
fof(f700,plain,(
  ( ! [X47,X48] : (sP2('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X47,X48)) )),
  inference(resolution,[],[f319,f349])).
fof(f349,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f225])).
fof(f225,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(rectify,[],[f224])).
fof(f224,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(flattening,[],[f223])).
fof(f223,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(nnf_transformation,[],[f196])).
fof(f13897,plain,(
  ( ! [X4,X5] : (sP3('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X4,X5) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(subsumption_resolution,[],[f13888,f351])).
fof(f351,plain,(
  ( ! [X2,X0,X1] : (sP3(X2,X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f227])).
fof(f13888,plain,(
  ( ! [X4,X5] : (sP3('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X4,X5) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X5))) )),
  inference(resolution,[],[f700,f343])).
fof(f343,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~sP2(X0,X1,X2) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f222])).
fof(f13887,plain,(
  ( ! [X2,X3] : (sP3('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X2,X3) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f700,f547])).
fof(f547,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~sP2(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(duplicate_literal_removal,[],[f342])).
fof(f342,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~sP2(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f222])).
fof(f13886,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(resolution,[],[f700,f346])).
fof(f346,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f225])).
fof(f5161,plain,(
  ( ! [X2,X3] : (sP1('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X2,X3)) ) | $spl70),
  inference(cnf_transformation,[],[f5161_D])).
fof(f5161_D,plain,(
  ( ! [X2,X3] : (sP1('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X2,X3)) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f13885,plain,(
  ( ! [X6,X7] : (sP1('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X6,X7) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(subsumption_resolution,[],[f13876,f335])).
fof(f335,plain,(
  ( ! [X2,X0,X1] : (sP1(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f219])).
fof(f219,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP1(X2,X0,X1)) & (~sP1(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f218])).
fof(f218,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP1(X2,X0,X1)) & (~sP1(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP1(X2,X0,X1))),
  inference(definition_folding,[],[f161,f194,f193])).
fof(f193,plain,(
  ! [X2,X0,X1] : (sP0(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f194,plain,(
  ! [X2,X0,X1] : (sP1(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP0(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f161,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f73])).
fof(f73,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f58])).
fof(f58,axiom,(
  ! [X19,X20,X21] : ((p(s(bool,X21)) <=> (p(s(bool,X20)) => p(s(bool,X19)))) <=> ((p(s(bool,X21)) | p(s(bool,X20))) & (p(s(bool,X21)) | ~p(s(bool,X19))) & (~p(s(bool,X20)) | p(s(bool,X19)) | ~p(s(bool,X21)))))),
  file('thm.relation.RC_RTC','thm.sat.sat_27_prove_ltb')).
fof(f13876,plain,(
  ( ! [X6,X7] : (sP1('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X6,X7) | p(s(bool,X6)) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f684,f327])).
fof(f327,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | ~sP0(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f214])).
fof(f214,plain,(
  ! [X0,X1,X2] : ((~sP1(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP0(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP0(X0,X1,X2) | sP1(X0,X1,X2)))),
  inference(rectify,[],[f213])).
fof(f213,plain,(
  ! [X2,X0,X1] : ((~sP1(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP0(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP0(X2,X0,X1) | sP1(X2,X0,X1)))),
  inference(flattening,[],[f212])).
fof(f212,plain,(
  ! [X2,X0,X1] : ((~sP1(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP0(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP0(X2,X0,X1)) | sP1(X2,X0,X1)))),
  inference(nnf_transformation,[],[f194])).
fof(f684,plain,(
  ( ! [X15,X16] : (sP0('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X15,X16)) )),
  inference(resolution,[],[f319,f332])).
fof(f332,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f217])).
fof(f217,plain,(
  ! [X0,X1,X2] : ((~sP0(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP0(X0,X1,X2)))),
  inference(rectify,[],[f216])).
fof(f216,plain,(
  ! [X2,X0,X1] : ((~sP0(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP0(X2,X0,X1)))),
  inference(flattening,[],[f215])).
fof(f215,plain,(
  ! [X2,X0,X1] : ((~sP0(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP0(X2,X0,X1)))),
  inference(nnf_transformation,[],[f193])).
fof(f13884,plain,(
  ( ! [X4,X5] : (sP1('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X4,X5) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(subsumption_resolution,[],[f13875,f334])).
fof(f334,plain,(
  ( ! [X2,X0,X1] : (sP1(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f219])).
fof(f13875,plain,(
  ( ! [X4,X5] : (sP1('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X4,X5) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~p(s(bool,X5))) )),
  inference(resolution,[],[f684,f326])).
fof(f326,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | ~sP0(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f214])).
fof(f13874,plain,(
  ( ! [X2,X3] : (sP1('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X2,X3) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f684,f546])).
fof(f546,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | ~sP0(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(duplicate_literal_removal,[],[f325])).
fof(f325,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | ~sP0(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f214])).
fof(f13873,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(resolution,[],[f684,f329])).
fof(f329,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP0(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f217])).
fof(f5149,plain,(
  ( ! [X8,X9] : (sP1(X8,X9,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | p(s(bool,X9))) ) | $spl68),
  inference(cnf_transformation,[],[f5149_D])).
fof(f5149_D,plain,(
  ( ! [X8,X9] : (sP1(X8,X9,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | p(s(bool,X9))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f13864,plain,(
  ( ! [X8,X9] : (sP1(X8,X9,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | p(s(bool,X9)) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f683,f328])).
fof(f328,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | ~sP0(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f214])).
fof(f683,plain,(
  ( ! [X14,X13] : (sP0(X13,X14,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(resolution,[],[f319,f330])).
fof(f330,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f217])).
fof(f13860,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f683,f329])).
fof(f13851,plain,(
  ( ! [X15] : (p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X15))) )),
  inference(resolution,[],[f610,f1864])).
fof(f1864,plain,(
  ( ! [X4,X5] : (~sP6(X4,'F_0',X5) | p(s(bool,X4)) | p(s(bool,X5))) )),
  inference(subsumption_resolution,[],[f1858,f507])).
fof(f1858,plain,(
  ( ! [X4,X5] : (~sP6(X4,'F_0',X5) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,'F_0') != s(bool,X4)) )),
  inference(resolution,[],[f1842,f395])).
fof(f1842,plain,(
  ( ! [X2,X3] : (sP9(X2,'F_0',X3) | ~sP6(X3,'F_0',X2) | p(s(bool,X3))) )),
  inference(resolution,[],[f1822,f382])).
fof(f1822,plain,(
  ( ! [X4,X5] : (~sP8(X4,X5,'F_0') | sP9(X4,'F_0',X5) | ~sP6(X5,'F_0',X4)) )),
  inference(subsumption_resolution,[],[f1820,f525])).
fof(f1820,plain,(
  ( ! [X4,X5] : (sP9(X4,'F_0',X5) | p(s(bool,'F_0')) | ~sP8(X4,X5,'F_0') | ~sP6(X5,'F_0',X4)) )),
  inference(resolution,[],[f1817,f378])).
fof(f610,plain,(
  ( ! [X112,X113] : (sP6('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X112,X113)) )),
  inference(resolution,[],[f318,f391])).
fof(f5063,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | $spl45),
  inference(cnf_transformation,[],[f5063_D])).
fof(f5063_D,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) <=> ~$spl45),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl45])])).
fof(f13850,plain,(
  ( ! [X14] : (p(s(bool,X14)) | s(bool,'F_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(resolution,[],[f610,f1855])).
fof(f9024,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) | $spl132),
  inference(cnf_transformation,[],[f9024_D])).
fof(f9024_D,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f13848,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | s(bool,'F_0') != s(bool,X11) | s(bool,X11) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) | $spl6),
  inference(resolution,[],[f610,f4784])).
fof(f5131,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,X0) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) | $spl64),
  inference(cnf_transformation,[],[f5131_D])).
fof(f5131_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,X0) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f13847,plain,(
  ( ! [X8,X9] : (p(s(bool,X8)) | p(s(bool,X9)) | s(bool,X8) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) | $spl6),
  inference(resolution,[],[f610,f4655])).
fof(f13846,plain,(
  ( ! [X6,X7] : (p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X6)) | p(s(bool,X7))) )),
  inference(resolution,[],[f610,f388])).
fof(f388,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f245])).
fof(f13791,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | $spl148),
  inference(cnf_transformation,[],[f13791_D])).
fof(f13791_D,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f13842,plain,(
  ( ! [X12,X13] : (sP9(X12,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X13) | p(s(bool,X12))) ) | $spl152),
  inference(cnf_transformation,[],[f13842_D])).
fof(f13842_D,plain,(
  ( ! [X12,X13] : (sP9(X12,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X13) | p(s(bool,X12))) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f13832,plain,(
  ( ! [X12,X13] : (sP9(X12,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X13) | p(s(bool,X12)) | s(bool,'T_0') = s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) | $spl140),
  inference(resolution,[],[f609,f13495])).
fof(f609,plain,(
  ( ! [X111,X110] : (sP6(X110,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X111)) )),
  inference(resolution,[],[f318,f390])).
fof(f5113,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) | $spl60),
  inference(cnf_transformation,[],[f5113_D])).
fof(f5113_D,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f13831,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | s(bool,'F_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | s(bool,X11) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) | $spl6),
  inference(resolution,[],[f609,f4784])).
fof(f13830,plain,(
  ( ! [X8,X9] : (p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X8)) | s(bool,X9) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) | $spl6),
  inference(resolution,[],[f609,f4655])).
fof(f13829,plain,(
  ( ! [X6,X7] : (p(s(bool,X6)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X7))) )),
  inference(resolution,[],[f609,f388])).
fof(f5116,plain,(
  ( ! [X6,X7] : (sP9(X6,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X7)) ) | $spl62),
  inference(cnf_transformation,[],[f5116_D])).
fof(f5116_D,plain,(
  ( ! [X6,X7] : (sP9(X6,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X7)) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f13840,plain,(
  ( ! [X2,X3] : (sP9(X2,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X3) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(subsumption_resolution,[],[f13839,f604])).
fof(f604,plain,(
  ( ! [X101,X100] : (sP8(X100,X101,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(resolution,[],[f318,f381])).
fof(f13839,plain,(
  ( ! [X2,X3] : (sP9(X2,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X3) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~sP8(X2,X3,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(subsumption_resolution,[],[f13827,f387])).
fof(f13827,plain,(
  ( ! [X2,X3] : (sP9(X2,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X3) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~sP8(X2,X3,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | ~sP7('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X3,X2)) )),
  inference(resolution,[],[f609,f378])).
fof(f13815,plain,(
  ( ! [X15] : (p(s(bool,X15)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f608,f1864])).
fof(f608,plain,(
  ( ! [X109,X108] : (sP6(X108,X109,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(resolution,[],[f318,f389])).
fof(f13814,plain,(
  ( ! [X14] : (p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | s(bool,'F_0') != s(bool,X14)) )),
  inference(resolution,[],[f608,f1855])).
fof(f13825,plain,(
  ( ! [X12,X13] : (sP9('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X12,X13) | s(bool,'T_0') = s(bool,X12)) ) | $spl150),
  inference(cnf_transformation,[],[f13825_D])).
fof(f13825_D,plain,(
  ( ! [X12,X13] : (sP9('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X12,X13) | s(bool,'T_0') = s(bool,X12)) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f13813,plain,(
  ( ! [X12,X13] : (sP9('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X12,X13) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | s(bool,'T_0') = s(bool,X12)) ) | $spl140),
  inference(resolution,[],[f608,f13495])).
fof(f8835,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,X0) | s(bool,X0) != s(bool,X1)) ) | $spl120),
  inference(cnf_transformation,[],[f8835_D])).
fof(f8835_D,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,X0) | s(bool,X0) != s(bool,X1)) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f13812,plain,(
  ( ! [X10,X11] : (p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | s(bool,'F_0') != s(bool,X10) | s(bool,X10) != s(bool,X11)) ) | $spl6),
  inference(resolution,[],[f608,f4784])).
fof(f13811,plain,(
  ( ! [X8,X9] : (p(s(bool,X8)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | s(bool,X8) != s(bool,X9)) ) | $spl6),
  inference(resolution,[],[f608,f4655])).
fof(f13810,plain,(
  ( ! [X6,X7] : (p(s(bool,X6)) | p(s(bool,X7)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f608,f388])).
fof(f5100,plain,(
  ( ! [X8,X9] : (sP9('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X8,X9)) ) | $spl56),
  inference(cnf_transformation,[],[f5100_D])).
fof(f5100_D,plain,(
  ( ! [X8,X9] : (sP9('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X8,X9)) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f13824,plain,(
  ( ! [X4,X5] : (sP9('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X4,X5) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(subsumption_resolution,[],[f13823,f607])).
fof(f607,plain,(
  ( ! [X107,X106] : (sP7(X106,X107,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(resolution,[],[f318,f385])).
fof(f13823,plain,(
  ( ! [X4,X5] : (sP9('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X4,X5) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~sP7(X4,X5,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(subsumption_resolution,[],[f13809,f383])).
fof(f13809,plain,(
  ( ! [X4,X5] : (sP9('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X4,X5) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~sP8('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X5,X4) | ~sP7(X4,X5,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(resolution,[],[f608,f379])).
fof(f5082,plain,(
  ( ! [X12,X13] : (sP9('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X12,X13) | ~p(s(bool,X13))) ) | $spl52),
  inference(cnf_transformation,[],[f5082_D])).
fof(f5082_D,plain,(
  ( ! [X12,X13] : (sP9('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X12,X13) | ~p(s(bool,X13))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f13798,plain,(
  ( ! [X12,X13] : (sP9('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X12,X13) | ~p(s(bool,X13)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) ) | $spl6),
  inference(resolution,[],[f607,f4398])).
fof(f5081,plain,(
  ( ! [X8,X9] : (sP9('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X8,X9) | s(bool,'T_0') != s(bool,X9)) ) | $spl50),
  inference(cnf_transformation,[],[f5081_D])).
fof(f5081_D,plain,(
  ( ! [X8,X9] : (sP9('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X8,X9) | s(bool,'T_0') != s(bool,X9)) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f13796,plain,(
  ( ! [X8,X9] : (sP9('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X8,X9) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | s(bool,'T_0') != s(bool,X9)) ) | $spl6),
  inference(resolution,[],[f607,f4378])).
fof(f13795,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f607,f384])).
fof(f5080,plain,(
  ( ! [X4,X5] : (sP9('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X4,X5) | ~sP6(X5,X4,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) | $spl48),
  inference(cnf_transformation,[],[f5080_D])).
fof(f5080_D,plain,(
  ( ! [X4,X5] : (sP9('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X4,X5) | ~sP6(X5,X4,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f13806,plain,(
  ( ! [X4,X5] : (sP9('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X4,X5) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~sP6(X5,X4,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(subsumption_resolution,[],[f13794,f383])).
fof(f13794,plain,(
  ( ! [X4,X5] : (sP9('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X4,X5) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~sP8('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X5,X4) | ~sP6(X5,X4,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(resolution,[],[f607,f379])).
fof(f5064,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X11) | ~p(s(bool,X11))) ) | $spl46),
  inference(cnf_transformation,[],[f5064_D])).
fof(f5064_D,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X11) | ~p(s(bool,X11))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f13781,plain,(
  ( ! [X21,X20] : (sP9(X20,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X21) | ~p(s(bool,X21)) | s(bool,'T_0') = s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) | $spl140),
  inference(resolution,[],[f604,f13494])).
fof(f13780,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'F_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | s(bool,'T_0') != s(bool,X19)) ) | $spl6),
  inference(resolution,[],[f604,f8227])).
fof(f13779,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | ~p(s(bool,X16)) | ~p(s(bool,X17))) ) | $spl6),
  inference(resolution,[],[f604,f8224])).
fof(f5060,plain,(
  ( ! [X2,X3] : (sP9(X2,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X3) | ~sP6(X3,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X2)) ) | $spl40),
  inference(cnf_transformation,[],[f5060_D])).
fof(f5060_D,plain,(
  ( ! [X2,X3] : (sP9(X2,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X3) | ~sP6(X3,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X2)) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f13778,plain,(
  ( ! [X14,X15] : (p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | sP9(X14,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X15) | ~sP6(X15,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X14)) ) | $spl6),
  inference(resolution,[],[f604,f4575])).
fof(f13776,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X11) | ~p(s(bool,X11)) | s(bool,'F_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) | $spl6),
  inference(resolution,[],[f604,f4304])).
fof(f5061,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X9) | s(bool,'T_0') != s(bool,X9)) ) | $spl42),
  inference(cnf_transformation,[],[f5061_D])).
fof(f5061_D,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X9) | s(bool,'T_0') != s(bool,X9)) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f13775,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X9) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | s(bool,'T_0') != s(bool,X9)) ) | $spl6),
  inference(resolution,[],[f604,f4298])).
fof(f13774,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f604,f380])).
fof(f13789,plain,(
  ( ! [X2,X3] : (sP9(X2,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X3) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~sP6(X3,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X2)) )),
  inference(subsumption_resolution,[],[f13772,f387])).
fof(f13772,plain,(
  ( ! [X2,X3] : (sP9(X2,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X3) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~sP7('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X3,X2) | ~sP6(X3,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X2)) )),
  inference(resolution,[],[f604,f378])).
fof(f5042,plain,(
  ( ! [X4,X5] : (sP5('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X4,X5)) ) | $spl38),
  inference(cnf_transformation,[],[f5042_D])).
fof(f5042_D,plain,(
  ( ! [X4,X5] : (sP5('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X4,X5)) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f13762,plain,(
  ( ! [X8,X9] : (sP5('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X8,X9) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f591,f548])).
fof(f548,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | p(s(bool,X0)) | ~sP4(X1,X2,X0)) )),
  inference(duplicate_literal_removal,[],[f362])).
fof(f362,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X0)) | ~sP4(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f230])).
fof(f591,plain,(
  ( ! [X74,X75] : (sP4(X74,X75,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(resolution,[],[f318,f364])).
fof(f364,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f232])).
fof(f13770,plain,(
  ( ! [X6,X7] : (sP5('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X6,X7) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(subsumption_resolution,[],[f13761,f367])).
fof(f367,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f234])).
fof(f234,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP5(X2,X0,X1)) & (~sP5(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f233])).
fof(f233,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP5(X2,X0,X1)) & (~sP5(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f201])).
fof(f201,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP5(X2,X0,X1))),
  inference(definition_folding,[],[f78,f200,f199])).
fof(f78,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f77])).
fof(f77,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f56])).
fof(f56,axiom,(
  ! [X19,X20,X21] : ((p(s(bool,X21)) <=> (p(s(bool,X20)) & p(s(bool,X19)))) <=> ((p(s(bool,X21)) | ~p(s(bool,X20)) | ~p(s(bool,X19))) & (p(s(bool,X20)) | ~p(s(bool,X21))) & (p(s(bool,X19)) | ~p(s(bool,X21)))))),
  file('thm.relation.RC_RTC','thm.sat.sat_23_prove_ltb')).
fof(f13761,plain,(
  ( ! [X6,X7] : (sP5('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X6,X7) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~p(s(bool,X7))) )),
  inference(resolution,[],[f591,f361])).
fof(f361,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X2)) | ~sP4(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f230])).
fof(f13769,plain,(
  ( ! [X4,X5] : (sP5('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X4,X5) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(subsumption_resolution,[],[f13760,f368])).
fof(f368,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f234])).
fof(f13760,plain,(
  ( ! [X4,X5] : (sP5('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X4,X5) | ~p(s(bool,X4)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f591,f360])).
fof(f360,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~sP4(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f230])).
fof(f13758,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f591,f363])).
fof(f5029,plain,(
  ( ! [X8,X9] : (sP3(X8,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X9) | p(s(bool,X9))) ) | $spl36),
  inference(cnf_transformation,[],[f5029_D])).
fof(f5029_D,plain,(
  ( ! [X8,X9] : (sP3(X8,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X9) | p(s(bool,X9))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f13749,plain,(
  ( ! [X8,X9] : (sP3(X8,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X9) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X9))) )),
  inference(resolution,[],[f575,f345])).
fof(f345,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~sP2(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f222])).
fof(f575,plain,(
  ( ! [X43,X42] : (sP2(X42,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X43)) )),
  inference(resolution,[],[f318,f348])).
fof(f348,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f225])).
fof(f13745,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X1))) )),
  inference(resolution,[],[f575,f346])).
fof(f5015,plain,(
  ( ! [X8,X9] : (sP3(X8,X9,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | p(s(bool,X9))) ) | $spl34),
  inference(cnf_transformation,[],[f5015_D])).
fof(f5015_D,plain,(
  ( ! [X8,X9] : (sP3(X8,X9,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | p(s(bool,X9))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f13736,plain,(
  ( ! [X8,X9] : (sP3(X8,X9,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | p(s(bool,X9)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f574,f345])).
fof(f574,plain,(
  ( ! [X41,X40] : (sP2(X40,X41,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(resolution,[],[f318,f347])).
fof(f347,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f225])).
fof(f13732,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f574,f346])).
fof(f5001,plain,(
  ( ! [X8,X9] : (sP1(X8,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X9) | ~p(s(bool,X9))) ) | $spl32),
  inference(cnf_transformation,[],[f5001_D])).
fof(f5001_D,plain,(
  ( ! [X8,X9] : (sP1(X8,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X9) | ~p(s(bool,X9))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f13723,plain,(
  ( ! [X8,X9] : (sP1(X8,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X9) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~p(s(bool,X9))) )),
  inference(resolution,[],[f561,f328])).
fof(f561,plain,(
  ( ! [X14,X15] : (sP0(X14,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X15)) )),
  inference(resolution,[],[f318,f331])).
fof(f331,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f217])).
fof(f13719,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~p(s(bool,X1))) )),
  inference(resolution,[],[f561,f329])).
fof(f13718,plain,(
  ( ! [X1] : (p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X1))) )),
  inference(subsumption_resolution,[],[f13710,f480])).
fof(f13710,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X1))) )),
  inference(resolution,[],[f622,f432])).
fof(f622,plain,(
  ( ! [X127] : (sP11('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X127)) )),
  inference(resolution,[],[f318,f434])).
fof(f434,plain,(
  ( ! [X0,X1] : (sP11(X0,X1) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f262])).
fof(f13717,plain,(
  ( ! [X0] : (p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f13709,f480])).
fof(f13709,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X0))) )),
  inference(resolution,[],[f622,f436])).
fof(f436,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP11(X0,X1)) )),
  inference(cnf_transformation,[],[f264])).
fof(f264,plain,(
  ! [X0,X1] : ((~sP11(X0,X1) | (p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP11(X0,X1)))),
  inference(flattening,[],[f263])).
fof(f263,plain,(
  ! [X0,X1] : ((~sP11(X0,X1) | ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0')))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP11(X0,X1)))),
  inference(nnf_transformation,[],[f210])).
fof(f210,plain,(
  ! [X0,X1] : (sP11(X0,X1) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(definition_folding,[],[f167,f209])).
fof(f167,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f166])).
fof(f166,plain,(
  ! [X0,X1] : (((p(s(bool,X1)) | p(s(bool,X0))) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | (p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(ennf_transformation,[],[f92])).
fof(f92,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(flattening,[],[f91])).
fof(f91,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(rectify,[],[f52])).
fof(f52,axiom,(
  ! [X16,X6] : ((~(p(s(bool,X6)) | p(s(bool,X16))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X6)) => p(s(bool,'F_0'))) => (~p(s(bool,X16)) => p(s(bool,'F_0')))))),
  file('thm.relation.RC_RTC','thm.sat.sat_9_prove_ltb')).
fof(f13708,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(subsumption_resolution,[],[f13700,f480])).
fof(f13700,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f621,f432])).
fof(f621,plain,(
  ( ! [X126] : (sP11(X126,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(resolution,[],[f318,f433])).
fof(f13707,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(subsumption_resolution,[],[f13699,f480])).
fof(f13699,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f621,f436])).
fof(f13693,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1))) ) | $spl0),
  inference(subsumption_resolution,[],[f13684,f321])).
fof(f13684,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | $spl0),
  inference(resolution,[],[f13590,f329])).
fof(f13590,plain,(
  ( ! [X14,X13] : (sP0(X13,X14,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | $spl0),
  inference(backward_demodulation,[],[f668,f683])).
fof(f668,plain,(
  s(sK12,sK14) = s(sK12,sK15) | $spl0),
  inference(cnf_transformation,[],[f668_D])).
fof(f668_D,plain,(
  s(sK12,sK14) = s(sK12,sK15) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f670,plain,(
  p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) | $spl2),
  inference(cnf_transformation,[],[f670_D])).
fof(f670_D,plain,(
  p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f9969,plain,(
  p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) | s(sK12,sK14) = s(sK12,sK15)),
  inference(subsumption_resolution,[],[f9709,f549])).
fof(f549,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f455])).
fof(f455,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f269])).
fof(f269,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f104])).
fof(f104,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f103])).
fof(f103,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X6] : ! [X8] : (s(X6,X8) = s(X6,X8) <=> p(s(bool,'T_0')))),
  file('thm.relation.RC_RTC','thm.bool.REFL_CLAUSE')).
fof(f9709,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) | s(sK12,sK14) = s(sK12,sK15)),
  inference(superposition,[],[f532,f636])).
fof(f532,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | s(X0,X2) = s(X0,X3) | ~p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f309])).
fof(f13522,plain,(
  ( ! [X37,X36] : (sP9(X36,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X37) | p(s(bool,X36)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) ) | ($spl2 | $spl140)),
  inference(resolution,[],[f13495,f3371])).
fof(f13502,plain,(
  ( ! [X17,X16] : (sP9(X16,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X17) | ~p(s(bool,X17)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) ) | ($spl2 | $spl140)),
  inference(resolution,[],[f13494,f3366])).
fof(f13415,plain,(
  ( ! [X116,X114,X115] : (p(s(bool,X114)) | ~p(s(bool,X115)) | sP9(X116,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X115) | p(s(bool,X116))) ) | $spl6),
  inference(subsumption_resolution,[],[f13369,f507])).
fof(f13369,plain,(
  ( ! [X116,X114,X115] : (s(bool,'F_0') != s(bool,X114) | p(s(bool,X114)) | ~p(s(bool,X115)) | sP9(X116,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X115) | p(s(bool,X116))) ) | $spl6),
  inference(superposition,[],[f2032,f758])).
fof(f13411,plain,(
  ( ! [X66,X64,X67] : (sP9(X67,X64,X66) | s(bool,'T_0') = s(bool,X64) | p(s(bool,X67)) | ~p(s(bool,X66))) ) | $spl140),
  inference(cnf_transformation,[],[f13411_D])).
fof(f13411_D,plain,(
  ( ! [X66,X64,X67] : (sP9(X67,X64,X66) | s(bool,'T_0') = s(bool,X64) | p(s(bool,X67)) | ~p(s(bool,X66))) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f13412,plain,(
  ( ! [X74,X72,X75,X73] : (p(s(bool,X73)) | ~p(s(bool,X74)) | sP9(X75,X72,X74) | p(s(bool,X75)) | s(bool,'T_0') = s(bool,X72)) ) | $spl6),
  inference(subsumption_resolution,[],[f13358,f507])).
fof(f13358,plain,(
  ( ! [X74,X72,X75,X73] : (s(bool,'F_0') != s(bool,X73) | p(s(bool,X73)) | ~p(s(bool,X74)) | sP9(X75,X72,X74) | p(s(bool,X75)) | s(bool,'T_0') = s(bool,X72)) ) | $spl6),
  inference(superposition,[],[f2032,f539])).
fof(f13410,plain,(
  ( ! [X66,X64,X67,X65] : (p(s(bool,X65)) | ~p(s(bool,X66)) | sP9(X67,X64,X66) | p(s(bool,X67)) | s(bool,'T_0') = s(bool,X64)) ) | $spl6),
  inference(subsumption_resolution,[],[f13356,f507])).
fof(f13356,plain,(
  ( ! [X66,X64,X67,X65] : (s(bool,'F_0') != s(bool,X65) | p(s(bool,X65)) | ~p(s(bool,X66)) | sP9(X67,X64,X66) | p(s(bool,X67)) | s(bool,'T_0') = s(bool,X64)) ) | $spl6),
  inference(superposition,[],[f2032,f538])).
fof(f13266,plain,(
  ( ! [X12] : (p(s(bool,X12)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) ) | $spl2),
  inference(resolution,[],[f3372,f1855])).
fof(f13261,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) ) | ($spl2 | $spl6)),
  inference(resolution,[],[f3372,f4655])).
fof(f13260,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'F_0') != s(bool,X1) | s(bool,X1) != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) ) | ($spl2 | $spl6)),
  inference(resolution,[],[f3372,f4784])).
fof(f13244,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))) | s(bool,X1) != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) ) | ($spl2 | $spl6)),
  inference(resolution,[],[f3371,f4784])).
fof(f13165,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))) | s(bool,'T_0') != s(bool,X19)) ) | ($spl2 | $spl6)),
  inference(resolution,[],[f3366,f8227])).
fof(f13164,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))) | ~p(s(bool,X16)) | ~p(s(bool,X17))) ) | ($spl2 | $spl6)),
  inference(resolution,[],[f3366,f8224])).
fof(f13161,plain,(
  ( ! [X10,X11] : (sP9(X10,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X11) | ~p(s(bool,X11)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) ) | ($spl2 | $spl6)),
  inference(resolution,[],[f3366,f4304])).
fof(f10144,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(duplicate_literal_removal,[],[f10135])).
fof(f10135,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~p(s(bool,X4))) )),
  inference(resolution,[],[f739,f417])).
fof(f739,plain,(
  ( ! [X122] : (sP10('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X122) | ~p(s(bool,X122))) )),
  inference(resolution,[],[f319,f421])).
fof(f10145,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(duplicate_literal_removal,[],[f10132])).
fof(f10132,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | ~p(s(bool,X1)) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f739,f424])).
fof(f10127,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(duplicate_literal_removal,[],[f10118])).
fof(f10118,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | ~p(s(bool,X4)) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f738,f417])).
fof(f10128,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(duplicate_literal_removal,[],[f10115])).
fof(f10115,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~p(s(bool,X1))) )),
  inference(resolution,[],[f738,f424])).
fof(f10111,plain,(
  ( ! [X2] : (p(s(bool,X2)) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(duplicate_literal_removal,[],[f10102])).
fof(f10102,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,X2)) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f737,f418])).
fof(f10112,plain,(
  ( ! [X1] : (p(s(bool,X1)) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(duplicate_literal_removal,[],[f10101])).
fof(f10101,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,X1)) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f737,f423])).
fof(f10097,plain,(
  ( ! [X3] : (p(s(bool,X3)) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(duplicate_literal_removal,[],[f10089])).
fof(f10089,plain,(
  ( ! [X3] : (p(s(bool,X3)) | p(s(bool,X3)) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f736,f421])).
fof(f10098,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(duplicate_literal_removal,[],[f10086])).
fof(f10086,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X0))) )),
  inference(resolution,[],[f736,f422])).
fof(f9687,plain,(
  ( ! [X3] : (p(s(bool,X3)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(duplicate_literal_removal,[],[f9678])).
fof(f9678,plain,(
  ( ! [X3] : (p(s(bool,X3)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X3))) )),
  inference(resolution,[],[f616,f416])).
fof(f616,plain,(
  ( ! [X121] : (sP10(X121,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | p(s(bool,X121))) )),
  inference(resolution,[],[f318,f421])).
fof(f9688,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(duplicate_literal_removal,[],[f9677])).
fof(f9677,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,X2)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f616,f425])).
fof(f425,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | ~sP10(X1,X0)) )),
  inference(cnf_transformation,[],[f258])).
fof(f9673,plain,(
  ( ! [X3] : (p(s(bool,X3)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(duplicate_literal_removal,[],[f9664])).
fof(f9664,plain,(
  ( ! [X3] : (p(s(bool,X3)) | p(s(bool,X3)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f615,f416])).
fof(f615,plain,(
  ( ! [X120] : (sP10('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X120) | p(s(bool,X120))) )),
  inference(resolution,[],[f318,f418])).
fof(f9674,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(duplicate_literal_removal,[],[f9663])).
fof(f9663,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X2))) )),
  inference(resolution,[],[f615,f425])).
fof(f9657,plain,(
  ( ! [X13] : (~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~p(s(bool,X13))) ) | ($spl6 | $spl12)),
  inference(trivial_inequality_removal,[],[f9650])).
fof(f9650,plain,(
  ( ! [X13] : (~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~p(s(bool,X13)) | s(bool,'F_0') != s(bool,'F_0')) ) | ($spl6 | $spl12)),
  inference(resolution,[],[f1940,f8236])).
fof(f8236,plain,(
  ( ! [X10,X11,X9] : (~sP9(X11,X9,X10) | ~p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,'F_0') != s(bool,X9)) ) | $spl6),
  inference(resolution,[],[f8224,f375])).
fof(f375,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X2,X1) | ~sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f237])).
fof(f9658,plain,(
  ( ! [X11] : (s(bool,'T_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | ~p(s(bool,X11))) ) | ($spl6 | $spl12)),
  inference(trivial_inequality_removal,[],[f9648])).
fof(f9648,plain,(
  ( ! [X11] : (s(bool,'T_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | ~p(s(bool,X11)) | s(bool,'F_0') != s(bool,'F_0')) ) | ($spl6 | $spl12)),
  inference(resolution,[],[f1940,f1373])).
fof(f9659,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | $spl139),
  inference(cnf_transformation,[],[f9659_D])).
fof(f9659_D,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) <=> ~$spl139),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl139])])).
fof(f9647,plain,(
  ( ! [X10] : (p(s(bool,X10)) | s(bool,'F_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) | $spl12),
  inference(resolution,[],[f1940,f395])).
fof(f9660,plain,(
  s(bool,'F_0') = s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | $spl138),
  inference(cnf_transformation,[],[f9660_D])).
fof(f9660_D,plain,(
  s(bool,'F_0') = s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f9646,plain,(
  ( ! [X9] : (s(bool,'F_0') = s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | ~p(s(bool,X9))) ) | $spl12),
  inference(resolution,[],[f1940,f394])).
fof(f9641,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) ) | $spl12),
  inference(resolution,[],[f1940,f1827])).
fof(f1827,plain,(
  ( ! [X6,X7] : (~sP9(X7,'F_0',X6) | ~p(s(bool,X7)) | ~p(s(bool,X6))) )),
  inference(resolution,[],[f1809,f375])).
fof(f9640,plain,(
  ( ! [X3] : (p(s(bool,X3)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) ) | $spl12),
  inference(resolution,[],[f1940,f1869])).
fof(f1869,plain,(
  ( ! [X8,X9] : (~sP9(X9,'F_0',X8) | p(s(bool,X9)) | p(s(bool,X8))) )),
  inference(resolution,[],[f1864,f373])).
fof(f9639,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,'T_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) | $spl12),
  inference(resolution,[],[f1940,f1135])).
fof(f9638,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) | ($spl6 | $spl12)),
  inference(resolution,[],[f1940,f1356])).
fof(f9630,plain,(
  ( ! [X12] : (p(s(bool,X12)) | s(bool,'F_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | $spl28),
  inference(resolution,[],[f9428,f1855])).
fof(f9625,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | ($spl6 | $spl28)),
  inference(resolution,[],[f9428,f4655])).
fof(f9624,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'F_0') != s(bool,X1) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | ($spl6 | $spl28)),
  inference(resolution,[],[f9428,f4784])).
fof(f9612,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'F_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | ($spl6 | $spl28)),
  inference(resolution,[],[f9427,f4784])).
fof(f9583,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(duplicate_literal_removal,[],[f9575])).
fof(f9575,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~p(s(bool,X2))) )),
  inference(resolution,[],[f614,f418])).
fof(f614,plain,(
  ( ! [X119] : (~sP10(X119,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | ~p(s(bool,X119))) )),
  inference(resolution,[],[f318,f417])).
fof(f9584,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(duplicate_literal_removal,[],[f9574])).
fof(f9574,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~p(s(bool,X1))) )),
  inference(resolution,[],[f614,f423])).
fof(f9566,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'F_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | s(bool,'T_0') != s(bool,X19)) ) | ($spl6 | $spl28)),
  inference(resolution,[],[f9422,f8227])).
fof(f9565,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | ~p(s(bool,X16)) | ~p(s(bool,X17))) ) | ($spl6 | $spl28)),
  inference(resolution,[],[f9422,f8224])).
fof(f9562,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X11) | ~p(s(bool,X11)) | s(bool,'F_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | ($spl6 | $spl28)),
  inference(resolution,[],[f9422,f4304])).
fof(f9514,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(duplicate_literal_removal,[],[f9507])).
fof(f9507,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~p(s(bool,X3))) )),
  inference(resolution,[],[f613,f421])).
fof(f9515,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(duplicate_literal_removal,[],[f9504])).
fof(f9504,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f613,f422])).
fof(f9361,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X11) | ~p(s(bool,X11)) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) ) | $spl6),
  inference(resolution,[],[f730,f4396])).
fof(f9359,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~p(s(bool,X6)) | p(s(bool,X7))) )),
  inference(resolution,[],[f730,f384])).
fof(f9346,plain,(
  ( ! [X12,X13] : (sP9(X12,X13,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X12))) ) | $spl6),
  inference(resolution,[],[f729,f4398])).
fof(f9345,plain,(
  ( ! [X10,X11] : (sP9(X10,X11,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~p(s(bool,X11))) ) | $spl6),
  inference(resolution,[],[f729,f4396])).
fof(f9344,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | p(s(bool,X8)) | s(bool,'T_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) | $spl6),
  inference(resolution,[],[f729,f4378])).
fof(f9343,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X7))) )),
  inference(resolution,[],[f729,f384])).
fof(f9354,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(subsumption_resolution,[],[f9353,f726])).
fof(f9353,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~sP8(X0,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X1)) )),
  inference(subsumption_resolution,[],[f9340,f391])).
fof(f9340,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~sP8(X0,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X1) | ~sP6('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X1,X0)) )),
  inference(resolution,[],[f729,f377])).
fof(f9330,plain,(
  ( ! [X28,X27] : (~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | s(bool,X27) != s(bool,X28) | s(bool,'T_0') != s(bool,X27)) ) | $spl6),
  inference(resolution,[],[f727,f4249])).
fof(f9328,plain,(
  ( ! [X24] : (s(bool,'T_0') != s(bool,X24) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f727,f1924])).
fof(f9325,plain,(
  ( ! [X21] : (~p(s(bool,X21)) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f727,f1809])).
fof(f9323,plain,(
  ( ! [X19,X18] : (~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | s(bool,'F_0') != s(bool,X18) | s(bool,'T_0') != s(bool,X19)) ) | $spl6),
  inference(resolution,[],[f727,f8227])).
fof(f9322,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,X16) | ~p(s(bool,X17)) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) ) | $spl6),
  inference(resolution,[],[f727,f8224])).
fof(f9317,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~p(s(bool,X6)) | p(s(bool,X7))) )),
  inference(resolution,[],[f727,f380])).
fof(f9316,plain,(
  ( ! [X4,X5] : (sP9('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X4,X5) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~sP7(X4,X5,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | ~sP6(X5,X4,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(resolution,[],[f727,f379])).
fof(f9309,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f9216,f1931])).
fof(f9216,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(resolution,[],[f1931,f528])).
fof(f528,plain,(
  p(s(bool,'F_0')) | p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f307])).
fof(f307,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f146])).
fof(f146,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f32])).
fof(f32,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.relation.RC_RTC','thm.bool.NOT_CLAUSES|split|1')).
fof(f9308,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f9215,f1931])).
fof(f9215,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(resolution,[],[f1931,f526])).
fof(f526,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f306])).
fof(f306,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f145])).
fof(f145,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f33])).
fof(f33,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.relation.RC_RTC','thm.bool.NOT_CLAUSES|split|2')).
fof(f9307,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f9210,f1931])).
fof(f9210,plain,(
  ( ! [X216] : (p(s(bool,X216))) ) | $spl8),
  inference(resolution,[],[f1931,f473])).
fof(f473,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f281])).
fof(f281,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,'T_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f280])).
fof(f280,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (~p(s(bool,'T_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f182])).
fof(f182,plain,(
  ! [X0] : ((~p(s(bool,'T_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(ennf_transformation,[],[f114])).
fof(f114,plain,(
  ! [X0] : ((p(s(bool,'T_0')) => p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X10] : ((p(s(bool,'T_0')) => p(s(bool,X10))) <=> p(s(bool,X10)))),
  file('thm.relation.RC_RTC','thm.bool.IMP_CLAUSES|split|0')).
fof(f9306,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f9209,f1931])).
fof(f9209,plain,(
  ( ! [X215] : (p(s(bool,X215))) ) | $spl8),
  inference(resolution,[],[f1931,f470])).
fof(f470,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f279])).
fof(f279,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,'T_0'))))),
  inference(flattening,[],[f278])).
fof(f278,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f181])).
fof(f181,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f113])).
fof(f113,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X10] : ((p(s(bool,X10)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  file('thm.relation.RC_RTC','thm.bool.IMP_CLAUSES|split|1')).
fof(f9305,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f9207,f1931])).
fof(f9207,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(resolution,[],[f1931,f467])).
fof(f467,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f277])).
fof(f277,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f276])).
fof(f276,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f180])).
fof(f180,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f112])).
fof(f112,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X10] : ((p(s(bool,'F_0')) => p(s(bool,X10))) <=> p(s(bool,'T_0')))),
  file('thm.relation.RC_RTC','thm.bool.IMP_CLAUSES|split|2')).
fof(f9304,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f9205,f1931])).
fof(f9205,plain,(
  ( ! [X212] : (p(s(bool,X212))) ) | $spl8),
  inference(resolution,[],[f1931,f464])).
fof(f464,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f275])).
fof(f275,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))))),
  inference(flattening,[],[f274])).
fof(f274,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f179])).
fof(f179,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f111])).
fof(f111,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X10] : ((p(s(bool,X10)) => p(s(bool,X10))) <=> p(s(bool,'T_0')))),
  file('thm.relation.RC_RTC','thm.bool.IMP_CLAUSES|split|3')).
fof(f9217,plain,(
  $false | $spl8),
  inference(resolution,[],[f1931,f530])).
fof(f530,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f8])).
fof(f8,axiom,(
  p(s(bool,'T_0'))),
  file('thm.relation.RC_RTC','thm.bool.TRUTH')).
fof(f9218,plain,(
  $false | $spl8),
  inference(resolution,[],[f1931,f531])).
fof(f531,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.relation.RC_RTC',aHL_TRUTH)).
fof(f9219,plain,(
  $false | $spl8),
  inference(resolution,[],[f1931,f549])).
fof(f9229,plain,(
  $false | $spl8),
  inference(resolution,[],[f1931,f321])).
fof(f9108,plain,(
  ( ! [X28,X27] : (~p(s(bool,X27)) | s(bool,X28) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | s(bool,'T_0') != s(bool,X28)) ) | $spl6),
  inference(resolution,[],[f726,f4249])).
fof(f9107,plain,(
  ( ! [X26,X25] : (s(bool,'T_0') != s(bool,X25) | sP9(X26,'F_0','const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | s(bool,X25) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) | $spl6),
  inference(resolution,[],[f726,f4191])).
fof(f9106,plain,(
  ( ! [X24] : (s(bool,'T_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | ~p(s(bool,X24))) )),
  inference(resolution,[],[f726,f1924])).
fof(f9103,plain,(
  ( ! [X21] : (~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~p(s(bool,X21))) )),
  inference(resolution,[],[f726,f1809])).
fof(f9102,plain,(
  ( ! [X20] : (sP9(X20,'F_0','const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | s(bool,'T_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(resolution,[],[f726,f1794])).
fof(f9101,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'F_0') != s(bool,X19) | s(bool,'T_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) | $spl6),
  inference(resolution,[],[f726,f8227])).
fof(f9100,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,X16) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~p(s(bool,X17))) ) | $spl6),
  inference(resolution,[],[f726,f8224])).
fof(f9098,plain,(
  ( ! [X12,X13] : (~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | sP9(X12,X13,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | ~sP7(X13,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X12)) ) | $spl6),
  inference(resolution,[],[f726,f4377])).
fof(f9097,plain,(
  ( ! [X10,X11] : (sP9(X10,X11,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | s(bool,'F_0') != s(bool,X11)) ) | $spl6),
  inference(resolution,[],[f726,f4304])).
fof(f9096,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | p(s(bool,X9)) | s(bool,'T_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) | $spl6),
  inference(resolution,[],[f726,f4298])).
fof(f9095,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X7))) )),
  inference(resolution,[],[f726,f380])).
fof(f9115,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~sP7(X1,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X0)) )),
  inference(subsumption_resolution,[],[f9092,f391])).
fof(f9092,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~sP7(X1,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X0) | ~sP6('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X1,X0)) )),
  inference(resolution,[],[f726,f377])).
fof(f9080,plain,(
  ( ! [X2,X3] : (sP5(X2,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X3) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~p(s(bool,X3))) )),
  inference(resolution,[],[f714,f359])).
fof(f9079,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(resolution,[],[f714,f363])).
fof(f9067,plain,(
  ( ! [X2,X3] : (sP5(X2,X3,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | ~p(s(bool,X3)) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f713,f359])).
fof(f9066,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X1))) )),
  inference(resolution,[],[f713,f363])).
fof(f9065,plain,(
  ( ! [X6,X7] : (sP3('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X6,X7) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(subsumption_resolution,[],[f9056,f352])).
fof(f9056,plain,(
  ( ! [X6,X7] : (sP3('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X6,X7) | p(s(bool,X6)) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f700,f344])).
fof(f9064,plain,(
  ( ! [X4,X5] : (sP3('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X4,X5) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(subsumption_resolution,[],[f9055,f351])).
fof(f9055,plain,(
  ( ! [X4,X5] : (sP3('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X4,X5) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X5))) )),
  inference(resolution,[],[f700,f343])).
fof(f9054,plain,(
  ( ! [X2,X3] : (sP3('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X2,X3) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f700,f547])).
fof(f9053,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(resolution,[],[f700,f346])).
fof(f9052,plain,(
  ( ! [X6,X7] : (sP1('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X6,X7) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(subsumption_resolution,[],[f9043,f335])).
fof(f9043,plain,(
  ( ! [X6,X7] : (sP1('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X6,X7) | p(s(bool,X6)) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f684,f327])).
fof(f9051,plain,(
  ( ! [X4,X5] : (sP1('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X4,X5) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(subsumption_resolution,[],[f9042,f334])).
fof(f9042,plain,(
  ( ! [X4,X5] : (sP1('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X4,X5) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~p(s(bool,X5))) )),
  inference(resolution,[],[f684,f326])).
fof(f9041,plain,(
  ( ! [X2,X3] : (sP1('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X2,X3) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f684,f546])).
fof(f9040,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(resolution,[],[f684,f329])).
fof(f9031,plain,(
  ( ! [X8,X9] : (sP1(X8,X9,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | p(s(bool,X9)) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f683,f328])).
fof(f9027,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f683,f329])).
fof(f9017,plain,(
  ( ! [X13] : (p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X13))) )),
  inference(resolution,[],[f610,f1864])).
fof(f9016,plain,(
  ( ! [X12] : (p(s(bool,X12)) | s(bool,'F_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(resolution,[],[f610,f1855])).
fof(f9012,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X4)) | p(s(bool,X5))) )),
  inference(resolution,[],[f610,f388])).
fof(f9011,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) | $spl6),
  inference(resolution,[],[f610,f4655])).
fof(f9010,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'F_0') != s(bool,X1) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) | $spl6),
  inference(resolution,[],[f610,f4784])).
fof(f9008,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X9) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(subsumption_resolution,[],[f9007,f604])).
fof(f9007,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X9) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~sP8(X8,X9,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(subsumption_resolution,[],[f8999,f387])).
fof(f8999,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X9) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~sP8(X8,X9,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | ~sP7('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X9,X8)) )),
  inference(resolution,[],[f609,f378])).
fof(f8997,plain,(
  ( ! [X4,X5] : (p(s(bool,X4)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X5))) )),
  inference(resolution,[],[f609,f388])).
fof(f8996,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X2)) | s(bool,X3) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) | $spl6),
  inference(resolution,[],[f609,f4655])).
fof(f8995,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'F_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) | $spl6),
  inference(resolution,[],[f609,f4784])).
fof(f8985,plain,(
  ( ! [X13] : (p(s(bool,X13)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f608,f1864])).
fof(f8984,plain,(
  ( ! [X12] : (p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | s(bool,'F_0') != s(bool,X12)) )),
  inference(resolution,[],[f608,f1855])).
fof(f8994,plain,(
  ( ! [X10,X11] : (sP9('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X10,X11) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(subsumption_resolution,[],[f8993,f607])).
fof(f8993,plain,(
  ( ! [X10,X11] : (sP9('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X10,X11) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~sP7(X10,X11,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(subsumption_resolution,[],[f8983,f383])).
fof(f8983,plain,(
  ( ! [X10,X11] : (sP9('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X10,X11) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~sP8('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X11,X10) | ~sP7(X10,X11,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(resolution,[],[f608,f379])).
fof(f8980,plain,(
  ( ! [X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f608,f388])).
fof(f8979,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | s(bool,X2) != s(bool,X3)) ) | $spl6),
  inference(resolution,[],[f608,f4655])).
fof(f8978,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | s(bool,'F_0') != s(bool,X0) | s(bool,X0) != s(bool,X1)) ) | $spl6),
  inference(resolution,[],[f608,f4784])).
fof(f8969,plain,(
  ( ! [X12,X13] : (sP9('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X12,X13) | ~p(s(bool,X13)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) ) | $spl6),
  inference(resolution,[],[f607,f4398])).
fof(f8967,plain,(
  ( ! [X8,X9] : (sP9('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X8,X9) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | s(bool,'T_0') != s(bool,X9)) ) | $spl6),
  inference(resolution,[],[f607,f4378])).
fof(f8966,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f607,f384])).
fof(f8977,plain,(
  ( ! [X4,X5] : (sP9('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X4,X5) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~sP6(X5,X4,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(subsumption_resolution,[],[f8965,f383])).
fof(f8965,plain,(
  ( ! [X4,X5] : (sP9('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X4,X5) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~sP8('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X5,X4) | ~sP6(X5,X4,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(resolution,[],[f607,f379])).
fof(f8954,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'F_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | s(bool,'T_0') != s(bool,X19)) ) | $spl6),
  inference(resolution,[],[f604,f8227])).
fof(f8953,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | ~p(s(bool,X16)) | ~p(s(bool,X17))) ) | $spl6),
  inference(resolution,[],[f604,f8224])).
fof(f8952,plain,(
  ( ! [X14,X15] : (p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | sP9(X14,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X15) | ~sP6(X15,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X14)) ) | $spl6),
  inference(resolution,[],[f604,f4575])).
fof(f8950,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X11) | ~p(s(bool,X11)) | s(bool,'F_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) | $spl6),
  inference(resolution,[],[f604,f4304])).
fof(f8949,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X9) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | s(bool,'T_0') != s(bool,X9)) ) | $spl6),
  inference(resolution,[],[f604,f4298])).
fof(f8948,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f604,f380])).
fof(f8962,plain,(
  ( ! [X2,X3] : (sP9(X2,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X3) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~sP6(X3,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X2)) )),
  inference(subsumption_resolution,[],[f8946,f387])).
fof(f8946,plain,(
  ( ! [X2,X3] : (sP9(X2,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X3) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~sP7('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X3,X2) | ~sP6(X3,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X2)) )),
  inference(resolution,[],[f604,f378])).
fof(f8936,plain,(
  ( ! [X8,X9] : (sP5('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X8,X9) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f591,f548])).
fof(f8944,plain,(
  ( ! [X6,X7] : (sP5('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X6,X7) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(subsumption_resolution,[],[f8935,f367])).
fof(f8935,plain,(
  ( ! [X6,X7] : (sP5('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X6,X7) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~p(s(bool,X7))) )),
  inference(resolution,[],[f591,f361])).
fof(f8943,plain,(
  ( ! [X4,X5] : (sP5('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X4,X5) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(subsumption_resolution,[],[f8934,f368])).
fof(f8934,plain,(
  ( ! [X4,X5] : (sP5('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X4,X5) | ~p(s(bool,X4)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f591,f360])).
fof(f8932,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f591,f363])).
fof(f8923,plain,(
  ( ! [X8,X9] : (sP3(X8,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X9) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X9))) )),
  inference(resolution,[],[f575,f345])).
fof(f8919,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X1))) )),
  inference(resolution,[],[f575,f346])).
fof(f8910,plain,(
  ( ! [X8,X9] : (sP3(X8,X9,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | p(s(bool,X9)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f574,f345])).
fof(f8906,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f574,f346])).
fof(f8897,plain,(
  ( ! [X8,X9] : (sP1(X8,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X9) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~p(s(bool,X9))) )),
  inference(resolution,[],[f561,f328])).
fof(f8893,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~p(s(bool,X1))) )),
  inference(resolution,[],[f561,f329])).
fof(f8892,plain,(
  ( ! [X1] : (p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X1))) )),
  inference(subsumption_resolution,[],[f8884,f480])).
fof(f8884,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X1))) )),
  inference(resolution,[],[f622,f432])).
fof(f8891,plain,(
  ( ! [X0] : (p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f8883,f480])).
fof(f8883,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X0))) )),
  inference(resolution,[],[f622,f436])).
fof(f8882,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(subsumption_resolution,[],[f8874,f480])).
fof(f8874,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f621,f432])).
fof(f8881,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(subsumption_resolution,[],[f8873,f480])).
fof(f8873,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f621,f436])).
fof(f4962,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | $spl28),
  inference(cnf_transformation,[],[f4962_D])).
fof(f4962_D,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f8862,plain,(
  ( ! [X13] : (p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | p(s(bool,X13))) ) | $spl0),
  inference(resolution,[],[f8613,f1864])).
fof(f8613,plain,(
  ( ! [X112,X113] : (sP6('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X112,X113)) ) | $spl0),
  inference(backward_demodulation,[],[f668,f610])).
fof(f8861,plain,(
  ( ! [X12] : (p(s(bool,X12)) | s(bool,'F_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | $spl0),
  inference(resolution,[],[f8613,f1855])).
fof(f8857,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | p(s(bool,X4)) | p(s(bool,X5))) ) | $spl0),
  inference(resolution,[],[f8613,f388])).
fof(f8856,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f8613,f4655])).
fof(f8855,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'F_0') != s(bool,X1) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f8613,f4784])).
fof(f8853,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X9)) ) | $spl126),
  inference(cnf_transformation,[],[f8853_D])).
fof(f8853_D,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X9)) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f8852,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X9) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | $spl0),
  inference(subsumption_resolution,[],[f8851,f8607])).
fof(f8607,plain,(
  ( ! [X101,X100] : (sP8(X100,X101,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | $spl0),
  inference(backward_demodulation,[],[f668,f604])).
fof(f8851,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X9) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | ~sP8(X8,X9,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | $spl0),
  inference(subsumption_resolution,[],[f8844,f387])).
fof(f8844,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X9) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | ~sP8(X8,X9,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | ~sP7('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X9,X8)) ) | $spl0),
  inference(resolution,[],[f8612,f378])).
fof(f8612,plain,(
  ( ! [X111,X110] : (sP6(X110,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X111)) ) | $spl0),
  inference(backward_demodulation,[],[f668,f609])).
fof(f8842,plain,(
  ( ! [X4,X5] : (p(s(bool,X4)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | p(s(bool,X5))) ) | $spl0),
  inference(resolution,[],[f8612,f388])).
fof(f8841,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | p(s(bool,X2)) | s(bool,X3) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f8612,f4655])).
fof(f8840,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'F_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f8612,f4784])).
fof(f8830,plain,(
  ( ! [X13] : (p(s(bool,X13)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | $spl0),
  inference(resolution,[],[f8611,f1864])).
fof(f8611,plain,(
  ( ! [X109,X108] : (sP6(X108,X109,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | $spl0),
  inference(backward_demodulation,[],[f668,f608])).
fof(f8829,plain,(
  ( ! [X12] : (p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | s(bool,'F_0') != s(bool,X12)) ) | $spl0),
  inference(resolution,[],[f8611,f1855])).
fof(f8839,plain,(
  ( ! [X10,X11] : (sP9('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X10,X11)) ) | $spl122),
  inference(cnf_transformation,[],[f8839_D])).
fof(f8839_D,plain,(
  ( ! [X10,X11] : (sP9('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X10,X11)) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f8838,plain,(
  ( ! [X10,X11] : (sP9('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X10,X11) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | $spl0),
  inference(subsumption_resolution,[],[f8837,f8610])).
fof(f8610,plain,(
  ( ! [X107,X106] : (sP7(X106,X107,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | $spl0),
  inference(backward_demodulation,[],[f668,f607])).
fof(f8837,plain,(
  ( ! [X10,X11] : (sP9('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X10,X11) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | ~sP7(X10,X11,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | $spl0),
  inference(subsumption_resolution,[],[f8828,f383])).
fof(f8828,plain,(
  ( ! [X10,X11] : (sP9('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X10,X11) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | ~sP8('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X11,X10) | ~sP7(X10,X11,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | $spl0),
  inference(resolution,[],[f8611,f379])).
fof(f8825,plain,(
  ( ! [X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | $spl0),
  inference(resolution,[],[f8611,f388])).
fof(f8824,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | s(bool,X2) != s(bool,X3)) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f8611,f4655])).
fof(f8823,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | s(bool,'F_0') != s(bool,X0) | s(bool,X0) != s(bool,X1)) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f8611,f4784])).
fof(f8822,plain,(
  ( ! [X12,X13] : (sP9('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X12,X13) | ~p(s(bool,X13))) ) | $spl118),
  inference(cnf_transformation,[],[f8822_D])).
fof(f8822_D,plain,(
  ( ! [X12,X13] : (sP9('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X12,X13) | ~p(s(bool,X13))) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f8813,plain,(
  ( ! [X12,X13] : (sP9('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X12,X13) | ~p(s(bool,X13)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f8610,f4398])).
fof(f8821,plain,(
  ( ! [X8,X9] : (sP9('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X8,X9) | s(bool,'T_0') != s(bool,X9)) ) | $spl116),
  inference(cnf_transformation,[],[f8821_D])).
fof(f8821_D,plain,(
  ( ! [X8,X9] : (sP9('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X8,X9) | s(bool,'T_0') != s(bool,X9)) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f8811,plain,(
  ( ! [X8,X9] : (sP9('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X8,X9) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | s(bool,'T_0') != s(bool,X9)) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f8610,f4378])).
fof(f8810,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | $spl0),
  inference(resolution,[],[f8610,f384])).
fof(f8820,plain,(
  ( ! [X4,X5] : (sP9('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X4,X5) | ~sP6(X5,X4,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | $spl114),
  inference(cnf_transformation,[],[f8820_D])).
fof(f8820_D,plain,(
  ( ! [X4,X5] : (sP9('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X4,X5) | ~sP6(X5,X4,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f8819,plain,(
  ( ! [X4,X5] : (sP9('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X4,X5) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | ~sP6(X5,X4,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | $spl0),
  inference(subsumption_resolution,[],[f8809,f383])).
fof(f8809,plain,(
  ( ! [X4,X5] : (sP9('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X4,X5) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | ~sP8('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X5,X4) | ~sP6(X5,X4,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | $spl0),
  inference(resolution,[],[f8610,f379])).
fof(f8795,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'F_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | s(bool,'T_0') != s(bool,X19)) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f8607,f8227])).
fof(f8794,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | ~p(s(bool,X16)) | ~p(s(bool,X17))) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f8607,f8224])).
fof(f8802,plain,(
  ( ! [X2,X3] : (sP9(X2,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X3) | ~sP6(X3,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X2)) ) | $spl106),
  inference(cnf_transformation,[],[f8802_D])).
fof(f8802_D,plain,(
  ( ! [X2,X3] : (sP9(X2,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X3) | ~sP6(X3,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X2)) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f8793,plain,(
  ( ! [X14,X15] : (p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | sP9(X14,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X15) | ~sP6(X15,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X14)) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f8607,f4575])).
fof(f8791,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X11) | ~p(s(bool,X11)) | s(bool,'F_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f8607,f4304])).
fof(f8803,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X9) | s(bool,'T_0') != s(bool,X9)) ) | $spl108),
  inference(cnf_transformation,[],[f8803_D])).
fof(f8803_D,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X9) | s(bool,'T_0') != s(bool,X9)) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f8790,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X9) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | s(bool,'T_0') != s(bool,X9)) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f8607,f4298])).
fof(f8789,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | $spl0),
  inference(resolution,[],[f8607,f380])).
fof(f8801,plain,(
  ( ! [X2,X3] : (sP9(X2,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X3) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | ~sP6(X3,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X2)) ) | $spl0),
  inference(subsumption_resolution,[],[f8787,f387])).
fof(f8787,plain,(
  ( ! [X2,X3] : (sP9(X2,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X3) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | ~sP7('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X3,X2) | ~sP6(X3,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X2)) ) | $spl0),
  inference(resolution,[],[f8607,f378])).
fof(f8784,plain,(
  ( ! [X4,X5] : (sP5('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X4,X5)) ) | $spl104),
  inference(cnf_transformation,[],[f8784_D])).
fof(f8784_D,plain,(
  ( ! [X4,X5] : (sP5('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X4,X5)) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f8778,plain,(
  ( ! [X8,X9] : (sP5('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X8,X9) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | $spl0),
  inference(resolution,[],[f8597,f548])).
fof(f8597,plain,(
  ( ! [X74,X75] : (sP4(X74,X75,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | $spl0),
  inference(backward_demodulation,[],[f668,f591])).
fof(f8785,plain,(
  ( ! [X6,X7] : (sP5('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X6,X7) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | $spl0),
  inference(subsumption_resolution,[],[f8777,f367])).
fof(f8777,plain,(
  ( ! [X6,X7] : (sP5('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X6,X7) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | ~p(s(bool,X7))) ) | $spl0),
  inference(resolution,[],[f8597,f361])).
fof(f8783,plain,(
  ( ! [X4,X5] : (sP5('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X4,X5) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | $spl0),
  inference(subsumption_resolution,[],[f8776,f368])).
fof(f8776,plain,(
  ( ! [X4,X5] : (sP5('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X4,X5) | ~p(s(bool,X4)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | $spl0),
  inference(resolution,[],[f8597,f360])).
fof(f8774,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | $spl0),
  inference(resolution,[],[f8597,f363])).
fof(f8773,plain,(
  ( ! [X8,X9] : (sP3(X8,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X9) | p(s(bool,X9))) ) | $spl102),
  inference(cnf_transformation,[],[f8773_D])).
fof(f8773_D,plain,(
  ( ! [X8,X9] : (sP3(X8,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X9) | p(s(bool,X9))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f8766,plain,(
  ( ! [X8,X9] : (sP3(X8,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X9) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | p(s(bool,X9))) ) | $spl0),
  inference(resolution,[],[f8587,f345])).
fof(f8587,plain,(
  ( ! [X43,X42] : (sP2(X42,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X43)) ) | $spl0),
  inference(backward_demodulation,[],[f668,f575])).
fof(f8762,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | p(s(bool,X1))) ) | $spl0),
  inference(resolution,[],[f8587,f346])).
fof(f8761,plain,(
  ( ! [X8,X9] : (sP3(X8,X9,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | p(s(bool,X9))) ) | $spl100),
  inference(cnf_transformation,[],[f8761_D])).
fof(f8761_D,plain,(
  ( ! [X8,X9] : (sP3(X8,X9,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | p(s(bool,X9))) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f8754,plain,(
  ( ! [X8,X9] : (sP3(X8,X9,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | p(s(bool,X9)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | $spl0),
  inference(resolution,[],[f8586,f345])).
fof(f8586,plain,(
  ( ! [X41,X40] : (sP2(X40,X41,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | $spl0),
  inference(backward_demodulation,[],[f668,f574])).
fof(f8750,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | $spl0),
  inference(resolution,[],[f8586,f346])).
fof(f8749,plain,(
  ( ! [X8,X9] : (sP1(X8,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X9) | ~p(s(bool,X9))) ) | $spl98),
  inference(cnf_transformation,[],[f8749_D])).
fof(f8749_D,plain,(
  ( ! [X8,X9] : (sP1(X8,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X9) | ~p(s(bool,X9))) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f8742,plain,(
  ( ! [X8,X9] : (sP1(X8,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X9) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | ~p(s(bool,X9))) ) | $spl0),
  inference(resolution,[],[f8576,f328])).
fof(f8576,plain,(
  ( ! [X14,X15] : (sP0(X14,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X15)) ) | $spl0),
  inference(backward_demodulation,[],[f668,f561])).
fof(f8738,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | ~p(s(bool,X1))) ) | $spl0),
  inference(resolution,[],[f8576,f329])).
fof(f8737,plain,(
  ( ! [X1] : (p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | p(s(bool,X1))) ) | $spl0),
  inference(subsumption_resolution,[],[f8731,f480])).
fof(f8731,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | p(s(bool,X1))) ) | $spl0),
  inference(resolution,[],[f8621,f432])).
fof(f8621,plain,(
  ( ! [X127] : (sP11('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X127)) ) | $spl0),
  inference(backward_demodulation,[],[f668,f622])).
fof(f8736,plain,(
  ( ! [X0] : (p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | p(s(bool,X0))) ) | $spl0),
  inference(subsumption_resolution,[],[f8730,f480])).
fof(f8730,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | p(s(bool,X0))) ) | $spl0),
  inference(resolution,[],[f8621,f436])).
fof(f8729,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | $spl0),
  inference(subsumption_resolution,[],[f8723,f480])).
fof(f8723,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | $spl0),
  inference(resolution,[],[f8620,f432])).
fof(f8620,plain,(
  ( ! [X126] : (sP11(X126,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | $spl0),
  inference(backward_demodulation,[],[f668,f621])).
fof(f8728,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | $spl0),
  inference(subsumption_resolution,[],[f8722,f480])).
fof(f8722,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | $spl0),
  inference(resolution,[],[f8620,f436])).
fof(f5287,plain,(
  p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) | s(sK12,sK14) = s(sK12,sK15) | $spl30),
  inference(resolution,[],[f4976,f532])).
fof(f5862,plain,(
  p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) | s(sK12,sK14) = s(sK12,sK15)),
  inference(subsumption_resolution,[],[f5604,f549])).
fof(f5604,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) | s(sK12,sK14) = s(sK12,sK15)),
  inference(superposition,[],[f532,f636])).
fof(f8562,plain,(
  ( ! [X41,X42] : (p(s(bool,X41)) | s(bool,'F_0') != s(bool,X42) | s(bool,X42) != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) ) | ($spl2 | $spl6)),
  inference(resolution,[],[f4784,f3372])).
fof(f8558,plain,(
  ( ! [X33,X32] : (p(s(bool,X32)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))) | s(bool,X33) != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) ) | ($spl2 | $spl6)),
  inference(resolution,[],[f4784,f3371])).
fof(f8485,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))) | s(bool,'T_0') != s(bool,X15)) ) | ($spl2 | $spl6)),
  inference(resolution,[],[f8227,f3366])).
fof(f8238,plain,(
  ( ! [X14,X15] : (s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))) | ~p(s(bool,X14)) | ~p(s(bool,X15))) ) | ($spl2 | $spl6)),
  inference(resolution,[],[f8224,f3366])).
fof(f8059,plain,(
  ( ! [X10] : (p(s(bool,X10)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) ) | $spl2),
  inference(resolution,[],[f3372,f1855])).
fof(f8054,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) ) | ($spl2 | $spl6)),
  inference(resolution,[],[f3372,f4655])).
fof(f7975,plain,(
  ( ! [X10,X11] : (sP9(X10,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X11) | ~p(s(bool,X11)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) ) | ($spl2 | $spl6)),
  inference(resolution,[],[f3366,f4304])).
fof(f5276,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X11) | ~p(s(bool,X11)) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) ) | $spl6),
  inference(resolution,[],[f730,f4396])).
fof(f5274,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~p(s(bool,X6)) | p(s(bool,X7))) )),
  inference(resolution,[],[f730,f384])).
fof(f5258,plain,(
  ( ! [X12,X13] : (sP9(X12,X13,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X12))) ) | $spl6),
  inference(resolution,[],[f729,f4398])).
fof(f5257,plain,(
  ( ! [X10,X11] : (sP9(X10,X11,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~p(s(bool,X11))) ) | $spl6),
  inference(resolution,[],[f729,f4396])).
fof(f5256,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | p(s(bool,X8)) | s(bool,'T_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) | $spl6),
  inference(resolution,[],[f729,f4378])).
fof(f5255,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X7))) )),
  inference(resolution,[],[f729,f384])).
fof(f5266,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(subsumption_resolution,[],[f5265,f726])).
fof(f5265,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~sP8(X0,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X1)) )),
  inference(subsumption_resolution,[],[f5252,f391])).
fof(f5252,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~sP8(X0,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X1) | ~sP6('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X1,X0)) )),
  inference(resolution,[],[f729,f377])).
fof(f5241,plain,(
  ( ! [X19] : (~p(s(bool,X19)) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f727,f1809])).
fof(f5239,plain,(
  ( ! [X17] : (s(bool,'T_0') != s(bool,X17) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f727,f1924])).
fof(f5233,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~p(s(bool,X6)) | p(s(bool,X7))) )),
  inference(resolution,[],[f727,f380])).
fof(f5232,plain,(
  ( ! [X4,X5] : (sP9('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X4,X5) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~sP7(X4,X5,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | ~sP6(X5,X4,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(resolution,[],[f727,f379])).
fof(f5217,plain,(
  ( ! [X19] : (~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~p(s(bool,X19))) )),
  inference(resolution,[],[f726,f1809])).
fof(f5215,plain,(
  ( ! [X17] : (s(bool,'T_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | ~p(s(bool,X17))) )),
  inference(resolution,[],[f726,f1924])).
fof(f5214,plain,(
  ( ! [X16] : (sP9(X16,'F_0','const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | s(bool,'T_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(resolution,[],[f726,f1794])).
fof(f5212,plain,(
  ( ! [X12,X13] : (~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | sP9(X12,X13,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | ~sP7(X13,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X12)) ) | $spl6),
  inference(resolution,[],[f726,f4377])).
fof(f5211,plain,(
  ( ! [X10,X11] : (sP9(X10,X11,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | s(bool,'F_0') != s(bool,X11)) ) | $spl6),
  inference(resolution,[],[f726,f4304])).
fof(f5210,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | p(s(bool,X9)) | s(bool,'T_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) | $spl6),
  inference(resolution,[],[f726,f4298])).
fof(f5209,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X7))) )),
  inference(resolution,[],[f726,f380])).
fof(f5225,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~sP7(X1,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X0)) )),
  inference(subsumption_resolution,[],[f5206,f391])).
fof(f5206,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~sP7(X1,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X0) | ~sP6('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X1,X0)) )),
  inference(resolution,[],[f726,f377])).
fof(f5193,plain,(
  ( ! [X2,X3] : (sP5(X2,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X3) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~p(s(bool,X3))) )),
  inference(resolution,[],[f714,f359])).
fof(f5192,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(resolution,[],[f714,f363])).
fof(f5179,plain,(
  ( ! [X2,X3] : (sP5(X2,X3,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | ~p(s(bool,X3)) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f713,f359])).
fof(f5178,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X1))) )),
  inference(resolution,[],[f713,f363])).
fof(f5177,plain,(
  ( ! [X6,X7] : (sP3('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X6,X7) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(subsumption_resolution,[],[f5167,f352])).
fof(f5167,plain,(
  ( ! [X6,X7] : (sP3('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X6,X7) | p(s(bool,X6)) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f700,f344])).
fof(f5176,plain,(
  ( ! [X4,X5] : (sP3('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X4,X5) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(subsumption_resolution,[],[f5166,f351])).
fof(f5166,plain,(
  ( ! [X4,X5] : (sP3('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X4,X5) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X5))) )),
  inference(resolution,[],[f700,f343])).
fof(f5165,plain,(
  ( ! [X2,X3] : (sP3('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X2,X3) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f700,f547])).
fof(f5164,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(resolution,[],[f700,f346])).
fof(f5163,plain,(
  ( ! [X6,X7] : (sP1('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X6,X7) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(subsumption_resolution,[],[f5153,f335])).
fof(f5153,plain,(
  ( ! [X6,X7] : (sP1('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X6,X7) | p(s(bool,X6)) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f684,f327])).
fof(f5162,plain,(
  ( ! [X4,X5] : (sP1('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X4,X5) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(subsumption_resolution,[],[f5152,f334])).
fof(f5152,plain,(
  ( ! [X4,X5] : (sP1('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X4,X5) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~p(s(bool,X5))) )),
  inference(resolution,[],[f684,f326])).
fof(f5151,plain,(
  ( ! [X2,X3] : (sP1('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X2,X3) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f684,f546])).
fof(f5150,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(resolution,[],[f684,f329])).
fof(f5138,plain,(
  ( ! [X8,X9] : (sP1(X8,X9,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | p(s(bool,X9)) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f683,f328])).
fof(f5134,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f683,f329])).
fof(f5124,plain,(
  ( ! [X11] : (p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X11))) )),
  inference(resolution,[],[f610,f1864])).
fof(f5123,plain,(
  ( ! [X10] : (p(s(bool,X10)) | s(bool,'F_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(resolution,[],[f610,f1855])).
fof(f5119,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X2)) | p(s(bool,X3))) )),
  inference(resolution,[],[f610,f388])).
fof(f5118,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) | $spl6),
  inference(resolution,[],[f610,f4655])).
fof(f5115,plain,(
  ( ! [X6,X7] : (sP9(X6,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X7) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(subsumption_resolution,[],[f5114,f604])).
fof(f5114,plain,(
  ( ! [X6,X7] : (sP9(X6,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X7) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~sP8(X6,X7,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(subsumption_resolution,[],[f5105,f387])).
fof(f5105,plain,(
  ( ! [X6,X7] : (sP9(X6,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X7) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~sP8(X6,X7,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | ~sP7('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X7,X6)) )),
  inference(resolution,[],[f609,f378])).
fof(f5103,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X3))) )),
  inference(resolution,[],[f609,f388])).
fof(f5102,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X0)) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) | $spl6),
  inference(resolution,[],[f609,f4655])).
fof(f5089,plain,(
  ( ! [X11] : (p(s(bool,X11)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f608,f1864])).
fof(f5088,plain,(
  ( ! [X10] : (p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | s(bool,'F_0') != s(bool,X10)) )),
  inference(resolution,[],[f608,f1855])).
fof(f5099,plain,(
  ( ! [X8,X9] : (sP9('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X8,X9) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(subsumption_resolution,[],[f5098,f607])).
fof(f5098,plain,(
  ( ! [X8,X9] : (sP9('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X8,X9) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~sP7(X8,X9,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(subsumption_resolution,[],[f5087,f383])).
fof(f5087,plain,(
  ( ! [X8,X9] : (sP9('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X8,X9) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~sP8('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X9,X8) | ~sP7(X8,X9,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(resolution,[],[f608,f379])).
fof(f5084,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | p(s(bool,X3)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f608,f388])).
fof(f5083,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | s(bool,X0) != s(bool,X1)) ) | $spl6),
  inference(resolution,[],[f608,f4655])).
fof(f5071,plain,(
  ( ! [X12,X13] : (sP9('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X12,X13) | ~p(s(bool,X13)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) ) | $spl6),
  inference(resolution,[],[f607,f4398])).
fof(f5069,plain,(
  ( ! [X8,X9] : (sP9('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X8,X9) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | s(bool,'T_0') != s(bool,X9)) ) | $spl6),
  inference(resolution,[],[f607,f4378])).
fof(f5068,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f607,f384])).
fof(f5079,plain,(
  ( ! [X4,X5] : (sP9('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X4,X5) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~sP6(X5,X4,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(subsumption_resolution,[],[f5067,f383])).
fof(f5067,plain,(
  ( ! [X4,X5] : (sP9('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X4,X5) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~sP8('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X5,X4) | ~sP6(X5,X4,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(resolution,[],[f607,f379])).
fof(f5051,plain,(
  ( ! [X14,X15] : (p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | sP9(X14,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X15) | ~sP6(X15,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X14)) ) | $spl6),
  inference(resolution,[],[f604,f4575])).
fof(f5049,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X11) | ~p(s(bool,X11)) | s(bool,'F_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) | $spl6),
  inference(resolution,[],[f604,f4304])).
fof(f5048,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X9) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | s(bool,'T_0') != s(bool,X9)) ) | $spl6),
  inference(resolution,[],[f604,f4298])).
fof(f5047,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f604,f380])).
fof(f5059,plain,(
  ( ! [X2,X3] : (sP9(X2,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X3) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~sP6(X3,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X2)) )),
  inference(subsumption_resolution,[],[f5045,f387])).
fof(f5045,plain,(
  ( ! [X2,X3] : (sP9(X2,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X3) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~sP7('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X3,X2) | ~sP6(X3,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X2)) )),
  inference(resolution,[],[f604,f378])).
fof(f5034,plain,(
  ( ! [X8,X9] : (sP5('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X8,X9) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f591,f548])).
fof(f5043,plain,(
  ( ! [X6,X7] : (sP5('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X6,X7) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(subsumption_resolution,[],[f5033,f367])).
fof(f5033,plain,(
  ( ! [X6,X7] : (sP5('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X6,X7) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~p(s(bool,X7))) )),
  inference(resolution,[],[f591,f361])).
fof(f5041,plain,(
  ( ! [X4,X5] : (sP5('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X4,X5) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(subsumption_resolution,[],[f5032,f368])).
fof(f5032,plain,(
  ( ! [X4,X5] : (sP5('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X4,X5) | ~p(s(bool,X4)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f591,f360])).
fof(f5030,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f591,f363])).
fof(f5020,plain,(
  ( ! [X8,X9] : (sP3(X8,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X9) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X9))) )),
  inference(resolution,[],[f575,f345])).
fof(f5016,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X1))) )),
  inference(resolution,[],[f575,f346])).
fof(f5006,plain,(
  ( ! [X8,X9] : (sP3(X8,X9,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | p(s(bool,X9)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f574,f345])).
fof(f5002,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f574,f346])).
fof(f4992,plain,(
  ( ! [X8,X9] : (sP1(X8,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X9) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~p(s(bool,X9))) )),
  inference(resolution,[],[f561,f328])).
fof(f4988,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~p(s(bool,X1))) )),
  inference(resolution,[],[f561,f329])).
fof(f4987,plain,(
  ( ! [X1] : (p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X1))) )),
  inference(subsumption_resolution,[],[f4979,f480])).
fof(f4979,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X1))) )),
  inference(resolution,[],[f622,f432])).
fof(f4986,plain,(
  ( ! [X0] : (p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f4978,f480])).
fof(f4978,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X0))) )),
  inference(resolution,[],[f622,f436])).
fof(f4977,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(subsumption_resolution,[],[f4967,f480])).
fof(f4967,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f621,f432])).
fof(f4974,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(subsumption_resolution,[],[f4966,f480])).
fof(f4966,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f621,f436])).
fof(f4963,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | $spl0),
  inference(subsumption_resolution,[],[f4955,f480])).
fof(f4955,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | $spl0),
  inference(resolution,[],[f4855,f432])).
fof(f4855,plain,(
  ( ! [X126] : (sP11(X126,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | $spl0),
  inference(backward_demodulation,[],[f668,f621])).
fof(f4960,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | $spl0),
  inference(subsumption_resolution,[],[f4954,f480])).
fof(f4954,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | $spl0),
  inference(resolution,[],[f4855,f436])).
fof(f4807,plain,(
  ( ! [X39,X38] : (p(s(bool,X38)) | p(s(bool,X39)) | s(bool,X38) != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) ) | ($spl2 | $spl6)),
  inference(resolution,[],[f4655,f3372])).
fof(f4562,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') != s(bool,X4) | s(bool,X3) != s(bool,X4) | ~p(s(bool,X3))) ) | $spl24),
  inference(cnf_transformation,[],[f4562_D])).
fof(f4562_D,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') != s(bool,X4) | s(bool,X3) != s(bool,X4) | ~p(s(bool,X3))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f4561,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | s(bool,'F_0') != s(bool,X4) | p(s(bool,X5)) | s(bool,X3) != s(bool,X4)) ) | $spl6),
  inference(duplicate_literal_removal,[],[f4542])).
fof(f4542,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | s(bool,'F_0') != s(bool,X4) | p(s(bool,X5)) | p(s(bool,X5)) | s(bool,X3) != s(bool,X4)) ) | $spl6),
  inference(resolution,[],[f4335,f395])).
fof(f4338,plain,(
  ( ! [X14,X15] : (sP9(X14,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X15) | ~p(s(bool,X15)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) ) | ($spl2 | $spl6)),
  inference(resolution,[],[f4304,f3366])).
fof(f4059,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X7) | s(bool,'T_0') != s(bool,X6) | s(bool,X6) != s(bool,X7)) ) | $spl20),
  inference(cnf_transformation,[],[f4059_D])).
fof(f4059_D,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X7) | s(bool,'T_0') != s(bool,X6) | s(bool,X6) != s(bool,X7)) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f4058,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X6) | s(bool,X6) != s(bool,X7) | p(s(bool,X8)) | s(bool,'F_0') != s(bool,X7)) ) | $spl6),
  inference(duplicate_literal_removal,[],[f4045])).
fof(f4045,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X6) | s(bool,X6) != s(bool,X7) | p(s(bool,X8)) | p(s(bool,X8)) | s(bool,'F_0') != s(bool,X7)) ) | $spl6),
  inference(resolution,[],[f1353,f395])).
fof(f1353,plain,(
  ( ! [X10,X8,X9] : (sP9(X10,X8,'F_0') | s(bool,'T_0') != s(bool,X9) | s(bool,X8) != s(bool,X9) | p(s(bool,X10))) ) | $spl6),
  inference(superposition,[],[f1160,f392])).
fof(f4027,plain,(
  ( ! [X8] : (p(s(bool,X8)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) ) | $spl2),
  inference(resolution,[],[f3372,f1855])).
fof(f2027,plain,(
  ( ! [X24,X22] : (s(bool,X22) != s(bool,X24) | ~p(s(bool,X22)) | p(s(bool,X24))) ) | $spl14),
  inference(cnf_transformation,[],[f2027_D])).
fof(f2027_D,plain,(
  ( ! [X24,X22] : (s(bool,X22) != s(bool,X24) | ~p(s(bool,X22)) | p(s(bool,X24))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f2029,plain,(
  ( ! [X26,X27,X25] : (s(bool,'T_0') != s(bool,X26) | s(bool,X25) != s(bool,X27) | p(s(bool,X27)) | ~p(s(bool,X25))) ) | $spl6),
  inference(subsumption_resolution,[],[f1993,f510])).
fof(f1993,plain,(
  ( ! [X26,X27,X25] : (s(bool,'T_0') != s(bool,X26) | s(bool,X25) != s(bool,X27) | p(s(bool,X27)) | ~p(s(bool,X26)) | ~p(s(bool,X25))) ) | $spl6),
  inference(superposition,[],[f1351,f452])).
fof(f452,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f177])).
fof(f2026,plain,(
  ( ! [X24,X23,X22] : (s(bool,'T_0') != s(bool,X23) | s(bool,X22) != s(bool,X24) | p(s(bool,X24)) | ~p(s(bool,X22))) ) | $spl6),
  inference(subsumption_resolution,[],[f1992,f510])).
fof(f1992,plain,(
  ( ! [X24,X23,X22] : (s(bool,'T_0') != s(bool,X23) | s(bool,X22) != s(bool,X24) | p(s(bool,X24)) | ~p(s(bool,X22)) | ~p(s(bool,X23))) ) | $spl6),
  inference(superposition,[],[f1351,f452])).
fof(f1938,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | s(bool,'T_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(resolution,[],[f1794,f726])).
fof(f1929,plain,(
  ( ! [X8] : (s(bool,'T_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | ~p(s(bool,X8))) )),
  inference(resolution,[],[f1924,f726])).
fof(f1671,plain,(
  p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) | s(sK12,sK14) = s(sK12,sK15)),
  inference(subsumption_resolution,[],[f1427,f549])).
fof(f1427,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) | s(sK12,sK14) = s(sK12,sK15)),
  inference(superposition,[],[f532,f636])).
fof(f1337,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f1265,f1159])).
fof(f1265,plain,(
  ~p(s(bool,'T_0')) | $spl4),
  inference(resolution,[],[f1159,f529])).
fof(f529,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f307])).
fof(f1336,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f1264,f1159])).
fof(f1264,plain,(
  ~p(s(bool,'T_0')) | $spl4),
  inference(resolution,[],[f1159,f527])).
fof(f527,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f306])).
fof(f1335,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f1261,f1159])).
fof(f1261,plain,(
  ( ! [X226] : (~p(s(bool,X226))) ) | $spl4),
  inference(resolution,[],[f1159,f521])).
fof(f521,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f188])).
fof(f188,plain,(
  ! [X0] : ((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f138])).
fof(f138,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) => ~p(s(bool,X0)))),
  inference(flattening,[],[f137])).
fof(f137,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) => ~p(s(bool,X0)))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X10] : ((p(s(bool,X10)) => p(s(bool,'F_0'))) => ~p(s(bool,X10)))),
  file('thm.relation.RC_RTC','thm.bool.IMP_F')).
fof(f1334,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f1255,f1159])).
fof(f1255,plain,(
  ( ! [X220] : (~p(s(bool,X220))) ) | $spl4),
  inference(resolution,[],[f1159,f462])).
fof(f1263,plain,(
  $false | $spl4),
  inference(resolution,[],[f1159,f525])).
fof(f1268,plain,(
  $false | $spl4),
  inference(resolution,[],[f1159,f319])).
fof(f1160,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | s(bool,'T_0') != s(bool,X2)) ) | $spl6),
  inference(cnf_transformation,[],[f1160_D])).
fof(f1160_D,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | s(bool,'T_0') != s(bool,X2)) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f1158,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) )),
  inference(duplicate_literal_removal,[],[f1153])).
fof(f1153,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) )),
  inference(resolution,[],[f1132,f395])).
fof(f1132,plain,(
  ( ! [X4,X5] : (sP9(X5,X4,'F_0') | s(bool,'T_0') != s(bool,X4) | p(s(bool,X5))) )),
  inference(superposition,[],[f1124,f392])).
fof(f553,plain,(
  p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) | s(sK12,sK14) = s(sK12,sK15)),
  inference(resolution,[],[f318,f532])).
% SZS output end Proof for thm.relation
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 796
% Passive clauses: 6223
% Generated clauses: 21508
% Final active clauses: 438
% Final passive clauses: 3720
% Input formulas: 65
% Initial clauses: 228
% 
% Selected by SInE selection: 62
% SInE iterations: 1
% 
% Duplicate literals: 1080
% Trivial inequalities: 187
% Fw subsumption resolutions: 2805
% Fw demodulations: 2098
% Bw demodulations: 267
% 
% Simple tautologies: 2323
% Equational tautologies: 43
% Forward subsumptions: 6807
% 
% Binary resolution: 7645
% Forward superposition: 4585
% Backward superposition: 2476
% Equality resolution: 17
% 
% Split clauses: 472
% Split components: 1054
% SAT solver clauses: 534
% SAT solver unit clauses: 20
% SAT solver binary clauses: 161
% 
% Sat splits: 472
% Sat splitting refutations: 21
% 
% Memory used [KB]: 5373
% Time elapsed: 0.462 s
% ------------------------------
% ------------------------------
% Success in time 12.146 s
