Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.relation
% (80014)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 91
% Passive clauses: 2789
% Generated clauses: 15517
% Final active clauses: 78
% Final passive clauses: 2581
% Input formulas: 6
% Initial clauses: 17
% 
% Duplicate literals: 951
% Trivial inequalities: 180
% Fw subsumption resolutions: 1787
% Fw demodulations: 1207
% Bw demodulations: 9
% 
% Simple tautologies: 289
% Equational tautologies: 95
% Forward subsumptions: 6094
% 
% Binary resolution: 84
% Forward superposition: 6538
% Backward superposition: 3968
% Self superposition: 349
% Equality factoring: 264
% Equality resolution: 9
% 
% Split clauses: 785
% Split components: 1596
% SAT solver clauses: 1897
% SAT solver unit clauses: 305
% SAT solver binary clauses: 519
% 
% Sat splits: 785
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 1893
% TWLsolver calls for satisfiability: 77
% 
% Memory used [KB]: 8059
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 294 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.relation
% (80030)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 113
% Passive clauses: 368
% Generated clauses: 10277
% Final active clauses: 86
% Final passive clauses: 137
% Input formulas: 6
% Initial clauses: 17
% 
% Duplicate literals: 1277
% Trivial inequalities: 57
% Fw subsumption resolutions: 462
% Fw demodulations: 756
% Bw demodulations: 135
% Condensations: 1464
% 
% Simple tautologies: 102
% Equational tautologies: 42
% Forward subsumptions: 3792
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 65
% Factoring: 1
% Forward superposition: 2958
% Backward superposition: 2449
% Self superposition: 133
% Equality factoring: 122
% 
% Split clauses: 272
% Split components: 546
% SAT solver clauses: 1010
% SAT solver unit clauses: 19
% SAT solver binary clauses: 42
% 
% Sat splits: 272
% Sat splitting refutations: 7
% 
% TWLsolver clauses: 1010
% TWLsolver calls for satisfiability: 210
% 
% Memory used [KB]: 8315
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 287 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.relation
% (80150)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 452
% Passive clauses: 5965
% Generated clauses: 116455
% Final active clauses: 76
% Final passive clauses: 855
% Input formulas: 6
% Initial clauses: 17
% Discarded non-redundant clauses: 23
% 
% Duplicate literals: 7080
% Trivial inequalities: 935
% Fw subsumption resolutions: 9268
% Fw demodulations: 12088
% Bw demodulations: 3569
% Condensations: 13060
% Global subsumptions: 718
% 
% Simple tautologies: 1048
% Equational tautologies: 1429
% Forward subsumptions: 63831
% Bw demodulations to eq. taut.: 269
% 
% Unit resulting resolution: 4009
% Factoring: 3
% Forward superposition: 31808
% Backward superposition: 19744
% Self superposition: 671
% Equality factoring: 435
% Equality resolution: 48
% 
% Split clauses: 380
% Split components: 792
% SAT solver clauses: 8174
% SAT solver unit clauses: 16
% SAT solver binary clauses: 296
% 
% Sat splits: 380
% Sat splitting refutations: 729
% 
% Memory used [KB]: 56161
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 183 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.relation
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.relation
% SZS output start Proof for thm.relation
fof(f5924,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f198,f210,f221,f249,f277,f327,f328,f202,f333,f334,f353,f361,f371,f396,f219,f218,f424,f229,f480,f605,f616,f627,f655,f683,f743,f746,f608,f749,f750,f751,f768,f775,f784,f813,f625,f624,f842,f635,f896,f1122,f1231,f1252,f1267,f1289,f1256,f1318,f1269,f1254,f1646,f1661,f1648,f1760,f1647,f1761,f1791,f1805,f3081,f3093,f3113,f3125,f3137,f3149,f3159,f3171,f3193,f3205,f3161,f3215,f3227,f3217,f4708,f4709,f4720,f4722,f2061,f2062,f2063,f2042,f2043,f1418,f1419,f1236,f4725,f4727,f1234,f4728,f4734,f4730,f4724,f4736,f4731,f4738,f4732,f1232,f4737,f4832,f4833,f4877,f4910,f4911,f4863,f4912,f4718,f4913,f4716,f4856,f4917,f4974,f5180,f5212,f5182,f5224,f4858,f5248,f4861,f5324,f1793,f5341,f4860,f1807,f636,f5614,f5616,f606,f752,f5618,f5627,f3139,f5728,f5704,f330,f5608,f335,f200,f5739,f5743,f5747,f3083,f5761,f5741,f5764,f5745,f5768,f5749,f230,f5796,f5840,f4880,f4879,f5914])).
fof(f5914,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f5913,f43])).
fof(f43,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f12])).
fof(f12,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.relation.TC_STRONG_INDUCT',aHL_FALSITY)).
fof(f5913,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(forward_demodulation,[],[f5912,f62])).
fof(f62,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK4))),s(sK1,sK5)))),
  inference(subsumption_resolution,[],[f60,f44])).
fof(f44,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.relation.TC_STRONG_INDUCT',aHL_TRUTH)).
fof(f60,plain,(
  ~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK4))),s(sK1,sK5)))),
  inference(superposition,[],[f33,f45])).
fof(f45,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.relation.TC_STRONG_INDUCT',aHL_BOOL_CASES)).
fof(f33,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK4))),s(sK1,sK5))))),
  inference(cnf_transformation,[],[f23])).
fof(f23,plain,(
  ! [X8,X9] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X8))),s(sK1,X9)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X8))),s(sK1,X9))))) & ! [X5,X6,X7] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X5))),s(sK1,X6)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X6))),s(sK1,X7)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X5),s(sK1,X6)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X6),s(sK1,X7)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X5))),s(sK1,X7))))) & (p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK4),s(sK1,sK5)))) & ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK4))),s(sK1,sK5)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK1,sK2,sK3,sK4,sK5])],[f22])).
fof(f22,plain,(
  ? [X0,X1,X2] : (! [X8,X9] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X8))),s(X0,X9)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X8))),s(X0,X9))))) & ! [X5,X6,X7] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X5))),s(X0,X6)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X6))),s(X0,X7)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X5),s(X0,X6)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X6),s(X0,X7)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X5))),s(X0,X7))))) & ? [X3,X4] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4))))))),
  inference(rectify,[],[f16])).
fof(f16,plain,(
  ? [X0,X1,X2] : (! [X6,X7] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X6))),s(X0,X7)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X6))),s(X0,X7))))) & ! [X3,X4,X5] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X4))),s(X0,X5)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X5)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X5))))) & ? [X8,X9] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X8),s(X0,X9)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X8))),s(X0,X9))))))),
  inference(flattening,[],[f15])).
fof(f15,plain,(
  ? [X0,X1,X2] : ((! [X6,X7] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X6))),s(X0,X7)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X6))),s(X0,X7))))) & ! [X3,X4,X5] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X4))),s(X0,X5)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X5))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X5)))))) & ? [X8,X9] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X8),s(X0,X9)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X8))),s(X0,X9))))))),
  inference(ennf_transformation,[],[f9])).
fof(f9,plain,(
  ~! [X0,X1,X2] : ((! [X6,X7] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X6))),s(X0,X7)))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X6))),s(X0,X7))))) & ! [X3,X4,X5] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X4))),s(X0,X5)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X5))))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X5)))))) => ! [X8,X9] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X8),s(X0,X9)))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X8))),s(X0,X9))))))),
  inference(flattening,[],[f8])).
fof(f8,plain,(
  ~! [X0] : ! [X1,X2] : ((! [X6,X7] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X6))),s(X0,X7)))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X6))),s(X0,X7))))) & ! [X3,X4,X5] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X4))),s(X0,X5)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X5))))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X5)))))) => ! [X8,X9] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X8),s(X0,X9)))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X8))),s(X0,X9))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : ((! [X9,X10] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X9))),s(X6,X10)))) => p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,X9))),s(X6,X10))))) & ! [X9,X10,X11] : ((p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,X9))),s(X6,X10)))) & p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,X10))),s(X6,X11)))) & p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X9),s(X6,X10)))) & p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X10),s(X6,X11))))) => p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,X9))),s(X6,X11)))))) => ! [X12,X13] : (p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X12),s(X6,X13)))) => p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,X12))),s(X6,X13))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : ((! [X9,X10] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X9))),s(X6,X10)))) => p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,X9))),s(X6,X10))))) & ! [X9,X10,X11] : ((p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,X9))),s(X6,X10)))) & p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,X10))),s(X6,X11)))) & p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X9),s(X6,X10)))) & p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X10),s(X6,X11))))) => p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,X9))),s(X6,X11)))))) => ! [X12,X13] : (p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X12),s(X6,X13)))) => p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,X12))),s(X6,X13))))))),
  file('thm.relation.TC_STRONG_INDUCT','thm.relation.TC_STRONG_INDUCT')).
fof(f5912,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK4))),s(sK1,sK5)))) | $spl14),
  inference(subsumption_resolution,[],[f5904,f44])).
fof(f5904,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK4))),s(sK1,sK5)))) | $spl14),
  inference(superposition,[],[f335,f53])).
fof(f53,plain,(
  s(bool,'T_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK4),s(sK1,sK5)))),
  inference(subsumption_resolution,[],[f52,f43])).
fof(f52,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK4),s(sK1,sK5)))),
  inference(superposition,[],[f32,f45])).
fof(f32,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK4),s(sK1,sK5))))),
  inference(cnf_transformation,[],[f23])).
fof(f4879,plain,(
  s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK2,sK1)),s(sK1,sK10(sK3,sK2,sK1)))) | $spl84),
  inference(cnf_transformation,[],[f4879_D])).
fof(f4879_D,plain,(
  s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK2,sK1)),s(sK1,sK10(sK3,sK2,sK1)))) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f4880,plain,(
  ( ! [X9] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK10(sK3,sK2,sK1)))),s(sK1,X9)))) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK10(sK3,sK2,sK1)),s(sK1,X9))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,X9))))) ) | $spl86),
  inference(cnf_transformation,[],[f4880_D])).
fof(f4880_D,plain,(
  ( ! [X9] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK10(sK3,sK2,sK1)))),s(sK1,X9)))) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK10(sK3,sK2,sK1)),s(sK1,X9))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,X9))))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f5840,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK10(sK3,sK2,sK1)))),s(sK1,X0)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,X0)))) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK10(sK3,sK2,sK1)),s(sK1,X0))) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK2,sK1)),s(sK1,sK10(sK3,sK2,sK1))))) ) | $spl12),
  inference(subsumption_resolution,[],[f5821,f44])).
fof(f5821,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK10(sK3,sK2,sK1)))),s(sK1,X0)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,X0)))) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK10(sK3,sK2,sK1)),s(sK1,X0))) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK2,sK1)),s(sK1,sK10(sK3,sK2,sK1))))) ) | $spl12),
  inference(superposition,[],[f1216,f442])).
fof(f442,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1)))) | $spl12),
  inference(subsumption_resolution,[],[f441,f43])).
fof(f441,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1)))) | $spl12),
  inference(superposition,[],[f329,f45])).
fof(f329,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1))))) | $spl12),
  inference(cnf_transformation,[],[f329_D])).
fof(f329_D,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1))))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f1216,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X5))),s(sK1,X6)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X6))),s(sK1,X7)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X5))),s(sK1,X7)))) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X6),s(sK1,X7))) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X5),s(sK1,X6)))) )),
  inference(subsumption_resolution,[],[f1200,f44])).
fof(f1200,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X5))),s(sK1,X7)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X6))),s(sK1,X7)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X5))),s(sK1,X6)))) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X6),s(sK1,X7))) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X5),s(sK1,X6)))) )),
  inference(superposition,[],[f171,f45])).
fof(f171,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X3),s(sK1,X1)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X3))),s(sK1,X2)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X1))),s(sK1,X2)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X3))),s(sK1,X1)))) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X1),s(sK1,X2)))) )),
  inference(subsumption_resolution,[],[f167,f44])).
fof(f167,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X3))),s(sK1,X2)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X3),s(sK1,X1)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X1))),s(sK1,X2)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X3))),s(sK1,X1)))) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X1),s(sK1,X2)))) )),
  inference(superposition,[],[f31,f45])).
fof(f31,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X6),s(sK1,X7)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X5))),s(sK1,X7)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X5),s(sK1,X6)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X6))),s(sK1,X7)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X5))),s(sK1,X6))))) )),
  inference(cnf_transformation,[],[f23])).
fof(f5796,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK10(sK3,sK2,sK1)))),s(sK1,X0)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,X0)))) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK10(sK3,sK2,sK1)),s(sK1,X0))) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK2,sK1)),s(sK1,sK10(sK3,sK2,sK1))))) ) | $spl12),
  inference(resolution,[],[f329,f1216])).
fof(f230,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK8(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0),s(sK1,sK7(sK2,sK3,sK1)))))) ) | $spl10),
  inference(cnf_transformation,[],[f230_D])).
fof(f230_D,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK8(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0),s(sK1,sK7(sK2,sK3,sK1)))))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f5749,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1)))) | $spl102),
  inference(cnf_transformation,[],[f5749_D])).
fof(f5749_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1)))) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f5768,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK3,sK1)))))) ) | ($spl2 | $spl43)),
  inference(subsumption_resolution,[],[f5767,f44])).
fof(f5767,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK3,sK1)))))) ) | ($spl2 | $spl43)),
  inference(forward_demodulation,[],[f5766,f469])).
fof(f469,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK3,sK1)))),s(sK1,sK8(sK2,sK3,sK1)))) | $spl2),
  inference(subsumption_resolution,[],[f468,f43])).
fof(f468,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK3,sK1)))),s(sK1,sK8(sK2,sK3,sK1)))) | $spl2),
  inference(superposition,[],[f357,f45])).
fof(f357,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK3,sK1)))),s(sK1,sK8(sK2,sK3,sK1))))) | $spl2),
  inference(resolution,[],[f202,f35])).
fof(f35,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK7(X2,X1,X0)))),s(X0,sK8(X2,X1,X0))))) | ~sP0(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f26])).
fof(f26,plain,(
  ! [X0,X1,X2] : (~sP0(X0,X1,X2) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK6(X2,X1,X0)))),s(X0,sK7(X2,X1,X0))))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK7(X2,X1,X0)))),s(X0,sK8(X2,X1,X0))))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X2),s(X0,sK6(X2,X1,X0)),s(X0,sK7(X2,X1,X0))))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X2),s(X0,sK7(X2,X1,X0)),s(X0,sK8(X2,X1,X0))))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK6(X2,X1,X0)))),s(X0,sK8(X2,X1,X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK6,sK7,sK8])],[f25])).
fof(f25,plain,(
  ! [X0,X1,X2] : (~sP0(X0,X1,X2) | ? [X3,X4,X5] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X4)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X5)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X2),s(X0,X3),s(X0,X4)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X2),s(X0,X4),s(X0,X5)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X5))))))),
  inference(rectify,[],[f24])).
fof(f24,plain,(
  ! [X0,X2,X1] : (~sP0(X0,X2,X1) | ? [X3,X4,X5] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X4))),s(X0,X5)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X5)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X5))))))),
  inference(nnf_transformation,[],[f20])).
fof(f20,plain,(
  ! [X0,X2,X1] : (~sP0(X0,X2,X1) | ? [X3,X4,X5] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X4))),s(X0,X5)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X5)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X5))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f5766,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK3,sK1)))),s(sK1,sK8(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK3,sK1)))))) ) | $spl43),
  inference(subsumption_resolution,[],[f5765,f1647])).
fof(f5765,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK3,sK1)))),s(sK1,sK8(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK3,sK1)))))) )),
  inference(forward_demodulation,[],[f1624,f53])).
fof(f1624,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK4),s(sK1,sK5))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK3,sK1)))),s(sK1,sK8(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK3,sK1)))))) )),
  inference(resolution,[],[f1237,f162])).
fof(f162,plain,(
  ( ! [X0,X1] : (~sP0(sK1,X1,sK2) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0),s(sK1,sK7(sK2,X1,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,X1,sK1)))),s(sK1,sK8(sK2,X1,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK7(sK2,X1,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK8(sK2,X1,sK1)))))) )),
  inference(resolution,[],[f31,f37])).
fof(f37,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X2),s(X0,sK7(X2,X1,X0)),s(X0,sK8(X2,X1,X0))))) | ~sP0(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f26])).
fof(f1237,plain,(
  ( ! [X1] : (sP0(sK1,sK3,X1) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X1),s(sK1,sK4),s(sK1,sK5))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),X1),s(sK1,sK9(sK3,X1,sK1)))),s(sK1,sK10(sK3,X1,sK1))))) )),
  inference(subsumption_resolution,[],[f1230,f43])).
fof(f1230,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | sP0(sK1,sK3,X1) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X1),s(sK1,sK4),s(sK1,sK5))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),X1),s(sK1,sK9(sK3,X1,sK1)))),s(sK1,sK10(sK3,X1,sK1))))) )),
  inference(superposition,[],[f203,f45])).
fof(f203,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK9(sK3,X0,sK1)))),s(sK1,sK10(sK3,X0,sK1))))) | sP0(sK1,sK3,X0) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK4),s(sK1,sK5)))) )),
  inference(subsumption_resolution,[],[f196,f44])).
fof(f196,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | sP0(sK1,sK3,X0) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK9(sK3,X0,sK1)))),s(sK1,sK10(sK3,X0,sK1))))) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK4),s(sK1,sK5)))) )),
  inference(superposition,[],[f54,f45])).
fof(f54,plain,(
  ( ! [X0] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK4),s(sK1,sK5)))) | sP0(sK1,sK3,X0) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK9(sK3,X0,sK1)))),s(sK1,sK10(sK3,X0,sK1)))))) )),
  inference(resolution,[],[f33,f39])).
fof(f39,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | sP0(X0,X2,X1) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK9(X2,X1,X0)))),s(X0,sK10(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f28])).
fof(f28,plain,(
  ! [X0,X1,X2] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK9(X2,X1,X0)))),s(X0,sK10(X2,X1,X0))))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK9(X2,X1,X0)))),s(X0,sK10(X2,X1,X0)))))) | sP0(X0,X2,X1) | ! [X3,X4] : (~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK9,sK10])],[f27])).
fof(f27,plain,(
  ! [X0,X1,X2] : (? [X5,X6] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X6)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X5))),s(X0,X6))))) | sP0(X0,X2,X1) | ! [X3,X4] : (~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))))))),
  inference(rectify,[],[f21])).
fof(f21,plain,(
  ! [X0,X1,X2] : (? [X6,X7] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X6))),s(X0,X7)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X6))),s(X0,X7))))) | sP0(X0,X2,X1) | ! [X8,X9] : (~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X8),s(X0,X9)))) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X8))),s(X0,X9)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X8),s(X0,X9)))))))),
  inference(definition_folding,[],[f18,f20])).
fof(f18,plain,(
  ! [X0,X1,X2] : (? [X6,X7] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X6))),s(X0,X7)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X6))),s(X0,X7))))) | ? [X3,X4,X5] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X4))),s(X0,X5)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X5)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X5))))) | ! [X8,X9] : (~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X8),s(X0,X9)))) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X8))),s(X0,X9)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X8),s(X0,X9)))))))),
  inference(flattening,[],[f17])).
fof(f17,plain,(
  ! [X0,X1,X2] : ((? [X6,X7] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X6))),s(X0,X7)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X6))),s(X0,X7))))) | ? [X3,X4,X5] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X4))),s(X0,X5)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X5))))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X5)))))) | ! [X8,X9] : (~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X8),s(X0,X9)))) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X8))),s(X0,X9)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X8),s(X0,X9)))))))),
  inference(ennf_transformation,[],[f11])).
fof(f11,plain,(
  ! [X0,X1,X2] : ((! [X6,X7] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X6))),s(X0,X7)))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X6))),s(X0,X7))))) & ! [X3,X4,X5] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X4))),s(X0,X5)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X5))))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X5)))))) => ! [X8,X9] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X8),s(X0,X9)))) => (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X8))),s(X0,X9)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X8),s(X0,X9)))))))),
  inference(flattening,[],[f10])).
fof(f10,plain,(
  ! [X0] : ! [X1,X2] : ((! [X6,X7] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X6))),s(X0,X7)))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X6))),s(X0,X7))))) & ! [X3,X4,X5] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X4))),s(X0,X5)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X5))))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X5)))))) => ! [X8,X9] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X8),s(X0,X9)))) => (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X8))),s(X0,X9)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X8),s(X0,X9)))))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6] : ! [X7,X8] : ((! [X9,X10] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X9))),s(X6,X10)))) => p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,X9))),s(X6,X10))))) & ! [X9,X10,X11] : ((p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,X9))),s(X6,X10)))) & p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,X10))),s(X6,X11)))) & p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X9),s(X6,X10)))) & p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X10),s(X6,X11))))) => p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,X9))),s(X6,X11)))))) => ! [X12,X13] : (p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X12),s(X6,X13)))) => (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,X12))),s(X6,X13)))) & p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X12),s(X6,X13)))))))),
  file('thm.relation.TC_STRONG_INDUCT','thm.relation.relation_47_prove_ltb')).
fof(f5745,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1)))) | $spl100),
  inference(cnf_transformation,[],[f5745_D])).
fof(f5745_D,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1)))) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f5764,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK3,sK1)))))) ) | ($spl2 | $spl43)),
  inference(subsumption_resolution,[],[f5763,f44])).
fof(f5763,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK3,sK1)))))) ) | ($spl2 | $spl43)),
  inference(forward_demodulation,[],[f5762,f469])).
fof(f5762,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK3,sK1)))),s(sK1,sK8(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK3,sK1)))))) ) | $spl43),
  inference(subsumption_resolution,[],[f1642,f1647])).
fof(f1642,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK3,sK1)))),s(sK1,sK8(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK3,sK1)))))) )),
  inference(forward_demodulation,[],[f1636,f53])).
fof(f1636,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK4),s(sK1,sK5))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK3,sK1)))),s(sK1,sK8(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK3,sK1)))))) )),
  inference(resolution,[],[f1420,f162])).
fof(f1420,plain,(
  ( ! [X0] : (sP0(sK1,sK3,X0) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK4),s(sK1,sK5))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,X0,sK1)))),s(sK1,sK10(sK3,X0,sK1))))) )),
  inference(subsumption_resolution,[],[f1415,f44])).
fof(f1415,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | sP0(sK1,sK3,X0) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK4),s(sK1,sK5))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,X0,sK1)))),s(sK1,sK10(sK3,X0,sK1))))) )),
  inference(superposition,[],[f293,f190])).
fof(f190,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,X1))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0))),s(sK1,X1)))) )),
  inference(subsumption_resolution,[],[f188,f44])).
fof(f188,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,X1))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0))),s(sK1,X1)))) )),
  inference(superposition,[],[f125,f45])).
fof(f125,plain,(
  ( ! [X2,X3] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2))),s(sK1,X3)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,X3)))) )),
  inference(subsumption_resolution,[],[f123,f43])).
fof(f123,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2))),s(sK1,X3)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,X3)))) )),
  inference(superposition,[],[f30,f45])).
fof(f30,plain,(
  ( ! [X8,X9] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X8))),s(sK1,X9)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X8))),s(sK1,X9))))) )),
  inference(cnf_transformation,[],[f23])).
fof(f293,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,X0,sK1)))),s(sK1,sK10(sK3,X0,sK1))))) | sP0(sK1,sK3,X0) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK4),s(sK1,sK5)))) )),
  inference(subsumption_resolution,[],[f291,f44])).
fof(f291,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | sP0(sK1,sK3,X0) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,X0,sK1)))),s(sK1,sK10(sK3,X0,sK1))))) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK4),s(sK1,sK5)))) )),
  inference(superposition,[],[f55,f45])).
fof(f55,plain,(
  ( ! [X1] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X1),s(sK1,sK4),s(sK1,sK5)))) | sP0(sK1,sK3,X1) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,X1,sK1)))),s(sK1,sK10(sK3,X1,sK1)))))) )),
  inference(resolution,[],[f33,f40])).
fof(f40,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | sP0(X0,X2,X1) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK9(X2,X1,X0)))),s(X0,sK10(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f28])).
fof(f5741,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1)))) | $spl98),
  inference(cnf_transformation,[],[f5741_D])).
fof(f5741_D,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1)))) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f5761,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK3,sK1)))))) ) | ($spl2 | $spl43)),
  inference(subsumption_resolution,[],[f5760,f44])).
fof(f5760,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK3,sK1)))))) ) | ($spl2 | $spl43)),
  inference(forward_demodulation,[],[f5759,f469])).
fof(f5759,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK3,sK1)))),s(sK1,sK8(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK3,sK1)))))) ) | $spl43),
  inference(subsumption_resolution,[],[f1657,f1647])).
fof(f1657,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK3,sK1)))),s(sK1,sK8(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK3,sK1)))))) )),
  inference(forward_demodulation,[],[f1651,f53])).
fof(f1651,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK4),s(sK1,sK5))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK3,sK1)))),s(sK1,sK8(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK3,sK1)))))) )),
  inference(resolution,[],[f1421,f162])).
fof(f1421,plain,(
  ( ! [X1] : (sP0(sK1,sK3,X1) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X1),s(sK1,sK4),s(sK1,sK5))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,X1,sK1)))),s(sK1,sK10(sK3,X1,sK1))))) )),
  inference(subsumption_resolution,[],[f1416,f44])).
fof(f1416,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | sP0(sK1,sK3,X1) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X1),s(sK1,sK4),s(sK1,sK5))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,X1,sK1)))),s(sK1,sK10(sK3,X1,sK1))))) )),
  inference(superposition,[],[f293,f45])).
fof(f3083,plain,(
  s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1)))) | $spl50),
  inference(cnf_transformation,[],[f3083_D])).
fof(f3083_D,plain,(
  s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1)))) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f5747,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK3,sK1)))))) ) | $spl2),
  inference(subsumption_resolution,[],[f5746,f44])).
fof(f5746,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK3,sK1)))))) ) | $spl2),
  inference(forward_demodulation,[],[f3096,f469])).
fof(f3096,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK3,sK1)))),s(sK1,sK8(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK3,sK1)))))) ) | $spl2),
  inference(resolution,[],[f2046,f162])).
fof(f2046,plain,(
  ( ! [X1] : (sP0(sK1,sK3,X1) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X1),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),X1),s(sK1,sK9(sK3,X1,sK1)))),s(sK1,sK10(sK3,X1,sK1))))) ) | $spl2),
  inference(subsumption_resolution,[],[f2041,f43])).
fof(f2041,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | sP0(sK1,sK3,X1) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X1),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),X1),s(sK1,sK9(sK3,X1,sK1)))),s(sK1,sK10(sK3,X1,sK1))))) ) | $spl2),
  inference(superposition,[],[f1676,f45])).
fof(f1676,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK9(sK3,X0,sK1)))),s(sK1,sK10(sK3,X0,sK1))))) | sP0(sK1,sK3,X0) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1))))) ) | $spl2),
  inference(subsumption_resolution,[],[f1674,f44])).
fof(f1674,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | sP0(sK1,sK3,X0) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK9(sK3,X0,sK1)))),s(sK1,sK10(sK3,X0,sK1))))) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1))))) ) | $spl2),
  inference(superposition,[],[f471,f45])).
fof(f471,plain,(
  ( ! [X0] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1))))) | sP0(sK1,sK3,X0) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK9(sK3,X0,sK1)))),s(sK1,sK10(sK3,X0,sK1)))))) ) | $spl2),
  inference(resolution,[],[f360,f39])).
fof(f360,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK3,sK1)))),s(sK1,sK8(sK2,sK3,sK1))))) | $spl2),
  inference(resolution,[],[f202,f38])).
fof(f38,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK6(X2,X1,X0)))),s(X0,sK8(X2,X1,X0))))) | ~sP0(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f26])).
fof(f5743,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK3,sK1)))))) ) | $spl2),
  inference(subsumption_resolution,[],[f5742,f44])).
fof(f5742,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK3,sK1)))))) ) | $spl2),
  inference(forward_demodulation,[],[f3104,f469])).
fof(f3104,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK3,sK1)))),s(sK1,sK8(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK3,sK1)))))) ) | $spl2),
  inference(resolution,[],[f2064,f162])).
fof(f2064,plain,(
  ( ! [X0] : (sP0(sK1,sK3,X0) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,X0,sK1)))),s(sK1,sK10(sK3,X0,sK1))))) ) | $spl2),
  inference(subsumption_resolution,[],[f2058,f44])).
fof(f2058,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | sP0(sK1,sK3,X0) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,X0,sK1)))),s(sK1,sK10(sK3,X0,sK1))))) ) | $spl2),
  inference(superposition,[],[f1685,f190])).
fof(f1685,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,X0,sK1)))),s(sK1,sK10(sK3,X0,sK1))))) | sP0(sK1,sK3,X0) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1))))) ) | $spl2),
  inference(subsumption_resolution,[],[f1683,f44])).
fof(f1683,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | sP0(sK1,sK3,X0) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,X0,sK1)))),s(sK1,sK10(sK3,X0,sK1))))) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1))))) ) | $spl2),
  inference(superposition,[],[f472,f45])).
fof(f472,plain,(
  ( ! [X1] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X1),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1))))) | sP0(sK1,sK3,X1) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,X1,sK1)))),s(sK1,sK10(sK3,X1,sK1)))))) ) | $spl2),
  inference(resolution,[],[f360,f40])).
fof(f5739,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK3,sK1)))))) ) | $spl2),
  inference(subsumption_resolution,[],[f5738,f44])).
fof(f5738,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK3,sK1)))))) ) | $spl2),
  inference(forward_demodulation,[],[f3116,f469])).
fof(f3116,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK3,sK1)))),s(sK1,sK8(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK3,sK1)))))) ) | $spl2),
  inference(resolution,[],[f2065,f162])).
fof(f2065,plain,(
  ( ! [X1] : (sP0(sK1,sK3,X1) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X1),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,X1,sK1)))),s(sK1,sK10(sK3,X1,sK1))))) ) | $spl2),
  inference(subsumption_resolution,[],[f2059,f44])).
fof(f2059,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | sP0(sK1,sK3,X1) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X1),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,X1,sK1)))),s(sK1,sK10(sK3,X1,sK1))))) ) | $spl2),
  inference(superposition,[],[f1685,f45])).
fof(f200,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1))))) | $spl0),
  inference(cnf_transformation,[],[f200_D])).
fof(f200_D,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f335,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0),s(sK1,X1)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,X1))))) ) | $spl14),
  inference(cnf_transformation,[],[f335_D])).
fof(f335_D,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0),s(sK1,X1)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,X1))))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f5608,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,X1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0),s(sK1,X1)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1)))))) ) | $spl2),
  inference(subsumption_resolution,[],[f5607,f43])).
fof(f5607,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,X1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0),s(sK1,X1)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1)))))) ) | $spl2),
  inference(forward_demodulation,[],[f5606,f479])).
fof(f479,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK3,sK1)))),s(sK1,sK8(sK2,sK3,sK1)))) | $spl2),
  inference(subsumption_resolution,[],[f477,f44])).
fof(f477,plain,(
  ~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK3,sK1)))),s(sK1,sK8(sK2,sK3,sK1)))) | $spl2),
  inference(superposition,[],[f360,f45])).
fof(f5606,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK3,sK1)))),s(sK1,sK8(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,X1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0),s(sK1,X1)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1)))))) ) | $spl2),
  inference(subsumption_resolution,[],[f5605,f44])).
fof(f5605,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK3,sK1)))),s(sK1,sK8(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,X1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0),s(sK1,X1)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1)))))) ) | $spl2),
  inference(forward_demodulation,[],[f5604,f462])).
fof(f462,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK3,sK1)))),s(sK1,sK7(sK2,sK3,sK1)))) | $spl2),
  inference(subsumption_resolution,[],[f461,f43])).
fof(f461,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK3,sK1)))),s(sK1,sK7(sK2,sK3,sK1)))) | $spl2),
  inference(superposition,[],[f356,f45])).
fof(f356,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK3,sK1)))),s(sK1,sK7(sK2,sK3,sK1))))) | $spl2),
  inference(resolution,[],[f202,f34])).
fof(f34,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK6(X2,X1,X0)))),s(X0,sK7(X2,X1,X0))))) | ~sP0(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f26])).
fof(f5604,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK3,sK1)))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK3,sK1)))),s(sK1,sK8(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,X1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0),s(sK1,X1)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1)))))) ) | $spl2),
  inference(subsumption_resolution,[],[f5603,f44])).
fof(f5603,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK3,sK1)))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK3,sK1)))),s(sK1,sK8(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,X1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0),s(sK1,X1)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1)))))) ) | $spl2),
  inference(forward_demodulation,[],[f5602,f469])).
fof(f5602,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK3,sK1)))),s(sK1,sK8(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK3,sK1)))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK3,sK1)))),s(sK1,sK8(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,X1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0),s(sK1,X1)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1)))))) ) | $spl2),
  inference(subsumption_resolution,[],[f5597,f44])).
fof(f5597,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK3,sK1)))),s(sK1,sK8(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK3,sK1)))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK3,sK1)))),s(sK1,sK8(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,X1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0),s(sK1,X1)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1)))))) ) | $spl2),
  inference(superposition,[],[f1132,f368])).
fof(f368,plain,(
  s(bool,'T_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK7(sK2,sK3,sK1)))) | $spl2),
  inference(subsumption_resolution,[],[f367,f43])).
fof(f367,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK7(sK2,sK3,sK1)))) | $spl2),
  inference(superposition,[],[f358,f45])).
fof(f358,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK7(sK2,sK3,sK1))))) | $spl2),
  inference(resolution,[],[f202,f36])).
fof(f36,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X2),s(X0,sK6(X2,X1,X0)),s(X0,sK7(X2,X1,X0))))) | ~sP0(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f26])).
fof(f1132,plain,(
  ( ! [X4,X2,X5,X3] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,X3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,X3,sK1)))),s(sK1,sK8(sK2,X3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,X3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,X3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),X3),s(sK1,X4))),s(sK1,X5)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X4),s(sK1,X5)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(X3,sK2,sK1)))),s(sK1,sK10(X3,sK2,sK1)))))) )),
  inference(resolution,[],[f162,f39])).
fof(f330,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1))))) | $spl13),
  inference(cnf_transformation,[],[f330_D])).
fof(f330_D,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1))))) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f5704,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,X1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0),s(sK1,X1)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1)))))) ) | $spl2),
  inference(subsumption_resolution,[],[f5703,f43])).
fof(f5703,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,X1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0),s(sK1,X1)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1)))))) ) | $spl2),
  inference(forward_demodulation,[],[f5702,f479])).
fof(f5702,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK3,sK1)))),s(sK1,sK8(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,X1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0),s(sK1,X1)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1)))))) ) | $spl2),
  inference(subsumption_resolution,[],[f5701,f44])).
fof(f5701,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK3,sK1)))),s(sK1,sK8(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,X1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0),s(sK1,X1)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1)))))) ) | $spl2),
  inference(forward_demodulation,[],[f5700,f462])).
fof(f5700,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK3,sK1)))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK3,sK1)))),s(sK1,sK8(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,X1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0),s(sK1,X1)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1)))))) ) | $spl2),
  inference(subsumption_resolution,[],[f5699,f44])).
fof(f5699,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK3,sK1)))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK3,sK1)))),s(sK1,sK8(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,X1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0),s(sK1,X1)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1)))))) ) | $spl2),
  inference(forward_demodulation,[],[f5698,f469])).
fof(f5698,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK3,sK1)))),s(sK1,sK8(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK3,sK1)))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK3,sK1)))),s(sK1,sK8(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,X1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0),s(sK1,X1)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1)))))) ) | $spl2),
  inference(subsumption_resolution,[],[f5693,f44])).
fof(f5693,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK3,sK1)))),s(sK1,sK8(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK3,sK1)))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK3,sK1)))),s(sK1,sK8(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,X1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0),s(sK1,X1)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1)))))) ) | $spl2),
  inference(superposition,[],[f1133,f368])).
fof(f1133,plain,(
  ( ! [X6,X8,X7,X9] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X6),s(sK1,sK7(sK2,X7,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,X7,sK1)))),s(sK1,sK8(sK2,X7,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X6))),s(sK1,sK7(sK2,X7,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X6))),s(sK1,sK8(sK2,X7,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),X7),s(sK1,X8))),s(sK1,X9)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X8),s(sK1,X9)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),X7),s(sK1,sK9(X7,sK2,sK1)))),s(sK1,sK10(X7,sK2,sK1)))))) )),
  inference(resolution,[],[f162,f40])).
fof(f5728,plain,(
  $false | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f5727,f43])).
fof(f5727,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f5726,f62])).
fof(f5726,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK4))),s(sK1,sK5)))) | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f5718,f44])).
fof(f5718,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK4))),s(sK1,sK5)))) | ($spl0 | $spl2)),
  inference(superposition,[],[f5706,f53])).
fof(f5706,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0),s(sK1,X1)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,X1))))) ) | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f5705,f44])).
fof(f5705,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,X1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0),s(sK1,X1))))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f5704,f336])).
fof(f336,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1)))) | $spl0),
  inference(resolution,[],[f200,f125])).
fof(f3139,plain,(
  s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK2,sK2,sK1)),s(sK1,sK8(sK2,sK2,sK1)))) | $spl52),
  inference(cnf_transformation,[],[f3139_D])).
fof(f3139_D,plain,(
  s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK2,sK2,sK1)),s(sK1,sK8(sK2,sK2,sK1)))) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f5627,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK8(sK2,sK2,sK1)))),s(sK1,X0)))) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8(sK2,sK2,sK1)),s(sK1,X0))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK2,sK1)))),s(sK1,X0))))) ) | $spl96),
  inference(cnf_transformation,[],[f5627_D])).
fof(f5627_D,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK8(sK2,sK2,sK1)))),s(sK1,X0)))) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8(sK2,sK2,sK1)),s(sK1,X0))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK2,sK1)))),s(sK1,X0))))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f5618,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK8(sK2,sK2,sK1)))),s(sK1,X0)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK2,sK1)))),s(sK1,X0)))) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8(sK2,sK2,sK1)),s(sK1,X0))) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK2,sK2,sK1)),s(sK1,sK8(sK2,sK2,sK1))))) ) | $spl94),
  inference(resolution,[],[f5616,f1216])).
fof(f752,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK2,sK1)))),s(sK1,sK10(sK2,sK2,sK1))))) | $spl28),
  inference(cnf_transformation,[],[f752_D])).
fof(f752_D,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK2,sK1)))),s(sK1,sK10(sK2,sK2,sK1))))) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f606,plain,(
  ( ! [X2,X1] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X1),s(sK1,X2)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X1))),s(sK1,X2))))) ) | $spl16),
  inference(cnf_transformation,[],[f606_D])).
fof(f606_D,plain,(
  ( ! [X2,X1] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X1),s(sK1,X2)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X1))),s(sK1,X2))))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f5616,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK2,sK1)))),s(sK1,sK8(sK2,sK2,sK1))))) | $spl94),
  inference(cnf_transformation,[],[f5616_D])).
fof(f5616_D,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK2,sK1)))),s(sK1,sK8(sK2,sK2,sK1))))) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f5614,plain,(
  ( ! [X2,X3] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK2,sK1)))),s(sK1,sK8(sK2,sK2,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2))),s(sK1,X3)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,X3)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK2,sK1)))),s(sK1,sK10(sK2,sK2,sK1)))))) ) | $spl18),
  inference(subsumption_resolution,[],[f5613,f44])).
fof(f5613,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK2,sK1)))),s(sK1,sK8(sK2,sK2,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2))),s(sK1,X3)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,X3)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK2,sK1)))),s(sK1,sK10(sK2,sK2,sK1)))))) ) | $spl18),
  inference(forward_demodulation,[],[f5612,f769])).
fof(f769,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK2,sK1)))),s(sK1,sK7(sK2,sK2,sK1)))) | $spl18),
  inference(resolution,[],[f608,f180])).
fof(f180,plain,(
  ( ! [X6] : (~sP0(sK1,sK2,X6) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(X6,sK2,sK1)))),s(sK1,sK7(X6,sK2,sK1))))) )),
  inference(resolution,[],[f125,f34])).
fof(f5612,plain,(
  ( ! [X2,X3] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK2,sK1)))),s(sK1,sK7(sK2,sK2,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK2,sK1)))),s(sK1,sK8(sK2,sK2,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2))),s(sK1,X3)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,X3)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK2,sK1)))),s(sK1,sK10(sK2,sK2,sK1)))))) ) | $spl18),
  inference(subsumption_resolution,[],[f5611,f44])).
fof(f5611,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK2,sK1)))),s(sK1,sK7(sK2,sK2,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK2,sK1)))),s(sK1,sK8(sK2,sK2,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2))),s(sK1,X3)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,X3)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK2,sK1)))),s(sK1,sK10(sK2,sK2,sK1)))))) ) | $spl18),
  inference(forward_demodulation,[],[f5610,f861])).
fof(f861,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK2,sK1)))),s(sK1,sK8(sK2,sK2,sK1)))) | $spl18),
  inference(resolution,[],[f771,f125])).
fof(f771,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7(sK2,sK2,sK1)))),s(sK1,sK8(sK2,sK2,sK1))))) | $spl18),
  inference(resolution,[],[f608,f35])).
fof(f5610,plain,(
  ( ! [X2,X3] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK2,sK1)))),s(sK1,sK8(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK2,sK1)))),s(sK1,sK7(sK2,sK2,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK2,sK1)))),s(sK1,sK8(sK2,sK2,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2))),s(sK1,X3)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,X3)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK2,sK1)))),s(sK1,sK10(sK2,sK2,sK1)))))) ) | $spl18),
  inference(subsumption_resolution,[],[f5598,f44])).
fof(f5598,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK2,sK1)))),s(sK1,sK8(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK2,sK1)))),s(sK1,sK7(sK2,sK2,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK2,sK1)))),s(sK1,sK8(sK2,sK2,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2))),s(sK1,X3)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,X3)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK2,sK1)))),s(sK1,sK10(sK2,sK2,sK1)))))) ) | $spl18),
  inference(superposition,[],[f1132,f783])).
fof(f783,plain,(
  s(bool,'T_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK2,sK2,sK1)),s(sK1,sK7(sK2,sK2,sK1)))) | $spl18),
  inference(subsumption_resolution,[],[f782,f43])).
fof(f782,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK2,sK2,sK1)),s(sK1,sK7(sK2,sK2,sK1)))) | $spl18),
  inference(superposition,[],[f772,f45])).
fof(f772,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK2,sK2,sK1)),s(sK1,sK7(sK2,sK2,sK1))))) | $spl18),
  inference(resolution,[],[f608,f36])).
fof(f636,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK8(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK7(sK2,sK2,sK1))))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0),s(sK1,sK7(sK2,sK2,sK1)))))) ) | $spl26),
  inference(cnf_transformation,[],[f636_D])).
fof(f636_D,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK8(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK7(sK2,sK2,sK1))))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0),s(sK1,sK7(sK2,sK2,sK1)))))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f1807,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK2,sK1)))),s(sK1,sK10(sK2,sK2,sK1)))) | $spl48),
  inference(cnf_transformation,[],[f1807_D])).
fof(f1807_D,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK2,sK1)))),s(sK1,sK10(sK2,sK2,sK1)))) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f4860,plain,(
  s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK3,sK1)),s(sK1,sK10(sK2,sK3,sK1)))) | $spl78),
  inference(cnf_transformation,[],[f4860_D])).
fof(f4860_D,plain,(
  s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK3,sK1)),s(sK1,sK10(sK2,sK3,sK1)))) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f5341,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK3,sK1)),s(sK1,sK10(sK2,sK3,sK1)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK2,sK1)))),s(sK1,sK10(sK2,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK2,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK2,sK1)))))) ) | ($spl18 | $spl88)),
  inference(subsumption_resolution,[],[f5340,f44])).
fof(f5340,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK3,sK1)),s(sK1,sK10(sK2,sK3,sK1)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK2,sK1)))),s(sK1,sK10(sK2,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK2,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK2,sK1)))))) ) | ($spl18 | $spl88)),
  inference(forward_demodulation,[],[f5332,f861])).
fof(f5332,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK3,sK1)),s(sK1,sK10(sK2,sK3,sK1)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK2,sK1)))),s(sK1,sK10(sK2,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK2,sK1)))),s(sK1,sK8(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK2,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK2,sK1)))))) ) | $spl88),
  inference(resolution,[],[f5291,f162])).
fof(f5291,plain,(
  ( ! [X0] : (sP0(sK1,sK2,X0) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK9(sK2,sK3,sK1)),s(sK1,sK10(sK2,sK3,sK1)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,X0,sK1)))),s(sK1,sK10(sK2,X0,sK1))))) ) | $spl88),
  inference(subsumption_resolution,[],[f5274,f44])).
fof(f5274,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | sP0(sK1,sK2,X0) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK9(sK2,sK3,sK1)),s(sK1,sK10(sK2,sK3,sK1)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,X0,sK1)))),s(sK1,sK10(sK2,X0,sK1))))) ) | $spl88),
  inference(superposition,[],[f5198,f45])).
fof(f5198,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,X0,sK1)))),s(sK1,sK10(sK2,X0,sK1))))) | sP0(sK1,sK2,X0) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK9(sK2,sK3,sK1)),s(sK1,sK10(sK2,sK3,sK1))))) ) | $spl88),
  inference(subsumption_resolution,[],[f5196,f44])).
fof(f5196,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | sP0(sK1,sK2,X0) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,X0,sK1)))),s(sK1,sK10(sK2,X0,sK1))))) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK9(sK2,sK3,sK1)),s(sK1,sK10(sK2,sK3,sK1))))) ) | $spl88),
  inference(superposition,[],[f4942,f45])).
fof(f4942,plain,(
  ( ! [X5] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X5),s(sK1,sK9(sK2,sK3,sK1)),s(sK1,sK10(sK2,sK3,sK1))))) | sP0(sK1,sK2,X5) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,X5,sK1)))),s(sK1,sK10(sK2,X5,sK1)))))) ) | $spl88),
  inference(subsumption_resolution,[],[f4931,f43])).
fof(f4931,plain,(
  ( ! [X5] : (p(s(bool,'F_0')) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X5),s(sK1,sK9(sK2,sK3,sK1)),s(sK1,sK10(sK2,sK3,sK1))))) | sP0(sK1,sK2,X5) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,X5,sK1)))),s(sK1,sK10(sK2,X5,sK1)))))) ) | $spl88),
  inference(superposition,[],[f40,f4917])).
fof(f1793,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK2,sK1)))),s(sK1,sK10(sK2,sK2,sK1)))) | $spl46),
  inference(cnf_transformation,[],[f1793_D])).
fof(f1793_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK2,sK1)))),s(sK1,sK10(sK2,sK2,sK1)))) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f5324,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK3,sK1)),s(sK1,sK10(sK2,sK3,sK1)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK2,sK1)))),s(sK1,sK10(sK2,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK2,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK2,sK1)))))) ) | ($spl18 | $spl88)),
  inference(subsumption_resolution,[],[f5323,f44])).
fof(f5323,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK3,sK1)),s(sK1,sK10(sK2,sK3,sK1)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK2,sK1)))),s(sK1,sK10(sK2,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK2,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK2,sK1)))))) ) | ($spl18 | $spl88)),
  inference(forward_demodulation,[],[f5315,f861])).
fof(f5315,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK3,sK1)),s(sK1,sK10(sK2,sK3,sK1)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK2,sK1)))),s(sK1,sK10(sK2,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK2,sK1)))),s(sK1,sK8(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK2,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK2,sK1)))))) ) | $spl88),
  inference(resolution,[],[f5249,f162])).
fof(f5249,plain,(
  ( ! [X1] : (sP0(sK1,sK2,X1) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X1),s(sK1,sK9(sK2,sK3,sK1)),s(sK1,sK10(sK2,sK3,sK1)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),X1),s(sK1,sK9(sK2,X1,sK1)))),s(sK1,sK10(sK2,X1,sK1))))) ) | $spl88),
  inference(subsumption_resolution,[],[f5234,f43])).
fof(f5234,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | sP0(sK1,sK2,X1) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X1),s(sK1,sK9(sK2,sK3,sK1)),s(sK1,sK10(sK2,sK3,sK1)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),X1),s(sK1,sK9(sK2,X1,sK1)))),s(sK1,sK10(sK2,X1,sK1))))) ) | $spl88),
  inference(superposition,[],[f5193,f45])).
fof(f5193,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK9(sK2,X0,sK1)))),s(sK1,sK10(sK2,X0,sK1))))) | sP0(sK1,sK2,X0) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK9(sK2,sK3,sK1)),s(sK1,sK10(sK2,sK3,sK1))))) ) | $spl88),
  inference(subsumption_resolution,[],[f5191,f44])).
fof(f5191,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | sP0(sK1,sK2,X0) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK9(sK2,X0,sK1)))),s(sK1,sK10(sK2,X0,sK1))))) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK9(sK2,sK3,sK1)),s(sK1,sK10(sK2,sK3,sK1))))) ) | $spl88),
  inference(superposition,[],[f4941,f45])).
fof(f4941,plain,(
  ( ! [X4] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X4),s(sK1,sK9(sK2,sK3,sK1)),s(sK1,sK10(sK2,sK3,sK1))))) | sP0(sK1,sK2,X4) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),X4),s(sK1,sK9(sK2,X4,sK1)))),s(sK1,sK10(sK2,X4,sK1)))))) ) | $spl88),
  inference(subsumption_resolution,[],[f4930,f43])).
fof(f4930,plain,(
  ( ! [X4] : (p(s(bool,'F_0')) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X4),s(sK1,sK9(sK2,sK3,sK1)),s(sK1,sK10(sK2,sK3,sK1))))) | sP0(sK1,sK2,X4) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),X4),s(sK1,sK9(sK2,X4,sK1)))),s(sK1,sK10(sK2,X4,sK1)))))) ) | $spl88),
  inference(superposition,[],[f39,f4917])).
fof(f4861,plain,(
  ( ! [X30] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK10(sK2,sK3,sK1)))),s(sK1,X30)))) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK10(sK2,sK3,sK1)),s(sK1,X30))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK2,sK3,sK1)))),s(sK1,X30))))) ) | $spl80),
  inference(cnf_transformation,[],[f4861_D])).
fof(f4861_D,plain,(
  ( ! [X30] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK10(sK2,sK3,sK1)))),s(sK1,X30)))) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK10(sK2,sK3,sK1)),s(sK1,X30))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK2,sK3,sK1)))),s(sK1,X30))))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f5248,plain,(
  s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK2,sK3,sK1)),s(sK1,sK10(sK2,sK3,sK1)))) | $spl92),
  inference(cnf_transformation,[],[f5248_D])).
fof(f5248_D,plain,(
  s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK2,sK3,sK1)),s(sK1,sK10(sK2,sK3,sK1)))) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f4858,plain,(
  sP0(sK1,sK2,sK3) | $spl76),
  inference(cnf_transformation,[],[f4858_D])).
fof(f4858_D,plain,(
  sP0(sK1,sK2,sK3) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f5224,plain,(
  ( ! [X2] : (sP0(sK1,sK2,sK3) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK2,sK3,sK1)),s(sK1,sK10(sK2,sK3,sK1)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK10(sK2,sK3,sK1)))),s(sK1,X2)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK2,sK3,sK1)))),s(sK1,X2)))) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK10(sK2,sK3,sK1)),s(sK1,X2))) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK3,sK1)),s(sK1,sK10(sK2,sK3,sK1))))) ) | $spl88),
  inference(resolution,[],[f5193,f1216])).
fof(f5182,plain,(
  s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK2,sK1)),s(sK1,sK10(sK3,sK2,sK1)))) | $spl90),
  inference(cnf_transformation,[],[f5182_D])).
fof(f5182_D,plain,(
  s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK2,sK1)),s(sK1,sK10(sK3,sK2,sK1)))) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f5212,plain,(
  ( ! [X3] : (s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK2,sK1)),s(sK1,sK10(sK3,sK2,sK1)))) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X3),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X3))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X3))),s(sK1,sK8(sK2,sK3,sK1)))))) ) | ($spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5211,f44])).
fof(f5211,plain,(
  ( ! [X3] : (~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK2,sK1)),s(sK1,sK10(sK3,sK2,sK1)))) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X3),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X3))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X3))),s(sK1,sK8(sK2,sK3,sK1)))))) ) | ($spl2 | $spl30)),
  inference(forward_demodulation,[],[f5202,f469])).
fof(f5202,plain,(
  ( ! [X3] : (s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK2,sK1)),s(sK1,sK10(sK3,sK2,sK1)))) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X3),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK3,sK1)))),s(sK1,sK8(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X3))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X3))),s(sK1,sK8(sK2,sK3,sK1)))))) ) | ($spl2 | $spl30)),
  inference(resolution,[],[f5008,f162])).
fof(f5008,plain,(
  ( ! [X9] : (sP0(sK1,sK3,X9) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,X9,sK1)),s(sK1,sK10(sK3,X9,sK1)))) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X9),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1))))) ) | ($spl2 | $spl30)),
  inference(resolution,[],[f4807,f1685])).
fof(f4807,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,X1)))) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0),s(sK1,X1)))) ) | $spl30),
  inference(subsumption_resolution,[],[f4805,f44])).
fof(f4805,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,X1)))) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0),s(sK1,X1)))) ) | $spl30),
  inference(superposition,[],[f1232,f45])).
fof(f5180,plain,(
  ( ! [X3] : (s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK2,sK1)),s(sK1,sK10(sK3,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X3),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X3))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X3))),s(sK1,sK8(sK2,sK3,sK1)))))) ) | ($spl2 | $spl30 | $spl43)),
  inference(subsumption_resolution,[],[f5179,f44])).
fof(f5179,plain,(
  ( ! [X3] : (~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK2,sK1)),s(sK1,sK10(sK3,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X3),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X3))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X3))),s(sK1,sK8(sK2,sK3,sK1)))))) ) | ($spl2 | $spl30 | $spl43)),
  inference(forward_demodulation,[],[f5178,f469])).
fof(f5178,plain,(
  ( ! [X3] : (s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK2,sK1)),s(sK1,sK10(sK3,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X3),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK3,sK1)))),s(sK1,sK8(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X3))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X3))),s(sK1,sK8(sK2,sK3,sK1)))))) ) | ($spl30 | $spl43)),
  inference(subsumption_resolution,[],[f5177,f1647])).
fof(f5177,plain,(
  ( ! [X3] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK2,sK1)),s(sK1,sK10(sK3,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X3),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK3,sK1)))),s(sK1,sK8(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X3))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X3))),s(sK1,sK8(sK2,sK3,sK1)))))) ) | $spl30),
  inference(forward_demodulation,[],[f5168,f53])).
fof(f5168,plain,(
  ( ! [X3] : (s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK2,sK1)),s(sK1,sK10(sK3,sK2,sK1)))) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK4),s(sK1,sK5))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X3),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK3,sK1)))),s(sK1,sK8(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X3))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X3))),s(sK1,sK8(sK2,sK3,sK1)))))) ) | $spl30),
  inference(resolution,[],[f5009,f162])).
fof(f5009,plain,(
  ( ! [X10] : (sP0(sK1,sK3,X10) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,X10,sK1)),s(sK1,sK10(sK3,X10,sK1)))) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X10),s(sK1,sK4),s(sK1,sK5)))) ) | $spl30),
  inference(resolution,[],[f4807,f293])).
fof(f4974,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK10(sK2,sK3,sK1)))),s(sK1,X0)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK2,sK3,sK1)))),s(sK1,X0)))) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK10(sK2,sK3,sK1)),s(sK1,X0))) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK3,sK1)),s(sK1,sK10(sK2,sK3,sK1))))) ) | $spl60),
  inference(subsumption_resolution,[],[f4956,f44])).
fof(f4956,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK10(sK2,sK3,sK1)))),s(sK1,X0)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK2,sK3,sK1)))),s(sK1,X0)))) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK10(sK2,sK3,sK1)),s(sK1,X0))) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK3,sK1)),s(sK1,sK10(sK2,sK3,sK1))))) ) | $spl60),
  inference(superposition,[],[f1216,f4718])).
fof(f4917,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK3,sK1)))),s(sK1,sK10(sK2,sK3,sK1)))) | $spl88),
  inference(cnf_transformation,[],[f4917_D])).
fof(f4917_D,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK3,sK1)))),s(sK1,sK10(sK2,sK3,sK1)))) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f4856,plain,(
  s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK3,sK1)),s(sK1,sK10(sK3,sK3,sK1)))) | $spl74),
  inference(cnf_transformation,[],[f4856_D])).
fof(f4856_D,plain,(
  s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK3,sK1)),s(sK1,sK10(sK3,sK3,sK1)))) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f4716,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK3,sK2,sK1)))),s(sK1,sK8(sK3,sK2,sK1)))) | $spl58),
  inference(cnf_transformation,[],[f4716_D])).
fof(f4716_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK3,sK2,sK1)))),s(sK1,sK8(sK3,sK2,sK1)))) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f4913,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK3,sK2,sK1)))),s(sK1,sK8(sK3,sK2,sK1)))) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK3,sK1)),s(sK1,sK10(sK3,sK3,sK1)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK3,sK1)))),s(sK1,sK10(sK2,sK3,sK1)))) | ($spl30 | $spl37)),
  inference(resolution,[],[f4890,f2257])).
fof(f2257,plain,(
  ( ! [X0] : (sP0(sK1,sK2,X0) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK9(sK3,sK3,sK1)),s(sK1,sK10(sK3,sK3,sK1)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,X0,sK1)))),s(sK1,sK10(sK2,X0,sK1))))) ) | $spl37),
  inference(subsumption_resolution,[],[f2248,f44])).
fof(f2248,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | sP0(sK1,sK2,X0) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK9(sK3,sK3,sK1)),s(sK1,sK10(sK3,sK3,sK1)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,X0,sK1)))),s(sK1,sK10(sK2,X0,sK1))))) ) | $spl37),
  inference(superposition,[],[f1830,f45])).
fof(f1830,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,X0,sK1)))),s(sK1,sK10(sK2,X0,sK1))))) | sP0(sK1,sK2,X0) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK9(sK3,sK3,sK1)),s(sK1,sK10(sK3,sK3,sK1))))) ) | $spl37),
  inference(subsumption_resolution,[],[f1828,f44])).
fof(f1828,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | sP0(sK1,sK2,X0) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,X0,sK1)))),s(sK1,sK10(sK2,X0,sK1))))) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK9(sK3,sK3,sK1)),s(sK1,sK10(sK3,sK3,sK1))))) ) | $spl37),
  inference(superposition,[],[f1369,f45])).
fof(f1369,plain,(
  ( ! [X1] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X1),s(sK1,sK9(sK3,sK3,sK1)),s(sK1,sK10(sK3,sK3,sK1))))) | sP0(sK1,sK2,X1) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,X1,sK1)))),s(sK1,sK10(sK2,X1,sK1)))))) ) | $spl37),
  inference(resolution,[],[f1254,f40])).
fof(f4890,plain,(
  ( ! [X2] : (~sP0(sK1,X2,sK3) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK3,X2,sK1)))),s(sK1,sK8(sK3,X2,sK1))))) ) | $spl30),
  inference(subsumption_resolution,[],[f4889,f43])).
fof(f4889,plain,(
  ( ! [X2] : (p(s(bool,'F_0')) | ~sP0(sK1,X2,sK3) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK3,X2,sK1)))),s(sK1,sK8(sK3,X2,sK1))))) ) | $spl30),
  inference(superposition,[],[f4799,f45])).
fof(f4799,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK3,X0,sK1)))),s(sK1,sK8(sK3,X0,sK1))))) | ~sP0(sK1,X0,sK3)) ) | $spl30),
  inference(resolution,[],[f1232,f37])).
fof(f4718,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK2,sK3,sK1)))),s(sK1,sK10(sK2,sK3,sK1)))) | $spl60),
  inference(cnf_transformation,[],[f4718_D])).
fof(f4718_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK2,sK3,sK1)))),s(sK1,sK10(sK2,sK3,sK1)))) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f4912,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK3,sK2,sK1)))),s(sK1,sK8(sK3,sK2,sK1)))) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK3,sK1)),s(sK1,sK10(sK3,sK3,sK1)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK2,sK3,sK1)))),s(sK1,sK10(sK2,sK3,sK1)))) | ($spl30 | $spl37)),
  inference(resolution,[],[f4890,f2234])).
fof(f2234,plain,(
  ( ! [X1] : (sP0(sK1,sK2,X1) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X1),s(sK1,sK9(sK3,sK3,sK1)),s(sK1,sK10(sK3,sK3,sK1)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),X1),s(sK1,sK9(sK2,X1,sK1)))),s(sK1,sK10(sK2,X1,sK1))))) ) | $spl37),
  inference(subsumption_resolution,[],[f2226,f43])).
fof(f2226,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | sP0(sK1,sK2,X1) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X1),s(sK1,sK9(sK3,sK3,sK1)),s(sK1,sK10(sK3,sK3,sK1)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),X1),s(sK1,sK9(sK2,X1,sK1)))),s(sK1,sK10(sK2,X1,sK1))))) ) | $spl37),
  inference(superposition,[],[f1825,f45])).
fof(f1825,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK9(sK2,X0,sK1)))),s(sK1,sK10(sK2,X0,sK1))))) | sP0(sK1,sK2,X0) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK9(sK3,sK3,sK1)),s(sK1,sK10(sK3,sK3,sK1))))) ) | $spl37),
  inference(subsumption_resolution,[],[f1823,f44])).
fof(f1823,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | sP0(sK1,sK2,X0) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK9(sK2,X0,sK1)))),s(sK1,sK10(sK2,X0,sK1))))) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK9(sK3,sK3,sK1)),s(sK1,sK10(sK3,sK3,sK1))))) ) | $spl37),
  inference(superposition,[],[f1368,f45])).
fof(f1368,plain,(
  ( ! [X0] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK9(sK3,sK3,sK1)),s(sK1,sK10(sK3,sK3,sK1))))) | sP0(sK1,sK2,X0) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK9(sK2,X0,sK1)))),s(sK1,sK10(sK2,X0,sK1)))))) ) | $spl37),
  inference(resolution,[],[f1254,f39])).
fof(f4863,plain,(
  s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK2,sK1)),s(sK1,sK8(sK2,sK2,sK1)))) | $spl82),
  inference(cnf_transformation,[],[f4863_D])).
fof(f4863_D,plain,(
  s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK2,sK1)),s(sK1,sK8(sK2,sK2,sK1)))) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f4911,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK3,sK2,sK1)))),s(sK1,sK8(sK3,sK2,sK1)))) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK2,sK1)),s(sK1,sK8(sK2,sK2,sK1)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK3,sK1)))),s(sK1,sK10(sK2,sK3,sK1)))) | ($spl18 | $spl30)),
  inference(resolution,[],[f4890,f2105])).
fof(f2105,plain,(
  ( ! [X0] : (sP0(sK1,sK2,X0) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK6(sK2,sK2,sK1)),s(sK1,sK8(sK2,sK2,sK1)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,X0,sK1)))),s(sK1,sK10(sK2,X0,sK1))))) ) | $spl18),
  inference(subsumption_resolution,[],[f2098,f44])).
fof(f2098,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | sP0(sK1,sK2,X0) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK6(sK2,sK2,sK1)),s(sK1,sK8(sK2,sK2,sK1)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,X0,sK1)))),s(sK1,sK10(sK2,X0,sK1))))) ) | $spl18),
  inference(superposition,[],[f1695,f45])).
fof(f1695,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,X0,sK1)))),s(sK1,sK10(sK2,X0,sK1))))) | sP0(sK1,sK2,X0) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK6(sK2,sK2,sK1)),s(sK1,sK8(sK2,sK2,sK1))))) ) | $spl18),
  inference(subsumption_resolution,[],[f1693,f44])).
fof(f1693,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | sP0(sK1,sK2,X0) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,X0,sK1)))),s(sK1,sK10(sK2,X0,sK1))))) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK6(sK2,sK2,sK1)),s(sK1,sK8(sK2,sK2,sK1))))) ) | $spl18),
  inference(superposition,[],[f870,f45])).
fof(f870,plain,(
  ( ! [X1] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X1),s(sK1,sK6(sK2,sK2,sK1)),s(sK1,sK8(sK2,sK2,sK1))))) | sP0(sK1,sK2,X1) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,X1,sK1)))),s(sK1,sK10(sK2,X1,sK1)))))) ) | $spl18),
  inference(resolution,[],[f774,f40])).
fof(f774,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK2,sK2,sK1)))),s(sK1,sK8(sK2,sK2,sK1))))) | $spl18),
  inference(resolution,[],[f608,f38])).
fof(f4910,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK3,sK2,sK1)))),s(sK1,sK8(sK3,sK2,sK1)))) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK2,sK1)),s(sK1,sK8(sK2,sK2,sK1)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK2,sK3,sK1)))),s(sK1,sK10(sK2,sK3,sK1)))) | ($spl18 | $spl30)),
  inference(resolution,[],[f4890,f2086])).
fof(f2086,plain,(
  ( ! [X1] : (sP0(sK1,sK2,X1) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X1),s(sK1,sK6(sK2,sK2,sK1)),s(sK1,sK8(sK2,sK2,sK1)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),X1),s(sK1,sK9(sK2,X1,sK1)))),s(sK1,sK10(sK2,X1,sK1))))) ) | $spl18),
  inference(subsumption_resolution,[],[f2080,f43])).
fof(f2080,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | sP0(sK1,sK2,X1) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X1),s(sK1,sK6(sK2,sK2,sK1)),s(sK1,sK8(sK2,sK2,sK1)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),X1),s(sK1,sK9(sK2,X1,sK1)))),s(sK1,sK10(sK2,X1,sK1))))) ) | $spl18),
  inference(superposition,[],[f1690,f45])).
fof(f1690,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK9(sK2,X0,sK1)))),s(sK1,sK10(sK2,X0,sK1))))) | sP0(sK1,sK2,X0) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK6(sK2,sK2,sK1)),s(sK1,sK8(sK2,sK2,sK1))))) ) | $spl18),
  inference(subsumption_resolution,[],[f1688,f44])).
fof(f1688,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | sP0(sK1,sK2,X0) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK9(sK2,X0,sK1)))),s(sK1,sK10(sK2,X0,sK1))))) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK6(sK2,sK2,sK1)),s(sK1,sK8(sK2,sK2,sK1))))) ) | $spl18),
  inference(superposition,[],[f869,f45])).
fof(f869,plain,(
  ( ! [X0] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK6(sK2,sK2,sK1)),s(sK1,sK8(sK2,sK2,sK1))))) | sP0(sK1,sK2,X0) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK9(sK2,X0,sK1)))),s(sK1,sK10(sK2,X0,sK1)))))) ) | $spl18),
  inference(resolution,[],[f774,f39])).
fof(f4877,plain,(
  ( ! [X9] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK10(sK3,sK2,sK1)))),s(sK1,X9)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,X9)))) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK10(sK3,sK2,sK1)),s(sK1,X9))) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK2,sK1)),s(sK1,sK10(sK3,sK2,sK1))))) ) | $spl0),
  inference(subsumption_resolution,[],[f4852,f44])).
fof(f4852,plain,(
  ( ! [X9] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK10(sK3,sK2,sK1)))),s(sK1,X9)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,X9)))) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK10(sK3,sK2,sK1)),s(sK1,X9))) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK2,sK1)),s(sK1,sK10(sK3,sK2,sK1))))) ) | $spl0),
  inference(superposition,[],[f1216,f336])).
fof(f4833,plain,(
  ( ! [X31] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK10(sK2,sK3,sK1)))),s(sK1,X31)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK2,sK3,sK1)))),s(sK1,X31)))) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK10(sK2,sK3,sK1)),s(sK1,X31))) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK3,sK1)),s(sK1,sK10(sK2,sK3,sK1)))) | sP0(sK1,sK2,sK3) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK2,sK1)),s(sK1,sK8(sK2,sK2,sK1))))) ) | $spl18),
  inference(resolution,[],[f1216,f1690])).
fof(f4832,plain,(
  ( ! [X30] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK10(sK2,sK3,sK1)))),s(sK1,X30)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK2,sK3,sK1)))),s(sK1,X30)))) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK10(sK2,sK3,sK1)),s(sK1,X30))) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK3,sK1)),s(sK1,sK10(sK2,sK3,sK1)))) | sP0(sK1,sK2,sK3) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK3,sK1)),s(sK1,sK10(sK3,sK3,sK1))))) ) | $spl37),
  inference(resolution,[],[f1216,f1825])).
fof(f4737,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK3,sK1)))),s(sK1,sK10(sK3,sK3,sK1))))) | $spl73),
  inference(cnf_transformation,[],[f4737_D])).
fof(f4737_D,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK3,sK1)))),s(sK1,sK10(sK3,sK3,sK1))))) <=> ~$spl73),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl73])])).
fof(f1232,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0),s(sK1,X1)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,X1))))) ) | $spl30),
  inference(cnf_transformation,[],[f1232_D])).
fof(f1232_D,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0),s(sK1,X1)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,X1))))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f4732,plain,(
  ( ! [X2,X3] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,X3)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2),s(sK1,X3)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK3,sK1)))),s(sK1,sK10(sK3,sK3,sK1)))))) ) | $spl35),
  inference(resolution,[],[f1235,f40])).
fof(f1235,plain,(
  ~sP0(sK1,sK3,sK3) | $spl35),
  inference(cnf_transformation,[],[f1235_D])).
fof(f1235_D,plain,(
  ~sP0(sK1,sK3,sK3) <=> ~$spl35),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl35])])).
fof(f4738,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK3,sK1)))),s(sK1,sK10(sK3,sK3,sK1))))) | $spl72),
  inference(cnf_transformation,[],[f4738_D])).
fof(f4738_D,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK3,sK1)))),s(sK1,sK10(sK3,sK3,sK1))))) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f4731,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,X1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0),s(sK1,X1)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK3,sK1)))),s(sK1,sK10(sK3,sK3,sK1)))))) ) | $spl35),
  inference(resolution,[],[f1235,f39])).
fof(f4736,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK3,sK1)))),s(sK1,sK10(sK3,sK3,sK1)))) | $spl70),
  inference(cnf_transformation,[],[f4736_D])).
fof(f4736_D,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK3,sK1)))),s(sK1,sK10(sK3,sK3,sK1)))) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f4724,plain,(
  s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1)))) | $spl66),
  inference(cnf_transformation,[],[f4724_D])).
fof(f4724_D,plain,(
  s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1)))) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f4730,plain,(
  s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK3,sK1)))),s(sK1,sK10(sK3,sK3,sK1)))) | ($spl2 | $spl35)),
  inference(resolution,[],[f1235,f2065])).
fof(f4734,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK3,sK1)))),s(sK1,sK10(sK3,sK3,sK1)))) | $spl68),
  inference(cnf_transformation,[],[f4734_D])).
fof(f4734_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK3,sK1)))),s(sK1,sK10(sK3,sK3,sK1)))) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f4728,plain,(
  s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK3,sK1)))),s(sK1,sK10(sK3,sK3,sK1)))) | ($spl2 | $spl35)),
  inference(resolution,[],[f1235,f2046])).
fof(f1234,plain,(
  s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK4),s(sK1,sK5))) | $spl32),
  inference(cnf_transformation,[],[f1234_D])).
fof(f1234_D,plain,(
  s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK4),s(sK1,sK5))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f4727,plain,(
  s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK4),s(sK1,sK5))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK3,sK1)))),s(sK1,sK10(sK3,sK3,sK1)))) | $spl35),
  inference(resolution,[],[f1235,f1421])).
fof(f4725,plain,(
  s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK4),s(sK1,sK5))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK3,sK1)))),s(sK1,sK10(sK3,sK3,sK1)))) | $spl35),
  inference(resolution,[],[f1235,f1237])).
fof(f1236,plain,(
  sP0(sK1,sK3,sK3) | $spl34),
  inference(cnf_transformation,[],[f1236_D])).
fof(f1236_D,plain,(
  sP0(sK1,sK3,sK3) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f1419,plain,(
  sP0(sK1,sK3,sK3) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK4),s(sK1,sK5)))),
  inference(duplicate_literal_removal,[],[f1405])).
fof(f1405,plain,(
  sP0(sK1,sK3,sK3) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK4),s(sK1,sK5))) | sP0(sK1,sK3,sK3) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK4),s(sK1,sK5)))),
  inference(resolution,[],[f293,f203])).
fof(f1418,plain,(
  ( ! [X2,X1] : (sP0(sK1,sK3,sK3) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK4),s(sK1,sK5))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X1))),s(sK1,X2)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X1),s(sK1,X2))))) )),
  inference(duplicate_literal_removal,[],[f1407])).
fof(f1407,plain,(
  ( ! [X2,X1] : (sP0(sK1,sK3,sK3) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK4),s(sK1,sK5))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X1))),s(sK1,X2)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X1),s(sK1,X2)))) | sP0(sK1,sK3,sK3)) )),
  inference(resolution,[],[f293,f39])).
fof(f2043,plain,(
  sP0(sK1,sK3,sK3) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1)))) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK4),s(sK1,sK5))) | $spl2),
  inference(duplicate_literal_removal,[],[f2032])).
fof(f2032,plain,(
  sP0(sK1,sK3,sK3) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1)))) | sP0(sK1,sK3,sK3) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK4),s(sK1,sK5))) | $spl2),
  inference(resolution,[],[f1676,f293])).
fof(f2042,plain,(
  ( ! [X2,X3] : (sP0(sK1,sK3,sK3) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,X3)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2),s(sK1,X3))))) ) | $spl2),
  inference(duplicate_literal_removal,[],[f2033])).
fof(f2033,plain,(
  ( ! [X2,X3] : (sP0(sK1,sK3,sK3) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,X3)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2),s(sK1,X3)))) | sP0(sK1,sK3,sK3)) ) | $spl2),
  inference(resolution,[],[f1676,f40])).
fof(f2063,plain,(
  sP0(sK1,sK3,sK3) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1)))) | $spl2),
  inference(duplicate_literal_removal,[],[f2047])).
fof(f2047,plain,(
  sP0(sK1,sK3,sK3) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1)))) | sP0(sK1,sK3,sK3) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1)))) | $spl2),
  inference(resolution,[],[f1685,f1676])).
fof(f2062,plain,(
  sP0(sK1,sK3,sK3) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1)))) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK4),s(sK1,sK5))) | $spl2),
  inference(duplicate_literal_removal,[],[f2048])).
fof(f2048,plain,(
  sP0(sK1,sK3,sK3) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1)))) | sP0(sK1,sK3,sK3) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK4),s(sK1,sK5))) | $spl2),
  inference(resolution,[],[f1685,f203])).
fof(f2061,plain,(
  ( ! [X2,X1] : (sP0(sK1,sK3,sK3) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X1))),s(sK1,X2)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X1),s(sK1,X2))))) ) | $spl2),
  inference(duplicate_literal_removal,[],[f2050])).
fof(f2050,plain,(
  ( ! [X2,X1] : (sP0(sK1,sK3,sK3) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X1))),s(sK1,X2)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X1),s(sK1,X2)))) | sP0(sK1,sK3,sK3)) ) | $spl2),
  inference(resolution,[],[f1685,f39])).
fof(f4722,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK3,sK2,sK1)))),s(sK1,sK7(sK3,sK2,sK1)))) | $spl64),
  inference(cnf_transformation,[],[f4722_D])).
fof(f4722_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK3,sK2,sK1)))),s(sK1,sK7(sK3,sK2,sK1)))) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f4720,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK3,sK3,sK1)))),s(sK1,sK7(sK3,sK3,sK1)))) | $spl62),
  inference(cnf_transformation,[],[f4720_D])).
fof(f4720_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK3,sK3,sK1)))),s(sK1,sK7(sK3,sK3,sK1)))) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f4709,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK3,sK3,sK1)))),s(sK1,sK7(sK3,sK3,sK1)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK2,sK3,sK1)))),s(sK1,sK10(sK2,sK3,sK1)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK3,sK2,sK1)))),s(sK1,sK7(sK3,sK2,sK1)))) | $spl34),
  inference(resolution,[],[f3962,f180])).
fof(f3962,plain,(
  ( ! [X2] : (sP0(sK1,X2,sK3) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),X2),s(sK1,sK6(sK3,sK3,sK1)))),s(sK1,sK7(sK3,sK3,sK1)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(X2,sK3,sK1)))),s(sK1,sK10(X2,sK3,sK1))))) ) | $spl34),
  inference(subsumption_resolution,[],[f3955,f43])).
fof(f3955,plain,(
  ( ! [X2] : (p(s(bool,'F_0')) | sP0(sK1,X2,sK3) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),X2),s(sK1,sK6(sK3,sK3,sK1)))),s(sK1,sK7(sK3,sK3,sK1)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(X2,sK3,sK1)))),s(sK1,sK10(X2,sK3,sK1))))) ) | $spl34),
  inference(superposition,[],[f2882,f45])).
fof(f2882,plain,(
  ( ! [X1] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(X1,sK3,sK1)))),s(sK1,sK10(X1,sK3,sK1))))) | sP0(sK1,X1,sK3) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),X1),s(sK1,sK6(sK3,sK3,sK1)))),s(sK1,sK7(sK3,sK3,sK1))))) ) | $spl34),
  inference(subsumption_resolution,[],[f2862,f43])).
fof(f2862,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(X1,sK3,sK1)))),s(sK1,sK10(X1,sK3,sK1))))) | sP0(sK1,X1,sK3) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),X1),s(sK1,sK6(sK3,sK3,sK1)))),s(sK1,sK7(sK3,sK3,sK1))))) ) | $spl34),
  inference(superposition,[],[f1246,f45])).
fof(f1246,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK6(sK3,sK3,sK1)))),s(sK1,sK7(sK3,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(X0,sK3,sK1)))),s(sK1,sK10(X0,sK3,sK1))))) | sP0(sK1,X0,sK3)) ) | $spl34),
  inference(resolution,[],[f1242,f39])).
fof(f1242,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK3,sK3,sK1)),s(sK1,sK7(sK3,sK3,sK1))))) | $spl34),
  inference(resolution,[],[f1236,f36])).
fof(f4708,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK3,sK3,sK1)))),s(sK1,sK7(sK3,sK3,sK1)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK2,sK3,sK1)))),s(sK1,sK10(sK2,sK3,sK1)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK3,sK2,sK1)))),s(sK1,sK8(sK3,sK2,sK1)))) | $spl34),
  inference(resolution,[],[f3962,f181])).
fof(f181,plain,(
  ( ! [X7] : (~sP0(sK1,sK2,X7) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(X7,sK2,sK1)))),s(sK1,sK8(X7,sK2,sK1))))) )),
  inference(resolution,[],[f125,f35])).
fof(f3217,plain,(
  s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK3,sK1)),s(sK1,sK10(sK3,sK3,sK1)))) | $spl56),
  inference(cnf_transformation,[],[f3217_D])).
fof(f3217_D,plain,(
  s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK3,sK1)),s(sK1,sK10(sK3,sK3,sK1)))) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f3227,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK3,sK1)),s(sK1,sK10(sK3,sK3,sK1)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK2,sK1)))),s(sK1,sK10(sK2,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK2,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK2,sK1)))))) ) | ($spl18 | $spl37)),
  inference(subsumption_resolution,[],[f3226,f44])).
fof(f3226,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK3,sK1)),s(sK1,sK10(sK3,sK3,sK1)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK2,sK1)))),s(sK1,sK10(sK2,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK2,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK2,sK1)))))) ) | ($spl18 | $spl37)),
  inference(forward_demodulation,[],[f3220,f861])).
fof(f3220,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK3,sK1)),s(sK1,sK10(sK3,sK3,sK1)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK2,sK1)))),s(sK1,sK10(sK2,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK2,sK1)))),s(sK1,sK8(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK2,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK2,sK1)))))) ) | $spl37),
  inference(resolution,[],[f2257,f162])).
fof(f3215,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK3,sK1)),s(sK1,sK10(sK3,sK3,sK1)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK2,sK1)))),s(sK1,sK10(sK2,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK2,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK2,sK1)))))) ) | ($spl18 | $spl37)),
  inference(subsumption_resolution,[],[f3214,f44])).
fof(f3214,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK3,sK1)),s(sK1,sK10(sK3,sK3,sK1)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK2,sK1)))),s(sK1,sK10(sK2,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK2,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK2,sK1)))))) ) | ($spl18 | $spl37)),
  inference(forward_demodulation,[],[f3208,f861])).
fof(f3208,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK3,sK1)),s(sK1,sK10(sK3,sK3,sK1)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK2,sK1)))),s(sK1,sK10(sK2,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK2,sK1)))),s(sK1,sK8(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK2,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK2,sK1)))))) ) | $spl37),
  inference(resolution,[],[f2234,f162])).
fof(f3161,plain,(
  s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK3,sK3,sK1)),s(sK1,sK8(sK3,sK3,sK1)))) | $spl54),
  inference(cnf_transformation,[],[f3161_D])).
fof(f3161_D,plain,(
  s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK3,sK3,sK1)),s(sK1,sK8(sK3,sK3,sK1)))) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f3205,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK3,sK3,sK1)),s(sK1,sK8(sK3,sK3,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK3,sK1)))))) ) | ($spl0 | $spl2 | $spl34 | $spl43)),
  inference(subsumption_resolution,[],[f3204,f44])).
fof(f3204,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK3,sK3,sK1)),s(sK1,sK8(sK3,sK3,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK3,sK1)))))) ) | ($spl0 | $spl2 | $spl34 | $spl43)),
  inference(forward_demodulation,[],[f3203,f469])).
fof(f3203,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK3,sK3,sK1)),s(sK1,sK8(sK3,sK3,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK3,sK1)))),s(sK1,sK8(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK3,sK1)))))) ) | ($spl0 | $spl34 | $spl43)),
  inference(subsumption_resolution,[],[f3202,f1647])).
fof(f3202,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK3,sK3,sK1)),s(sK1,sK8(sK3,sK3,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK3,sK1)))),s(sK1,sK8(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK3,sK1)))))) ) | ($spl0 | $spl34)),
  inference(forward_demodulation,[],[f3196,f336])).
fof(f3196,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK3,sK3,sK1)),s(sK1,sK8(sK3,sK3,sK1)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK3,sK1)))),s(sK1,sK8(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK3,sK1)))))) ) | $spl34),
  inference(resolution,[],[f2190,f162])).
fof(f2190,plain,(
  ( ! [X1] : (sP0(sK1,sK3,X1) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X1),s(sK1,sK6(sK3,sK3,sK1)),s(sK1,sK8(sK3,sK3,sK1)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,X1,sK1)))),s(sK1,sK10(sK3,X1,sK1))))) ) | $spl34),
  inference(subsumption_resolution,[],[f2183,f44])).
fof(f2183,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | sP0(sK1,sK3,X1) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X1),s(sK1,sK6(sK3,sK3,sK1)),s(sK1,sK8(sK3,sK3,sK1)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,X1,sK1)))),s(sK1,sK10(sK3,X1,sK1))))) ) | $spl34),
  inference(superposition,[],[f1774,f45])).
fof(f1774,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,X0,sK1)))),s(sK1,sK10(sK3,X0,sK1))))) | sP0(sK1,sK3,X0) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK6(sK3,sK3,sK1)),s(sK1,sK8(sK3,sK3,sK1))))) ) | $spl34),
  inference(subsumption_resolution,[],[f1772,f44])).
fof(f1772,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | sP0(sK1,sK3,X0) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,X0,sK1)))),s(sK1,sK10(sK3,X0,sK1))))) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK6(sK3,sK3,sK1)),s(sK1,sK8(sK3,sK3,sK1))))) ) | $spl34),
  inference(superposition,[],[f1358,f45])).
fof(f1358,plain,(
  ( ! [X1] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X1),s(sK1,sK6(sK3,sK3,sK1)),s(sK1,sK8(sK3,sK3,sK1))))) | sP0(sK1,sK3,X1) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,X1,sK1)))),s(sK1,sK10(sK3,X1,sK1)))))) ) | $spl34),
  inference(resolution,[],[f1244,f40])).
fof(f1244,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK3,sK3,sK1)))),s(sK1,sK8(sK3,sK3,sK1))))) | $spl34),
  inference(resolution,[],[f1236,f38])).
fof(f3193,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK3,sK3,sK1)),s(sK1,sK8(sK3,sK3,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK3,sK1)))))) ) | ($spl0 | $spl2 | $spl34 | $spl43)),
  inference(subsumption_resolution,[],[f3192,f44])).
fof(f3192,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK3,sK3,sK1)),s(sK1,sK8(sK3,sK3,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK3,sK1)))))) ) | ($spl0 | $spl2 | $spl34 | $spl43)),
  inference(forward_demodulation,[],[f3191,f469])).
fof(f3191,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK3,sK3,sK1)),s(sK1,sK8(sK3,sK3,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK3,sK1)))),s(sK1,sK8(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK3,sK1)))))) ) | ($spl0 | $spl34 | $spl43)),
  inference(subsumption_resolution,[],[f3190,f1647])).
fof(f3190,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK3,sK3,sK1)),s(sK1,sK8(sK3,sK3,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK3,sK1)))),s(sK1,sK8(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK3,sK1)))))) ) | ($spl0 | $spl34)),
  inference(forward_demodulation,[],[f3184,f343])).
fof(f343,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1)))) | $spl0),
  inference(subsumption_resolution,[],[f342,f43])).
fof(f342,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1)))) | $spl0),
  inference(superposition,[],[f200,f45])).
fof(f3184,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK3,sK3,sK1)),s(sK1,sK8(sK3,sK3,sK1)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK3,sK1)))),s(sK1,sK8(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK3,sK1)))))) ) | $spl34),
  inference(resolution,[],[f2189,f162])).
fof(f2189,plain,(
  ( ! [X0] : (sP0(sK1,sK3,X0) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK6(sK3,sK3,sK1)),s(sK1,sK8(sK3,sK3,sK1)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,X0,sK1)))),s(sK1,sK10(sK3,X0,sK1))))) ) | $spl34),
  inference(subsumption_resolution,[],[f2182,f44])).
fof(f2182,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | sP0(sK1,sK3,X0) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK6(sK3,sK3,sK1)),s(sK1,sK8(sK3,sK3,sK1)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,X0,sK1)))),s(sK1,sK10(sK3,X0,sK1))))) ) | $spl34),
  inference(superposition,[],[f1774,f190])).
fof(f3171,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK3,sK3,sK1)),s(sK1,sK8(sK3,sK3,sK1)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK2,sK1)))),s(sK1,sK10(sK2,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK2,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK2,sK1)))))) ) | ($spl18 | $spl34)),
  inference(subsumption_resolution,[],[f3170,f44])).
fof(f3170,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK3,sK3,sK1)),s(sK1,sK8(sK3,sK3,sK1)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK2,sK1)))),s(sK1,sK10(sK2,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK2,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK2,sK1)))))) ) | ($spl18 | $spl34)),
  inference(forward_demodulation,[],[f3164,f861])).
fof(f3164,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK3,sK3,sK1)),s(sK1,sK8(sK3,sK3,sK1)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK2,sK1)))),s(sK1,sK10(sK2,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK2,sK1)))),s(sK1,sK8(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK2,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK2,sK1)))))) ) | $spl34),
  inference(resolution,[],[f2147,f162])).
fof(f2147,plain,(
  ( ! [X0] : (sP0(sK1,sK2,X0) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK6(sK3,sK3,sK1)),s(sK1,sK8(sK3,sK3,sK1)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,X0,sK1)))),s(sK1,sK10(sK2,X0,sK1))))) ) | $spl34),
  inference(subsumption_resolution,[],[f2139,f44])).
fof(f2139,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | sP0(sK1,sK2,X0) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK6(sK3,sK3,sK1)),s(sK1,sK8(sK3,sK3,sK1)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,X0,sK1)))),s(sK1,sK10(sK2,X0,sK1))))) ) | $spl34),
  inference(superposition,[],[f1754,f45])).
fof(f1754,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,X0,sK1)))),s(sK1,sK10(sK2,X0,sK1))))) | sP0(sK1,sK2,X0) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK6(sK3,sK3,sK1)),s(sK1,sK8(sK3,sK3,sK1))))) ) | $spl34),
  inference(subsumption_resolution,[],[f1752,f44])).
fof(f1752,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | sP0(sK1,sK2,X0) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,X0,sK1)))),s(sK1,sK10(sK2,X0,sK1))))) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK6(sK3,sK3,sK1)),s(sK1,sK8(sK3,sK3,sK1))))) ) | $spl34),
  inference(superposition,[],[f1330,f45])).
fof(f1330,plain,(
  ( ! [X1] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X1),s(sK1,sK6(sK3,sK3,sK1)),s(sK1,sK8(sK3,sK3,sK1))))) | sP0(sK1,sK2,X1) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,X1,sK1)))),s(sK1,sK10(sK2,X1,sK1)))))) ) | $spl34),
  inference(resolution,[],[f1239,f40])).
fof(f1239,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK3,sK3,sK1)))),s(sK1,sK8(sK3,sK3,sK1))))) | $spl34),
  inference(resolution,[],[f1236,f115])).
fof(f115,plain,(
  ( ! [X0] : (~sP0(sK1,sK3,X0) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(X0,sK3,sK1)))),s(sK1,sK8(X0,sK3,sK1)))))) )),
  inference(resolution,[],[f30,f38])).
fof(f3159,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK3,sK3,sK1)),s(sK1,sK8(sK3,sK3,sK1)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK2,sK1)))),s(sK1,sK10(sK2,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK2,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK2,sK1)))))) ) | ($spl18 | $spl34)),
  inference(subsumption_resolution,[],[f3158,f44])).
fof(f3158,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK3,sK3,sK1)),s(sK1,sK8(sK3,sK3,sK1)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK2,sK1)))),s(sK1,sK10(sK2,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK2,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK2,sK1)))))) ) | ($spl18 | $spl34)),
  inference(forward_demodulation,[],[f3152,f861])).
fof(f3152,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK3,sK3,sK1)),s(sK1,sK8(sK3,sK3,sK1)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK2,sK1)))),s(sK1,sK10(sK2,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK2,sK1)))),s(sK1,sK8(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK2,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK2,sK1)))))) ) | $spl34),
  inference(resolution,[],[f2126,f162])).
fof(f2126,plain,(
  ( ! [X1] : (sP0(sK1,sK2,X1) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X1),s(sK1,sK6(sK3,sK3,sK1)),s(sK1,sK8(sK3,sK3,sK1)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),X1),s(sK1,sK9(sK2,X1,sK1)))),s(sK1,sK10(sK2,X1,sK1))))) ) | $spl34),
  inference(subsumption_resolution,[],[f2119,f43])).
fof(f2119,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | sP0(sK1,sK2,X1) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X1),s(sK1,sK6(sK3,sK3,sK1)),s(sK1,sK8(sK3,sK3,sK1)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),X1),s(sK1,sK9(sK2,X1,sK1)))),s(sK1,sK10(sK2,X1,sK1))))) ) | $spl34),
  inference(superposition,[],[f1749,f45])).
fof(f1749,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK9(sK2,X0,sK1)))),s(sK1,sK10(sK2,X0,sK1))))) | sP0(sK1,sK2,X0) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK6(sK3,sK3,sK1)),s(sK1,sK8(sK3,sK3,sK1))))) ) | $spl34),
  inference(subsumption_resolution,[],[f1747,f44])).
fof(f1747,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | sP0(sK1,sK2,X0) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK9(sK2,X0,sK1)))),s(sK1,sK10(sK2,X0,sK1))))) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK6(sK3,sK3,sK1)),s(sK1,sK8(sK3,sK3,sK1))))) ) | $spl34),
  inference(superposition,[],[f1329,f45])).
fof(f1329,plain,(
  ( ! [X0] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK6(sK3,sK3,sK1)),s(sK1,sK8(sK3,sK3,sK1))))) | sP0(sK1,sK2,X0) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK9(sK2,X0,sK1)))),s(sK1,sK10(sK2,X0,sK1)))))) ) | $spl34),
  inference(resolution,[],[f1239,f39])).
fof(f3149,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK2,sK2,sK1)),s(sK1,sK8(sK2,sK2,sK1)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK2,sK1)))),s(sK1,sK10(sK2,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK2,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK2,sK1)))))) ) | $spl18),
  inference(subsumption_resolution,[],[f3148,f44])).
fof(f3148,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK2,sK2,sK1)),s(sK1,sK8(sK2,sK2,sK1)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK2,sK1)))),s(sK1,sK10(sK2,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK2,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK2,sK1)))))) ) | $spl18),
  inference(forward_demodulation,[],[f3142,f861])).
fof(f3142,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK2,sK2,sK1)),s(sK1,sK8(sK2,sK2,sK1)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK2,sK1)))),s(sK1,sK10(sK2,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK2,sK1)))),s(sK1,sK8(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK2,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK2,sK1)))))) ) | $spl18),
  inference(resolution,[],[f2105,f162])).
fof(f3137,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK2,sK2,sK1)),s(sK1,sK8(sK2,sK2,sK1)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK2,sK1)))),s(sK1,sK10(sK2,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK2,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK2,sK1)))))) ) | $spl18),
  inference(subsumption_resolution,[],[f3136,f44])).
fof(f3136,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK2,sK2,sK1)),s(sK1,sK8(sK2,sK2,sK1)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK2,sK1)))),s(sK1,sK10(sK2,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK2,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK2,sK1)))))) ) | $spl18),
  inference(forward_demodulation,[],[f3130,f861])).
fof(f3130,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK2,sK2,sK1)),s(sK1,sK8(sK2,sK2,sK1)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK2,sK1)))),s(sK1,sK10(sK2,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK2,sK1)))),s(sK1,sK8(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK2,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK2,sK1)))))) ) | $spl18),
  inference(resolution,[],[f2086,f162])).
fof(f3125,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK3,sK1)))))) ) | ($spl0 | $spl2 | $spl43)),
  inference(subsumption_resolution,[],[f3124,f44])).
fof(f3124,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK3,sK1)))))) ) | ($spl0 | $spl2 | $spl43)),
  inference(forward_demodulation,[],[f3123,f469])).
fof(f3123,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK3,sK1)))),s(sK1,sK8(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK3,sK1)))))) ) | ($spl0 | $spl2 | $spl43)),
  inference(subsumption_resolution,[],[f3122,f1647])).
fof(f3122,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK3,sK1)))),s(sK1,sK8(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK3,sK1)))))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f3116,f336])).
fof(f3113,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK3,sK1)))))) ) | ($spl0 | $spl2 | $spl43)),
  inference(subsumption_resolution,[],[f3112,f44])).
fof(f3112,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK3,sK1)))))) ) | ($spl0 | $spl2 | $spl43)),
  inference(forward_demodulation,[],[f3111,f469])).
fof(f3111,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK3,sK1)))),s(sK1,sK8(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK3,sK1)))))) ) | ($spl0 | $spl2 | $spl43)),
  inference(subsumption_resolution,[],[f3110,f1647])).
fof(f3110,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK3,sK1)))),s(sK1,sK8(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK3,sK1)))))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f3104,f343])).
fof(f3093,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK2,sK1)))),s(sK1,sK10(sK2,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK2,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK2,sK1)))))) ) | ($spl2 | $spl18)),
  inference(subsumption_resolution,[],[f3092,f44])).
fof(f3092,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK2,sK1)))),s(sK1,sK10(sK2,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK2,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK2,sK1)))))) ) | ($spl2 | $spl18)),
  inference(forward_demodulation,[],[f3086,f861])).
fof(f3086,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK2,sK1)))),s(sK1,sK10(sK2,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK2,sK1)))),s(sK1,sK8(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK2,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK2,sK1)))))) ) | $spl2),
  inference(resolution,[],[f2028,f162])).
fof(f2028,plain,(
  ( ! [X0] : (sP0(sK1,sK2,X0) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,X0,sK1)))),s(sK1,sK10(sK2,X0,sK1))))) ) | $spl2),
  inference(subsumption_resolution,[],[f2022,f44])).
fof(f2022,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | sP0(sK1,sK2,X0) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,X0,sK1)))),s(sK1,sK10(sK2,X0,sK1))))) ) | $spl2),
  inference(superposition,[],[f1671,f45])).
fof(f1671,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,X0,sK1)))),s(sK1,sK10(sK2,X0,sK1))))) | sP0(sK1,sK2,X0) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1))))) ) | $spl2),
  inference(subsumption_resolution,[],[f1669,f44])).
fof(f1669,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | sP0(sK1,sK2,X0) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,X0,sK1)))),s(sK1,sK10(sK2,X0,sK1))))) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1))))) ) | $spl2),
  inference(superposition,[],[f444,f45])).
fof(f444,plain,(
  ( ! [X1] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X1),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1))))) | sP0(sK1,sK2,X1) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,X1,sK1)))),s(sK1,sK10(sK2,X1,sK1)))))) ) | $spl2),
  inference(resolution,[],[f355,f40])).
fof(f355,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK2,sK3,sK1)))),s(sK1,sK8(sK2,sK3,sK1))))) | $spl2),
  inference(resolution,[],[f202,f115])).
fof(f3081,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK2,sK1)))),s(sK1,sK10(sK2,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK2,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK2,sK1)))))) ) | ($spl2 | $spl18)),
  inference(subsumption_resolution,[],[f3080,f44])).
fof(f3080,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK2,sK1)))),s(sK1,sK10(sK2,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK2,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK2,sK1)))))) ) | ($spl2 | $spl18)),
  inference(forward_demodulation,[],[f3074,f861])).
fof(f3074,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK2,sK1)))),s(sK1,sK10(sK2,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK2,sK1)))),s(sK1,sK8(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK2,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK2,sK1)))))) ) | $spl2),
  inference(resolution,[],[f2011,f162])).
fof(f2011,plain,(
  ( ! [X1] : (sP0(sK1,sK2,X1) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X1),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),X1),s(sK1,sK9(sK2,X1,sK1)))),s(sK1,sK10(sK2,X1,sK1))))) ) | $spl2),
  inference(subsumption_resolution,[],[f2006,f43])).
fof(f2006,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | sP0(sK1,sK2,X1) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X1),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),X1),s(sK1,sK9(sK2,X1,sK1)))),s(sK1,sK10(sK2,X1,sK1))))) ) | $spl2),
  inference(superposition,[],[f1666,f45])).
fof(f1666,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK9(sK2,X0,sK1)))),s(sK1,sK10(sK2,X0,sK1))))) | sP0(sK1,sK2,X0) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1))))) ) | $spl2),
  inference(subsumption_resolution,[],[f1664,f44])).
fof(f1664,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | sP0(sK1,sK2,X0) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK9(sK2,X0,sK1)))),s(sK1,sK10(sK2,X0,sK1))))) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1))))) ) | $spl2),
  inference(superposition,[],[f443,f45])).
fof(f443,plain,(
  ( ! [X0] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1))))) | sP0(sK1,sK2,X0) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK9(sK2,X0,sK1)))),s(sK1,sK10(sK2,X0,sK1)))))) ) | $spl2),
  inference(resolution,[],[f355,f39])).
fof(f1805,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK2,sK1)))),s(sK1,sK10(sK2,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK2,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK2,sK1)))))) ) | ($spl18 | $spl43)),
  inference(subsumption_resolution,[],[f1804,f44])).
fof(f1804,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK2,sK1)))),s(sK1,sK10(sK2,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK2,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK2,sK1)))))) ) | ($spl18 | $spl43)),
  inference(forward_demodulation,[],[f1803,f861])).
fof(f1803,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK2,sK1)))),s(sK1,sK10(sK2,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK2,sK1)))),s(sK1,sK8(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK2,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK2,sK1)))))) ) | $spl43),
  inference(subsumption_resolution,[],[f1802,f1647])).
fof(f1802,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK2,sK1)))),s(sK1,sK10(sK2,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK2,sK1)))),s(sK1,sK8(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK2,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK2,sK1)))))) )),
  inference(forward_demodulation,[],[f1796,f53])).
fof(f1796,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK4),s(sK1,sK5))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK2,sK1)))),s(sK1,sK10(sK2,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK2,sK1)))),s(sK1,sK8(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK2,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK2,sK1)))))) )),
  inference(resolution,[],[f1599,f162])).
fof(f1599,plain,(
  ( ! [X0] : (sP0(sK1,sK2,X0) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK4),s(sK1,sK5))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,X0,sK1)))),s(sK1,sK10(sK2,X0,sK1))))) )),
  inference(subsumption_resolution,[],[f1594,f44])).
fof(f1594,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | sP0(sK1,sK2,X0) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK4),s(sK1,sK5))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,X0,sK1)))),s(sK1,sK10(sK2,X0,sK1))))) )),
  inference(superposition,[],[f1158,f45])).
fof(f1158,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,X0,sK1)))),s(sK1,sK10(sK2,X0,sK1))))) | sP0(sK1,sK2,X0) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK4),s(sK1,sK5)))) )),
  inference(subsumption_resolution,[],[f1156,f44])).
fof(f1156,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | sP0(sK1,sK2,X0) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,X0,sK1)))),s(sK1,sK10(sK2,X0,sK1))))) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK4),s(sK1,sK5)))) )),
  inference(superposition,[],[f127,f45])).
fof(f127,plain,(
  ( ! [X1] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X1),s(sK1,sK4),s(sK1,sK5)))) | sP0(sK1,sK2,X1) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,X1,sK1)))),s(sK1,sK10(sK2,X1,sK1)))))) )),
  inference(resolution,[],[f124,f40])).
fof(f124,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK4))),s(sK1,sK5))))),
  inference(subsumption_resolution,[],[f121,f43])).
fof(f121,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK4))),s(sK1,sK5))))),
  inference(superposition,[],[f30,f62])).
fof(f1791,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK2,sK1)))),s(sK1,sK10(sK2,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK2,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK2,sK1)))))) ) | ($spl18 | $spl43)),
  inference(subsumption_resolution,[],[f1790,f44])).
fof(f1790,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK2,sK1)))),s(sK1,sK10(sK2,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK2,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK2,sK1)))))) ) | ($spl18 | $spl43)),
  inference(forward_demodulation,[],[f1789,f861])).
fof(f1789,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK2,sK1)))),s(sK1,sK10(sK2,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK2,sK1)))),s(sK1,sK8(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK2,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK2,sK1)))))) ) | $spl43),
  inference(subsumption_resolution,[],[f1788,f1647])).
fof(f1788,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK2,sK1)))),s(sK1,sK10(sK2,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK2,sK1)))),s(sK1,sK8(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK2,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK2,sK1)))))) )),
  inference(forward_demodulation,[],[f1782,f53])).
fof(f1782,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK4),s(sK1,sK5))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK2,sK1)))),s(sK1,sK10(sK2,sK2,sK1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK2,sK1)))),s(sK1,sK8(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK2,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK2,sK1)))))) )),
  inference(resolution,[],[f1582,f162])).
fof(f1582,plain,(
  ( ! [X1] : (sP0(sK1,sK2,X1) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X1),s(sK1,sK4),s(sK1,sK5))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),X1),s(sK1,sK9(sK2,X1,sK1)))),s(sK1,sK10(sK2,X1,sK1))))) )),
  inference(subsumption_resolution,[],[f1579,f43])).
fof(f1579,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | sP0(sK1,sK2,X1) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X1),s(sK1,sK4),s(sK1,sK5))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),X1),s(sK1,sK9(sK2,X1,sK1)))),s(sK1,sK10(sK2,X1,sK1))))) )),
  inference(superposition,[],[f1147,f45])).
fof(f1147,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK9(sK2,X0,sK1)))),s(sK1,sK10(sK2,X0,sK1))))) | sP0(sK1,sK2,X0) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK4),s(sK1,sK5)))) )),
  inference(subsumption_resolution,[],[f1145,f44])).
fof(f1145,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | sP0(sK1,sK2,X0) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK9(sK2,X0,sK1)))),s(sK1,sK10(sK2,X0,sK1))))) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK4),s(sK1,sK5)))) )),
  inference(superposition,[],[f126,f45])).
fof(f126,plain,(
  ( ! [X0] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK4),s(sK1,sK5)))) | sP0(sK1,sK2,X0) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK9(sK2,X0,sK1)))),s(sK1,sK10(sK2,X0,sK1)))))) )),
  inference(resolution,[],[f124,f39])).
fof(f1761,plain,(
  ( ! [X1] : (~p(s(bool,happ(s(fun(sK1,bool),X1),s(sK1,sK5)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK11(happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK4)),X1,bool,sK1))))) ) | $spl44),
  inference(cnf_transformation,[],[f1761_D])).
fof(f1761_D,plain,(
  ( ! [X1] : (~p(s(bool,happ(s(fun(sK1,bool),X1),s(sK1,sK5)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK11(happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK4)),X1,bool,sK1))))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f1647,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl43),
  inference(cnf_transformation,[],[f1647_D])).
fof(f1647_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl43),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl43])])).
fof(f1760,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,happ(s(fun(sK1,bool),X1),s(sK1,sK5)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK11(happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK4)),X1,bool,sK1))))) )),
  inference(subsumption_resolution,[],[f1758,f454])).
fof(f454,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(sK1,sK11(happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK4)),X0,bool,sK1)))) | ~p(s(bool,happ(s(fun(sK1,bool),X0),s(sK1,sK5)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK11(happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK4)),X0,bool,sK1))))) )),
  inference(superposition,[],[f58,f45])).
fof(f58,plain,(
  ( ! [X0] : (s(bool,f(s(sK1,sK11(happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK4)),X0,bool,sK1)))) != s(bool,g(s(sK1,sK11(happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK4)),X0,bool,sK1)))) | ~p(s(bool,happ(s(fun(sK1,bool),X0),s(sK1,sK5))))) )),
  inference(superposition,[],[f33,f46])).
fof(f46,plain,(
  ( ! [X2,X0,X3,X1] : (s(fun(X0,X1),X2) = s(fun(X0,X1),X3) | s(X1,f(s(X0,sK11(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK11(X3,X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f29])).
fof(f29,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK11(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK11(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK11])],[f19])).
fof(f19,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f14])).
fof(f14,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f13])).
fof(f13,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.relation.TC_STRONG_INDUCT',aHL_EXT)).
fof(f1758,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,happ(s(fun(sK1,bool),X1),s(sK1,sK5)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK11(happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK4)),X1,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK11(happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK4)),X1,bool,sK1))))) )),
  inference(superposition,[],[f454,f45])).
fof(f1648,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl42),
  inference(cnf_transformation,[],[f1648_D])).
fof(f1648_D,plain,(
  s(bool,'T_0') = s(bool,'F_0') <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f1661,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,'F_0') | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK3,sK1)))))) ) | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f1660,f44])).
fof(f1660,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'F_0') | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK3,sK1)))))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f1659,f469])).
fof(f1659,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,'F_0') | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK3,sK1)))),s(sK1,sK8(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK3,sK1)))))) ) | $spl0),
  inference(duplicate_literal_removal,[],[f1658])).
fof(f1658,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,'F_0') | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK3,sK1)))),s(sK1,sK8(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK3,sK1)))))) ) | $spl0),
  inference(forward_demodulation,[],[f1657,f336])).
fof(f1646,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,'F_0') | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK3,sK1)))))) ) | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f1645,f44])).
fof(f1645,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'F_0') | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK3,sK1)))))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f1644,f469])).
fof(f1644,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,'F_0') | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK3,sK1)))),s(sK1,sK8(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK3,sK1)))))) ) | $spl0),
  inference(duplicate_literal_removal,[],[f1643])).
fof(f1643,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,'F_0') | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK3,sK1)))),s(sK1,sK8(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK7(sK2,sK3,sK1))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,sK8(sK2,sK3,sK1)))))) ) | $spl0),
  inference(forward_demodulation,[],[f1642,f343])).
fof(f1254,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK3,sK1)))),s(sK1,sK10(sK3,sK3,sK1))))) | $spl37),
  inference(cnf_transformation,[],[f1254_D])).
fof(f1254_D,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK3,sK1)))),s(sK1,sK10(sK3,sK3,sK1))))) <=> ~$spl37),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl37])])).
fof(f1269,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK3,sK3,sK1)))),s(sK1,sK8(sK3,sK3,sK1))))) | $spl40),
  inference(cnf_transformation,[],[f1269_D])).
fof(f1269_D,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK3,sK3,sK1)))),s(sK1,sK8(sK3,sK3,sK1))))) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f1318,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK3,sK3,sK1)))),s(sK1,sK8(sK3,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK3,sK1)))),s(sK1,sK10(sK3,sK3,sK1))))) | $spl34),
  inference(subsumption_resolution,[],[f1317,f44])).
fof(f1317,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK3,sK3,sK1)))),s(sK1,sK8(sK3,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK3,sK1)))),s(sK1,sK10(sK3,sK3,sK1))))) | $spl34),
  inference(subsumption_resolution,[],[f1304,f35])).
fof(f1304,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK3,sK3,sK1)))),s(sK1,sK8(sK3,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK3,sK1)))),s(sK1,sK10(sK3,sK3,sK1))))) | sP0(sK1,sK3,sK3) | $spl34),
  inference(superposition,[],[f116,f1270])).
fof(f1270,plain,(
  s(bool,'T_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK3,sK3,sK1)),s(sK1,sK8(sK3,sK3,sK1)))) | $spl34),
  inference(subsumption_resolution,[],[f1266,f43])).
fof(f1266,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK3,sK3,sK1)),s(sK1,sK8(sK3,sK3,sK1)))) | $spl34),
  inference(superposition,[],[f1243,f45])).
fof(f1243,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK3,sK3,sK1)),s(sK1,sK8(sK3,sK3,sK1))))) | $spl34),
  inference(resolution,[],[f1236,f37])).
fof(f116,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),X1),s(sK1,X2),s(sK1,X3)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,X3)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,X1,sK1)))),s(sK1,sK10(sK3,X1,sK1))))) | sP0(sK1,sK3,X1)) )),
  inference(resolution,[],[f30,f40])).
fof(f1256,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK3,sK3,sK1)))),s(sK1,sK7(sK3,sK3,sK1))))) | $spl38),
  inference(cnf_transformation,[],[f1256_D])).
fof(f1256_D,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK3,sK3,sK1)))),s(sK1,sK7(sK3,sK3,sK1))))) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f1289,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK3,sK3,sK1)))),s(sK1,sK7(sK3,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK3,sK1)))),s(sK1,sK10(sK3,sK3,sK1))))) | $spl34),
  inference(subsumption_resolution,[],[f1288,f44])).
fof(f1288,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK3,sK3,sK1)))),s(sK1,sK7(sK3,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK3,sK1)))),s(sK1,sK10(sK3,sK3,sK1))))) | $spl34),
  inference(subsumption_resolution,[],[f1275,f34])).
fof(f1275,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK3,sK3,sK1)))),s(sK1,sK7(sK3,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK3,sK1)))),s(sK1,sK10(sK3,sK3,sK1))))) | sP0(sK1,sK3,sK3) | $spl34),
  inference(superposition,[],[f116,f1257])).
fof(f1257,plain,(
  s(bool,'T_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK3,sK3,sK1)),s(sK1,sK7(sK3,sK3,sK1)))) | $spl34),
  inference(subsumption_resolution,[],[f1251,f43])).
fof(f1251,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK3,sK3,sK1)),s(sK1,sK7(sK3,sK3,sK1)))) | $spl34),
  inference(superposition,[],[f1242,f45])).
fof(f1267,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK3,sK3,sK1)))),s(sK1,sK8(sK3,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK3,sK1)))),s(sK1,sK10(sK3,sK3,sK1))))) | $spl34),
  inference(subsumption_resolution,[],[f1260,f35])).
fof(f1260,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK3,sK3,sK1)))),s(sK1,sK8(sK3,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK3,sK1)))),s(sK1,sK10(sK3,sK3,sK1))))) | sP0(sK1,sK3,sK3) | $spl34),
  inference(resolution,[],[f1243,f116])).
fof(f1252,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK3,sK3,sK1)))),s(sK1,sK7(sK3,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK3,sK1)))),s(sK1,sK10(sK3,sK3,sK1))))) | $spl34),
  inference(subsumption_resolution,[],[f1245,f34])).
fof(f1245,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK3,sK3,sK1)))),s(sK1,sK7(sK3,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK3,sK1)))),s(sK1,sK10(sK3,sK3,sK1))))) | sP0(sK1,sK3,sK3) | $spl34),
  inference(resolution,[],[f1242,f116])).
fof(f1231,plain,(
  ( ! [X0,X1] : (sP0(sK1,sK3,sK3) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK4),s(sK1,sK5))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,X1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0),s(sK1,X1))))) )),
  inference(duplicate_literal_removal,[],[f1222])).
fof(f1222,plain,(
  ( ! [X0,X1] : (sP0(sK1,sK3,sK3) | s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK4),s(sK1,sK5))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,X1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0),s(sK1,X1)))) | sP0(sK1,sK3,sK3)) )),
  inference(resolution,[],[f203,f40])).
fof(f1122,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f1121,f43])).
fof(f1121,plain,(
  p(s(bool,'F_0')) | $spl16),
  inference(forward_demodulation,[],[f1120,f134])).
fof(f134,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK4))),s(sK1,sK5)))),
  inference(subsumption_resolution,[],[f132,f44])).
fof(f132,plain,(
  ~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK4))),s(sK1,sK5)))),
  inference(superposition,[],[f124,f45])).
fof(f1120,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK4))),s(sK1,sK5)))) | $spl16),
  inference(subsumption_resolution,[],[f1113,f44])).
fof(f1113,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK4))),s(sK1,sK5)))) | $spl16),
  inference(superposition,[],[f606,f53])).
fof(f896,plain,(
  $false | ($spl18 | $spl25)),
  inference(subsumption_resolution,[],[f887,f771])).
fof(f887,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7(sK2,sK2,sK1)))),s(sK1,sK8(sK2,sK2,sK1))))) | $spl25),
  inference(resolution,[],[f635,f30])).
fof(f635,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK2,sK1)))),s(sK1,sK8(sK2,sK2,sK1))))) | $spl25),
  inference(cnf_transformation,[],[f635_D])).
fof(f635_D,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK2,sK1)))),s(sK1,sK8(sK2,sK2,sK1))))) <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f842,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK8(sK2,sK2,sK1))))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0),s(sK1,sK7(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK2,sK1)))),s(sK1,sK8(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK7(sK2,sK2,sK1)))))) ) | $spl18),
  inference(subsumption_resolution,[],[f828,f44])).
fof(f828,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK8(sK2,sK2,sK1))))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0),s(sK1,sK7(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK2,sK1)))),s(sK1,sK8(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK7(sK2,sK2,sK1)))))) ) | $spl18),
  inference(superposition,[],[f31,f792])).
fof(f792,plain,(
  s(bool,'T_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7(sK2,sK2,sK1)),s(sK1,sK8(sK2,sK2,sK1)))) | $spl18),
  inference(subsumption_resolution,[],[f791,f43])).
fof(f791,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7(sK2,sK2,sK1)),s(sK1,sK8(sK2,sK2,sK1)))) | $spl18),
  inference(superposition,[],[f773,f45])).
fof(f773,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7(sK2,sK2,sK1)),s(sK1,sK8(sK2,sK2,sK1))))) | $spl18),
  inference(resolution,[],[f608,f37])).
fof(f624,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK2,sK1)))),s(sK1,sK7(sK2,sK2,sK1))))) | $spl21),
  inference(cnf_transformation,[],[f624_D])).
fof(f624_D,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK2,sK1)))),s(sK1,sK7(sK2,sK2,sK1))))) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f625,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK7(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK6(sK2,sK2,sK1))))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0),s(sK1,sK6(sK2,sK2,sK1)))))) ) | $spl22),
  inference(cnf_transformation,[],[f625_D])).
fof(f625_D,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK7(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK6(sK2,sK2,sK1))))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0),s(sK1,sK6(sK2,sK2,sK1)))))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f813,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK7(sK2,sK2,sK1))))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0),s(sK1,sK6(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK2,sK1)))),s(sK1,sK7(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK6(sK2,sK2,sK1)))))) ) | $spl18),
  inference(subsumption_resolution,[],[f799,f44])).
fof(f799,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK7(sK2,sK2,sK1))))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0),s(sK1,sK6(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK2,sK1)))),s(sK1,sK7(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK6(sK2,sK2,sK1)))))) ) | $spl18),
  inference(superposition,[],[f31,f783])).
fof(f784,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK8(sK2,sK2,sK1))))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0),s(sK1,sK7(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK2,sK1)))),s(sK1,sK8(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK7(sK2,sK2,sK1)))))) ) | $spl18),
  inference(resolution,[],[f773,f31])).
fof(f775,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK7(sK2,sK2,sK1))))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0),s(sK1,sK6(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK2,sK1)))),s(sK1,sK7(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK6(sK2,sK2,sK1)))))) ) | $spl18),
  inference(resolution,[],[f772,f31])).
fof(f768,plain,(
  $false | $spl19),
  inference(subsumption_resolution,[],[f748,f607])).
fof(f607,plain,(
  ~sP0(sK1,sK2,sK2) | $spl19),
  inference(cnf_transformation,[],[f607_D])).
fof(f607_D,plain,(
  ~sP0(sK1,sK2,sK2) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f748,plain,(
  sP0(sK1,sK2,sK2)),
  inference(subsumption_resolution,[],[f747,f43])).
fof(f747,plain,(
  p(s(bool,'F_0')) | sP0(sK1,sK2,sK2)),
  inference(forward_demodulation,[],[f707,f134])).
fof(f707,plain,(
  sP0(sK1,sK2,sK2) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK4))),s(sK1,sK5))))),
  inference(duplicate_literal_removal,[],[f694])).
fof(f694,plain,(
  sP0(sK1,sK2,sK2) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK4))),s(sK1,sK5)))) | sP0(sK1,sK2,sK2) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK4))),s(sK1,sK5))))),
  inference(resolution,[],[f48,f47])).
fof(f47,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(X0,sK2,sK1)))),s(sK1,sK10(X0,sK2,sK1))))) | sP0(sK1,X0,sK2) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),X0),s(sK1,sK4))),s(sK1,sK5))))) )),
  inference(resolution,[],[f32,f39])).
fof(f48,plain,(
  ( ! [X1] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),X1),s(sK1,sK9(X1,sK2,sK1)))),s(sK1,sK10(X1,sK2,sK1))))) | sP0(sK1,X1,sK2) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),X1),s(sK1,sK4))),s(sK1,sK5))))) )),
  inference(resolution,[],[f32,f40])).
fof(f751,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK2,sK1)))),s(sK1,sK10(sK2,sK2,sK1))))) | $spl29),
  inference(cnf_transformation,[],[f751_D])).
fof(f751_D,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK2,sK1)))),s(sK1,sK10(sK2,sK2,sK1))))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f750,plain,(
  ( ! [X2,X3] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2))),s(sK1,X3)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,X3)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK2,sK1)))),s(sK1,sK10(sK2,sK2,sK1)))))) ) | $spl19),
  inference(resolution,[],[f607,f40])).
fof(f749,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0))),s(sK1,X1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0),s(sK1,X1)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK2,sK2,sK1)))),s(sK1,sK10(sK2,sK2,sK1)))))) ) | $spl19),
  inference(resolution,[],[f607,f39])).
fof(f608,plain,(
  sP0(sK1,sK2,sK2) | $spl18),
  inference(cnf_transformation,[],[f608_D])).
fof(f608_D,plain,(
  sP0(sK1,sK2,sK2) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f746,plain,(
  ( ! [X0,X1] : (sP0(sK1,sK2,sK2) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0))),s(sK1,X1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0),s(sK1,X1))))) )),
  inference(subsumption_resolution,[],[f745,f43])).
fof(f745,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | sP0(sK1,sK2,sK2) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0))),s(sK1,X1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0),s(sK1,X1))))) )),
  inference(forward_demodulation,[],[f706,f134])).
fof(f706,plain,(
  ( ! [X0,X1] : (sP0(sK1,sK2,sK2) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK4))),s(sK1,sK5)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0))),s(sK1,X1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0),s(sK1,X1))))) )),
  inference(duplicate_literal_removal,[],[f696])).
fof(f696,plain,(
  ( ! [X0,X1] : (sP0(sK1,sK2,sK2) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK4))),s(sK1,sK5)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0))),s(sK1,X1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0),s(sK1,X1)))) | sP0(sK1,sK2,sK2)) )),
  inference(resolution,[],[f48,f39])).
fof(f743,plain,(
  $false | ($spl18 | $spl21)),
  inference(subsumption_resolution,[],[f734,f611])).
fof(f611,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK2,sK2,sK1)))),s(sK1,sK7(sK2,sK2,sK1))))) | $spl18),
  inference(resolution,[],[f608,f34])).
fof(f734,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK2,sK2,sK1)))),s(sK1,sK7(sK2,sK2,sK1))))) | $spl21),
  inference(resolution,[],[f624,f30])).
fof(f683,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK8(sK2,sK2,sK1))))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0),s(sK1,sK7(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK2,sK1)))),s(sK1,sK8(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK7(sK2,sK2,sK1)))))) ) | $spl18),
  inference(subsumption_resolution,[],[f670,f44])).
fof(f670,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK8(sK2,sK2,sK1))))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0),s(sK1,sK7(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK2,sK1)))),s(sK1,sK8(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK7(sK2,sK2,sK1)))))) ) | $spl18),
  inference(superposition,[],[f31,f637])).
fof(f637,plain,(
  s(bool,'T_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7(sK2,sK2,sK1)),s(sK1,sK8(sK2,sK2,sK1)))) | $spl18),
  inference(subsumption_resolution,[],[f633,f43])).
fof(f633,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7(sK2,sK2,sK1)),s(sK1,sK8(sK2,sK2,sK1)))) | $spl18),
  inference(superposition,[],[f614,f45])).
fof(f614,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7(sK2,sK2,sK1)),s(sK1,sK8(sK2,sK2,sK1))))) | $spl18),
  inference(resolution,[],[f608,f37])).
fof(f655,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK7(sK2,sK2,sK1))))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0),s(sK1,sK6(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK2,sK1)))),s(sK1,sK7(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK6(sK2,sK2,sK1)))))) ) | $spl18),
  inference(subsumption_resolution,[],[f642,f44])).
fof(f642,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK7(sK2,sK2,sK1))))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0),s(sK1,sK6(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK2,sK1)))),s(sK1,sK7(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK6(sK2,sK2,sK1)))))) ) | $spl18),
  inference(superposition,[],[f31,f626])).
fof(f626,plain,(
  s(bool,'T_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK2,sK2,sK1)),s(sK1,sK7(sK2,sK2,sK1)))) | $spl18),
  inference(subsumption_resolution,[],[f622,f43])).
fof(f622,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK2,sK2,sK1)),s(sK1,sK7(sK2,sK2,sK1)))) | $spl18),
  inference(superposition,[],[f613,f45])).
fof(f613,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK2,sK2,sK1)),s(sK1,sK7(sK2,sK2,sK1))))) | $spl18),
  inference(resolution,[],[f608,f36])).
fof(f627,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK8(sK2,sK2,sK1))))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0),s(sK1,sK7(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK2,sK1)))),s(sK1,sK8(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK7(sK2,sK2,sK1)))))) ) | $spl18),
  inference(resolution,[],[f614,f31])).
fof(f616,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK7(sK2,sK2,sK1))))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0),s(sK1,sK6(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK2,sK1)))),s(sK1,sK7(sK2,sK2,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK6(sK2,sK2,sK1)))))) ) | $spl18),
  inference(resolution,[],[f613,f31])).
fof(f605,plain,(
  ( ! [X2,X1] : (sP0(sK1,sK2,sK2) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X1))),s(sK1,X2)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X1),s(sK1,X2))))) )),
  inference(subsumption_resolution,[],[f604,f43])).
fof(f604,plain,(
  ( ! [X2,X1] : (p(s(bool,'F_0')) | sP0(sK1,sK2,sK2) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X1))),s(sK1,X2)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X1),s(sK1,X2))))) )),
  inference(forward_demodulation,[],[f603,f134])).
fof(f603,plain,(
  ( ! [X2,X1] : (sP0(sK1,sK2,sK2) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK4))),s(sK1,sK5)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X1))),s(sK1,X2)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X1),s(sK1,X2))))) )),
  inference(duplicate_literal_removal,[],[f595])).
fof(f595,plain,(
  ( ! [X2,X1] : (sP0(sK1,sK2,sK2) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK4))),s(sK1,sK5)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X1))),s(sK1,X2)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X1),s(sK1,X2)))) | sP0(sK1,sK2,sK2)) )),
  inference(resolution,[],[f47,f40])).
fof(f480,plain,(
  $false | ($spl2 | $spl9)),
  inference(subsumption_resolution,[],[f229,f357])).
fof(f229,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK3,sK1)))),s(sK1,sK8(sK2,sK3,sK1))))) | $spl9),
  inference(cnf_transformation,[],[f229_D])).
fof(f229_D,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK3,sK1)))),s(sK1,sK8(sK2,sK3,sK1))))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f424,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK8(sK2,sK3,sK1))))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK3,sK1)))),s(sK1,sK8(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK7(sK2,sK3,sK1)))))) ) | $spl2),
  inference(subsumption_resolution,[],[f411,f44])).
fof(f411,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK8(sK2,sK3,sK1))))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK3,sK1)))),s(sK1,sK8(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK7(sK2,sK3,sK1)))))) ) | $spl2),
  inference(superposition,[],[f31,f378])).
fof(f378,plain,(
  s(bool,'T_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1)))) | $spl2),
  inference(subsumption_resolution,[],[f377,f43])).
fof(f377,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1)))) | $spl2),
  inference(superposition,[],[f359,f45])).
fof(f359,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1))))) | $spl2),
  inference(resolution,[],[f202,f37])).
fof(f218,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK3,sK1)))),s(sK1,sK7(sK2,sK3,sK1))))) | $spl5),
  inference(cnf_transformation,[],[f218_D])).
fof(f218_D,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK3,sK1)))),s(sK1,sK7(sK2,sK3,sK1))))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f219,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK6(sK2,sK3,sK1))))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0),s(sK1,sK6(sK2,sK3,sK1)))))) ) | $spl6),
  inference(cnf_transformation,[],[f219_D])).
fof(f219_D,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK6(sK2,sK3,sK1))))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0),s(sK1,sK6(sK2,sK3,sK1)))))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f396,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0),s(sK1,sK6(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK3,sK1)))),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK6(sK2,sK3,sK1)))))) ) | $spl2),
  inference(subsumption_resolution,[],[f383,f44])).
fof(f383,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0),s(sK1,sK6(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK3,sK1)))),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK6(sK2,sK3,sK1)))))) ) | $spl2),
  inference(superposition,[],[f31,f368])).
fof(f371,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK8(sK2,sK3,sK1))))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK3,sK1)))),s(sK1,sK8(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK7(sK2,sK3,sK1)))))) ) | $spl2),
  inference(resolution,[],[f359,f31])).
fof(f361,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0),s(sK1,sK6(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK3,sK1)))),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK6(sK2,sK3,sK1)))))) ) | $spl2),
  inference(resolution,[],[f358,f31])).
fof(f353,plain,(
  $false | ($spl0 | $spl13)),
  inference(subsumption_resolution,[],[f344,f200])).
fof(f344,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1))))) | $spl13),
  inference(resolution,[],[f330,f30])).
fof(f334,plain,(
  ( ! [X2,X3] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,X3)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2),s(sK1,X3)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1)))))) ) | $spl3),
  inference(resolution,[],[f201,f40])).
fof(f201,plain,(
  ~sP0(sK1,sK3,sK2) | $spl3),
  inference(cnf_transformation,[],[f201_D])).
fof(f201_D,plain,(
  ~sP0(sK1,sK3,sK2) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f333,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,X1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0),s(sK1,X1)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1)))))) ) | $spl3),
  inference(resolution,[],[f201,f39])).
fof(f202,plain,(
  sP0(sK1,sK3,sK2) | $spl2),
  inference(cnf_transformation,[],[f202_D])).
fof(f202_D,plain,(
  sP0(sK1,sK3,sK2) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f328,plain,(
  sP0(sK1,sK3,sK2) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1)))))),
  inference(subsumption_resolution,[],[f290,f44])).
fof(f290,plain,(
  ~p(s(bool,'T_0')) | sP0(sK1,sK3,sK2) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1)))))),
  inference(superposition,[],[f55,f53])).
fof(f327,plain,(
  $false | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f218,f205])).
fof(f205,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK3,sK1)))),s(sK1,sK7(sK2,sK3,sK1))))) | $spl2),
  inference(resolution,[],[f202,f34])).
fof(f277,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK8(sK2,sK3,sK1))))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK3,sK1)))),s(sK1,sK8(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK7(sK2,sK3,sK1)))))) ) | $spl2),
  inference(subsumption_resolution,[],[f264,f44])).
fof(f264,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK8(sK2,sK3,sK1))))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK3,sK1)))),s(sK1,sK8(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK7(sK2,sK3,sK1)))))) ) | $spl2),
  inference(superposition,[],[f31,f231])).
fof(f231,plain,(
  s(bool,'T_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1)))) | $spl2),
  inference(subsumption_resolution,[],[f227,f43])).
fof(f227,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1)))) | $spl2),
  inference(superposition,[],[f208,f45])).
fof(f208,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1))))) | $spl2),
  inference(resolution,[],[f202,f37])).
fof(f249,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0),s(sK1,sK6(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK3,sK1)))),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK6(sK2,sK3,sK1)))))) ) | $spl2),
  inference(subsumption_resolution,[],[f236,f44])).
fof(f236,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0),s(sK1,sK6(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK3,sK1)))),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK6(sK2,sK3,sK1)))))) ) | $spl2),
  inference(superposition,[],[f31,f220])).
fof(f220,plain,(
  s(bool,'T_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK7(sK2,sK3,sK1)))) | $spl2),
  inference(subsumption_resolution,[],[f216,f43])).
fof(f216,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK7(sK2,sK3,sK1)))) | $spl2),
  inference(superposition,[],[f207,f45])).
fof(f207,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(sK2,sK3,sK1)),s(sK1,sK7(sK2,sK3,sK1))))) | $spl2),
  inference(resolution,[],[f202,f36])).
fof(f221,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK8(sK2,sK3,sK1))))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK3,sK1)))),s(sK1,sK8(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK7(sK2,sK3,sK1)))))) ) | $spl2),
  inference(resolution,[],[f208,f31])).
fof(f210,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0),s(sK1,sK6(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK6(sK2,sK3,sK1)))),s(sK1,sK7(sK2,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,sK6(sK2,sK3,sK1)))))) ) | $spl2),
  inference(resolution,[],[f207,f31])).
fof(f198,plain,(
  sP0(sK1,sK3,sK2) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1)))))),
  inference(subsumption_resolution,[],[f195,f44])).
fof(f195,plain,(
  ~p(s(bool,'T_0')) | sP0(sK1,sK3,sK2) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1)))))),
  inference(superposition,[],[f54,f53])).
% SZS output end Proof for thm.relation
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 547
% Passive clauses: 2394
% Generated clauses: 5875
% Final active clauses: 409
% Final passive clauses: 1419
% Input formulas: 6
% Initial clauses: 17
% 
% Selected by SInE selection: 6
% SInE iterations: 1
% 
% Duplicate literals: 386
% Trivial inequalities: 20
% Fw subsumption resolutions: 441
% Fw demodulations: 429
% Bw demodulations: 32
% 
% Simple tautologies: 424
% Forward subsumptions: 1748
% 
% Binary resolution: 1233
% Forward superposition: 2307
% Backward superposition: 946
% Equality resolution: 8
% 
% Split clauses: 86
% Split components: 205
% SAT solver clauses: 96
% SAT solver unit clauses: 3
% SAT solver binary clauses: 12
% 
% Sat splits: 86
% Sat splitting refutations: 9
% 
% Memory used [KB]: 4605
% Time elapsed: 0.301 s
% ------------------------------
% ------------------------------
% Success in time 11.94 s
