Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.relation
% (107602)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 961
% Passive clauses: 7542
% Generated clauses: 25871
% Final active clauses: 899
% Final passive clauses: 6042
% Input formulas: 61
% Initial clauses: 220
% 
% Duplicate literals: 752
% Trivial inequalities: 88
% Fw subsumption resolutions: 2978
% Fw demodulations: 2142
% Bw demodulations: 58
% 
% Simple tautologies: 306
% Equational tautologies: 48
% Forward subsumptions: 11487
% 
% Binary resolution: 8524
% Factoring: 1432
% Forward superposition: 4363
% Backward superposition: 4967
% Self superposition: 37
% Equality factoring: 30
% Equality resolution: 12
% 
% Split clauses: 572
% Split components: 1152
% SAT solver clauses: 977
% SAT solver unit clauses: 18
% SAT solver binary clauses: 162
% 
% Sat splits: 572
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 975
% TWLsolver calls for satisfiability: 439
% 
% Memory used [KB]: 11385
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.relation
% (107634)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 398
% Passive clauses: 1973
% Generated clauses: 13614
% Final active clauses: 387
% Final passive clauses: 1557
% Input formulas: 61
% Initial clauses: 220
% 
% Duplicate literals: 296
% Trivial inequalities: 96
% Fw subsumption resolutions: 1185
% Fw demodulations: 1631
% Bw demodulations: 26
% Condensations: 358
% Global subsumptions: 22
% 
% Simple tautologies: 326
% Equational tautologies: 21
% Forward subsumptions: 4994
% 
% Binary resolution: 2469
% Factoring: 133
% Forward superposition: 2886
% Backward superposition: 2221
% Self superposition: 43
% Equality factoring: 32
% Equality resolution: 11
% 
% Split clauses: 75
% Split components: 155
% SAT solver clauses: 2955
% SAT solver unit clauses: 222
% SAT solver binary clauses: 83
% SAT solver learnt clauses: 44
% SAT solver learnt literals: 10
% 
% Sat splits: 75
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 2867
% TWLsolver calls for satisfiability: 1745
% 
% Memory used [KB]: 13688
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.relation
% (107658)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 2543
% Passive clauses: 10244
% Generated clauses: 200249
% Final active clauses: 1726
% Final passive clauses: 4006
% Input formulas: 61
% Initial clauses: 220
% 
% Duplicate literals: 9504
% Trivial inequalities: 263
% Fw subsumption resolutions: 7871
% Fw demodulations: 27933
% Bw demodulations: 2026
% Condensations: 18577
% Global subsumptions: 781
% 
% Simple tautologies: 309
% Equational tautologies: 1712
% Forward subsumptions: 118812
% Fw demodulations to eq. taut.: 3
% Bw demodulations to eq. taut.: 3
% 
% Unit resulting resolution: 38468
% Factoring: 1
% Forward superposition: 54040
% Backward superposition: 36088
% Self superposition: 217
% Equality factoring: 177
% Equality resolution: 5
% 
% Split clauses: 427
% Split components: 860
% SAT solver clauses: 14523
% SAT solver unit clauses: 1824
% SAT solver binary clauses: 3564
% 
% Sat splits: 427
% Sat splitting refutations: 2641
% 
% Memory used [KB]: 103495
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.relation
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.relation
% SZS output start Proof for thm.relation
fof(f37463,plain,(
  $false),
  inference(subsumption_resolution,[],[f37462,f515])).
fof(f515,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f8])).
fof(f8,axiom,(
  p(s(bool,'T_0'))),
  file('thm.relation.TFL_INDUCTIVE_INVARIANT_ON_WFREC','thm.bool.TRUTH')).
fof(f37462,plain,(
  ~p(s(bool,'T_0'))),
  inference(forward_demodulation,[],[f37461,f721])).
fof(f721,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK13,bool),sK16),s(sK13,sK19)))),
  inference(resolution,[],[f305,f495])).
fof(f495,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f297])).
fof(f297,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f128])).
fof(f128,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X14] : (s(bool,'T_0') = s(bool,X14) <=> p(s(bool,X14)))),
  file('thm.relation.TFL_INDUCTIVE_INVARIANT_ON_WFREC','thm.bool.EQ_CLAUSES|split|0')).
fof(f305,plain,(
  p(s(bool,happ(s(fun(sK13,bool),sK16),s(sK13,sK19))))),
  inference(cnf_transformation,[],[f203])).
fof(f203,plain,(
  s(fun(sK13,sK12),sK14) = s(fun(sK13,sK12),'const.relation.WFREC_2'(s(fun(sK13,fun(sK13,bool)),sK15),s(fun(fun(sK13,sK12),fun(sK13,sK12)),sK18))) & p(s(bool,'const.relation.WF_1'(s(fun(sK13,fun(sK13,bool)),sK15)))) & p(s(bool,'const.relation.INDUCTIVE_INVARIANT_ON_4'(s(fun(sK13,fun(sK13,bool)),sK15),s(fun(sK13,bool),sK16),s(fun(sK13,fun(sK12,bool)),sK17),s(fun(fun(sK13,sK12),fun(sK13,sK12)),sK18)))) & p(s(bool,happ(s(fun(sK13,bool),sK16),s(sK13,sK19)))) & ~p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK13,fun(sK12,bool)),sK17),s(sK13,sK19))),s(sK12,happ(s(fun(sK13,sK12),sK14),s(sK13,sK19))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13,sK14,sK15,sK16,sK17,sK18,sK19])],[f150])).
fof(f150,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7] : (s(fun(X1,X0),X2) = s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X6))) & p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X3)))) & p(s(bool,'const.relation.INDUCTIVE_INVARIANT_ON_4'(s(fun(X1,fun(X1,bool)),X3),s(fun(X1,bool),X4),s(fun(X1,fun(X0,bool)),X5),s(fun(fun(X1,X0),fun(X1,X0)),X6)))) & p(s(bool,happ(s(fun(X1,bool),X4),s(X1,X7)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X5),s(X1,X7))),s(X0,happ(s(fun(X1,X0),X2),s(X1,X7)))))))),
  inference(flattening,[],[f149])).
fof(f149,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7] : ((s(fun(X1,X0),X2) = s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X6))) & p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X3)))) & p(s(bool,'const.relation.INDUCTIVE_INVARIANT_ON_4'(s(fun(X1,fun(X1,bool)),X3),s(fun(X1,bool),X4),s(fun(X1,fun(X0,bool)),X5),s(fun(fun(X1,X0),fun(X1,X0)),X6)))) & p(s(bool,happ(s(fun(X1,bool),X4),s(X1,X7))))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X5),s(X1,X7))),s(X0,happ(s(fun(X1,X0),X2),s(X1,X7)))))))),
  inference(ennf_transformation,[],[f64])).
fof(f64,plain,(
  ~! [X0,X1,X2,X3,X4,X5,X6,X7] : ((s(fun(X1,X0),X2) = s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X6))) & p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X3)))) & p(s(bool,'const.relation.INDUCTIVE_INVARIANT_ON_4'(s(fun(X1,fun(X1,bool)),X3),s(fun(X1,bool),X4),s(fun(X1,fun(X0,bool)),X5),s(fun(fun(X1,X0),fun(X1,X0)),X6)))) & p(s(bool,happ(s(fun(X1,bool),X4),s(X1,X7))))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X5),s(X1,X7))),s(X0,happ(s(fun(X1,X0),X2),s(X1,X7)))))))),
  inference(flattening,[],[f63])).
fof(f63,plain,(
  ~! [X0,X1] : ! [X2,X3,X4,X5,X6,X7] : ((s(fun(X1,X0),X2) = s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X6))) & p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X3)))) & p(s(bool,'const.relation.INDUCTIVE_INVARIANT_ON_4'(s(fun(X1,fun(X1,bool)),X3),s(fun(X1,bool),X4),s(fun(X1,fun(X0,bool)),X5),s(fun(fun(X1,X0),fun(X1,X0)),X6)))) & p(s(bool,happ(s(fun(X1,bool),X4),s(X1,X7))))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X5),s(X1,X7))),s(X0,happ(s(fun(X1,X0),X2),s(X1,X7)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10,X11,X12,X13] : ((s(fun(X7,X6),X8) = s(fun(X7,X6),'const.relation.WFREC_2'(s(fun(X7,fun(X7,bool)),X9),s(fun(fun(X7,X6),fun(X7,X6)),X12))) & p(s(bool,'const.relation.WF_1'(s(fun(X7,fun(X7,bool)),X9)))) & p(s(bool,'const.relation.INDUCTIVE_INVARIANT_ON_4'(s(fun(X7,fun(X7,bool)),X9),s(fun(X7,bool),X10),s(fun(X7,fun(X6,bool)),X11),s(fun(fun(X7,X6),fun(X7,X6)),X12)))) & p(s(bool,happ(s(fun(X7,bool),X10),s(X7,X13))))) => p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X7,fun(X6,bool)),X11),s(X7,X13))),s(X6,happ(s(fun(X7,X6),X8),s(X7,X13)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10,X11,X12,X13] : ((s(fun(X7,X6),X8) = s(fun(X7,X6),'const.relation.WFREC_2'(s(fun(X7,fun(X7,bool)),X9),s(fun(fun(X7,X6),fun(X7,X6)),X12))) & p(s(bool,'const.relation.WF_1'(s(fun(X7,fun(X7,bool)),X9)))) & p(s(bool,'const.relation.INDUCTIVE_INVARIANT_ON_4'(s(fun(X7,fun(X7,bool)),X9),s(fun(X7,bool),X10),s(fun(X7,fun(X6,bool)),X11),s(fun(fun(X7,X6),fun(X7,X6)),X12)))) & p(s(bool,happ(s(fun(X7,bool),X10),s(X7,X13))))) => p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X7,fun(X6,bool)),X11),s(X7,X13))),s(X6,happ(s(fun(X7,X6),X8),s(X7,X13)))))))),
  file('thm.relation.TFL_INDUCTIVE_INVARIANT_ON_WFREC','thm.relation.TFL_INDUCTIVE_INVARIANT_ON_WFREC')).
fof(f37461,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),sK16),s(sK13,sK19))))),
  inference(subsumption_resolution,[],[f37438,f505])).
fof(f505,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f181])).
fof(f181,plain,(
  ! [X0] : ((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f137])).
fof(f137,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) => ~p(s(bool,X0)))),
  inference(flattening,[],[f136])).
fof(f136,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) => ~p(s(bool,X0)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X14] : ((p(s(bool,X14)) => p(s(bool,'F_0'))) => ~p(s(bool,X14)))),
  file('thm.relation.TFL_INDUCTIVE_INVARIANT_ON_WFREC','thm.bool.IMP_F')).
fof(f37438,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,happ(s(fun(sK13,bool),sK16),s(sK13,sK19))))),
  inference(superposition,[],[f37023,f886])).
fof(f886,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK13,fun(sK12,bool)),sK17),s(sK13,sK19))),s(sK12,happ(s(fun(sK13,sK12),sK14),s(sK13,sK19)))))),
  inference(resolution,[],[f306,f491])).
fof(f491,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f295])).
fof(f295,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f124])).
fof(f124,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X14] : (s(bool,'F_0') = s(bool,X14) <=> ~p(s(bool,X14)))),
  file('thm.relation.TFL_INDUCTIVE_INVARIANT_ON_WFREC','thm.bool.EQ_CLAUSES|split|2')).
fof(f306,plain,(
  ~p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK13,fun(sK12,bool)),sK17),s(sK13,sK19))),s(sK12,happ(s(fun(sK13,sK12),sK14),s(sK13,sK19))))))),
  inference(cnf_transformation,[],[f203])).
fof(f37023,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK13,fun(sK12,bool)),sK17),s(sK13,X0))),s(sK12,happ(s(fun(sK13,sK12),sK14),s(sK13,X0)))))) | ~p(s(bool,happ(s(fun(sK13,bool),sK16),s(sK13,X0))))) )),
  inference(forward_demodulation,[],[f1053,f763])).
fof(f763,plain,(
  ( ! [X0] : (s(sK12,happ(s(fun(sK13,sK12),sK14),s(sK13,X0))) = s(sK12,'const.relation.WFREC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(fun(fun(sK13,sK12),fun(sK13,sK12)),sK18),s(sK13,X0)))) )),
  inference(superposition,[],[f308,f302])).
fof(f302,plain,(
  s(fun(sK13,sK12),sK14) = s(fun(sK13,sK12),'const.relation.WFREC_2'(s(fun(sK13,fun(sK13,bool)),sK15),s(fun(fun(sK13,sK12),fun(sK13,sK12)),sK18)))),
  inference(cnf_transformation,[],[f203])).
fof(f308,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.relation.WFREC_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X3),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X3))),s(X1,X4)))) )),
  inference(cnf_transformation,[],[f68])).
fof(f68,plain,(
  ! [X0,X1,X2,X3,X4] : s(X0,'const.relation.WFREC_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X3),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X3))),s(X1,X4)))),
  inference(flattening,[],[f67])).
fof(f67,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : s(X0,'const.relation.WFREC_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X3),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X3))),s(X1,X4)))),
  inference(rectify,[],[f62])).
fof(f62,axiom,(
  ! [X6,X7] : ! [X22,X23,X24] : s(X6,'const.relation.WFREC_3'(s(fun(X7,fun(X7,bool)),X22),s(fun(fun(X7,X6),fun(X7,X6)),X23),s(X7,X24))) = s(X6,happ(s(fun(X7,X6),'const.relation.WFREC_2'(s(fun(X7,fun(X7,bool)),X22),s(fun(fun(X7,X6),fun(X7,X6)),X23))),s(X7,X24)))),
  file('thm.relation.TFL_INDUCTIVE_INVARIANT_ON_WFREC','ARITY_2_3_const.relation.WFREC')).
fof(f1053,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK13,fun(sK12,bool)),sK17),s(sK13,X0))),s(sK12,'const.relation.WFREC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(fun(fun(sK13,sK12),fun(sK13,sK12)),sK18),s(sK13,X0)))))) | ~p(s(bool,happ(s(fun(sK13,bool),sK16),s(sK13,X0))))) )),
  inference(subsumption_resolution,[],[f947,f303])).
fof(f303,plain,(
  p(s(bool,'const.relation.WF_1'(s(fun(sK13,fun(sK13,bool)),sK15))))),
  inference(cnf_transformation,[],[f203])).
fof(f947,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK13,fun(sK12,bool)),sK17),s(sK13,X0))),s(sK12,'const.relation.WFREC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(fun(fun(sK13,sK12),fun(sK13,sK12)),sK18),s(sK13,X0)))))) | ~p(s(bool,happ(s(fun(sK13,bool),sK16),s(sK13,X0)))) | ~p(s(bool,'const.relation.WF_1'(s(fun(sK13,fun(sK13,bool)),sK15))))) )),
  inference(resolution,[],[f304,f307])).
fof(f307,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X3),s(X1,X6))),s(X0,'const.relation.WFREC_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X4),s(X1,X6)))))) | ~p(s(bool,happ(s(fun(X1,bool),X5),s(X1,X6)))) | ~p(s(bool,'const.relation.INDUCTIVE_INVARIANT_ON_4'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,bool),X5),s(fun(X1,fun(X0,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X4)))) | ~p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X2))))) )),
  inference(cnf_transformation,[],[f152])).
fof(f152,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : (~p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X2)))) | ~p(s(bool,'const.relation.INDUCTIVE_INVARIANT_ON_4'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,bool),X5),s(fun(X1,fun(X0,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X4)))) | ~p(s(bool,happ(s(fun(X1,bool),X5),s(X1,X6)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X3),s(X1,X6))),s(X0,'const.relation.WFREC_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X4),s(X1,X6)))))))),
  inference(flattening,[],[f151])).
fof(f151,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ((~p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X2)))) | ~p(s(bool,'const.relation.INDUCTIVE_INVARIANT_ON_4'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,bool),X5),s(fun(X1,fun(X0,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X4)))) | ~p(s(bool,happ(s(fun(X1,bool),X5),s(X1,X6))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X3),s(X1,X6))),s(X0,'const.relation.WFREC_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X4),s(X1,X6)))))))),
  inference(ennf_transformation,[],[f66])).
fof(f66,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ((p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X2)))) & p(s(bool,'const.relation.INDUCTIVE_INVARIANT_ON_4'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,bool),X5),s(fun(X1,fun(X0,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X4)))) & p(s(bool,happ(s(fun(X1,bool),X5),s(X1,X6))))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X3),s(X1,X6))),s(X0,'const.relation.WFREC_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X4),s(X1,X6)))))))),
  inference(flattening,[],[f65])).
fof(f65,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6] : ((p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X2)))) & p(s(bool,'const.relation.INDUCTIVE_INVARIANT_ON_4'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,bool),X5),s(fun(X1,fun(X0,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X4)))) & p(s(bool,happ(s(fun(X1,bool),X5),s(X1,X6))))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X3),s(X1,X6))),s(X0,'const.relation.WFREC_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X4),s(X1,X6)))))))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X6,X7] : ! [X9,X11,X12,X10,X13] : ((p(s(bool,'const.relation.WF_1'(s(fun(X7,fun(X7,bool)),X9)))) & p(s(bool,'const.relation.INDUCTIVE_INVARIANT_ON_4'(s(fun(X7,fun(X7,bool)),X9),s(fun(X7,bool),X10),s(fun(X7,fun(X6,bool)),X11),s(fun(fun(X7,X6),fun(X7,X6)),X12)))) & p(s(bool,happ(s(fun(X7,bool),X10),s(X7,X13))))) => p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X7,fun(X6,bool)),X11),s(X7,X13))),s(X6,'const.relation.WFREC_3'(s(fun(X7,fun(X7,bool)),X9),s(fun(fun(X7,X6),fun(X7,X6)),X12),s(X7,X13)))))))),
  file('thm.relation.TFL_INDUCTIVE_INVARIANT_ON_WFREC','thm.relation.INDUCTIVE_INVARIANT_ON_WFREC')).
fof(f304,plain,(
  p(s(bool,'const.relation.INDUCTIVE_INVARIANT_ON_4'(s(fun(sK13,fun(sK13,bool)),sK15),s(fun(sK13,bool),sK16),s(fun(sK13,fun(sK12,bool)),sK17),s(fun(fun(sK13,sK12),fun(sK13,sK12)),sK18))))),
  inference(cnf_transformation,[],[f203])).
% SZS output end Proof for thm.relation
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 911
% Passive clauses: 7537
% Generated clauses: 36876
% Final active clauses: 896
% Final passive clauses: 6233
% Input formulas: 61
% Initial clauses: 218
% 
% Selected by SInE selection: 59
% SInE iterations: 1
% 
% Duplicate literals: 1488
% Trivial inequalities: 90
% Fw subsumption resolutions: 4167
% Fw demodulations: 1888
% Bw demodulations: 7
% 
% Simple tautologies: 2151
% Equational tautologies: 110
% Forward subsumptions: 17767
% 
% Binary resolution: 25541
% Factoring: 117
% Forward superposition: 2007
% Backward superposition: 1339
% Self superposition: 2
% Equality factoring: 2
% Equality resolution: 11
% 
% Split clauses: 2171
% Split components: 5161
% SAT solver clauses: 2418
% SAT solver unit clauses: 27
% SAT solver binary clauses: 115
% 
% Sat splits: 2171
% Sat splitting refutations: 34
% 
% Memory used [KB]: 10106
% Time elapsed: 0.562 s
% ------------------------------
% ------------------------------
% Success in time 12.289 s
