Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 44
dis+11_7_41 on thm.extreal
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.extreal
% SZS output start Proof for thm.extreal
fof(f51224,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f51223,f2161,f2944,f2942,f1947,f2943])).
fof(f2943,plain,(
  ~p(s(bool,'F_0')) | $spl1),
  inference(cnf_transformation,[],[f2943_D])).
fof(f2943_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f1947,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f523])).
fof(f523,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.extreal.lt_infty|split|3',aHL_FALSITY)).
fof(f2942,plain,(
  p(s(bool,'F_0')) | $spl0),
  inference(cnf_transformation,[],[f2942_D])).
fof(f2942_D,plain,(
  p(s(bool,'F_0')) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f2944,plain,(
  p(s(bool,'T_0')) | $spl2),
  inference(cnf_transformation,[],[f2944_D])).
fof(f2944_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f2161,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f1532])).
fof(f1532,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f1531])).
fof(f1531,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f1278])).
fof(f1278,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f683])).
fof(f683,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f54])).
fof(f54,axiom,(
  ! [X15] : ((p(s(bool,'F_0')) => p(s(bool,X15))) <=> p(s(bool,'T_0')))),
  file('thm.extreal.lt_infty|split|3','thm.bool.IMP_CLAUSES|split|2')).
fof(f51223,plain,(
  $false | $spl2),
  inference(subsumption_resolution,[],[f51220,f2944])).
fof(f51220,plain,(
  ~p(s(bool,'T_0')) | $spl2),
  inference(superposition,[],[f50996,f3327])).
fof(f3327,plain,(
  s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal','const.extreal.NegInf_0')))),
  inference(resolution,[],[f2102,f1908])).
fof(f1908,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal','const.extreal.NegInf_0'))))),
  inference(cnf_transformation,[],[f1461])).
fof(f1461,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal','const.extreal.NegInf_0'))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK17])],[f1235])).
fof(f1235,plain,(
  ? [X0] : p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.NegInf_0'))))),
  inference(ennf_transformation,[],[f520])).
fof(f520,plain,(
  ~! [X0] : ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.NegInf_0'))))),
  inference(flattening,[],[f519])).
fof(f519,plain,(
  ~! [X0] : ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.NegInf_0'))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal','const.extreal.NegInf_0'))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal','const.extreal.NegInf_0'))))),
  file('thm.extreal.lt_infty|split|3','thm.extreal.lt_infty|split|3')).
fof(f2102,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f1496])).
fof(f1496,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f655])).
fof(f655,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X15] : (s(bool,'T_0') = s(bool,X15) <=> p(s(bool,X15)))),
  file('thm.extreal.lt_infty|split|3','thm.bool.EQ_CLAUSES|split|1')).
fof(f50996,plain,(
  ( ! [X8] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X8),s('type.extreal.extreal','const.extreal.NegInf_0'))))) ) | $spl2),
  inference(forward_demodulation,[],[f50995,f1929])).
fof(f1929,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))),
  inference(cnf_transformation,[],[f112])).
fof(f112,axiom,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))),
  file('thm.extreal.lt_infty|split|3','thm.extreal.extreal_mul_def|split|2')).
fof(f50995,plain,(
  ( ! [X8] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X8),s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))))))) ) | $spl2),
  inference(subsumption_resolution,[],[f50994,f2944])).
fof(f50994,plain,(
  ( ! [X8] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X8),s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))))))) )),
  inference(forward_demodulation,[],[f50888,f1976])).
fof(f1976,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)))) )),
  inference(cnf_transformation,[],[f552])).
fof(f552,plain,(
  ! [X0] : s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)))),
  inference(rectify,[],[f82])).
fof(f82,axiom,(
  ! [X16] : s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X16)))),
  file('thm.extreal.lt_infty|split|3','thm.extreal.extreal_le_def|split|1')).
fof(f50888,plain,(
  ( ! [X8] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X8)))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X8),s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))))))) )),
  inference(superposition,[],[f2553,f1929])).
fof(f2553,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))) )),
  inference(cnf_transformation,[],[f1739])).
fof(f1739,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0))))) & (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))))),
  inference(nnf_transformation,[],[f1044])).
fof(f1044,plain,(
  ! [X0,X1] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) <=> ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))))),
  inference(flattening,[],[f1043])).
fof(f1043,plain,(
  ! [X0,X1] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) <=> ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6,X7] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X7)))) <=> ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X7),s('type.extreal.extreal',X6)))))),
  file('thm.extreal.lt_infty|split|3','thm.extreal.extreal_lt_def')).
% SZS output end Proof for thm.extreal
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 1378
% Passive clauses: 8906
% Generated clauses: 49287
% Final active clauses: 1281
% Final passive clauses: 6982
% Input formulas: 517
% Initial clauses: 1023
% 
% Duplicate literals: 212
% Trivial inequalities: 93
% Fw subsumption resolutions: 2467
% Fw demodulations: 4250
% Bw demodulations: 51
% 
% Simple tautologies: 421
% Equational tautologies: 20
% Forward subsumptions: 32823
% 
% Binary resolution: 7518
% Factoring: 450
% Forward superposition: 16145
% Backward superposition: 16815
% Self superposition: 58
% Equality factoring: 2
% Equality resolution: 11
% 
% Split clauses: 400
% Split components: 802
% SAT solver clauses: 608
% SAT solver unit clauses: 47
% SAT solver binary clauses: 98
% SAT solver learnt clauses: 1
% 
% Sat splits: 400
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 606
% TWLsolver calls for satisfiability: 454
% 
% Memory used [KB]: 28656
% Time elapsed: 0.785 s
% ------------------------------
% ------------------------------
% Success in time 0.81 s
