Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.pred_set
% (43674)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1039
% Passive clauses: 10044
% Generated clauses: 32159
% Final active clauses: 1004
% Final passive clauses: 8646
% Input formulas: 260
% Initial clauses: 562
% 
% Duplicate literals: 201
% Trivial inequalities: 62
% Fw subsumption resolutions: 1795
% Fw demodulations: 2994
% Bw demodulations: 14
% 
% Simple tautologies: 279
% Equational tautologies: 91
% Forward subsumptions: 16045
% 
% Binary resolution: 11797
% Factoring: 543
% Forward superposition: 6421
% Backward superposition: 7569
% Self superposition: 34
% Equality factoring: 2
% Equality resolution: 12
% 
% Split clauses: 663
% Split components: 1328
% SAT solver clauses: 943
% SAT solver unit clauses: 12
% SAT solver binary clauses: 97
% 
% Sat splits: 663
% Sat splitting refutations: 8
% 
% TWLsolver clauses: 943
% TWLsolver calls for satisfiability: 319
% 
% Memory used [KB]: 18677
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.pred_set
% (43683)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 504
% Passive clauses: 2859
% Generated clauses: 20673
% Final active clauses: 434
% Final passive clauses: 2027
% Input formulas: 260
% Initial clauses: 562
% 
% Duplicate literals: 109
% Trivial inequalities: 27
% Fw subsumption resolutions: 452
% Fw demodulations: 3898
% Bw demodulations: 368
% Condensations: 95
% Global subsumptions: 23
% 
% Simple tautologies: 311
% Equational tautologies: 8
% Forward subsumptions: 9638
% Bw demodulations to eq. taut.: 29
% 
% Binary resolution: 2515
% Factoring: 25
% Forward superposition: 6667
% Backward superposition: 2997
% Self superposition: 41
% Equality factoring: 2
% Equality resolution: 5
% 
% Split clauses: 33
% Split components: 66
% SAT solver clauses: 3827
% SAT solver unit clauses: 641
% SAT solver binary clauses: 195
% SAT solver learnt clauses: 46
% SAT solver learnt literals: 13
% 
% Sat splits: 33
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 3736
% TWLsolver calls for satisfiability: 4119
% 
% Memory used [KB]: 21364
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f95875,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f1103,f1693,f1691,f1102,f1694,f1690,f1109,f1698,f1696,f1108,f1697,f4284,f4286,f4300,f4301,f4285,f4633,f4636,f4656,f4667,f4669,f4677,f4679,f7866,f7868,f8355,f8357,f9222,f9233,f9236,f9713,f9719,f9717,f9826,f9724,f10161,f10163,f27235,f27239,f28119,f28120,f28121,f28122,f28123,f28124,f28125,f28126,f28129,f28130,f28131,f28132,f27968,f28136,f28134,f28271,f28273,f28277,f28279,f28281,f28285,f28287,f28289,f28292,f28294,f28296,f28298,f28301,f28303,f28305,f28310,f28313,f28318,f28322,f28346,f28348,f28357,f28359,f28369,f28371,f28373,f28405,f28407,f28411,f28413,f28418,f28420,f28422,f28463,f28467,f28474,f28477,f28486,f28490,f28498,f28501,f28506,f28512,f28516,f28523,f28526,f28535,f28539,f28546,f28549,f28554,f28556,f28558,f28560,f28612,f28614,f28616,f28618,f28624,f28626,f28628,f28630,f28633,f28636,f28638,f28640,f28643,f28645,f28648,f28650,f28652,f28655,f28657,f28659,f28674,f28676,f28679,f28681,f28688,f28690,f28693,f28695,f28702,f28705,f28728,f28730,f28733,f28735,f28742,f28744,f28747,f28749,f28756,f28759,f28774,f28778,f28785,f28788,f28797,f28801,f28808,f28811,f28816,f28823,f28826,f28830,f28833,f28842,f28845,f28849,f28852,f28861,f28864,f28869,f28872,f28876,f28879,f28887,f28890,f28894,f28897,f28904,f28906,f28912,f28914,f28922,f28926,f28930,f28933,f28939,f28943,f28947,f28950,f28956,f28959,f28963,f28965,f28969,f28971,f28974,f28976,f28982,f28984,f28986,f28988,f28992,f28994,f28998,f29000,f29006,f29008,f29010,f29012,f29016,f29019,f29028,f29030,f29041,f29046,f29050,f29053,f29062,f29065,f29074,f29076,f29084,f29086,f29095,f29097,f29103,f29106,f29113,f29117,f29121,f29124,f29146,f29154,f29163,f29165,f29180,f29182,f29185,f29187,f29191,f29194,f29204,f29208,f29213,f29216,f29221,f29225,f29228,f29232,f29234,f29236,f29238,f29242,f29244,f29251,f29253,f29259,f29261,f29263,f29265,f29270,f29273,f29279,f29282,f29288,f29290,f29292,f29294,f29296,f29298,f29300,f29302,f29304,f29310,f29312,f29314,f29316,f29344,f29345,f29346,f29347,f29348,f29349,f29350,f29351,f29352,f29353,f29354,f29355,f29356,f29357,f29358,f29359,f29360,f29361,f29362,f29363,f29364,f29365,f29366,f29367,f29368,f29369,f29370,f29371,f29372,f29373,f29374,f29375,f29376,f29377,f29378,f29379,f29380,f29381,f29382,f29383,f29384,f29385,f29386,f29387,f29388,f29389,f29390,f29391,f29392,f29393,f29394,f29395,f29396,f29397,f29398,f29399,f29400,f29401,f29402,f29403,f29404,f29405,f29406,f29407,f29408,f29409,f29410,f29411,f29412,f29413,f29414,f29415,f29416,f29417,f29418,f29419,f29420,f29421,f29422,f29423,f29424,f29425,f29426,f29427,f29428,f29429,f29430,f29431,f29432,f29433,f29434,f29435,f29436,f29437,f29438,f29439,f29440,f29441,f29442,f29443,f29444,f29445,f29446,f29447,f29448,f29449,f29450,f29451,f29452,f29453,f29454,f29455,f29456,f29457,f29458,f29459,f29460,f29461,f29462,f29463,f29464,f29465,f29466,f29467,f29468,f29469,f29470,f29471,f29472,f29473,f29474,f29475,f29476,f29477,f29478,f29479,f29480,f29481,f29482,f29483,f29484,f29485,f29486,f29487,f29488,f29489,f29490,f29491,f29492,f29493,f29494,f29495,f29496,f29497,f29498,f29499,f29500,f29501,f29502,f29503,f29504,f29505,f29506,f29507,f29508,f29509,f29510,f29511,f29512,f29513,f29514,f29515,f29516,f29517,f29518,f29519,f29520,f29521,f29522,f29523,f29524,f29525,f29526,f29527,f29528,f29529,f29530,f29531,f29532,f29533,f29534,f29535,f29536,f29537,f29538,f29539,f29540,f29541,f29542,f29543,f29544,f29545,f29546,f29547,f29548,f29549,f29550,f29551,f29552,f29553,f29554,f29555,f29556,f29557,f29558,f29559,f29560,f29561,f29562,f29563,f29564,f29565,f29566,f29567,f29568,f29569,f29570,f29571,f29572,f29573,f29574,f29575,f29576,f29577,f29578,f29579,f29580,f29581,f29582,f29583,f29584,f29585,f29586,f29587,f29588,f29589,f29590,f29591,f29592,f29593,f29594,f29595,f29596,f29597,f29598,f29599,f29600,f29601,f29602,f29603,f29604,f29605,f29606,f29607,f29608,f29609,f29610,f29611,f29612,f29613,f29614,f29615,f29616,f29617,f29618,f29619,f29620,f29621,f29622,f29623,f29624,f29625,f29626,f29627,f29628,f29703,f29704,f29705,f29706,f29707,f29708,f29709,f29710,f29711,f29712,f29713,f30411,f30412,f31112,f31113,f31212,f31213,f31294,f31295,f31384,f31385,f31470,f31471,f50287,f50290,f50289,f81232,f81236,f81234,f81250,f81254,f81252,f89821,f89822,f89824,f7869,f89825,f89844,f89845,f89847,f89848,f89975,f89981,f89979,f90023,f90024,f90053,f90054,f90366,f9234,f90367,f93265,f93267,f93268,f1699,f93270,f95866,f95843,f95847,f95851,f95857,f95861,f95867,f95868,f95869,f95870,f95871,f95872,f95873,f95874])).
fof(f95874,plain,(
  $false | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f95844,f9726])).
fof(f9726,plain,(
  ( ! [X0] : (sP5('T_0','const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK14)),'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1104,f1694,f1560])).
fof(f1560,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f1076])).
fof(f1076,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP5(X2,X1,X0)) & (~sP5(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f829])).
fof(f829,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP5(X2,X1,X0))),
  inference(definition_folding,[],[f631,f828,f827,f826,f825])).
fof(f825,plain,(
  ! [X0,X1,X2] : (sP2(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f826,plain,(
  ! [X1,X0,X2] : (sP3(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f827,plain,(
  ! [X2,X0,X1] : (sP4(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f828,plain,(
  ! [X2,X1,X0] : (sP5(X2,X1,X0) <=> (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f631,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f630])).
fof(f630,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f53])).
fof(f53,axiom,(
  ! [X33,X34,X23] : ((p(s(bool,X23)) <=> s(bool,X33) = s(bool,X34)) <=> ((p(s(bool,X23)) | p(s(bool,X34)) | p(s(bool,X33))) & (p(s(bool,X23)) | ~p(s(bool,X33)) | ~p(s(bool,X34))) & (p(s(bool,X34)) | ~p(s(bool,X33)) | ~p(s(bool,X23))) & (p(s(bool,X33)) | ~p(s(bool,X34)) | ~p(s(bool,X23)))))),
  file('thm.pred_set.EXTENSION','thm.sat.sat_2|49|_prove_ltb')).
fof(f1104,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f11])).
fof(f11,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pred_set.EXTENSION','thm.bool.TRUTH')).
fof(f95844,plain,(
  ~sP5('T_0','const.bool.IN_2'(s(sK12,sK81(sK13,sK14,sK12,bool)),s(fun(sK12,bool),sK14)),'const.bool.IN_2'(s(sK12,sK81(sK13,sK14,sK12,bool)),s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f1104,f27233,f1561])).
fof(f1561,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | ~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f1076])).
fof(f27233,plain,(
  s(bool,'const.bool.IN_2'(s(sK12,sK81(sK13,sK14,sK12,bool)),s(fun(sK12,bool),sK13))) != s(bool,'const.bool.IN_2'(s(sK12,sK81(sK13,sK14,sK12,bool)),s(fun(sK12,bool),sK14))) | $spl1),
  inference(forward_demodulation,[],[f27232,f1257])).
fof(f1257,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2))) = s(bool,happ(s(fun(X0,bool),X2),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f398])).
fof(f398,plain,(
  ! [X0,X1,X2] : s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2))) = s(bool,happ(s(fun(X0,bool),X2),s(X0,X1)))),
  inference(flattening,[],[f397])).
fof(f397,plain,(
  ! [X0] : ! [X1,X2] : s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2))) = s(bool,happ(s(fun(X0,bool),X2),s(X0,X1)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X6] : ! [X9,X17] : s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X17))) = s(bool,happ(s(fun(X6,bool),X17),s(X6,X9)))),
  file('thm.pred_set.EXTENSION','thm.pred_set.IN_ABS')).
fof(f27232,plain,(
  s(bool,'const.bool.IN_2'(s(sK12,sK81(sK13,sK14,sK12,bool)),s(fun(sK12,bool),sK14))) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK81(sK13,sK14,sK12,bool)))) | $spl1),
  inference(forward_demodulation,[],[f27139,f1257])).
fof(f27139,plain,(
  s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK81(sK13,sK14,sK12,bool)))) != s(bool,happ(s(fun(sK12,bool),sK14),s(sK12,sK81(sK13,sK14,sK12,bool)))) | $spl1),
  inference(unit_resulting_resolution,[],[f1691,f1460])).
fof(f1460,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,sK81(X3,X2,X1,X0)))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,sK81(X3,X2,X1,X0)))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)) )),
  inference(cnf_transformation,[],[f1021])).
fof(f1021,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X5] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))) & (s(X0,happ(s(fun(X1,X0),X2),s(X1,sK81(X3,X2,X1,X0)))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,sK81(X3,X2,X1,X0)))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK81])],[f1020])).
fof(f1020,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X5] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))) & (? [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(rectify,[],[f1019])).
fof(f1019,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4)))) & (? [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(nnf_transformation,[],[f557])).
fof(f557,plain,(
  ! [X0,X1,X2,X3] : (s(fun(X1,X0),X2) = s(fun(X1,X0),X3) <=> ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))))),
  inference(flattening,[],[f556])).
fof(f556,plain,(
  ! [X0,X1] : ! [X2,X3] : (s(fun(X1,X0),X2) = s(fun(X1,X0),X3) <=> ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))))),
  inference(rectify,[],[f81])).
fof(f81,axiom,(
  ! [X20,X6] : ! [X15,X14] : (s(fun(X6,X20),X14) = s(fun(X6,X20),X15) <=> ! [X9] : s(X20,happ(s(fun(X6,X20),X14),s(X6,X9))) = s(X20,happ(s(fun(X6,X20),X15),s(X6,X9))))),
  file('thm.pred_set.EXTENSION','thm.bool.FUN_EQ_THM')).
fof(f95873,plain,(
  $false | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f95846,f1104])).
fof(f95846,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl4)),
  inference(unit_resulting_resolution,[],[f9726,f27233,f1561])).
fof(f95872,plain,(
  $false | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f95848,f9725])).
fof(f9725,plain,(
  ( ! [X0] : (~sP5('F_0','const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK14)),'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1700,f1694,f1562])).
fof(f1562,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f1076])).
fof(f1700,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f1111,f1104])).
fof(f1111,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f845])).
fof(f845,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f265])).
fof(f265,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f60])).
fof(f60,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.pred_set.EXTENSION','thm.bool.NOT_CLAUSES|split|2')).
fof(f95848,plain,(
  sP5('F_0','const.bool.IN_2'(s(sK12,sK81(sK13,sK14,sK12,bool)),s(fun(sK12,bool),sK14)),'const.bool.IN_2'(s(sK12,sK81(sK13,sK14,sK12,bool)),s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f1700,f27233,f1559])).
fof(f1559,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f1076])).
fof(f95871,plain,(
  $false | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f95850,f1700])).
fof(f95850,plain,(
  p(s(bool,'F_0')) | ($spl1 | $spl4)),
  inference(unit_resulting_resolution,[],[f9725,f27233,f1559])).
fof(f95870,plain,(
  $false | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f95854,f9731])).
fof(f9731,plain,(
  ( ! [X0] : (sP5('T_0','const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13)),'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1104,f1694,f1560])).
fof(f95854,plain,(
  ~sP5('T_0','const.bool.IN_2'(s(sK12,sK81(sK13,sK14,sK12,bool)),s(fun(sK12,bool),sK13)),'const.bool.IN_2'(s(sK12,sK81(sK13,sK14,sK12,bool)),s(fun(sK12,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f1104,f27233,f1561])).
fof(f95869,plain,(
  $false | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f95856,f1104])).
fof(f95856,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl4)),
  inference(unit_resulting_resolution,[],[f9731,f27233,f1561])).
fof(f95868,plain,(
  $false | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f95858,f9730])).
fof(f9730,plain,(
  ( ! [X0] : (~sP5('F_0','const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13)),'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1700,f1694,f1562])).
fof(f95858,plain,(
  sP5('F_0','const.bool.IN_2'(s(sK12,sK81(sK13,sK14,sK12,bool)),s(fun(sK12,bool),sK13)),'const.bool.IN_2'(s(sK12,sK81(sK13,sK14,sK12,bool)),s(fun(sK12,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f1700,f27233,f1559])).
fof(f95867,plain,(
  $false | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f95860,f1700])).
fof(f95860,plain,(
  p(s(bool,'F_0')) | ($spl1 | $spl4)),
  inference(unit_resulting_resolution,[],[f9730,f27233,f1559])).
fof(f95861,plain,(
  $false | ($spl1 | $spl4)),
  inference(unit_resulting_resolution,[],[f1700,f9730,f27233,f1559])).
fof(f95857,plain,(
  $false | ($spl1 | $spl4)),
  inference(unit_resulting_resolution,[],[f1104,f9731,f27233,f1561])).
fof(f95851,plain,(
  $false | ($spl1 | $spl4)),
  inference(unit_resulting_resolution,[],[f1700,f9725,f27233,f1559])).
fof(f95847,plain,(
  $false | ($spl1 | $spl4)),
  inference(unit_resulting_resolution,[],[f1104,f9726,f27233,f1561])).
fof(f95843,plain,(
  $false | ($spl1 | $spl4)),
  inference(unit_resulting_resolution,[],[f1694,f27233])).
fof(f95866,plain,(
  $false | ($spl1 | $spl4)),
  inference(trivial_inequality_removal,[],[f95864])).
fof(f95864,plain,(
  s(bool,'const.bool.IN_2'(s(sK12,sK81(sK13,sK14,sK12,bool)),s(fun(sK12,bool),sK13))) != s(bool,'const.bool.IN_2'(s(sK12,sK81(sK13,sK14,sK12,bool)),s(fun(sK12,bool),sK13))) | ($spl1 | $spl4)),
  inference(superposition,[],[f27233,f1694])).
fof(f93270,plain,(
  ~p(s(bool,sK29('_i_0_0',bool))) | $spl95),
  inference(cnf_transformation,[],[f93270_D])).
fof(f93270_D,plain,(
  ~p(s(bool,sK29('_i_0_0',bool))) <=> ~$spl95),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl95])])).
fof(f1699,plain,(
  ( ! [X1] : (p(s(bool,X1))) ) | $spl10),
  inference(cnf_transformation,[],[f1699_D])).
fof(f1699_D,plain,(
  ( ! [X1] : (p(s(bool,X1))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f93268,plain,(
  ( ! [X1] : (p(s(bool,X1)) | ~p(s(bool,sK29('_i_0_0',bool)))) )),
  inference(forward_demodulation,[],[f93158,f6735])).
fof(f6735,plain,(
  ( ! [X6] : (s(bool,X6) = s(bool,'const.bool.IN_2'(s(bool,X6),s(fun(bool,bool),'_i_0_0')))) )),
  inference(superposition,[],[f1257,f1189])).
fof(f1189,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f315])).
fof(f315,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(flattening,[],[f314])).
fof(f314,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(rectify,[],[f105])).
fof(f105,axiom,(
  ! [X38] : ! [X40] : s(X38,X40) = s(X38,happ(s(fun(X38,X38),'_i_0_0'),s(X38,X40)))),
  file('thm.pred_set.EXTENSION','_i0_ax')).
fof(f93158,plain,(
  ( ! [X1] : (~p(s(bool,sK29('_i_0_0',bool))) | p(s(bool,'const.bool.IN_2'(s(bool,X1),s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f1730,f6735])).
fof(f1730,plain,(
  ( ! [X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,sK29(X1,X0)),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))))) )),
  inference(forward_demodulation,[],[f1722,f1257])).
fof(f1722,plain,(
  ( ! [X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,sK29(X1,X0)),s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) )),
  inference(backward_demodulation,[],[f1257,f1250])).
fof(f1250,plain,(
  ( ! [X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK29(X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f912])).
fof(f912,plain,(
  ! [X0,X1] : ((! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK28(X1,X0)))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK29(X1,X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK28,sK29])],[f911])).
fof(f911,plain,(
  ! [X0,X1] : ((! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(rectify,[],[f910])).
fof(f910,plain,(
  ! [X0,X1] : ((! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(nnf_transformation,[],[f751])).
fof(f751,plain,(
  ! [X0,X1] : (? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(ennf_transformation,[],[f387])).
fof(f387,plain,(
  ! [X0,X1] : (~! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(flattening,[],[f386])).
fof(f386,plain,(
  ! [X0] : ! [X1] : (~! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(rectify,[],[f88])).
fof(f88,axiom,(
  ! [X6] : ! [X17] : (~! [X9] : p(s(bool,happ(s(fun(X6,bool),X17),s(X6,X9)))) <=> ? [X9] : ~p(s(bool,happ(s(fun(X6,bool),X17),s(X6,X9)))))),
  file('thm.pred_set.EXTENSION','thm.bool.NOT_FORALL_THM')).
fof(f93267,plain,(
  ~p(s(bool,sK28('_i_0_0',bool))) | $spl93),
  inference(cnf_transformation,[],[f93267_D])).
fof(f93267_D,plain,(
  ~p(s(bool,sK28('_i_0_0',bool))) <=> ~$spl93),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl93])])).
fof(f93265,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,sK28('_i_0_0',bool)))) )),
  inference(forward_demodulation,[],[f93157,f6735])).
fof(f93157,plain,(
  ( ! [X0] : (~p(s(bool,sK28('_i_0_0',bool))) | p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f1731,f6735])).
fof(f1731,plain,(
  ( ! [X0,X5,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,sK28(X1,X0)),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),X1))))) )),
  inference(forward_demodulation,[],[f1723,f1257])).
fof(f1723,plain,(
  ( ! [X0,X5,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),X1)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK28(X1,X0)))))) )),
  inference(backward_demodulation,[],[f1257,f1249])).
fof(f1249,plain,(
  ( ! [X0,X5,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK28(X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))) )),
  inference(cnf_transformation,[],[f912])).
fof(f90367,plain,(
  ( ! [X10,X9] : (sP11('const.combin.I_1'(s(bool,'T_0')),X9,X10) | ~p(s(bool,X9))) ) | $spl90),
  inference(cnf_transformation,[],[f90367_D])).
fof(f90367_D,plain,(
  ( ! [X10,X9] : (sP11('const.combin.I_1'(s(bool,'T_0')),X9,X10) | ~p(s(bool,X9))) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f9234,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13))))) ) | $spl26),
  inference(cnf_transformation,[],[f9234_D])).
fof(f9234_D,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13))))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f90366,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X8),s(fun(sK12,bool),sK13)))) | sP11('const.combin.I_1'(s(bool,'T_0')),X9,X10) | ~p(s(bool,X9))) ) | $spl4),
  inference(subsumption_resolution,[],[f90365,f1137])).
fof(f1137,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f849])).
fof(f849,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f291])).
fof(f291,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f290])).
fof(f290,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f51])).
fof(f51,axiom,(
  ! [X8] : (s(bool,'F_0') = s(bool,X8) <=> ~p(s(bool,X8)))),
  file('thm.pred_set.EXTENSION','thm.bool.EQ_CLAUSES|split|2')).
fof(f90365,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X8),s(fun(sK12,bool),sK13)))) | sP11('const.combin.I_1'(s(bool,'T_0')),X9,X10) | ~p(s(bool,X9)) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,X8),s(fun(sK12,bool),sK13)))) ) | $spl4),
  inference(forward_demodulation,[],[f90311,f1694])).
fof(f90311,plain,(
  ( ! [X10,X8,X9] : (sP11('const.combin.I_1'(s(bool,'T_0')),X9,X10) | ~p(s(bool,X9)) | ~p(s(bool,'const.bool.IN_2'(s(sK12,X8),s(fun(sK12,bool),sK14)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,X8),s(fun(sK12,bool),sK13)))) ) | $spl4),
  inference(superposition,[],[f2068,f9743])).
fof(f9743,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,X1),s(fun(sK12,bool),sK14))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,X1),s(fun(sK12,bool),sK13)))) ) | $spl4),
  inference(superposition,[],[f1694,f1120])).
fof(f1120,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f274])).
fof(f274,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  inference(rectify,[],[f64])).
fof(f64,axiom,(
  ! [X8] : (s(bool,'T_0') = s(bool,X8) | s(bool,'F_0') = s(bool,X8))),
  file('thm.pred_set.EXTENSION','thm.bool.BOOL_CASES_AX')).
fof(f2068,plain,(
  ( ! [X6,X7,X5] : (sP11('const.combin.I_1'(s(bool,X5)),X6,X7) | ~p(s(bool,X6)) | ~p(s(bool,X5))) )),
  inference(superposition,[],[f1610,f1188])).
fof(f1188,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))) )),
  inference(cnf_transformation,[],[f313])).
fof(f313,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))),
  inference(flattening,[],[f312])).
fof(f312,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))),
  inference(rectify,[],[f148])).
fof(f148,axiom,(
  ! [X6] : ! [X9] : s(X6,X9) = s(X6,'const.combin.I_1'(s(X6,X9)))),
  file('thm.pred_set.EXTENSION','thm.combin.I_THM')).
fof(f1610,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f1099])).
fof(f1099,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f1098])).
fof(f1098,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f838])).
fof(f838,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP11(X2,X0,X1))),
  inference(definition_folding,[],[f818,f837,f836])).
fof(f836,plain,(
  ! [X2,X0,X1] : (sP10(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f837,plain,(
  ! [X2,X0,X1] : (sP11(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f818,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f637])).
fof(f637,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f636])).
fof(f636,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f82])).
fof(f82,axiom,(
  ! [X33,X34,X23] : ((p(s(bool,X23)) <=> (p(s(bool,X34)) => p(s(bool,X33)))) <=> ((p(s(bool,X23)) | p(s(bool,X34))) & (p(s(bool,X23)) | ~p(s(bool,X33))) & (~p(s(bool,X34)) | p(s(bool,X33)) | ~p(s(bool,X23)))))),
  file('thm.pred_set.EXTENSION','thm.sat.sat_27_prove_ltb')).
fof(f90054,plain,(
  ( ! [X4,X3] : (sP11('const.combin.I_1'(s(bool,'T_0')),X3,X4) | p(s(bool,X4))) ) | $spl88),
  inference(cnf_transformation,[],[f90054_D])).
fof(f90054_D,plain,(
  ( ! [X4,X3] : (sP11('const.combin.I_1'(s(bool,'T_0')),X3,X4) | p(s(bool,X4))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f90053,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X2),s(fun(sK12,bool),sK13)))) | sP11('const.combin.I_1'(s(bool,'T_0')),X3,X4) | p(s(bool,X4))) ) | $spl4),
  inference(subsumption_resolution,[],[f90052,f1137])).
fof(f90052,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X2),s(fun(sK12,bool),sK13)))) | sP11('const.combin.I_1'(s(bool,'T_0')),X3,X4) | p(s(bool,X4)) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,X2),s(fun(sK12,bool),sK13)))) ) | $spl4),
  inference(forward_demodulation,[],[f90037,f1694])).
fof(f90037,plain,(
  ( ! [X4,X2,X3] : (sP11('const.combin.I_1'(s(bool,'T_0')),X3,X4) | p(s(bool,X4)) | ~p(s(bool,'const.bool.IN_2'(s(sK12,X2),s(fun(sK12,bool),sK14)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,X2),s(fun(sK12,bool),sK13)))) ) | $spl4),
  inference(superposition,[],[f2059,f9743])).
fof(f2059,plain,(
  ( ! [X6,X7,X5] : (sP11('const.combin.I_1'(s(bool,X5)),X7,X6) | p(s(bool,X6)) | ~p(s(bool,X5))) )),
  inference(superposition,[],[f1609,f1188])).
fof(f1609,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f1099])).
fof(f90024,plain,(
  ( ! [X4,X3] : (sP9('const.combin.I_1'(s(bool,'T_0')),X3,X4) | ~p(s(bool,X3))) ) | $spl86),
  inference(cnf_transformation,[],[f90024_D])).
fof(f90024_D,plain,(
  ( ! [X4,X3] : (sP9('const.combin.I_1'(s(bool,'T_0')),X3,X4) | ~p(s(bool,X3))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f90023,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X2),s(fun(sK12,bool),sK13)))) | sP9('const.combin.I_1'(s(bool,'T_0')),X3,X4) | ~p(s(bool,X3))) ) | $spl4),
  inference(subsumption_resolution,[],[f90022,f1137])).
fof(f90022,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X2),s(fun(sK12,bool),sK13)))) | sP9('const.combin.I_1'(s(bool,'T_0')),X3,X4) | ~p(s(bool,X3)) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,X2),s(fun(sK12,bool),sK13)))) ) | $spl4),
  inference(forward_demodulation,[],[f90007,f1694])).
fof(f90007,plain,(
  ( ! [X4,X2,X3] : (sP9('const.combin.I_1'(s(bool,'T_0')),X3,X4) | ~p(s(bool,X3)) | ~p(s(bool,'const.bool.IN_2'(s(sK12,X2),s(fun(sK12,bool),sK14)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,X2),s(fun(sK12,bool),sK13)))) ) | $spl4),
  inference(superposition,[],[f2048,f9743])).
fof(f2048,plain,(
  ( ! [X6,X7,X5] : (sP9('const.combin.I_1'(s(bool,X5)),X6,X7) | ~p(s(bool,X6)) | ~p(s(bool,X5))) )),
  inference(superposition,[],[f1593,f1188])).
fof(f1593,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f1091])).
fof(f1091,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f1090])).
fof(f1090,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f835])).
fof(f835,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP9(X2,X0,X1))),
  inference(definition_folding,[],[f635,f834,f833])).
fof(f833,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f834,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f635,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f634])).
fof(f634,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f59])).
fof(f59,axiom,(
  ! [X33,X34,X23] : ((p(s(bool,X23)) <=> (p(s(bool,X34)) | p(s(bool,X33)))) <=> ((p(s(bool,X23)) | ~p(s(bool,X34))) & (p(s(bool,X23)) | ~p(s(bool,X33))) & (p(s(bool,X34)) | p(s(bool,X33)) | ~p(s(bool,X23)))))),
  file('thm.pred_set.EXTENSION','thm.sat.sat_25_prove_ltb')).
fof(f89979,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK12,sK61(sK14,sK14,sK12)),s(fun(sK12,bool),sK13)))) | $spl82),
  inference(cnf_transformation,[],[f89979_D])).
fof(f89979_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK12,sK61(sK14,sK14,sK12)),s(fun(sK12,bool),sK13)))) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f89981,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK12,sK60(sK14,sK14,sK12)),s(fun(sK12,bool),sK13)))) | $spl84),
  inference(cnf_transformation,[],[f89981_D])).
fof(f89981_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK12,sK60(sK14,sK14,sK12)),s(fun(sK12,bool),sK13)))) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f89975,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK12,sK60(sK14,sK14,sK12)),s(fun(sK12,bool),sK13)))) | p(s(bool,'const.bool.IN_2'(s(sK12,sK61(sK14,sK14,sK12)),s(fun(sK12,bool),sK13)))) | ($spl4 | $spl58)),
  inference(superposition,[],[f31385,f1694])).
fof(f89848,plain,(
  ( ! [X85] : (sP1(X85,'const.combin.I_1'(s(bool,'T_0'))) | p(s(bool,X85))) ) | $spl80),
  inference(cnf_transformation,[],[f89848_D])).
fof(f89848_D,plain,(
  ( ! [X85] : (sP1(X85,'const.combin.I_1'(s(bool,'T_0'))) | p(s(bool,X85))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f89847,plain,(
  ( ! [X85,X84] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X84),s(fun(sK12,bool),sK13)))) | sP1(X85,'const.combin.I_1'(s(bool,'T_0'))) | p(s(bool,X85))) ) | $spl4),
  inference(subsumption_resolution,[],[f89846,f1137])).
fof(f89846,plain,(
  ( ! [X85,X84] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X84),s(fun(sK12,bool),sK13)))) | sP1(X85,'const.combin.I_1'(s(bool,'T_0'))) | p(s(bool,X85)) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,X84),s(fun(sK12,bool),sK13)))) ) | $spl4),
  inference(forward_demodulation,[],[f89709,f1694])).
fof(f89709,plain,(
  ( ! [X85,X84] : (sP1(X85,'const.combin.I_1'(s(bool,'T_0'))) | p(s(bool,X85)) | ~p(s(bool,'const.bool.IN_2'(s(sK12,X84),s(fun(sK12,bool),sK14)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,X84),s(fun(sK12,bool),sK13)))) ) | $spl4),
  inference(superposition,[],[f1987,f9743])).
fof(f1987,plain,(
  ( ! [X2,X1] : (sP1(X2,'const.combin.I_1'(s(bool,X1))) | p(s(bool,X2)) | ~p(s(bool,X1))) )),
  inference(superposition,[],[f1416,f1188])).
fof(f1416,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f1000])).
fof(f1000,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1)))),
  inference(rectify,[],[f999])).
fof(f999,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0)))),
  inference(flattening,[],[f998])).
fof(f998,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP1(X1,X0)))),
  inference(nnf_transformation,[],[f823])).
fof(f823,plain,(
  ! [X1,X0] : (sP1(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f89845,plain,(
  ( ! [X83] : (sP1('const.combin.I_1'(s(bool,'T_0')),X83) | p(s(bool,X83))) ) | $spl78),
  inference(cnf_transformation,[],[f89845_D])).
fof(f89845_D,plain,(
  ( ! [X83] : (sP1('const.combin.I_1'(s(bool,'T_0')),X83) | p(s(bool,X83))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f89844,plain,(
  ( ! [X83,X82] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X82),s(fun(sK12,bool),sK13)))) | sP1('const.combin.I_1'(s(bool,'T_0')),X83) | p(s(bool,X83))) ) | $spl4),
  inference(subsumption_resolution,[],[f89843,f1137])).
fof(f89843,plain,(
  ( ! [X83,X82] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X82),s(fun(sK12,bool),sK13)))) | sP1('const.combin.I_1'(s(bool,'T_0')),X83) | p(s(bool,X83)) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,X82),s(fun(sK12,bool),sK13)))) ) | $spl4),
  inference(forward_demodulation,[],[f89708,f1694])).
fof(f89708,plain,(
  ( ! [X83,X82] : (sP1('const.combin.I_1'(s(bool,'T_0')),X83) | p(s(bool,X83)) | ~p(s(bool,'const.bool.IN_2'(s(sK12,X82),s(fun(sK12,bool),sK14)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,X82),s(fun(sK12,bool),sK13)))) ) | $spl4),
  inference(superposition,[],[f1979,f9743])).
fof(f1979,plain,(
  ( ! [X2,X1] : (sP1('const.combin.I_1'(s(bool,X1)),X2) | p(s(bool,X2)) | ~p(s(bool,X1))) )),
  inference(superposition,[],[f1413,f1188])).
fof(f1413,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f1000])).
fof(f89825,plain,(
  ( ! [X19] : (p(s(bool,X19)) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK55(sK14,X19,sK12)),s(fun(sK12,bool),sK13)))) ) | $spl76),
  inference(cnf_transformation,[],[f89825_D])).
fof(f89825_D,plain,(
  ( ! [X19] : (p(s(bool,X19)) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK55(sK14,X19,sK12)),s(fun(sK12,bool),sK13)))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f7869,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13))))) ) | $spl20),
  inference(cnf_transformation,[],[f7869_D])).
fof(f7869_D,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13))))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f89824,plain,(
  ( ! [X19,X20] : (p(s(bool,'const.bool.IN_2'(s(sK12,X20),s(fun(sK12,bool),sK13)))) | p(s(bool,X19)) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK55(sK14,X19,sK12)),s(fun(sK12,bool),sK13)))) ) | $spl4),
  inference(forward_demodulation,[],[f89823,f1694])).
fof(f89823,plain,(
  ( ! [X19,X20] : (p(s(bool,'const.bool.IN_2'(s(sK12,X20),s(fun(sK12,bool),sK14)))) | p(s(bool,X19)) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK55(sK14,X19,sK12)),s(fun(sK12,bool),sK13)))) ) | $spl4),
  inference(subsumption_resolution,[],[f89677,f1104])).
fof(f89677,plain,(
  ( ! [X19,X20] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK12,X20),s(fun(sK12,bool),sK14)))) | p(s(bool,X19)) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK55(sK14,X19,sK12)),s(fun(sK12,bool),sK13)))) ) | $spl4),
  inference(superposition,[],[f1802,f9743])).
fof(f1802,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,sK55(X2,X1,X0)),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) | p(s(bool,X1))) )),
  inference(forward_demodulation,[],[f1801,f1257])).
fof(f1801,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) | p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK55(X2,X1,X0)))))) )),
  inference(forward_demodulation,[],[f1306,f1257])).
fof(f1306,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK55(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f949])).
fof(f949,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK54(X2,X1,X0)))))) | p(s(bool,X1)) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5))))) & ((~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK55(X2,X1,X0)))))) | ! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK54,sK55])],[f948])).
fof(f948,plain,(
  ! [X0,X1,X2] : ((? [X6] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | p(s(bool,X1)) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5))))) & ((~p(s(bool,X1)) & ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(rectify,[],[f947])).
fof(f947,plain,(
  ! [X0,X1,X2] : ((? [X3] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | p(s(bool,X1)) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) & ((~p(s(bool,X1)) & ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(flattening,[],[f946])).
fof(f946,plain,(
  ! [X0,X1,X2] : ((? [X3] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (p(s(bool,X1)) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & ((~p(s(bool,X1)) & ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f434])).
fof(f434,plain,(
  ! [X0,X1,X2] : (! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (p(s(bool,X1)) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f433])).
fof(f433,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (p(s(bool,X1)) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f94])).
fof(f94,axiom,(
  ! [X6] : ! [X17,X16] : (! [X9] : (p(s(bool,X17)) | p(s(bool,happ(s(fun(X6,bool),X16),s(X6,X9))))) <=> (p(s(bool,X17)) | ! [X9] : p(s(bool,happ(s(fun(X6,bool),X16),s(X6,X9))))))),
  file('thm.pred_set.EXTENSION','thm.bool.RIGHT_FORALL_OR_THM')).
fof(f89822,plain,(
  ( ! [X17] : (p(s(bool,X17)) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK54(sK14,X17,sK12)),s(fun(sK12,bool),sK13)))) ) | $spl74),
  inference(cnf_transformation,[],[f89822_D])).
fof(f89822_D,plain,(
  ( ! [X17] : (p(s(bool,X17)) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK54(sK14,X17,sK12)),s(fun(sK12,bool),sK13)))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f89821,plain,(
  ( ! [X17,X18] : (p(s(bool,'const.bool.IN_2'(s(sK12,X18),s(fun(sK12,bool),sK13)))) | p(s(bool,X17)) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK54(sK14,X17,sK12)),s(fun(sK12,bool),sK13)))) ) | $spl4),
  inference(forward_demodulation,[],[f89820,f1694])).
fof(f89820,plain,(
  ( ! [X17,X18] : (p(s(bool,'const.bool.IN_2'(s(sK12,X18),s(fun(sK12,bool),sK14)))) | p(s(bool,X17)) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK54(sK14,X17,sK12)),s(fun(sK12,bool),sK13)))) ) | $spl4),
  inference(subsumption_resolution,[],[f89676,f1104])).
fof(f89676,plain,(
  ( ! [X17,X18] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK12,X18),s(fun(sK12,bool),sK14)))) | p(s(bool,X17)) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK54(sK14,X17,sK12)),s(fun(sK12,bool),sK13)))) ) | $spl4),
  inference(superposition,[],[f1804,f9743])).
fof(f1804,plain,(
  ( ! [X2,X0,X5,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,sK54(X2,X1,X0)),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),X2)))) | p(s(bool,X1))) )),
  inference(forward_demodulation,[],[f1803,f1257])).
fof(f1803,plain,(
  ( ! [X2,X0,X5,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),X2)))) | p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK54(X2,X1,X0)))))) )),
  inference(forward_demodulation,[],[f1304,f1257])).
fof(f1304,plain,(
  ( ! [X2,X0,X5,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) | p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK54(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f949])).
fof(f81252,plain,(
  s(bool,'T_0') = s(bool,g(s(sK12,sK80(sK13,sK14,bool,sK12)))) | $spl70),
  inference(cnf_transformation,[],[f81252_D])).
fof(f81252_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK12,sK80(sK13,sK14,bool,sK12)))) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f81254,plain,(
  s(bool,'F_0') != s(bool,f(s(sK12,sK80(sK13,sK14,bool,sK12)))) | $spl73),
  inference(cnf_transformation,[],[f81254_D])).
fof(f81254_D,plain,(
  s(bool,'F_0') != s(bool,f(s(sK12,sK80(sK13,sK14,bool,sK12)))) <=> ~$spl73),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl73])])).
fof(f81250,plain,(
  s(bool,'F_0') != s(bool,f(s(sK12,sK80(sK13,sK14,bool,sK12)))) | s(bool,'T_0') = s(bool,g(s(sK12,sK80(sK13,sK14,bool,sK12)))) | $spl1),
  inference(superposition,[],[f14063,f1120])).
fof(f14063,plain,(
  s(bool,f(s(sK12,sK80(sK13,sK14,bool,sK12)))) != s(bool,g(s(sK12,sK80(sK13,sK14,bool,sK12)))) | $spl1),
  inference(unit_resulting_resolution,[],[f1691,f1458])).
fof(f1458,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK80(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK80(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f1018])).
fof(f1018,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK80(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK80(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK80])],[f797])).
fof(f797,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f555])).
fof(f555,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f554])).
fof(f554,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.pred_set.EXTENSION',aHL_EXT)).
fof(f81234,plain,(
  s(bool,'T_0') = s(bool,g(s(sK12,sK80(sK14,sK13,bool,sK12)))) | $spl66),
  inference(cnf_transformation,[],[f81234_D])).
fof(f81234_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK12,sK80(sK14,sK13,bool,sK12)))) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f81236,plain,(
  s(bool,'F_0') != s(bool,f(s(sK12,sK80(sK14,sK13,bool,sK12)))) | $spl69),
  inference(cnf_transformation,[],[f81236_D])).
fof(f81236_D,plain,(
  s(bool,'F_0') != s(bool,f(s(sK12,sK80(sK14,sK13,bool,sK12)))) <=> ~$spl69),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl69])])).
fof(f81232,plain,(
  s(bool,'F_0') != s(bool,f(s(sK12,sK80(sK14,sK13,bool,sK12)))) | s(bool,'T_0') = s(bool,g(s(sK12,sK80(sK14,sK13,bool,sK12)))) | $spl1),
  inference(superposition,[],[f14062,f1120])).
fof(f14062,plain,(
  s(bool,f(s(sK12,sK80(sK14,sK13,bool,sK12)))) != s(bool,g(s(sK12,sK80(sK14,sK13,bool,sK12)))) | $spl1),
  inference(unit_resulting_resolution,[],[f1691,f1458])).
fof(f50289,plain,(
  s(bool,'T_0') = s(bool,'const.bool.ARB_0') | $spl62),
  inference(cnf_transformation,[],[f50289_D])).
fof(f50289_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.ARB_0') <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f50290,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'const.relation.RESTRICT_4'(s(fun(X0,bool),X1),s(fun(X0,fun(X0,bool)),X2),s(X0,X3),s(X0,X4))) = s(bool,'const.bool.COND_3'(s(bool,'const.relation.inv_3'(s(fun(X0,fun(X0,bool)),X2),s(X0,X3),s(X0,X4))),s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))),s(bool,'F_0')))) ) | $spl64),
  inference(cnf_transformation,[],[f50290_D])).
fof(f50290_D,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'const.relation.RESTRICT_4'(s(fun(X0,bool),X1),s(fun(X0,fun(X0,bool)),X2),s(X0,X3),s(X0,X4))) = s(bool,'const.bool.COND_3'(s(bool,'const.relation.inv_3'(s(fun(X0,fun(X0,bool)),X2),s(X0,X3),s(X0,X4))),s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))),s(bool,'F_0')))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f50287,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'const.relation.RESTRICT_4'(s(fun(X0,bool),X1),s(fun(X0,fun(X0,bool)),X2),s(X0,X3),s(X0,X4))) = s(bool,'const.bool.COND_3'(s(bool,'const.relation.inv_3'(s(fun(X0,fun(X0,bool)),X2),s(X0,X3),s(X0,X4))),s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.bool.ARB_0')) )),
  inference(forward_demodulation,[],[f49874,f1257])).
fof(f49874,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'const.relation.RESTRICT_4'(s(fun(X0,bool),X1),s(fun(X0,fun(X0,bool)),X2),s(X0,X3),s(X0,X4))) = s(bool,'const.bool.COND_3'(s(bool,'const.relation.inv_3'(s(fun(X0,fun(X0,bool)),X2),s(X0,X3),s(X0,X4))),s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.bool.ARB_0')) )),
  inference(superposition,[],[f18664,f1120])).
fof(f18664,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X1,'const.relation.RESTRICT_4'(s(fun(X0,X1),X2),s(fun(X0,fun(X0,bool)),X3),s(X0,X4),s(X0,X5))) = s(X1,'const.bool.COND_3'(s(bool,'const.relation.inv_3'(s(fun(X0,fun(X0,bool)),X3),s(X0,X4),s(X0,X5))),s(X1,happ(s(fun(X0,X1),X2),s(X0,X5))),s(X1,'const.bool.ARB_0')))) )),
  inference(backward_demodulation,[],[f18484,f1879])).
fof(f1879,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X1,'const.relation.RESTRICT_4'(s(fun(X0,X1),X2),s(fun(X0,fun(X0,bool)),X3),s(X0,X4),s(X0,X5))) = s(X1,'const.bool.COND_3'(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,X5))))),s(X1,happ(s(fun(X0,X1),X2),s(X0,X5))),s(X1,'const.bool.ARB_0')))) )),
  inference(forward_demodulation,[],[f1497,f1257])).
fof(f1497,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X1,'const.relation.RESTRICT_4'(s(fun(X0,X1),X2),s(fun(X0,fun(X0,bool)),X3),s(X0,X4),s(X0,X5))) = s(X1,'const.bool.COND_3'(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,X5))),s(X0,X4))),s(X1,happ(s(fun(X0,X1),X2),s(X0,X5))),s(X1,'const.bool.ARB_0')))) )),
  inference(cnf_transformation,[],[f601])).
fof(f601,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X1,'const.relation.RESTRICT_4'(s(fun(X0,X1),X2),s(fun(X0,fun(X0,bool)),X3),s(X0,X4),s(X0,X5))) = s(X1,'const.bool.COND_3'(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,X5))),s(X0,X4))),s(X1,happ(s(fun(X0,X1),X2),s(X0,X5))),s(X1,'const.bool.ARB_0')))),
  inference(flattening,[],[f600])).
fof(f600,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5] : s(X1,'const.relation.RESTRICT_4'(s(fun(X0,X1),X2),s(fun(X0,fun(X0,bool)),X3),s(X0,X4),s(X0,X5))) = s(X1,'const.bool.COND_3'(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,X5))),s(X0,X4))),s(X1,happ(s(fun(X0,X1),X2),s(X0,X5))),s(X1,'const.bool.ARB_0')))),
  inference(rectify,[],[f131])).
fof(f131,axiom,(
  ! [X6,X20] : ! [X15,X18,X9,X5] : s(X20,'const.relation.RESTRICT_4'(s(fun(X6,X20),X15),s(fun(X6,fun(X6,bool)),X18),s(X6,X9),s(X6,X5))) = s(X20,'const.bool.COND_3'(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X18),s(X6,X5))),s(X6,X9))),s(X20,happ(s(fun(X6,X20),X15),s(X6,X5))),s(X20,'const.bool.ARB_0')))),
  file('thm.pred_set.EXTENSION','thm.relation.RESTRICT_DEF')).
fof(f18484,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,'const.relation.inv_3'(s(fun(X3,fun(X3,bool)),X4),s(X3,X6),s(X3,X5))) = s(bool,'const.bool.IN_2'(s(X3,X6),s(fun(X3,bool),happ(s(fun(X3,fun(X3,bool)),X4),s(X3,X5)))))) )),
  inference(superposition,[],[f1319,f1257])).
fof(f1319,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'const.relation.inv_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f442])).
fof(f442,plain,(
  ! [X0,X1,X2,X3] : s(bool,'const.relation.inv_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2)))),
  inference(flattening,[],[f441])).
fof(f441,plain,(
  ! [X0] : ! [X1,X2,X3] : s(bool,'const.relation.inv_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2)))),
  inference(rectify,[],[f156])).
fof(f156,axiom,(
  ! [X6] : ! [X18,X9,X19] : s(bool,'const.relation.inv_3'(s(fun(X6,fun(X6,bool)),X18),s(X6,X9),s(X6,X19))) = s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X18),s(X6,X19))),s(X6,X9)))),
  file('thm.pred_set.EXTENSION','thm.relation.inv_DEF')).
fof(f31471,plain,(
  ( ! [X1] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK61(X1,sK14,sK12)),s(fun(sK12,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(sK12,sK60(X1,sK14,sK12)),s(fun(sK12,bool),sK13))))) ) | $spl60),
  inference(cnf_transformation,[],[f31471_D])).
fof(f31471_D,plain,(
  ( ! [X1] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK61(X1,sK14,sK12)),s(fun(sK12,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(sK12,sK60(X1,sK14,sK12)),s(fun(sK12,bool),sK13))))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f31470,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK60(X1,sK14,sK12)),s(fun(sK12,bool),sK13)))) | ~p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13)))) | p(s(bool,'const.bool.IN_2'(s(sK12,sK61(X1,sK14,sK12)),s(fun(sK12,bool),X1))))) ) | $spl4),
  inference(forward_demodulation,[],[f31457,f1694])).
fof(f31457,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13)))) | p(s(bool,'const.bool.IN_2'(s(sK12,sK60(X1,sK14,sK12)),s(fun(sK12,bool),sK14)))) | p(s(bool,'const.bool.IN_2'(s(sK12,sK61(X1,sK14,sK12)),s(fun(sK12,bool),X1))))) ) | $spl4),
  inference(superposition,[],[f1824,f1694])).
fof(f1824,plain,(
  ( ! [X2,X0,X8,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X8),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK60(X2,X1,X0)),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK61(X2,X1,X0)),s(fun(X0,bool),X2))))) )),
  inference(forward_demodulation,[],[f1823,f1257])).
fof(f1823,plain,(
  ( ! [X2,X0,X8,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK60(X2,X1,X0)),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK61(X2,X1,X0)),s(fun(X0,bool),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X8))))) )),
  inference(forward_demodulation,[],[f1822,f1257])).
fof(f1822,plain,(
  ( ! [X2,X0,X8,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK61(X2,X1,X0)),s(fun(X0,bool),X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK60(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X8))))) )),
  inference(forward_demodulation,[],[f1315,f1257])).
fof(f1315,plain,(
  ( ! [X2,X0,X8,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK61(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK60(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X8))))) )),
  inference(cnf_transformation,[],[f961])).
fof(f961,plain,(
  ! [X0,X1,X2] : ((! [X8] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X8)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X8))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK60(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK61(X2,X1,X0)))))) & ((! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK62(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK62(X2,X1,X0))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK60,sK61,sK62])],[f960])).
fof(f960,plain,(
  ! [X0,X1,X2] : ((! [X8] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X8)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X8))))) | ? [X7] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X7)))) | ? [X6] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) & ((! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(rectify,[],[f959])).
fof(f959,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) & ((! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(flattening,[],[f958])).
fof(f958,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & ((! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f440])).
fof(f440,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f439])).
fof(f439,plain,(
  ! [X0] : ! [X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X6] : ! [X17,X16] : (? [X9] : (p(s(bool,happ(s(fun(X6,bool),X17),s(X6,X9)))) | p(s(bool,happ(s(fun(X6,bool),X16),s(X6,X9))))) <=> (? [X9] : p(s(bool,happ(s(fun(X6,bool),X17),s(X6,X9)))) | ? [X9] : p(s(bool,happ(s(fun(X6,bool),X16),s(X6,X9))))))),
  file('thm.pred_set.EXTENSION','thm.bool.EXISTS_OR_THM')).
fof(f31385,plain,(
  ( ! [X1] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK61(sK14,X1,sK12)),s(fun(sK12,bool),sK13)))) | p(s(bool,'const.bool.IN_2'(s(sK12,sK60(sK14,X1,sK12)),s(fun(sK12,bool),X1))))) ) | $spl58),
  inference(cnf_transformation,[],[f31385_D])).
fof(f31385_D,plain,(
  ( ! [X1] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK61(sK14,X1,sK12)),s(fun(sK12,bool),sK13)))) | p(s(bool,'const.bool.IN_2'(s(sK12,sK60(sK14,X1,sK12)),s(fun(sK12,bool),X1))))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f31384,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK61(sK14,X1,sK12)),s(fun(sK12,bool),sK13)))) | ~p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13)))) | p(s(bool,'const.bool.IN_2'(s(sK12,sK60(sK14,X1,sK12)),s(fun(sK12,bool),X1))))) ) | $spl4),
  inference(forward_demodulation,[],[f31371,f1694])).
fof(f31371,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13)))) | p(s(bool,'const.bool.IN_2'(s(sK12,sK60(sK14,X1,sK12)),s(fun(sK12,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(sK12,sK61(sK14,X1,sK12)),s(fun(sK12,bool),sK14))))) ) | $spl4),
  inference(superposition,[],[f1821,f1694])).
fof(f1821,plain,(
  ( ! [X2,X0,X8,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X8),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK60(X2,X1,X0)),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK61(X2,X1,X0)),s(fun(X0,bool),X2))))) )),
  inference(forward_demodulation,[],[f1820,f1257])).
fof(f1820,plain,(
  ( ! [X2,X0,X8,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK60(X2,X1,X0)),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK61(X2,X1,X0)),s(fun(X0,bool),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X8))))) )),
  inference(forward_demodulation,[],[f1819,f1257])).
fof(f1819,plain,(
  ( ! [X2,X0,X8,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK61(X2,X1,X0)),s(fun(X0,bool),X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK60(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X8))))) )),
  inference(forward_demodulation,[],[f1316,f1257])).
fof(f1316,plain,(
  ( ! [X2,X0,X8,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK61(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK60(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X8))))) )),
  inference(cnf_transformation,[],[f961])).
fof(f31295,plain,(
  ( ! [X1] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK62(X1,sK14,sK12)),s(fun(sK12,bool),sK13)))) | p(s(bool,'const.bool.IN_2'(s(sK12,sK62(X1,sK14,sK12)),s(fun(sK12,bool),X1))))) ) | $spl56),
  inference(cnf_transformation,[],[f31295_D])).
fof(f31295_D,plain,(
  ( ! [X1] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK62(X1,sK14,sK12)),s(fun(sK12,bool),sK13)))) | p(s(bool,'const.bool.IN_2'(s(sK12,sK62(X1,sK14,sK12)),s(fun(sK12,bool),X1))))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f31294,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK62(X1,sK14,sK12)),s(fun(sK12,bool),sK13)))) | ~p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13)))) | p(s(bool,'const.bool.IN_2'(s(sK12,sK62(X1,sK14,sK12)),s(fun(sK12,bool),X1))))) ) | $spl4),
  inference(forward_demodulation,[],[f31281,f1694])).
fof(f31281,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13)))) | p(s(bool,'const.bool.IN_2'(s(sK12,sK62(X1,sK14,sK12)),s(fun(sK12,bool),sK14)))) | p(s(bool,'const.bool.IN_2'(s(sK12,sK62(X1,sK14,sK12)),s(fun(sK12,bool),X1))))) ) | $spl4),
  inference(superposition,[],[f1818,f1694])).
fof(f1818,plain,(
  ( ! [X2,X0,X5,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK62(X2,X1,X0)),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK62(X2,X1,X0)),s(fun(X0,bool),X2))))) )),
  inference(forward_demodulation,[],[f1817,f1257])).
fof(f1817,plain,(
  ( ! [X2,X0,X5,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK62(X2,X1,X0)),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK62(X2,X1,X0)),s(fun(X0,bool),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))) )),
  inference(forward_demodulation,[],[f1816,f1257])).
fof(f1816,plain,(
  ( ! [X2,X0,X5,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK62(X2,X1,X0)),s(fun(X0,bool),X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK62(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))) )),
  inference(forward_demodulation,[],[f1317,f1257])).
fof(f1317,plain,(
  ( ! [X2,X0,X5,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK62(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK62(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))) )),
  inference(cnf_transformation,[],[f961])).
fof(f31213,plain,(
  ( ! [X1] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK62(sK14,X1,sK12)),s(fun(sK12,bool),sK13)))) | p(s(bool,'const.bool.IN_2'(s(sK12,sK62(sK14,X1,sK12)),s(fun(sK12,bool),X1))))) ) | $spl54),
  inference(cnf_transformation,[],[f31213_D])).
fof(f31213_D,plain,(
  ( ! [X1] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK62(sK14,X1,sK12)),s(fun(sK12,bool),sK13)))) | p(s(bool,'const.bool.IN_2'(s(sK12,sK62(sK14,X1,sK12)),s(fun(sK12,bool),X1))))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f31212,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK62(sK14,X1,sK12)),s(fun(sK12,bool),sK13)))) | ~p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13)))) | p(s(bool,'const.bool.IN_2'(s(sK12,sK62(sK14,X1,sK12)),s(fun(sK12,bool),X1))))) ) | $spl4),
  inference(forward_demodulation,[],[f31199,f1694])).
fof(f31199,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13)))) | p(s(bool,'const.bool.IN_2'(s(sK12,sK62(sK14,X1,sK12)),s(fun(sK12,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(sK12,sK62(sK14,X1,sK12)),s(fun(sK12,bool),sK14))))) ) | $spl4),
  inference(superposition,[],[f1815,f1694])).
fof(f1815,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK62(X2,X1,X0)),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK62(X2,X1,X0)),s(fun(X0,bool),X2))))) )),
  inference(forward_demodulation,[],[f1814,f1257])).
fof(f1814,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK62(X2,X1,X0)),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK62(X2,X1,X0)),s(fun(X0,bool),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) )),
  inference(forward_demodulation,[],[f1813,f1257])).
fof(f1813,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK62(X2,X1,X0)),s(fun(X0,bool),X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK62(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) )),
  inference(forward_demodulation,[],[f1318,f1257])).
fof(f1318,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK62(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK62(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) )),
  inference(cnf_transformation,[],[f961])).
fof(f31113,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK51(sK14,X0,sK12)),s(fun(sK12,bool),sK13)))) | ~p(s(bool,'const.bool.IN_2'(s(sK12,sK51(sK14,X0,sK12)),s(fun(sK12,bool),X0))))) ) | $spl52),
  inference(cnf_transformation,[],[f31113_D])).
fof(f31113_D,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK51(sK14,X0,sK12)),s(fun(sK12,bool),sK13)))) | ~p(s(bool,'const.bool.IN_2'(s(sK12,sK51(sK14,X0,sK12)),s(fun(sK12,bool),X0))))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f31112,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(sK12,X1),s(fun(sK12,bool),sK13)))) | ~p(s(bool,'const.bool.IN_2'(s(sK12,sK51(sK14,X0,sK12)),s(fun(sK12,bool),sK13)))) | ~p(s(bool,'const.bool.IN_2'(s(sK12,sK51(sK14,X0,sK12)),s(fun(sK12,bool),X0))))) ) | $spl4),
  inference(forward_demodulation,[],[f31099,f1694])).
fof(f31099,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK51(sK14,X0,sK12)),s(fun(sK12,bool),sK13)))) | ~p(s(bool,'const.bool.IN_2'(s(sK12,sK51(sK14,X0,sK12)),s(fun(sK12,bool),X0)))) | p(s(bool,'const.bool.IN_2'(s(sK12,X1),s(fun(sK12,bool),sK14))))) ) | $spl4),
  inference(superposition,[],[f1797,f1694])).
fof(f1797,plain,(
  ( ! [X6,X2,X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,sK51(X2,X1,X0)),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,sK51(X2,X1,X0)),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X2))))) )),
  inference(forward_demodulation,[],[f1796,f1257])).
fof(f1796,plain,(
  ( ! [X6,X2,X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,sK51(X2,X1,X0)),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK51(X2,X1,X0)))))) )),
  inference(forward_demodulation,[],[f1795,f1257])).
fof(f1795,plain,(
  ( ! [X6,X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK51(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK51(X2,X1,X0)))))) )),
  inference(forward_demodulation,[],[f1300,f1257])).
fof(f1300,plain,(
  ( ! [X6,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK51(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK51(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f945])).
fof(f945,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK51(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK51(X2,X1,X0)))))) | (! [X7] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X7)))) & ! [X6] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK52(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK53(X2,X1,X0))))) | ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK51,sK52,sK53])],[f944])).
fof(f944,plain,(
  ! [X0,X1,X2] : ((? [X8] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X8)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X8))))) | (! [X7] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X7)))) & ! [X6] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))))) & (? [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(rectify,[],[f943])).
fof(f943,plain,(
  ! [X0,X1,X2] : ((? [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (? [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(flattening,[],[f942])).
fof(f942,plain,(
  ! [X0,X1,X2] : ((? [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & ((? [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f432])).
fof(f432,plain,(
  ! [X0,X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f431])).
fof(f431,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f93])).
fof(f93,axiom,(
  ! [X6] : ! [X17,X16] : (! [X9] : (p(s(bool,happ(s(fun(X6,bool),X17),s(X6,X9)))) & p(s(bool,happ(s(fun(X6,bool),X16),s(X6,X9))))) <=> (! [X9] : p(s(bool,happ(s(fun(X6,bool),X17),s(X6,X9)))) & ! [X9] : p(s(bool,happ(s(fun(X6,bool),X16),s(X6,X9))))))),
  file('thm.pred_set.EXTENSION','thm.bool.FORALL_AND_THM')).
fof(f30412,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK53(sK14,X0,sK12)),s(fun(sK12,bool),sK13)))) | ~p(s(bool,'const.bool.IN_2'(s(sK12,sK52(sK14,X0,sK12)),s(fun(sK12,bool),X0))))) ) | $spl50),
  inference(cnf_transformation,[],[f30412_D])).
fof(f30412_D,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK53(sK14,X0,sK12)),s(fun(sK12,bool),sK13)))) | ~p(s(bool,'const.bool.IN_2'(s(sK12,sK52(sK14,X0,sK12)),s(fun(sK12,bool),X0))))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f30411,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(sK12,X1),s(fun(sK12,bool),sK13)))) | ~p(s(bool,'const.bool.IN_2'(s(sK12,sK53(sK14,X0,sK12)),s(fun(sK12,bool),sK13)))) | ~p(s(bool,'const.bool.IN_2'(s(sK12,sK52(sK14,X0,sK12)),s(fun(sK12,bool),X0))))) ) | $spl4),
  inference(forward_demodulation,[],[f30398,f1694])).
fof(f30398,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK53(sK14,X0,sK12)),s(fun(sK12,bool),sK13)))) | ~p(s(bool,'const.bool.IN_2'(s(sK12,sK52(sK14,X0,sK12)),s(fun(sK12,bool),X0)))) | p(s(bool,'const.bool.IN_2'(s(sK12,X1),s(fun(sK12,bool),sK14))))) ) | $spl4),
  inference(superposition,[],[f1791,f1694])).
fof(f1791,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,sK53(X2,X1,X0)),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,sK52(X2,X1,X0)),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) )),
  inference(forward_demodulation,[],[f1790,f1257])).
fof(f1790,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,sK53(X2,X1,X0)),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK52(X2,X1,X0)))))) )),
  inference(forward_demodulation,[],[f1789,f1257])).
fof(f1789,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK53(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK52(X2,X1,X0)))))) )),
  inference(forward_demodulation,[],[f1302,f1257])).
fof(f1302,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK53(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK52(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f945])).
fof(f29713,plain,(
  $false | ($spl1 | $spl4 | $spl20)),
  inference(subsumption_resolution,[],[f29690,f29629])).
fof(f29629,plain,(
  ( ! [X4] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,X4),s(fun(sK12,bool),sK14)))) ) | ($spl4 | $spl20)),
  inference(backward_demodulation,[],[f28258,f1694])).
fof(f28258,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13)))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f1104,f7869,f1364])).
fof(f1364,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f774])).
fof(f774,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f773])).
fof(f773,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f491])).
fof(f491,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X10,X11] : ((p(s(bool,X10)) => p(s(bool,X11))) => ((p(s(bool,X11)) => p(s(bool,X10))) => s(bool,X10) = s(bool,X11)))),
  file('thm.pred_set.EXTENSION','thm.bool.IMP_ANTISYM_AX')).
fof(f29690,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK81(sK13,sK14,sK12,bool)),s(fun(sK12,bool),sK14))) | ($spl1 | $spl20)),
  inference(backward_demodulation,[],[f28258,f27233])).
fof(f29712,plain,(
  $false | ($spl1 | $spl4 | $spl20)),
  inference(subsumption_resolution,[],[f29689,f29629])).
fof(f29689,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK81(sK14,sK13,sK12,bool)),s(fun(sK12,bool),sK14))) | ($spl1 | $spl20)),
  inference(backward_demodulation,[],[f28258,f27237])).
fof(f27237,plain,(
  s(bool,'const.bool.IN_2'(s(sK12,sK81(sK14,sK13,sK12,bool)),s(fun(sK12,bool),sK13))) != s(bool,'const.bool.IN_2'(s(sK12,sK81(sK14,sK13,sK12,bool)),s(fun(sK12,bool),sK14))) | $spl1),
  inference(forward_demodulation,[],[f27236,f1257])).
fof(f27236,plain,(
  s(bool,'const.bool.IN_2'(s(sK12,sK81(sK14,sK13,sK12,bool)),s(fun(sK12,bool),sK14))) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK81(sK14,sK13,sK12,bool)))) | $spl1),
  inference(forward_demodulation,[],[f27138,f1257])).
fof(f27138,plain,(
  s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK81(sK14,sK13,sK12,bool)))) != s(bool,happ(s(fun(sK12,bool),sK14),s(sK12,sK81(sK14,sK13,sK12,bool)))) | $spl1),
  inference(unit_resulting_resolution,[],[f1691,f1460])).
fof(f29711,plain,(
  $false | ($spl20 | $spl49)),
  inference(subsumption_resolution,[],[f29688,f1104])).
fof(f29688,plain,(
  ~p(s(bool,'T_0')) | ($spl20 | $spl49)),
  inference(backward_demodulation,[],[f28258,f28136])).
fof(f29710,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f29687,f1104])).
fof(f29687,plain,(
  ~p(s(bool,'T_0')) | ($spl20 | $spl23)),
  inference(backward_demodulation,[],[f28258,f27280])).
fof(f27280,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK55(sK13,'F_0',sK12)),s(fun(sK12,bool),sK13)))) | $spl23),
  inference(unit_resulting_resolution,[],[f1700,f8468,f1802])).
fof(f8468,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK28(sK13,sK12)),s(fun(sK12,bool),sK13)))) | $spl23),
  inference(unit_resulting_resolution,[],[f8357,f1731])).
fof(f29709,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f29686,f1104])).
fof(f29686,plain,(
  ~p(s(bool,'T_0')) | ($spl20 | $spl23)),
  inference(backward_demodulation,[],[f28258,f27737])).
fof(f27737,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK55(sK13,'const.bool.IN_2'(s(sK12,sK28(sK13,sK12)),s(fun(sK12,bool),sK13)),sK12)),s(fun(sK12,bool),sK13)))) | $spl23),
  inference(unit_resulting_resolution,[],[f8468,f8468,f1802])).
fof(f29708,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f29685,f1104])).
fof(f29685,plain,(
  ~p(s(bool,'T_0')) | ($spl20 | $spl23)),
  inference(backward_demodulation,[],[f28258,f27282])).
fof(f27282,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK54(sK13,'F_0',sK12)),s(fun(sK12,bool),sK13)))) | $spl23),
  inference(unit_resulting_resolution,[],[f1700,f8468,f1804])).
fof(f29707,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f29684,f1104])).
fof(f29684,plain,(
  ~p(s(bool,'T_0')) | ($spl20 | $spl23)),
  inference(backward_demodulation,[],[f28258,f27738])).
fof(f27738,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK54(sK13,'const.bool.IN_2'(s(sK12,sK28(sK13,sK12)),s(fun(sK12,bool),sK13)),sK12)),s(fun(sK12,bool),sK13)))) | $spl23),
  inference(unit_resulting_resolution,[],[f8468,f8468,f1804])).
fof(f29706,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f29675,f1115])).
fof(f1115,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(cnf_transformation,[],[f268])).
fof(f268,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(flattening,[],[f73])).
fof(f73,axiom,(
  ~s(bool,'T_0') = s(bool,'F_0')),
  file('thm.pred_set.EXTENSION','thm.bool.BOOL_EQ_DISTINCT|split|1')).
fof(f29675,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl19 | $spl20)),
  inference(backward_demodulation,[],[f28258,f9247])).
fof(f9247,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK29(sK13,sK12)),s(fun(sK12,bool),sK13))) | $spl19),
  inference(unit_resulting_resolution,[],[f1700,f7870,f1361])).
fof(f1361,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f774])).
fof(f7870,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK29(sK13,sK12)),s(fun(sK12,bool),sK13)))) | $spl19),
  inference(unit_resulting_resolution,[],[f7868,f1730])).
fof(f29705,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f29674,f1115])).
fof(f29674,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl19 | $spl20)),
  inference(backward_demodulation,[],[f28258,f7875])).
fof(f7875,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13))) | $spl19),
  inference(unit_resulting_resolution,[],[f1700,f7868,f1361])).
fof(f29704,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f29673,f1115])).
fof(f29673,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl20 | $spl23)),
  inference(backward_demodulation,[],[f28258,f27290])).
fof(f27290,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK28(sK13,sK12)),s(fun(sK12,bool),sK13))) | $spl23),
  inference(unit_resulting_resolution,[],[f1700,f8468,f1361])).
fof(f29703,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f29672,f1115])).
fof(f29672,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl20 | $spl23)),
  inference(backward_demodulation,[],[f28258,f8474])).
fof(f8474,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13))) | $spl23),
  inference(unit_resulting_resolution,[],[f1700,f8357,f1361])).
fof(f29628,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28276,f1104])).
fof(f28276,plain,(
  ~p(s(bool,'T_0')) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f7891,f7869,f1412])).
fof(f1412,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f1000])).
fof(f7891,plain,(
  sP1('const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)),'T_0') | $spl19),
  inference(unit_resulting_resolution,[],[f1104,f7868,f1416])).
fof(f29627,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28278,f1104])).
fof(f28278,plain,(
  ~p(s(bool,'T_0')) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f9262,f7869,f1412])).
fof(f9262,plain,(
  sP1('const.bool.IN_2'(s(sK12,sK29(sK13,sK12)),s(fun(sK12,bool),sK13)),'T_0') | $spl19),
  inference(unit_resulting_resolution,[],[f1104,f7870,f1416])).
fof(f29626,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28280,f1104])).
fof(f28280,plain,(
  ~p(s(bool,'T_0')) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8490,f7869,f1412])).
fof(f8490,plain,(
  sP1('const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),'T_0') | $spl23),
  inference(unit_resulting_resolution,[],[f1104,f8357,f1416])).
fof(f29625,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28284,f1104])).
fof(f28284,plain,(
  ~p(s(bool,'T_0')) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f7890,f7869,f1412])).
fof(f7890,plain,(
  sP1('T_0','const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13))) | $spl19),
  inference(unit_resulting_resolution,[],[f1104,f7868,f1413])).
fof(f29624,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28286,f1104])).
fof(f28286,plain,(
  ~p(s(bool,'T_0')) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f9261,f7869,f1412])).
fof(f9261,plain,(
  sP1('T_0','const.bool.IN_2'(s(sK12,sK29(sK13,sK12)),s(fun(sK12,bool),sK13))) | $spl19),
  inference(unit_resulting_resolution,[],[f1104,f7870,f1413])).
fof(f29623,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28288,f1104])).
fof(f28288,plain,(
  ~p(s(bool,'T_0')) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8489,f7869,f1412])).
fof(f8489,plain,(
  sP1('T_0','const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13))) | $spl23),
  inference(unit_resulting_resolution,[],[f1104,f8357,f1413])).
fof(f29622,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28309,f1104])).
fof(f28309,plain,(
  ~p(s(bool,'T_0')) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f8579,f7869,f1543])).
fof(f1543,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f1067])).
fof(f1067,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | (sP2(X2,X1,X0) & sP3(X1,X2,X0) & sP4(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP2(X2,X1,X0) | ~sP3(X1,X2,X0) | ~sP4(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f1066])).
fof(f1066,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP5(X2,X1,X0)))),
  inference(flattening,[],[f1065])).
fof(f1065,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP5(X2,X1,X0)))),
  inference(nnf_transformation,[],[f828])).
fof(f8579,plain,(
  sP5('const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),'T_0','F_0') | $spl23),
  inference(unit_resulting_resolution,[],[f1115,f8357,f1559])).
fof(f29621,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28312,f1104])).
fof(f28312,plain,(
  ~p(s(bool,'T_0')) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f7980,f7869,f1543])).
fof(f7980,plain,(
  sP5('const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)),'T_0','F_0') | $spl19),
  inference(unit_resulting_resolution,[],[f1115,f7868,f1559])).
fof(f29620,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28317,f1104])).
fof(f28317,plain,(
  ~p(s(bool,'T_0')) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f7895,f7869,f1543])).
fof(f7895,plain,(
  sP5('T_0','const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)),'F_0') | $spl19),
  inference(unit_resulting_resolution,[],[f1913,f1908,f1916,f7868,f1545])).
fof(f1545,plain,(
  ( ! [X2,X0,X1] : (~sP2(X2,X1,X0) | p(s(bool,X1)) | ~sP4(X0,X2,X1) | ~sP3(X1,X2,X0) | sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f1067])).
fof(f1916,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f1104,f1556])).
fof(f1556,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f1075])).
fof(f1075,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(flattening,[],[f1074])).
fof(f1074,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(nnf_transformation,[],[f825])).
fof(f1908,plain,(
  ( ! [X0,X1] : (sP4(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f1700,f1549])).
fof(f1549,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f1070])).
fof(f1070,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f1069])).
fof(f1069,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(flattening,[],[f1068])).
fof(f1068,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(nnf_transformation,[],[f827])).
fof(f1913,plain,(
  ( ! [X0,X1] : (sP3(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f1700,f1553])).
fof(f1553,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f1073])).
fof(f1073,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f1072])).
fof(f1072,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(flattening,[],[f1071])).
fof(f1071,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(nnf_transformation,[],[f826])).
fof(f29619,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28321,f1104])).
fof(f28321,plain,(
  ~p(s(bool,'T_0')) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f8494,f7869,f1543])).
fof(f8494,plain,(
  sP5('T_0','const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),'F_0') | $spl23),
  inference(unit_resulting_resolution,[],[f1913,f1908,f1916,f8357,f1545])).
fof(f29618,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28325,f1916])).
fof(f28325,plain,(
  ~sP2('const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)),'T_0','T_0') | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f7892,f7946,f7935,f7869,f1544])).
fof(f1544,plain,(
  ( ! [X2,X0,X1] : (~sP2(X2,X1,X0) | ~p(s(bool,X2)) | ~sP4(X0,X2,X1) | ~sP3(X1,X2,X0) | sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f1067])).
fof(f7935,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)),X1)) ) | $spl19),
  inference(unit_resulting_resolution,[],[f7868,f1549])).
fof(f7946,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)),X1)) ) | $spl19),
  inference(unit_resulting_resolution,[],[f7868,f1553])).
fof(f7892,plain,(
  ~sP5('T_0','T_0','const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13))) | $spl19),
  inference(unit_resulting_resolution,[],[f1104,f1104,f7868,f1543])).
fof(f29617,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28326,f28559])).
fof(f28559,plain,(
  ( ! [X2,X0,X1] : (sP2('const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13)),X1,X2)) ) | $spl20),
  inference(unit_resulting_resolution,[],[f7869,f1558])).
fof(f1558,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f1075])).
fof(f28326,plain,(
  ~sP2('const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)),'const.combin.I_1'(s(bool,'const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)))),'F_0') | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f2781,f7946,f7935,f7869,f1544])).
fof(f2781,plain,(
  ( ! [X0] : (~sP5('F_0','const.combin.I_1'(s(bool,X0)),X0)) )),
  inference(unit_resulting_resolution,[],[f1700,f1188,f1562])).
fof(f29616,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28327,f1916])).
fof(f28327,plain,(
  ~sP2('const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)),'T_0','T_0') | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f7892,f1911,f7935,f7869,f1544])).
fof(f1911,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f1104,f1552])).
fof(f1552,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f1073])).
fof(f29615,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28328,f1916])).
fof(f28328,plain,(
  ~sP2('const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),'T_0','T_0') | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8491,f8545,f8534,f7869,f1544])).
fof(f8534,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),X1)) ) | $spl23),
  inference(unit_resulting_resolution,[],[f8357,f1549])).
fof(f8545,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),X1)) ) | $spl23),
  inference(unit_resulting_resolution,[],[f8357,f1553])).
fof(f8491,plain,(
  ~sP5('T_0','T_0','const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13))) | $spl23),
  inference(unit_resulting_resolution,[],[f1104,f1104,f8357,f1543])).
fof(f29614,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28329,f28559])).
fof(f28329,plain,(
  ~sP2('const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),'const.combin.I_1'(s(bool,'const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)))),'F_0') | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f2781,f8545,f8534,f7869,f1544])).
fof(f29613,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28330,f1916])).
fof(f28330,plain,(
  ~sP2('const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),'T_0','T_0') | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8491,f1911,f8534,f7869,f1544])).
fof(f29612,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28331,f1916])).
fof(f28331,plain,(
  ~sP2('const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)),'T_0','T_0') | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f7892,f7946,f1906,f7869,f1544])).
fof(f1906,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f1104,f1548])).
fof(f1548,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f1070])).
fof(f29611,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28332,f1916])).
fof(f28332,plain,(
  ~sP2('const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),'T_0','T_0') | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8491,f8545,f1906,f7869,f1544])).
fof(f29610,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28333,f1916])).
fof(f28333,plain,(
  ~sP2('const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)),'T_0','T_0') | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f7892,f1911,f1906,f7869,f1544])).
fof(f29609,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28334,f1916])).
fof(f28334,plain,(
  ~sP2('const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),'T_0','T_0') | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8491,f1911,f1906,f7869,f1544])).
fof(f29608,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28335,f28559])).
fof(f28335,plain,(
  ~sP2('const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)),'const.combin.I_1'(s(bool,'const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)))),'F_0') | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f2781,f7946,f1910,f7869,f1544])).
fof(f1910,plain,(
  ( ! [X0,X1] : (sP4('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f1700,f1550])).
fof(f1550,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f1070])).
fof(f29607,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28336,f28559])).
fof(f28336,plain,(
  ~sP2('const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),'const.combin.I_1'(s(bool,'const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)))),'F_0') | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f2781,f8545,f1910,f7869,f1544])).
fof(f29606,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28337,f1906])).
fof(f28337,plain,(
  ~sP4('T_0','const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)),'T_0') | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f7892,f7946,f1916,f7869,f1544])).
fof(f29605,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28338,f1906])).
fof(f28338,plain,(
  ~sP4('T_0','const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),'T_0') | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8491,f8545,f1916,f7869,f1544])).
fof(f29604,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28339,f1906])).
fof(f28339,plain,(
  ~sP4('T_0','const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)),'T_0') | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f7892,f1911,f1916,f7869,f1544])).
fof(f29603,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28340,f1906])).
fof(f28340,plain,(
  ~sP4('T_0','const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),'T_0') | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8491,f1911,f1916,f7869,f1544])).
fof(f29602,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28344,f1911])).
fof(f28344,plain,(
  ~sP3('T_0','const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)),'T_0') | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f7892,f7935,f1916,f7869,f1544])).
fof(f29601,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28355,f1911])).
fof(f28355,plain,(
  ~sP3('T_0','const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),'T_0') | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8491,f8534,f1916,f7869,f1544])).
fof(f29600,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28366,f1911])).
fof(f28366,plain,(
  ~sP3('T_0','const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)),'T_0') | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f7892,f1906,f1916,f7869,f1544])).
fof(f29599,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28367,f1911])).
fof(f28367,plain,(
  ~sP3('T_0','const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),'T_0') | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8491,f1906,f1916,f7869,f1544])).
fof(f29598,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28399,f1906])).
fof(f28399,plain,(
  ~sP4('T_0','const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)),'T_0') | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f7892,f7946,f1917,f7869,f1544])).
fof(f1917,plain,(
  ( ! [X0,X1] : (sP2(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f1104,f1557])).
fof(f1557,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f1075])).
fof(f29597,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28400,f1906])).
fof(f28400,plain,(
  ~sP4('T_0','const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),'T_0') | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8491,f8545,f1917,f7869,f1544])).
fof(f29596,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28401,f1906])).
fof(f28401,plain,(
  ~sP4('T_0','const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)),'T_0') | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f7892,f1911,f1917,f7869,f1544])).
fof(f29595,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28402,f1906])).
fof(f28402,plain,(
  ~sP4('T_0','const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),'T_0') | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8491,f1911,f1917,f7869,f1544])).
fof(f29594,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28403,f1911])).
fof(f28403,plain,(
  ~sP3('T_0','const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)),'T_0') | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f7892,f7935,f1917,f7869,f1544])).
fof(f29593,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28409,f1911])).
fof(f28409,plain,(
  ~sP3('T_0','const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),'T_0') | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8491,f8534,f1917,f7869,f1544])).
fof(f29592,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28415,f1911])).
fof(f28415,plain,(
  ~sP3('T_0','const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)),'T_0') | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f7892,f1906,f1917,f7869,f1544])).
fof(f29591,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28416,f1911])).
fof(f28416,plain,(
  ~sP3('T_0','const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),'T_0') | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8491,f1906,f1917,f7869,f1544])).
fof(f29590,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28462,f29272])).
fof(f29272,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1104,f8187,f7869,f1611])).
fof(f1611,plain,(
  ( ! [X2,X0,X1] : (~sP11(X2,X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f1099])).
fof(f8187,plain,(
  ( ! [X0] : (sP11('T_0',X0,'const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)))) ) | $spl19),
  inference(unit_resulting_resolution,[],[f1104,f7868,f1609])).
fof(f28462,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f7936,f7869,f1547])).
fof(f1547,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f1070])).
fof(f7936,plain,(
  ( ! [X0,X1] : (sP4('const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)),X0,X1)) ) | $spl19),
  inference(unit_resulting_resolution,[],[f7868,f1550])).
fof(f29589,plain,(
  $false | ($spl19 | $spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28466,f29272])).
fof(f28466,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f8535,f7869,f1547])).
fof(f8535,plain,(
  ( ! [X0,X1] : (sP4('const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),X0,X1)) ) | $spl23),
  inference(unit_resulting_resolution,[],[f8357,f1550])).
fof(f29588,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28472,f7869])).
fof(f28472,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)))) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f7935,f7869,f1547])).
fof(f29587,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28475,f7869])).
fof(f28475,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)))) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f8534,f7869,f1547])).
fof(f29586,plain,(
  $false | ($spl7 | $spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28478,f29272])).
fof(f28478,plain,(
  ~p(s(bool,sK16)) | ($spl7 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f1907,f7869,f1547])).
fof(f1907,plain,(
  ( ! [X0,X1] : (sP4(X0,sK16,X1)) ) | $spl7),
  inference(unit_resulting_resolution,[],[f1696,f1549])).
fof(f29585,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28479,f29272])).
fof(f28479,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK63(X0,X1,'T_0',X2)))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f1700,f2481,f7869,f1547])).
fof(f2481,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP4(X0,sK63(X1,X2,'T_0',X3),X4)) )),
  inference(unit_resulting_resolution,[],[f2471,f1549])).
fof(f2471,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK63(X0,X1,'T_0',X2)))) )),
  inference(unit_resulting_resolution,[],[f1115,f1344])).
fof(f1344,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,sK63(X3,X2,X1,X0))) | s(bool,'F_0') = s(bool,X1)) )),
  inference(cnf_transformation,[],[f972])).
fof(f972,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,sK63(X3,X2,X1,X0))) | s(bool,'F_0') = s(bool,X1)) & (s(bool,'F_0') != s(bool,X1) | p(s(bool,sK63(X3,X2,X1,X0)))) & ((~p(s(bool,sK64(X3,X2,X1,X0))) | s(bool,'T_0') = s(bool,X1)) & (s(bool,'T_0') != s(bool,X1) | p(s(bool,sK64(X3,X2,X1,X0)))) & s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.min.@_1'(s(fun(X0,bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(X0,bool),'_dst_x0x1_2'(s(fun(bool,bool),'==>_1'(s(bool,sK64(X3,X2,X1,X0)))),s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X2))))),s(fun(X0,bool),'_dst_x0x1_2'(s(fun(bool,bool),'==>_1'(s(bool,sK63(X3,X2,X1,X0)))),s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X3)))))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK63,sK64])],[f971])).
fof(f971,plain,(
  ! [X0,X1,X2,X3] : ? [X4] : ((~p(s(bool,X4)) | s(bool,'F_0') = s(bool,X1)) & (s(bool,'F_0') != s(bool,X1) | p(s(bool,X4))) & ? [X5] : ((~p(s(bool,X5)) | s(bool,'T_0') = s(bool,X1)) & (s(bool,'T_0') != s(bool,X1) | p(s(bool,X5))) & s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.min.@_1'(s(fun(X0,bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(X0,bool),'_dst_x0x1_2'(s(fun(bool,bool),'==>_1'(s(bool,X5))),s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X2))))),s(fun(X0,bool),'_dst_x0x1_2'(s(fun(bool,bool),'==>_1'(s(bool,X4))),s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X3)))))))))))),
  inference(flattening,[],[f970])).
fof(f970,plain,(
  ! [X0,X1,X2,X3] : ? [X4] : (((~p(s(bool,X4)) | s(bool,'F_0') = s(bool,X1)) & (s(bool,'F_0') != s(bool,X1) | p(s(bool,X4)))) & ? [X5] : (((~p(s(bool,X5)) | s(bool,'T_0') = s(bool,X1)) & (s(bool,'T_0') != s(bool,X1) | p(s(bool,X5)))) & s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.min.@_1'(s(fun(X0,bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(X0,bool),'_dst_x0x1_2'(s(fun(bool,bool),'==>_1'(s(bool,X5))),s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X2))))),s(fun(X0,bool),'_dst_x0x1_2'(s(fun(bool,bool),'==>_1'(s(bool,X4))),s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X3)))))))))))),
  inference(nnf_transformation,[],[f476])).
fof(f476,plain,(
  ! [X0,X1,X2,X3] : ? [X4] : ((p(s(bool,X4)) <=> s(bool,'F_0') = s(bool,X1)) & ? [X5] : ((p(s(bool,X5)) <=> s(bool,'T_0') = s(bool,X1)) & s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.min.@_1'(s(fun(X0,bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(X0,bool),'_dst_x0x1_2'(s(fun(bool,bool),'==>_1'(s(bool,X5))),s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X2))))),s(fun(X0,bool),'_dst_x0x1_2'(s(fun(bool,bool),'==>_1'(s(bool,X4))),s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X3)))))))))))),
  inference(flattening,[],[f475])).
fof(f475,plain,(
  ! [X0] : ! [X1,X2,X3] : ? [X4] : ((p(s(bool,X4)) <=> s(bool,'F_0') = s(bool,X1)) & ? [X5] : ((p(s(bool,X5)) <=> s(bool,'T_0') = s(bool,X1)) & s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.min.@_1'(s(fun(X0,bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(X0,bool),'_dst_x0x1_2'(s(fun(bool,bool),'==>_1'(s(bool,X5))),s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X2))))),s(fun(X0,bool),'_dst_x0x1_2'(s(fun(bool,bool),'==>_1'(s(bool,X4))),s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X3)))))))))))),
  inference(rectify,[],[f70])).
fof(f70,axiom,(
  ! [X6] : ! [X5,X13,X12] : ? [X46] : ((p(s(bool,X46)) <=> s(bool,'F_0') = s(bool,X5)) & ? [X47] : ((p(s(bool,X47)) <=> s(bool,'T_0') = s(bool,X5)) & s(X6,'const.bool.COND_3'(s(bool,X5),s(X6,X13),s(X6,X12))) = s(X6,'const.min.@_1'(s(fun(X6,bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(X6,bool),'_dst_x0x1_2'(s(fun(bool,bool),'==>_1'(s(bool,X47))),s(fun(X6,bool),'_dst_x1x0_2'(s(fun(X6,fun(X6,bool)),'=_0'),s(X6,X13))))),s(fun(X6,bool),'_dst_x0x1_2'(s(fun(bool,bool),'==>_1'(s(bool,X46))),s(fun(X6,bool),'_dst_x1x0_2'(s(fun(X6,fun(X6,bool)),'=_0'),s(X6,X12)))))))))))),
  file('thm.pred_set.EXTENSION','thm.bool.COND_DEF')).
fof(f29584,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28480,f29272])).
fof(f28480,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK64(X0,X1,'F_0',X2)))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f1700,f2584,f7869,f1547])).
fof(f2584,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP4(X0,sK64(X1,X2,'F_0',X3),X4)) )),
  inference(unit_resulting_resolution,[],[f2552,f1549])).
fof(f2552,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK64(X0,X1,'F_0',X2)))) )),
  inference(unit_resulting_resolution,[],[f1115,f1346])).
fof(f1346,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,sK64(X3,X2,X1,X0))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(cnf_transformation,[],[f972])).
fof(f29583,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28485,f29272])).
fof(f28485,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f7935,f7869,f1547])).
fof(f29582,plain,(
  $false | ($spl19 | $spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28489,f29272])).
fof(f28489,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f8534,f7869,f1547])).
fof(f29581,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28496,f7869])).
fof(f28496,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)))) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f7936,f7869,f1547])).
fof(f29580,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28499,f7869])).
fof(f28499,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)))) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f8535,f7869,f1547])).
fof(f29579,plain,(
  $false | ($spl7 | $spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28502,f29272])).
fof(f28502,plain,(
  ~p(s(bool,sK16)) | ($spl7 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f1909,f7869,f1547])).
fof(f1909,plain,(
  ( ! [X0,X1] : (sP4(sK16,X0,X1)) ) | $spl7),
  inference(unit_resulting_resolution,[],[f1696,f1550])).
fof(f29578,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28503,f29272])).
fof(f28503,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK63(X0,X1,'T_0',X2)))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f1700,f2482,f7869,f1547])).
fof(f2482,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP4(sK63(X0,X1,'T_0',X2),X3,X4)) )),
  inference(unit_resulting_resolution,[],[f2471,f1550])).
fof(f29577,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28504,f29272])).
fof(f28504,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK64(X0,X1,'F_0',X2)))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f1700,f2585,f7869,f1547])).
fof(f2585,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP4(sK64(X0,X1,'F_0',X2),X3,X4)) )),
  inference(unit_resulting_resolution,[],[f2552,f1550])).
fof(f29576,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28511,f29272])).
fof(f28511,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f7947,f7869,f1551])).
fof(f1551,plain,(
  ( ! [X2,X0,X1] : (~sP3(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f1073])).
fof(f7947,plain,(
  ( ! [X0,X1] : (sP3('const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)),X0,X1)) ) | $spl19),
  inference(unit_resulting_resolution,[],[f7868,f1554])).
fof(f1554,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f1073])).
fof(f29575,plain,(
  $false | ($spl19 | $spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28515,f29272])).
fof(f28515,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f8546,f7869,f1551])).
fof(f8546,plain,(
  ( ! [X0,X1] : (sP3('const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),X0,X1)) ) | $spl23),
  inference(unit_resulting_resolution,[],[f8357,f1554])).
fof(f29574,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28521,f7869])).
fof(f28521,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)))) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f7946,f7869,f1551])).
fof(f29573,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28524,f7869])).
fof(f28524,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)))) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f8545,f7869,f1551])).
fof(f29572,plain,(
  $false | ($spl7 | $spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28527,f29272])).
fof(f28527,plain,(
  ~p(s(bool,sK16)) | ($spl7 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f1912,f7869,f1551])).
fof(f1912,plain,(
  ( ! [X0,X1] : (sP3(X0,sK16,X1)) ) | $spl7),
  inference(unit_resulting_resolution,[],[f1696,f1553])).
fof(f29571,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28528,f29272])).
fof(f28528,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK63(X0,X1,'T_0',X2)))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f1700,f2483,f7869,f1551])).
fof(f2483,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP3(X0,sK63(X1,X2,'T_0',X3),X4)) )),
  inference(unit_resulting_resolution,[],[f2471,f1553])).
fof(f29570,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28529,f29272])).
fof(f28529,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK64(X0,X1,'F_0',X2)))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f1700,f2586,f7869,f1551])).
fof(f2586,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP3(X0,sK64(X1,X2,'F_0',X3),X4)) )),
  inference(unit_resulting_resolution,[],[f2552,f1553])).
fof(f29569,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28534,f29272])).
fof(f28534,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f7946,f7869,f1551])).
fof(f29568,plain,(
  $false | ($spl19 | $spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28538,f29272])).
fof(f28538,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f8545,f7869,f1551])).
fof(f29567,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28544,f7869])).
fof(f28544,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)))) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f7947,f7869,f1551])).
fof(f29566,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28547,f7869])).
fof(f28547,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)))) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f8546,f7869,f1551])).
fof(f29565,plain,(
  $false | ($spl7 | $spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28550,f29272])).
fof(f28550,plain,(
  ~p(s(bool,sK16)) | ($spl7 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f1914,f7869,f1551])).
fof(f1914,plain,(
  ( ! [X0,X1] : (sP3(sK16,X0,X1)) ) | $spl7),
  inference(unit_resulting_resolution,[],[f1696,f1554])).
fof(f29564,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28551,f29272])).
fof(f28551,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK63(X0,X1,'T_0',X2)))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f1700,f2484,f7869,f1551])).
fof(f2484,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP3(sK63(X0,X1,'T_0',X2),X3,X4)) )),
  inference(unit_resulting_resolution,[],[f2471,f1554])).
fof(f29563,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28552,f29272])).
fof(f28552,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK64(X0,X1,'F_0',X2)))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f1700,f2587,f7869,f1551])).
fof(f2587,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP3(sK64(X0,X1,'F_0',X2),X3,X4)) )),
  inference(unit_resulting_resolution,[],[f2552,f1554])).
fof(f29562,plain,(
  $false | ($spl7 | $spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28611,f1936])).
fof(f1936,plain,(
  s(bool,'F_0') = s(bool,sK16) | $spl7),
  inference(unit_resulting_resolution,[],[f1696,f1130])).
fof(f1130,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f730])).
fof(f730,plain,(
  ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0))),
  inference(ennf_transformation,[],[f283])).
fof(f283,plain,(
  ! [X0] : (~p(s(bool,X0)) => s(bool,'F_0') = s(bool,X0))),
  inference(flattening,[],[f282])).
fof(f282,plain,(
  ! [X0] : (~p(s(bool,X0)) => s(bool,'F_0') = s(bool,X0))),
  inference(rectify,[],[f76])).
fof(f76,axiom,(
  ! [X8] : (~p(s(bool,X8)) => s(bool,'F_0') = s(bool,X8))),
  file('thm.pred_set.EXTENSION','thm.bool.NOT_F')).
fof(f28611,plain,(
  s(bool,'F_0') != s(bool,sK16) | ($spl7 | $spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8583,f7869,f1560])).
fof(f8583,plain,(
  ~sP5('const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),sK16,'F_0') | ($spl7 | $spl23)),
  inference(unit_resulting_resolution,[],[f1936,f8357,f1562])).
fof(f29561,plain,(
  $false | ($spl7 | $spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28613,f1936])).
fof(f28613,plain,(
  s(bool,'F_0') != s(bool,sK16) | ($spl7 | $spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8598,f7869,f1560])).
fof(f8598,plain,(
  ~sP5('const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),'F_0',sK16) | ($spl7 | $spl23)),
  inference(unit_resulting_resolution,[],[f1936,f8357,f1562])).
fof(f29560,plain,(
  $false | ($spl7 | $spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28615,f1936])).
fof(f28615,plain,(
  s(bool,'F_0') != s(bool,sK16) | ($spl7 | $spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f7984,f7869,f1560])).
fof(f7984,plain,(
  ~sP5('const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)),sK16,'F_0') | ($spl7 | $spl19)),
  inference(unit_resulting_resolution,[],[f1936,f7868,f1562])).
fof(f29559,plain,(
  $false | ($spl7 | $spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28617,f1936])).
fof(f28617,plain,(
  s(bool,'F_0') != s(bool,sK16) | ($spl7 | $spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f7999,f7869,f1560])).
fof(f7999,plain,(
  ~sP5('const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)),'F_0',sK16) | ($spl7 | $spl19)),
  inference(unit_resulting_resolution,[],[f1936,f7868,f1562])).
fof(f29558,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28623,f1115])).
fof(f28623,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8577,f7869,f1561])).
fof(f8577,plain,(
  sP5('const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),'F_0','T_0') | $spl23),
  inference(unit_resulting_resolution,[],[f1115,f8357,f1559])).
fof(f29557,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28625,f1115])).
fof(f28625,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8579,f7869,f1561])).
fof(f29556,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28627,f1115])).
fof(f28627,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f7978,f7869,f1561])).
fof(f7978,plain,(
  sP5('const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)),'F_0','T_0') | $spl19),
  inference(unit_resulting_resolution,[],[f1115,f7868,f1559])).
fof(f29555,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28629,f1115])).
fof(f28629,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f7980,f7869,f1561])).
fof(f29554,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28671,f1919])).
fof(f1919,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f1104,f1571])).
fof(f1571,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f1081])).
fof(f1081,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(flattening,[],[f1080])).
fof(f1080,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(nnf_transformation,[],[f830])).
fof(f830,plain,(
  ! [X0,X1,X2] : (sP6(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f28671,plain,(
  ~sP6('const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)),'T_0','T_0') | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1104,f8021,f7869,f1566])).
fof(f1566,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f1079])).
fof(f1079,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | (sP6(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP6(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f1078])).
fof(f1078,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP7(X2,X0,X1)))),
  inference(flattening,[],[f1077])).
fof(f1077,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP7(X2,X0,X1)))),
  inference(nnf_transformation,[],[f831])).
fof(f831,plain,(
  ! [X2,X0,X1] : (sP7(X2,X0,X1) <=> (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f8021,plain,(
  ( ! [X0] : (~sP7('T_0','const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)),X0)) ) | $spl19),
  inference(unit_resulting_resolution,[],[f1104,f7868,f1565])).
fof(f1565,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f1079])).
fof(f29553,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28672,f1919])).
fof(f28672,plain,(
  ( ! [X0] : (~sP6('const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)),'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13)),'T_0')) ) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f7869,f8021,f7869,f1566])).
fof(f29552,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28673,f29272])).
fof(f28673,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f8040,f8021,f7869,f1566])).
fof(f8040,plain,(
  ( ! [X0,X1] : (sP6('const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)),X0,X1)) ) | $spl19),
  inference(unit_resulting_resolution,[],[f7868,f1573])).
fof(f1573,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f1081])).
fof(f29551,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28675,f29272])).
fof(f28675,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1919,f8021,f7869,f1566])).
fof(f29550,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28678,f7869])).
fof(f28678,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)))) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f8039,f8021,f7869,f1566])).
fof(f8039,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)),X1)) ) | $spl19),
  inference(unit_resulting_resolution,[],[f7868,f1572])).
fof(f1572,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f1081])).
fof(f29549,plain,(
  $false | ($spl19 | $spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28680,f7869])).
fof(f28680,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)))) | ($spl19 | $spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8638,f8021,f7869,f1566])).
fof(f8638,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),X1)) ) | $spl23),
  inference(unit_resulting_resolution,[],[f8357,f1572])).
fof(f29548,plain,(
  $false | ($spl7 | $spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28682,f29272])).
fof(f28682,plain,(
  ~p(s(bool,sK16)) | ($spl7 | $spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1920,f8021,f7869,f1566])).
fof(f1920,plain,(
  ( ! [X0,X1] : (sP6(X0,sK16,X1)) ) | $spl7),
  inference(unit_resulting_resolution,[],[f1696,f1572])).
fof(f29547,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28683,f29272])).
fof(f28683,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK63(X0,X1,'T_0',X2)))) ) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f2491,f8021,f7869,f1566])).
fof(f2491,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP6(X0,sK63(X1,X2,'T_0',X3),X4)) )),
  inference(unit_resulting_resolution,[],[f2471,f1572])).
fof(f29546,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28684,f29272])).
fof(f28684,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK64(X0,X1,'F_0',X2)))) ) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f2594,f8021,f7869,f1566])).
fof(f2594,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP6(X0,sK64(X1,X2,'F_0',X3),X4)) )),
  inference(unit_resulting_resolution,[],[f2552,f1572])).
fof(f29545,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28685,f1919])).
fof(f28685,plain,(
  ~sP6('const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),'T_0','T_0') | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1104,f8620,f7869,f1566])).
fof(f8620,plain,(
  ( ! [X0] : (~sP7('T_0','const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),X0)) ) | $spl23),
  inference(unit_resulting_resolution,[],[f1104,f8357,f1565])).
fof(f29544,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28686,f1919])).
fof(f28686,plain,(
  ( ! [X0] : (~sP6('const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13)),'T_0')) ) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f7869,f8620,f7869,f1566])).
fof(f29543,plain,(
  $false | ($spl19 | $spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28687,f29272])).
fof(f28687,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8639,f8620,f7869,f1566])).
fof(f8639,plain,(
  ( ! [X0,X1] : (sP6('const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),X0,X1)) ) | $spl23),
  inference(unit_resulting_resolution,[],[f8357,f1573])).
fof(f29542,plain,(
  $false | ($spl19 | $spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28689,f29272])).
fof(f28689,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1919,f8620,f7869,f1566])).
fof(f29541,plain,(
  $false | ($spl19 | $spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28692,f7869])).
fof(f28692,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)))) | ($spl19 | $spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8039,f8620,f7869,f1566])).
fof(f29540,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28694,f7869])).
fof(f28694,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)))) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8638,f8620,f7869,f1566])).
fof(f29539,plain,(
  $false | ($spl7 | $spl19 | $spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28696,f29272])).
fof(f28696,plain,(
  ~p(s(bool,sK16)) | ($spl7 | $spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1920,f8620,f7869,f1566])).
fof(f29538,plain,(
  $false | ($spl19 | $spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28697,f29272])).
fof(f28697,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK63(X0,X1,'T_0',X2)))) ) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f2491,f8620,f7869,f1566])).
fof(f29537,plain,(
  $false | ($spl19 | $spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28698,f29272])).
fof(f28698,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK64(X0,X1,'F_0',X2)))) ) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f2594,f8620,f7869,f1566])).
fof(f29536,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28700,f1919])).
fof(f28700,plain,(
  ( ! [X0] : (~sP6('const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13)),'const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),'T_0')) ) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f7869,f8611,f7869,f1566])).
fof(f8611,plain,(
  ( ! [X0] : (~sP7('T_0',X0,'const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)))) ) | $spl23),
  inference(unit_resulting_resolution,[],[f1104,f8357,f1564])).
fof(f1564,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f1079])).
fof(f29535,plain,(
  $false | ($spl19 | $spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28701,f7869])).
fof(f28701,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)))) | ($spl19 | $spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8040,f8611,f7869,f1566])).
fof(f29534,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28703,f1919])).
fof(f28703,plain,(
  ( ! [X0] : (~sP6('const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13)),'const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)),'T_0')) ) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f7869,f8012,f7869,f1566])).
fof(f8012,plain,(
  ( ! [X0] : (~sP7('T_0',X0,'const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)))) ) | $spl19),
  inference(unit_resulting_resolution,[],[f1104,f7868,f1564])).
fof(f29533,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28704,f7869])).
fof(f28704,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)))) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f8040,f8012,f7869,f1566])).
fof(f29532,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28706,f29272])).
fof(f28706,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK63(X0,X1,'T_0',X2)))) ) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f8040,f2485,f7869,f1566])).
fof(f2485,plain,(
  ( ! [X2,X0,X3,X1] : (~sP7('T_0',X0,sK63(X1,X2,'T_0',X3))) )),
  inference(unit_resulting_resolution,[],[f1104,f2471,f1564])).
fof(f29531,plain,(
  $false | ($spl19 | $spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28707,f29272])).
fof(f28707,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK63(X0,X1,'T_0',X2)))) ) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8639,f2485,f7869,f1566])).
fof(f29530,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28708,f29272])).
fof(f28708,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK63(X0,X1,'T_0',X2)))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f2491,f2485,f7869,f1566])).
fof(f29529,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28709,f29272])).
fof(f28709,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK63(X0,X1,'T_0',X2)))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f1919,f2485,f7869,f1566])).
fof(f29528,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28710,f29272])).
fof(f28710,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK64(X0,X1,'F_0',X2)))) ) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f8040,f2588,f7869,f1566])).
fof(f2588,plain,(
  ( ! [X2,X0,X3,X1] : (~sP7('T_0',X0,sK64(X1,X2,'F_0',X3))) )),
  inference(unit_resulting_resolution,[],[f1104,f2552,f1564])).
fof(f29527,plain,(
  $false | ($spl19 | $spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28711,f29272])).
fof(f28711,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK64(X0,X1,'F_0',X2)))) ) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8639,f2588,f7869,f1566])).
fof(f29526,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28712,f29272])).
fof(f28712,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK64(X0,X1,'F_0',X2)))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f2594,f2588,f7869,f1566])).
fof(f29525,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28713,f29272])).
fof(f28713,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK64(X0,X1,'F_0',X2)))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f1919,f2588,f7869,f1566])).
fof(f29524,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28725,f1919])).
fof(f28725,plain,(
  ~sP6('T_0','const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),'T_0') | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1104,f8611,f7869,f1566])).
fof(f29523,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28726,f1919])).
fof(f28726,plain,(
  ( ! [X0] : (~sP6('const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13)),'const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),'T_0')) ) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f7869,f8611,f7869,f1566])).
fof(f29522,plain,(
  $false | ($spl19 | $spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28727,f29272])).
fof(f28727,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8638,f8611,f7869,f1566])).
fof(f29521,plain,(
  $false | ($spl19 | $spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28729,f29272])).
fof(f28729,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1919,f8611,f7869,f1566])).
fof(f29520,plain,(
  $false | ($spl19 | $spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28732,f7869])).
fof(f28732,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)))) | ($spl19 | $spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8040,f8611,f7869,f1566])).
fof(f29519,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28734,f7869])).
fof(f28734,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)))) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8639,f8611,f7869,f1566])).
fof(f29518,plain,(
  $false | ($spl7 | $spl19 | $spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28736,f29272])).
fof(f28736,plain,(
  ~p(s(bool,sK16)) | ($spl7 | $spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1922,f8611,f7869,f1566])).
fof(f1922,plain,(
  ( ! [X0,X1] : (sP6(sK16,X0,X1)) ) | $spl7),
  inference(unit_resulting_resolution,[],[f1696,f1573])).
fof(f29517,plain,(
  $false | ($spl19 | $spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28737,f29272])).
fof(f28737,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK63(X0,X1,'T_0',X2)))) ) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f2492,f8611,f7869,f1566])).
fof(f2492,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP6(sK63(X0,X1,'T_0',X2),X3,X4)) )),
  inference(unit_resulting_resolution,[],[f2471,f1573])).
fof(f29516,plain,(
  $false | ($spl19 | $spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28738,f29272])).
fof(f28738,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK64(X0,X1,'F_0',X2)))) ) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f2595,f8611,f7869,f1566])).
fof(f2595,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP6(sK64(X0,X1,'F_0',X2),X3,X4)) )),
  inference(unit_resulting_resolution,[],[f2552,f1573])).
fof(f29515,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28739,f1919])).
fof(f28739,plain,(
  ~sP6('T_0','const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)),'T_0') | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1104,f8012,f7869,f1566])).
fof(f29514,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28740,f1919])).
fof(f28740,plain,(
  ( ! [X0] : (~sP6('const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13)),'const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)),'T_0')) ) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f7869,f8012,f7869,f1566])).
fof(f29513,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28741,f29272])).
fof(f28741,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f8039,f8012,f7869,f1566])).
fof(f29512,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28743,f29272])).
fof(f28743,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1919,f8012,f7869,f1566])).
fof(f29511,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28746,f7869])).
fof(f28746,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)))) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f8040,f8012,f7869,f1566])).
fof(f29510,plain,(
  $false | ($spl19 | $spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28748,f7869])).
fof(f28748,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)))) | ($spl19 | $spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8639,f8012,f7869,f1566])).
fof(f29509,plain,(
  $false | ($spl7 | $spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28750,f29272])).
fof(f28750,plain,(
  ~p(s(bool,sK16)) | ($spl7 | $spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1922,f8012,f7869,f1566])).
fof(f29508,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28751,f29272])).
fof(f28751,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK63(X0,X1,'T_0',X2)))) ) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f2492,f8012,f7869,f1566])).
fof(f29507,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28752,f29272])).
fof(f28752,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK64(X0,X1,'F_0',X2)))) ) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f2595,f8012,f7869,f1566])).
fof(f29506,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28754,f1919])).
fof(f28754,plain,(
  ( ! [X0] : (~sP6('const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)),'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13)),'T_0')) ) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f7869,f8021,f7869,f1566])).
fof(f29505,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28755,f7869])).
fof(f28755,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)))) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f8040,f8021,f7869,f1566])).
fof(f29504,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28757,f1919])).
fof(f28757,plain,(
  ( ! [X0] : (~sP6('const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13)),'T_0')) ) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f7869,f8620,f7869,f1566])).
fof(f29503,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28758,f7869])).
fof(f28758,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)))) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8639,f8620,f7869,f1566])).
fof(f29502,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28760,f29272])).
fof(f28760,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK63(X0,X1,'T_0',X2)))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f2492,f2488,f7869,f1566])).
fof(f2488,plain,(
  ( ! [X2,X0,X3,X1] : (~sP7('T_0',sK63(X0,X1,'T_0',X2),X3)) )),
  inference(unit_resulting_resolution,[],[f1104,f2471,f1565])).
fof(f29501,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28761,f29272])).
fof(f28761,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK63(X0,X1,'T_0',X2)))) ) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f8039,f2488,f7869,f1566])).
fof(f29500,plain,(
  $false | ($spl19 | $spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28762,f29272])).
fof(f28762,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK63(X0,X1,'T_0',X2)))) ) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8638,f2488,f7869,f1566])).
fof(f29499,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28763,f29272])).
fof(f28763,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK63(X0,X1,'T_0',X2)))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f1919,f2488,f7869,f1566])).
fof(f29498,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28764,f29272])).
fof(f28764,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK64(X0,X1,'F_0',X2)))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f2595,f2591,f7869,f1566])).
fof(f2591,plain,(
  ( ! [X2,X0,X3,X1] : (~sP7('T_0',sK64(X0,X1,'F_0',X2),X3)) )),
  inference(unit_resulting_resolution,[],[f1104,f2552,f1565])).
fof(f29497,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28765,f29272])).
fof(f28765,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK64(X0,X1,'F_0',X2)))) ) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f8039,f2591,f7869,f1566])).
fof(f29496,plain,(
  $false | ($spl19 | $spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28766,f29272])).
fof(f28766,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK64(X0,X1,'F_0',X2)))) ) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8638,f2591,f7869,f1566])).
fof(f29495,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28767,f29272])).
fof(f28767,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK64(X0,X1,'F_0',X2)))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f1919,f2591,f7869,f1566])).
fof(f29494,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28773,f29272])).
fof(f28773,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f8040,f7869,f1570])).
fof(f1570,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f1081])).
fof(f29493,plain,(
  $false | ($spl19 | $spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28777,f29272])).
fof(f28777,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f8639,f7869,f1570])).
fof(f29492,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28783,f7869])).
fof(f28783,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)))) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f8039,f7869,f1570])).
fof(f29491,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28786,f7869])).
fof(f28786,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)))) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f8638,f7869,f1570])).
fof(f29490,plain,(
  $false | ($spl7 | $spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28789,f29272])).
fof(f28789,plain,(
  ~p(s(bool,sK16)) | ($spl7 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f1920,f7869,f1570])).
fof(f29489,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28790,f29272])).
fof(f28790,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK63(X0,X1,'T_0',X2)))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f1700,f2491,f7869,f1570])).
fof(f29488,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28791,f29272])).
fof(f28791,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK64(X0,X1,'F_0',X2)))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f1700,f2594,f7869,f1570])).
fof(f29487,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28796,f29272])).
fof(f28796,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f8039,f7869,f1570])).
fof(f29486,plain,(
  $false | ($spl19 | $spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28800,f29272])).
fof(f28800,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f8638,f7869,f1570])).
fof(f29485,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28806,f7869])).
fof(f28806,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)))) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f8040,f7869,f1570])).
fof(f29484,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28809,f7869])).
fof(f28809,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)))) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f8639,f7869,f1570])).
fof(f29483,plain,(
  $false | ($spl7 | $spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28812,f29272])).
fof(f28812,plain,(
  ~p(s(bool,sK16)) | ($spl7 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f1922,f7869,f1570])).
fof(f29482,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28813,f29272])).
fof(f28813,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK63(X0,X1,'T_0',X2)))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f1700,f2492,f7869,f1570])).
fof(f29481,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28814,f29272])).
fof(f28814,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK64(X0,X1,'F_0',X2)))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f1700,f2595,f7869,f1570])).
fof(f29480,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28821,f1104])).
fof(f28821,plain,(
  ~p(s(bool,'T_0')) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1104,f8021,f7869,f1576])).
fof(f1576,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f1083])).
fof(f1083,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f1082])).
fof(f1082,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f832])).
fof(f832,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP7(X2,X0,X1))),
  inference(definition_folding,[],[f633,f831,f830])).
fof(f633,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f632])).
fof(f632,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f52])).
fof(f52,axiom,(
  ! [X33,X34,X23] : ((p(s(bool,X23)) <=> (p(s(bool,X34)) & p(s(bool,X33)))) <=> ((p(s(bool,X23)) | ~p(s(bool,X34)) | ~p(s(bool,X33))) & (p(s(bool,X34)) | ~p(s(bool,X23))) & (p(s(bool,X33)) | ~p(s(bool,X23)))))),
  file('thm.pred_set.EXTENSION','thm.sat.sat_23_prove_ltb')).
fof(f29479,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28822,f29272])).
fof(f28822,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1104,f8021,f7869,f1576])).
fof(f29478,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28824,f1104])).
fof(f28824,plain,(
  ~p(s(bool,'T_0')) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1104,f8620,f7869,f1576])).
fof(f29477,plain,(
  $false | ($spl19 | $spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28825,f29272])).
fof(f28825,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1104,f8620,f7869,f1576])).
fof(f29476,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28828,f7869])).
fof(f28828,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)))) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1104,f8611,f7869,f1576])).
fof(f29475,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28829,f1104])).
fof(f28829,plain,(
  ~p(s(bool,'T_0')) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f7869,f8611,f7869,f1576])).
fof(f29474,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28831,f7869])).
fof(f28831,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)))) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1104,f8012,f7869,f1576])).
fof(f29473,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28832,f1104])).
fof(f28832,plain,(
  ~p(s(bool,'T_0')) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f7869,f8012,f7869,f1576])).
fof(f29472,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28834,f29272])).
fof(f28834,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK63(X0,X1,'T_0',X2)))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f1104,f2485,f7869,f1576])).
fof(f29471,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28835,f29272])).
fof(f28835,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK64(X0,X1,'F_0',X2)))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f1104,f2588,f7869,f1576])).
fof(f29470,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28840,f1104])).
fof(f28840,plain,(
  ~p(s(bool,'T_0')) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1104,f8611,f7869,f1576])).
fof(f29469,plain,(
  $false | ($spl19 | $spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28841,f29272])).
fof(f28841,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1104,f8611,f7869,f1576])).
fof(f29468,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28843,f1104])).
fof(f28843,plain,(
  ~p(s(bool,'T_0')) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1104,f8012,f7869,f1576])).
fof(f29467,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28844,f29272])).
fof(f28844,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1104,f8012,f7869,f1576])).
fof(f29466,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28847,f7869])).
fof(f28847,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)))) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1104,f8021,f7869,f1576])).
fof(f29465,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28848,f1104])).
fof(f28848,plain,(
  ~p(s(bool,'T_0')) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f7869,f8021,f7869,f1576])).
fof(f29464,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28850,f7869])).
fof(f28850,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)))) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1104,f8620,f7869,f1576])).
fof(f29463,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28851,f1104])).
fof(f28851,plain,(
  ~p(s(bool,'T_0')) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f7869,f8620,f7869,f1576])).
fof(f29462,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28853,f29272])).
fof(f28853,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK63(X0,X1,'T_0',X2)))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f1104,f2488,f7869,f1576])).
fof(f29461,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28854,f29272])).
fof(f28854,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK64(X0,X1,'F_0',X2)))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f1104,f2591,f7869,f1576])).
fof(f29460,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28859,f1104])).
fof(f28859,plain,(
  ~p(s(bool,'T_0')) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1104,f8656,f7869,f1576])).
fof(f8656,plain,(
  ~sP7('const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),'T_0','T_0') | $spl23),
  inference(unit_resulting_resolution,[],[f1104,f1104,f8357,f1579])).
fof(f1579,plain,(
  ( ! [X2,X0,X1] : (~sP7(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f1083])).
fof(f29459,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28860,f1104])).
fof(f28860,plain,(
  ~p(s(bool,'T_0')) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1104,f8656,f7869,f1576])).
fof(f29458,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28862,f1104])).
fof(f28862,plain,(
  ~p(s(bool,'T_0')) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1104,f8057,f7869,f1576])).
fof(f8057,plain,(
  ~sP7('const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)),'T_0','T_0') | $spl19),
  inference(unit_resulting_resolution,[],[f1104,f1104,f7868,f1579])).
fof(f29457,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28863,f1104])).
fof(f28863,plain,(
  ~p(s(bool,'T_0')) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1104,f8057,f7869,f1576])).
fof(f29456,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28868,f29272])).
fof(f28868,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f8049,f7869,f1579])).
fof(f8049,plain,(
  ( ! [X0] : (sP7('F_0','const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)),X0)) ) | $spl19),
  inference(unit_resulting_resolution,[],[f1700,f7868,f1575])).
fof(f1575,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f1083])).
fof(f29455,plain,(
  $false | ($spl19 | $spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28871,f29272])).
fof(f28871,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f8648,f7869,f1579])).
fof(f8648,plain,(
  ( ! [X0] : (sP7('F_0','const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),X0)) ) | $spl23),
  inference(unit_resulting_resolution,[],[f1700,f8357,f1575])).
fof(f29454,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28874,f7869])).
fof(f28874,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)))) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f8640,f7869,f1579])).
fof(f8640,plain,(
  ( ! [X0] : (sP7('F_0',X0,'const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)))) ) | $spl23),
  inference(unit_resulting_resolution,[],[f1700,f8357,f1574])).
fof(f1574,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f1083])).
fof(f29453,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28877,f7869])).
fof(f28877,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)))) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f8041,f7869,f1579])).
fof(f8041,plain,(
  ( ! [X0] : (sP7('F_0',X0,'const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)))) ) | $spl19),
  inference(unit_resulting_resolution,[],[f1700,f7868,f1574])).
fof(f29452,plain,(
  $false | ($spl7 | $spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28880,f29272])).
fof(f28880,plain,(
  ~p(s(bool,sK16)) | ($spl7 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f2098,f7869,f1579])).
fof(f2098,plain,(
  ( ! [X0] : (sP7('F_0',X0,sK16)) ) | $spl7),
  inference(unit_resulting_resolution,[],[f1920,f1700,f1678])).
fof(f1678,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0)) | ~sP6(X1,X2,X0)) )),
  inference(duplicate_literal_removal,[],[f1569])).
fof(f1569,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X0)) | ~sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f1079])).
fof(f29451,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28881,f29272])).
fof(f28881,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK63(X0,X1,'T_0',X2)))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f1700,f2493,f7869,f1579])).
fof(f2493,plain,(
  ( ! [X2,X0,X3,X1] : (sP7('F_0',X0,sK63(X1,X2,'T_0',X3))) )),
  inference(unit_resulting_resolution,[],[f1700,f2471,f1574])).
fof(f29450,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28882,f29272])).
fof(f28882,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK64(X0,X1,'F_0',X2)))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f1700,f2596,f7869,f1579])).
fof(f2596,plain,(
  ( ! [X2,X0,X3,X1] : (sP7('F_0',X0,sK64(X1,X2,'F_0',X3))) )),
  inference(unit_resulting_resolution,[],[f1700,f2552,f1574])).
fof(f29449,plain,(
  $false | ($spl19 | $spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28886,f29272])).
fof(f28886,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f8640,f7869,f1579])).
fof(f29448,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28889,f29272])).
fof(f28889,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f8041,f7869,f1579])).
fof(f29447,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28892,f7869])).
fof(f28892,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)))) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f8049,f7869,f1579])).
fof(f29446,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28895,f7869])).
fof(f28895,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)))) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f8648,f7869,f1579])).
fof(f29445,plain,(
  $false | ($spl7 | $spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28898,f29272])).
fof(f28898,plain,(
  ~p(s(bool,sK16)) | ($spl7 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f2100,f7869,f1579])).
fof(f2100,plain,(
  ( ! [X0] : (sP7('F_0',sK16,X0)) ) | $spl7),
  inference(unit_resulting_resolution,[],[f1922,f1700,f1678])).
fof(f29444,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28899,f29272])).
fof(f28899,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK63(X0,X1,'T_0',X2)))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f1700,f2498,f7869,f1579])).
fof(f2498,plain,(
  ( ! [X2,X0,X3,X1] : (sP7('F_0',sK63(X0,X1,'T_0',X2),X3)) )),
  inference(unit_resulting_resolution,[],[f1700,f2471,f1575])).
fof(f29443,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28900,f29272])).
fof(f28900,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK64(X0,X1,'F_0',X2)))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f1700,f2601,f7869,f1579])).
fof(f2601,plain,(
  ( ! [X2,X0,X3,X1] : (sP7('F_0',sK64(X0,X1,'F_0',X2),X3)) )),
  inference(unit_resulting_resolution,[],[f1700,f2552,f1575])).
fof(f29442,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28917,f1925])).
fof(f1925,plain,(
  ( ! [X0,X1] : (sP8(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f1104,f1589])).
fof(f1589,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f1089])).
fof(f1089,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f1088])).
fof(f1088,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f1087])).
fof(f1087,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f833])).
fof(f28917,plain,(
  ~sP8('const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),'T_0','F_0') | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f8658,f7869,f1584])).
fof(f1584,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f1086])).
fof(f1086,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP8(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f1085])).
fof(f1085,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1) | sP9(X2,X0,X1)))),
  inference(flattening,[],[f1084])).
fof(f1084,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1)) | sP9(X2,X0,X1)))),
  inference(nnf_transformation,[],[f834])).
fof(f8658,plain,(
  ( ! [X0] : (~sP9('const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),'T_0',X0)) ) | $spl23),
  inference(unit_resulting_resolution,[],[f1104,f8357,f1581])).
fof(f1581,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f1086])).
fof(f29441,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28918,f1925])).
fof(f28918,plain,(
  ~sP8('const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)),'T_0','F_0') | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f8059,f7869,f1584])).
fof(f8059,plain,(
  ( ! [X0] : (~sP9('const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)),'T_0',X0)) ) | $spl19),
  inference(unit_resulting_resolution,[],[f1104,f7868,f1581])).
fof(f29440,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28935,f1924])).
fof(f1924,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f1104,f1588])).
fof(f1588,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f1089])).
fof(f28935,plain,(
  ~sP8('const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),'F_0','T_0') | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f8657,f7869,f1585])).
fof(f1585,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f1086])).
fof(f8657,plain,(
  ( ! [X0] : (~sP9('const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),X0,'T_0')) ) | $spl23),
  inference(unit_resulting_resolution,[],[f1104,f8357,f1580])).
fof(f1580,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f1086])).
fof(f29439,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28936,f1924])).
fof(f28936,plain,(
  ~sP8('const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)),'F_0','T_0') | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f8058,f7869,f1585])).
fof(f8058,plain,(
  ( ! [X0] : (~sP9('const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)),X0,'T_0')) ) | $spl19),
  inference(unit_resulting_resolution,[],[f1104,f7868,f1580])).
fof(f29438,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28968,f1104])).
fof(f28968,plain,(
  ~p(s(bool,'T_0')) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8735,f7869,f1592])).
fof(f1592,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f1091])).
fof(f8735,plain,(
  ~sP9('T_0','F_0','const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13))) | $spl23),
  inference(unit_resulting_resolution,[],[f1700,f1104,f8357,f1594])).
fof(f1594,plain,(
  ( ! [X2,X0,X1] : (~sP9(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f1091])).
fof(f29437,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28970,f1104])).
fof(f28970,plain,(
  ~p(s(bool,'T_0')) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f8136,f7869,f1592])).
fof(f8136,plain,(
  ~sP9('T_0','F_0','const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13))) | $spl19),
  inference(unit_resulting_resolution,[],[f1700,f1104,f7868,f1594])).
fof(f29436,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28973,f7869])).
fof(f28973,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)))) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8658,f7869,f1592])).
fof(f29435,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28975,f7869])).
fof(f28975,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)))) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f8059,f7869,f1592])).
fof(f29434,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28977,f29272])).
fof(f28977,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK63(X0,X1,'T_0',X2)))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f2504,f7869,f1592])).
fof(f2504,plain,(
  ( ! [X2,X0,X3,X1] : (~sP9(sK63(X0,X1,'T_0',X2),'T_0',X3)) )),
  inference(unit_resulting_resolution,[],[f1104,f2471,f1581])).
fof(f29433,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28978,f29272])).
fof(f28978,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK64(X0,X1,'F_0',X2)))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f2607,f7869,f1592])).
fof(f2607,plain,(
  ( ! [X2,X0,X3,X1] : (~sP9(sK64(X0,X1,'F_0',X2),'T_0',X3)) )),
  inference(unit_resulting_resolution,[],[f1104,f2552,f1581])).
fof(f29432,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28981,f1104])).
fof(f28981,plain,(
  ~p(s(bool,'T_0')) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8657,f7869,f1592])).
fof(f29431,plain,(
  $false | ($spl19 | $spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28983,f29272])).
fof(f28983,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8658,f7869,f1592])).
fof(f29430,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28985,f1104])).
fof(f28985,plain,(
  ~p(s(bool,'T_0')) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f8058,f7869,f1592])).
fof(f29429,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28987,f29272])).
fof(f28987,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f8059,f7869,f1592])).
fof(f29428,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28991,f1104])).
fof(f28991,plain,(
  ~p(s(bool,'T_0')) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8732,f7869,f1593])).
fof(f8732,plain,(
  ~sP9('T_0','const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),'F_0') | $spl23),
  inference(unit_resulting_resolution,[],[f1700,f1104,f8357,f1594])).
fof(f29427,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28993,f1104])).
fof(f28993,plain,(
  ~p(s(bool,'T_0')) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f8133,f7869,f1593])).
fof(f8133,plain,(
  ~sP9('T_0','const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)),'F_0') | $spl19),
  inference(unit_resulting_resolution,[],[f1700,f1104,f7868,f1594])).
fof(f29426,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f28997,f7869])).
fof(f28997,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)))) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8657,f7869,f1593])).
fof(f29425,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f28999,f7869])).
fof(f28999,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)))) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f8058,f7869,f1593])).
fof(f29424,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f29001,f29272])).
fof(f29001,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK63(X0,X1,'T_0',X2)))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f2503,f7869,f1593])).
fof(f2503,plain,(
  ( ! [X2,X0,X3,X1] : (~sP9(sK63(X0,X1,'T_0',X2),X3,'T_0')) )),
  inference(unit_resulting_resolution,[],[f1104,f2471,f1580])).
fof(f29423,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f29002,f29272])).
fof(f29002,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK64(X0,X1,'F_0',X2)))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f2606,f7869,f1593])).
fof(f2606,plain,(
  ( ! [X2,X0,X3,X1] : (~sP9(sK64(X0,X1,'F_0',X2),X3,'T_0')) )),
  inference(unit_resulting_resolution,[],[f1104,f2552,f1580])).
fof(f29422,plain,(
  $false | ($spl19 | $spl20 | $spl23)),
  inference(subsumption_resolution,[],[f29005,f29272])).
fof(f29005,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8657,f7869,f1593])).
fof(f29421,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f29007,f1104])).
fof(f29007,plain,(
  ~p(s(bool,'T_0')) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8658,f7869,f1593])).
fof(f29420,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f29009,f29272])).
fof(f29009,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f8058,f7869,f1593])).
fof(f29419,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f29011,f1104])).
fof(f29011,plain,(
  ~p(s(bool,'T_0')) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f8059,f7869,f1593])).
fof(f29418,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f29033,f1930])).
fof(f1930,plain,(
  ( ! [X0,X1] : (sP10(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f1104,f1606])).
fof(f1606,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f1097])).
fof(f1097,plain,(
  ! [X0,X1,X2] : ((~sP10(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP10(X0,X1,X2)))),
  inference(rectify,[],[f1096])).
fof(f1096,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(flattening,[],[f1095])).
fof(f1095,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(nnf_transformation,[],[f836])).
fof(f29033,plain,(
  ~sP10('const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),'T_0','T_0') | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1104,f8743,f7869,f1601])).
fof(f1601,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f1094])).
fof(f1094,plain,(
  ! [X0,X1,X2] : ((~sP11(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP10(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP10(X0,X1,X2) | sP11(X0,X1,X2)))),
  inference(rectify,[],[f1093])).
fof(f1093,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1) | sP11(X2,X0,X1)))),
  inference(flattening,[],[f1092])).
fof(f1092,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1)) | sP11(X2,X0,X1)))),
  inference(nnf_transformation,[],[f837])).
fof(f8743,plain,(
  ( ! [X0] : (~sP11('const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),'T_0',X0)) ) | $spl23),
  inference(unit_resulting_resolution,[],[f1104,f8357,f1598])).
fof(f1598,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f1094])).
fof(f29417,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f29034,f1930])).
fof(f29034,plain,(
  ( ! [X0] : (~sP10('const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),'T_0','const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13)))) ) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f7869,f8743,f7869,f1601])).
fof(f29416,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f29035,f1930])).
fof(f29035,plain,(
  ~sP10('const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)),'T_0','T_0') | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1104,f8144,f7869,f1601])).
fof(f8144,plain,(
  ( ! [X0] : (~sP11('const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)),'T_0',X0)) ) | $spl19),
  inference(unit_resulting_resolution,[],[f1104,f7868,f1598])).
fof(f29415,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f29036,f1930])).
fof(f29036,plain,(
  ( ! [X0] : (~sP10('const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)),'T_0','const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13)))) ) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f7869,f8144,f7869,f1601])).
fof(f29414,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f29040,f29272])).
fof(f29040,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f8144,f8179,f7869,f1601])).
fof(f8179,plain,(
  ( ! [X0,X1] : (sP10('const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)),X0,X1)) ) | $spl19),
  inference(unit_resulting_resolution,[],[f7868,f1607])).
fof(f1607,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f1097])).
fof(f29413,plain,(
  $false | ($spl19 | $spl20 | $spl23)),
  inference(subsumption_resolution,[],[f29045,f29272])).
fof(f29045,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8743,f8778,f7869,f1601])).
fof(f8778,plain,(
  ( ! [X0,X1] : (sP10('const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),X0,X1)) ) | $spl23),
  inference(unit_resulting_resolution,[],[f8357,f1607])).
fof(f29412,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f29049,f7869])).
fof(f29049,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)))) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8743,f8777,f7869,f1601])).
fof(f8777,plain,(
  ( ! [X0,X1] : (sP10(X0,X1,'const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)))) ) | $spl23),
  inference(unit_resulting_resolution,[],[f8357,f1605])).
fof(f1605,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f1097])).
fof(f29411,plain,(
  $false | ($spl19 | $spl20 | $spl23)),
  inference(subsumption_resolution,[],[f29052,f7869])).
fof(f29052,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)))) | ($spl19 | $spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8743,f8178,f7869,f1601])).
fof(f8178,plain,(
  ( ! [X0,X1] : (sP10(X0,X1,'const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)))) ) | $spl19),
  inference(unit_resulting_resolution,[],[f7868,f1605])).
fof(f29410,plain,(
  $false | ($spl7 | $spl19 | $spl20 | $spl23)),
  inference(subsumption_resolution,[],[f29054,f29272])).
fof(f29054,plain,(
  ~p(s(bool,sK16)) | ($spl7 | $spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8743,f1928,f7869,f1601])).
fof(f1928,plain,(
  ( ! [X0,X1] : (sP10(X0,X1,sK16)) ) | $spl7),
  inference(unit_resulting_resolution,[],[f1696,f1605])).
fof(f29409,plain,(
  $false | ($spl19 | $spl20 | $spl23)),
  inference(subsumption_resolution,[],[f29055,f29272])).
fof(f29055,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK63(X0,X1,'T_0',X2)))) ) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8743,f2512,f7869,f1601])).
fof(f2512,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10(X0,X1,sK63(X2,X3,'T_0',X4))) )),
  inference(unit_resulting_resolution,[],[f2471,f1605])).
fof(f29408,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f29056,f29272])).
fof(f29056,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK63(X0,X1,'T_0',X2)))) ) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f8144,f2512,f7869,f1601])).
fof(f29407,plain,(
  $false | ($spl19 | $spl20 | $spl23)),
  inference(subsumption_resolution,[],[f29057,f29272])).
fof(f29057,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK64(X0,X1,'F_0',X2)))) ) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8743,f2615,f7869,f1601])).
fof(f2615,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10(X0,X1,sK64(X2,X3,'F_0',X4))) )),
  inference(unit_resulting_resolution,[],[f2552,f1605])).
fof(f29406,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f29058,f29272])).
fof(f29058,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK64(X0,X1,'F_0',X2)))) ) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f8144,f2615,f7869,f1601])).
fof(f29405,plain,(
  $false | ($spl19 | $spl20 | $spl23)),
  inference(subsumption_resolution,[],[f29061,f29272])).
fof(f29061,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8743,f1930,f7869,f1601])).
fof(f29404,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f29064,f29272])).
fof(f29064,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f8144,f1930,f7869,f1601])).
fof(f29403,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f29067,f1930])).
fof(f29067,plain,(
  ( ! [X0] : (~sP10('const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),'T_0','const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13)))) ) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f7869,f8743,f7869,f1601])).
fof(f29402,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f29068,f1930])).
fof(f29068,plain,(
  ( ! [X0] : (~sP10('const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)),'T_0','const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13)))) ) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f7869,f8144,f7869,f1601])).
fof(f29401,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f29073,f7869])).
fof(f29073,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)))) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8743,f8777,f7869,f1601])).
fof(f29400,plain,(
  $false | ($spl19 | $spl20 | $spl23)),
  inference(subsumption_resolution,[],[f29075,f7869])).
fof(f29075,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)))) | ($spl19 | $spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8144,f8777,f7869,f1601])).
fof(f29399,plain,(
  $false | ($spl19 | $spl20 | $spl23)),
  inference(subsumption_resolution,[],[f29077,f29272])).
fof(f29077,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK63(X0,X1,'T_0',X2)))) ) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f2511,f8777,f7869,f1601])).
fof(f2511,plain,(
  ( ! [X2,X0,X3,X1] : (~sP11(sK63(X0,X1,'T_0',X2),'T_0',X3)) )),
  inference(unit_resulting_resolution,[],[f1104,f2471,f1598])).
fof(f29398,plain,(
  $false | ($spl19 | $spl20 | $spl23)),
  inference(subsumption_resolution,[],[f29078,f29272])).
fof(f29078,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK64(X0,X1,'F_0',X2)))) ) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f2614,f8777,f7869,f1601])).
fof(f2614,plain,(
  ( ! [X2,X0,X3,X1] : (~sP11(sK64(X0,X1,'F_0',X2),'T_0',X3)) )),
  inference(unit_resulting_resolution,[],[f1104,f2552,f1598])).
fof(f29397,plain,(
  $false | ($spl19 | $spl20 | $spl23)),
  inference(subsumption_resolution,[],[f29083,f7869])).
fof(f29083,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)))) | ($spl19 | $spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8743,f8178,f7869,f1601])).
fof(f29396,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f29085,f7869])).
fof(f29085,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)))) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f8144,f8178,f7869,f1601])).
fof(f29395,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f29087,f29272])).
fof(f29087,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK63(X0,X1,'T_0',X2)))) ) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f2511,f8178,f7869,f1601])).
fof(f29394,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f29088,f29272])).
fof(f29088,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK64(X0,X1,'F_0',X2)))) ) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f2614,f8178,f7869,f1601])).
fof(f29393,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f29094,f7869])).
fof(f29094,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)))) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8743,f1930,f7869,f1601])).
fof(f29392,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f29096,f7869])).
fof(f29096,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)))) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f8144,f1930,f7869,f1601])).
fof(f29391,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f29098,f29272])).
fof(f29098,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK63(X0,X1,'T_0',X2)))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f2511,f1930,f7869,f1601])).
fof(f29390,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f29099,f29272])).
fof(f29099,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK64(X0,X1,'F_0',X2)))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f2614,f1930,f7869,f1601])).
fof(f29389,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f29102,f7869])).
fof(f29102,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)))) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f8144,f8179,f7869,f1601])).
fof(f29388,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f29105,f7869])).
fof(f29105,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)))) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8743,f8778,f7869,f1601])).
fof(f29387,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f29107,f29272])).
fof(f29107,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK63(X0,X1,'T_0',X2)))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f2511,f2513,f7869,f1601])).
fof(f2513,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10(sK63(X0,X1,'T_0',X2),X3,X4)) )),
  inference(unit_resulting_resolution,[],[f2471,f1607])).
fof(f29386,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f29108,f29272])).
fof(f29108,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK64(X0,X1,'F_0',X2)))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f2614,f2616,f7869,f1601])).
fof(f2616,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10(sK64(X0,X1,'F_0',X2),X3,X4)) )),
  inference(unit_resulting_resolution,[],[f2552,f1607])).
fof(f29385,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f29109,f1929])).
fof(f1929,plain,(
  ( ! [X0,X1] : (sP10(X0,X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f1700,f1605])).
fof(f29109,plain,(
  ~sP10('const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),'F_0','F_0') | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f8738,f7869,f1602])).
fof(f1602,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f1094])).
fof(f8738,plain,(
  ( ! [X0] : (~sP11('const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),X0,'F_0')) ) | $spl23),
  inference(unit_resulting_resolution,[],[f1700,f8357,f1597])).
fof(f1597,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f1094])).
fof(f29384,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f29110,f1929])).
fof(f29110,plain,(
  ~sP10('const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)),'F_0','F_0') | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f8139,f7869,f1602])).
fof(f8139,plain,(
  ( ! [X0] : (~sP11('const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)),X0,'F_0')) ) | $spl19),
  inference(unit_resulting_resolution,[],[f1700,f7868,f1597])).
fof(f29383,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f29142,f8777])).
fof(f29142,plain,(
  ~sP10('F_0','F_0','const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13))) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f8740,f7869,f1603])).
fof(f1603,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f1094])).
fof(f8740,plain,(
  ( ! [X0] : (~sP11('F_0',X0,'const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)))) ) | $spl23),
  inference(unit_resulting_resolution,[],[f1700,f8357,f1597])).
fof(f29382,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f29143,f8178])).
fof(f29143,plain,(
  ~sP10('F_0','F_0','const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13))) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f8141,f7869,f1603])).
fof(f8141,plain,(
  ( ! [X0] : (~sP11('F_0',X0,'const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)))) ) | $spl19),
  inference(unit_resulting_resolution,[],[f1700,f7868,f1597])).
fof(f29381,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f29178,f29272])).
fof(f29178,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f8179,f7869,f1604])).
fof(f1604,plain,(
  ( ! [X2,X0,X1] : (~sP10(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f1097])).
fof(f29380,plain,(
  $false | ($spl19 | $spl20 | $spl23)),
  inference(subsumption_resolution,[],[f29183,f29272])).
fof(f29183,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f8778,f7869,f1604])).
fof(f29379,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f29189,f7869])).
fof(f29189,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)))) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f8777,f7869,f1604])).
fof(f29378,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f29192,f7869])).
fof(f29192,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)))) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f8178,f7869,f1604])).
fof(f29377,plain,(
  $false | ($spl7 | $spl19 | $spl20)),
  inference(subsumption_resolution,[],[f29195,f29272])).
fof(f29195,plain,(
  ~p(s(bool,sK16)) | ($spl7 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f1928,f7869,f1604])).
fof(f29376,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f29196,f29272])).
fof(f29196,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK63(X0,X1,'T_0',X2)))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f1700,f2512,f7869,f1604])).
fof(f29375,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f29197,f29272])).
fof(f29197,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK64(X0,X1,'F_0',X2)))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f1700,f2615,f7869,f1604])).
fof(f29374,plain,(
  $false | ($spl19 | $spl20 | $spl23)),
  inference(subsumption_resolution,[],[f29203,f29272])).
fof(f29203,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f8777,f7869,f1604])).
fof(f29373,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f29207,f29272])).
fof(f29207,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f8178,f7869,f1604])).
fof(f29372,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f29211,f7869])).
fof(f29211,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)))) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f8179,f7869,f1604])).
fof(f29371,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f29214,f7869])).
fof(f29214,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)))) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f8778,f7869,f1604])).
fof(f29370,plain,(
  $false | ($spl7 | $spl19 | $spl20)),
  inference(subsumption_resolution,[],[f29217,f29272])).
fof(f29217,plain,(
  ~p(s(bool,sK16)) | ($spl7 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f1931,f7869,f1604])).
fof(f1931,plain,(
  ( ! [X0,X1] : (sP10(sK16,X0,X1)) ) | $spl7),
  inference(unit_resulting_resolution,[],[f1696,f1607])).
fof(f29369,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f29218,f29272])).
fof(f29218,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK63(X0,X1,'T_0',X2)))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f1700,f2513,f7869,f1604])).
fof(f29368,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f29219,f29272])).
fof(f29219,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK64(X0,X1,'F_0',X2)))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f1700,f2616,f7869,f1604])).
fof(f29367,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f29241,f1104])).
fof(f29241,plain,(
  ~p(s(bool,'T_0')) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8790,f7869,f1610])).
fof(f8790,plain,(
  ~sP11('T_0','const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),'T_0') | $spl23),
  inference(unit_resulting_resolution,[],[f1104,f1104,f8357,f1611])).
fof(f29366,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f29243,f1104])).
fof(f29243,plain,(
  ~p(s(bool,'T_0')) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f8191,f7869,f1610])).
fof(f8191,plain,(
  ~sP11('T_0','const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)),'T_0') | $spl19),
  inference(unit_resulting_resolution,[],[f1104,f1104,f7868,f1611])).
fof(f29365,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f29250,f7869])).
fof(f29250,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)))) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8738,f7869,f1610])).
fof(f29364,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f29252,f7869])).
fof(f29252,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)))) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f8139,f7869,f1610])).
fof(f29363,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f29254,f29272])).
fof(f29254,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK63(X0,X1,'T_0',X2)))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f2506,f7869,f1610])).
fof(f2506,plain,(
  ( ! [X2,X0,X3,X1] : (~sP11(sK63(X0,X1,'T_0',X2),X3,'F_0')) )),
  inference(unit_resulting_resolution,[],[f1700,f2471,f1597])).
fof(f29362,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f29255,f29272])).
fof(f29255,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK64(X0,X1,'F_0',X2)))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f2609,f7869,f1610])).
fof(f2609,plain,(
  ( ! [X2,X0,X3,X1] : (~sP11(sK64(X0,X1,'F_0',X2),X3,'F_0')) )),
  inference(unit_resulting_resolution,[],[f1700,f2552,f1597])).
fof(f29361,plain,(
  $false | ($spl19 | $spl20 | $spl23)),
  inference(subsumption_resolution,[],[f29258,f29272])).
fof(f29258,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8738,f7869,f1610])).
fof(f29360,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f29260,f1104])).
fof(f29260,plain,(
  ~p(s(bool,'T_0')) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8743,f7869,f1610])).
fof(f29359,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f29262,f29272])).
fof(f29262,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f8139,f7869,f1610])).
fof(f29358,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f29264,f1104])).
fof(f29264,plain,(
  ~p(s(bool,'T_0')) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f8144,f7869,f1610])).
fof(f29357,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f29268,f1104])).
fof(f29268,plain,(
  ~p(s(bool,'T_0')) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f8786,f7869,f1611])).
fof(f8786,plain,(
  ( ! [X0] : (sP11('T_0',X0,'const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)))) ) | $spl23),
  inference(unit_resulting_resolution,[],[f1104,f8357,f1609])).
fof(f29356,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f29271,f1104])).
fof(f29271,plain,(
  ~p(s(bool,'T_0')) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f8187,f7869,f1611])).
fof(f29355,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f29277,f1104])).
fof(f29277,plain,(
  ~p(s(bool,'T_0')) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f8784,f7869,f1611])).
fof(f8784,plain,(
  sP11('const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),'F_0','T_0') | $spl23),
  inference(unit_resulting_resolution,[],[f1700,f1104,f8357,f1608])).
fof(f1608,plain,(
  ( ! [X2,X0,X1] : (sP11(X2,X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f1099])).
fof(f29354,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f29278,f1700])).
fof(f29278,plain,(
  p(s(bool,'F_0')) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1104,f8784,f7869,f1611])).
fof(f29353,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f29280,f1104])).
fof(f29280,plain,(
  ~p(s(bool,'T_0')) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f8185,f7869,f1611])).
fof(f8185,plain,(
  sP11('const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)),'F_0','T_0') | $spl19),
  inference(unit_resulting_resolution,[],[f1700,f1104,f7868,f1608])).
fof(f29352,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f29281,f1700])).
fof(f29281,plain,(
  p(s(bool,'F_0')) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1104,f8185,f7869,f1611])).
fof(f29351,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f29283,f1929])).
fof(f29283,plain,(
  ( ! [X0] : (~sP10('const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),X0,'F_0')) ) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8738,f7869,f1676])).
fof(f1676,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP10(X0,X1,X2) | sP11(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f1600])).
fof(f1600,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f1094])).
fof(f29350,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f29284,f1930])).
fof(f29284,plain,(
  ( ! [X0] : (~sP10('const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),'T_0',X0)) ) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8743,f7869,f1676])).
fof(f29349,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f29285,f1929])).
fof(f29285,plain,(
  ( ! [X0] : (~sP10('const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)),X0,'F_0')) ) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f8139,f7869,f1676])).
fof(f29348,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f29286,f1930])).
fof(f29286,plain,(
  ( ! [X0] : (~sP10('const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)),'T_0',X0)) ) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f8144,f7869,f1676])).
fof(f29347,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f29305,f1924])).
fof(f29305,plain,(
  ( ! [X0] : (~sP8('const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),X0,'T_0')) ) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8657,f7869,f1677])).
fof(f1677,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f1583])).
fof(f1583,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f1086])).
fof(f29346,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f29306,f1925])).
fof(f29306,plain,(
  ( ! [X0] : (~sP8('const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),'T_0',X0)) ) | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8658,f7869,f1677])).
fof(f29345,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f29307,f1924])).
fof(f29307,plain,(
  ( ! [X0] : (~sP8('const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)),X0,'T_0')) ) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f8058,f7869,f1677])).
fof(f29344,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f29308,f1925])).
fof(f29308,plain,(
  ( ! [X0] : (~sP8('const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)),'T_0',X0)) ) | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f8059,f7869,f1677])).
fof(f29316,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8658,f1925,f7869,f1677])).
fof(f29314,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8657,f1924,f7869,f1677])).
fof(f29312,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8657,f8716,f7869,f1677])).
fof(f8716,plain,(
  ( ! [X0,X1] : (sP8('const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),X0,X1)) ) | $spl23),
  inference(unit_resulting_resolution,[],[f8357,f1590])).
fof(f1590,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f1089])).
fof(f29310,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f8058,f8117,f7869,f1677])).
fof(f8117,plain,(
  ( ! [X0,X1] : (sP8('const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)),X0,X1)) ) | $spl19),
  inference(unit_resulting_resolution,[],[f7868,f1590])).
fof(f29304,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8743,f1930,f7869,f1676])).
fof(f29302,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8743,f2615,f7869,f1676])).
fof(f29300,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8743,f2512,f7869,f1676])).
fof(f29298,plain,(
  $false | ($spl7 | $spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8743,f1928,f7869,f1676])).
fof(f29296,plain,(
  $false | ($spl19 | $spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8743,f8178,f7869,f1676])).
fof(f29294,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8743,f8777,f7869,f1676])).
fof(f29292,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8738,f1929,f7869,f1676])).
fof(f29290,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8738,f8778,f7869,f1676])).
fof(f29288,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f8139,f8179,f7869,f1676])).
fof(f29282,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f1104,f8185,f7869,f1611])).
fof(f29279,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f1104,f8784,f7869,f1611])).
fof(f29273,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f1104,f8187,f7869,f1611])).
fof(f29270,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f1104,f8786,f7869,f1611])).
fof(f29265,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1104,f8144,f7869,f1610])).
fof(f29263,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1104,f8139,f7869,f1610])).
fof(f29261,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1104,f8743,f7869,f1610])).
fof(f29259,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1104,f8738,f7869,f1610])).
fof(f29253,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f7869,f8139,f7869,f1610])).
fof(f29251,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f7869,f8738,f7869,f1610])).
fof(f29244,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1104,f8191,f7869,f1610])).
fof(f29242,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1104,f8790,f7869,f1610])).
fof(f29238,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f8144,f7869,f1609])).
fof(f29236,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f8139,f7869,f1609])).
fof(f29234,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f8743,f7869,f1609])).
fof(f29232,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f8738,f7869,f1609])).
fof(f29228,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f1700,f8141,f7869,f1608])).
fof(f29225,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f1700,f8740,f7869,f1608])).
fof(f29221,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8768,f7869,f1606])).
fof(f8768,plain,(
  ~sP10('T_0','const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),'T_0') | $spl23),
  inference(unit_resulting_resolution,[],[f1104,f1104,f8357,f1604])).
fof(f29216,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f7869,f8778,f7869,f1604])).
fof(f29213,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f7869,f8179,f7869,f1604])).
fof(f29208,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1104,f1700,f8178,f7869,f1604])).
fof(f29204,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1104,f1700,f8777,f7869,f1604])).
fof(f29194,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f7869,f8178,f7869,f1604])).
fof(f29191,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f7869,f8777,f7869,f1604])).
fof(f29187,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f7869,f8778,f7869,f1604])).
fof(f29185,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f1104,f8778,f7869,f1604])).
fof(f29182,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f7869,f8179,f7869,f1604])).
fof(f29180,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f1104,f8179,f7869,f1604])).
fof(f29165,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f8141,f1932,f7869,f1603])).
fof(f1932,plain,(
  ( ! [X0,X1] : (sP10('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f1700,f1607])).
fof(f29163,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f8740,f1932,f7869,f1603])).
fof(f29154,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f8141,f8178,f7869,f1603])).
fof(f29146,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f8740,f8777,f7869,f1603])).
fof(f29124,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f8139,f1929,f7869,f1602])).
fof(f29121,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f8738,f1929,f7869,f1602])).
fof(f29117,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f8738,f8778,f7869,f1602])).
fof(f29113,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f8139,f8179,f7869,f1602])).
fof(f29106,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f7869,f8743,f8778,f7869,f1601])).
fof(f29103,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f7869,f8144,f8179,f7869,f1601])).
fof(f29097,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f7869,f8144,f1930,f7869,f1601])).
fof(f29095,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f7869,f8743,f1930,f7869,f1601])).
fof(f29086,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f7869,f8144,f8178,f7869,f1601])).
fof(f29084,plain,(
  $false | ($spl19 | $spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f7869,f8743,f8178,f7869,f1601])).
fof(f29076,plain,(
  $false | ($spl19 | $spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f7869,f8144,f8777,f7869,f1601])).
fof(f29074,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f7869,f8743,f8777,f7869,f1601])).
fof(f29065,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1104,f8144,f1930,f7869,f1601])).
fof(f29062,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1104,f8743,f1930,f7869,f1601])).
fof(f29053,plain,(
  $false | ($spl19 | $spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f7869,f8743,f8178,f7869,f1601])).
fof(f29050,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f7869,f8743,f8777,f7869,f1601])).
fof(f29046,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1104,f8743,f8778,f7869,f1601])).
fof(f29041,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1104,f8144,f8179,f7869,f1601])).
fof(f29030,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f8180,f7869,f1598])).
fof(f8180,plain,(
  sP11('F_0','const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)),'T_0') | $spl19),
  inference(unit_resulting_resolution,[],[f1700,f1104,f7868,f1608])).
fof(f29028,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f8779,f7869,f1598])).
fof(f8779,plain,(
  sP11('F_0','const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),'T_0') | $spl23),
  inference(unit_resulting_resolution,[],[f1700,f1104,f8357,f1608])).
fof(f29019,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f1700,f8129,f7869,f1594])).
fof(f8129,plain,(
  sP9('const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)),'F_0','F_0') | $spl19),
  inference(unit_resulting_resolution,[],[f1700,f1700,f7868,f1591])).
fof(f1591,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f1091])).
fof(f29016,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f1700,f8728,f7869,f1594])).
fof(f8728,plain,(
  sP9('const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),'F_0','F_0') | $spl23),
  inference(unit_resulting_resolution,[],[f1700,f1700,f8357,f1591])).
fof(f29012,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1104,f8059,f7869,f1593])).
fof(f29010,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1104,f8058,f7869,f1593])).
fof(f29008,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1104,f8658,f7869,f1593])).
fof(f29006,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1104,f8657,f7869,f1593])).
fof(f29000,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f7869,f8058,f7869,f1593])).
fof(f28998,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f7869,f8657,f7869,f1593])).
fof(f28994,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1104,f8133,f7869,f1593])).
fof(f28992,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1104,f8732,f7869,f1593])).
fof(f28988,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1104,f8059,f7869,f1592])).
fof(f28986,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1104,f8058,f7869,f1592])).
fof(f28984,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1104,f8658,f7869,f1592])).
fof(f28982,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1104,f8657,f7869,f1592])).
fof(f28976,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f7869,f8059,f7869,f1592])).
fof(f28974,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f7869,f8658,f7869,f1592])).
fof(f28971,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1104,f8136,f7869,f1592])).
fof(f28969,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1104,f8735,f7869,f1592])).
fof(f28965,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8698,f7869,f1589])).
fof(f8698,plain,(
  ~sP8('T_0','const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),'F_0') | $spl23),
  inference(unit_resulting_resolution,[],[f1104,f1700,f8357,f1587])).
fof(f1587,plain,(
  ( ! [X2,X0,X1] : (~sP8(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f1089])).
fof(f28963,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8707,f7869,f1588])).
fof(f8707,plain,(
  ~sP8('T_0','F_0','const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13))) | $spl23),
  inference(unit_resulting_resolution,[],[f1104,f1700,f8357,f1587])).
fof(f28959,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f1700,f8716,f7869,f1587])).
fof(f28956,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f1700,f8117,f7869,f1587])).
fof(f28950,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f8058,f1924,f7869,f1585])).
fof(f28947,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f8657,f1924,f7869,f1585])).
fof(f28943,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f8657,f8716,f7869,f1585])).
fof(f28939,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f8058,f8117,f7869,f1585])).
fof(f28933,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f8059,f1925,f7869,f1584])).
fof(f28930,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f8658,f1925,f7869,f1584])).
fof(f28926,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f8658,f8716,f7869,f1584])).
fof(f28922,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f8059,f8117,f7869,f1584])).
fof(f28914,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f8118,f7869,f1581])).
fof(f8118,plain,(
  sP9('F_0','const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)),'F_0') | $spl19),
  inference(unit_resulting_resolution,[],[f1700,f1700,f7868,f1591])).
fof(f28912,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f8717,f7869,f1581])).
fof(f8717,plain,(
  sP9('F_0','const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),'F_0') | $spl23),
  inference(unit_resulting_resolution,[],[f1700,f1700,f8357,f1591])).
fof(f28906,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f8124,f7869,f1580])).
fof(f8124,plain,(
  sP9('F_0','F_0','const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13))) | $spl19),
  inference(unit_resulting_resolution,[],[f1700,f1700,f7868,f1591])).
fof(f28904,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f8723,f7869,f1580])).
fof(f8723,plain,(
  sP9('F_0','F_0','const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13))) | $spl23),
  inference(unit_resulting_resolution,[],[f1700,f1700,f8357,f1591])).
fof(f28897,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f7869,f8648,f7869,f1579])).
fof(f28894,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f7869,f8049,f7869,f1579])).
fof(f28890,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1104,f1700,f8041,f7869,f1579])).
fof(f28887,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1104,f1700,f8640,f7869,f1579])).
fof(f28879,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f7869,f8041,f7869,f1579])).
fof(f28876,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f7869,f8640,f7869,f1579])).
fof(f28872,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1104,f1700,f8648,f7869,f1579])).
fof(f28869,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1104,f1700,f8049,f7869,f1579])).
fof(f28864,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1104,f1104,f8057,f7869,f1576])).
fof(f28861,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1104,f1104,f8656,f7869,f1576])).
fof(f28852,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1104,f7869,f8620,f7869,f1576])).
fof(f28849,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1104,f7869,f8021,f7869,f1576])).
fof(f28845,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1104,f1104,f8012,f7869,f1576])).
fof(f28842,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1104,f1104,f8611,f7869,f1576])).
fof(f28833,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1104,f7869,f8012,f7869,f1576])).
fof(f28830,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1104,f7869,f8611,f7869,f1576])).
fof(f28826,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1104,f1104,f8620,f7869,f1576])).
fof(f28823,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1104,f1104,f8021,f7869,f1576])).
fof(f28816,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8629,f7869,f1571])).
fof(f8629,plain,(
  ~sP6('T_0','T_0','const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13))) | $spl23),
  inference(unit_resulting_resolution,[],[f1104,f1104,f8357,f1570])).
fof(f28811,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f7869,f8639,f7869,f1570])).
fof(f28808,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f7869,f8040,f7869,f1570])).
fof(f28801,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1104,f1700,f8638,f7869,f1570])).
fof(f28797,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1104,f1700,f8039,f7869,f1570])).
fof(f28788,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f7869,f8638,f7869,f1570])).
fof(f28785,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f7869,f8039,f7869,f1570])).
fof(f28778,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1104,f1700,f8639,f7869,f1570])).
fof(f28774,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1104,f1700,f8040,f7869,f1570])).
fof(f28759,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f7869,f8639,f8620,f7869,f1566])).
fof(f28756,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f7869,f8040,f8021,f7869,f1566])).
fof(f28749,plain,(
  $false | ($spl19 | $spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f7869,f8639,f8012,f7869,f1566])).
fof(f28747,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f7869,f8040,f8012,f7869,f1566])).
fof(f28744,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1104,f1919,f8012,f7869,f1566])).
fof(f28742,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1104,f8039,f8012,f7869,f1566])).
fof(f28735,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f7869,f8639,f8611,f7869,f1566])).
fof(f28733,plain,(
  $false | ($spl19 | $spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f7869,f8040,f8611,f7869,f1566])).
fof(f28730,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1104,f1919,f8611,f7869,f1566])).
fof(f28728,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1104,f8638,f8611,f7869,f1566])).
fof(f28705,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f7869,f8040,f8012,f7869,f1566])).
fof(f28702,plain,(
  $false | ($spl19 | $spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f7869,f8040,f8611,f7869,f1566])).
fof(f28695,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f7869,f8638,f8620,f7869,f1566])).
fof(f28693,plain,(
  $false | ($spl19 | $spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f7869,f8039,f8620,f7869,f1566])).
fof(f28690,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1104,f1919,f8620,f7869,f1566])).
fof(f28688,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1104,f8639,f8620,f7869,f1566])).
fof(f28681,plain,(
  $false | ($spl19 | $spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f7869,f8638,f8021,f7869,f1566])).
fof(f28679,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f7869,f8039,f8021,f7869,f1566])).
fof(f28676,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1104,f1919,f8021,f7869,f1566])).
fof(f28674,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1104,f8040,f8021,f7869,f1566])).
fof(f28659,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f8200,f7869,f1565])).
fof(f8200,plain,(
  ( ! [X0] : (sP7('const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)),'F_0',X0)) ) | $spl19),
  inference(unit_resulting_resolution,[],[f1923,f7868,f1678])).
fof(f1923,plain,(
  ( ! [X0,X1] : (sP6('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f1700,f1573])).
fof(f28657,plain,(
  $false | ($spl7 | $spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f8197,f7869,f1565])).
fof(f8197,plain,(
  ( ! [X0] : (sP7('const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)),X0,sK16)) ) | ($spl7 | $spl19)),
  inference(unit_resulting_resolution,[],[f1920,f7868,f1678])).
fof(f28655,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f8196,f7869,f1565])).
fof(f8196,plain,(
  ( ! [X0] : (sP7('const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)),X0,'F_0')) ) | $spl19),
  inference(unit_resulting_resolution,[],[f1921,f7868,f1678])).
fof(f1921,plain,(
  ( ! [X0,X1] : (sP6(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f1700,f1572])).
fof(f28652,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f8799,f7869,f1565])).
fof(f8799,plain,(
  ( ! [X0] : (sP7('const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),'F_0',X0)) ) | $spl23),
  inference(unit_resulting_resolution,[],[f1923,f8357,f1678])).
fof(f28650,plain,(
  $false | ($spl7 | $spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f8796,f7869,f1565])).
fof(f8796,plain,(
  ( ! [X0] : (sP7('const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),X0,sK16)) ) | ($spl7 | $spl23)),
  inference(unit_resulting_resolution,[],[f1920,f8357,f1678])).
fof(f28648,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f8795,f7869,f1565])).
fof(f8795,plain,(
  ( ! [X0] : (sP7('const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),X0,'F_0')) ) | $spl23),
  inference(unit_resulting_resolution,[],[f1921,f8357,f1678])).
fof(f28645,plain,(
  $false | ($spl7 | $spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f8201,f7869,f1564])).
fof(f8201,plain,(
  ( ! [X0] : (sP7('const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)),sK16,X0)) ) | ($spl7 | $spl19)),
  inference(unit_resulting_resolution,[],[f1922,f7868,f1678])).
fof(f28643,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f8200,f7869,f1564])).
fof(f28640,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f8196,f7869,f1564])).
fof(f28638,plain,(
  $false | ($spl7 | $spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f8800,f7869,f1564])).
fof(f8800,plain,(
  ( ! [X0] : (sP7('const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),sK16,X0)) ) | ($spl7 | $spl23)),
  inference(unit_resulting_resolution,[],[f1922,f8357,f1678])).
fof(f28636,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f8799,f7869,f1564])).
fof(f28633,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f8795,f7869,f1564])).
fof(f28630,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1115,f7980,f7869,f1561])).
fof(f28628,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1115,f7978,f7869,f1561])).
fof(f28626,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1115,f8579,f7869,f1561])).
fof(f28624,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1115,f8577,f7869,f1561])).
fof(f28618,plain,(
  $false | ($spl7 | $spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1936,f7999,f7869,f1560])).
fof(f28616,plain,(
  $false | ($spl7 | $spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1936,f7984,f7869,f1560])).
fof(f28614,plain,(
  $false | ($spl7 | $spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1936,f8598,f7869,f1560])).
fof(f28612,plain,(
  $false | ($spl7 | $spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1936,f8583,f7869,f1560])).
fof(f28560,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8547,f7869,f1558])).
fof(f8547,plain,(
  ~sP2('const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),'F_0','F_0') | $spl23),
  inference(unit_resulting_resolution,[],[f1700,f1700,f8357,f1555])).
fof(f1555,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f1075])).
fof(f28558,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8557,f7869,f1557])).
fof(f8557,plain,(
  ~sP2('F_0','const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),'F_0') | $spl23),
  inference(unit_resulting_resolution,[],[f1700,f1700,f8357,f1555])).
fof(f28556,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8567,f7869,f1556])).
fof(f8567,plain,(
  ~sP2('F_0','F_0','const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13))) | $spl23),
  inference(unit_resulting_resolution,[],[f1700,f1700,f8357,f1555])).
fof(f28554,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f7937,f7869,f1552])).
fof(f7937,plain,(
  ~sP3('T_0','T_0','const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13))) | $spl19),
  inference(unit_resulting_resolution,[],[f1104,f1104,f7868,f1551])).
fof(f28549,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f7869,f8546,f7869,f1551])).
fof(f28546,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f7869,f7947,f7869,f1551])).
fof(f28539,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1104,f1700,f8545,f7869,f1551])).
fof(f28535,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1104,f1700,f7946,f7869,f1551])).
fof(f28526,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f7869,f8545,f7869,f1551])).
fof(f28523,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f7869,f7946,f7869,f1551])).
fof(f28516,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1104,f1700,f8546,f7869,f1551])).
fof(f28512,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1104,f1700,f7947,f7869,f1551])).
fof(f28506,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f7925,f7869,f1548])).
fof(f7925,plain,(
  ~sP4('T_0','T_0','const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13))) | $spl19),
  inference(unit_resulting_resolution,[],[f1104,f1104,f7868,f1547])).
fof(f28501,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f7869,f8535,f7869,f1547])).
fof(f28498,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f7869,f7936,f7869,f1547])).
fof(f28490,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1104,f1700,f8534,f7869,f1547])).
fof(f28486,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1104,f1700,f7935,f7869,f1547])).
fof(f28477,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f7869,f8534,f7869,f1547])).
fof(f28474,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f7869,f7935,f7869,f1547])).
fof(f28467,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1104,f1700,f8535,f7869,f1547])).
fof(f28463,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1104,f1700,f7936,f7869,f1547])).
fof(f28422,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f7892,f1911,f1906,f1917,f7869,f1544])).
fof(f28420,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8491,f8545,f1906,f1917,f7869,f1544])).
fof(f28418,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f7892,f7946,f1906,f1917,f7869,f1544])).
fof(f28413,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8491,f1911,f8534,f1917,f7869,f1544])).
fof(f28411,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8491,f8545,f8534,f1917,f7869,f1544])).
fof(f28407,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f7892,f1911,f7935,f1917,f7869,f1544])).
fof(f28405,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f7892,f7946,f7935,f1917,f7869,f1544])).
fof(f28373,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f7892,f1911,f1906,f1916,f7869,f1544])).
fof(f28371,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8491,f8545,f1906,f1916,f7869,f1544])).
fof(f28369,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f7892,f7946,f1906,f1916,f7869,f1544])).
fof(f28359,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8491,f1911,f8534,f1916,f7869,f1544])).
fof(f28357,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f8491,f8545,f8534,f1916,f7869,f1544])).
fof(f28348,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f7892,f1911,f7935,f1916,f7869,f1544])).
fof(f28346,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f7892,f7946,f7935,f1916,f7869,f1544])).
fof(f28322,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1104,f1700,f8494,f7869,f1543])).
fof(f28318,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1104,f1700,f7895,f7869,f1543])).
fof(f28313,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1104,f1700,f7980,f7869,f1543])).
fof(f28310,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1104,f1700,f8579,f7869,f1543])).
fof(f28305,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f8487,f7869,f1416])).
fof(f8487,plain,(
  ~sP1('F_0','const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13))) | $spl23),
  inference(unit_resulting_resolution,[],[f1700,f8357,f1411])).
fof(f1411,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f1000])).
fof(f28303,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f9259,f7869,f1416])).
fof(f9259,plain,(
  ~sP1('F_0','const.bool.IN_2'(s(sK12,sK29(sK13,sK12)),s(fun(sK12,bool),sK13))) | $spl19),
  inference(unit_resulting_resolution,[],[f1700,f7870,f1411])).
fof(f28301,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f7888,f7869,f1416])).
fof(f7888,plain,(
  ~sP1('F_0','const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13))) | $spl19),
  inference(unit_resulting_resolution,[],[f1700,f7868,f1411])).
fof(f28298,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f27301,f7869,f1413])).
fof(f27301,plain,(
  ~sP1('const.bool.IN_2'(s(sK12,sK28(sK13,sK12)),s(fun(sK12,bool),sK13)),'F_0') | $spl23),
  inference(unit_resulting_resolution,[],[f1700,f8468,f1411])).
fof(f28296,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1700,f8485,f7869,f1413])).
fof(f8485,plain,(
  ~sP1('const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)),'F_0') | $spl23),
  inference(unit_resulting_resolution,[],[f1700,f8357,f1411])).
fof(f28294,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f9257,f7869,f1413])).
fof(f9257,plain,(
  ~sP1('const.bool.IN_2'(s(sK12,sK29(sK13,sK12)),s(fun(sK12,bool),sK13)),'F_0') | $spl19),
  inference(unit_resulting_resolution,[],[f1700,f7870,f1411])).
fof(f28292,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1700,f7886,f7869,f1413])).
fof(f7886,plain,(
  ~sP1('const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)),'F_0') | $spl19),
  inference(unit_resulting_resolution,[],[f1700,f7868,f1411])).
fof(f28289,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1104,f8489,f7869,f1412])).
fof(f28287,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1104,f9261,f7869,f1412])).
fof(f28285,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1104,f7890,f7869,f1412])).
fof(f28281,plain,(
  $false | ($spl20 | $spl23)),
  inference(unit_resulting_resolution,[],[f1104,f8490,f7869,f1412])).
fof(f28279,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1104,f9262,f7869,f1412])).
fof(f28277,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f1104,f7891,f7869,f1412])).
fof(f28273,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f8210,f7869,f1398])).
fof(f1398,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f993])).
fof(f993,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f992])).
fof(f992,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f821])).
fof(f821,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f8210,plain,(
  ~sP0('const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)),'F_0') | $spl19),
  inference(unit_resulting_resolution,[],[f1700,f7868,f1851])).
fof(f1851,plain,(
  ( ! [X0,X1] : (~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f1396,f1700])).
fof(f1396,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f993])).
fof(f28271,plain,(
  $false | ($spl19 | $spl20)),
  inference(unit_resulting_resolution,[],[f8218,f7869,f1397])).
fof(f1397,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f993])).
fof(f8218,plain,(
  ~sP0('F_0','const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13))) | $spl19),
  inference(unit_resulting_resolution,[],[f1700,f7868,f1851])).
fof(f28134,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),sK14),s(fun(fun(sK12,bool),bool),'!_1')))) | $spl46),
  inference(cnf_transformation,[],[f28134_D])).
fof(f28134_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),sK14),s(fun(fun(sK12,bool),bool),'!_1')))) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f28136,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK68(sK14,sK12)),s(fun(sK12,bool),sK13)))) | $spl49),
  inference(cnf_transformation,[],[f28136_D])).
fof(f28136_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK68(sK14,sK12)),s(fun(sK12,bool),sK13)))) <=> ~$spl49),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl49])])).
fof(f27968,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK68(sK14,sK12)),s(fun(sK12,bool),sK13)))) | p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),sK14),s(fun(fun(sK12,bool),bool),'!_1')))) | $spl4),
  inference(superposition,[],[f1841,f1694])).
fof(f1841,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,sK68(X1,X0)),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X1),s(fun(fun(X0,bool),bool),'!_1'))))) )),
  inference(forward_demodulation,[],[f1840,f1257])).
fof(f1840,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X1),s(fun(fun(X0,bool),bool),'!_1')))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK68(X1,X0)))))) )),
  inference(forward_demodulation,[],[f1386,f1257])).
fof(f1386,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK68(X1,X0)))))) )),
  inference(cnf_transformation,[],[f984])).
fof(f984,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1)))) | ! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK68(X1,X0))))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK68])],[f983])).
fof(f983,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1)))) | ! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1))))))),
  inference(rectify,[],[f982])).
fof(f982,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1)))) | ! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1))))))),
  inference(nnf_transformation,[],[f502])).
fof(f502,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1)))) <=> ! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(rectify,[],[f124])).
fof(f124,axiom,(
  ! [X43,X60] : (p(s(bool,happ(s(fun(fun(X43,bool),bool),'!_1'),s(fun(X43,bool),X60)))) <=> ! [X44] : p(s(bool,happ(s(fun(X43,bool),X60),s(X43,X44)))))),
  file('thm.pred_set.EXTENSION','!_HL_QUANT')).
fof(f28132,plain,(
  ( ! [X25] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK55(sK14,X25,sK12)),s(fun(sK12,bool),sK13)))) | p(s(bool,X25))) ) | $spl44),
  inference(cnf_transformation,[],[f28132_D])).
fof(f28132_D,plain,(
  ( ! [X25] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK55(sK14,X25,sK12)),s(fun(sK12,bool),sK13)))) | p(s(bool,X25))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f28131,plain,(
  ( ! [X26,X25] : (p(s(bool,'const.bool.IN_2'(s(sK12,X26),s(fun(sK12,bool),sK13)))) | ~p(s(bool,'const.bool.IN_2'(s(sK12,sK55(sK14,X25,sK12)),s(fun(sK12,bool),sK13)))) | p(s(bool,X25))) ) | $spl4),
  inference(forward_demodulation,[],[f27967,f1694])).
fof(f27967,plain,(
  ( ! [X26,X25] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK55(sK14,X25,sK12)),s(fun(sK12,bool),sK13)))) | p(s(bool,'const.bool.IN_2'(s(sK12,X26),s(fun(sK12,bool),sK14)))) | p(s(bool,X25))) ) | $spl4),
  inference(superposition,[],[f1802,f1694])).
fof(f28130,plain,(
  ( ! [X23] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK54(sK14,X23,sK12)),s(fun(sK12,bool),sK13)))) | p(s(bool,X23))) ) | $spl42),
  inference(cnf_transformation,[],[f28130_D])).
fof(f28130_D,plain,(
  ( ! [X23] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK54(sK14,X23,sK12)),s(fun(sK12,bool),sK13)))) | p(s(bool,X23))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f28129,plain,(
  ( ! [X24,X23] : (p(s(bool,'const.bool.IN_2'(s(sK12,X24),s(fun(sK12,bool),sK13)))) | ~p(s(bool,'const.bool.IN_2'(s(sK12,sK54(sK14,X23,sK12)),s(fun(sK12,bool),sK13)))) | p(s(bool,X23))) ) | $spl4),
  inference(forward_demodulation,[],[f27966,f1694])).
fof(f27966,plain,(
  ( ! [X24,X23] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK54(sK14,X23,sK12)),s(fun(sK12,bool),sK13)))) | p(s(bool,'const.bool.IN_2'(s(sK12,X24),s(fun(sK12,bool),sK14)))) | p(s(bool,X23))) ) | $spl4),
  inference(superposition,[],[f1804,f1694])).
fof(f28126,plain,(
  ( ! [X15] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK49(X15,sK14,sK12)),s(fun(sK12,bool),sK13)))) | p(s(bool,X15))) ) | $spl40),
  inference(cnf_transformation,[],[f28126_D])).
fof(f28126_D,plain,(
  ( ! [X15] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK49(X15,sK14,sK12)),s(fun(sK12,bool),sK13)))) | p(s(bool,X15))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f28125,plain,(
  ( ! [X14,X15] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK49(X15,sK14,sK12)),s(fun(sK12,bool),sK13)))) | ~p(s(bool,'const.bool.IN_2'(s(sK12,X14),s(fun(sK12,bool),sK13)))) | p(s(bool,X15))) ) | $spl4),
  inference(forward_demodulation,[],[f27958,f1694])).
fof(f27958,plain,(
  ( ! [X14,X15] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X14),s(fun(sK12,bool),sK13)))) | p(s(bool,'const.bool.IN_2'(s(sK12,sK49(X15,sK14,sK12)),s(fun(sK12,bool),sK14)))) | p(s(bool,X15))) ) | $spl4),
  inference(superposition,[],[f1788,f1694])).
fof(f1788,plain,(
  ( ! [X6,X2,X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK49(X2,X1,X0)),s(fun(X0,bool),X1)))) | p(s(bool,X2))) )),
  inference(forward_demodulation,[],[f1787,f1257])).
fof(f1787,plain,(
  ( ! [X6,X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK49(X2,X1,X0)),s(fun(X0,bool),X1)))) | p(s(bool,X2)) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6))))) )),
  inference(forward_demodulation,[],[f1295,f1257])).
fof(f1295,plain,(
  ( ! [X6,X2,X0,X1] : (p(s(bool,X2)) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK49(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6))))) )),
  inference(cnf_transformation,[],[f941])).
fof(f941,plain,(
  ! [X0,X1,X2] : (((! [X6] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) & ~p(s(bool,X2))) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK49(X2,X1,X0))))) | p(s(bool,X2)))) & (! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ~p(s(bool,X2))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK50(X2,X1,X0))))) | p(s(bool,X2))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK49,sK50])],[f940])).
fof(f940,plain,(
  ! [X0,X1,X2] : (((! [X6] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) & ~p(s(bool,X2))) | ? [X5] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | p(s(bool,X2)))) & (! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ~p(s(bool,X2))) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2))))),
  inference(rectify,[],[f939])).
fof(f939,plain,(
  ! [X0,X1,X2] : (((! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & ~p(s(bool,X2))) | ? [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,X2)))) & (! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ~p(s(bool,X2))) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2))))),
  inference(flattening,[],[f938])).
fof(f938,plain,(
  ! [X0,X1,X2] : (((! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & ~p(s(bool,X2))) | ? [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,X2)))) & (! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ~p(s(bool,X2))) | (? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2)))))),
  inference(nnf_transformation,[],[f430])).
fof(f430,plain,(
  ! [X0,X1,X2] : ((? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2))) <=> ? [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,X2))))),
  inference(flattening,[],[f429])).
fof(f429,plain,(
  ! [X0] : ! [X1,X2] : ((? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2))) <=> ? [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,X2))))),
  inference(rectify,[],[f99])).
fof(f99,axiom,(
  ! [X6] : ! [X17,X16] : ((? [X9] : p(s(bool,happ(s(fun(X6,bool),X17),s(X6,X9)))) | p(s(bool,X16))) <=> ? [X9] : (p(s(bool,happ(s(fun(X6,bool),X17),s(X6,X9)))) | p(s(bool,X16))))),
  file('thm.pred_set.EXTENSION','thm.bool.LEFT_OR_EXISTS_THM')).
fof(f28124,plain,(
  ( ! [X13] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK50(X13,sK14,sK12)),s(fun(sK12,bool),sK13)))) | p(s(bool,X13))) ) | $spl38),
  inference(cnf_transformation,[],[f28124_D])).
fof(f28124_D,plain,(
  ( ! [X13] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK50(X13,sK14,sK12)),s(fun(sK12,bool),sK13)))) | p(s(bool,X13))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f28123,plain,(
  ( ! [X12,X13] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK50(X13,sK14,sK12)),s(fun(sK12,bool),sK13)))) | ~p(s(bool,'const.bool.IN_2'(s(sK12,X12),s(fun(sK12,bool),sK13)))) | p(s(bool,X13))) ) | $spl4),
  inference(forward_demodulation,[],[f27957,f1694])).
fof(f27957,plain,(
  ( ! [X12,X13] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X12),s(fun(sK12,bool),sK13)))) | p(s(bool,'const.bool.IN_2'(s(sK12,sK50(X13,sK14,sK12)),s(fun(sK12,bool),sK14)))) | p(s(bool,X13))) ) | $spl4),
  inference(superposition,[],[f1786,f1694])).
fof(f1786,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK50(X2,X1,X0)),s(fun(X0,bool),X1)))) | p(s(bool,X2))) )),
  inference(forward_demodulation,[],[f1785,f1257])).
fof(f1785,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK50(X2,X1,X0)),s(fun(X0,bool),X1)))) | p(s(bool,X2)) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) )),
  inference(forward_demodulation,[],[f1297,f1257])).
fof(f1297,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,X2)) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK50(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) )),
  inference(cnf_transformation,[],[f941])).
fof(f28122,plain,(
  ( ! [X11] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK47(sK14,X11,sK12)),s(fun(sK12,bool),sK13)))) | p(s(bool,X11))) ) | $spl36),
  inference(cnf_transformation,[],[f28122_D])).
fof(f28122_D,plain,(
  ( ! [X11] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK47(sK14,X11,sK12)),s(fun(sK12,bool),sK13)))) | p(s(bool,X11))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f28121,plain,(
  ( ! [X10,X11] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK47(sK14,X11,sK12)),s(fun(sK12,bool),sK13)))) | ~p(s(bool,'const.bool.IN_2'(s(sK12,X10),s(fun(sK12,bool),sK13)))) | p(s(bool,X11))) ) | $spl4),
  inference(forward_demodulation,[],[f27956,f1694])).
fof(f27956,plain,(
  ( ! [X10,X11] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X10),s(fun(sK12,bool),sK13)))) | p(s(bool,'const.bool.IN_2'(s(sK12,sK47(sK14,X11,sK12)),s(fun(sK12,bool),sK14)))) | p(s(bool,X11))) ) | $spl4),
  inference(superposition,[],[f1784,f1694])).
fof(f1784,plain,(
  ( ! [X6,X2,X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK47(X2,X1,X0)),s(fun(X0,bool),X2)))) | p(s(bool,X1))) )),
  inference(forward_demodulation,[],[f1783,f1257])).
fof(f1783,plain,(
  ( ! [X6,X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK47(X2,X1,X0)),s(fun(X0,bool),X2)))) | p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) )),
  inference(forward_demodulation,[],[f1292,f1257])).
fof(f1292,plain,(
  ( ! [X6,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK47(X2,X1,X0))))) | p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) )),
  inference(cnf_transformation,[],[f937])).
fof(f937,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X6] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK47(X2,X1,X0))))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK48(X2,X1,X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK47,sK48])],[f936])).
fof(f936,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X6] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | ? [X5] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(rectify,[],[f935])).
fof(f935,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(flattening,[],[f934])).
fof(f934,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f428])).
fof(f428,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f427])).
fof(f427,plain,(
  ! [X0] : ! [X1,X2] : ((p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f97])).
fof(f97,axiom,(
  ! [X6] : ! [X17,X16] : ((p(s(bool,X17)) | ? [X9] : p(s(bool,happ(s(fun(X6,bool),X16),s(X6,X9))))) <=> ? [X9] : (p(s(bool,X17)) | p(s(bool,happ(s(fun(X6,bool),X16),s(X6,X9))))))),
  file('thm.pred_set.EXTENSION','thm.bool.RIGHT_OR_EXISTS_THM')).
fof(f28120,plain,(
  ( ! [X9] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK48(sK14,X9,sK12)),s(fun(sK12,bool),sK13)))) | p(s(bool,X9))) ) | $spl34),
  inference(cnf_transformation,[],[f28120_D])).
fof(f28120_D,plain,(
  ( ! [X9] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK48(sK14,X9,sK12)),s(fun(sK12,bool),sK13)))) | p(s(bool,X9))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f28119,plain,(
  ( ! [X8,X9] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK48(sK14,X9,sK12)),s(fun(sK12,bool),sK13)))) | ~p(s(bool,'const.bool.IN_2'(s(sK12,X8),s(fun(sK12,bool),sK13)))) | p(s(bool,X9))) ) | $spl4),
  inference(forward_demodulation,[],[f27955,f1694])).
fof(f27955,plain,(
  ( ! [X8,X9] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X8),s(fun(sK12,bool),sK13)))) | p(s(bool,'const.bool.IN_2'(s(sK12,sK48(sK14,X9,sK12)),s(fun(sK12,bool),sK14)))) | p(s(bool,X9))) ) | $spl4),
  inference(superposition,[],[f1782,f1694])).
fof(f1782,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK48(X2,X1,X0)),s(fun(X0,bool),X2)))) | p(s(bool,X1))) )),
  inference(forward_demodulation,[],[f1781,f1257])).
fof(f1781,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK48(X2,X1,X0)),s(fun(X0,bool),X2)))) | p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) )),
  inference(forward_demodulation,[],[f1294,f1257])).
fof(f1294,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK48(X2,X1,X0))))) | p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) )),
  inference(cnf_transformation,[],[f937])).
fof(f27239,plain,(
  $false | ($spl1 | $spl4 | $spl26)),
  inference(subsumption_resolution,[],[f27238,f10177])).
fof(f10177,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13)))) ) | $spl26),
  inference(unit_resulting_resolution,[],[f1700,f9234,f1361])).
fof(f27238,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK81(sK14,sK13,sK12,bool)),s(fun(sK12,bool),sK13))) | ($spl1 | $spl4 | $spl26)),
  inference(forward_demodulation,[],[f27237,f10683])).
fof(f10683,plain,(
  ( ! [X4] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,X4),s(fun(sK12,bool),sK14)))) ) | ($spl4 | $spl26)),
  inference(backward_demodulation,[],[f10177,f1694])).
fof(f27235,plain,(
  $false | ($spl1 | $spl4 | $spl26)),
  inference(subsumption_resolution,[],[f27234,f10177])).
fof(f27234,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK81(sK13,sK14,sK12,bool)),s(fun(sK12,bool),sK13))) | ($spl1 | $spl4 | $spl26)),
  inference(forward_demodulation,[],[f27233,f10683])).
fof(f10163,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK12,'const.min.@_1'(s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK13)))) | $spl32),
  inference(cnf_transformation,[],[f10163_D])).
fof(f10163_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK12,'const.min.@_1'(s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK13)))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f10161,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK12,'const.min.@_1'(s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK13)))) | ~p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13))))) ) | $spl4),
  inference(forward_demodulation,[],[f10151,f1694])).
fof(f10151,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13)))) | p(s(bool,'const.bool.IN_2'(s(sK12,'const.min.@_1'(s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK14))))) ) | $spl4),
  inference(superposition,[],[f1734,f1694])).
fof(f1734,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,'const.min.@_1'(s(fun(X0,bool),X1))),s(fun(X0,bool),X1))))) )),
  inference(forward_demodulation,[],[f1733,f1257])).
fof(f1733,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,'const.min.@_1'(s(fun(X0,bool),X1))),s(fun(X0,bool),X1)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) )),
  inference(forward_demodulation,[],[f1263,f1257])).
fof(f1263,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) )),
  inference(cnf_transformation,[],[f752])).
fof(f752,plain,(
  ! [X0,X1,X2] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(ennf_transformation,[],[f410])).
fof(f410,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(flattening,[],[f409])).
fof(f409,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(rectify,[],[f74])).
fof(f74,axiom,(
  ! [X6] : ! [X17,X9] : (p(s(bool,happ(s(fun(X6,bool),X17),s(X6,X9)))) => p(s(bool,happ(s(fun(X6,bool),X17),s(X6,'const.min.@_1'(s(fun(X6,bool),X17)))))))),
  file('thm.pred_set.EXTENSION','thm.bool.SELECT_AX')).
fof(f9724,plain,(
  $false | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f1693,f1694])).
fof(f9826,plain,(
  $false | ($spl3 | $spl4)),
  inference(trivial_inequality_removal,[],[f9745])).
fof(f9745,plain,(
  s(bool,'const.bool.IN_2'(s(sK12,sK15),s(fun(sK12,bool),sK13))) != s(bool,'const.bool.IN_2'(s(sK12,sK15),s(fun(sK12,bool),sK13))) | ($spl3 | $spl4)),
  inference(superposition,[],[f1693,f1694])).
fof(f9717,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15),s(fun(sK12,bool),sK14))) | $spl28),
  inference(cnf_transformation,[],[f9717_D])).
fof(f9717_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15),s(fun(sK12,bool),sK14))) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f9719,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK15),s(fun(sK12,bool),sK13))) | $spl31),
  inference(cnf_transformation,[],[f9719_D])).
fof(f9719_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK15),s(fun(sK12,bool),sK13))) <=> ~$spl31),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl31])])).
fof(f9713,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK15),s(fun(sK12,bool),sK13))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15),s(fun(sK12,bool),sK14))) | $spl3),
  inference(superposition,[],[f1693,f1120])).
fof(f9236,plain,(
  $false | ($spl0 | $spl3)),
  inference(trivial_inequality_removal,[],[f9235])).
fof(f9235,plain,(
  s(bool,'const.bool.IN_2'(s(sK12,sK15),s(fun(sK12,bool),sK13))) != s(bool,'const.bool.IN_2'(s(sK12,sK15),s(fun(sK12,bool),sK13))) | ($spl0 | $spl3)),
  inference(backward_demodulation,[],[f1690,f1693])).
fof(f9233,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),sK14),s(fun(fun(sK12,bool),bool),'?_1')))) | $spl24),
  inference(cnf_transformation,[],[f9233_D])).
fof(f9233_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),sK14),s(fun(fun(sK12,bool),bool),'?_1')))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f9222,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13)))) | p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),sK14),s(fun(fun(sK12,bool),bool),'?_1'))))) ) | $spl4),
  inference(superposition,[],[f1845,f1694])).
fof(f1845,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X1),s(fun(fun(X0,bool),bool),'?_1'))))) )),
  inference(forward_demodulation,[],[f1844,f1257])).
fof(f1844,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X1),s(fun(fun(X0,bool),bool),'?_1')))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) )),
  inference(forward_demodulation,[],[f1388,f1257])).
fof(f1388,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) )),
  inference(cnf_transformation,[],[f987])).
fof(f987,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK69(X1,X0)))))) & (! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK69])],[f986])).
fof(f986,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1)))) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1))))))),
  inference(rectify,[],[f985])).
fof(f985,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1)))) | ? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1))))))),
  inference(nnf_transformation,[],[f503])).
fof(f503,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1)))) <=> ? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(rectify,[],[f123])).
fof(f123,axiom,(
  ! [X43,X60] : (p(s(bool,happ(s(fun(fun(X43,bool),bool),'?_1'),s(fun(X43,bool),X60)))) <=> ? [X44] : p(s(bool,happ(s(fun(X43,bool),X60),s(X43,X44)))))),
  file('thm.pred_set.EXTENSION','?_HL_QUANT')).
fof(f8357,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)))) | $spl23),
  inference(cnf_transformation,[],[f8357_D])).
fof(f8357_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)))) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f8355,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13)))) | ~p(s(bool,'const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13))))) ) | $spl4),
  inference(forward_demodulation,[],[f8345,f1694])).
fof(f8345,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK28(sK14,sK12)),s(fun(sK12,bool),sK13)))) | p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK14))))) ) | $spl4),
  inference(superposition,[],[f1731,f1694])).
fof(f7868,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)))) | $spl19),
  inference(cnf_transformation,[],[f7868_D])).
fof(f7868_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)))) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f7866,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13)))) | ~p(s(bool,'const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13))))) ) | $spl4),
  inference(forward_demodulation,[],[f7856,f1694])).
fof(f7856,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK29(sK14,sK12)),s(fun(sK12,bool),sK13)))) | p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK14))))) ) | $spl4),
  inference(superposition,[],[f1730,f1694])).
fof(f4679,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f4678,f4665])).
fof(f4665,plain,(
  ( ! [X2,X0,X3] : (s(X0,X2) = s(X0,X3)) ) | $spl12),
  inference(subsumption_resolution,[],[f4664,f4647])).
fof(f4647,plain,(
  ( ! [X86] : (p(s(bool,X86))) ) | $spl12),
  inference(global_subsumption,[],[f1103,f1102,f1104,f1106,f1107,f1109,f1108,f1700,f1115,f1117,f1118,f1120,f1130,f1134,f1133,f1137,f1188,f1189,f1701,f1191,f1192,f1193,f1194,f1195,f1196,f1197,f1198,f1199,f1200,f1202,f1203,f1204,f1206,f1207,f1208,f1209,f1210,f1211,f1212,f1215,f1214,f1220,f1219,f1218,f1224,f1223,f1222,f1221,f1228,f1227,f1226,f1225,f1230,f1229,f1232,f1231,f1234,f1233,f1238,f1237,f1236,f1235,f1242,f1241,f1240,f1239,f1246,f1245,f1244,f1243,f1248,f1247,f1250,f1249,f1253,f1254,f1255,f1256,f1257,f1703,f1724,f1725,f1726,f1707,f1708,f1728,f1710,f1711,f1712,f1713,f1729,f1715,f1716,f1717,f1718,f1719,f1720,f1721,f1730,f1731,f1732,f1262,f1734,f1737,f1740,f1742,f1745,f1748,f1751,f1754,f1757,f1760,f1763,f1766,f1769,f1772,f1773,f1280,f1774,f1775,f1283,f1777,f1778,f1285,f1780,f1288,f1782,f1784,f1786,f1788,f1791,f1794,f1797,f1800,f1802,f1804,f1806,f1808,f1810,f1812,f1815,f1818,f1821,f1824,f1319,f1320,f1321,f1322,f1323,f1324,f1325,f1326,f1327,f1328,f1329,f1330,f1825,f1334,f1333,f1332,f1826,f1336,f1827,f1828,f1339,f1829,f1830,f1831,f1833,f1348,f1347,f1346,f1345,f1344,f1350,f1349,f1352,f1351,f1353,f1834,f1355,f1356,f1357,f1358,f1359,f1360,f1364,f1361,f1370,f1381,f1835,f1836,f1679,f1837,f1838,f1839,f1841,f1843,f1845,f1847,f1848,f1849,f1850,f1398,f1397,f1851,f1416,f1413,f1412,f1411,f1429,f1430,f1431,f1852,f1853,f1432,f1436,f1435,f1438,f1437,f1855,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1447,f1448,f1864,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1460,f1459,f1865,f1866,f1867,f1868,f1466,f1869,f1870,f1872,f1873,f1874,f1875,f1876,f1475,f1476,f1477,f1877,f1478,f1479,f1480,f1878,f1483,f1482,f1486,f1485,f1484,f1489,f1488,f1487,f1492,f1491,f1490,f1493,f1494,f1495,f1496,f1879,f1498,f1880,f1500,f1881,f1882,f1503,f1883,f1884,f1885,f1887,f1888,f1889,f1512,f1890,f1891,f1893,f1894,f1895,f1897,f1899,f1518,f1519,f1900,f1901,f1546,f1545,f1544,f1543,f1542,f1541,f1540,f1550,f1549,f1548,f1547,f1554,f1553,f1552,f1551,f1558,f1557,f1556,f1555,f1562,f1561,f1560,f1559,f1678,f1566,f1565,f1564,f1563,f1573,f1572,f1571,f1570,f1579,f1576,f1575,f1574,f1586,f1585,f1584,f1677,f1582,f1581,f1580,f1590,f1589,f1588,f1587,f1594,f1593,f1592,f1591,f1603,f1602,f1601,f1676,f1599,f1598,f1597,f1607,f1606,f1605,f1604,f1611,f1610,f1609,f1608,f1614,f1903,f1615,f1616,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1634,f1635,f1636,f1637,f1638,f1639,f1640,f1641,f1642,f1643,f1644,f1645,f1646,f1647,f1648,f1649,f1650,f1651,f1652,f1653,f1654,f1655,f1656,f1657,f1658,f1659,f1660,f1661,f1662,f1663,f1904,f1905,f1906,f1908,f1910,f1911,f1913,f1915,f1916,f1917,f1918,f1919,f1921,f1923,f1924,f1925,f1927,f1929,f1930,f1932,f1934,f1935,f1946,f1957,f1974,f1975,f1977,f1979,f1985,f1987,f1994,f2019,f2022,f2025,f2028,f2035,f2036,f2038,f2045,f2048,f2052,f2053,f2056,f2059,f2065,f2068,f2080,f2095,f2168,f2169,f2170,f2171,f2172,f2173,f2174,f2175,f2176,f2179,f2183,f2203,f2204,f2271,f2334,f2333,f2332,f2331,f2330,f2329,f2328,f2326,f2322,f2320,f2317,f2316,f2312,f2311,f2307,f2306,f2305,f2303,f2301,f2300,f2299,f2298,f2297,f2296,f2295,f2294,f2293,f2292,f2291,f2290,f2289,f2287,f2286,f2285,f2284,f2283,f2365,f2405,f2420,f2421,f2422,f2423,f2424,f2425,f2426,f2427,f2428,f2471,f2527,f2526,f2524,f2519,f2517,f2515,f2513,f2512,f2511,f2510,f2509,f2506,f2505,f2504,f2503,f2502,f2498,f2497,f2493,f2492,f2491,f2488,f2485,f2484,f2483,f2482,f2481,f2480,f2479,f2478,f2477,f2475,f2474,f2473,f2545,f2544,f2543,f2547,f2552,f2561,f2560,f2559,f2568,f2569,f2570,f2630,f2629,f2627,f2622,f2620,f2618,f2616,f2615,f2614,f2613,f2612,f2609,f2608,f2607,f2606,f2605,f2601,f2600,f2596,f2595,f2594,f2591,f2588,f2587,f2586,f2585,f2584,f2583,f2582,f2581,f2580,f2578,f2577,f2576,f2663,f2739,f2738,f2757,f2756,f2755,f2753,f2751,f2750,f2749,f2747,f2769,f2767,f2789,f2788,f2787,f2785,f2782,f2781,f2780,f2778,f2802,f2801,f2799,f2798,f2796,f2794,f2813,f2829,f2834,f2835,f2836,f2837,f2838,f2839,f2840,f2841,f2842,f2875,f2878,f2879,f2881,f2882,f2884,f2885,f2886,f2889,f2892,f2893,f2895,f2896,f2899,f2900,f2909,f2908,f2906,f2905,f2903,f2901,f2911,f2933,f2950,f2955,f2956,f2957,f2958,f2959,f2960,f2961,f2962,f2963,f2973,f2974,f2998,f2999,f3001,f3002,f3004,f3005,f3006,f3007,f3010,f3011,f3013,f3014,f3017,f3018,f3019,f3028,f3027,f3025,f3024,f3022,f3020,f3052,f3071,f3076,f3077,f3078,f3079,f3080,f3081,f3082,f3083,f3084,f3094,f3095,f3096,f3097,f3119,f3122,f3123,f3125,f3126,f3127,f3128,f3129,f3138,f3139,f3141,f3142,f3145,f3146,f3147,f3148,f3177,f3193,f3201,f3205,f3210,f3214,f3335,f3334,f3333,f3332,f3331,f3330,f3329,f3328,f3326,f3323,f3320,f3319,f3316,f3313,f3311,f3307,f3306,f3305,f3304,f3302,f3300,f3299,f3298,f3294,f3293,f3292,f3291,f3290,f3289,f3288,f3287,f3286,f3285,f3284,f3282,f3280,f3279,f3278,f3277,f3276,f3275,f3273,f3271,f3270,f3269,f3268,f3264,f3263,f3259,f3256,f3251,f3250,f3246,f3242,f3241,f3238,f3237,f3236,f3235,f3233,f3232,f3231,f3225,f3224,f3223,f3361,f3369,f3409,f3416,f3423,f3488,f3496,f3504,f3546,f3767,f3766,f3765,f3764,f3763,f3762,f3761,f3760,f3758,f3757,f3754,f3753,f3748,f3745,f3744,f3740,f3739,f3734,f3719,f3718,f3717,f3712,f3711,f3706,f3705,f3704,f3701,f3697,f3695,f3691,f3690,f3685,f3684,f3683,f3682,f3681,f3677,f3675,f3669,f3668,f3663,f3662,f3656,f3645,f3644,f3643,f3642,f3641,f3640,f3639,f3638,f3637,f3636,f3634,f3632,f3631,f3630,f3629,f3628,f3627,f3625,f3623,f3622,f3621,f3620,f3615,f3614,f3608,f3607,f3601,f3600,f3594,f3592,f3591,f3588,f3587,f3586,f3585,f3583,f3582,f3581,f3575,f3574,f3573,f4065,f4064,f4063,f4062,f4061,f4060,f4059,f4058,f4056,f4055,f4052,f4051,f4045,f4042,f4041,f4037,f4036,f4031,f4015,f4014,f4013,f4008,f4007,f4002,f4001,f4000,f3997,f3993,f3991,f3987,f3986,f3981,f3980,f3979,f3978,f3977,f3973,f3971,f3965,f3964,f3959,f3958,f3952,f3941,f3940,f3939,f3938,f3937,f3936,f3935,f3934,f3933,f3932,f3930,f3928,f3927,f3926,f3925,f3924,f3923,f3921,f3919,f3918,f3917,f3916,f3911,f3910,f3904,f3903,f3897,f3896,f3890,f3888,f3887,f3884,f3883,f3882,f3881,f3879,f3878,f3877,f3871,f3870,f3869,f4134,f4154,f4153,f4151,f4150,f4148,f4146,f4180,f4201,f4206,f4207,f4208,f4209,f4210,f4211,f4212,f4213,f4214,f4224,f4225,f4226,f4227,f4228,f4229,f4258,f4284,f4287,f4300,f4302,f4285,f4319,f4336,f4339,f4343,f4344,f4345,f4346,f4349,f4644,f4350,f4645,f4351,f4646])).
fof(f4646,plain,(
  ( ! [X85,X86,X84] : (p(s(bool,X86)) | sP1('const.bool.COND_3'(s(bool,'T_0'),s(bool,X84),s(bool,X85)),X86)) ) | $spl12),
  inference(global_subsumption,[],[f1103,f1102,f1104,f1106,f1107,f1109,f1108,f1700,f1115,f1117,f1118,f1120,f1130,f1134,f1133,f1137,f1188,f1189,f1701,f1191,f1192,f1193,f1194,f1195,f1196,f1197,f1198,f1199,f1200,f1202,f1203,f1204,f1206,f1207,f1208,f1209,f1210,f1211,f1212,f1215,f1214,f1220,f1219,f1218,f1224,f1223,f1222,f1221,f1228,f1227,f1226,f1225,f1230,f1229,f1232,f1231,f1234,f1233,f1238,f1237,f1236,f1235,f1242,f1241,f1240,f1239,f1246,f1245,f1244,f1243,f1248,f1247,f1250,f1249,f1253,f1254,f1255,f1256,f1257,f1703,f1724,f1725,f1726,f1707,f1708,f1728,f1710,f1711,f1712,f1713,f1729,f1715,f1716,f1717,f1718,f1719,f1720,f1721,f1730,f1731,f1732,f1262,f1734,f1737,f1740,f1742,f1745,f1748,f1751,f1754,f1757,f1760,f1763,f1766,f1769,f1772,f1773,f1280,f1774,f1775,f1283,f1777,f1778,f1285,f1780,f1288,f1782,f1784,f1786,f1788,f1791,f1794,f1797,f1800,f1802,f1804,f1806,f1808,f1810,f1812,f1815,f1818,f1821,f1824,f1319,f1320,f1321,f1322,f1323,f1324,f1325,f1326,f1327,f1328,f1329,f1330,f1825,f1334,f1333,f1332,f1826,f1336,f1827,f1828,f1339,f1829,f1830,f1831,f1833,f1348,f1347,f1346,f1345,f1344,f1350,f1349,f1352,f1351,f1353,f1834,f1355,f1356,f1357,f1358,f1359,f1360,f1364,f1361,f1370,f1381,f1835,f1836,f1679,f1837,f1838,f1839,f1841,f1843,f1845,f1847,f1848,f1849,f1850,f1398,f1397,f1851,f1416,f1413,f1412,f1411,f1429,f1430,f1431,f1852,f1853,f1432,f1436,f1435,f1438,f1437,f1855,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1447,f1448,f1864,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1460,f1459,f1865,f1866,f1867,f1868,f1466,f1869,f1870,f1872,f1873,f1874,f1875,f1876,f1475,f1476,f1477,f1877,f1478,f1479,f1480,f1878,f1483,f1482,f1486,f1485,f1484,f1489,f1488,f1487,f1492,f1491,f1490,f1493,f1494,f1495,f1496,f1879,f1498,f1880,f1500,f1881,f1882,f1503,f1883,f1884,f1885,f1887,f1888,f1889,f1512,f1890,f1891,f1893,f1894,f1895,f1897,f1899,f1518,f1519,f1900,f1901,f1546,f1545,f1544,f1543,f1542,f1541,f1540,f1550,f1549,f1548,f1547,f1554,f1553,f1552,f1551,f1558,f1557,f1556,f1555,f1562,f1561,f1560,f1559,f1678,f1566,f1565,f1564,f1563,f1573,f1572,f1571,f1570,f1579,f1576,f1575,f1574,f1586,f1585,f1584,f1677,f1582,f1581,f1580,f1590,f1589,f1588,f1587,f1594,f1593,f1592,f1591,f1603,f1602,f1601,f1676,f1599,f1598,f1597,f1607,f1606,f1605,f1604,f1611,f1610,f1609,f1608,f1614,f1903,f1615,f1616,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1634,f1635,f1636,f1637,f1638,f1639,f1640,f1641,f1642,f1643,f1644,f1645,f1646,f1647,f1648,f1649,f1650,f1651,f1652,f1653,f1654,f1655,f1656,f1657,f1658,f1659,f1660,f1661,f1662,f1663,f1904,f1905,f1906,f1908,f1910,f1911,f1913,f1915,f1916,f1917,f1918,f1919,f1921,f1923,f1924,f1925,f1927,f1929,f1930,f1932,f1934,f1935,f1946,f1957,f1974,f1975,f1977,f1979,f1985,f1987,f1994,f2019,f2022,f2025,f2028,f2035,f2036,f2038,f2045,f2048,f2052,f2053,f2056,f2059,f2065,f2068,f2080,f2095,f2168,f2169,f2170,f2171,f2172,f2173,f2174,f2175,f2176,f2179,f2183,f2203,f2204,f2271,f2334,f2333,f2332,f2331,f2330,f2329,f2328,f2326,f2322,f2320,f2317,f2316,f2312,f2311,f2307,f2306,f2305,f2303,f2301,f2300,f2299,f2298,f2297,f2296,f2295,f2294,f2293,f2292,f2291,f2290,f2289,f2287,f2286,f2285,f2284,f2283,f2365,f2405,f2420,f2421,f2422,f2423,f2424,f2425,f2426,f2427,f2428,f2471,f2527,f2526,f2524,f2519,f2517,f2515,f2513,f2512,f2511,f2510,f2509,f2506,f2505,f2504,f2503,f2502,f2498,f2497,f2493,f2492,f2491,f2488,f2485,f2484,f2483,f2482,f2481,f2480,f2479,f2478,f2477,f2475,f2474,f2473,f2545,f2544,f2543,f2547,f2552,f2561,f2560,f2559,f2568,f2569,f2570,f2630,f2629,f2627,f2622,f2620,f2618,f2616,f2615,f2614,f2613,f2612,f2609,f2608,f2607,f2606,f2605,f2601,f2600,f2596,f2595,f2594,f2591,f2588,f2587,f2586,f2585,f2584,f2583,f2582,f2581,f2580,f2578,f2577,f2576,f2663,f2739,f2738,f2757,f2756,f2755,f2753,f2751,f2750,f2749,f2747,f2769,f2767,f2789,f2788,f2787,f2785,f2782,f2781,f2780,f2778,f2802,f2801,f2799,f2798,f2796,f2794,f2813,f2829,f2834,f2835,f2836,f2837,f2838,f2839,f2840,f2841,f2842,f2875,f2878,f2879,f2881,f2882,f2884,f2885,f2886,f2889,f2892,f2893,f2895,f2896,f2899,f2900,f2909,f2908,f2906,f2905,f2903,f2901,f2911,f2933,f2950,f2955,f2956,f2957,f2958,f2959,f2960,f2961,f2962,f2963,f2973,f2974,f2998,f2999,f3001,f3002,f3004,f3005,f3006,f3007,f3010,f3011,f3013,f3014,f3017,f3018,f3019,f3028,f3027,f3025,f3024,f3022,f3020,f3052,f3071,f3076,f3077,f3078,f3079,f3080,f3081,f3082,f3083,f3084,f3094,f3095,f3096,f3097,f3119,f3122,f3123,f3125,f3126,f3127,f3128,f3129,f3138,f3139,f3141,f3142,f3145,f3146,f3147,f3148,f3177,f3193,f3201,f3205,f3210,f3214,f3335,f3334,f3333,f3332,f3331,f3330,f3329,f3328,f3326,f3323,f3320,f3319,f3316,f3313,f3311,f3307,f3306,f3305,f3304,f3302,f3300,f3299,f3298,f3294,f3293,f3292,f3291,f3290,f3289,f3288,f3287,f3286,f3285,f3284,f3282,f3280,f3279,f3278,f3277,f3276,f3275,f3273,f3271,f3270,f3269,f3268,f3264,f3263,f3259,f3256,f3251,f3250,f3246,f3242,f3241,f3238,f3237,f3236,f3235,f3233,f3232,f3231,f3225,f3224,f3223,f3361,f3369,f3409,f3416,f3423,f3488,f3496,f3504,f3546,f3767,f3766,f3765,f3764,f3763,f3762,f3761,f3760,f3758,f3757,f3754,f3753,f3748,f3745,f3744,f3740,f3739,f3734,f3719,f3718,f3717,f3712,f3711,f3706,f3705,f3704,f3701,f3697,f3695,f3691,f3690,f3685,f3684,f3683,f3682,f3681,f3677,f3675,f3669,f3668,f3663,f3662,f3656,f3645,f3644,f3643,f3642,f3641,f3640,f3639,f3638,f3637,f3636,f3634,f3632,f3631,f3630,f3629,f3628,f3627,f3625,f3623,f3622,f3621,f3620,f3615,f3614,f3608,f3607,f3601,f3600,f3594,f3592,f3591,f3588,f3587,f3586,f3585,f3583,f3582,f3581,f3575,f3574,f3573,f4065,f4064,f4063,f4062,f4061,f4060,f4059,f4058,f4056,f4055,f4052,f4051,f4045,f4042,f4041,f4037,f4036,f4031,f4015,f4014,f4013,f4008,f4007,f4002,f4001,f4000,f3997,f3993,f3991,f3987,f3986,f3981,f3980,f3979,f3978,f3977,f3973,f3971,f3965,f3964,f3959,f3958,f3952,f3941,f3940,f3939,f3938,f3937,f3936,f3935,f3934,f3933,f3932,f3930,f3928,f3927,f3926,f3925,f3924,f3923,f3921,f3919,f3918,f3917,f3916,f3911,f3910,f3904,f3903,f3897,f3896,f3890,f3888,f3887,f3884,f3883,f3882,f3881,f3879,f3878,f3877,f3871,f3870,f3869,f4134,f4154,f4153,f4151,f4150,f4148,f4146,f4180,f4201,f4206,f4207,f4208,f4209,f4210,f4211,f4212,f4213,f4214,f4224,f4225,f4226,f4227,f4228,f4229,f4258,f4284,f4287,f4300,f4302,f4285,f4319,f4336,f4339,f4343,f4344,f4345,f4346,f4349,f4644,f4350,f4645,f4351])).
fof(f4351,plain,(
  ( ! [X85,X86,X84] : (sP1('const.bool.COND_3'(s(bool,'T_0'),s(bool,X84),s(bool,X85)),X86) | ~p(s(bool,X85)) | p(s(bool,X86))) ) | $spl12),
  inference(backward_demodulation,[],[f4285,f2957])).
fof(f4645,plain,(
  ( ! [X83,X81,X82] : (sP0('const.bool.COND_3'(s(bool,'T_0'),s(bool,X81),s(bool,X82)),X83)) ) | $spl12),
  inference(global_subsumption,[],[f1103,f1102,f1104,f1106,f1107,f1109,f1108,f1700,f1115,f1117,f1118,f1120,f1130,f1134,f1133,f1137,f1188,f1189,f1701,f1191,f1192,f1193,f1194,f1195,f1196,f1197,f1198,f1199,f1200,f1202,f1203,f1204,f1206,f1207,f1208,f1209,f1210,f1211,f1212,f1215,f1214,f1220,f1219,f1218,f1224,f1223,f1222,f1221,f1228,f1227,f1226,f1225,f1230,f1229,f1232,f1231,f1234,f1233,f1238,f1237,f1236,f1235,f1242,f1241,f1240,f1239,f1246,f1245,f1244,f1243,f1248,f1247,f1250,f1249,f1253,f1254,f1255,f1256,f1257,f1703,f1724,f1725,f1726,f1707,f1708,f1728,f1710,f1711,f1712,f1713,f1729,f1715,f1716,f1717,f1718,f1719,f1720,f1721,f1730,f1731,f1732,f1262,f1734,f1737,f1740,f1742,f1745,f1748,f1751,f1754,f1757,f1760,f1763,f1766,f1769,f1772,f1773,f1280,f1774,f1775,f1283,f1777,f1778,f1285,f1780,f1288,f1782,f1784,f1786,f1788,f1791,f1794,f1797,f1800,f1802,f1804,f1806,f1808,f1810,f1812,f1815,f1818,f1821,f1824,f1319,f1320,f1321,f1322,f1323,f1324,f1325,f1326,f1327,f1328,f1329,f1330,f1825,f1334,f1333,f1332,f1826,f1336,f1827,f1828,f1339,f1829,f1830,f1831,f1833,f1348,f1347,f1346,f1345,f1344,f1350,f1349,f1352,f1351,f1353,f1834,f1355,f1356,f1357,f1358,f1359,f1360,f1364,f1361,f1370,f1381,f1835,f1836,f1679,f1837,f1838,f1839,f1841,f1843,f1845,f1847,f1848,f1849,f1850,f1398,f1397,f1851,f1416,f1413,f1412,f1411,f1429,f1430,f1431,f1852,f1853,f1432,f1436,f1435,f1438,f1437,f1855,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1447,f1448,f1864,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1460,f1459,f1865,f1866,f1867,f1868,f1466,f1869,f1870,f1872,f1873,f1874,f1875,f1876,f1475,f1476,f1477,f1877,f1478,f1479,f1480,f1878,f1483,f1482,f1486,f1485,f1484,f1489,f1488,f1487,f1492,f1491,f1490,f1493,f1494,f1495,f1496,f1879,f1498,f1880,f1500,f1881,f1882,f1503,f1883,f1884,f1885,f1887,f1888,f1889,f1512,f1890,f1891,f1893,f1894,f1895,f1897,f1899,f1518,f1519,f1900,f1901,f1546,f1545,f1544,f1543,f1542,f1541,f1540,f1550,f1549,f1548,f1547,f1554,f1553,f1552,f1551,f1558,f1557,f1556,f1555,f1562,f1561,f1560,f1559,f1678,f1566,f1565,f1564,f1563,f1573,f1572,f1571,f1570,f1579,f1576,f1575,f1574,f1586,f1585,f1584,f1677,f1582,f1581,f1580,f1590,f1589,f1588,f1587,f1594,f1593,f1592,f1591,f1603,f1602,f1601,f1676,f1599,f1598,f1597,f1607,f1606,f1605,f1604,f1611,f1610,f1609,f1608,f1614,f1903,f1615,f1616,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1634,f1635,f1636,f1637,f1638,f1639,f1640,f1641,f1642,f1643,f1644,f1645,f1646,f1647,f1648,f1649,f1650,f1651,f1652,f1653,f1654,f1655,f1656,f1657,f1658,f1659,f1660,f1661,f1662,f1663,f1904,f1905,f1906,f1908,f1910,f1911,f1913,f1915,f1916,f1917,f1918,f1919,f1921,f1923,f1924,f1925,f1927,f1929,f1930,f1932,f1934,f1935,f1946,f1957,f1974,f1975,f1977,f1979,f1985,f1987,f1994,f2019,f2022,f2025,f2028,f2035,f2036,f2038,f2045,f2048,f2052,f2053,f2056,f2059,f2065,f2068,f2080,f2095,f2168,f2169,f2170,f2171,f2172,f2173,f2174,f2175,f2176,f2179,f2183,f2203,f2204,f2271,f2334,f2333,f2332,f2331,f2330,f2329,f2328,f2326,f2322,f2320,f2317,f2316,f2312,f2311,f2307,f2306,f2305,f2303,f2301,f2300,f2299,f2298,f2297,f2296,f2295,f2294,f2293,f2292,f2291,f2290,f2289,f2287,f2286,f2285,f2284,f2283,f2365,f2405,f2420,f2421,f2422,f2423,f2424,f2425,f2426,f2427,f2428,f2471,f2527,f2526,f2524,f2519,f2517,f2515,f2513,f2512,f2511,f2510,f2509,f2506,f2505,f2504,f2503,f2502,f2498,f2497,f2493,f2492,f2491,f2488,f2485,f2484,f2483,f2482,f2481,f2480,f2479,f2478,f2477,f2475,f2474,f2473,f2545,f2544,f2543,f2547,f2552,f2561,f2560,f2559,f2568,f2569,f2570,f2630,f2629,f2627,f2622,f2620,f2618,f2616,f2615,f2614,f2613,f2612,f2609,f2608,f2607,f2606,f2605,f2601,f2600,f2596,f2595,f2594,f2591,f2588,f2587,f2586,f2585,f2584,f2583,f2582,f2581,f2580,f2578,f2577,f2576,f2663,f2739,f2738,f2757,f2756,f2755,f2753,f2751,f2750,f2749,f2747,f2769,f2767,f2789,f2788,f2787,f2785,f2782,f2781,f2780,f2778,f2802,f2801,f2799,f2798,f2796,f2794,f2813,f2829,f2834,f2835,f2836,f2837,f2838,f2839,f2840,f2841,f2842,f2875,f2878,f2879,f2881,f2882,f2884,f2885,f2886,f2889,f2892,f2893,f2895,f2896,f2899,f2900,f2909,f2908,f2906,f2905,f2903,f2901,f2911,f2933,f2950,f2955,f2956,f2957,f2958,f2959,f2960,f2961,f2962,f2963,f2973,f2974,f2998,f2999,f3001,f3002,f3004,f3005,f3006,f3007,f3010,f3011,f3013,f3014,f3017,f3018,f3019,f3028,f3027,f3025,f3024,f3022,f3020,f3052,f3071,f3076,f3077,f3078,f3079,f3080,f3081,f3082,f3083,f3084,f3094,f3095,f3096,f3097,f3119,f3122,f3123,f3125,f3126,f3127,f3128,f3129,f3138,f3139,f3141,f3142,f3145,f3146,f3147,f3148,f3177,f3193,f3201,f3205,f3210,f3214,f3335,f3334,f3333,f3332,f3331,f3330,f3329,f3328,f3326,f3323,f3320,f3319,f3316,f3313,f3311,f3307,f3306,f3305,f3304,f3302,f3300,f3299,f3298,f3294,f3293,f3292,f3291,f3290,f3289,f3288,f3287,f3286,f3285,f3284,f3282,f3280,f3279,f3278,f3277,f3276,f3275,f3273,f3271,f3270,f3269,f3268,f3264,f3263,f3259,f3256,f3251,f3250,f3246,f3242,f3241,f3238,f3237,f3236,f3235,f3233,f3232,f3231,f3225,f3224,f3223,f3361,f3369,f3409,f3416,f3423,f3488,f3496,f3504,f3546,f3767,f3766,f3765,f3764,f3763,f3762,f3761,f3760,f3758,f3757,f3754,f3753,f3748,f3745,f3744,f3740,f3739,f3734,f3719,f3718,f3717,f3712,f3711,f3706,f3705,f3704,f3701,f3697,f3695,f3691,f3690,f3685,f3684,f3683,f3682,f3681,f3677,f3675,f3669,f3668,f3663,f3662,f3656,f3645,f3644,f3643,f3642,f3641,f3640,f3639,f3638,f3637,f3636,f3634,f3632,f3631,f3630,f3629,f3628,f3627,f3625,f3623,f3622,f3621,f3620,f3615,f3614,f3608,f3607,f3601,f3600,f3594,f3592,f3591,f3588,f3587,f3586,f3585,f3583,f3582,f3581,f3575,f3574,f3573,f4065,f4064,f4063,f4062,f4061,f4060,f4059,f4058,f4056,f4055,f4052,f4051,f4045,f4042,f4041,f4037,f4036,f4031,f4015,f4014,f4013,f4008,f4007,f4002,f4001,f4000,f3997,f3993,f3991,f3987,f3986,f3981,f3980,f3979,f3978,f3977,f3973,f3971,f3965,f3964,f3959,f3958,f3952,f3941,f3940,f3939,f3938,f3937,f3936,f3935,f3934,f3933,f3932,f3930,f3928,f3927,f3926,f3925,f3924,f3923,f3921,f3919,f3918,f3917,f3916,f3911,f3910,f3904,f3903,f3897,f3896,f3890,f3888,f3887,f3884,f3883,f3882,f3881,f3879,f3878,f3877,f3871,f3870,f3869,f4134,f4154,f4153,f4151,f4150,f4148,f4146,f4180,f4201,f4206,f4207,f4208,f4209,f4210,f4211,f4212,f4213,f4214,f4224,f4225,f4226,f4227,f4228,f4229,f4258,f4284,f4287,f4300,f4302,f4285,f4319,f4336,f4339,f4343,f4344,f4345,f4346,f4349,f4644,f4350])).
fof(f4350,plain,(
  ( ! [X83,X81,X82] : (sP0('const.bool.COND_3'(s(bool,'T_0'),s(bool,X81),s(bool,X82)),X83) | ~p(s(bool,X82))) ) | $spl12),
  inference(backward_demodulation,[],[f4285,f2956])).
fof(f4644,plain,(
  ( ! [X80,X78,X79] : (sP0(X80,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X78),s(bool,X79)))) ) | $spl12),
  inference(global_subsumption,[],[f1103,f1102,f1104,f1106,f1107,f1109,f1108,f1700,f1115,f1117,f1118,f1120,f1130,f1134,f1133,f1137,f1188,f1189,f1701,f1191,f1192,f1193,f1194,f1195,f1196,f1197,f1198,f1199,f1200,f1202,f1203,f1204,f1206,f1207,f1208,f1209,f1210,f1211,f1212,f1215,f1214,f1220,f1219,f1218,f1224,f1223,f1222,f1221,f1228,f1227,f1226,f1225,f1230,f1229,f1232,f1231,f1234,f1233,f1238,f1237,f1236,f1235,f1242,f1241,f1240,f1239,f1246,f1245,f1244,f1243,f1248,f1247,f1250,f1249,f1253,f1254,f1255,f1256,f1257,f1703,f1724,f1725,f1726,f1707,f1708,f1728,f1710,f1711,f1712,f1713,f1729,f1715,f1716,f1717,f1718,f1719,f1720,f1721,f1730,f1731,f1732,f1262,f1734,f1737,f1740,f1742,f1745,f1748,f1751,f1754,f1757,f1760,f1763,f1766,f1769,f1772,f1773,f1280,f1774,f1775,f1283,f1777,f1778,f1285,f1780,f1288,f1782,f1784,f1786,f1788,f1791,f1794,f1797,f1800,f1802,f1804,f1806,f1808,f1810,f1812,f1815,f1818,f1821,f1824,f1319,f1320,f1321,f1322,f1323,f1324,f1325,f1326,f1327,f1328,f1329,f1330,f1825,f1334,f1333,f1332,f1826,f1336,f1827,f1828,f1339,f1829,f1830,f1831,f1833,f1348,f1347,f1346,f1345,f1344,f1350,f1349,f1352,f1351,f1353,f1834,f1355,f1356,f1357,f1358,f1359,f1360,f1364,f1361,f1370,f1381,f1835,f1836,f1679,f1837,f1838,f1839,f1841,f1843,f1845,f1847,f1848,f1849,f1850,f1398,f1397,f1851,f1416,f1413,f1412,f1411,f1429,f1430,f1431,f1852,f1853,f1432,f1436,f1435,f1438,f1437,f1855,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1447,f1448,f1864,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1460,f1459,f1865,f1866,f1867,f1868,f1466,f1869,f1870,f1872,f1873,f1874,f1875,f1876,f1475,f1476,f1477,f1877,f1478,f1479,f1480,f1878,f1483,f1482,f1486,f1485,f1484,f1489,f1488,f1487,f1492,f1491,f1490,f1493,f1494,f1495,f1496,f1879,f1498,f1880,f1500,f1881,f1882,f1503,f1883,f1884,f1885,f1887,f1888,f1889,f1512,f1890,f1891,f1893,f1894,f1895,f1897,f1899,f1518,f1519,f1900,f1901,f1546,f1545,f1544,f1543,f1542,f1541,f1540,f1550,f1549,f1548,f1547,f1554,f1553,f1552,f1551,f1558,f1557,f1556,f1555,f1562,f1561,f1560,f1559,f1678,f1566,f1565,f1564,f1563,f1573,f1572,f1571,f1570,f1579,f1576,f1575,f1574,f1586,f1585,f1584,f1677,f1582,f1581,f1580,f1590,f1589,f1588,f1587,f1594,f1593,f1592,f1591,f1603,f1602,f1601,f1676,f1599,f1598,f1597,f1607,f1606,f1605,f1604,f1611,f1610,f1609,f1608,f1614,f1903,f1615,f1616,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1634,f1635,f1636,f1637,f1638,f1639,f1640,f1641,f1642,f1643,f1644,f1645,f1646,f1647,f1648,f1649,f1650,f1651,f1652,f1653,f1654,f1655,f1656,f1657,f1658,f1659,f1660,f1661,f1662,f1663,f1904,f1905,f1906,f1908,f1910,f1911,f1913,f1915,f1916,f1917,f1918,f1919,f1921,f1923,f1924,f1925,f1927,f1929,f1930,f1932,f1934,f1935,f1946,f1957,f1974,f1975,f1977,f1979,f1985,f1987,f1994,f2019,f2022,f2025,f2028,f2035,f2036,f2038,f2045,f2048,f2052,f2053,f2056,f2059,f2065,f2068,f2080,f2095,f2168,f2169,f2170,f2171,f2172,f2173,f2174,f2175,f2176,f2179,f2183,f2203,f2204,f2271,f2334,f2333,f2332,f2331,f2330,f2329,f2328,f2326,f2322,f2320,f2317,f2316,f2312,f2311,f2307,f2306,f2305,f2303,f2301,f2300,f2299,f2298,f2297,f2296,f2295,f2294,f2293,f2292,f2291,f2290,f2289,f2287,f2286,f2285,f2284,f2283,f2365,f2405,f2420,f2421,f2422,f2423,f2424,f2425,f2426,f2427,f2428,f2471,f2527,f2526,f2524,f2519,f2517,f2515,f2513,f2512,f2511,f2510,f2509,f2506,f2505,f2504,f2503,f2502,f2498,f2497,f2493,f2492,f2491,f2488,f2485,f2484,f2483,f2482,f2481,f2480,f2479,f2478,f2477,f2475,f2474,f2473,f2545,f2544,f2543,f2547,f2552,f2561,f2560,f2559,f2568,f2569,f2570,f2630,f2629,f2627,f2622,f2620,f2618,f2616,f2615,f2614,f2613,f2612,f2609,f2608,f2607,f2606,f2605,f2601,f2600,f2596,f2595,f2594,f2591,f2588,f2587,f2586,f2585,f2584,f2583,f2582,f2581,f2580,f2578,f2577,f2576,f2663,f2739,f2738,f2757,f2756,f2755,f2753,f2751,f2750,f2749,f2747,f2769,f2767,f2789,f2788,f2787,f2785,f2782,f2781,f2780,f2778,f2802,f2801,f2799,f2798,f2796,f2794,f2813,f2829,f2834,f2835,f2836,f2837,f2838,f2839,f2840,f2841,f2842,f2875,f2878,f2879,f2881,f2882,f2884,f2885,f2886,f2889,f2892,f2893,f2895,f2896,f2899,f2900,f2909,f2908,f2906,f2905,f2903,f2901,f2911,f2933,f2950,f2955,f2956,f2957,f2958,f2959,f2960,f2961,f2962,f2963,f2973,f2974,f2998,f2999,f3001,f3002,f3004,f3005,f3006,f3007,f3010,f3011,f3013,f3014,f3017,f3018,f3019,f3028,f3027,f3025,f3024,f3022,f3020,f3052,f3071,f3076,f3077,f3078,f3079,f3080,f3081,f3082,f3083,f3084,f3094,f3095,f3096,f3097,f3119,f3122,f3123,f3125,f3126,f3127,f3128,f3129,f3138,f3139,f3141,f3142,f3145,f3146,f3147,f3148,f3177,f3193,f3201,f3205,f3210,f3214,f3335,f3334,f3333,f3332,f3331,f3330,f3329,f3328,f3326,f3323,f3320,f3319,f3316,f3313,f3311,f3307,f3306,f3305,f3304,f3302,f3300,f3299,f3298,f3294,f3293,f3292,f3291,f3290,f3289,f3288,f3287,f3286,f3285,f3284,f3282,f3280,f3279,f3278,f3277,f3276,f3275,f3273,f3271,f3270,f3269,f3268,f3264,f3263,f3259,f3256,f3251,f3250,f3246,f3242,f3241,f3238,f3237,f3236,f3235,f3233,f3232,f3231,f3225,f3224,f3223,f3361,f3369,f3409,f3416,f3423,f3488,f3496,f3504,f3546,f3767,f3766,f3765,f3764,f3763,f3762,f3761,f3760,f3758,f3757,f3754,f3753,f3748,f3745,f3744,f3740,f3739,f3734,f3719,f3718,f3717,f3712,f3711,f3706,f3705,f3704,f3701,f3697,f3695,f3691,f3690,f3685,f3684,f3683,f3682,f3681,f3677,f3675,f3669,f3668,f3663,f3662,f3656,f3645,f3644,f3643,f3642,f3641,f3640,f3639,f3638,f3637,f3636,f3634,f3632,f3631,f3630,f3629,f3628,f3627,f3625,f3623,f3622,f3621,f3620,f3615,f3614,f3608,f3607,f3601,f3600,f3594,f3592,f3591,f3588,f3587,f3586,f3585,f3583,f3582,f3581,f3575,f3574,f3573,f4065,f4064,f4063,f4062,f4061,f4060,f4059,f4058,f4056,f4055,f4052,f4051,f4045,f4042,f4041,f4037,f4036,f4031,f4015,f4014,f4013,f4008,f4007,f4002,f4001,f4000,f3997,f3993,f3991,f3987,f3986,f3981,f3980,f3979,f3978,f3977,f3973,f3971,f3965,f3964,f3959,f3958,f3952,f3941,f3940,f3939,f3938,f3937,f3936,f3935,f3934,f3933,f3932,f3930,f3928,f3927,f3926,f3925,f3924,f3923,f3921,f3919,f3918,f3917,f3916,f3911,f3910,f3904,f3903,f3897,f3896,f3890,f3888,f3887,f3884,f3883,f3882,f3881,f3879,f3878,f3877,f3871,f3870,f3869,f4134,f4154,f4153,f4151,f4150,f4148,f4146,f4180,f4201,f4206,f4207,f4208,f4209,f4210,f4211,f4212,f4213,f4214,f4224,f4225,f4226,f4227,f4228,f4229,f4258,f4284,f4287,f4300,f4302,f4285,f4319,f4336,f4339,f4343,f4344,f4345,f4346,f4349])).
fof(f4349,plain,(
  ( ! [X80,X78,X79] : (sP0(X80,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X78),s(bool,X79))) | ~p(s(bool,X79))) ) | $spl12),
  inference(backward_demodulation,[],[f4285,f2955])).
fof(f4346,plain,(
  ( ! [X0,X1] : (sP5('T_0',X0,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X0)))) ) | $spl12),
  inference(backward_demodulation,[],[f4285,f2909])).
fof(f4345,plain,(
  ( ! [X0,X1] : (~sP5('F_0',X0,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X0)))) ) | $spl12),
  inference(backward_demodulation,[],[f4285,f2908])).
fof(f4344,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)),X1)) ) | $spl12),
  inference(backward_demodulation,[],[f4285,f2906])).
fof(f4343,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)),X1)) ) | $spl12),
  inference(backward_demodulation,[],[f4285,f2905])).
fof(f4339,plain,(
  ( ! [X8,X7,X9] : (sP5('T_0','const.relation.EMPTY_REL_2'(s(X7,X8),s(X7,X9)),'const.combin.I_1'(s(bool,'T_0')))) ) | $spl12),
  inference(backward_demodulation,[],[f4285,f2896])).
fof(f4336,plain,(
  ( ! [X8,X7,X9] : (sP5('T_0','const.combin.I_1'(s(bool,'T_0')),'const.relation.EMPTY_REL_2'(s(X7,X8),s(X7,X9)))) ) | $spl12),
  inference(backward_demodulation,[],[f4285,f2882])).
fof(f4319,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) ) | $spl12),
  inference(backward_demodulation,[],[f4285,f1254])).
fof(f4302,plain,(
  ( ! [X2,X3] : (~p(s(bool,'|dollar|or_2'(s(bool,X3),s(bool,X2)))) | p(s(bool,X3)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f4272,f1700])).
fof(f4272,plain,(
  ( ! [X2,X3] : (~p(s(bool,'|dollar|or_2'(s(bool,X3),s(bool,X2)))) | p(s(bool,X3)) | p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1679,f1120])).
fof(f4287,plain,(
  ( ! [X2,X3] : (~p(s(bool,'|dollar|or_2'(s(bool,X2),s(bool,X3)))) | p(s(bool,X3)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f4260,f1700])).
fof(f4260,plain,(
  ( ! [X2,X3] : (~p(s(bool,'|dollar|or_2'(s(bool,X2),s(bool,X3)))) | p(s(bool,'F_0')) | p(s(bool,X3)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1679,f1120])).
fof(f4258,plain,(
  ~p(s(bool,'|dollar|or_2'(s(bool,'F_0'),s(bool,'F_0'))))),
  inference(unit_resulting_resolution,[],[f1700,f1700,f1679])).
fof(f4229,plain,(
  ( ! [X210,X208,X209] : (sP4('T_0','const.combin.I_1'(s(bool,X209)),happ(s(fun(X208,bool),'_dst_x0_1'(s(bool,X209))),s(X208,X210)))) )),
  inference(superposition,[],[f2889,f1448])).
fof(f4228,plain,(
  ( ! [X206,X207,X205] : (sP4('T_0',happ(s(fun(X205,bool),'_dst_x0_1'(s(bool,X206))),s(X205,X207)),'const.combin.I_1'(s(bool,X206)))) )),
  inference(superposition,[],[f2875,f1448])).
fof(f4227,plain,(
  ( ! [X204,X202,X203] : (~sP5('F_0',happ(s(fun(X202,bool),'_dst_x0_1'(s(bool,X203))),s(X202,X204)),'const.combin.I_1'(s(bool,X203)))) )),
  inference(superposition,[],[f2788,f1448])).
fof(f4226,plain,(
  ( ! [X200,X199,X201] : (~sP5('F_0','const.combin.I_1'(s(bool,X200)),happ(s(fun(X199,bool),'_dst_x0_1'(s(bool,X200))),s(X199,X201)))) )),
  inference(superposition,[],[f2781,f1448])).
fof(f4225,plain,(
  ( ! [X198,X196,X197] : (sP5('T_0',happ(s(fun(X196,bool),'_dst_x0_1'(s(bool,X197))),s(X196,X198)),'const.combin.I_1'(s(bool,X197)))) )),
  inference(superposition,[],[f2756,f1448])).
fof(f4224,plain,(
  ( ! [X194,X195,X193] : (sP5('T_0','const.combin.I_1'(s(bool,X194)),happ(s(fun(X193,bool),'_dst_x0_1'(s(bool,X194))),s(X193,X195)))) )),
  inference(superposition,[],[f2750,f1448])).
fof(f4214,plain,(
  ( ! [X158,X156,X159,X157,X155] : (~p(s(bool,X156)) | ~sP8(happ(s(fun(X155,bool),'_dst_x0_1'(s(bool,X156))),s(X155,X157)),X158,X159) | sP9(happ(s(fun(X155,bool),'_dst_x0_1'(s(bool,X156))),s(X155,X157)),X158,X159)) )),
  inference(superposition,[],[f1677,f1448])).
fof(f4213,plain,(
  ( ! [X154,X152,X151,X153,X150] : (~p(s(bool,X151)) | ~sP10(happ(s(fun(X150,bool),'_dst_x0_1'(s(bool,X151))),s(X150,X152)),X153,X154) | sP11(happ(s(fun(X150,bool),'_dst_x0_1'(s(bool,X151))),s(X150,X152)),X153,X154)) )),
  inference(superposition,[],[f1676,f1448])).
fof(f4212,plain,(
  ( ! [X146,X149,X147,X145,X148] : (~p(s(bool,X146)) | ~p(s(bool,X148)) | sP11(happ(s(fun(X145,bool),'_dst_x0_1'(s(bool,X146))),s(X145,X147)),X148,X149)) )),
  inference(superposition,[],[f1610,f1448])).
fof(f4211,plain,(
  ( ! [X144,X142,X140,X143,X141] : (~p(s(bool,X141)) | p(s(bool,X143)) | sP11(happ(s(fun(X140,bool),'_dst_x0_1'(s(bool,X141))),s(X140,X142)),X144,X143)) )),
  inference(superposition,[],[f1609,f1448])).
fof(f4210,plain,(
  ( ! [X138,X136,X135,X139,X137] : (~p(s(bool,X136)) | ~p(s(bool,X138)) | sP9(happ(s(fun(X135,bool),'_dst_x0_1'(s(bool,X136))),s(X135,X137)),X138,X139)) )),
  inference(superposition,[],[f1593,f1448])).
fof(f4209,plain,(
  ( ! [X134,X132,X133,X131] : (~p(s(bool,X132)) | p(s(bool,X134)) | sP1(X134,happ(s(fun(X131,bool),'_dst_x0_1'(s(bool,X132))),s(X131,X133)))) )),
  inference(superposition,[],[f1416,f1448])).
fof(f4208,plain,(
  ( ! [X127,X130,X128,X129] : (~p(s(bool,X128)) | p(s(bool,X130)) | sP1(happ(s(fun(X127,bool),'_dst_x0_1'(s(bool,X128))),s(X127,X129)),X130)) )),
  inference(superposition,[],[f1413,f1448])).
fof(f4207,plain,(
  ( ! [X125,X123,X126,X124] : (~p(s(bool,X124)) | sP0(happ(s(fun(X123,bool),'_dst_x0_1'(s(bool,X124))),s(X123,X125)),X126)) )),
  inference(superposition,[],[f1398,f1448])).
fof(f4206,plain,(
  ( ! [X121,X122,X120,X119] : (~p(s(bool,X120)) | sP0(X122,happ(s(fun(X119,bool),'_dst_x0_1'(s(bool,X120))),s(X119,X121)))) )),
  inference(superposition,[],[f1397,f1448])).
fof(f4201,plain,(
  ( ! [X101,X99,X97,X102,X100,X98] : (s(bool,'T_0') != s(bool,X98) | p(s(bool,sK64(X100,X101,happ(s(fun(X97,bool),'_dst_x0_1'(s(bool,X98))),s(X97,X99)),X102)))) )),
  inference(superposition,[],[f1347,f1448])).
fof(f4180,plain,(
  ( ! [X6,X4,X5] : (s(bool,'F_0') = s(bool,X5) | s(bool,'T_0') = s(bool,happ(s(fun(X4,bool),'_dst_x0_1'(s(bool,X5))),s(X4,X6)))) )),
  inference(superposition,[],[f1448,f1120])).
fof(f4146,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,sK64(X0,X1,happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)),X4)))) )),
  inference(unit_resulting_resolution,[],[f1448,f1347])).
fof(f4148,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,sK63(X0,X1,happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X2,X3)),X4)))) )),
  inference(unit_resulting_resolution,[],[f1448,f1345])).
fof(f4150,plain,(
  ( ! [X2,X0,X1] : (~sP5('F_0',happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,X1))),s(X0,X2)),X1)) )),
  inference(unit_resulting_resolution,[],[f1700,f1448,f1562])).
fof(f4151,plain,(
  ( ! [X2,X0,X1] : (sP5('T_0',happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,X1))),s(X0,X2)),X1)) )),
  inference(unit_resulting_resolution,[],[f1104,f1448,f1560])).
fof(f4153,plain,(
  ( ! [X2,X0,X1] : (~sP5('F_0',X0,happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,X0))),s(X1,X2)))) )),
  inference(unit_resulting_resolution,[],[f1700,f1448,f1562])).
fof(f4154,plain,(
  ( ! [X2,X0,X1] : (sP5('T_0',X0,happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,X0))),s(X1,X2)))) )),
  inference(unit_resulting_resolution,[],[f1104,f1448,f1560])).
fof(f4134,plain,(
  ( ! [X2,X0,X1] : (sP11(sK63(X0,X1,'T_0',X2),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f1104,f1700,f2513,f1603])).
fof(f3869,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))))) )),
  inference(unit_resulting_resolution,[],[f1104,f1191,f1364])).
fof(f3870,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1))))) = s(bool,'const.relation.symmetric_1'(s(fun(X2,fun(X2,bool)),'const.relation.SC_1'(s(fun(X2,fun(X2,bool)),X3)))))) )),
  inference(unit_resulting_resolution,[],[f1191,f1191,f1364])).
fof(f3871,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))))) )),
  inference(unit_resulting_resolution,[],[f1700,f1191,f1370])).
fof(f3877,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,'const.relation.symmetric_1'(s(fun(X1,fun(X1,bool)),'const.relation.SC_1'(s(fun(X1,fun(X1,bool)),X2)))))) )),
  inference(unit_resulting_resolution,[],[f1191,f1397])).
fof(f3878,plain,(
  ( ! [X2,X0,X1] : (sP0('const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))),X2)) )),
  inference(unit_resulting_resolution,[],[f1191,f1398])).
fof(f3879,plain,(
  ( ! [X0,X1] : (~sP1('const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f1104,f1191,f1412])).
fof(f3881,plain,(
  ( ! [X0,X1] : (~sP1('T_0','const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))))) )),
  inference(unit_resulting_resolution,[],[f1104,f1191,f1412])).
fof(f3882,plain,(
  ( ! [X2,X0,X3,X1] : (~sP1('const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))),'const.relation.symmetric_1'(s(fun(X2,fun(X2,bool)),'const.relation.SC_1'(s(fun(X2,fun(X2,bool)),X3)))))) )),
  inference(unit_resulting_resolution,[],[f1191,f1191,f1412])).
fof(f3883,plain,(
  ( ! [X0,X1] : (sP1('const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1700,f1191,f1413])).
fof(f3884,plain,(
  ( ! [X0,X1] : (sP1('F_0','const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))))) )),
  inference(unit_resulting_resolution,[],[f1700,f1191,f1416])).
fof(f3887,plain,(
  ( ! [X0,X1] : (~sP5('T_0','const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1104,f1700,f1191,f1543])).
fof(f3888,plain,(
  ( ! [X2,X0,X3,X1] : (~sP5('const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))),'const.relation.symmetric_1'(s(fun(X2,fun(X2,bool)),'const.relation.SC_1'(s(fun(X2,fun(X2,bool)),X3)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1191,f1700,f1191,f1543])).
fof(f3890,plain,(
  ( ! [X0,X1] : (~sP4('const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f1104,f1700,f1191,f1547])).
fof(f3896,plain,(
  ( ! [X0,X1] : (~sP4('T_0','const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1104,f1700,f1191,f1547])).
fof(f3897,plain,(
  ( ! [X2,X0,X3,X1] : (~sP4('const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))),'const.relation.symmetric_1'(s(fun(X2,fun(X2,bool)),'const.relation.SC_1'(s(fun(X2,fun(X2,bool)),X3)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1191,f1700,f1191,f1547])).
fof(f3903,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(X0,X1,'const.relation.symmetric_1'(s(fun(X2,fun(X2,bool)),'const.relation.SC_1'(s(fun(X2,fun(X2,bool)),X3)))))) )),
  inference(unit_resulting_resolution,[],[f1191,f1548])).
fof(f3904,plain,(
  ( ! [X0,X1] : (~sP3('const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f1104,f1700,f1191,f1551])).
fof(f3910,plain,(
  ( ! [X0,X1] : (~sP3('T_0','const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1104,f1700,f1191,f1551])).
fof(f3911,plain,(
  ( ! [X2,X0,X3,X1] : (~sP3('const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))),'const.relation.symmetric_1'(s(fun(X2,fun(X2,bool)),'const.relation.SC_1'(s(fun(X2,fun(X2,bool)),X3)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1191,f1700,f1191,f1551])).
fof(f3916,plain,(
  ( ! [X2,X0,X3,X1] : (sP3(X0,X1,'const.relation.symmetric_1'(s(fun(X2,fun(X2,bool)),'const.relation.SC_1'(s(fun(X2,fun(X2,bool)),X3)))))) )),
  inference(unit_resulting_resolution,[],[f1191,f1552])).
fof(f3917,plain,(
  ( ! [X2,X0,X3,X1] : (sP2(X0,X1,'const.relation.symmetric_1'(s(fun(X2,fun(X2,bool)),'const.relation.SC_1'(s(fun(X2,fun(X2,bool)),X3)))))) )),
  inference(unit_resulting_resolution,[],[f1191,f1556])).
fof(f3918,plain,(
  ( ! [X2,X0,X3,X1] : (sP2(X0,'const.relation.symmetric_1'(s(fun(X1,fun(X1,bool)),'const.relation.SC_1'(s(fun(X1,fun(X1,bool)),X2)))),X3)) )),
  inference(unit_resulting_resolution,[],[f1191,f1557])).
fof(f3919,plain,(
  ( ! [X2,X0,X3,X1] : (sP2('const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))),X2,X3)) )),
  inference(unit_resulting_resolution,[],[f1191,f1558])).
fof(f3921,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP5('const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))),'const.relation.EMPTY_REL_2'(s(X2,X3),s(X2,X4)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1253,f1191,f1560])).
fof(f3923,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP5('const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))),'const.relation.RUNIV_2'(s(X2,X3),s(X4,X5)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f1447,f1191,f1560])).
fof(f3924,plain,(
  ( ! [X2,X0,X1] : (sP5('const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))),'const.combin.I_1'(s(bool,X2)),X2)) )),
  inference(unit_resulting_resolution,[],[f1188,f1191,f1560])).
fof(f3925,plain,(
  ( ! [X2,X0,X1] : (sP5('const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))),happ(s(fun(bool,bool),'_i_0_0'),s(bool,X2)),X2)) )),
  inference(unit_resulting_resolution,[],[f1189,f1191,f1560])).
fof(f3926,plain,(
  ( ! [X2,X0,X1] : (sP5('const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))),happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X2)),X2)) )),
  inference(unit_resulting_resolution,[],[f1701,f1191,f1560])).
fof(f3927,plain,(
  ( ! [X2,X0,X3,X1] : (sP5('const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,X3)),X3)) )),
  inference(unit_resulting_resolution,[],[f1254,f1191,f1560])).
fof(f3928,plain,(
  ( ! [X2,X0,X3,X1] : (sP5('const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X2),s(bool,X3)),X2)) )),
  inference(unit_resulting_resolution,[],[f1255,f1191,f1560])).
fof(f3930,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP5('const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))),'F_0','const.relation.EMPTY_REL_2'(s(X2,X3),s(X2,X4)))) )),
  inference(unit_resulting_resolution,[],[f1253,f1191,f1560])).
fof(f3932,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP5('const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))),'T_0','const.relation.RUNIV_2'(s(X2,X3),s(X4,X5)))) )),
  inference(unit_resulting_resolution,[],[f1447,f1191,f1560])).
fof(f3933,plain,(
  ( ! [X2,X0,X1] : (sP5('const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))),X2,'const.combin.I_1'(s(bool,X2)))) )),
  inference(unit_resulting_resolution,[],[f1188,f1191,f1560])).
fof(f3934,plain,(
  ( ! [X2,X0,X1] : (sP5('const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))),X2,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X2)))) )),
  inference(unit_resulting_resolution,[],[f1189,f1191,f1560])).
fof(f3935,plain,(
  ( ! [X2,X0,X1] : (sP5('const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))),X2,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X2)))) )),
  inference(unit_resulting_resolution,[],[f1701,f1191,f1560])).
fof(f3936,plain,(
  ( ! [X2,X0,X3,X1] : (sP5('const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))),X2,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,X2)))) )),
  inference(unit_resulting_resolution,[],[f1254,f1191,f1560])).
fof(f3937,plain,(
  ( ! [X2,X0,X3,X1] : (sP5('const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))),X2,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X2),s(bool,X3)))) )),
  inference(unit_resulting_resolution,[],[f1255,f1191,f1560])).
fof(f3938,plain,(
  ( ! [X0,X1] : (~sP5('const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f1115,f1191,f1561])).
fof(f3939,plain,(
  ( ! [X0,X1] : (~sP5('const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f1115,f1191,f1561])).
fof(f3940,plain,(
  ( ! [X2,X0,X1] : (~sP7('const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))),X2,'F_0')) )),
  inference(unit_resulting_resolution,[],[f1700,f1191,f1564])).
fof(f3941,plain,(
  ( ! [X2,X0,X1] : (~sP7('const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))),'F_0',X2)) )),
  inference(unit_resulting_resolution,[],[f1700,f1191,f1565])).
fof(f3952,plain,(
  ( ! [X0,X1] : (~sP6('const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f1104,f1700,f1191,f1570])).
fof(f3958,plain,(
  ( ! [X0,X1] : (~sP6('T_0','const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1104,f1700,f1191,f1570])).
fof(f3959,plain,(
  ( ! [X2,X0,X3,X1] : (~sP6('const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))),'const.relation.symmetric_1'(s(fun(X2,fun(X2,bool)),'const.relation.SC_1'(s(fun(X2,fun(X2,bool)),X3)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1191,f1700,f1191,f1570])).
fof(f3964,plain,(
  ( ! [X2,X0,X3,X1] : (sP6(X0,X1,'const.relation.symmetric_1'(s(fun(X2,fun(X2,bool)),'const.relation.SC_1'(s(fun(X2,fun(X2,bool)),X3)))))) )),
  inference(unit_resulting_resolution,[],[f1191,f1571])).
fof(f3965,plain,(
  ( ! [X0,X1] : (sP7('T_0','const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f1104,f1104,f1191,f1576])).
fof(f3971,plain,(
  ( ! [X0,X1] : (sP7('T_0','T_0','const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))))) )),
  inference(unit_resulting_resolution,[],[f1104,f1104,f1191,f1576])).
fof(f3973,plain,(
  ( ! [X2,X0,X3,X1] : (sP7('T_0','const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))),'const.relation.symmetric_1'(s(fun(X2,fun(X2,bool)),'const.relation.SC_1'(s(fun(X2,fun(X2,bool)),X3)))))) )),
  inference(unit_resulting_resolution,[],[f1104,f1191,f1191,f1576])).
fof(f3977,plain,(
  ( ! [X0,X1] : (sP7('const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))),'T_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f1104,f1104,f1191,f1576])).
fof(f3978,plain,(
  ( ! [X2,X0,X3,X1] : (sP7('const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))),'const.relation.symmetric_1'(s(fun(X2,fun(X2,bool)),'const.relation.SC_1'(s(fun(X2,fun(X2,bool)),X3)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f1191,f1104,f1191,f1576])).
fof(f3979,plain,(
  ( ! [X2,X0,X3,X1] : (sP7('const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))),'T_0','const.relation.symmetric_1'(s(fun(X2,fun(X2,bool)),'const.relation.SC_1'(s(fun(X2,fun(X2,bool)),X3)))))) )),
  inference(unit_resulting_resolution,[],[f1104,f1191,f1191,f1576])).
fof(f3980,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP7('const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))),'const.relation.symmetric_1'(s(fun(X2,fun(X2,bool)),'const.relation.SC_1'(s(fun(X2,fun(X2,bool)),X3)))),'const.relation.symmetric_1'(s(fun(X4,fun(X4,bool)),'const.relation.SC_1'(s(fun(X4,fun(X4,bool)),X5)))))) )),
  inference(unit_resulting_resolution,[],[f1191,f1191,f1191,f1576])).
fof(f3981,plain,(
  ( ! [X0,X1] : (~sP7('F_0','const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f1700,f1104,f1191,f1579])).
fof(f3986,plain,(
  ( ! [X0,X1] : (~sP7('F_0','T_0','const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))))) )),
  inference(unit_resulting_resolution,[],[f1700,f1104,f1191,f1579])).
fof(f3987,plain,(
  ( ! [X2,X0,X3,X1] : (~sP7('F_0','const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))),'const.relation.symmetric_1'(s(fun(X2,fun(X2,bool)),'const.relation.SC_1'(s(fun(X2,fun(X2,bool)),X3)))))) )),
  inference(unit_resulting_resolution,[],[f1700,f1191,f1191,f1579])).
fof(f3991,plain,(
  ( ! [X2,X0,X1] : (~sP9('F_0',X0,'const.relation.symmetric_1'(s(fun(X1,fun(X1,bool)),'const.relation.SC_1'(s(fun(X1,fun(X1,bool)),X2)))))) )),
  inference(unit_resulting_resolution,[],[f1700,f1191,f1580])).
fof(f3993,plain,(
  ( ! [X2,X0,X1] : (~sP9('F_0','const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))),X2)) )),
  inference(unit_resulting_resolution,[],[f1700,f1191,f1581])).
fof(f3997,plain,(
  ( ! [X0,X1] : (~sP8('const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f1700,f1700,f1191,f1587])).
fof(f4000,plain,(
  ( ! [X2,X0,X3,X1] : (sP8(X0,X1,'const.relation.symmetric_1'(s(fun(X2,fun(X2,bool)),'const.relation.SC_1'(s(fun(X2,fun(X2,bool)),X3)))))) )),
  inference(unit_resulting_resolution,[],[f1191,f1588])).
fof(f4001,plain,(
  ( ! [X2,X0,X3,X1] : (sP8(X0,'const.relation.symmetric_1'(s(fun(X1,fun(X1,bool)),'const.relation.SC_1'(s(fun(X1,fun(X1,bool)),X2)))),X3)) )),
  inference(unit_resulting_resolution,[],[f1191,f1589])).
fof(f4002,plain,(
  ( ! [X2,X0,X1] : (sP9('T_0',X0,'const.relation.symmetric_1'(s(fun(X1,fun(X1,bool)),'const.relation.SC_1'(s(fun(X1,fun(X1,bool)),X2)))))) )),
  inference(unit_resulting_resolution,[],[f1104,f1191,f1592])).
fof(f4007,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP9('const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))),X2,'const.relation.symmetric_1'(s(fun(X3,fun(X3,bool)),'const.relation.SC_1'(s(fun(X3,fun(X3,bool)),X4)))))) )),
  inference(unit_resulting_resolution,[],[f1191,f1191,f1592])).
fof(f4008,plain,(
  ( ! [X2,X0,X1] : (sP9('T_0','const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))),X2)) )),
  inference(unit_resulting_resolution,[],[f1104,f1191,f1593])).
fof(f4013,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP9('const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))),'const.relation.symmetric_1'(s(fun(X2,fun(X2,bool)),'const.relation.SC_1'(s(fun(X2,fun(X2,bool)),X3)))),X4)) )),
  inference(unit_resulting_resolution,[],[f1191,f1191,f1593])).
fof(f4014,plain,(
  ( ! [X0,X1] : (~sP9('const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f1700,f1700,f1191,f1594])).
fof(f4015,plain,(
  ( ! [X2,X0,X1] : (~sP11('F_0','const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))),X2)) )),
  inference(unit_resulting_resolution,[],[f1700,f1191,f1598])).
fof(f4031,plain,(
  ( ! [X0,X1] : (~sP10('const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f1700,f1104,f1191,f1604])).
fof(f4036,plain,(
  ( ! [X0,X1] : (~sP10('T_0','F_0','const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))))) )),
  inference(unit_resulting_resolution,[],[f1104,f1700,f1191,f1604])).
fof(f4037,plain,(
  ( ! [X2,X0,X3,X1] : (~sP10('const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))),'F_0','const.relation.symmetric_1'(s(fun(X2,fun(X2,bool)),'const.relation.SC_1'(s(fun(X2,fun(X2,bool)),X3)))))) )),
  inference(unit_resulting_resolution,[],[f1191,f1700,f1191,f1604])).
fof(f4041,plain,(
  ( ! [X2,X0,X3,X1] : (sP10(X0,'const.relation.symmetric_1'(s(fun(X1,fun(X1,bool)),'const.relation.SC_1'(s(fun(X1,fun(X1,bool)),X2)))),X3)) )),
  inference(unit_resulting_resolution,[],[f1191,f1606])).
fof(f4042,plain,(
  ( ! [X0,X1] : (sP11('F_0','F_0','const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))))) )),
  inference(unit_resulting_resolution,[],[f1700,f1700,f1191,f1608])).
fof(f4045,plain,(
  ( ! [X2,X0,X1] : (sP11('T_0','const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))),X2)) )),
  inference(unit_resulting_resolution,[],[f1104,f1191,f1610])).
fof(f4051,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP11('const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))),'const.relation.symmetric_1'(s(fun(X2,fun(X2,bool)),'const.relation.SC_1'(s(fun(X2,fun(X2,bool)),X3)))),X4)) )),
  inference(unit_resulting_resolution,[],[f1191,f1191,f1610])).
fof(f4052,plain,(
  ( ! [X0,X1] : (~sP11('T_0','F_0','const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))))) )),
  inference(unit_resulting_resolution,[],[f1700,f1104,f1191,f1611])).
fof(f4055,plain,(
  ( ! [X0,X1] : (~sP11('const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f1700,f1104,f1191,f1611])).
fof(f4056,plain,(
  ( ! [X2,X0,X3,X1] : (~sP11('const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))),'F_0','const.relation.symmetric_1'(s(fun(X2,fun(X2,bool)),'const.relation.SC_1'(s(fun(X2,fun(X2,bool)),X3)))))) )),
  inference(unit_resulting_resolution,[],[f1700,f1191,f1191,f1611])).
fof(f4058,plain,(
  ( ! [X2,X0,X1] : (sP11('const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))),X2,'F_0')) )),
  inference(unit_resulting_resolution,[],[f1929,f1191,f1676])).
fof(f4059,plain,(
  ( ! [X2,X0,X1] : (sP11('const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))),'T_0',X2)) )),
  inference(unit_resulting_resolution,[],[f1930,f1191,f1676])).
fof(f4060,plain,(
  ( ! [X2,X0,X1] : (sP9('const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))),X2,'T_0')) )),
  inference(unit_resulting_resolution,[],[f1924,f1191,f1677])).
fof(f4061,plain,(
  ( ! [X2,X0,X1] : (sP9('const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))),'T_0',X2)) )),
  inference(unit_resulting_resolution,[],[f1925,f1191,f1677])).
fof(f4062,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'|dollar|or_2'(s(bool,X0),s(bool,'const.relation.symmetric_1'(s(fun(X1,fun(X1,bool)),'const.relation.SC_1'(s(fun(X1,fun(X1,bool)),X2))))))))) )),
  inference(unit_resulting_resolution,[],[f1191,f1835])).
fof(f4063,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'|dollar|or_2'(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1))))),s(bool,X2))))) )),
  inference(unit_resulting_resolution,[],[f1191,f1836])).
fof(f4064,plain,(
  ( ! [X2,X0,X1] : (sP0('const.combin.I_1'(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))))),X2)) )),
  inference(unit_resulting_resolution,[],[f1191,f1934])).
fof(f4065,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,'const.combin.I_1'(s(bool,'const.relation.symmetric_1'(s(fun(X1,fun(X1,bool)),'const.relation.SC_1'(s(fun(X1,fun(X1,bool)),X2)))))))) )),
  inference(unit_resulting_resolution,[],[f1191,f1935])).
fof(f3573,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)))) )),
  inference(unit_resulting_resolution,[],[f1104,f1256,f1364])).
fof(f3574,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2))) = s(bool,'const.relation.RTC_3'(s(fun(X3,fun(X3,bool)),X4),s(X3,X5),s(X3,X5)))) )),
  inference(unit_resulting_resolution,[],[f1256,f1256,f1364])).
fof(f3575,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)))) )),
  inference(unit_resulting_resolution,[],[f1700,f1256,f1370])).
fof(f3581,plain,(
  ( ! [X2,X0,X3,X1] : (sP0(X0,'const.relation.RTC_3'(s(fun(X1,fun(X1,bool)),X2),s(X1,X3),s(X1,X3)))) )),
  inference(unit_resulting_resolution,[],[f1256,f1397])).
fof(f3582,plain,(
  ( ! [X2,X0,X3,X1] : (sP0('const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)),X3)) )),
  inference(unit_resulting_resolution,[],[f1256,f1398])).
fof(f3583,plain,(
  ( ! [X2,X0,X1] : (~sP1('const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f1104,f1256,f1412])).
fof(f3585,plain,(
  ( ! [X2,X0,X1] : (~sP1('T_0','const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)))) )),
  inference(unit_resulting_resolution,[],[f1104,f1256,f1412])).
fof(f3586,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~sP1('const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)),'const.relation.RTC_3'(s(fun(X3,fun(X3,bool)),X4),s(X3,X5),s(X3,X5)))) )),
  inference(unit_resulting_resolution,[],[f1256,f1256,f1412])).
fof(f3587,plain,(
  ( ! [X2,X0,X1] : (sP1('const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1700,f1256,f1413])).
fof(f3588,plain,(
  ( ! [X2,X0,X1] : (sP1('F_0','const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)))) )),
  inference(unit_resulting_resolution,[],[f1700,f1256,f1416])).
fof(f3591,plain,(
  ( ! [X2,X0,X1] : (~sP5('T_0','const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1104,f1700,f1256,f1543])).
fof(f3592,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~sP5('const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)),'const.relation.RTC_3'(s(fun(X3,fun(X3,bool)),X4),s(X3,X5),s(X3,X5)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1256,f1700,f1256,f1543])).
fof(f3594,plain,(
  ( ! [X2,X0,X1] : (~sP4('const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f1104,f1700,f1256,f1547])).
fof(f3600,plain,(
  ( ! [X2,X0,X1] : (~sP4('T_0','const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1104,f1700,f1256,f1547])).
fof(f3601,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~sP4('const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)),'const.relation.RTC_3'(s(fun(X3,fun(X3,bool)),X4),s(X3,X5),s(X3,X5)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1256,f1700,f1256,f1547])).
fof(f3607,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP4(X0,X1,'const.relation.RTC_3'(s(fun(X2,fun(X2,bool)),X3),s(X2,X4),s(X2,X4)))) )),
  inference(unit_resulting_resolution,[],[f1256,f1548])).
fof(f3608,plain,(
  ( ! [X2,X0,X1] : (~sP3('const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f1104,f1700,f1256,f1551])).
fof(f3614,plain,(
  ( ! [X2,X0,X1] : (~sP3('T_0','const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1104,f1700,f1256,f1551])).
fof(f3615,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~sP3('const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)),'const.relation.RTC_3'(s(fun(X3,fun(X3,bool)),X4),s(X3,X5),s(X3,X5)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1256,f1700,f1256,f1551])).
fof(f3620,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP3(X0,X1,'const.relation.RTC_3'(s(fun(X2,fun(X2,bool)),X3),s(X2,X4),s(X2,X4)))) )),
  inference(unit_resulting_resolution,[],[f1256,f1552])).
fof(f3621,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP2(X0,X1,'const.relation.RTC_3'(s(fun(X2,fun(X2,bool)),X3),s(X2,X4),s(X2,X4)))) )),
  inference(unit_resulting_resolution,[],[f1256,f1556])).
fof(f3622,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP2(X0,'const.relation.RTC_3'(s(fun(X1,fun(X1,bool)),X2),s(X1,X3),s(X1,X3)),X4)) )),
  inference(unit_resulting_resolution,[],[f1256,f1557])).
fof(f3623,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP2('const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)),X3,X4)) )),
  inference(unit_resulting_resolution,[],[f1256,f1558])).
fof(f3625,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP5('const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)),'const.relation.EMPTY_REL_2'(s(X3,X4),s(X3,X5)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1253,f1256,f1560])).
fof(f3627,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (sP5('const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)),'const.relation.RUNIV_2'(s(X3,X4),s(X5,X6)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f1447,f1256,f1560])).
fof(f3628,plain,(
  ( ! [X2,X0,X3,X1] : (sP5('const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)),'const.combin.I_1'(s(bool,X3)),X3)) )),
  inference(unit_resulting_resolution,[],[f1188,f1256,f1560])).
fof(f3629,plain,(
  ( ! [X2,X0,X3,X1] : (sP5('const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)),happ(s(fun(bool,bool),'_i_0_0'),s(bool,X3)),X3)) )),
  inference(unit_resulting_resolution,[],[f1189,f1256,f1560])).
fof(f3630,plain,(
  ( ! [X2,X0,X3,X1] : (sP5('const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)),happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X3)),X3)) )),
  inference(unit_resulting_resolution,[],[f1701,f1256,f1560])).
fof(f3631,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP5('const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,X4)),X4)) )),
  inference(unit_resulting_resolution,[],[f1254,f1256,f1560])).
fof(f3632,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP5('const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X3),s(bool,X4)),X3)) )),
  inference(unit_resulting_resolution,[],[f1255,f1256,f1560])).
fof(f3634,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP5('const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)),'F_0','const.relation.EMPTY_REL_2'(s(X3,X4),s(X3,X5)))) )),
  inference(unit_resulting_resolution,[],[f1253,f1256,f1560])).
fof(f3636,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (sP5('const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)),'T_0','const.relation.RUNIV_2'(s(X3,X4),s(X5,X6)))) )),
  inference(unit_resulting_resolution,[],[f1447,f1256,f1560])).
fof(f3637,plain,(
  ( ! [X2,X0,X3,X1] : (sP5('const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)),X3,'const.combin.I_1'(s(bool,X3)))) )),
  inference(unit_resulting_resolution,[],[f1188,f1256,f1560])).
fof(f3638,plain,(
  ( ! [X2,X0,X3,X1] : (sP5('const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)),X3,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X3)))) )),
  inference(unit_resulting_resolution,[],[f1189,f1256,f1560])).
fof(f3639,plain,(
  ( ! [X2,X0,X3,X1] : (sP5('const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)),X3,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X3)))) )),
  inference(unit_resulting_resolution,[],[f1701,f1256,f1560])).
fof(f3640,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP5('const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)),X3,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X4),s(bool,X3)))) )),
  inference(unit_resulting_resolution,[],[f1254,f1256,f1560])).
fof(f3641,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP5('const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)),X3,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X3),s(bool,X4)))) )),
  inference(unit_resulting_resolution,[],[f1255,f1256,f1560])).
fof(f3642,plain,(
  ( ! [X2,X0,X1] : (~sP5('const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f1115,f1256,f1561])).
fof(f3643,plain,(
  ( ! [X2,X0,X1] : (~sP5('const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f1115,f1256,f1561])).
fof(f3644,plain,(
  ( ! [X2,X0,X3,X1] : (~sP7('const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)),X3,'F_0')) )),
  inference(unit_resulting_resolution,[],[f1700,f1256,f1564])).
fof(f3645,plain,(
  ( ! [X2,X0,X3,X1] : (~sP7('const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)),'F_0',X3)) )),
  inference(unit_resulting_resolution,[],[f1700,f1256,f1565])).
fof(f3656,plain,(
  ( ! [X2,X0,X1] : (~sP6('const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f1104,f1700,f1256,f1570])).
fof(f3662,plain,(
  ( ! [X2,X0,X1] : (~sP6('T_0','const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1104,f1700,f1256,f1570])).
fof(f3663,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~sP6('const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)),'const.relation.RTC_3'(s(fun(X3,fun(X3,bool)),X4),s(X3,X5),s(X3,X5)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1256,f1700,f1256,f1570])).
fof(f3668,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP6(X0,X1,'const.relation.RTC_3'(s(fun(X2,fun(X2,bool)),X3),s(X2,X4),s(X2,X4)))) )),
  inference(unit_resulting_resolution,[],[f1256,f1571])).
fof(f3669,plain,(
  ( ! [X2,X0,X1] : (sP7('T_0','const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f1104,f1104,f1256,f1576])).
fof(f3675,plain,(
  ( ! [X2,X0,X1] : (sP7('T_0','T_0','const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)))) )),
  inference(unit_resulting_resolution,[],[f1104,f1104,f1256,f1576])).
fof(f3677,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP7('T_0','const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)),'const.relation.RTC_3'(s(fun(X3,fun(X3,bool)),X4),s(X3,X5),s(X3,X5)))) )),
  inference(unit_resulting_resolution,[],[f1104,f1256,f1256,f1576])).
fof(f3681,plain,(
  ( ! [X2,X0,X1] : (sP7('const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)),'T_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f1104,f1104,f1256,f1576])).
fof(f3682,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP7('const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)),'const.relation.RTC_3'(s(fun(X3,fun(X3,bool)),X4),s(X3,X5),s(X3,X5)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f1256,f1104,f1256,f1576])).
fof(f3683,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP7('const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)),'T_0','const.relation.RTC_3'(s(fun(X3,fun(X3,bool)),X4),s(X3,X5),s(X3,X5)))) )),
  inference(unit_resulting_resolution,[],[f1104,f1256,f1256,f1576])).
fof(f3684,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (sP7('const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)),'const.relation.RTC_3'(s(fun(X3,fun(X3,bool)),X4),s(X3,X5),s(X3,X5)),'const.relation.RTC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X8),s(X6,X8)))) )),
  inference(unit_resulting_resolution,[],[f1256,f1256,f1256,f1576])).
fof(f3685,plain,(
  ( ! [X2,X0,X1] : (~sP7('F_0','const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f1700,f1104,f1256,f1579])).
fof(f3690,plain,(
  ( ! [X2,X0,X1] : (~sP7('F_0','T_0','const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)))) )),
  inference(unit_resulting_resolution,[],[f1700,f1104,f1256,f1579])).
fof(f3691,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~sP7('F_0','const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)),'const.relation.RTC_3'(s(fun(X3,fun(X3,bool)),X4),s(X3,X5),s(X3,X5)))) )),
  inference(unit_resulting_resolution,[],[f1700,f1256,f1256,f1579])).
fof(f3695,plain,(
  ( ! [X2,X0,X3,X1] : (~sP9('F_0',X0,'const.relation.RTC_3'(s(fun(X1,fun(X1,bool)),X2),s(X1,X3),s(X1,X3)))) )),
  inference(unit_resulting_resolution,[],[f1700,f1256,f1580])).
fof(f3697,plain,(
  ( ! [X2,X0,X3,X1] : (~sP9('F_0','const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)),X3)) )),
  inference(unit_resulting_resolution,[],[f1700,f1256,f1581])).
fof(f3701,plain,(
  ( ! [X2,X0,X1] : (~sP8('const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f1700,f1700,f1256,f1587])).
fof(f3704,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP8(X0,X1,'const.relation.RTC_3'(s(fun(X2,fun(X2,bool)),X3),s(X2,X4),s(X2,X4)))) )),
  inference(unit_resulting_resolution,[],[f1256,f1588])).
fof(f3705,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP8(X0,'const.relation.RTC_3'(s(fun(X1,fun(X1,bool)),X2),s(X1,X3),s(X1,X3)),X4)) )),
  inference(unit_resulting_resolution,[],[f1256,f1589])).
fof(f3706,plain,(
  ( ! [X2,X0,X3,X1] : (sP9('T_0',X0,'const.relation.RTC_3'(s(fun(X1,fun(X1,bool)),X2),s(X1,X3),s(X1,X3)))) )),
  inference(unit_resulting_resolution,[],[f1104,f1256,f1592])).
fof(f3711,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (sP9('const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)),X3,'const.relation.RTC_3'(s(fun(X4,fun(X4,bool)),X5),s(X4,X6),s(X4,X6)))) )),
  inference(unit_resulting_resolution,[],[f1256,f1256,f1592])).
fof(f3712,plain,(
  ( ! [X2,X0,X3,X1] : (sP9('T_0','const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)),X3)) )),
  inference(unit_resulting_resolution,[],[f1104,f1256,f1593])).
fof(f3717,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (sP9('const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)),'const.relation.RTC_3'(s(fun(X3,fun(X3,bool)),X4),s(X3,X5),s(X3,X5)),X6)) )),
  inference(unit_resulting_resolution,[],[f1256,f1256,f1593])).
fof(f3718,plain,(
  ( ! [X2,X0,X1] : (~sP9('const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f1700,f1700,f1256,f1594])).
fof(f3719,plain,(
  ( ! [X2,X0,X3,X1] : (~sP11('F_0','const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)),X3)) )),
  inference(unit_resulting_resolution,[],[f1700,f1256,f1598])).
fof(f3734,plain,(
  ( ! [X2,X0,X1] : (~sP10('const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f1700,f1104,f1256,f1604])).
fof(f3739,plain,(
  ( ! [X2,X0,X1] : (~sP10('T_0','F_0','const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)))) )),
  inference(unit_resulting_resolution,[],[f1104,f1700,f1256,f1604])).
fof(f3740,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~sP10('const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)),'F_0','const.relation.RTC_3'(s(fun(X3,fun(X3,bool)),X4),s(X3,X5),s(X3,X5)))) )),
  inference(unit_resulting_resolution,[],[f1256,f1700,f1256,f1604])).
fof(f3744,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10(X0,'const.relation.RTC_3'(s(fun(X1,fun(X1,bool)),X2),s(X1,X3),s(X1,X3)),X4)) )),
  inference(unit_resulting_resolution,[],[f1256,f1606])).
fof(f3745,plain,(
  ( ! [X2,X0,X1] : (sP11('F_0','F_0','const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)))) )),
  inference(unit_resulting_resolution,[],[f1700,f1700,f1256,f1608])).
fof(f3748,plain,(
  ( ! [X2,X0,X3,X1] : (sP11('T_0','const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)),X3)) )),
  inference(unit_resulting_resolution,[],[f1104,f1256,f1610])).
fof(f3753,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (sP11('const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)),'const.relation.RTC_3'(s(fun(X3,fun(X3,bool)),X4),s(X3,X5),s(X3,X5)),X6)) )),
  inference(unit_resulting_resolution,[],[f1256,f1256,f1610])).
fof(f3754,plain,(
  ( ! [X2,X0,X1] : (~sP11('T_0','F_0','const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)))) )),
  inference(unit_resulting_resolution,[],[f1700,f1104,f1256,f1611])).
fof(f3757,plain,(
  ( ! [X2,X0,X1] : (~sP11('const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f1700,f1104,f1256,f1611])).
fof(f3758,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~sP11('const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)),'F_0','const.relation.RTC_3'(s(fun(X3,fun(X3,bool)),X4),s(X3,X5),s(X3,X5)))) )),
  inference(unit_resulting_resolution,[],[f1700,f1256,f1256,f1611])).
fof(f3760,plain,(
  ( ! [X2,X0,X3,X1] : (sP11('const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)),X3,'F_0')) )),
  inference(unit_resulting_resolution,[],[f1929,f1256,f1676])).
fof(f3761,plain,(
  ( ! [X2,X0,X3,X1] : (sP11('const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)),'T_0',X3)) )),
  inference(unit_resulting_resolution,[],[f1930,f1256,f1676])).
fof(f3762,plain,(
  ( ! [X2,X0,X3,X1] : (sP9('const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)),X3,'T_0')) )),
  inference(unit_resulting_resolution,[],[f1924,f1256,f1677])).
fof(f3763,plain,(
  ( ! [X2,X0,X3,X1] : (sP9('const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)),'T_0',X3)) )),
  inference(unit_resulting_resolution,[],[f1925,f1256,f1677])).
fof(f3764,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'|dollar|or_2'(s(bool,X0),s(bool,'const.relation.RTC_3'(s(fun(X1,fun(X1,bool)),X2),s(X1,X3),s(X1,X3))))))) )),
  inference(unit_resulting_resolution,[],[f1256,f1835])).
fof(f3765,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'|dollar|or_2'(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2))),s(bool,X3))))) )),
  inference(unit_resulting_resolution,[],[f1256,f1836])).
fof(f3766,plain,(
  ( ! [X2,X0,X3,X1] : (sP0('const.combin.I_1'(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)))),X3)) )),
  inference(unit_resulting_resolution,[],[f1256,f1934])).
fof(f3767,plain,(
  ( ! [X2,X0,X3,X1] : (sP0(X0,'const.combin.I_1'(s(bool,'const.relation.RTC_3'(s(fun(X1,fun(X1,bool)),X2),s(X1,X3),s(X1,X3)))))) )),
  inference(unit_resulting_resolution,[],[f1256,f1935])).
fof(f3546,plain,(
  ( ! [X2,X0,X1] : (sP9(sK63(X0,X1,'T_0',X2),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f1700,f1700,f2505,f1586])).
fof(f3504,plain,(
  ~sP11('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f3496,f1599])).
fof(f3496,plain,(
  ~sP10('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f1104,f1104,f1700,f1604])).
fof(f3488,plain,(
  sP11('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f1932,f1104,f1700,f1603])).
fof(f3423,plain,(
  ~sP9('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f3416,f1582])).
fof(f3416,plain,(
  ~sP8('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f1104,f1700,f1700,f1587])).
fof(f3409,plain,(
  sP9('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f1927,f1700,f1700,f1586])).
fof(f3369,plain,(
  ~sP7('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f3361,f1563])).
fof(f3361,plain,(
  ~sP6('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f1104,f1700,f1104,f1570])).
fof(f3223,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,sK64(X0,X1,'T_0',X2))) )),
  inference(unit_resulting_resolution,[],[f1104,f2570,f1364])).
fof(f3224,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,sK64(X0,X1,'T_0',X2)) = s(bool,sK64(X3,X4,'T_0',X5))) )),
  inference(unit_resulting_resolution,[],[f2570,f2570,f1364])).
fof(f3225,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,sK64(X0,X1,'T_0',X2))) )),
  inference(unit_resulting_resolution,[],[f1700,f2570,f1370])).
fof(f3231,plain,(
  ( ! [X2,X0,X3,X1] : (sP0(X0,sK64(X1,X2,'T_0',X3))) )),
  inference(unit_resulting_resolution,[],[f2570,f1397])).
fof(f3232,plain,(
  ( ! [X2,X0,X3,X1] : (sP0(sK64(X0,X1,'T_0',X2),X3)) )),
  inference(unit_resulting_resolution,[],[f2570,f1398])).
fof(f3233,plain,(
  ( ! [X2,X0,X1] : (~sP1(sK64(X0,X1,'T_0',X2),'T_0')) )),
  inference(unit_resulting_resolution,[],[f1104,f2570,f1412])).
fof(f3235,plain,(
  ( ! [X2,X0,X1] : (~sP1('T_0',sK64(X0,X1,'T_0',X2))) )),
  inference(unit_resulting_resolution,[],[f1104,f2570,f1412])).
fof(f3236,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~sP1(sK64(X0,X1,'T_0',X2),sK64(X3,X4,'T_0',X5))) )),
  inference(unit_resulting_resolution,[],[f2570,f2570,f1412])).
fof(f3237,plain,(
  ( ! [X2,X0,X1] : (sP1(sK64(X0,X1,'T_0',X2),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1700,f2570,f1413])).
fof(f3238,plain,(
  ( ! [X2,X0,X1] : (sP1('F_0',sK64(X0,X1,'T_0',X2))) )),
  inference(unit_resulting_resolution,[],[f1700,f2570,f1416])).
fof(f3241,plain,(
  ( ! [X2,X0,X1] : (~sP5('T_0',sK64(X0,X1,'T_0',X2),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1104,f1700,f2570,f1543])).
fof(f3242,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~sP5(sK64(X0,X1,'T_0',X2),sK64(X3,X4,'T_0',X5),'F_0')) )),
  inference(unit_resulting_resolution,[],[f2570,f1700,f2570,f1543])).
fof(f3246,plain,(
  ( ! [X2,X0,X1] : (~sP4(sK64(X0,X1,'T_0',X2),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f1700,f1104,f2570,f1547])).
fof(f3250,plain,(
  ( ! [X2,X0,X1] : (~sP4('T_0',sK64(X0,X1,'T_0',X2),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1104,f1700,f2570,f1547])).
fof(f3251,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~sP4(sK64(X0,X1,'T_0',X2),sK64(X3,X4,'T_0',X5),'F_0')) )),
  inference(unit_resulting_resolution,[],[f2570,f1700,f2570,f1547])).
fof(f3256,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP4(X0,X1,sK64(X2,X3,'T_0',X4))) )),
  inference(unit_resulting_resolution,[],[f2570,f1548])).
fof(f3259,plain,(
  ( ! [X2,X0,X1] : (~sP3(sK64(X0,X1,'T_0',X2),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f1700,f1104,f2570,f1551])).
fof(f3263,plain,(
  ( ! [X2,X0,X1] : (~sP3('T_0',sK64(X0,X1,'T_0',X2),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1104,f1700,f2570,f1551])).
fof(f3264,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~sP3(sK64(X0,X1,'T_0',X2),sK64(X3,X4,'T_0',X5),'F_0')) )),
  inference(unit_resulting_resolution,[],[f2570,f1700,f2570,f1551])).
fof(f3268,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP3(X0,X1,sK64(X2,X3,'T_0',X4))) )),
  inference(unit_resulting_resolution,[],[f2570,f1552])).
fof(f3269,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP2(X0,X1,sK64(X2,X3,'T_0',X4))) )),
  inference(unit_resulting_resolution,[],[f2570,f1556])).
fof(f3270,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP2(X0,sK64(X1,X2,'T_0',X3),X4)) )),
  inference(unit_resulting_resolution,[],[f2570,f1557])).
fof(f3271,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP2(sK64(X0,X1,'T_0',X2),X3,X4)) )),
  inference(unit_resulting_resolution,[],[f2570,f1558])).
fof(f3273,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP5(sK64(X0,X1,'T_0',X2),'const.relation.EMPTY_REL_2'(s(X3,X4),s(X3,X5)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1253,f2570,f1560])).
fof(f3275,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (sP5(sK64(X0,X1,'T_0',X2),'const.relation.RUNIV_2'(s(X3,X4),s(X5,X6)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f1447,f2570,f1560])).
fof(f3276,plain,(
  ( ! [X2,X0,X3,X1] : (sP5(sK64(X0,X1,'T_0',X2),'const.combin.I_1'(s(bool,X3)),X3)) )),
  inference(unit_resulting_resolution,[],[f1188,f2570,f1560])).
fof(f3277,plain,(
  ( ! [X2,X0,X3,X1] : (sP5(sK64(X0,X1,'T_0',X2),happ(s(fun(bool,bool),'_i_0_0'),s(bool,X3)),X3)) )),
  inference(unit_resulting_resolution,[],[f1189,f2570,f1560])).
fof(f3278,plain,(
  ( ! [X2,X0,X3,X1] : (sP5(sK64(X0,X1,'T_0',X2),happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X3)),X3)) )),
  inference(unit_resulting_resolution,[],[f1701,f2570,f1560])).
fof(f3279,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP5(sK64(X0,X1,'T_0',X2),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,X4)),X4)) )),
  inference(unit_resulting_resolution,[],[f1254,f2570,f1560])).
fof(f3280,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP5(sK64(X0,X1,'T_0',X2),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X3),s(bool,X4)),X3)) )),
  inference(unit_resulting_resolution,[],[f1255,f2570,f1560])).
fof(f3282,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP5(sK64(X0,X1,'T_0',X2),'F_0','const.relation.EMPTY_REL_2'(s(X3,X4),s(X3,X5)))) )),
  inference(unit_resulting_resolution,[],[f1253,f2570,f1560])).
fof(f3284,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (sP5(sK64(X0,X1,'T_0',X2),'T_0','const.relation.RUNIV_2'(s(X3,X4),s(X5,X6)))) )),
  inference(unit_resulting_resolution,[],[f1447,f2570,f1560])).
fof(f3285,plain,(
  ( ! [X2,X0,X3,X1] : (sP5(sK64(X0,X1,'T_0',X2),X3,'const.combin.I_1'(s(bool,X3)))) )),
  inference(unit_resulting_resolution,[],[f1188,f2570,f1560])).
fof(f3286,plain,(
  ( ! [X2,X0,X3,X1] : (sP5(sK64(X0,X1,'T_0',X2),X3,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X3)))) )),
  inference(unit_resulting_resolution,[],[f1189,f2570,f1560])).
fof(f3287,plain,(
  ( ! [X2,X0,X3,X1] : (sP5(sK64(X0,X1,'T_0',X2),X3,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X3)))) )),
  inference(unit_resulting_resolution,[],[f1701,f2570,f1560])).
fof(f3288,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP5(sK64(X0,X1,'T_0',X2),X3,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X4),s(bool,X3)))) )),
  inference(unit_resulting_resolution,[],[f1254,f2570,f1560])).
fof(f3289,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP5(sK64(X0,X1,'T_0',X2),X3,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X3),s(bool,X4)))) )),
  inference(unit_resulting_resolution,[],[f1255,f2570,f1560])).
fof(f3290,plain,(
  ( ! [X2,X0,X1] : (~sP5(sK64(X0,X1,'T_0',X2),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f1115,f2570,f1561])).
fof(f3291,plain,(
  ( ! [X2,X0,X1] : (~sP5(sK64(X0,X1,'T_0',X2),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f1115,f2570,f1561])).
fof(f3292,plain,(
  ( ! [X2,X0,X3,X1] : (~sP7(sK64(X0,X1,'T_0',X2),X3,'F_0')) )),
  inference(unit_resulting_resolution,[],[f1700,f2570,f1564])).
fof(f3293,plain,(
  ( ! [X2,X0,X3,X1] : (~sP7(sK64(X0,X1,'T_0',X2),'F_0',X3)) )),
  inference(unit_resulting_resolution,[],[f1700,f2570,f1565])).
fof(f3294,plain,(
  ( ! [X2,X0,X1] : (sP7('T_0',sK64(X0,X1,'T_0',X2),'T_0')) )),
  inference(unit_resulting_resolution,[],[f1919,f1104,f2570,f1566])).
fof(f3298,plain,(
  ( ! [X2,X0,X1] : (sP7('T_0','T_0',sK64(X0,X1,'T_0',X2))) )),
  inference(unit_resulting_resolution,[],[f1919,f1104,f2570,f1566])).
fof(f3299,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP7('T_0',sK64(X0,X1,'T_0',X2),sK64(X3,X4,'T_0',X5))) )),
  inference(unit_resulting_resolution,[],[f1919,f2570,f2570,f1566])).
fof(f3300,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP6(X0,X1,sK64(X2,X3,'T_0',X4))) )),
  inference(unit_resulting_resolution,[],[f2570,f1571])).
fof(f3302,plain,(
  ( ! [X2,X0,X3,X1] : (~sP9('F_0',X0,sK64(X1,X2,'T_0',X3))) )),
  inference(unit_resulting_resolution,[],[f1700,f2570,f1580])).
fof(f3304,plain,(
  ( ! [X2,X0,X3,X1] : (~sP9('F_0',sK64(X0,X1,'T_0',X2),X3)) )),
  inference(unit_resulting_resolution,[],[f1700,f2570,f1581])).
fof(f3305,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP8(X0,X1,sK64(X2,X3,'T_0',X4))) )),
  inference(unit_resulting_resolution,[],[f2570,f1588])).
fof(f3306,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP8(X0,sK64(X1,X2,'T_0',X3),X4)) )),
  inference(unit_resulting_resolution,[],[f2570,f1589])).
fof(f3307,plain,(
  ( ! [X2,X0,X3,X1] : (sP9('T_0',X0,sK64(X1,X2,'T_0',X3))) )),
  inference(unit_resulting_resolution,[],[f1104,f2570,f1592])).
fof(f3311,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (sP9(sK64(X0,X1,'T_0',X2),X3,sK64(X4,X5,'T_0',X6))) )),
  inference(unit_resulting_resolution,[],[f2570,f2570,f1592])).
fof(f3313,plain,(
  ( ! [X2,X0,X3,X1] : (sP9('T_0',sK64(X0,X1,'T_0',X2),X3)) )),
  inference(unit_resulting_resolution,[],[f1104,f2570,f1593])).
fof(f3316,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (sP9(sK64(X0,X1,'T_0',X2),sK64(X3,X4,'T_0',X5),X6)) )),
  inference(unit_resulting_resolution,[],[f2570,f2570,f1593])).
fof(f3319,plain,(
  ( ! [X2,X0,X3,X1] : (~sP11('F_0',sK64(X0,X1,'T_0',X2),X3)) )),
  inference(unit_resulting_resolution,[],[f1700,f2570,f1598])).
fof(f3320,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10(X0,sK64(X1,X2,'T_0',X3),X4)) )),
  inference(unit_resulting_resolution,[],[f2570,f1606])).
fof(f3323,plain,(
  ( ! [X2,X0,X3,X1] : (sP11('T_0',sK64(X0,X1,'T_0',X2),X3)) )),
  inference(unit_resulting_resolution,[],[f1104,f2570,f1610])).
fof(f3326,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (sP11(sK64(X0,X1,'T_0',X2),sK64(X3,X4,'T_0',X5),X6)) )),
  inference(unit_resulting_resolution,[],[f2570,f2570,f1610])).
fof(f3328,plain,(
  ( ! [X2,X0,X3,X1] : (sP11(sK64(X0,X1,'T_0',X2),X3,'F_0')) )),
  inference(unit_resulting_resolution,[],[f1929,f2570,f1676])).
fof(f3329,plain,(
  ( ! [X2,X0,X3,X1] : (sP11(sK64(X0,X1,'T_0',X2),'T_0',X3)) )),
  inference(unit_resulting_resolution,[],[f1930,f2570,f1676])).
fof(f3330,plain,(
  ( ! [X2,X0,X3,X1] : (sP9(sK64(X0,X1,'T_0',X2),X3,'T_0')) )),
  inference(unit_resulting_resolution,[],[f1924,f2570,f1677])).
fof(f3331,plain,(
  ( ! [X2,X0,X3,X1] : (sP9(sK64(X0,X1,'T_0',X2),'T_0',X3)) )),
  inference(unit_resulting_resolution,[],[f1925,f2570,f1677])).
fof(f3332,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'|dollar|or_2'(s(bool,X0),s(bool,sK64(X1,X2,'T_0',X3)))))) )),
  inference(unit_resulting_resolution,[],[f2570,f1835])).
fof(f3333,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'|dollar|or_2'(s(bool,sK64(X0,X1,'T_0',X2)),s(bool,X3))))) )),
  inference(unit_resulting_resolution,[],[f2570,f1836])).
fof(f3334,plain,(
  ( ! [X2,X0,X3,X1] : (sP0('const.combin.I_1'(s(bool,sK64(X0,X1,'T_0',X2))),X3)) )),
  inference(unit_resulting_resolution,[],[f2570,f1934])).
fof(f3335,plain,(
  ( ! [X2,X0,X3,X1] : (sP0(X0,'const.combin.I_1'(s(bool,sK64(X1,X2,'T_0',X3))))) )),
  inference(unit_resulting_resolution,[],[f2570,f1935])).
fof(f3214,plain,(
  sP7('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f1919,f1104,f1104,f1566])).
fof(f3210,plain,(
  ~sP5('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f3205,f1540])).
fof(f3205,plain,(
  ~sP2('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f1700,f1700,f1700,f1555])).
fof(f3201,plain,(
  ~sP5('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f3193,f1541])).
fof(f3193,plain,(
  ~sP3('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f1104,f1700,f1104,f1551])).
fof(f3177,plain,(
  ~sP4('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f1104,f1700,f1104,f1547])).
fof(f3148,plain,(
  ( ! [X15,X16] : (sP4('T_0','const.combin.I_1'(s(bool,X15)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X15),s(bool,X16)))) )),
  inference(superposition,[],[f2889,f1255])).
fof(f3147,plain,(
  ( ! [X14,X13] : (sP4('T_0','const.combin.I_1'(s(bool,X14)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X13),s(bool,X14)))) )),
  inference(superposition,[],[f2889,f1254])).
fof(f3146,plain,(
  ( ! [X12] : (sP4('T_0','const.combin.I_1'(s(bool,X12)),happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X12)))) )),
  inference(superposition,[],[f2889,f1701])).
fof(f3145,plain,(
  ( ! [X11] : (sP4('T_0','const.combin.I_1'(s(bool,X11)),happ(s(fun(bool,bool),'_i_0_0'),s(bool,X11)))) )),
  inference(superposition,[],[f2889,f1189])).
fof(f3142,plain,(
  ( ! [X8,X7,X9] : (sP4('T_0','const.combin.I_1'(s(bool,'F_0')),'const.relation.EMPTY_REL_2'(s(X7,X8),s(X7,X9)))) )),
  inference(superposition,[],[f2889,f1253])).
fof(f3141,plain,(
  ( ! [X6,X4,X5,X3] : (sP4('T_0','const.combin.I_1'(s(bool,'T_0')),'const.relation.RUNIV_2'(s(X3,X4),s(X5,X6)))) )),
  inference(superposition,[],[f2889,f1447])).
fof(f3139,plain,(
  ( ! [X1] : (sP4('T_0','const.combin.I_1'(s(bool,X1)),'F_0') | s(bool,'T_0') = s(bool,X1)) )),
  inference(superposition,[],[f2889,f1120])).
fof(f3138,plain,(
  ( ! [X0] : (sP4('T_0','const.combin.I_1'(s(bool,'F_0')),X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2889,f1120])).
fof(f3129,plain,(
  ( ! [X15,X16] : (sP4('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X15),s(bool,X16)),'const.combin.I_1'(s(bool,X15)))) )),
  inference(superposition,[],[f2875,f1255])).
fof(f3128,plain,(
  ( ! [X14,X13] : (sP4('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X13),s(bool,X14)),'const.combin.I_1'(s(bool,X14)))) )),
  inference(superposition,[],[f2875,f1254])).
fof(f3127,plain,(
  ( ! [X12] : (sP4('T_0',happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X12)),'const.combin.I_1'(s(bool,X12)))) )),
  inference(superposition,[],[f2875,f1701])).
fof(f3126,plain,(
  ( ! [X11] : (sP4('T_0',happ(s(fun(bool,bool),'_i_0_0'),s(bool,X11)),'const.combin.I_1'(s(bool,X11)))) )),
  inference(superposition,[],[f2875,f1189])).
fof(f3125,plain,(
  ( ! [X10] : (sP4('T_0','const.combin.I_1'(s(bool,X10)),'const.combin.I_1'(s(bool,X10)))) )),
  inference(superposition,[],[f2875,f1188])).
fof(f3123,plain,(
  ( ! [X8,X7,X9] : (sP4('T_0','const.relation.EMPTY_REL_2'(s(X7,X8),s(X7,X9)),'const.combin.I_1'(s(bool,'F_0')))) )),
  inference(superposition,[],[f2875,f1253])).
fof(f3122,plain,(
  ( ! [X6,X4,X5,X3] : (sP4('T_0','const.relation.RUNIV_2'(s(X3,X4),s(X5,X6)),'const.combin.I_1'(s(bool,'T_0')))) )),
  inference(superposition,[],[f2875,f1447])).
fof(f3119,plain,(
  ( ! [X0] : (sP4('T_0',X0,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2875,f1120])).
fof(f3097,plain,(
  ( ! [X149,X148] : (~sP5('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X148),s(bool,X149)),'const.combin.I_1'(s(bool,X148)))) )),
  inference(superposition,[],[f2788,f1255])).
fof(f3096,plain,(
  ( ! [X146,X147] : (~sP5('F_0','const.combin.I_1'(s(bool,X146)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X146),s(bool,X147)))) )),
  inference(superposition,[],[f2781,f1255])).
fof(f3095,plain,(
  ( ! [X144,X145] : (sP5('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X144),s(bool,X145)),'const.combin.I_1'(s(bool,X144)))) )),
  inference(superposition,[],[f2756,f1255])).
fof(f3094,plain,(
  ( ! [X142,X143] : (sP5('T_0','const.combin.I_1'(s(bool,X142)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X142),s(bool,X143)))) )),
  inference(superposition,[],[f2750,f1255])).
fof(f3084,plain,(
  ( ! [X116,X114,X117,X115] : (~p(s(bool,X114)) | ~sP8('const.bool.COND_3'(s(bool,'T_0'),s(bool,X114),s(bool,X115)),X116,X117) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X114),s(bool,X115)),X116,X117)) )),
  inference(superposition,[],[f1677,f1255])).
fof(f3083,plain,(
  ( ! [X111,X112,X110,X113] : (~p(s(bool,X110)) | ~sP10('const.bool.COND_3'(s(bool,'T_0'),s(bool,X110),s(bool,X111)),X112,X113) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X110),s(bool,X111)),X112,X113)) )),
  inference(superposition,[],[f1676,f1255])).
fof(f3082,plain,(
  ( ! [X109,X107,X108,X106] : (~p(s(bool,X106)) | ~p(s(bool,X108)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X106),s(bool,X107)),X108,X109)) )),
  inference(superposition,[],[f1610,f1255])).
fof(f3081,plain,(
  ( ! [X103,X105,X102,X104] : (~p(s(bool,X102)) | p(s(bool,X104)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X102),s(bool,X103)),X105,X104)) )),
  inference(superposition,[],[f1609,f1255])).
fof(f3080,plain,(
  ( ! [X101,X99,X100,X98] : (~p(s(bool,X98)) | ~p(s(bool,X100)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X98),s(bool,X99)),X100,X101)) )),
  inference(superposition,[],[f1593,f1255])).
fof(f3079,plain,(
  ( ! [X97,X95,X96] : (~p(s(bool,X95)) | p(s(bool,X97)) | sP1(X97,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X95),s(bool,X96)))) )),
  inference(superposition,[],[f1416,f1255])).
fof(f3078,plain,(
  ( ! [X94,X92,X93] : (~p(s(bool,X92)) | p(s(bool,X94)) | sP1('const.bool.COND_3'(s(bool,'T_0'),s(bool,X92),s(bool,X93)),X94)) )),
  inference(superposition,[],[f1413,f1255])).
fof(f3077,plain,(
  ( ! [X90,X91,X89] : (~p(s(bool,X89)) | sP0('const.bool.COND_3'(s(bool,'T_0'),s(bool,X89),s(bool,X90)),X91)) )),
  inference(superposition,[],[f1398,f1255])).
fof(f3076,plain,(
  ( ! [X88,X87,X86] : (~p(s(bool,X86)) | sP0(X88,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X86),s(bool,X87)))) )),
  inference(superposition,[],[f1397,f1255])).
fof(f3071,plain,(
  ( ! [X70,X72,X71,X69,X73] : (s(bool,'T_0') != s(bool,X69) | p(s(bool,sK64(X71,X72,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X69),s(bool,X70)),X73)))) )),
  inference(superposition,[],[f1347,f1255])).
fof(f3052,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') = s(bool,X3) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X3),s(bool,X4)))) )),
  inference(superposition,[],[f1255,f1120])).
fof(f3020,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,sK64(X0,X1,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X2)),X3)))) )),
  inference(unit_resulting_resolution,[],[f1255,f1347])).
fof(f3022,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,sK63(X0,X1,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X2)),X3)))) )),
  inference(unit_resulting_resolution,[],[f1255,f1345])).
fof(f3024,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)),X0)) )),
  inference(unit_resulting_resolution,[],[f1700,f1255,f1562])).
fof(f3025,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)),X0)) )),
  inference(unit_resulting_resolution,[],[f1104,f1255,f1560])).
fof(f3027,plain,(
  ( ! [X0,X1] : (~sP5('F_0',X0,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)))) )),
  inference(unit_resulting_resolution,[],[f1700,f1255,f1562])).
fof(f3028,plain,(
  ( ! [X0,X1] : (sP5('T_0',X0,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)))) )),
  inference(unit_resulting_resolution,[],[f1104,f1255,f1560])).
fof(f3019,plain,(
  ( ! [X14,X13] : (~sP5('F_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X13),s(bool,X14)),'const.combin.I_1'(s(bool,X14)))) )),
  inference(superposition,[],[f2788,f1254])).
fof(f3018,plain,(
  ( ! [X12] : (~sP5('F_0',happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X12)),'const.combin.I_1'(s(bool,X12)))) )),
  inference(superposition,[],[f2788,f1701])).
fof(f3017,plain,(
  ( ! [X11] : (~sP5('F_0',happ(s(fun(bool,bool),'_i_0_0'),s(bool,X11)),'const.combin.I_1'(s(bool,X11)))) )),
  inference(superposition,[],[f2788,f1189])).
fof(f3014,plain,(
  ( ! [X8,X7,X9] : (~sP5('F_0','const.relation.EMPTY_REL_2'(s(X7,X8),s(X7,X9)),'const.combin.I_1'(s(bool,'F_0')))) )),
  inference(superposition,[],[f2788,f1253])).
fof(f3013,plain,(
  ( ! [X6,X4,X5,X3] : (~sP5('F_0','const.relation.RUNIV_2'(s(X3,X4),s(X5,X6)),'const.combin.I_1'(s(bool,'T_0')))) )),
  inference(superposition,[],[f2788,f1447])).
fof(f3011,plain,(
  ( ! [X1] : (~sP5('F_0','F_0','const.combin.I_1'(s(bool,X1))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(superposition,[],[f2788,f1120])).
fof(f3010,plain,(
  ( ! [X0] : (~sP5('F_0',X0,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2788,f1120])).
fof(f3007,plain,(
  ( ! [X14,X13] : (~sP5('F_0','const.combin.I_1'(s(bool,X14)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X13),s(bool,X14)))) )),
  inference(superposition,[],[f2781,f1254])).
fof(f3006,plain,(
  ( ! [X12] : (~sP5('F_0','const.combin.I_1'(s(bool,X12)),happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X12)))) )),
  inference(superposition,[],[f2781,f1701])).
fof(f3005,plain,(
  ( ! [X11] : (~sP5('F_0','const.combin.I_1'(s(bool,X11)),happ(s(fun(bool,bool),'_i_0_0'),s(bool,X11)))) )),
  inference(superposition,[],[f2781,f1189])).
fof(f3004,plain,(
  ( ! [X10] : (~sP5('F_0','const.combin.I_1'(s(bool,X10)),'const.combin.I_1'(s(bool,X10)))) )),
  inference(superposition,[],[f2781,f1188])).
fof(f3002,plain,(
  ( ! [X8,X7,X9] : (~sP5('F_0','const.combin.I_1'(s(bool,'F_0')),'const.relation.EMPTY_REL_2'(s(X7,X8),s(X7,X9)))) )),
  inference(superposition,[],[f2781,f1253])).
fof(f3001,plain,(
  ( ! [X6,X4,X5,X3] : (~sP5('F_0','const.combin.I_1'(s(bool,'T_0')),'const.relation.RUNIV_2'(s(X3,X4),s(X5,X6)))) )),
  inference(superposition,[],[f2781,f1447])).
fof(f2999,plain,(
  ( ! [X1] : (~sP5('F_0','const.combin.I_1'(s(bool,X1)),'F_0') | s(bool,'T_0') = s(bool,X1)) )),
  inference(superposition,[],[f2781,f1120])).
fof(f2998,plain,(
  ( ! [X0] : (~sP5('F_0','const.combin.I_1'(s(bool,'F_0')),X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2781,f1120])).
fof(f2974,plain,(
  ( ! [X136,X137] : (sP5('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X136),s(bool,X137)),'const.combin.I_1'(s(bool,X137)))) )),
  inference(superposition,[],[f2756,f1254])).
fof(f2973,plain,(
  ( ! [X134,X135] : (sP5('T_0','const.combin.I_1'(s(bool,X135)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X134),s(bool,X135)))) )),
  inference(superposition,[],[f2750,f1254])).
fof(f2963,plain,(
  ( ! [X109,X107,X108,X106] : (~p(s(bool,X107)) | ~sP8('const.bool.COND_3'(s(bool,'F_0'),s(bool,X106),s(bool,X107)),X108,X109) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X106),s(bool,X107)),X108,X109)) )),
  inference(superposition,[],[f1677,f1254])).
fof(f2962,plain,(
  ( ! [X103,X105,X102,X104] : (~p(s(bool,X103)) | ~sP10('const.bool.COND_3'(s(bool,'F_0'),s(bool,X102),s(bool,X103)),X104,X105) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X102),s(bool,X103)),X104,X105)) )),
  inference(superposition,[],[f1676,f1254])).
fof(f2961,plain,(
  ( ! [X101,X99,X100,X98] : (~p(s(bool,X99)) | ~p(s(bool,X100)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X98),s(bool,X99)),X100,X101)) )),
  inference(superposition,[],[f1610,f1254])).
fof(f2960,plain,(
  ( ! [X94,X97,X95,X96] : (~p(s(bool,X95)) | p(s(bool,X96)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X94),s(bool,X95)),X97,X96)) )),
  inference(superposition,[],[f1609,f1254])).
fof(f2959,plain,(
  ( ! [X92,X90,X93,X91] : (~p(s(bool,X91)) | ~p(s(bool,X92)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X90),s(bool,X91)),X92,X93)) )),
  inference(superposition,[],[f1593,f1254])).
fof(f2958,plain,(
  ( ! [X88,X87,X89] : (~p(s(bool,X88)) | p(s(bool,X89)) | sP1(X89,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X87),s(bool,X88)))) )),
  inference(superposition,[],[f1416,f1254])).
fof(f2957,plain,(
  ( ! [X85,X86,X84] : (~p(s(bool,X85)) | p(s(bool,X86)) | sP1('const.bool.COND_3'(s(bool,'F_0'),s(bool,X84),s(bool,X85)),X86)) )),
  inference(superposition,[],[f1413,f1254])).
fof(f2956,plain,(
  ( ! [X83,X81,X82] : (~p(s(bool,X82)) | sP0('const.bool.COND_3'(s(bool,'F_0'),s(bool,X81),s(bool,X82)),X83)) )),
  inference(superposition,[],[f1398,f1254])).
fof(f2955,plain,(
  ( ! [X80,X78,X79] : (~p(s(bool,X79)) | sP0(X80,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X78),s(bool,X79)))) )),
  inference(superposition,[],[f1397,f1254])).
fof(f2950,plain,(
  ( ! [X61,X64,X62,X65,X63] : (s(bool,'T_0') != s(bool,X62) | p(s(bool,sK64(X63,X64,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X61),s(bool,X62)),X65)))) )),
  inference(superposition,[],[f1347,f1254])).
fof(f2933,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') = s(bool,X4) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,X4)))) )),
  inference(superposition,[],[f1254,f1120])).
fof(f2911,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.bool.COND_3'(s(bool,X0),s(X1,X3),s(X1,X2))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f1254,f1120])).
fof(f2901,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,sK64(X0,X1,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,'T_0')),X3)))) )),
  inference(unit_resulting_resolution,[],[f1254,f1347])).
fof(f2903,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,sK63(X0,X1,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,'F_0')),X3)))) )),
  inference(unit_resulting_resolution,[],[f1254,f1345])).
fof(f2905,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,X1)),X1)) )),
  inference(unit_resulting_resolution,[],[f1700,f1254,f1562])).
fof(f2906,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,X1)),X1)) )),
  inference(unit_resulting_resolution,[],[f1104,f1254,f1560])).
fof(f2908,plain,(
  ( ! [X0,X1] : (~sP5('F_0',X0,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f1700,f1254,f1562])).
fof(f2909,plain,(
  ( ! [X0,X1] : (sP5('T_0',X0,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f1104,f1254,f1560])).
fof(f2900,plain,(
  ( ! [X12] : (sP5('T_0',happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X12)),'const.combin.I_1'(s(bool,X12)))) )),
  inference(superposition,[],[f2756,f1701])).
fof(f2899,plain,(
  ( ! [X11] : (sP5('T_0',happ(s(fun(bool,bool),'_i_0_0'),s(bool,X11)),'const.combin.I_1'(s(bool,X11)))) )),
  inference(superposition,[],[f2756,f1189])).
fof(f2896,plain,(
  ( ! [X8,X7,X9] : (sP5('T_0','const.relation.EMPTY_REL_2'(s(X7,X8),s(X7,X9)),'const.combin.I_1'(s(bool,'F_0')))) )),
  inference(superposition,[],[f2756,f1253])).
fof(f2895,plain,(
  ( ! [X6,X4,X5,X3] : (sP5('T_0','const.relation.RUNIV_2'(s(X3,X4),s(X5,X6)),'const.combin.I_1'(s(bool,'T_0')))) )),
  inference(superposition,[],[f2756,f1447])).
fof(f2893,plain,(
  ( ! [X1] : (sP5('T_0','F_0','const.combin.I_1'(s(bool,X1))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(superposition,[],[f2756,f1120])).
fof(f2892,plain,(
  ( ! [X0] : (sP5('T_0',X0,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2756,f1120])).
fof(f2889,plain,(
  ( ! [X0] : (sP4('T_0','const.combin.I_1'(s(bool,X0)),X0)) )),
  inference(unit_resulting_resolution,[],[f2756,f1542])).
fof(f2886,plain,(
  ( ! [X12] : (sP5('T_0','const.combin.I_1'(s(bool,X12)),happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X12)))) )),
  inference(superposition,[],[f2750,f1701])).
fof(f2885,plain,(
  ( ! [X11] : (sP5('T_0','const.combin.I_1'(s(bool,X11)),happ(s(fun(bool,bool),'_i_0_0'),s(bool,X11)))) )),
  inference(superposition,[],[f2750,f1189])).
fof(f2884,plain,(
  ( ! [X10] : (sP5('T_0','const.combin.I_1'(s(bool,X10)),'const.combin.I_1'(s(bool,X10)))) )),
  inference(superposition,[],[f2750,f1188])).
fof(f2882,plain,(
  ( ! [X8,X7,X9] : (sP5('T_0','const.combin.I_1'(s(bool,'F_0')),'const.relation.EMPTY_REL_2'(s(X7,X8),s(X7,X9)))) )),
  inference(superposition,[],[f2750,f1253])).
fof(f2881,plain,(
  ( ! [X6,X4,X5,X3] : (sP5('T_0','const.combin.I_1'(s(bool,'T_0')),'const.relation.RUNIV_2'(s(X3,X4),s(X5,X6)))) )),
  inference(superposition,[],[f2750,f1447])).
fof(f2879,plain,(
  ( ! [X1] : (sP5('T_0','const.combin.I_1'(s(bool,X1)),'F_0') | s(bool,'T_0') = s(bool,X1)) )),
  inference(superposition,[],[f2750,f1120])).
fof(f2878,plain,(
  ( ! [X0] : (sP5('T_0','const.combin.I_1'(s(bool,'F_0')),X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2750,f1120])).
fof(f2875,plain,(
  ( ! [X0] : (sP4('T_0',X0,'const.combin.I_1'(s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f2750,f1542])).
fof(f2842,plain,(
  ( ! [X76,X77,X75] : (~p(s(bool,X75)) | ~sP8(happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X75)),X76,X77) | sP9(happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X75)),X76,X77)) )),
  inference(superposition,[],[f1677,f1701])).
fof(f2841,plain,(
  ( ! [X74,X72,X73] : (~p(s(bool,X72)) | ~sP10(happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X72)),X73,X74) | sP11(happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X72)),X73,X74)) )),
  inference(superposition,[],[f1676,f1701])).
fof(f2840,plain,(
  ( ! [X70,X71,X69] : (~p(s(bool,X69)) | ~p(s(bool,X70)) | sP11(happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X69)),X70,X71)) )),
  inference(superposition,[],[f1610,f1701])).
fof(f2839,plain,(
  ( ! [X68,X66,X67] : (~p(s(bool,X66)) | p(s(bool,X67)) | sP11(happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X66)),X68,X67)) )),
  inference(superposition,[],[f1609,f1701])).
fof(f2838,plain,(
  ( ! [X64,X65,X63] : (~p(s(bool,X63)) | ~p(s(bool,X64)) | sP9(happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X63)),X64,X65)) )),
  inference(superposition,[],[f1593,f1701])).
fof(f2837,plain,(
  ( ! [X61,X62] : (~p(s(bool,X61)) | p(s(bool,X62)) | sP1(X62,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X61)))) )),
  inference(superposition,[],[f1416,f1701])).
fof(f2836,plain,(
  ( ! [X59,X60] : (~p(s(bool,X59)) | p(s(bool,X60)) | sP1(happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X59)),X60)) )),
  inference(superposition,[],[f1413,f1701])).
fof(f2835,plain,(
  ( ! [X57,X58] : (sP0(happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X57)),X58) | ~p(s(bool,X57))) )),
  inference(superposition,[],[f1398,f1701])).
fof(f2834,plain,(
  ( ! [X56,X55] : (sP0(X56,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X55))) | ~p(s(bool,X55))) )),
  inference(superposition,[],[f1397,f1701])).
fof(f2829,plain,(
  ( ! [X45,X43,X46,X44] : (s(bool,'T_0') != s(bool,X43) | p(s(bool,sK64(X44,X45,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X43)),X46)))) )),
  inference(superposition,[],[f1347,f1701])).
fof(f2813,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X2)))) )),
  inference(superposition,[],[f1701,f1120])).
fof(f2794,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,sK64(X0,X1,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'T_0')),X2)))) )),
  inference(unit_resulting_resolution,[],[f1701,f1347])).
fof(f2796,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,sK63(X0,X1,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0')),X2)))) )),
  inference(unit_resulting_resolution,[],[f1701,f1345])).
fof(f2798,plain,(
  ( ! [X0] : (~sP5('F_0',happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X0)),X0)) )),
  inference(unit_resulting_resolution,[],[f1700,f1701,f1562])).
fof(f2799,plain,(
  ( ! [X0] : (sP5('T_0',happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X0)),X0)) )),
  inference(unit_resulting_resolution,[],[f1104,f1701,f1560])).
fof(f2801,plain,(
  ( ! [X0] : (~sP5('F_0',X0,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f1700,f1701,f1562])).
fof(f2802,plain,(
  ( ! [X0] : (sP5('T_0',X0,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f1104,f1701,f1560])).
fof(f2778,plain,(
  ( ! [X2,X0,X1] : (~sP5('F_0','const.relation.EMPTY_REL_2'(s(X0,X1),s(X0,X2)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1700,f1253,f1562])).
fof(f2780,plain,(
  ( ! [X2,X0,X3,X1] : (~sP5('F_0','const.relation.RUNIV_2'(s(X0,X1),s(X2,X3)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f1700,f1447,f1562])).
fof(f2782,plain,(
  ( ! [X0] : (~sP5('F_0',happ(s(fun(bool,bool),'_i_0_0'),s(bool,X0)),X0)) )),
  inference(unit_resulting_resolution,[],[f1700,f1189,f1562])).
fof(f2785,plain,(
  ( ! [X2,X0,X1] : (~sP5('F_0','F_0','const.relation.EMPTY_REL_2'(s(X0,X1),s(X0,X2)))) )),
  inference(unit_resulting_resolution,[],[f1700,f1253,f1562])).
fof(f2787,plain,(
  ( ! [X2,X0,X3,X1] : (~sP5('F_0','T_0','const.relation.RUNIV_2'(s(X0,X1),s(X2,X3)))) )),
  inference(unit_resulting_resolution,[],[f1700,f1447,f1562])).
fof(f2788,plain,(
  ( ! [X0] : (~sP5('F_0',X0,'const.combin.I_1'(s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f1700,f1188,f1562])).
fof(f2789,plain,(
  ( ! [X0] : (~sP5('F_0',X0,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f1700,f1189,f1562])).
fof(f2767,plain,(
  ~sP5('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f1104,f1115,f1561])).
fof(f2769,plain,(
  ~sP5('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f1104,f1115,f1561])).
fof(f2747,plain,(
  ( ! [X2,X0,X1] : (sP5('T_0','const.relation.EMPTY_REL_2'(s(X0,X1),s(X0,X2)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1104,f1253,f1560])).
fof(f2749,plain,(
  ( ! [X2,X0,X3,X1] : (sP5('T_0','const.relation.RUNIV_2'(s(X0,X1),s(X2,X3)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f1104,f1447,f1560])).
fof(f2750,plain,(
  ( ! [X0] : (sP5('T_0','const.combin.I_1'(s(bool,X0)),X0)) )),
  inference(unit_resulting_resolution,[],[f1104,f1188,f1560])).
fof(f2751,plain,(
  ( ! [X0] : (sP5('T_0',happ(s(fun(bool,bool),'_i_0_0'),s(bool,X0)),X0)) )),
  inference(unit_resulting_resolution,[],[f1104,f1189,f1560])).
fof(f2753,plain,(
  ( ! [X2,X0,X1] : (sP5('T_0','F_0','const.relation.EMPTY_REL_2'(s(X0,X1),s(X0,X2)))) )),
  inference(unit_resulting_resolution,[],[f1104,f1253,f1560])).
fof(f2755,plain,(
  ( ! [X2,X0,X3,X1] : (sP5('T_0','T_0','const.relation.RUNIV_2'(s(X0,X1),s(X2,X3)))) )),
  inference(unit_resulting_resolution,[],[f1104,f1447,f1560])).
fof(f2756,plain,(
  ( ! [X0] : (sP5('T_0',X0,'const.combin.I_1'(s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f1104,f1188,f1560])).
fof(f2757,plain,(
  ( ! [X0] : (sP5('T_0',X0,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f1104,f1189,f1560])).
fof(f2738,plain,(
  sP5('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f1700,f1115,f1559])).
fof(f2739,plain,(
  sP5('F_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f1700,f1115,f1559])).
fof(f2663,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f2647,f1700])).
fof(f2647,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f1361,f1120])).
fof(f2576,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,sK64(X0,X1,'F_0',X2))) )),
  inference(unit_resulting_resolution,[],[f2552,f1130])).
fof(f2577,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,sK64(X0,X1,'F_0',X2))) )),
  inference(unit_resulting_resolution,[],[f2552,f1133])).
fof(f2578,plain,(
  ( ! [X2,X0,X1] : (~sP1(sK64(X0,X1,'F_0',X2),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1700,f2552,f1411])).
fof(f2580,plain,(
  ( ! [X2,X0,X1] : (~sP1('F_0',sK64(X0,X1,'F_0',X2))) )),
  inference(unit_resulting_resolution,[],[f1700,f2552,f1411])).
fof(f2581,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~sP1(sK64(X0,X1,'F_0',X2),sK64(X3,X4,'F_0',X5))) )),
  inference(unit_resulting_resolution,[],[f2552,f2552,f1411])).
fof(f2582,plain,(
  ( ! [X2,X0,X1] : (sP1('T_0',sK64(X0,X1,'F_0',X2))) )),
  inference(unit_resulting_resolution,[],[f1104,f2552,f1413])).
fof(f2583,plain,(
  ( ! [X2,X0,X1] : (sP1(sK64(X0,X1,'F_0',X2),'T_0')) )),
  inference(unit_resulting_resolution,[],[f1104,f2552,f1416])).
fof(f2600,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (sP7(sK64(X0,X1,'F_0',X2),X3,sK64(X4,X5,'F_0',X6))) )),
  inference(unit_resulting_resolution,[],[f2552,f2552,f1574])).
fof(f2605,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (sP7(sK64(X0,X1,'F_0',X2),sK64(X3,X4,'F_0',X5),X6)) )),
  inference(unit_resulting_resolution,[],[f2552,f2552,f1575])).
fof(f2608,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP8(sK64(X0,X1,'F_0',X2),X3,X4)) )),
  inference(unit_resulting_resolution,[],[f2552,f1590])).
fof(f2612,plain,(
  ( ! [X2,X0,X3,X1] : (~sP11('F_0',X0,sK64(X1,X2,'F_0',X3))) )),
  inference(unit_resulting_resolution,[],[f1700,f2552,f1597])).
fof(f2613,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~sP11(sK64(X0,X1,'F_0',X2),X3,sK64(X4,X5,'F_0',X6))) )),
  inference(unit_resulting_resolution,[],[f2552,f2552,f1597])).
fof(f2618,plain,(
  ( ! [X2,X0,X3,X1] : (sP11('T_0',X0,sK64(X1,X2,'F_0',X3))) )),
  inference(unit_resulting_resolution,[],[f1104,f2552,f1609])).
fof(f2620,plain,(
  ( ! [X2,X0,X3,X1] : (sP7(sK64(X0,X1,'F_0',X2),X3,'F_0')) )),
  inference(unit_resulting_resolution,[],[f1921,f2552,f1678])).
fof(f2622,plain,(
  ( ! [X2,X0,X3,X1] : (sP7(sK64(X0,X1,'F_0',X2),'F_0',X3)) )),
  inference(unit_resulting_resolution,[],[f1923,f2552,f1678])).
fof(f2627,plain,(
  ( ! [X2,X0,X1] : (~sP0(sK64(X0,X1,'F_0',X2),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1700,f2552,f1851])).
fof(f2629,plain,(
  ( ! [X2,X0,X1] : (~sP0('F_0',sK64(X0,X1,'F_0',X2))) )),
  inference(unit_resulting_resolution,[],[f1700,f2552,f1851])).
fof(f2630,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~sP0(sK64(X0,X1,'F_0',X2),sK64(X3,X4,'F_0',X5))) )),
  inference(unit_resulting_resolution,[],[f2552,f2552,f1851])).
fof(f2570,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,sK64(X0,X1,'T_0',X2)))) )),
  inference(equality_resolution,[],[f1347])).
fof(f2569,plain,(
  ( ! [X35,X33,X34,X32] : (s(bool,'T_0') != s(bool,X32) | p(s(bool,sK64(X33,X34,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X32)),X35)))) )),
  inference(superposition,[],[f1347,f1189])).
fof(f2568,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'T_0') != s(bool,X28) | p(s(bool,sK64(X29,X30,'const.combin.I_1'(s(bool,X28)),X31)))) )),
  inference(superposition,[],[f1347,f1188])).
fof(f2559,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,sK64(X0,X1,'const.relation.RUNIV_2'(s(X2,X3),s(X4,X5)),X6)))) )),
  inference(unit_resulting_resolution,[],[f1447,f1347])).
fof(f2560,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,sK64(X0,X1,'const.combin.I_1'(s(bool,'T_0')),X2)))) )),
  inference(unit_resulting_resolution,[],[f1188,f1347])).
fof(f2561,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,sK64(X0,X1,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'T_0')),X2)))) )),
  inference(unit_resulting_resolution,[],[f1189,f1347])).
fof(f2547,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') != s(bool,X2) | s(bool,'T_0') = s(bool,sK63(X0,X1,X2,X3))) )),
  inference(subsumption_resolution,[],[f2546,f1700])).
fof(f2546,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'F_0')) | s(bool,'F_0') != s(bool,X2) | s(bool,'T_0') = s(bool,sK63(X0,X1,X2,X3))) )),
  inference(superposition,[],[f1345,f1120])).
fof(f2543,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,sK63(X0,X1,'const.relation.EMPTY_REL_2'(s(X2,X3),s(X2,X4)),X5)))) )),
  inference(unit_resulting_resolution,[],[f1253,f1345])).
fof(f2544,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,sK63(X0,X1,'const.combin.I_1'(s(bool,'F_0')),X2)))) )),
  inference(unit_resulting_resolution,[],[f1188,f1345])).
fof(f2545,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,sK63(X0,X1,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0')),X2)))) )),
  inference(unit_resulting_resolution,[],[f1189,f1345])).
fof(f2473,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,sK63(X0,X1,'T_0',X2))) )),
  inference(unit_resulting_resolution,[],[f2471,f1130])).
fof(f2474,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,sK63(X0,X1,'T_0',X2))) )),
  inference(unit_resulting_resolution,[],[f2471,f1133])).
fof(f2475,plain,(
  ( ! [X2,X0,X1] : (~sP1(sK63(X0,X1,'T_0',X2),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1700,f2471,f1411])).
fof(f2477,plain,(
  ( ! [X2,X0,X1] : (~sP1('F_0',sK63(X0,X1,'T_0',X2))) )),
  inference(unit_resulting_resolution,[],[f1700,f2471,f1411])).
fof(f2478,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~sP1(sK63(X0,X1,'T_0',X2),sK63(X3,X4,'T_0',X5))) )),
  inference(unit_resulting_resolution,[],[f2471,f2471,f1411])).
fof(f2479,plain,(
  ( ! [X2,X0,X1] : (sP1('T_0',sK63(X0,X1,'T_0',X2))) )),
  inference(unit_resulting_resolution,[],[f1104,f2471,f1413])).
fof(f2480,plain,(
  ( ! [X2,X0,X1] : (sP1(sK63(X0,X1,'T_0',X2),'T_0')) )),
  inference(unit_resulting_resolution,[],[f1104,f2471,f1416])).
fof(f2497,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (sP7(sK63(X0,X1,'T_0',X2),X3,sK63(X4,X5,'T_0',X6))) )),
  inference(unit_resulting_resolution,[],[f2471,f2471,f1574])).
fof(f2502,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (sP7(sK63(X0,X1,'T_0',X2),sK63(X3,X4,'T_0',X5),X6)) )),
  inference(unit_resulting_resolution,[],[f2471,f2471,f1575])).
fof(f2505,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP8(sK63(X0,X1,'T_0',X2),X3,X4)) )),
  inference(unit_resulting_resolution,[],[f2471,f1590])).
fof(f2509,plain,(
  ( ! [X2,X0,X3,X1] : (~sP11('F_0',X0,sK63(X1,X2,'T_0',X3))) )),
  inference(unit_resulting_resolution,[],[f1700,f2471,f1597])).
fof(f2510,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~sP11(sK63(X0,X1,'T_0',X2),X3,sK63(X4,X5,'T_0',X6))) )),
  inference(unit_resulting_resolution,[],[f2471,f2471,f1597])).
fof(f2515,plain,(
  ( ! [X2,X0,X3,X1] : (sP11('T_0',X0,sK63(X1,X2,'T_0',X3))) )),
  inference(unit_resulting_resolution,[],[f1104,f2471,f1609])).
fof(f2517,plain,(
  ( ! [X2,X0,X3,X1] : (sP7(sK63(X0,X1,'T_0',X2),X3,'F_0')) )),
  inference(unit_resulting_resolution,[],[f1921,f2471,f1678])).
fof(f2519,plain,(
  ( ! [X2,X0,X3,X1] : (sP7(sK63(X0,X1,'T_0',X2),'F_0',X3)) )),
  inference(unit_resulting_resolution,[],[f1923,f2471,f1678])).
fof(f2524,plain,(
  ( ! [X2,X0,X1] : (~sP0(sK63(X0,X1,'T_0',X2),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1700,f2471,f1851])).
fof(f2526,plain,(
  ( ! [X2,X0,X1] : (~sP0('F_0',sK63(X0,X1,'T_0',X2))) )),
  inference(unit_resulting_resolution,[],[f1700,f2471,f1851])).
fof(f2527,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~sP0(sK63(X0,X1,'T_0',X2),sK63(X3,X4,'T_0',X5))) )),
  inference(unit_resulting_resolution,[],[f2471,f2471,f1851])).
fof(f2428,plain,(
  ( ! [X61,X62,X63] : (~p(s(bool,X61)) | ~sP8(happ(s(fun(bool,bool),'_i_0_0'),s(bool,X61)),X62,X63) | sP9(happ(s(fun(bool,bool),'_i_0_0'),s(bool,X61)),X62,X63)) )),
  inference(superposition,[],[f1677,f1189])).
fof(f2427,plain,(
  ( ! [X59,X60,X58] : (~p(s(bool,X58)) | ~sP10(happ(s(fun(bool,bool),'_i_0_0'),s(bool,X58)),X59,X60) | sP11(happ(s(fun(bool,bool),'_i_0_0'),s(bool,X58)),X59,X60)) )),
  inference(superposition,[],[f1676,f1189])).
fof(f2426,plain,(
  ( ! [X57,X56,X55] : (~p(s(bool,X55)) | ~p(s(bool,X56)) | sP11(happ(s(fun(bool,bool),'_i_0_0'),s(bool,X55)),X56,X57)) )),
  inference(superposition,[],[f1610,f1189])).
fof(f2425,plain,(
  ( ! [X54,X52,X53] : (~p(s(bool,X52)) | p(s(bool,X53)) | sP11(happ(s(fun(bool,bool),'_i_0_0'),s(bool,X52)),X54,X53)) )),
  inference(superposition,[],[f1609,f1189])).
fof(f2424,plain,(
  ( ! [X50,X51,X49] : (~p(s(bool,X49)) | ~p(s(bool,X50)) | sP9(happ(s(fun(bool,bool),'_i_0_0'),s(bool,X49)),X50,X51)) )),
  inference(superposition,[],[f1593,f1189])).
fof(f2423,plain,(
  ( ! [X47,X48] : (~p(s(bool,X47)) | p(s(bool,X48)) | sP1(X48,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X47)))) )),
  inference(superposition,[],[f1416,f1189])).
fof(f2422,plain,(
  ( ! [X45,X46] : (~p(s(bool,X45)) | p(s(bool,X46)) | sP1(happ(s(fun(bool,bool),'_i_0_0'),s(bool,X45)),X46)) )),
  inference(superposition,[],[f1413,f1189])).
fof(f2421,plain,(
  ( ! [X43,X44] : (sP0(happ(s(fun(bool,bool),'_i_0_0'),s(bool,X43)),X44) | ~p(s(bool,X43))) )),
  inference(superposition,[],[f1398,f1189])).
fof(f2420,plain,(
  ( ! [X41,X42] : (sP0(X42,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X41))) | ~p(s(bool,X41))) )),
  inference(superposition,[],[f1397,f1189])).
fof(f2405,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X2)))) )),
  inference(superposition,[],[f1189,f1120])).
fof(f2365,plain,(
  ( ! [X0] : (p(s(bool,'|dollar|or_2'(s(bool,'T_0'),s(bool,X0))))) )),
  inference(unit_resulting_resolution,[],[f1104,f1836])).
fof(f2283,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'|dollar|or_2'(s(bool,X0),s(bool,'T_0')))) )),
  inference(unit_resulting_resolution,[],[f2271,f1134])).
fof(f2284,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,'|dollar|or_2'(s(bool,X0),s(bool,'T_0')))) )),
  inference(unit_resulting_resolution,[],[f2271,f1137])).
fof(f2285,plain,(
  ( ! [X0,X1] : (sP0(X0,'|dollar|or_2'(s(bool,X1),s(bool,'T_0')))) )),
  inference(unit_resulting_resolution,[],[f2271,f1397])).
fof(f2286,plain,(
  ( ! [X0,X1] : (sP0('|dollar|or_2'(s(bool,X0),s(bool,'T_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f2271,f1398])).
fof(f2287,plain,(
  ( ! [X0] : (~sP1('|dollar|or_2'(s(bool,X0),s(bool,'T_0')),'T_0')) )),
  inference(unit_resulting_resolution,[],[f1104,f2271,f1412])).
fof(f2289,plain,(
  ( ! [X0] : (~sP1('T_0','|dollar|or_2'(s(bool,X0),s(bool,'T_0')))) )),
  inference(unit_resulting_resolution,[],[f1104,f2271,f1412])).
fof(f2290,plain,(
  ( ! [X0,X1] : (~sP1('|dollar|or_2'(s(bool,X0),s(bool,'T_0')),'|dollar|or_2'(s(bool,X1),s(bool,'T_0')))) )),
  inference(unit_resulting_resolution,[],[f2271,f2271,f1412])).
fof(f2291,plain,(
  ( ! [X0] : (sP1('|dollar|or_2'(s(bool,X0),s(bool,'T_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1700,f2271,f1413])).
fof(f2292,plain,(
  ( ! [X0] : (sP1('F_0','|dollar|or_2'(s(bool,X0),s(bool,'T_0')))) )),
  inference(unit_resulting_resolution,[],[f1700,f2271,f1416])).
fof(f2293,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,'|dollar|or_2'(s(bool,X2),s(bool,'T_0')))) )),
  inference(unit_resulting_resolution,[],[f2271,f1548])).
fof(f2294,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,'|dollar|or_2'(s(bool,X2),s(bool,'T_0')))) )),
  inference(unit_resulting_resolution,[],[f2271,f1552])).
fof(f2295,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,'|dollar|or_2'(s(bool,X2),s(bool,'T_0')))) )),
  inference(unit_resulting_resolution,[],[f2271,f1556])).
fof(f2296,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,'|dollar|or_2'(s(bool,X1),s(bool,'T_0')),X2)) )),
  inference(unit_resulting_resolution,[],[f2271,f1557])).
fof(f2297,plain,(
  ( ! [X2,X0,X1] : (sP2('|dollar|or_2'(s(bool,X0),s(bool,'T_0')),X1,X2)) )),
  inference(unit_resulting_resolution,[],[f2271,f1558])).
fof(f2298,plain,(
  ( ! [X0,X1] : (~sP7('|dollar|or_2'(s(bool,X0),s(bool,'T_0')),X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f1700,f2271,f1564])).
fof(f2299,plain,(
  ( ! [X0,X1] : (~sP7('|dollar|or_2'(s(bool,X0),s(bool,'T_0')),'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f1700,f2271,f1565])).
fof(f2300,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,'|dollar|or_2'(s(bool,X2),s(bool,'T_0')))) )),
  inference(unit_resulting_resolution,[],[f2271,f1571])).
fof(f2301,plain,(
  ( ! [X0,X1] : (~sP9('F_0',X0,'|dollar|or_2'(s(bool,X1),s(bool,'T_0')))) )),
  inference(unit_resulting_resolution,[],[f1700,f2271,f1580])).
fof(f2303,plain,(
  ( ! [X0,X1] : (~sP9('F_0','|dollar|or_2'(s(bool,X0),s(bool,'T_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f1700,f2271,f1581])).
fof(f2305,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,'|dollar|or_2'(s(bool,X2),s(bool,'T_0')))) )),
  inference(unit_resulting_resolution,[],[f2271,f1588])).
fof(f2306,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,'|dollar|or_2'(s(bool,X1),s(bool,'T_0')),X2)) )),
  inference(unit_resulting_resolution,[],[f2271,f1589])).
fof(f2307,plain,(
  ( ! [X0,X1] : (sP9('T_0',X0,'|dollar|or_2'(s(bool,X1),s(bool,'T_0')))) )),
  inference(unit_resulting_resolution,[],[f1104,f2271,f1592])).
fof(f2311,plain,(
  ( ! [X2,X0,X1] : (sP9('|dollar|or_2'(s(bool,X0),s(bool,'T_0')),X1,'|dollar|or_2'(s(bool,X2),s(bool,'T_0')))) )),
  inference(unit_resulting_resolution,[],[f2271,f2271,f1592])).
fof(f2312,plain,(
  ( ! [X0,X1] : (sP9('T_0','|dollar|or_2'(s(bool,X0),s(bool,'T_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f1104,f2271,f1593])).
fof(f2316,plain,(
  ( ! [X2,X0,X1] : (sP9('|dollar|or_2'(s(bool,X0),s(bool,'T_0')),'|dollar|or_2'(s(bool,X1),s(bool,'T_0')),X2)) )),
  inference(unit_resulting_resolution,[],[f2271,f2271,f1593])).
fof(f2317,plain,(
  ( ! [X0,X1] : (~sP11('F_0','|dollar|or_2'(s(bool,X0),s(bool,'T_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f1700,f2271,f1598])).
fof(f2320,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,'|dollar|or_2'(s(bool,X1),s(bool,'T_0')),X2)) )),
  inference(unit_resulting_resolution,[],[f2271,f1606])).
fof(f2322,plain,(
  ( ! [X0,X1] : (sP11('T_0','|dollar|or_2'(s(bool,X0),s(bool,'T_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f1104,f2271,f1610])).
fof(f2326,plain,(
  ( ! [X2,X0,X1] : (sP11('|dollar|or_2'(s(bool,X0),s(bool,'T_0')),'|dollar|or_2'(s(bool,X1),s(bool,'T_0')),X2)) )),
  inference(unit_resulting_resolution,[],[f2271,f2271,f1610])).
fof(f2328,plain,(
  ( ! [X0,X1] : (sP11('|dollar|or_2'(s(bool,X0),s(bool,'T_0')),X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f1929,f2271,f1676])).
fof(f2329,plain,(
  ( ! [X0,X1] : (sP11('|dollar|or_2'(s(bool,X0),s(bool,'T_0')),'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f1930,f2271,f1676])).
fof(f2330,plain,(
  ( ! [X0,X1] : (sP9('|dollar|or_2'(s(bool,X0),s(bool,'T_0')),X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f1924,f2271,f1677])).
fof(f2331,plain,(
  ( ! [X0,X1] : (sP9('|dollar|or_2'(s(bool,X0),s(bool,'T_0')),'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f1925,f2271,f1677])).
fof(f2332,plain,(
  ( ! [X0,X1] : (p(s(bool,'|dollar|or_2'(s(bool,X0),s(bool,'|dollar|or_2'(s(bool,X1),s(bool,'T_0'))))))) )),
  inference(unit_resulting_resolution,[],[f2271,f1835])).
fof(f2333,plain,(
  ( ! [X0,X1] : (sP0('const.combin.I_1'(s(bool,'|dollar|or_2'(s(bool,X0),s(bool,'T_0')))),X1)) )),
  inference(unit_resulting_resolution,[],[f2271,f1934])).
fof(f2334,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.combin.I_1'(s(bool,'|dollar|or_2'(s(bool,X1),s(bool,'T_0')))))) )),
  inference(unit_resulting_resolution,[],[f2271,f1935])).
fof(f2271,plain,(
  ( ! [X0] : (p(s(bool,'|dollar|or_2'(s(bool,X0),s(bool,'T_0'))))) )),
  inference(unit_resulting_resolution,[],[f1104,f1835])).
fof(f2204,plain,(
  ( ! [X41,X42] : (s(bool,X41) = s(bool,'const.relation.EMPTY_REL_2'(s(bool,X41),s(bool,X42))) | s(bool,'T_0') = s(bool,X41)) )),
  inference(superposition,[],[f1253,f1120])).
fof(f2203,plain,(
  ( ! [X39,X40] : (s(bool,X39) = s(bool,'const.relation.EMPTY_REL_2'(s(bool,X40),s(bool,X39))) | s(bool,'T_0') = s(bool,X39)) )),
  inference(superposition,[],[f1253,f1120])).
fof(f2183,plain,(
  ( ! [X10] : (s(bool,'F_0') = s(bool,X10) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X10)))) )),
  inference(superposition,[],[f1120,f1188])).
fof(f2179,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f1120,f1120])).
fof(f2176,plain,(
  ( ! [X61,X64,X62,X60,X65,X63] : (~sP8('const.relation.RUNIV_2'(s(X60,X61),s(X62,X63)),X64,X65) | sP9('const.relation.RUNIV_2'(s(X60,X61),s(X62,X63)),X64,X65)) )),
  inference(subsumption_resolution,[],[f2157,f1104])).
fof(f2157,plain,(
  ( ! [X61,X64,X62,X60,X65,X63] : (~p(s(bool,'T_0')) | ~sP8('const.relation.RUNIV_2'(s(X60,X61),s(X62,X63)),X64,X65) | sP9('const.relation.RUNIV_2'(s(X60,X61),s(X62,X63)),X64,X65)) )),
  inference(superposition,[],[f1677,f1447])).
fof(f2175,plain,(
  ( ! [X59,X57,X54,X58,X56,X55] : (~sP10('const.relation.RUNIV_2'(s(X54,X55),s(X56,X57)),X58,X59) | sP11('const.relation.RUNIV_2'(s(X54,X55),s(X56,X57)),X58,X59)) )),
  inference(subsumption_resolution,[],[f2156,f1104])).
fof(f2156,plain,(
  ( ! [X59,X57,X54,X58,X56,X55] : (~p(s(bool,'T_0')) | ~sP10('const.relation.RUNIV_2'(s(X54,X55),s(X56,X57)),X58,X59) | sP11('const.relation.RUNIV_2'(s(X54,X55),s(X56,X57)),X58,X59)) )),
  inference(superposition,[],[f1676,f1447])).
fof(f2174,plain,(
  ( ! [X52,X50,X48,X53,X51,X49] : (~p(s(bool,X52)) | sP11('const.relation.RUNIV_2'(s(X48,X49),s(X50,X51)),X52,X53)) )),
  inference(subsumption_resolution,[],[f2155,f1104])).
fof(f2155,plain,(
  ( ! [X52,X50,X48,X53,X51,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,X52)) | sP11('const.relation.RUNIV_2'(s(X48,X49),s(X50,X51)),X52,X53)) )),
  inference(superposition,[],[f1610,f1447])).
fof(f2173,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (sP11('const.relation.RUNIV_2'(s(X42,X43),s(X44,X45)),X47,X46) | p(s(bool,X46))) )),
  inference(subsumption_resolution,[],[f2154,f1104])).
fof(f2154,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (~p(s(bool,'T_0')) | p(s(bool,X46)) | sP11('const.relation.RUNIV_2'(s(X42,X43),s(X44,X45)),X47,X46)) )),
  inference(superposition,[],[f1609,f1447])).
fof(f2172,plain,(
  ( ! [X39,X37,X41,X38,X36,X40] : (~p(s(bool,X40)) | sP9('const.relation.RUNIV_2'(s(X36,X37),s(X38,X39)),X40,X41)) )),
  inference(subsumption_resolution,[],[f2153,f1104])).
fof(f2153,plain,(
  ( ! [X39,X37,X41,X38,X36,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,X40)) | sP9('const.relation.RUNIV_2'(s(X36,X37),s(X38,X39)),X40,X41)) )),
  inference(superposition,[],[f1593,f1447])).
fof(f2171,plain,(
  ( ! [X35,X33,X31,X34,X32] : (sP1(X35,'const.relation.RUNIV_2'(s(X31,X32),s(X33,X34))) | p(s(bool,X35))) )),
  inference(subsumption_resolution,[],[f2152,f1104])).
fof(f2152,plain,(
  ( ! [X35,X33,X31,X34,X32] : (~p(s(bool,'T_0')) | p(s(bool,X35)) | sP1(X35,'const.relation.RUNIV_2'(s(X31,X32),s(X33,X34)))) )),
  inference(superposition,[],[f1416,f1447])).
fof(f2170,plain,(
  ( ! [X30,X28,X26,X29,X27] : (sP1('const.relation.RUNIV_2'(s(X26,X27),s(X28,X29)),X30) | p(s(bool,X30))) )),
  inference(subsumption_resolution,[],[f2151,f1104])).
fof(f2151,plain,(
  ( ! [X30,X28,X26,X29,X27] : (~p(s(bool,'T_0')) | p(s(bool,X30)) | sP1('const.relation.RUNIV_2'(s(X26,X27),s(X28,X29)),X30)) )),
  inference(superposition,[],[f1413,f1447])).
fof(f2169,plain,(
  ( ! [X24,X23,X21,X25,X22] : (sP0('const.relation.RUNIV_2'(s(X21,X22),s(X23,X24)),X25)) )),
  inference(subsumption_resolution,[],[f2150,f1104])).
fof(f2150,plain,(
  ( ! [X24,X23,X21,X25,X22] : (~p(s(bool,'T_0')) | sP0('const.relation.RUNIV_2'(s(X21,X22),s(X23,X24)),X25)) )),
  inference(superposition,[],[f1398,f1447])).
fof(f2168,plain,(
  ( ! [X19,X17,X20,X18,X16] : (sP0(X20,'const.relation.RUNIV_2'(s(X16,X17),s(X18,X19)))) )),
  inference(subsumption_resolution,[],[f2149,f1104])).
fof(f2149,plain,(
  ( ! [X19,X17,X20,X18,X16] : (~p(s(bool,'T_0')) | sP0(X20,'const.relation.RUNIV_2'(s(X16,X17),s(X18,X19)))) )),
  inference(superposition,[],[f1397,f1447])).
fof(f2095,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,X5)) | ~sP8('const.combin.I_1'(s(bool,X5)),X6,X7) | sP9('const.combin.I_1'(s(bool,X5)),X6,X7)) )),
  inference(superposition,[],[f1677,f1188])).
fof(f2080,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,X5)) | ~sP10('const.combin.I_1'(s(bool,X5)),X6,X7) | sP11('const.combin.I_1'(s(bool,X5)),X6,X7)) )),
  inference(superposition,[],[f1676,f1188])).
fof(f2065,plain,(
  ( ! [X0] : (sP11('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f1104,f1104,f1610])).
fof(f2056,plain,(
  ( ! [X0] : (sP11('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f1700,f1104,f1609])).
fof(f2053,plain,(
  ( ! [X0] : (~sP11('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f1104,f1700,f1598])).
fof(f2052,plain,(
  ( ! [X0] : (~sP11('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f1700,f1700,f1597])).
fof(f2045,plain,(
  ( ! [X0] : (sP9('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f1104,f1104,f1593])).
fof(f2038,plain,(
  ( ! [X0] : (sP9('T_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f1104,f1104,f1592])).
fof(f2036,plain,(
  ( ! [X0] : (~sP9('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f1104,f1700,f1581])).
fof(f2035,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f1104,f1700,f1580])).
fof(f2028,plain,(
  ( ! [X0] : (sP7('F_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f1700,f1700,f1575])).
fof(f2025,plain,(
  ( ! [X0] : (sP7('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f1700,f1700,f1574])).
fof(f2022,plain,(
  ( ! [X0] : (~sP7('T_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f1104,f1700,f1565])).
fof(f2019,plain,(
  ( ! [X0] : (~sP7('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f1104,f1700,f1564])).
fof(f1994,plain,(
  ~sP0('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f1700,f1700,f1851])).
fof(f1985,plain,(
  sP1('F_0','T_0')),
  inference(unit_resulting_resolution,[],[f1700,f1104,f1416])).
fof(f1977,plain,(
  sP1('T_0','F_0')),
  inference(unit_resulting_resolution,[],[f1700,f1104,f1413])).
fof(f1975,plain,(
  ~sP1('T_0','T_0')),
  inference(unit_resulting_resolution,[],[f1104,f1104,f1412])).
fof(f1974,plain,(
  ~sP1('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f1700,f1700,f1411])).
fof(f1957,plain,(
  ( ! [X0] : (sP0(X0,'const.combin.I_1'(s(bool,'T_0')))) )),
  inference(unit_resulting_resolution,[],[f1104,f1935])).
fof(f1946,plain,(
  ( ! [X0] : (sP0('const.combin.I_1'(s(bool,'T_0')),X0)) )),
  inference(unit_resulting_resolution,[],[f1104,f1934])).
fof(f1935,plain,(
  ( ! [X2,X3] : (sP0(X3,'const.combin.I_1'(s(bool,X2))) | ~p(s(bool,X2))) )),
  inference(superposition,[],[f1397,f1188])).
fof(f1934,plain,(
  ( ! [X0,X1] : (sP0('const.combin.I_1'(s(bool,X0)),X1) | ~p(s(bool,X0))) )),
  inference(superposition,[],[f1398,f1188])).
fof(f1927,plain,(
  ( ! [X0,X1] : (sP8('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f1700,f1590])).
fof(f1918,plain,(
  ( ! [X0,X1] : (sP2('T_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f1104,f1558])).
fof(f1915,plain,(
  ( ! [X0,X1] : (sP3('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f1700,f1554])).
fof(f1905,plain,(
  ( ! [X0] : (sP0('T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f1104,f1398])).
fof(f1904,plain,(
  ( ! [X0] : (sP0(X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f1104,f1397])).
fof(f1663,plain,(
  ( ! [X6,X4,X2,X0,X12,X10,X8,X7,X5,X3,X1,X11,X9] : (s(fun(X3,fun(X4,fun(X5,X0))),'_dst_x000x011x110x111_4'(s(fun(X1,fun(X2,fun(X6,X0))),X8),s(fun(X4,fun(X5,X1)),X9),s(fun(X3,fun(X4,X2)),X10),s(fun(X3,fun(X4,fun(X5,X6))),happ(s(fun(X7,fun(X3,fun(X4,fun(X5,X6)))),X11),s(X7,X12))))) = s(fun(X3,fun(X4,fun(X5,X0))),happ(s(fun(X7,fun(X3,fun(X4,fun(X5,X0)))),'_dst_x0000x0011x0110x1111_4'(s(fun(X1,fun(X2,fun(X6,X0))),X8),s(fun(X4,fun(X5,X1)),X9),s(fun(X3,fun(X4,X2)),X10),s(fun(X7,fun(X3,fun(X4,fun(X5,X6)))),X11))),s(X7,X12)))) )),
  inference(cnf_transformation,[],[f722])).
fof(f722,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12] : s(fun(X3,fun(X4,fun(X5,X0))),'_dst_x000x011x110x111_4'(s(fun(X1,fun(X2,fun(X6,X0))),X8),s(fun(X4,fun(X5,X1)),X9),s(fun(X3,fun(X4,X2)),X10),s(fun(X3,fun(X4,fun(X5,X6))),happ(s(fun(X7,fun(X3,fun(X4,fun(X5,X6)))),X11),s(X7,X12))))) = s(fun(X3,fun(X4,fun(X5,X0))),happ(s(fun(X7,fun(X3,fun(X4,fun(X5,X0)))),'_dst_x0000x0011x0110x1111_4'(s(fun(X1,fun(X2,fun(X6,X0))),X8),s(fun(X4,fun(X5,X1)),X9),s(fun(X3,fun(X4,X2)),X10),s(fun(X7,fun(X3,fun(X4,fun(X5,X6)))),X11))),s(X7,X12)))),
  inference(flattening,[],[f721])).
fof(f721,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ! [X8,X9,X10,X11,X12] : s(fun(X3,fun(X4,fun(X5,X0))),'_dst_x000x011x110x111_4'(s(fun(X1,fun(X2,fun(X6,X0))),X8),s(fun(X4,fun(X5,X1)),X9),s(fun(X3,fun(X4,X2)),X10),s(fun(X3,fun(X4,fun(X5,X6))),happ(s(fun(X7,fun(X3,fun(X4,fun(X5,X6)))),X11),s(X7,X12))))) = s(fun(X3,fun(X4,fun(X5,X0))),happ(s(fun(X7,fun(X3,fun(X4,fun(X5,X0)))),'_dst_x0000x0011x0110x1111_4'(s(fun(X1,fun(X2,fun(X6,X0))),X8),s(fun(X4,fun(X5,X1)),X9),s(fun(X3,fun(X4,X2)),X10),s(fun(X7,fun(X3,fun(X4,fun(X5,X6)))),X11))),s(X7,X12)))),
  inference(rectify,[],[f183])).
fof(f183,axiom,(
  ! [X35,X57,X58,X36,X37,X56,X59,X38] : ! [X39,X40,X41,X63,X42] : s(fun(X36,fun(X37,fun(X56,X35))),'_dst_x000x011x110x111_4'(s(fun(X57,fun(X58,fun(X59,X35))),X39),s(fun(X37,fun(X56,X57)),X40),s(fun(X36,fun(X37,X58)),X41),s(fun(X36,fun(X37,fun(X56,X59))),happ(s(fun(X38,fun(X36,fun(X37,fun(X56,X59)))),X63),s(X38,X42))))) = s(fun(X36,fun(X37,fun(X56,X35))),happ(s(fun(X38,fun(X36,fun(X37,fun(X56,X35)))),'_dst_x0000x0011x0110x1111_4'(s(fun(X57,fun(X58,fun(X59,X35))),X39),s(fun(X37,fun(X56,X57)),X40),s(fun(X36,fun(X37,X58)),X41),s(fun(X38,fun(X36,fun(X37,fun(X56,X59)))),X63))),s(X38,X42)))),
  file('thm.pred_set.EXTENSION','_dst_x0000x0011x0110x1111_ax')).
fof(f1662,plain,(
  ( ! [X6,X4,X2,X0,X12,X10,X8,X7,X5,X3,X1,X11,X9] : (s(fun(X3,fun(X4,fun(X5,X0))),'_dst_x000x011x100x111_4'(s(fun(X1,fun(X2,fun(X6,X0))),X8),s(fun(X4,fun(X5,X1)),X9),s(fun(X3,X2),X10),s(fun(X3,fun(X4,fun(X5,X6))),happ(s(fun(X7,fun(X3,fun(X4,fun(X5,X6)))),X11),s(X7,X12))))) = s(fun(X3,fun(X4,fun(X5,X0))),happ(s(fun(X7,fun(X3,fun(X4,fun(X5,X0)))),'_dst_x0000x0011x0100x1111_4'(s(fun(X1,fun(X2,fun(X6,X0))),X8),s(fun(X4,fun(X5,X1)),X9),s(fun(X3,X2),X10),s(fun(X7,fun(X3,fun(X4,fun(X5,X6)))),X11))),s(X7,X12)))) )),
  inference(cnf_transformation,[],[f720])).
fof(f720,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12] : s(fun(X3,fun(X4,fun(X5,X0))),'_dst_x000x011x100x111_4'(s(fun(X1,fun(X2,fun(X6,X0))),X8),s(fun(X4,fun(X5,X1)),X9),s(fun(X3,X2),X10),s(fun(X3,fun(X4,fun(X5,X6))),happ(s(fun(X7,fun(X3,fun(X4,fun(X5,X6)))),X11),s(X7,X12))))) = s(fun(X3,fun(X4,fun(X5,X0))),happ(s(fun(X7,fun(X3,fun(X4,fun(X5,X0)))),'_dst_x0000x0011x0100x1111_4'(s(fun(X1,fun(X2,fun(X6,X0))),X8),s(fun(X4,fun(X5,X1)),X9),s(fun(X3,X2),X10),s(fun(X7,fun(X3,fun(X4,fun(X5,X6)))),X11))),s(X7,X12)))),
  inference(flattening,[],[f719])).
fof(f719,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ! [X8,X9,X10,X11,X12] : s(fun(X3,fun(X4,fun(X5,X0))),'_dst_x000x011x100x111_4'(s(fun(X1,fun(X2,fun(X6,X0))),X8),s(fun(X4,fun(X5,X1)),X9),s(fun(X3,X2),X10),s(fun(X3,fun(X4,fun(X5,X6))),happ(s(fun(X7,fun(X3,fun(X4,fun(X5,X6)))),X11),s(X7,X12))))) = s(fun(X3,fun(X4,fun(X5,X0))),happ(s(fun(X7,fun(X3,fun(X4,fun(X5,X0)))),'_dst_x0000x0011x0100x1111_4'(s(fun(X1,fun(X2,fun(X6,X0))),X8),s(fun(X4,fun(X5,X1)),X9),s(fun(X3,X2),X10),s(fun(X7,fun(X3,fun(X4,fun(X5,X6)))),X11))),s(X7,X12)))),
  inference(rectify,[],[f174])).
fof(f174,axiom,(
  ! [X35,X57,X58,X36,X37,X56,X59,X38] : ! [X39,X40,X41,X63,X42] : s(fun(X36,fun(X37,fun(X56,X35))),'_dst_x000x011x100x111_4'(s(fun(X57,fun(X58,fun(X59,X35))),X39),s(fun(X37,fun(X56,X57)),X40),s(fun(X36,X58),X41),s(fun(X36,fun(X37,fun(X56,X59))),happ(s(fun(X38,fun(X36,fun(X37,fun(X56,X59)))),X63),s(X38,X42))))) = s(fun(X36,fun(X37,fun(X56,X35))),happ(s(fun(X38,fun(X36,fun(X37,fun(X56,X35)))),'_dst_x0000x0011x0100x1111_4'(s(fun(X57,fun(X58,fun(X59,X35))),X39),s(fun(X37,fun(X56,X57)),X40),s(fun(X36,X58),X41),s(fun(X38,fun(X36,fun(X37,fun(X56,X59)))),X63))),s(X38,X42)))),
  file('thm.pred_set.EXTENSION','_dst_x0000x0011x0100x1111_ax')).
fof(f1661,plain,(
  ( ! [X6,X4,X2,X0,X12,X10,X8,X7,X5,X3,X1,X11,X9] : (s(fun(X3,fun(X4,fun(X5,X0))),'_dst_x000x001x100x111_4'(s(fun(X1,fun(X2,fun(X6,X0))),X8),s(fun(X5,X1),X9),s(fun(X3,X2),X10),s(fun(X3,fun(X4,fun(X5,X6))),happ(s(fun(X7,fun(X3,fun(X4,fun(X5,X6)))),X11),s(X7,X12))))) = s(fun(X3,fun(X4,fun(X5,X0))),happ(s(fun(X7,fun(X3,fun(X4,fun(X5,X0)))),'_dst_x0000x0001x0100x1111_4'(s(fun(X1,fun(X2,fun(X6,X0))),X8),s(fun(X5,X1),X9),s(fun(X3,X2),X10),s(fun(X7,fun(X3,fun(X4,fun(X5,X6)))),X11))),s(X7,X12)))) )),
  inference(cnf_transformation,[],[f718])).
fof(f718,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12] : s(fun(X3,fun(X4,fun(X5,X0))),'_dst_x000x001x100x111_4'(s(fun(X1,fun(X2,fun(X6,X0))),X8),s(fun(X5,X1),X9),s(fun(X3,X2),X10),s(fun(X3,fun(X4,fun(X5,X6))),happ(s(fun(X7,fun(X3,fun(X4,fun(X5,X6)))),X11),s(X7,X12))))) = s(fun(X3,fun(X4,fun(X5,X0))),happ(s(fun(X7,fun(X3,fun(X4,fun(X5,X0)))),'_dst_x0000x0001x0100x1111_4'(s(fun(X1,fun(X2,fun(X6,X0))),X8),s(fun(X5,X1),X9),s(fun(X3,X2),X10),s(fun(X7,fun(X3,fun(X4,fun(X5,X6)))),X11))),s(X7,X12)))),
  inference(flattening,[],[f717])).
fof(f717,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ! [X8,X9,X10,X11,X12] : s(fun(X3,fun(X4,fun(X5,X0))),'_dst_x000x001x100x111_4'(s(fun(X1,fun(X2,fun(X6,X0))),X8),s(fun(X5,X1),X9),s(fun(X3,X2),X10),s(fun(X3,fun(X4,fun(X5,X6))),happ(s(fun(X7,fun(X3,fun(X4,fun(X5,X6)))),X11),s(X7,X12))))) = s(fun(X3,fun(X4,fun(X5,X0))),happ(s(fun(X7,fun(X3,fun(X4,fun(X5,X0)))),'_dst_x0000x0001x0100x1111_4'(s(fun(X1,fun(X2,fun(X6,X0))),X8),s(fun(X5,X1),X9),s(fun(X3,X2),X10),s(fun(X7,fun(X3,fun(X4,fun(X5,X6)))),X11))),s(X7,X12)))),
  inference(rectify,[],[f177])).
fof(f177,axiom,(
  ! [X35,X57,X58,X36,X37,X56,X59,X38] : ! [X39,X40,X41,X63,X42] : s(fun(X36,fun(X37,fun(X56,X35))),'_dst_x000x001x100x111_4'(s(fun(X57,fun(X58,fun(X59,X35))),X39),s(fun(X56,X57),X40),s(fun(X36,X58),X41),s(fun(X36,fun(X37,fun(X56,X59))),happ(s(fun(X38,fun(X36,fun(X37,fun(X56,X59)))),X63),s(X38,X42))))) = s(fun(X36,fun(X37,fun(X56,X35))),happ(s(fun(X38,fun(X36,fun(X37,fun(X56,X35)))),'_dst_x0000x0001x0100x1111_4'(s(fun(X57,fun(X58,fun(X59,X35))),X39),s(fun(X56,X57),X40),s(fun(X36,X58),X41),s(fun(X38,fun(X36,fun(X37,fun(X56,X59)))),X63))),s(X38,X42)))),
  file('thm.pred_set.EXTENSION','_dst_x0000x0001x0100x1111_ax')).
fof(f1660,plain,(
  ( ! [X6,X4,X2,X0,X12,X10,X8,X7,X5,X3,X1,X11,X9] : (s(fun(X3,fun(X4,fun(X5,X0))),'_dst_x000x001x000x111_4'(s(fun(X1,fun(X2,fun(X6,X0))),X8),s(fun(X5,X1),X9),s(X2,X10),s(fun(X3,fun(X4,fun(X5,X6))),happ(s(fun(X7,fun(X3,fun(X4,fun(X5,X6)))),X11),s(X7,X12))))) = s(fun(X3,fun(X4,fun(X5,X0))),happ(s(fun(X7,fun(X3,fun(X4,fun(X5,X0)))),'_dst_x0000x0001x0000x1111_4'(s(fun(X1,fun(X2,fun(X6,X0))),X8),s(fun(X5,X1),X9),s(X2,X10),s(fun(X7,fun(X3,fun(X4,fun(X5,X6)))),X11))),s(X7,X12)))) )),
  inference(cnf_transformation,[],[f716])).
fof(f716,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12] : s(fun(X3,fun(X4,fun(X5,X0))),'_dst_x000x001x000x111_4'(s(fun(X1,fun(X2,fun(X6,X0))),X8),s(fun(X5,X1),X9),s(X2,X10),s(fun(X3,fun(X4,fun(X5,X6))),happ(s(fun(X7,fun(X3,fun(X4,fun(X5,X6)))),X11),s(X7,X12))))) = s(fun(X3,fun(X4,fun(X5,X0))),happ(s(fun(X7,fun(X3,fun(X4,fun(X5,X0)))),'_dst_x0000x0001x0000x1111_4'(s(fun(X1,fun(X2,fun(X6,X0))),X8),s(fun(X5,X1),X9),s(X2,X10),s(fun(X7,fun(X3,fun(X4,fun(X5,X6)))),X11))),s(X7,X12)))),
  inference(flattening,[],[f715])).
fof(f715,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ! [X8,X9,X10,X11,X12] : s(fun(X3,fun(X4,fun(X5,X0))),'_dst_x000x001x000x111_4'(s(fun(X1,fun(X2,fun(X6,X0))),X8),s(fun(X5,X1),X9),s(X2,X10),s(fun(X3,fun(X4,fun(X5,X6))),happ(s(fun(X7,fun(X3,fun(X4,fun(X5,X6)))),X11),s(X7,X12))))) = s(fun(X3,fun(X4,fun(X5,X0))),happ(s(fun(X7,fun(X3,fun(X4,fun(X5,X0)))),'_dst_x0000x0001x0000x1111_4'(s(fun(X1,fun(X2,fun(X6,X0))),X8),s(fun(X5,X1),X9),s(X2,X10),s(fun(X7,fun(X3,fun(X4,fun(X5,X6)))),X11))),s(X7,X12)))),
  inference(rectify,[],[f201])).
fof(f201,axiom,(
  ! [X35,X57,X58,X36,X37,X56,X59,X38] : ! [X39,X40,X41,X63,X42] : s(fun(X36,fun(X37,fun(X56,X35))),'_dst_x000x001x000x111_4'(s(fun(X57,fun(X58,fun(X59,X35))),X39),s(fun(X56,X57),X40),s(X58,X41),s(fun(X36,fun(X37,fun(X56,X59))),happ(s(fun(X38,fun(X36,fun(X37,fun(X56,X59)))),X63),s(X38,X42))))) = s(fun(X36,fun(X37,fun(X56,X35))),happ(s(fun(X38,fun(X36,fun(X37,fun(X56,X35)))),'_dst_x0000x0001x0000x1111_4'(s(fun(X57,fun(X58,fun(X59,X35))),X39),s(fun(X56,X57),X40),s(X58,X41),s(fun(X38,fun(X36,fun(X37,fun(X56,X59)))),X63))),s(X38,X42)))),
  file('thm.pred_set.EXTENSION','_dst_x0000x0001x0000x1111_ax')).
fof(f1659,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (s(fun(X3,fun(X4,fun(X5,fun(X6,X0)))),'_dst_x0000x0011x1111_3'(s(fun(X1,fun(X7,X0)),X8),s(fun(X5,fun(X6,X1)),happ(s(fun(X2,fun(X5,fun(X6,X1))),X9),s(X2,X11))),s(fun(X3,fun(X4,fun(X5,fun(X6,X7)))),X10))) = s(fun(X3,fun(X4,fun(X5,fun(X6,X0)))),happ(s(fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,X0))))),'_dst_x00000x10011x01111_3'(s(fun(X1,fun(X7,X0)),X8),s(fun(X2,fun(X5,fun(X6,X1))),X9),s(fun(X3,fun(X4,fun(X5,fun(X6,X7)))),X10))),s(X2,X11)))) )),
  inference(cnf_transformation,[],[f714])).
fof(f714,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : s(fun(X3,fun(X4,fun(X5,fun(X6,X0)))),'_dst_x0000x0011x1111_3'(s(fun(X1,fun(X7,X0)),X8),s(fun(X5,fun(X6,X1)),happ(s(fun(X2,fun(X5,fun(X6,X1))),X9),s(X2,X11))),s(fun(X3,fun(X4,fun(X5,fun(X6,X7)))),X10))) = s(fun(X3,fun(X4,fun(X5,fun(X6,X0)))),happ(s(fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,X0))))),'_dst_x00000x10011x01111_3'(s(fun(X1,fun(X7,X0)),X8),s(fun(X2,fun(X5,fun(X6,X1))),X9),s(fun(X3,fun(X4,fun(X5,fun(X6,X7)))),X10))),s(X2,X11)))),
  inference(flattening,[],[f713])).
fof(f713,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ! [X8,X9,X10,X11] : s(fun(X3,fun(X4,fun(X5,fun(X6,X0)))),'_dst_x0000x0011x1111_3'(s(fun(X1,fun(X7,X0)),X8),s(fun(X5,fun(X6,X1)),happ(s(fun(X2,fun(X5,fun(X6,X1))),X9),s(X2,X11))),s(fun(X3,fun(X4,fun(X5,fun(X6,X7)))),X10))) = s(fun(X3,fun(X4,fun(X5,fun(X6,X0)))),happ(s(fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,X0))))),'_dst_x00000x10011x01111_3'(s(fun(X1,fun(X7,X0)),X8),s(fun(X2,fun(X5,fun(X6,X1))),X9),s(fun(X3,fun(X4,fun(X5,fun(X6,X7)))),X10))),s(X2,X11)))),
  inference(rectify,[],[f117])).
fof(f117,axiom,(
  ! [X35,X58,X38,X36,X37,X56,X57,X59] : ! [X39,X40,X41,X42] : s(fun(X36,fun(X37,fun(X56,fun(X57,X35)))),'_dst_x0000x0011x1111_3'(s(fun(X58,fun(X59,X35)),X39),s(fun(X56,fun(X57,X58)),happ(s(fun(X38,fun(X56,fun(X57,X58))),X40),s(X38,X42))),s(fun(X36,fun(X37,fun(X56,fun(X57,X59)))),X41))) = s(fun(X36,fun(X37,fun(X56,fun(X57,X35)))),happ(s(fun(X38,fun(X36,fun(X37,fun(X56,fun(X57,X35))))),'_dst_x00000x10011x01111_3'(s(fun(X58,fun(X59,X35)),X39),s(fun(X38,fun(X56,fun(X57,X58))),X40),s(fun(X36,fun(X37,fun(X56,fun(X57,X59)))),X41))),s(X38,X42)))),
  file('thm.pred_set.EXTENSION','_dst_x00000x10011x01111_ax')).
fof(f1658,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (s(fun(X3,fun(X4,X0)),happ(s(fun(X6,fun(X3,fun(X4,X0))),'_dst_x000x011x110x111_4'(s(fun(X1,fun(X2,fun(X5,X0))),X7),s(fun(X3,fun(X4,X1)),X8),s(fun(X6,fun(X3,X2)),X9),s(fun(X6,fun(X3,fun(X4,X5))),X10))),s(X6,X11))) = s(fun(X3,fun(X4,X0)),'_dst_x00x11x10x11_4'(s(fun(X1,fun(X2,fun(X5,X0))),X7),s(fun(X3,fun(X4,X1)),X8),s(fun(X3,X2),happ(s(fun(X6,fun(X3,X2)),X9),s(X6,X11))),s(fun(X3,fun(X4,X5)),happ(s(fun(X6,fun(X3,fun(X4,X5))),X10),s(X6,X11)))))) )),
  inference(cnf_transformation,[],[f712])).
fof(f712,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : s(fun(X3,fun(X4,X0)),happ(s(fun(X6,fun(X3,fun(X4,X0))),'_dst_x000x011x110x111_4'(s(fun(X1,fun(X2,fun(X5,X0))),X7),s(fun(X3,fun(X4,X1)),X8),s(fun(X6,fun(X3,X2)),X9),s(fun(X6,fun(X3,fun(X4,X5))),X10))),s(X6,X11))) = s(fun(X3,fun(X4,X0)),'_dst_x00x11x10x11_4'(s(fun(X1,fun(X2,fun(X5,X0))),X7),s(fun(X3,fun(X4,X1)),X8),s(fun(X3,X2),happ(s(fun(X6,fun(X3,X2)),X9),s(X6,X11))),s(fun(X3,fun(X4,X5)),happ(s(fun(X6,fun(X3,fun(X4,X5))),X10),s(X6,X11)))))),
  inference(flattening,[],[f711])).
fof(f711,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ! [X7,X8,X9,X10,X11] : s(fun(X3,fun(X4,X0)),happ(s(fun(X6,fun(X3,fun(X4,X0))),'_dst_x000x011x110x111_4'(s(fun(X1,fun(X2,fun(X5,X0))),X7),s(fun(X3,fun(X4,X1)),X8),s(fun(X6,fun(X3,X2)),X9),s(fun(X6,fun(X3,fun(X4,X5))),X10))),s(X6,X11))) = s(fun(X3,fun(X4,X0)),'_dst_x00x11x10x11_4'(s(fun(X1,fun(X2,fun(X5,X0))),X7),s(fun(X3,fun(X4,X1)),X8),s(fun(X3,X2),happ(s(fun(X6,fun(X3,X2)),X9),s(X6,X11))),s(fun(X3,fun(X4,X5)),happ(s(fun(X6,fun(X3,fun(X4,X5))),X10),s(X6,X11)))))),
  inference(rectify,[],[f187])).
fof(f187,axiom,(
  ! [X35,X56,X57,X36,X37,X58,X38] : ! [X39,X40,X41,X63,X42] : s(fun(X36,fun(X37,X35)),happ(s(fun(X38,fun(X36,fun(X37,X35))),'_dst_x000x011x110x111_4'(s(fun(X56,fun(X57,fun(X58,X35))),X39),s(fun(X36,fun(X37,X56)),X40),s(fun(X38,fun(X36,X57)),X41),s(fun(X38,fun(X36,fun(X37,X58))),X63))),s(X38,X42))) = s(fun(X36,fun(X37,X35)),'_dst_x00x11x10x11_4'(s(fun(X56,fun(X57,fun(X58,X35))),X39),s(fun(X36,fun(X37,X56)),X40),s(fun(X36,X57),happ(s(fun(X38,fun(X36,X57)),X41),s(X38,X42))),s(fun(X36,fun(X37,X58)),happ(s(fun(X38,fun(X36,fun(X37,X58))),X63),s(X38,X42)))))),
  file('thm.pred_set.EXTENSION','_dst_x000x011x110x111_ax')).
fof(f1657,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (s(fun(X3,fun(X4,X0)),happ(s(fun(X6,fun(X3,fun(X4,X0))),'_dst_x000x011x100x111_4'(s(fun(X1,fun(X2,fun(X5,X0))),X7),s(fun(X3,fun(X4,X1)),X8),s(fun(X6,X2),X9),s(fun(X6,fun(X3,fun(X4,X5))),X10))),s(X6,X11))) = s(fun(X3,fun(X4,X0)),'_dst_x00x11x00x11_4'(s(fun(X1,fun(X2,fun(X5,X0))),X7),s(fun(X3,fun(X4,X1)),X8),s(X2,happ(s(fun(X6,X2),X9),s(X6,X11))),s(fun(X3,fun(X4,X5)),happ(s(fun(X6,fun(X3,fun(X4,X5))),X10),s(X6,X11)))))) )),
  inference(cnf_transformation,[],[f710])).
fof(f710,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : s(fun(X3,fun(X4,X0)),happ(s(fun(X6,fun(X3,fun(X4,X0))),'_dst_x000x011x100x111_4'(s(fun(X1,fun(X2,fun(X5,X0))),X7),s(fun(X3,fun(X4,X1)),X8),s(fun(X6,X2),X9),s(fun(X6,fun(X3,fun(X4,X5))),X10))),s(X6,X11))) = s(fun(X3,fun(X4,X0)),'_dst_x00x11x00x11_4'(s(fun(X1,fun(X2,fun(X5,X0))),X7),s(fun(X3,fun(X4,X1)),X8),s(X2,happ(s(fun(X6,X2),X9),s(X6,X11))),s(fun(X3,fun(X4,X5)),happ(s(fun(X6,fun(X3,fun(X4,X5))),X10),s(X6,X11)))))),
  inference(flattening,[],[f709])).
fof(f709,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ! [X7,X8,X9,X10,X11] : s(fun(X3,fun(X4,X0)),happ(s(fun(X6,fun(X3,fun(X4,X0))),'_dst_x000x011x100x111_4'(s(fun(X1,fun(X2,fun(X5,X0))),X7),s(fun(X3,fun(X4,X1)),X8),s(fun(X6,X2),X9),s(fun(X6,fun(X3,fun(X4,X5))),X10))),s(X6,X11))) = s(fun(X3,fun(X4,X0)),'_dst_x00x11x00x11_4'(s(fun(X1,fun(X2,fun(X5,X0))),X7),s(fun(X3,fun(X4,X1)),X8),s(X2,happ(s(fun(X6,X2),X9),s(X6,X11))),s(fun(X3,fun(X4,X5)),happ(s(fun(X6,fun(X3,fun(X4,X5))),X10),s(X6,X11)))))),
  inference(rectify,[],[f171])).
fof(f171,axiom,(
  ! [X35,X56,X57,X36,X37,X58,X38] : ! [X39,X40,X41,X63,X42] : s(fun(X36,fun(X37,X35)),happ(s(fun(X38,fun(X36,fun(X37,X35))),'_dst_x000x011x100x111_4'(s(fun(X56,fun(X57,fun(X58,X35))),X39),s(fun(X36,fun(X37,X56)),X40),s(fun(X38,X57),X41),s(fun(X38,fun(X36,fun(X37,X58))),X63))),s(X38,X42))) = s(fun(X36,fun(X37,X35)),'_dst_x00x11x00x11_4'(s(fun(X56,fun(X57,fun(X58,X35))),X39),s(fun(X36,fun(X37,X56)),X40),s(X57,happ(s(fun(X38,X57),X41),s(X38,X42))),s(fun(X36,fun(X37,X58)),happ(s(fun(X38,fun(X36,fun(X37,X58))),X63),s(X38,X42)))))),
  file('thm.pred_set.EXTENSION','_dst_x000x011x100x111_ax')).
fof(f1656,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (s(fun(X3,fun(X4,X0)),happ(s(fun(X6,fun(X3,fun(X4,X0))),'_dst_x000x001x100x111_4'(s(fun(X1,fun(X2,fun(X5,X0))),X7),s(fun(X4,X1),X8),s(fun(X6,X2),X9),s(fun(X6,fun(X3,fun(X4,X5))),X10))),s(X6,X11))) = s(fun(X3,fun(X4,X0)),'_dst_x00x01x00x11_4'(s(fun(X1,fun(X2,fun(X5,X0))),X7),s(fun(X4,X1),X8),s(X2,happ(s(fun(X6,X2),X9),s(X6,X11))),s(fun(X3,fun(X4,X5)),happ(s(fun(X6,fun(X3,fun(X4,X5))),X10),s(X6,X11)))))) )),
  inference(cnf_transformation,[],[f708])).
fof(f708,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : s(fun(X3,fun(X4,X0)),happ(s(fun(X6,fun(X3,fun(X4,X0))),'_dst_x000x001x100x111_4'(s(fun(X1,fun(X2,fun(X5,X0))),X7),s(fun(X4,X1),X8),s(fun(X6,X2),X9),s(fun(X6,fun(X3,fun(X4,X5))),X10))),s(X6,X11))) = s(fun(X3,fun(X4,X0)),'_dst_x00x01x00x11_4'(s(fun(X1,fun(X2,fun(X5,X0))),X7),s(fun(X4,X1),X8),s(X2,happ(s(fun(X6,X2),X9),s(X6,X11))),s(fun(X3,fun(X4,X5)),happ(s(fun(X6,fun(X3,fun(X4,X5))),X10),s(X6,X11)))))),
  inference(flattening,[],[f707])).
fof(f707,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ! [X7,X8,X9,X10,X11] : s(fun(X3,fun(X4,X0)),happ(s(fun(X6,fun(X3,fun(X4,X0))),'_dst_x000x001x100x111_4'(s(fun(X1,fun(X2,fun(X5,X0))),X7),s(fun(X4,X1),X8),s(fun(X6,X2),X9),s(fun(X6,fun(X3,fun(X4,X5))),X10))),s(X6,X11))) = s(fun(X3,fun(X4,X0)),'_dst_x00x01x00x11_4'(s(fun(X1,fun(X2,fun(X5,X0))),X7),s(fun(X4,X1),X8),s(X2,happ(s(fun(X6,X2),X9),s(X6,X11))),s(fun(X3,fun(X4,X5)),happ(s(fun(X6,fun(X3,fun(X4,X5))),X10),s(X6,X11)))))),
  inference(rectify,[],[f170])).
fof(f170,axiom,(
  ! [X35,X56,X57,X36,X37,X58,X38] : ! [X39,X40,X41,X63,X42] : s(fun(X36,fun(X37,X35)),happ(s(fun(X38,fun(X36,fun(X37,X35))),'_dst_x000x001x100x111_4'(s(fun(X56,fun(X57,fun(X58,X35))),X39),s(fun(X37,X56),X40),s(fun(X38,X57),X41),s(fun(X38,fun(X36,fun(X37,X58))),X63))),s(X38,X42))) = s(fun(X36,fun(X37,X35)),'_dst_x00x01x00x11_4'(s(fun(X56,fun(X57,fun(X58,X35))),X39),s(fun(X37,X56),X40),s(X57,happ(s(fun(X38,X57),X41),s(X38,X42))),s(fun(X36,fun(X37,X58)),happ(s(fun(X38,fun(X36,fun(X37,X58))),X63),s(X38,X42)))))),
  file('thm.pred_set.EXTENSION','_dst_x000x001x100x111_ax')).
fof(f1655,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (s(fun(X3,fun(X4,X0)),'_dst_x00x01x00x11_4'(s(fun(X1,fun(X2,fun(X5,X0))),X7),s(fun(X4,X1),X8),s(X2,X9),s(fun(X3,fun(X4,X5)),happ(s(fun(X6,fun(X3,fun(X4,X5))),X10),s(X6,X11))))) = s(fun(X3,fun(X4,X0)),happ(s(fun(X6,fun(X3,fun(X4,X0))),'_dst_x000x001x000x111_4'(s(fun(X1,fun(X2,fun(X5,X0))),X7),s(fun(X4,X1),X8),s(X2,X9),s(fun(X6,fun(X3,fun(X4,X5))),X10))),s(X6,X11)))) )),
  inference(cnf_transformation,[],[f706])).
fof(f706,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : s(fun(X3,fun(X4,X0)),'_dst_x00x01x00x11_4'(s(fun(X1,fun(X2,fun(X5,X0))),X7),s(fun(X4,X1),X8),s(X2,X9),s(fun(X3,fun(X4,X5)),happ(s(fun(X6,fun(X3,fun(X4,X5))),X10),s(X6,X11))))) = s(fun(X3,fun(X4,X0)),happ(s(fun(X6,fun(X3,fun(X4,X0))),'_dst_x000x001x000x111_4'(s(fun(X1,fun(X2,fun(X5,X0))),X7),s(fun(X4,X1),X8),s(X2,X9),s(fun(X6,fun(X3,fun(X4,X5))),X10))),s(X6,X11)))),
  inference(flattening,[],[f705])).
fof(f705,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ! [X7,X8,X9,X10,X11] : s(fun(X3,fun(X4,X0)),'_dst_x00x01x00x11_4'(s(fun(X1,fun(X2,fun(X5,X0))),X7),s(fun(X4,X1),X8),s(X2,X9),s(fun(X3,fun(X4,X5)),happ(s(fun(X6,fun(X3,fun(X4,X5))),X10),s(X6,X11))))) = s(fun(X3,fun(X4,X0)),happ(s(fun(X6,fun(X3,fun(X4,X0))),'_dst_x000x001x000x111_4'(s(fun(X1,fun(X2,fun(X5,X0))),X7),s(fun(X4,X1),X8),s(X2,X9),s(fun(X6,fun(X3,fun(X4,X5))),X10))),s(X6,X11)))),
  inference(rectify,[],[f188])).
fof(f188,axiom,(
  ! [X35,X56,X57,X36,X37,X58,X38] : ! [X39,X40,X41,X63,X42] : s(fun(X36,fun(X37,X35)),'_dst_x00x01x00x11_4'(s(fun(X56,fun(X57,fun(X58,X35))),X39),s(fun(X37,X56),X40),s(X57,X41),s(fun(X36,fun(X37,X58)),happ(s(fun(X38,fun(X36,fun(X37,X58))),X63),s(X38,X42))))) = s(fun(X36,fun(X37,X35)),happ(s(fun(X38,fun(X36,fun(X37,X35))),'_dst_x000x001x000x111_4'(s(fun(X56,fun(X57,fun(X58,X35))),X39),s(fun(X37,X56),X40),s(X57,X41),s(fun(X38,fun(X36,fun(X37,X58))),X63))),s(X38,X42)))),
  file('thm.pred_set.EXTENSION','_dst_x000x001x000x111_ax')).
fof(f1654,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (s(fun(X2,fun(X3,fun(X4,X0))),'_dst_x000x011x111_3'(s(fun(X1,fun(X5,X0)),X7),s(fun(X3,fun(X4,X1)),X8),s(fun(X2,fun(X3,fun(X4,X5))),happ(s(fun(X6,fun(X2,fun(X3,fun(X4,X5)))),X9),s(X6,X10))))) = s(fun(X2,fun(X3,fun(X4,X0))),happ(s(fun(X6,fun(X2,fun(X3,fun(X4,X0)))),'_dst_x0000x0011x1111_3'(s(fun(X1,fun(X5,X0)),X7),s(fun(X3,fun(X4,X1)),X8),s(fun(X6,fun(X2,fun(X3,fun(X4,X5)))),X9))),s(X6,X10)))) )),
  inference(cnf_transformation,[],[f704])).
fof(f704,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : s(fun(X2,fun(X3,fun(X4,X0))),'_dst_x000x011x111_3'(s(fun(X1,fun(X5,X0)),X7),s(fun(X3,fun(X4,X1)),X8),s(fun(X2,fun(X3,fun(X4,X5))),happ(s(fun(X6,fun(X2,fun(X3,fun(X4,X5)))),X9),s(X6,X10))))) = s(fun(X2,fun(X3,fun(X4,X0))),happ(s(fun(X6,fun(X2,fun(X3,fun(X4,X0)))),'_dst_x0000x0011x1111_3'(s(fun(X1,fun(X5,X0)),X7),s(fun(X3,fun(X4,X1)),X8),s(fun(X6,fun(X2,fun(X3,fun(X4,X5)))),X9))),s(X6,X10)))),
  inference(flattening,[],[f703])).
fof(f703,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ! [X7,X8,X9,X10] : s(fun(X2,fun(X3,fun(X4,X0))),'_dst_x000x011x111_3'(s(fun(X1,fun(X5,X0)),X7),s(fun(X3,fun(X4,X1)),X8),s(fun(X2,fun(X3,fun(X4,X5))),happ(s(fun(X6,fun(X2,fun(X3,fun(X4,X5)))),X9),s(X6,X10))))) = s(fun(X2,fun(X3,fun(X4,X0))),happ(s(fun(X6,fun(X2,fun(X3,fun(X4,X0)))),'_dst_x0000x0011x1111_3'(s(fun(X1,fun(X5,X0)),X7),s(fun(X3,fun(X4,X1)),X8),s(fun(X6,fun(X2,fun(X3,fun(X4,X5)))),X9))),s(X6,X10)))),
  inference(rectify,[],[f120])).
fof(f120,axiom,(
  ! [X35,X57,X36,X37,X56,X58,X38] : ! [X39,X40,X41,X42] : s(fun(X36,fun(X37,fun(X56,X35))),'_dst_x000x011x111_3'(s(fun(X57,fun(X58,X35)),X39),s(fun(X37,fun(X56,X57)),X40),s(fun(X36,fun(X37,fun(X56,X58))),happ(s(fun(X38,fun(X36,fun(X37,fun(X56,X58)))),X41),s(X38,X42))))) = s(fun(X36,fun(X37,fun(X56,X35))),happ(s(fun(X38,fun(X36,fun(X37,fun(X56,X35)))),'_dst_x0000x0011x1111_3'(s(fun(X57,fun(X58,X35)),X39),s(fun(X37,fun(X56,X57)),X40),s(fun(X38,fun(X36,fun(X37,fun(X56,X58)))),X41))),s(X38,X42)))),
  file('thm.pred_set.EXTENSION','_dst_x0000x0011x1111_ax')).
fof(f1653,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (s(fun(X1,fun(X4,fun(X5,X0))),'_dst_x000x111x011_3'(s(fun(X2,fun(X6,X0)),X7),s(fun(X1,fun(X4,fun(X5,X2))),happ(s(fun(X3,fun(X1,fun(X4,fun(X5,X2)))),X8),s(X3,X10))),s(fun(X4,fun(X5,X6)),X9))) = s(fun(X1,fun(X4,fun(X5,X0))),happ(s(fun(X3,fun(X1,fun(X4,fun(X5,X0)))),'_dst_x0000x1111x0011_3'(s(fun(X2,fun(X6,X0)),X7),s(fun(X3,fun(X1,fun(X4,fun(X5,X2)))),X8),s(fun(X4,fun(X5,X6)),X9))),s(X3,X10)))) )),
  inference(cnf_transformation,[],[f702])).
fof(f702,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : s(fun(X1,fun(X4,fun(X5,X0))),'_dst_x000x111x011_3'(s(fun(X2,fun(X6,X0)),X7),s(fun(X1,fun(X4,fun(X5,X2))),happ(s(fun(X3,fun(X1,fun(X4,fun(X5,X2)))),X8),s(X3,X10))),s(fun(X4,fun(X5,X6)),X9))) = s(fun(X1,fun(X4,fun(X5,X0))),happ(s(fun(X3,fun(X1,fun(X4,fun(X5,X0)))),'_dst_x0000x1111x0011_3'(s(fun(X2,fun(X6,X0)),X7),s(fun(X3,fun(X1,fun(X4,fun(X5,X2)))),X8),s(fun(X4,fun(X5,X6)),X9))),s(X3,X10)))),
  inference(flattening,[],[f701])).
fof(f701,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ! [X7,X8,X9,X10] : s(fun(X1,fun(X4,fun(X5,X0))),'_dst_x000x111x011_3'(s(fun(X2,fun(X6,X0)),X7),s(fun(X1,fun(X4,fun(X5,X2))),happ(s(fun(X3,fun(X1,fun(X4,fun(X5,X2)))),X8),s(X3,X10))),s(fun(X4,fun(X5,X6)),X9))) = s(fun(X1,fun(X4,fun(X5,X0))),happ(s(fun(X3,fun(X1,fun(X4,fun(X5,X0)))),'_dst_x0000x1111x0011_3'(s(fun(X2,fun(X6,X0)),X7),s(fun(X3,fun(X1,fun(X4,fun(X5,X2)))),X8),s(fun(X4,fun(X5,X6)),X9))),s(X3,X10)))),
  inference(rectify,[],[f185])).
fof(f185,axiom,(
  ! [X35,X36,X57,X38,X37,X56,X58] : ! [X39,X40,X41,X42] : s(fun(X36,fun(X37,fun(X56,X35))),'_dst_x000x111x011_3'(s(fun(X57,fun(X58,X35)),X39),s(fun(X36,fun(X37,fun(X56,X57))),happ(s(fun(X38,fun(X36,fun(X37,fun(X56,X57)))),X40),s(X38,X42))),s(fun(X37,fun(X56,X58)),X41))) = s(fun(X36,fun(X37,fun(X56,X35))),happ(s(fun(X38,fun(X36,fun(X37,fun(X56,X35)))),'_dst_x0000x1111x0011_3'(s(fun(X57,fun(X58,X35)),X39),s(fun(X38,fun(X36,fun(X37,fun(X56,X57)))),X40),s(fun(X37,fun(X56,X58)),X41))),s(X38,X42)))),
  file('thm.pred_set.EXTENSION','_dst_x0000x1111x0011_ax')).
fof(f1652,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (s(fun(X3,fun(X4,fun(X5,X0))),'_dst_x000x001x111_3'(s(fun(X1,fun(X6,X0)),X7),s(fun(X5,X1),happ(s(fun(X2,fun(X5,X1)),X8),s(X2,X10))),s(fun(X3,fun(X4,fun(X5,X6))),X9))) = s(fun(X3,fun(X4,fun(X5,X0))),happ(s(fun(X2,fun(X3,fun(X4,fun(X5,X0)))),'_dst_x0000x1001x0111_3'(s(fun(X1,fun(X6,X0)),X7),s(fun(X2,fun(X5,X1)),X8),s(fun(X3,fun(X4,fun(X5,X6))),X9))),s(X2,X10)))) )),
  inference(cnf_transformation,[],[f700])).
fof(f700,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : s(fun(X3,fun(X4,fun(X5,X0))),'_dst_x000x001x111_3'(s(fun(X1,fun(X6,X0)),X7),s(fun(X5,X1),happ(s(fun(X2,fun(X5,X1)),X8),s(X2,X10))),s(fun(X3,fun(X4,fun(X5,X6))),X9))) = s(fun(X3,fun(X4,fun(X5,X0))),happ(s(fun(X2,fun(X3,fun(X4,fun(X5,X0)))),'_dst_x0000x1001x0111_3'(s(fun(X1,fun(X6,X0)),X7),s(fun(X2,fun(X5,X1)),X8),s(fun(X3,fun(X4,fun(X5,X6))),X9))),s(X2,X10)))),
  inference(flattening,[],[f699])).
fof(f699,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ! [X7,X8,X9,X10] : s(fun(X3,fun(X4,fun(X5,X0))),'_dst_x000x001x111_3'(s(fun(X1,fun(X6,X0)),X7),s(fun(X5,X1),happ(s(fun(X2,fun(X5,X1)),X8),s(X2,X10))),s(fun(X3,fun(X4,fun(X5,X6))),X9))) = s(fun(X3,fun(X4,fun(X5,X0))),happ(s(fun(X2,fun(X3,fun(X4,fun(X5,X0)))),'_dst_x0000x1001x0111_3'(s(fun(X1,fun(X6,X0)),X7),s(fun(X2,fun(X5,X1)),X8),s(fun(X3,fun(X4,fun(X5,X6))),X9))),s(X2,X10)))),
  inference(rectify,[],[f184])).
fof(f184,axiom,(
  ! [X35,X57,X38,X36,X37,X56,X58] : ! [X39,X40,X41,X42] : s(fun(X36,fun(X37,fun(X56,X35))),'_dst_x000x001x111_3'(s(fun(X57,fun(X58,X35)),X39),s(fun(X56,X57),happ(s(fun(X38,fun(X56,X57)),X40),s(X38,X42))),s(fun(X36,fun(X37,fun(X56,X58))),X41))) = s(fun(X36,fun(X37,fun(X56,X35))),happ(s(fun(X38,fun(X36,fun(X37,fun(X56,X35)))),'_dst_x0000x1001x0111_3'(s(fun(X57,fun(X58,X35)),X39),s(fun(X38,fun(X56,X57)),X40),s(fun(X36,fun(X37,fun(X56,X58))),X41))),s(X38,X42)))),
  file('thm.pred_set.EXTENSION','_dst_x0000x1001x0111_ax')).
fof(f1651,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (s(fun(X1,fun(X2,fun(X5,X0))),'_dst_x000x110x001_3'(s(fun(X3,fun(X6,X0)),X7),s(fun(X1,fun(X2,X3)),happ(s(fun(X4,fun(X1,fun(X2,X3))),X8),s(X4,X10))),s(fun(X5,X6),X9))) = s(fun(X1,fun(X2,fun(X5,X0))),happ(s(fun(X4,fun(X1,fun(X2,fun(X5,X0)))),'_dst_x0000x1110x0001_3'(s(fun(X3,fun(X6,X0)),X7),s(fun(X4,fun(X1,fun(X2,X3))),X8),s(fun(X5,X6),X9))),s(X4,X10)))) )),
  inference(cnf_transformation,[],[f698])).
fof(f698,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : s(fun(X1,fun(X2,fun(X5,X0))),'_dst_x000x110x001_3'(s(fun(X3,fun(X6,X0)),X7),s(fun(X1,fun(X2,X3)),happ(s(fun(X4,fun(X1,fun(X2,X3))),X8),s(X4,X10))),s(fun(X5,X6),X9))) = s(fun(X1,fun(X2,fun(X5,X0))),happ(s(fun(X4,fun(X1,fun(X2,fun(X5,X0)))),'_dst_x0000x1110x0001_3'(s(fun(X3,fun(X6,X0)),X7),s(fun(X4,fun(X1,fun(X2,X3))),X8),s(fun(X5,X6),X9))),s(X4,X10)))),
  inference(flattening,[],[f697])).
fof(f697,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ! [X7,X8,X9,X10] : s(fun(X1,fun(X2,fun(X5,X0))),'_dst_x000x110x001_3'(s(fun(X3,fun(X6,X0)),X7),s(fun(X1,fun(X2,X3)),happ(s(fun(X4,fun(X1,fun(X2,X3))),X8),s(X4,X10))),s(fun(X5,X6),X9))) = s(fun(X1,fun(X2,fun(X5,X0))),happ(s(fun(X4,fun(X1,fun(X2,fun(X5,X0)))),'_dst_x0000x1110x0001_3'(s(fun(X3,fun(X6,X0)),X7),s(fun(X4,fun(X1,fun(X2,X3))),X8),s(fun(X5,X6),X9))),s(X4,X10)))),
  inference(rectify,[],[f172])).
fof(f172,axiom,(
  ! [X35,X36,X37,X57,X38,X56,X58] : ! [X39,X40,X41,X42] : s(fun(X36,fun(X37,fun(X56,X35))),'_dst_x000x110x001_3'(s(fun(X57,fun(X58,X35)),X39),s(fun(X36,fun(X37,X57)),happ(s(fun(X38,fun(X36,fun(X37,X57))),X40),s(X38,X42))),s(fun(X56,X58),X41))) = s(fun(X36,fun(X37,fun(X56,X35))),happ(s(fun(X38,fun(X36,fun(X37,fun(X56,X35)))),'_dst_x0000x1110x0001_3'(s(fun(X57,fun(X58,X35)),X39),s(fun(X38,fun(X36,fun(X37,X57))),X40),s(fun(X56,X58),X41))),s(X38,X42)))),
  file('thm.pred_set.EXTENSION','_dst_x0000x1110x0001_ax')).
fof(f1650,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (s(fun(X4,fun(X1,fun(X5,X0))),'_dst_x000x010x101_3'(s(fun(X2,fun(X6,X0)),X7),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X8),s(X3,X10))),s(fun(X4,fun(X5,X6)),X9))) = s(fun(X4,fun(X1,fun(X5,X0))),happ(s(fun(X3,fun(X4,fun(X1,fun(X5,X0)))),'_dst_x0000x1010x0101_3'(s(fun(X2,fun(X6,X0)),X7),s(fun(X3,fun(X1,X2)),X8),s(fun(X4,fun(X5,X6)),X9))),s(X3,X10)))) )),
  inference(cnf_transformation,[],[f696])).
fof(f696,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : s(fun(X4,fun(X1,fun(X5,X0))),'_dst_x000x010x101_3'(s(fun(X2,fun(X6,X0)),X7),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X8),s(X3,X10))),s(fun(X4,fun(X5,X6)),X9))) = s(fun(X4,fun(X1,fun(X5,X0))),happ(s(fun(X3,fun(X4,fun(X1,fun(X5,X0)))),'_dst_x0000x1010x0101_3'(s(fun(X2,fun(X6,X0)),X7),s(fun(X3,fun(X1,X2)),X8),s(fun(X4,fun(X5,X6)),X9))),s(X3,X10)))),
  inference(flattening,[],[f695])).
fof(f695,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ! [X7,X8,X9,X10] : s(fun(X4,fun(X1,fun(X5,X0))),'_dst_x000x010x101_3'(s(fun(X2,fun(X6,X0)),X7),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X8),s(X3,X10))),s(fun(X4,fun(X5,X6)),X9))) = s(fun(X4,fun(X1,fun(X5,X0))),happ(s(fun(X3,fun(X4,fun(X1,fun(X5,X0)))),'_dst_x0000x1010x0101_3'(s(fun(X2,fun(X6,X0)),X7),s(fun(X3,fun(X1,X2)),X8),s(fun(X4,fun(X5,X6)),X9))),s(X3,X10)))),
  inference(rectify,[],[f118])).
fof(f118,axiom,(
  ! [X35,X37,X57,X38,X36,X56,X58] : ! [X39,X40,X41,X42] : s(fun(X36,fun(X37,fun(X56,X35))),'_dst_x000x010x101_3'(s(fun(X57,fun(X58,X35)),X39),s(fun(X37,X57),happ(s(fun(X38,fun(X37,X57)),X40),s(X38,X42))),s(fun(X36,fun(X56,X58)),X41))) = s(fun(X36,fun(X37,fun(X56,X35))),happ(s(fun(X38,fun(X36,fun(X37,fun(X56,X35)))),'_dst_x0000x1010x0101_3'(s(fun(X57,fun(X58,X35)),X39),s(fun(X38,fun(X37,X57)),X40),s(fun(X36,fun(X56,X58)),X41))),s(X38,X42)))),
  file('thm.pred_set.EXTENSION','_dst_x0000x1010x0101_ax')).
fof(f1649,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (s(fun(X3,X0),happ(s(fun(X5,fun(X3,X0)),'_dst_x00x11x10x11_4'(s(fun(X1,fun(X2,fun(X4,X0))),X6),s(fun(X5,fun(X3,X1)),X7),s(fun(X5,X2),X8),s(fun(X5,fun(X3,X4)),X9))),s(X5,X10))) = s(fun(X3,X0),'_dst_x0x1x0x1_4'(s(fun(X1,fun(X2,fun(X4,X0))),X6),s(fun(X3,X1),happ(s(fun(X5,fun(X3,X1)),X7),s(X5,X10))),s(X2,happ(s(fun(X5,X2),X8),s(X5,X10))),s(fun(X3,X4),happ(s(fun(X5,fun(X3,X4)),X9),s(X5,X10)))))) )),
  inference(cnf_transformation,[],[f694])).
fof(f694,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : s(fun(X3,X0),happ(s(fun(X5,fun(X3,X0)),'_dst_x00x11x10x11_4'(s(fun(X1,fun(X2,fun(X4,X0))),X6),s(fun(X5,fun(X3,X1)),X7),s(fun(X5,X2),X8),s(fun(X5,fun(X3,X4)),X9))),s(X5,X10))) = s(fun(X3,X0),'_dst_x0x1x0x1_4'(s(fun(X1,fun(X2,fun(X4,X0))),X6),s(fun(X3,X1),happ(s(fun(X5,fun(X3,X1)),X7),s(X5,X10))),s(X2,happ(s(fun(X5,X2),X8),s(X5,X10))),s(fun(X3,X4),happ(s(fun(X5,fun(X3,X4)),X9),s(X5,X10)))))),
  inference(flattening,[],[f693])).
fof(f693,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ! [X6,X7,X8,X9,X10] : s(fun(X3,X0),happ(s(fun(X5,fun(X3,X0)),'_dst_x00x11x10x11_4'(s(fun(X1,fun(X2,fun(X4,X0))),X6),s(fun(X5,fun(X3,X1)),X7),s(fun(X5,X2),X8),s(fun(X5,fun(X3,X4)),X9))),s(X5,X10))) = s(fun(X3,X0),'_dst_x0x1x0x1_4'(s(fun(X1,fun(X2,fun(X4,X0))),X6),s(fun(X3,X1),happ(s(fun(X5,fun(X3,X1)),X7),s(X5,X10))),s(X2,happ(s(fun(X5,X2),X8),s(X5,X10))),s(fun(X3,X4),happ(s(fun(X5,fun(X3,X4)),X9),s(X5,X10)))))),
  inference(rectify,[],[f191])).
fof(f191,axiom,(
  ! [X35,X37,X56,X36,X57,X38] : ! [X39,X40,X41,X63,X42] : s(fun(X36,X35),happ(s(fun(X38,fun(X36,X35)),'_dst_x00x11x10x11_4'(s(fun(X37,fun(X56,fun(X57,X35))),X39),s(fun(X38,fun(X36,X37)),X40),s(fun(X38,X56),X41),s(fun(X38,fun(X36,X57)),X63))),s(X38,X42))) = s(fun(X36,X35),'_dst_x0x1x0x1_4'(s(fun(X37,fun(X56,fun(X57,X35))),X39),s(fun(X36,X37),happ(s(fun(X38,fun(X36,X37)),X40),s(X38,X42))),s(X56,happ(s(fun(X38,X56),X41),s(X38,X42))),s(fun(X36,X57),happ(s(fun(X38,fun(X36,X57)),X63),s(X38,X42)))))),
  file('thm.pred_set.EXTENSION','_dst_x00x11x10x11_ax')).
fof(f1648,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (s(fun(X3,X0),happ(s(fun(X5,fun(X3,X0)),'_dst_x00x11x00x11_4'(s(fun(X1,fun(X2,fun(X4,X0))),X6),s(fun(X5,fun(X3,X1)),X7),s(X2,X8),s(fun(X5,fun(X3,X4)),X9))),s(X5,X10))) = s(fun(X3,X0),'_dst_x0x1x0x1_4'(s(fun(X1,fun(X2,fun(X4,X0))),X6),s(fun(X3,X1),happ(s(fun(X5,fun(X3,X1)),X7),s(X5,X10))),s(X2,X8),s(fun(X3,X4),happ(s(fun(X5,fun(X3,X4)),X9),s(X5,X10)))))) )),
  inference(cnf_transformation,[],[f692])).
fof(f692,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : s(fun(X3,X0),happ(s(fun(X5,fun(X3,X0)),'_dst_x00x11x00x11_4'(s(fun(X1,fun(X2,fun(X4,X0))),X6),s(fun(X5,fun(X3,X1)),X7),s(X2,X8),s(fun(X5,fun(X3,X4)),X9))),s(X5,X10))) = s(fun(X3,X0),'_dst_x0x1x0x1_4'(s(fun(X1,fun(X2,fun(X4,X0))),X6),s(fun(X3,X1),happ(s(fun(X5,fun(X3,X1)),X7),s(X5,X10))),s(X2,X8),s(fun(X3,X4),happ(s(fun(X5,fun(X3,X4)),X9),s(X5,X10)))))),
  inference(flattening,[],[f691])).
fof(f691,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ! [X6,X7,X8,X9,X10] : s(fun(X3,X0),happ(s(fun(X5,fun(X3,X0)),'_dst_x00x11x00x11_4'(s(fun(X1,fun(X2,fun(X4,X0))),X6),s(fun(X5,fun(X3,X1)),X7),s(X2,X8),s(fun(X5,fun(X3,X4)),X9))),s(X5,X10))) = s(fun(X3,X0),'_dst_x0x1x0x1_4'(s(fun(X1,fun(X2,fun(X4,X0))),X6),s(fun(X3,X1),happ(s(fun(X5,fun(X3,X1)),X7),s(X5,X10))),s(X2,X8),s(fun(X3,X4),happ(s(fun(X5,fun(X3,X4)),X9),s(X5,X10)))))),
  inference(rectify,[],[f175])).
fof(f175,axiom,(
  ! [X35,X37,X56,X36,X57,X38] : ! [X39,X40,X41,X63,X42] : s(fun(X36,X35),happ(s(fun(X38,fun(X36,X35)),'_dst_x00x11x00x11_4'(s(fun(X37,fun(X56,fun(X57,X35))),X39),s(fun(X38,fun(X36,X37)),X40),s(X56,X41),s(fun(X38,fun(X36,X57)),X63))),s(X38,X42))) = s(fun(X36,X35),'_dst_x0x1x0x1_4'(s(fun(X37,fun(X56,fun(X57,X35))),X39),s(fun(X36,X37),happ(s(fun(X38,fun(X36,X37)),X40),s(X38,X42))),s(X56,X41),s(fun(X36,X57),happ(s(fun(X38,fun(X36,X57)),X63),s(X38,X42)))))),
  file('thm.pred_set.EXTENSION','_dst_x00x11x00x11_ax')).
fof(f1647,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (s(fun(X3,X0),'_dst_x0x1x0x1_4'(s(fun(X1,fun(X2,fun(X4,X0))),X6),s(fun(X3,X1),X7),s(X2,X8),s(fun(X3,X4),happ(s(fun(X5,fun(X3,X4)),X9),s(X5,X10))))) = s(fun(X3,X0),happ(s(fun(X5,fun(X3,X0)),'_dst_x00x01x00x11_4'(s(fun(X1,fun(X2,fun(X4,X0))),X6),s(fun(X3,X1),X7),s(X2,X8),s(fun(X5,fun(X3,X4)),X9))),s(X5,X10)))) )),
  inference(cnf_transformation,[],[f690])).
fof(f690,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : s(fun(X3,X0),'_dst_x0x1x0x1_4'(s(fun(X1,fun(X2,fun(X4,X0))),X6),s(fun(X3,X1),X7),s(X2,X8),s(fun(X3,X4),happ(s(fun(X5,fun(X3,X4)),X9),s(X5,X10))))) = s(fun(X3,X0),happ(s(fun(X5,fun(X3,X0)),'_dst_x00x01x00x11_4'(s(fun(X1,fun(X2,fun(X4,X0))),X6),s(fun(X3,X1),X7),s(X2,X8),s(fun(X5,fun(X3,X4)),X9))),s(X5,X10)))),
  inference(flattening,[],[f689])).
fof(f689,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ! [X6,X7,X8,X9,X10] : s(fun(X3,X0),'_dst_x0x1x0x1_4'(s(fun(X1,fun(X2,fun(X4,X0))),X6),s(fun(X3,X1),X7),s(X2,X8),s(fun(X3,X4),happ(s(fun(X5,fun(X3,X4)),X9),s(X5,X10))))) = s(fun(X3,X0),happ(s(fun(X5,fun(X3,X0)),'_dst_x00x01x00x11_4'(s(fun(X1,fun(X2,fun(X4,X0))),X6),s(fun(X3,X1),X7),s(X2,X8),s(fun(X5,fun(X3,X4)),X9))),s(X5,X10)))),
  inference(rectify,[],[f161])).
fof(f161,axiom,(
  ! [X35,X37,X56,X36,X57,X38] : ! [X39,X40,X41,X63,X42] : s(fun(X36,X35),'_dst_x0x1x0x1_4'(s(fun(X37,fun(X56,fun(X57,X35))),X39),s(fun(X36,X37),X40),s(X56,X41),s(fun(X36,X57),happ(s(fun(X38,fun(X36,X57)),X63),s(X38,X42))))) = s(fun(X36,X35),happ(s(fun(X38,fun(X36,X35)),'_dst_x00x01x00x11_4'(s(fun(X37,fun(X56,fun(X57,X35))),X39),s(fun(X36,X37),X40),s(X56,X41),s(fun(X38,fun(X36,X57)),X63))),s(X38,X42)))),
  file('thm.pred_set.EXTENSION','_dst_x00x01x00x11_ax')).
fof(f1646,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s(fun(X3,fun(X4,X0)),'_dst_x00x11x11_3'(s(fun(X1,fun(X5,X0)),X6),s(fun(X3,fun(X4,X1)),happ(s(fun(X2,fun(X3,fun(X4,X1))),X7),s(X2,X9))),s(fun(X3,fun(X4,X5)),X8))) = s(fun(X3,fun(X4,X0)),happ(s(fun(X2,fun(X3,fun(X4,X0))),'_dst_x000x111x011_3'(s(fun(X1,fun(X5,X0)),X6),s(fun(X2,fun(X3,fun(X4,X1))),X7),s(fun(X3,fun(X4,X5)),X8))),s(X2,X9)))) )),
  inference(cnf_transformation,[],[f688])).
fof(f688,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : s(fun(X3,fun(X4,X0)),'_dst_x00x11x11_3'(s(fun(X1,fun(X5,X0)),X6),s(fun(X3,fun(X4,X1)),happ(s(fun(X2,fun(X3,fun(X4,X1))),X7),s(X2,X9))),s(fun(X3,fun(X4,X5)),X8))) = s(fun(X3,fun(X4,X0)),happ(s(fun(X2,fun(X3,fun(X4,X0))),'_dst_x000x111x011_3'(s(fun(X1,fun(X5,X0)),X6),s(fun(X2,fun(X3,fun(X4,X1))),X7),s(fun(X3,fun(X4,X5)),X8))),s(X2,X9)))),
  inference(flattening,[],[f687])).
fof(f687,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ! [X6,X7,X8,X9] : s(fun(X3,fun(X4,X0)),'_dst_x00x11x11_3'(s(fun(X1,fun(X5,X0)),X6),s(fun(X3,fun(X4,X1)),happ(s(fun(X2,fun(X3,fun(X4,X1))),X7),s(X2,X9))),s(fun(X3,fun(X4,X5)),X8))) = s(fun(X3,fun(X4,X0)),happ(s(fun(X2,fun(X3,fun(X4,X0))),'_dst_x000x111x011_3'(s(fun(X1,fun(X5,X0)),X6),s(fun(X2,fun(X3,fun(X4,X1))),X7),s(fun(X3,fun(X4,X5)),X8))),s(X2,X9)))),
  inference(rectify,[],[f194])).
fof(f194,axiom,(
  ! [X35,X56,X38,X36,X37,X57] : ! [X39,X40,X41,X42] : s(fun(X36,fun(X37,X35)),'_dst_x00x11x11_3'(s(fun(X56,fun(X57,X35)),X39),s(fun(X36,fun(X37,X56)),happ(s(fun(X38,fun(X36,fun(X37,X56))),X40),s(X38,X42))),s(fun(X36,fun(X37,X57)),X41))) = s(fun(X36,fun(X37,X35)),happ(s(fun(X38,fun(X36,fun(X37,X35))),'_dst_x000x111x011_3'(s(fun(X56,fun(X57,X35)),X39),s(fun(X38,fun(X36,fun(X37,X56))),X40),s(fun(X36,fun(X37,X57)),X41))),s(X38,X42)))),
  file('thm.pred_set.EXTENSION','_dst_x000x111x011_ax')).
fof(f1645,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s(fun(X2,fun(X3,X0)),'_dst_x00x11x11_3'(s(fun(X1,fun(X4,X0)),X6),s(fun(X2,fun(X3,X1)),X7),s(fun(X2,fun(X3,X4)),happ(s(fun(X5,fun(X2,fun(X3,X4))),X8),s(X5,X9))))) = s(fun(X2,fun(X3,X0)),happ(s(fun(X5,fun(X2,fun(X3,X0))),'_dst_x000x011x111_3'(s(fun(X1,fun(X4,X0)),X6),s(fun(X2,fun(X3,X1)),X7),s(fun(X5,fun(X2,fun(X3,X4))),X8))),s(X5,X9)))) )),
  inference(cnf_transformation,[],[f686])).
fof(f686,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : s(fun(X2,fun(X3,X0)),'_dst_x00x11x11_3'(s(fun(X1,fun(X4,X0)),X6),s(fun(X2,fun(X3,X1)),X7),s(fun(X2,fun(X3,X4)),happ(s(fun(X5,fun(X2,fun(X3,X4))),X8),s(X5,X9))))) = s(fun(X2,fun(X3,X0)),happ(s(fun(X5,fun(X2,fun(X3,X0))),'_dst_x000x011x111_3'(s(fun(X1,fun(X4,X0)),X6),s(fun(X2,fun(X3,X1)),X7),s(fun(X5,fun(X2,fun(X3,X4))),X8))),s(X5,X9)))),
  inference(flattening,[],[f685])).
fof(f685,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ! [X6,X7,X8,X9] : s(fun(X2,fun(X3,X0)),'_dst_x00x11x11_3'(s(fun(X1,fun(X4,X0)),X6),s(fun(X2,fun(X3,X1)),X7),s(fun(X2,fun(X3,X4)),happ(s(fun(X5,fun(X2,fun(X3,X4))),X8),s(X5,X9))))) = s(fun(X2,fun(X3,X0)),happ(s(fun(X5,fun(X2,fun(X3,X0))),'_dst_x000x011x111_3'(s(fun(X1,fun(X4,X0)),X6),s(fun(X2,fun(X3,X1)),X7),s(fun(X5,fun(X2,fun(X3,X4))),X8))),s(X5,X9)))),
  inference(rectify,[],[f122])).
fof(f122,axiom,(
  ! [X35,X56,X36,X37,X57,X38] : ! [X39,X40,X41,X42] : s(fun(X36,fun(X37,X35)),'_dst_x00x11x11_3'(s(fun(X56,fun(X57,X35)),X39),s(fun(X36,fun(X37,X56)),X40),s(fun(X36,fun(X37,X57)),happ(s(fun(X38,fun(X36,fun(X37,X57))),X41),s(X38,X42))))) = s(fun(X36,fun(X37,X35)),happ(s(fun(X38,fun(X36,fun(X37,X35))),'_dst_x000x011x111_3'(s(fun(X56,fun(X57,X35)),X39),s(fun(X36,fun(X37,X56)),X40),s(fun(X38,fun(X36,fun(X37,X57))),X41))),s(X38,X42)))),
  file('thm.pred_set.EXTENSION','_dst_x000x011x111_ax')).
fof(f1644,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s(fun(X4,fun(X1,X0)),'_dst_x00x11x10_3'(s(fun(X2,fun(X5,X0)),X6),s(fun(X4,fun(X1,X2)),happ(s(fun(X3,fun(X4,fun(X1,X2))),X7),s(X3,X9))),s(fun(X4,X5),X8))) = s(fun(X4,fun(X1,X0)),happ(s(fun(X3,fun(X4,fun(X1,X0))),'_dst_x000x111x010_3'(s(fun(X2,fun(X5,X0)),X6),s(fun(X3,fun(X4,fun(X1,X2))),X7),s(fun(X4,X5),X8))),s(X3,X9)))) )),
  inference(cnf_transformation,[],[f684])).
fof(f684,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : s(fun(X4,fun(X1,X0)),'_dst_x00x11x10_3'(s(fun(X2,fun(X5,X0)),X6),s(fun(X4,fun(X1,X2)),happ(s(fun(X3,fun(X4,fun(X1,X2))),X7),s(X3,X9))),s(fun(X4,X5),X8))) = s(fun(X4,fun(X1,X0)),happ(s(fun(X3,fun(X4,fun(X1,X0))),'_dst_x000x111x010_3'(s(fun(X2,fun(X5,X0)),X6),s(fun(X3,fun(X4,fun(X1,X2))),X7),s(fun(X4,X5),X8))),s(X3,X9)))),
  inference(flattening,[],[f683])).
fof(f683,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ! [X6,X7,X8,X9] : s(fun(X4,fun(X1,X0)),'_dst_x00x11x10_3'(s(fun(X2,fun(X5,X0)),X6),s(fun(X4,fun(X1,X2)),happ(s(fun(X3,fun(X4,fun(X1,X2))),X7),s(X3,X9))),s(fun(X4,X5),X8))) = s(fun(X4,fun(X1,X0)),happ(s(fun(X3,fun(X4,fun(X1,X0))),'_dst_x000x111x010_3'(s(fun(X2,fun(X5,X0)),X6),s(fun(X3,fun(X4,fun(X1,X2))),X7),s(fun(X4,X5),X8))),s(X3,X9)))),
  inference(rectify,[],[f189])).
fof(f189,axiom,(
  ! [X35,X37,X56,X38,X36,X57] : ! [X39,X40,X41,X42] : s(fun(X36,fun(X37,X35)),'_dst_x00x11x10_3'(s(fun(X56,fun(X57,X35)),X39),s(fun(X36,fun(X37,X56)),happ(s(fun(X38,fun(X36,fun(X37,X56))),X40),s(X38,X42))),s(fun(X36,X57),X41))) = s(fun(X36,fun(X37,X35)),happ(s(fun(X38,fun(X36,fun(X37,X35))),'_dst_x000x111x010_3'(s(fun(X56,fun(X57,X35)),X39),s(fun(X38,fun(X36,fun(X37,X56))),X40),s(fun(X36,X57),X41))),s(X38,X42)))),
  file('thm.pred_set.EXTENSION','_dst_x000x111x010_ax')).
fof(f1643,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s(fun(X2,fun(X3,X0)),'_dst_x00x01x11_3'(s(fun(X1,fun(X4,X0)),X6),s(fun(X3,X1),X7),s(fun(X2,fun(X3,X4)),happ(s(fun(X5,fun(X2,fun(X3,X4))),X8),s(X5,X9))))) = s(fun(X2,fun(X3,X0)),happ(s(fun(X5,fun(X2,fun(X3,X0))),'_dst_x000x001x111_3'(s(fun(X1,fun(X4,X0)),X6),s(fun(X3,X1),X7),s(fun(X5,fun(X2,fun(X3,X4))),X8))),s(X5,X9)))) )),
  inference(cnf_transformation,[],[f682])).
fof(f682,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : s(fun(X2,fun(X3,X0)),'_dst_x00x01x11_3'(s(fun(X1,fun(X4,X0)),X6),s(fun(X3,X1),X7),s(fun(X2,fun(X3,X4)),happ(s(fun(X5,fun(X2,fun(X3,X4))),X8),s(X5,X9))))) = s(fun(X2,fun(X3,X0)),happ(s(fun(X5,fun(X2,fun(X3,X0))),'_dst_x000x001x111_3'(s(fun(X1,fun(X4,X0)),X6),s(fun(X3,X1),X7),s(fun(X5,fun(X2,fun(X3,X4))),X8))),s(X5,X9)))),
  inference(flattening,[],[f681])).
fof(f681,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ! [X6,X7,X8,X9] : s(fun(X2,fun(X3,X0)),'_dst_x00x01x11_3'(s(fun(X1,fun(X4,X0)),X6),s(fun(X3,X1),X7),s(fun(X2,fun(X3,X4)),happ(s(fun(X5,fun(X2,fun(X3,X4))),X8),s(X5,X9))))) = s(fun(X2,fun(X3,X0)),happ(s(fun(X5,fun(X2,fun(X3,X0))),'_dst_x000x001x111_3'(s(fun(X1,fun(X4,X0)),X6),s(fun(X3,X1),X7),s(fun(X5,fun(X2,fun(X3,X4))),X8))),s(X5,X9)))),
  inference(rectify,[],[f186])).
fof(f186,axiom,(
  ! [X35,X56,X36,X37,X57,X38] : ! [X39,X40,X41,X42] : s(fun(X36,fun(X37,X35)),'_dst_x00x01x11_3'(s(fun(X56,fun(X57,X35)),X39),s(fun(X37,X56),X40),s(fun(X36,fun(X37,X57)),happ(s(fun(X38,fun(X36,fun(X37,X57))),X41),s(X38,X42))))) = s(fun(X36,fun(X37,X35)),happ(s(fun(X38,fun(X36,fun(X37,X35))),'_dst_x000x001x111_3'(s(fun(X56,fun(X57,X35)),X39),s(fun(X37,X56),X40),s(fun(X38,fun(X36,fun(X37,X57))),X41))),s(X38,X42)))),
  file('thm.pred_set.EXTENSION','_dst_x000x001x111_ax')).
fof(f1642,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s(fun(X1,fun(X4,X0)),'_dst_x00x10x01_3'(s(fun(X2,fun(X5,X0)),X6),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X7),s(X3,X9))),s(fun(X4,X5),X8))) = s(fun(X1,fun(X4,X0)),happ(s(fun(X3,fun(X1,fun(X4,X0))),'_dst_x000x110x001_3'(s(fun(X2,fun(X5,X0)),X6),s(fun(X3,fun(X1,X2)),X7),s(fun(X4,X5),X8))),s(X3,X9)))) )),
  inference(cnf_transformation,[],[f680])).
fof(f680,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : s(fun(X1,fun(X4,X0)),'_dst_x00x10x01_3'(s(fun(X2,fun(X5,X0)),X6),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X7),s(X3,X9))),s(fun(X4,X5),X8))) = s(fun(X1,fun(X4,X0)),happ(s(fun(X3,fun(X1,fun(X4,X0))),'_dst_x000x110x001_3'(s(fun(X2,fun(X5,X0)),X6),s(fun(X3,fun(X1,X2)),X7),s(fun(X4,X5),X8))),s(X3,X9)))),
  inference(flattening,[],[f679])).
fof(f679,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ! [X6,X7,X8,X9] : s(fun(X1,fun(X4,X0)),'_dst_x00x10x01_3'(s(fun(X2,fun(X5,X0)),X6),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X7),s(X3,X9))),s(fun(X4,X5),X8))) = s(fun(X1,fun(X4,X0)),happ(s(fun(X3,fun(X1,fun(X4,X0))),'_dst_x000x110x001_3'(s(fun(X2,fun(X5,X0)),X6),s(fun(X3,fun(X1,X2)),X7),s(fun(X4,X5),X8))),s(X3,X9)))),
  inference(rectify,[],[f173])).
fof(f173,axiom,(
  ! [X35,X36,X56,X38,X37,X57] : ! [X39,X40,X41,X42] : s(fun(X36,fun(X37,X35)),'_dst_x00x10x01_3'(s(fun(X56,fun(X57,X35)),X39),s(fun(X36,X56),happ(s(fun(X38,fun(X36,X56)),X40),s(X38,X42))),s(fun(X37,X57),X41))) = s(fun(X36,fun(X37,X35)),happ(s(fun(X38,fun(X36,fun(X37,X35))),'_dst_x000x110x001_3'(s(fun(X56,fun(X57,X35)),X39),s(fun(X38,fun(X36,X56)),X40),s(fun(X37,X57),X41))),s(X38,X42)))),
  file('thm.pred_set.EXTENSION','_dst_x000x110x001_ax')).
fof(f1641,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s(fun(X1,fun(X3,X0)),'_dst_x00x10x01_3'(s(fun(X2,fun(X4,X0)),X6),s(fun(X1,X2),X7),s(fun(X3,X4),happ(s(fun(X5,fun(X3,X4)),X8),s(X5,X9))))) = s(fun(X1,fun(X3,X0)),happ(s(fun(X5,fun(X1,fun(X3,X0))),'_dst_x000x010x101_3'(s(fun(X2,fun(X4,X0)),X6),s(fun(X1,X2),X7),s(fun(X5,fun(X3,X4)),X8))),s(X5,X9)))) )),
  inference(cnf_transformation,[],[f678])).
fof(f678,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : s(fun(X1,fun(X3,X0)),'_dst_x00x10x01_3'(s(fun(X2,fun(X4,X0)),X6),s(fun(X1,X2),X7),s(fun(X3,X4),happ(s(fun(X5,fun(X3,X4)),X8),s(X5,X9))))) = s(fun(X1,fun(X3,X0)),happ(s(fun(X5,fun(X1,fun(X3,X0))),'_dst_x000x010x101_3'(s(fun(X2,fun(X4,X0)),X6),s(fun(X1,X2),X7),s(fun(X5,fun(X3,X4)),X8))),s(X5,X9)))),
  inference(flattening,[],[f677])).
fof(f677,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ! [X6,X7,X8,X9] : s(fun(X1,fun(X3,X0)),'_dst_x00x10x01_3'(s(fun(X2,fun(X4,X0)),X6),s(fun(X1,X2),X7),s(fun(X3,X4),happ(s(fun(X5,fun(X3,X4)),X8),s(X5,X9))))) = s(fun(X1,fun(X3,X0)),happ(s(fun(X5,fun(X1,fun(X3,X0))),'_dst_x000x010x101_3'(s(fun(X2,fun(X4,X0)),X6),s(fun(X1,X2),X7),s(fun(X5,fun(X3,X4)),X8))),s(X5,X9)))),
  inference(rectify,[],[f121])).
fof(f121,axiom,(
  ! [X35,X36,X56,X37,X57,X38] : ! [X39,X40,X41,X42] : s(fun(X36,fun(X37,X35)),'_dst_x00x10x01_3'(s(fun(X56,fun(X57,X35)),X39),s(fun(X36,X56),X40),s(fun(X37,X57),happ(s(fun(X38,fun(X37,X57)),X41),s(X38,X42))))) = s(fun(X36,fun(X37,X35)),happ(s(fun(X38,fun(X36,fun(X37,X35))),'_dst_x000x010x101_3'(s(fun(X56,fun(X57,X35)),X39),s(fun(X36,X56),X40),s(fun(X38,fun(X37,X57)),X41))),s(X38,X42)))),
  file('thm.pred_set.EXTENSION','_dst_x000x010x101_ax')).
fof(f1640,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s(fun(X3,fun(X4,X0)),'_dst_x00x00x11_3'(s(fun(X1,fun(X5,X0)),X6),s(X1,happ(s(fun(X2,X1),X7),s(X2,X9))),s(fun(X3,fun(X4,X5)),X8))) = s(fun(X3,fun(X4,X0)),happ(s(fun(X2,fun(X3,fun(X4,X0))),'_dst_x000x100x011_3'(s(fun(X1,fun(X5,X0)),X6),s(fun(X2,X1),X7),s(fun(X3,fun(X4,X5)),X8))),s(X2,X9)))) )),
  inference(cnf_transformation,[],[f676])).
fof(f676,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : s(fun(X3,fun(X4,X0)),'_dst_x00x00x11_3'(s(fun(X1,fun(X5,X0)),X6),s(X1,happ(s(fun(X2,X1),X7),s(X2,X9))),s(fun(X3,fun(X4,X5)),X8))) = s(fun(X3,fun(X4,X0)),happ(s(fun(X2,fun(X3,fun(X4,X0))),'_dst_x000x100x011_3'(s(fun(X1,fun(X5,X0)),X6),s(fun(X2,X1),X7),s(fun(X3,fun(X4,X5)),X8))),s(X2,X9)))),
  inference(flattening,[],[f675])).
fof(f675,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ! [X6,X7,X8,X9] : s(fun(X3,fun(X4,X0)),'_dst_x00x00x11_3'(s(fun(X1,fun(X5,X0)),X6),s(X1,happ(s(fun(X2,X1),X7),s(X2,X9))),s(fun(X3,fun(X4,X5)),X8))) = s(fun(X3,fun(X4,X0)),happ(s(fun(X2,fun(X3,fun(X4,X0))),'_dst_x000x100x011_3'(s(fun(X1,fun(X5,X0)),X6),s(fun(X2,X1),X7),s(fun(X3,fun(X4,X5)),X8))),s(X2,X9)))),
  inference(rectify,[],[f160])).
fof(f160,axiom,(
  ! [X35,X56,X38,X36,X37,X57] : ! [X39,X40,X41,X42] : s(fun(X36,fun(X37,X35)),'_dst_x00x00x11_3'(s(fun(X56,fun(X57,X35)),X39),s(X56,happ(s(fun(X38,X56),X40),s(X38,X42))),s(fun(X36,fun(X37,X57)),X41))) = s(fun(X36,fun(X37,X35)),happ(s(fun(X38,fun(X36,fun(X37,X35))),'_dst_x000x100x011_3'(s(fun(X56,fun(X57,X35)),X39),s(fun(X38,X56),X40),s(fun(X36,fun(X37,X57)),X41))),s(X38,X42)))),
  file('thm.pred_set.EXTENSION','_dst_x000x100x011_ax')).
fof(f1639,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(fun(X1,fun(X2,fun(X3,X0))),'_dst_x000x111_2'(s(fun(X4,X0),X6),s(fun(X1,fun(X2,fun(X3,X4))),happ(s(fun(X5,fun(X1,fun(X2,fun(X3,X4)))),X7),s(X5,X8))))) = s(fun(X1,fun(X2,fun(X3,X0))),happ(s(fun(X5,fun(X1,fun(X2,fun(X3,X0)))),'_dst_x0000x1111_2'(s(fun(X4,X0),X6),s(fun(X5,fun(X1,fun(X2,fun(X3,X4)))),X7))),s(X5,X8)))) )),
  inference(cnf_transformation,[],[f674])).
fof(f674,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : s(fun(X1,fun(X2,fun(X3,X0))),'_dst_x000x111_2'(s(fun(X4,X0),X6),s(fun(X1,fun(X2,fun(X3,X4))),happ(s(fun(X5,fun(X1,fun(X2,fun(X3,X4)))),X7),s(X5,X8))))) = s(fun(X1,fun(X2,fun(X3,X0))),happ(s(fun(X5,fun(X1,fun(X2,fun(X3,X0)))),'_dst_x0000x1111_2'(s(fun(X4,X0),X6),s(fun(X5,fun(X1,fun(X2,fun(X3,X4)))),X7))),s(X5,X8)))),
  inference(flattening,[],[f673])).
fof(f673,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ! [X6,X7,X8] : s(fun(X1,fun(X2,fun(X3,X0))),'_dst_x000x111_2'(s(fun(X4,X0),X6),s(fun(X1,fun(X2,fun(X3,X4))),happ(s(fun(X5,fun(X1,fun(X2,fun(X3,X4)))),X7),s(X5,X8))))) = s(fun(X1,fun(X2,fun(X3,X0))),happ(s(fun(X5,fun(X1,fun(X2,fun(X3,X0)))),'_dst_x0000x1111_2'(s(fun(X4,X0),X6),s(fun(X5,fun(X1,fun(X2,fun(X3,X4)))),X7))),s(X5,X8)))),
  inference(rectify,[],[f113])).
fof(f113,axiom,(
  ! [X35,X36,X37,X56,X57,X38] : ! [X39,X40,X42] : s(fun(X36,fun(X37,fun(X56,X35))),'_dst_x000x111_2'(s(fun(X57,X35),X39),s(fun(X36,fun(X37,fun(X56,X57))),happ(s(fun(X38,fun(X36,fun(X37,fun(X56,X57)))),X40),s(X38,X42))))) = s(fun(X36,fun(X37,fun(X56,X35))),happ(s(fun(X38,fun(X36,fun(X37,fun(X56,X35)))),'_dst_x0000x1111_2'(s(fun(X57,X35),X39),s(fun(X38,fun(X36,fun(X37,fun(X56,X57)))),X40))),s(X38,X42)))),
  file('thm.pred_set.EXTENSION','_dst_x0000x1111_ax')).
fof(f1638,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(fun(X1,fun(X2,fun(X3,fun(X4,X0)))),'_dst_x0000x1111_2'(s(fun(X5,X0),X7),s(fun(X1,fun(X2,fun(X3,fun(X4,X5)))),X6))) = s(fun(X1,fun(X2,fun(X3,fun(X4,X0)))),happ(s(fun(fun(X5,X0),fun(X1,fun(X2,fun(X3,fun(X4,X0))))),'_dstsel_0_x01111_1'(s(fun(X1,fun(X2,fun(X3,fun(X4,X5)))),X6))),s(fun(X5,X0),X7)))) )),
  inference(cnf_transformation,[],[f672])).
fof(f672,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(fun(X1,fun(X2,fun(X3,fun(X4,X0)))),'_dst_x0000x1111_2'(s(fun(X5,X0),X7),s(fun(X1,fun(X2,fun(X3,fun(X4,X5)))),X6))) = s(fun(X1,fun(X2,fun(X3,fun(X4,X0)))),happ(s(fun(fun(X5,X0),fun(X1,fun(X2,fun(X3,fun(X4,X0))))),'_dstsel_0_x01111_1'(s(fun(X1,fun(X2,fun(X3,fun(X4,X5)))),X6))),s(fun(X5,X0),X7)))),
  inference(flattening,[],[f671])).
fof(f671,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ! [X6,X7] : s(fun(X1,fun(X2,fun(X3,fun(X4,X0)))),'_dst_x0000x1111_2'(s(fun(X5,X0),X7),s(fun(X1,fun(X2,fun(X3,fun(X4,X5)))),X6))) = s(fun(X1,fun(X2,fun(X3,fun(X4,X0)))),happ(s(fun(fun(X5,X0),fun(X1,fun(X2,fun(X3,fun(X4,X0))))),'_dstsel_0_x01111_1'(s(fun(X1,fun(X2,fun(X3,fun(X4,X5)))),X6))),s(fun(X5,X0),X7)))),
  inference(rectify,[],[f195])).
fof(f195,axiom,(
  ! [X35,X36,X37,X56,X57,X38] : ! [X40,X42] : s(fun(X36,fun(X37,fun(X56,fun(X57,X35)))),'_dst_x0000x1111_2'(s(fun(X38,X35),X42),s(fun(X36,fun(X37,fun(X56,fun(X57,X38)))),X40))) = s(fun(X36,fun(X37,fun(X56,fun(X57,X35)))),happ(s(fun(fun(X38,X35),fun(X36,fun(X37,fun(X56,fun(X57,X35))))),'_dstsel_0_x01111_1'(s(fun(X36,fun(X37,fun(X56,fun(X57,X38)))),X40))),s(fun(X38,X35),X42)))),
  file('thm.pred_set.EXTENSION','_dstsel_0_x01111_ax')).
fof(f1637,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s(X0,happ(s(fun(X4,X0),'_dst_x0x1x0x1_4'(s(fun(X1,fun(X2,fun(X3,X0))),X5),s(fun(X4,X1),X6),s(X2,X7),s(fun(X4,X3),X8))),s(X4,X9))) = s(X0,happ(s(fun(X3,X0),happ(s(fun(X2,fun(X3,X0)),happ(s(fun(X1,fun(X2,fun(X3,X0))),X5),s(X1,happ(s(fun(X4,X1),X6),s(X4,X9))))),s(X2,X7))),s(X3,happ(s(fun(X4,X3),X8),s(X4,X9)))))) )),
  inference(cnf_transformation,[],[f670])).
fof(f670,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : s(X0,happ(s(fun(X4,X0),'_dst_x0x1x0x1_4'(s(fun(X1,fun(X2,fun(X3,X0))),X5),s(fun(X4,X1),X6),s(X2,X7),s(fun(X4,X3),X8))),s(X4,X9))) = s(X0,happ(s(fun(X3,X0),happ(s(fun(X2,fun(X3,X0)),happ(s(fun(X1,fun(X2,fun(X3,X0))),X5),s(X1,happ(s(fun(X4,X1),X6),s(X4,X9))))),s(X2,X7))),s(X3,happ(s(fun(X4,X3),X8),s(X4,X9)))))),
  inference(flattening,[],[f669])).
fof(f669,plain,(
  ! [X0,X1,X2,X3,X4] : ! [X5,X6,X7,X8,X9] : s(X0,happ(s(fun(X4,X0),'_dst_x0x1x0x1_4'(s(fun(X1,fun(X2,fun(X3,X0))),X5),s(fun(X4,X1),X6),s(X2,X7),s(fun(X4,X3),X8))),s(X4,X9))) = s(X0,happ(s(fun(X3,X0),happ(s(fun(X2,fun(X3,X0)),happ(s(fun(X1,fun(X2,fun(X3,X0))),X5),s(X1,happ(s(fun(X4,X1),X6),s(X4,X9))))),s(X2,X7))),s(X3,happ(s(fun(X4,X3),X8),s(X4,X9)))))),
  inference(rectify,[],[f163])).
fof(f163,axiom,(
  ! [X35,X36,X37,X56,X38] : ! [X39,X40,X41,X63,X42] : s(X35,happ(s(fun(X38,X35),'_dst_x0x1x0x1_4'(s(fun(X36,fun(X37,fun(X56,X35))),X39),s(fun(X38,X36),X40),s(X37,X41),s(fun(X38,X56),X63))),s(X38,X42))) = s(X35,happ(s(fun(X56,X35),happ(s(fun(X37,fun(X56,X35)),happ(s(fun(X36,fun(X37,fun(X56,X35))),X39),s(X36,happ(s(fun(X38,X36),X40),s(X38,X42))))),s(X37,X41))),s(X56,happ(s(fun(X38,X56),X63),s(X38,X42)))))),
  file('thm.pred_set.EXTENSION','_dst_x0x1x0x1_ax')).
fof(f1636,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(fun(X2,X0),happ(s(fun(X4,fun(X2,X0)),'_dst_x00x11x11_3'(s(fun(X1,fun(X3,X0)),X5),s(fun(X4,fun(X2,X1)),X6),s(fun(X4,fun(X2,X3)),X7))),s(X4,X8))) = s(fun(X2,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X3,X0)),X5),s(fun(X2,X1),happ(s(fun(X4,fun(X2,X1)),X6),s(X4,X8))),s(fun(X2,X3),happ(s(fun(X4,fun(X2,X3)),X7),s(X4,X8)))))) )),
  inference(cnf_transformation,[],[f668])).
fof(f668,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : s(fun(X2,X0),happ(s(fun(X4,fun(X2,X0)),'_dst_x00x11x11_3'(s(fun(X1,fun(X3,X0)),X5),s(fun(X4,fun(X2,X1)),X6),s(fun(X4,fun(X2,X3)),X7))),s(X4,X8))) = s(fun(X2,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X3,X0)),X5),s(fun(X2,X1),happ(s(fun(X4,fun(X2,X1)),X6),s(X4,X8))),s(fun(X2,X3),happ(s(fun(X4,fun(X2,X3)),X7),s(X4,X8)))))),
  inference(flattening,[],[f667])).
fof(f667,plain,(
  ! [X0,X1,X2,X3,X4] : ! [X5,X6,X7,X8] : s(fun(X2,X0),happ(s(fun(X4,fun(X2,X0)),'_dst_x00x11x11_3'(s(fun(X1,fun(X3,X0)),X5),s(fun(X4,fun(X2,X1)),X6),s(fun(X4,fun(X2,X3)),X7))),s(X4,X8))) = s(fun(X2,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X3,X0)),X5),s(fun(X2,X1),happ(s(fun(X4,fun(X2,X1)),X6),s(X4,X8))),s(fun(X2,X3),happ(s(fun(X4,fun(X2,X3)),X7),s(X4,X8)))))),
  inference(rectify,[],[f114])).
fof(f114,axiom,(
  ! [X35,X37,X36,X56,X38] : ! [X39,X40,X41,X42] : s(fun(X36,X35),happ(s(fun(X38,fun(X36,X35)),'_dst_x00x11x11_3'(s(fun(X37,fun(X56,X35)),X39),s(fun(X38,fun(X36,X37)),X40),s(fun(X38,fun(X36,X56)),X41))),s(X38,X42))) = s(fun(X36,X35),'_dst_x0x1x1_3'(s(fun(X37,fun(X56,X35)),X39),s(fun(X36,X37),happ(s(fun(X38,fun(X36,X37)),X40),s(X38,X42))),s(fun(X36,X56),happ(s(fun(X38,fun(X36,X56)),X41),s(X38,X42)))))),
  file('thm.pred_set.EXTENSION','_dst_x00x11x11_ax')).
fof(f1635,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(fun(X1,X0),happ(s(fun(X4,fun(X1,X0)),'_dst_x00x11x10_3'(s(fun(X2,fun(X3,X0)),X5),s(fun(X4,fun(X1,X2)),X6),s(fun(X4,X3),X7))),s(X4,X8))) = s(fun(X1,X0),'_dst_x0x1x0_3'(s(fun(X2,fun(X3,X0)),X5),s(fun(X1,X2),happ(s(fun(X4,fun(X1,X2)),X6),s(X4,X8))),s(X3,happ(s(fun(X4,X3),X7),s(X4,X8)))))) )),
  inference(cnf_transformation,[],[f666])).
fof(f666,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : s(fun(X1,X0),happ(s(fun(X4,fun(X1,X0)),'_dst_x00x11x10_3'(s(fun(X2,fun(X3,X0)),X5),s(fun(X4,fun(X1,X2)),X6),s(fun(X4,X3),X7))),s(X4,X8))) = s(fun(X1,X0),'_dst_x0x1x0_3'(s(fun(X2,fun(X3,X0)),X5),s(fun(X1,X2),happ(s(fun(X4,fun(X1,X2)),X6),s(X4,X8))),s(X3,happ(s(fun(X4,X3),X7),s(X4,X8)))))),
  inference(flattening,[],[f665])).
fof(f665,plain,(
  ! [X0,X1,X2,X3,X4] : ! [X5,X6,X7,X8] : s(fun(X1,X0),happ(s(fun(X4,fun(X1,X0)),'_dst_x00x11x10_3'(s(fun(X2,fun(X3,X0)),X5),s(fun(X4,fun(X1,X2)),X6),s(fun(X4,X3),X7))),s(X4,X8))) = s(fun(X1,X0),'_dst_x0x1x0_3'(s(fun(X2,fun(X3,X0)),X5),s(fun(X1,X2),happ(s(fun(X4,fun(X1,X2)),X6),s(X4,X8))),s(X3,happ(s(fun(X4,X3),X7),s(X4,X8)))))),
  inference(rectify,[],[f190])).
fof(f190,axiom,(
  ! [X35,X36,X37,X56,X38] : ! [X39,X40,X41,X42] : s(fun(X36,X35),happ(s(fun(X38,fun(X36,X35)),'_dst_x00x11x10_3'(s(fun(X37,fun(X56,X35)),X39),s(fun(X38,fun(X36,X37)),X40),s(fun(X38,X56),X41))),s(X38,X42))) = s(fun(X36,X35),'_dst_x0x1x0_3'(s(fun(X37,fun(X56,X35)),X39),s(fun(X36,X37),happ(s(fun(X38,fun(X36,X37)),X40),s(X38,X42))),s(X56,happ(s(fun(X38,X56),X41),s(X38,X42)))))),
  file('thm.pred_set.EXTENSION','_dst_x00x11x10_ax')).
fof(f1634,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(fun(X2,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X3,X0)),X5),s(fun(X2,X1),X6),s(fun(X2,X3),happ(s(fun(X4,fun(X2,X3)),X7),s(X4,X8))))) = s(fun(X2,X0),happ(s(fun(X4,fun(X2,X0)),'_dst_x00x01x11_3'(s(fun(X1,fun(X3,X0)),X5),s(fun(X2,X1),X6),s(fun(X4,fun(X2,X3)),X7))),s(X4,X8)))) )),
  inference(cnf_transformation,[],[f664])).
fof(f664,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : s(fun(X2,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X3,X0)),X5),s(fun(X2,X1),X6),s(fun(X2,X3),happ(s(fun(X4,fun(X2,X3)),X7),s(X4,X8))))) = s(fun(X2,X0),happ(s(fun(X4,fun(X2,X0)),'_dst_x00x01x11_3'(s(fun(X1,fun(X3,X0)),X5),s(fun(X2,X1),X6),s(fun(X4,fun(X2,X3)),X7))),s(X4,X8)))),
  inference(flattening,[],[f663])).
fof(f663,plain,(
  ! [X0,X1,X2,X3,X4] : ! [X5,X6,X7,X8] : s(fun(X2,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X3,X0)),X5),s(fun(X2,X1),X6),s(fun(X2,X3),happ(s(fun(X4,fun(X2,X3)),X7),s(X4,X8))))) = s(fun(X2,X0),happ(s(fun(X4,fun(X2,X0)),'_dst_x00x01x11_3'(s(fun(X1,fun(X3,X0)),X5),s(fun(X2,X1),X6),s(fun(X4,fun(X2,X3)),X7))),s(X4,X8)))),
  inference(rectify,[],[f198])).
fof(f198,axiom,(
  ! [X35,X37,X36,X56,X38] : ! [X39,X40,X41,X42] : s(fun(X36,X35),'_dst_x0x1x1_3'(s(fun(X37,fun(X56,X35)),X39),s(fun(X36,X37),X40),s(fun(X36,X56),happ(s(fun(X38,fun(X36,X56)),X41),s(X38,X42))))) = s(fun(X36,X35),happ(s(fun(X38,fun(X36,X35)),'_dst_x00x01x11_3'(s(fun(X37,fun(X56,X35)),X39),s(fun(X36,X37),X40),s(fun(X38,fun(X36,X56)),X41))),s(X38,X42)))),
  file('thm.pred_set.EXTENSION','_dst_x00x01x11_ax')).
fof(f1633,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(fun(X2,X0),'_dst_x0x0x1_3'(s(fun(X1,fun(X3,X0)),X5),s(X1,X6),s(fun(X2,X3),happ(s(fun(X4,fun(X2,X3)),X7),s(X4,X8))))) = s(fun(X2,X0),happ(s(fun(X4,fun(X2,X0)),'_dst_x00x00x11_3'(s(fun(X1,fun(X3,X0)),X5),s(X1,X6),s(fun(X4,fun(X2,X3)),X7))),s(X4,X8)))) )),
  inference(cnf_transformation,[],[f662])).
fof(f662,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : s(fun(X2,X0),'_dst_x0x0x1_3'(s(fun(X1,fun(X3,X0)),X5),s(X1,X6),s(fun(X2,X3),happ(s(fun(X4,fun(X2,X3)),X7),s(X4,X8))))) = s(fun(X2,X0),happ(s(fun(X4,fun(X2,X0)),'_dst_x00x00x11_3'(s(fun(X1,fun(X3,X0)),X5),s(X1,X6),s(fun(X4,fun(X2,X3)),X7))),s(X4,X8)))),
  inference(flattening,[],[f661])).
fof(f661,plain,(
  ! [X0,X1,X2,X3,X4] : ! [X5,X6,X7,X8] : s(fun(X2,X0),'_dst_x0x0x1_3'(s(fun(X1,fun(X3,X0)),X5),s(X1,X6),s(fun(X2,X3),happ(s(fun(X4,fun(X2,X3)),X7),s(X4,X8))))) = s(fun(X2,X0),happ(s(fun(X4,fun(X2,X0)),'_dst_x00x00x11_3'(s(fun(X1,fun(X3,X0)),X5),s(X1,X6),s(fun(X4,fun(X2,X3)),X7))),s(X4,X8)))),
  inference(rectify,[],[f162])).
fof(f162,axiom,(
  ! [X35,X37,X36,X56,X38] : ! [X39,X40,X41,X42] : s(fun(X36,X35),'_dst_x0x0x1_3'(s(fun(X37,fun(X56,X35)),X39),s(X37,X40),s(fun(X36,X56),happ(s(fun(X38,fun(X36,X56)),X41),s(X38,X42))))) = s(fun(X36,X35),happ(s(fun(X38,fun(X36,X35)),'_dst_x00x00x11_3'(s(fun(X37,fun(X56,X35)),X39),s(X37,X40),s(fun(X38,fun(X36,X56)),X41))),s(X38,X42)))),
  file('thm.pred_set.EXTENSION','_dst_x00x00x11_ax')).
fof(f1632,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(fun(X3,X0),'_dst_x0x0x1_3'(s(fun(X1,fun(X4,X0)),X5),s(X1,happ(s(fun(X2,X1),X6),s(X2,X8))),s(fun(X3,X4),X7))) = s(fun(X3,X0),happ(s(fun(X2,fun(X3,X0)),'_dst_x00x10x01_3'(s(fun(X1,fun(X4,X0)),X5),s(fun(X2,X1),X6),s(fun(X3,X4),X7))),s(X2,X8)))) )),
  inference(cnf_transformation,[],[f660])).
fof(f660,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : s(fun(X3,X0),'_dst_x0x0x1_3'(s(fun(X1,fun(X4,X0)),X5),s(X1,happ(s(fun(X2,X1),X6),s(X2,X8))),s(fun(X3,X4),X7))) = s(fun(X3,X0),happ(s(fun(X2,fun(X3,X0)),'_dst_x00x10x01_3'(s(fun(X1,fun(X4,X0)),X5),s(fun(X2,X1),X6),s(fun(X3,X4),X7))),s(X2,X8)))),
  inference(flattening,[],[f659])).
fof(f659,plain,(
  ! [X0,X1,X2,X3,X4] : ! [X5,X6,X7,X8] : s(fun(X3,X0),'_dst_x0x0x1_3'(s(fun(X1,fun(X4,X0)),X5),s(X1,happ(s(fun(X2,X1),X6),s(X2,X8))),s(fun(X3,X4),X7))) = s(fun(X3,X0),happ(s(fun(X2,fun(X3,X0)),'_dst_x00x10x01_3'(s(fun(X1,fun(X4,X0)),X5),s(fun(X2,X1),X6),s(fun(X3,X4),X7))),s(X2,X8)))),
  inference(rectify,[],[f112])).
fof(f112,axiom,(
  ! [X35,X37,X38,X36,X56] : ! [X39,X40,X41,X42] : s(fun(X36,X35),'_dst_x0x0x1_3'(s(fun(X37,fun(X56,X35)),X39),s(X37,happ(s(fun(X38,X37),X40),s(X38,X42))),s(fun(X36,X56),X41))) = s(fun(X36,X35),happ(s(fun(X38,fun(X36,X35)),'_dst_x00x10x01_3'(s(fun(X37,fun(X56,X35)),X39),s(fun(X38,X37),X40),s(fun(X36,X56),X41))),s(X38,X42)))),
  file('thm.pred_set.EXTENSION','_dst_x00x10x01_ax')).
fof(f1631,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(fun(X1,fun(X2,X0)),'_dst_x00x11_2'(s(fun(X3,X0),X5),s(fun(X1,fun(X2,X3)),happ(s(fun(X4,fun(X1,fun(X2,X3))),X6),s(X4,X7))))) = s(fun(X1,fun(X2,X0)),happ(s(fun(X4,fun(X1,fun(X2,X0))),'_dst_x000x111_2'(s(fun(X3,X0),X5),s(fun(X4,fun(X1,fun(X2,X3))),X6))),s(X4,X7)))) )),
  inference(cnf_transformation,[],[f658])).
fof(f658,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(fun(X1,fun(X2,X0)),'_dst_x00x11_2'(s(fun(X3,X0),X5),s(fun(X1,fun(X2,X3)),happ(s(fun(X4,fun(X1,fun(X2,X3))),X6),s(X4,X7))))) = s(fun(X1,fun(X2,X0)),happ(s(fun(X4,fun(X1,fun(X2,X0))),'_dst_x000x111_2'(s(fun(X3,X0),X5),s(fun(X4,fun(X1,fun(X2,X3))),X6))),s(X4,X7)))),
  inference(flattening,[],[f657])).
fof(f657,plain,(
  ! [X0,X1,X2,X3,X4] : ! [X5,X6,X7] : s(fun(X1,fun(X2,X0)),'_dst_x00x11_2'(s(fun(X3,X0),X5),s(fun(X1,fun(X2,X3)),happ(s(fun(X4,fun(X1,fun(X2,X3))),X6),s(X4,X7))))) = s(fun(X1,fun(X2,X0)),happ(s(fun(X4,fun(X1,fun(X2,X0))),'_dst_x000x111_2'(s(fun(X3,X0),X5),s(fun(X4,fun(X1,fun(X2,X3))),X6))),s(X4,X7)))),
  inference(rectify,[],[f109])).
fof(f109,axiom,(
  ! [X35,X36,X37,X56,X38] : ! [X39,X40,X42] : s(fun(X36,fun(X37,X35)),'_dst_x00x11_2'(s(fun(X56,X35),X39),s(fun(X36,fun(X37,X56)),happ(s(fun(X38,fun(X36,fun(X37,X56))),X40),s(X38,X42))))) = s(fun(X36,fun(X37,X35)),happ(s(fun(X38,fun(X36,fun(X37,X35))),'_dst_x000x111_2'(s(fun(X56,X35),X39),s(fun(X38,fun(X36,fun(X37,X56))),X40))),s(X38,X42)))),
  file('thm.pred_set.EXTENSION','_dst_x000x111_ax')).
fof(f1630,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(fun(X1,fun(X2,fun(X3,X0))),'_dst_x000x111_2'(s(fun(X4,X0),X6),s(fun(X1,fun(X2,fun(X3,X4))),X5))) = s(fun(X1,fun(X2,fun(X3,X0))),happ(s(fun(fun(X4,X0),fun(X1,fun(X2,fun(X3,X0)))),'_dstsel_0_x0111_1'(s(fun(X1,fun(X2,fun(X3,X4))),X5))),s(fun(X4,X0),X6)))) )),
  inference(cnf_transformation,[],[f656])).
fof(f656,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(fun(X1,fun(X2,fun(X3,X0))),'_dst_x000x111_2'(s(fun(X4,X0),X6),s(fun(X1,fun(X2,fun(X3,X4))),X5))) = s(fun(X1,fun(X2,fun(X3,X0))),happ(s(fun(fun(X4,X0),fun(X1,fun(X2,fun(X3,X0)))),'_dstsel_0_x0111_1'(s(fun(X1,fun(X2,fun(X3,X4))),X5))),s(fun(X4,X0),X6)))),
  inference(flattening,[],[f655])).
fof(f655,plain,(
  ! [X0,X1,X2,X3,X4] : ! [X5,X6] : s(fun(X1,fun(X2,fun(X3,X0))),'_dst_x000x111_2'(s(fun(X4,X0),X6),s(fun(X1,fun(X2,fun(X3,X4))),X5))) = s(fun(X1,fun(X2,fun(X3,X0))),happ(s(fun(fun(X4,X0),fun(X1,fun(X2,fun(X3,X0)))),'_dstsel_0_x0111_1'(s(fun(X1,fun(X2,fun(X3,X4))),X5))),s(fun(X4,X0),X6)))),
  inference(rectify,[],[f178])).
fof(f178,axiom,(
  ! [X35,X36,X37,X56,X38] : ! [X40,X42] : s(fun(X36,fun(X37,fun(X56,X35))),'_dst_x000x111_2'(s(fun(X38,X35),X42),s(fun(X36,fun(X37,fun(X56,X38))),X40))) = s(fun(X36,fun(X37,fun(X56,X35))),happ(s(fun(fun(X38,X35),fun(X36,fun(X37,fun(X56,X35)))),'_dstsel_0_x0111_1'(s(fun(X36,fun(X37,fun(X56,X38))),X40))),s(fun(X38,X35),X42)))),
  file('thm.pred_set.EXTENSION','_dstsel_0_x0111_ax')).
fof(f1629,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun(X3,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X2,X0)),X4),s(fun(X3,X1),X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,happ(s(fun(X3,X1),X5),s(X3,X7))))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))) )),
  inference(cnf_transformation,[],[f654])).
fof(f654,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun(X3,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X2,X0)),X4),s(fun(X3,X1),X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,happ(s(fun(X3,X1),X5),s(X3,X7))))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))),
  inference(flattening,[],[f653])).
fof(f653,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6,X7] : s(X0,happ(s(fun(X3,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X2,X0)),X4),s(fun(X3,X1),X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,happ(s(fun(X3,X1),X5),s(X3,X7))))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))),
  inference(rectify,[],[f62])).
fof(f62,axiom,(
  ! [X35,X36,X37,X38] : ! [X39,X40,X41,X42] : s(X35,happ(s(fun(X38,X35),'_dst_x0x1x1_3'(s(fun(X36,fun(X37,X35)),X39),s(fun(X38,X36),X40),s(fun(X38,X37),X41))),s(X38,X42))) = s(X35,happ(s(fun(X37,X35),happ(s(fun(X36,fun(X37,X35)),X39),s(X36,happ(s(fun(X38,X36),X40),s(X38,X42))))),s(X37,happ(s(fun(X38,X37),X41),s(X38,X42)))))),
  file('thm.pred_set.EXTENSION','_dst_x0x1x1_ax')).
fof(f1628,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun(X3,X0),'_dst_x0x0x1_3'(s(fun(X1,fun(X2,X0)),X4),s(X1,X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,X5))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))) )),
  inference(cnf_transformation,[],[f652])).
fof(f652,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun(X3,X0),'_dst_x0x0x1_3'(s(fun(X1,fun(X2,X0)),X4),s(X1,X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,X5))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))),
  inference(flattening,[],[f651])).
fof(f651,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6,X7] : s(X0,happ(s(fun(X3,X0),'_dst_x0x0x1_3'(s(fun(X1,fun(X2,X0)),X4),s(X1,X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,X5))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))),
  inference(rectify,[],[f111])).
fof(f111,axiom,(
  ! [X35,X36,X37,X38] : ! [X39,X40,X41,X42] : s(X35,happ(s(fun(X38,X35),'_dst_x0x0x1_3'(s(fun(X36,fun(X37,X35)),X39),s(X36,X40),s(fun(X38,X37),X41))),s(X38,X42))) = s(X35,happ(s(fun(X37,X35),happ(s(fun(X36,fun(X37,X35)),X39),s(X36,X40))),s(X37,happ(s(fun(X38,X37),X41),s(X38,X42)))))),
  file('thm.pred_set.EXTENSION','_dst_x0x0x1_ax')).
fof(f1627,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun(X2,X0),'_dst_x0x1x0_3'(s(fun(X1,fun(X3,X0)),X4),s(fun(X2,X1),X5),s(X3,X6))),s(X2,X7))) = s(X0,happ(s(fun(X3,X0),happ(s(fun(X1,fun(X3,X0)),X4),s(X1,happ(s(fun(X2,X1),X5),s(X2,X7))))),s(X3,X6)))) )),
  inference(cnf_transformation,[],[f650])).
fof(f650,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun(X2,X0),'_dst_x0x1x0_3'(s(fun(X1,fun(X3,X0)),X4),s(fun(X2,X1),X5),s(X3,X6))),s(X2,X7))) = s(X0,happ(s(fun(X3,X0),happ(s(fun(X1,fun(X3,X0)),X4),s(X1,happ(s(fun(X2,X1),X5),s(X2,X7))))),s(X3,X6)))),
  inference(flattening,[],[f649])).
fof(f649,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6,X7] : s(X0,happ(s(fun(X2,X0),'_dst_x0x1x0_3'(s(fun(X1,fun(X3,X0)),X4),s(fun(X2,X1),X5),s(X3,X6))),s(X2,X7))) = s(X0,happ(s(fun(X3,X0),happ(s(fun(X1,fun(X3,X0)),X4),s(X1,happ(s(fun(X2,X1),X5),s(X2,X7))))),s(X3,X6)))),
  inference(rectify,[],[f197])).
fof(f197,axiom,(
  ! [X35,X36,X38,X37] : ! [X39,X40,X41,X42] : s(X35,happ(s(fun(X38,X35),'_dst_x0x1x0_3'(s(fun(X36,fun(X37,X35)),X39),s(fun(X38,X36),X40),s(X37,X41))),s(X38,X42))) = s(X35,happ(s(fun(X37,X35),happ(s(fun(X36,fun(X37,X35)),X39),s(X36,happ(s(fun(X38,X36),X40),s(X38,X42))))),s(X37,X41)))),
  file('thm.pred_set.EXTENSION','_dst_x0x1x0_ax')).
fof(f1626,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X5),s(X3,X6))))) = s(fun(X1,X0),happ(s(fun(X3,fun(X1,X0)),'_dst_x00x11_2'(s(fun(X2,X0),X4),s(fun(X3,fun(X1,X2)),X5))),s(X3,X6)))) )),
  inference(cnf_transformation,[],[f648])).
fof(f648,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X5),s(X3,X6))))) = s(fun(X1,X0),happ(s(fun(X3,fun(X1,X0)),'_dst_x00x11_2'(s(fun(X2,X0),X4),s(fun(X3,fun(X1,X2)),X5))),s(X3,X6)))),
  inference(flattening,[],[f647])).
fof(f647,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6] : s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X5),s(X3,X6))))) = s(fun(X1,X0),happ(s(fun(X3,fun(X1,X0)),'_dst_x00x11_2'(s(fun(X2,X0),X4),s(fun(X3,fun(X1,X2)),X5))),s(X3,X6)))),
  inference(rectify,[],[f110])).
fof(f110,axiom,(
  ! [X35,X36,X37,X38] : ! [X39,X40,X42] : s(fun(X36,X35),'_dst_x0x1_2'(s(fun(X37,X35),X39),s(fun(X36,X37),happ(s(fun(X38,fun(X36,X37)),X40),s(X38,X42))))) = s(fun(X36,X35),happ(s(fun(X38,fun(X36,X35)),'_dst_x00x11_2'(s(fun(X37,X35),X39),s(fun(X38,fun(X36,X37)),X40))),s(X38,X42)))),
  file('thm.pred_set.EXTENSION','_dst_x00x11_ax')).
fof(f1625,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(fun(X0,X1),'_dst_x1x0_2'(s(fun(X0,fun(X2,X1)),X4),s(X2,happ(s(fun(X3,X2),X5),s(X3,X6))))) = s(fun(X0,X1),happ(s(fun(X3,fun(X0,X1)),'_dst_x01x10_2'(s(fun(X0,fun(X2,X1)),X4),s(fun(X3,X2),X5))),s(X3,X6)))) )),
  inference(cnf_transformation,[],[f646])).
fof(f646,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(fun(X0,X1),'_dst_x1x0_2'(s(fun(X0,fun(X2,X1)),X4),s(X2,happ(s(fun(X3,X2),X5),s(X3,X6))))) = s(fun(X0,X1),happ(s(fun(X3,fun(X0,X1)),'_dst_x01x10_2'(s(fun(X0,fun(X2,X1)),X4),s(fun(X3,X2),X5))),s(X3,X6)))),
  inference(flattening,[],[f645])).
fof(f645,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6] : s(fun(X0,X1),'_dst_x1x0_2'(s(fun(X0,fun(X2,X1)),X4),s(X2,happ(s(fun(X3,X2),X5),s(X3,X6))))) = s(fun(X0,X1),happ(s(fun(X3,fun(X0,X1)),'_dst_x01x10_2'(s(fun(X0,fun(X2,X1)),X4),s(fun(X3,X2),X5))),s(X3,X6)))),
  inference(rectify,[],[f119])).
fof(f119,axiom,(
  ! [X36,X35,X37,X38] : ! [X39,X40,X42] : s(fun(X36,X35),'_dst_x1x0_2'(s(fun(X36,fun(X37,X35)),X39),s(X37,happ(s(fun(X38,X37),X40),s(X38,X42))))) = s(fun(X36,X35),happ(s(fun(X38,fun(X36,X35)),'_dst_x01x10_2'(s(fun(X36,fun(X37,X35)),X39),s(fun(X38,X37),X40))),s(X38,X42)))),
  file('thm.pred_set.EXTENSION','_dst_x01x10_ax')).
fof(f1616,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),'_dst_x1x0_2'(s(fun(X1,fun(X2,X0)),X3),s(X2,X4))),s(X1,X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X4)))) )),
  inference(cnf_transformation,[],[f643])).
fof(f643,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X1,X0),'_dst_x1x0_2'(s(fun(X1,fun(X2,X0)),X3),s(X2,X4))),s(X1,X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X4)))),
  inference(flattening,[],[f642])).
fof(f642,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,happ(s(fun(X1,X0),'_dst_x1x0_2'(s(fun(X1,fun(X2,X0)),X3),s(X2,X4))),s(X1,X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X4)))),
  inference(rectify,[],[f66])).
fof(f66,axiom,(
  ! [X35,X38,X36] : ! [X39,X40,X42] : s(X35,happ(s(fun(X36,X35),happ(s(fun(X38,fun(X36,X35)),X39),s(X38,X42))),s(X36,X40))) = s(X35,happ(s(fun(X38,X35),'_dst_x1x0_2'(s(fun(X38,fun(X36,X35)),X39),s(X36,X40))),s(X38,X42)))),
  file('thm.pred_set.EXTENSION','_dst_x1x0_ax')).
fof(f1615,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))) )),
  inference(cnf_transformation,[],[f641])).
fof(f641,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(flattening,[],[f640])).
fof(f640,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(rectify,[],[f67])).
fof(f67,axiom,(
  ! [X35,X36,X38] : ! [X39,X40,X42] : s(X35,happ(s(fun(X36,X35),X39),s(X36,happ(s(fun(X38,X36),X40),s(X38,X42))))) = s(X35,happ(s(fun(X38,X35),'_dst_x0x1_2'(s(fun(X36,X35),X39),s(fun(X38,X36),X40))),s(X38,X42)))),
  file('thm.pred_set.EXTENSION','_dst_x0x1_ax')).
fof(f1903,plain,(
  s(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.num.num'),'const.arithmetic.findq_0') = s(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.num.num'),'const.relation.WFREC_2'(s(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),bool)),'const.prim_rec.measure_1'(s(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.num.num'),'const.pair.UNCURRY_1'(s(fun('type.num.num',fun('type.pair.prod'('type.num.num','type.num.num'),'type.num.num')),'_dst_x0_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.num.num'),'const.pair.UNCURRY_1'(s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.-_0'))))))))),s(fun(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.num.num'),fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.num.num')),'_dst_x0x1_2'(s(fun(fun('type.num.num',fun('type.pair.prod'('type.num.num','type.num.num'),'type.num.num')),fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.num.num')),'const.pair.UNCURRY_0'),s(fun(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.num.num'),fun('type.num.num',fun('type.pair.prod'('type.num.num','type.num.num'),'type.num.num'))),'_dst_x00x11_2'(s(fun(fun('type.num.num',fun('type.num.num','type.num.num')),fun('type.pair.prod'('type.num.num','type.num.num'),'type.num.num')),'const.pair.UNCURRY_0'),s(fun(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.num.num'),fun('type.num.num',fun('type.num.num',fun('type.num.num','type.num.num')))),'_dst_x0000x0001x0100x1111_4'(s(fun(bool,fun('type.num.num',fun('type.num.num','type.num.num'))),'const.bool.COND_0'),s(fun('type.num.num',bool),'_dst_x1x0_2'(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.num.|48|_0'))),s(fun('type.num.num','type.num.num'),'_i_0_0'),s(fun(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.num.num'),fun('type.num.num',fun('type.num.num',fun('type.num.num','type.num.num')))),'_dst_x0000x1110x0001_3'(s(fun(fun('type.num.num','type.num.num'),fun('type.num.num','type.num.num')),'const.bool.LET_0'),s(fun(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.num.num'),fun('type.num.num',fun('type.num.num',fun('type.num.num','type.num.num')))),'_dst_x0000x0011x0100x1111_4'(s(fun(bool,fun('type.num.num',fun('type.num.num','type.num.num'))),'const.bool.COND_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0'),s(fun('type.num.num','type.num.num'),'_i_0_0'),s(fun(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.num.num'),fun('type.num.num',fun('type.num.num',fun('type.num.num','type.num.num')))),'_dstsel_0_x0111_1'(s(fun('type.num.num',fun('type.num.num',fun('type.num.num','type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num'))))),'_dst_x000x100x011_3'(s(fun('type.num.num',fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')))),'const.pair.,_0'),s(fun('type.num.num','type.num.num'),'const.arithmetic.*_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s(fun('type.num.num',fun('type.num.num','type.pair.prod'('type.num.num','type.num.num'))),'const.pair.,_0'))))))),s(fun('type.num.num','type.num.num'),'const.arithmetic.*_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))))))))),
  inference(forward_demodulation,[],[f1902,f1614])).
fof(f1902,plain,(
  s(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.num.num'),'const.arithmetic.findq_0') = s(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.num.num'),'const.relation.WFREC_2'(s(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),bool)),'const.prim_rec.measure_1'(s(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.num.num'),'const.pair.UNCURRY_1'(s(fun('type.num.num',fun('type.pair.prod'('type.num.num','type.num.num'),'type.num.num')),'_dst_x0_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.num.num'),happ(s(fun(fun('type.num.num',fun('type.num.num','type.num.num')),fun('type.pair.prod'('type.num.num','type.num.num'),'type.num.num')),'const.pair.UNCURRY_0'),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.-_0'))))))))),s(fun(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.num.num'),fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.num.num')),'_dst_x0x1_2'(s(fun(fun('type.num.num',fun('type.pair.prod'('type.num.num','type.num.num'),'type.num.num')),fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.num.num')),'const.pair.UNCURRY_0'),s(fun(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.num.num'),fun('type.num.num',fun('type.pair.prod'('type.num.num','type.num.num'),'type.num.num'))),'_dst_x00x11_2'(s(fun(fun('type.num.num',fun('type.num.num','type.num.num')),fun('type.pair.prod'('type.num.num','type.num.num'),'type.num.num')),'const.pair.UNCURRY_0'),s(fun(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.num.num'),fun('type.num.num',fun('type.num.num',fun('type.num.num','type.num.num')))),'_dst_x0000x0001x0100x1111_4'(s(fun(bool,fun('type.num.num',fun('type.num.num','type.num.num'))),'const.bool.COND_0'),s(fun('type.num.num',bool),'_dst_x1x0_2'(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.num.|48|_0'))),s(fun('type.num.num','type.num.num'),'_i_0_0'),s(fun(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.num.num'),fun('type.num.num',fun('type.num.num',fun('type.num.num','type.num.num')))),'_dst_x0000x1110x0001_3'(s(fun(fun('type.num.num','type.num.num'),fun('type.num.num','type.num.num')),'const.bool.LET_0'),s(fun(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.num.num'),fun('type.num.num',fun('type.num.num',fun('type.num.num','type.num.num')))),'_dst_x0000x0011x0100x1111_4'(s(fun(bool,fun('type.num.num',fun('type.num.num','type.num.num'))),'const.bool.COND_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0'),s(fun('type.num.num','type.num.num'),'_i_0_0'),s(fun(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.num.num'),fun('type.num.num',fun('type.num.num',fun('type.num.num','type.num.num')))),'_dstsel_0_x0111_1'(s(fun('type.num.num',fun('type.num.num',fun('type.num.num','type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num'))))),'_dst_x000x100x011_3'(s(fun('type.num.num',fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')))),'const.pair.,_0'),s(fun('type.num.num','type.num.num'),'const.arithmetic.*_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s(fun('type.num.num',fun('type.num.num','type.pair.prod'('type.num.num','type.num.num'))),'const.pair.,_0'))))))),s(fun('type.num.num','type.num.num'),'const.arithmetic.*_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))))))))),
  inference(backward_demodulation,[],[f1614,f1106])).
fof(f1614,plain,(
  ( ! [X2,X0,X3,X1] : (s(fun('type.pair.prod'(X0,X1),X2),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,X2)),X3))) = s(fun('type.pair.prod'(X0,X1),X2),happ(s(fun(fun(X0,fun(X1,X2)),fun('type.pair.prod'(X0,X1),X2)),'const.pair.UNCURRY_0'),s(fun(X0,fun(X1,X2)),X3)))) )),
  inference(cnf_transformation,[],[f639])).
fof(f639,plain,(
  ! [X0,X1,X2,X3] : s(fun('type.pair.prod'(X0,X1),X2),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,X2)),X3))) = s(fun('type.pair.prod'(X0,X1),X2),happ(s(fun(fun(X0,fun(X1,X2)),fun('type.pair.prod'(X0,X1),X2)),'const.pair.UNCURRY_0'),s(fun(X0,fun(X1,X2)),X3)))),
  inference(flattening,[],[f638])).
fof(f638,plain,(
  ! [X0,X1,X2] : ! [X3] : s(fun('type.pair.prod'(X0,X1),X2),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,X2)),X3))) = s(fun('type.pair.prod'(X0,X1),X2),happ(s(fun(fun(X0,fun(X1,X2)),fun('type.pair.prod'(X0,X1),X2)),'const.pair.UNCURRY_0'),s(fun(X0,fun(X1,X2)),X3)))),
  inference(rectify,[],[f167])).
fof(f167,axiom,(
  ! [X6,X20,X48] : ! [X21] : s(fun('type.pair.prod'(X6,X20),X48),'const.pair.UNCURRY_1'(s(fun(X6,fun(X20,X48)),X21))) = s(fun('type.pair.prod'(X6,X20),X48),happ(s(fun(fun(X6,fun(X20,X48)),fun('type.pair.prod'(X6,X20),X48)),'const.pair.UNCURRY_0'),s(fun(X6,fun(X20,X48)),X21)))),
  file('thm.pred_set.EXTENSION','ARITY_0_1_const.pair.UNCURRY')).
fof(f1599,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f1094])).
fof(f1582,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f1086])).
fof(f1586,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f1086])).
fof(f1563,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f1079])).
fof(f1540,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f1067])).
fof(f1541,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP3(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f1067])).
fof(f1542,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP4(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f1067])).
fof(f1546,plain,(
  ( ! [X2,X0,X1] : (~sP2(X2,X1,X0) | p(s(bool,X0)) | ~sP4(X0,X2,X1) | ~sP3(X1,X2,X0) | sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f1067])).
fof(f1901,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1)))))) | s(X0,X1) = s(X0,X2)) )),
  inference(forward_demodulation,[],[f1520,f1257])).
fof(f1520,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,X2) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2))))) )),
  inference(cnf_transformation,[],[f1058])).
fof(f1058,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2))))))),
  inference(nnf_transformation,[],[f626])).
fof(f626,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f65])).
fof(f65,axiom,(
  ! [X43,X44,X45] : (p(s(bool,happ(s(fun(X43,bool),happ(s(fun(X43,fun(X43,bool)),'=_0'),s(X43,X44))),s(X43,X45)))) <=> s(X43,X44) = s(X43,X45))),
  file('thm.pred_set.EXTENSION',arity_0_HL_EQ)).
fof(f1900,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1)))))) | s(X0,X1) != s(X0,X2)) )),
  inference(forward_demodulation,[],[f1521,f1257])).
fof(f1521,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) != s(X0,X2)) )),
  inference(cnf_transformation,[],[f1058])).
fof(f1519,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun(fun(X0,fun(X1,bool)),bool),'_dst_x0x1_2'(s(fun(fun(X0,bool),bool),'?_0'),s(fun(fun(X0,fun(X1,bool)),fun(X0,bool)),'_dst_x00x11_2'(s(fun(fun(X1,bool),bool),'?_0'),s(fun(fun(X0,fun(X1,bool)),fun(X0,fun(X1,bool))),'_dst_x000x111_2'(s(fun(fun(X0,bool),bool),'!_0'),s(fun(fun(X0,fun(X1,bool)),fun(X0,fun(X1,fun(X0,bool)))),'_dst_x0000x1111_2'(s(fun(fun(X1,bool),bool),'!_0'),s(fun(fun(X0,fun(X1,bool)),fun(X0,fun(X1,fun(X0,fun(X1,bool))))),'_dst_x00000x10011x01111_3'(s(fun(bool,fun(bool,bool)),'=_0'),s(fun(fun(X0,fun(X1,bool)),fun(X0,fun(X1,bool))),'_i_0_0'),s(fun(X0,fun(X1,fun(X0,fun(X1,bool)))),'_dst_x0000x1010x0101_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(X0,fun(X0,bool)),'_dst_x01x10_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(fun(X0,X0),'_i_0_0'))),s(fun(X1,fun(X1,bool)),'_dst_x01x10_2'(s(fun(X1,fun(X1,bool)),'=_0'),s(fun(X1,X1),'_i_0_0'))))))))))))))),s(fun('type.pair.prod'(X0,X1),fun(X0,fun(X1,bool))),sK91(X1,X0)))))) )),
  inference(cnf_transformation,[],[f1057])).
fof(f1057,plain,(
  ! [X0,X1] : p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun(fun(X0,fun(X1,bool)),bool),'_dst_x0x1_2'(s(fun(fun(X0,bool),bool),'?_0'),s(fun(fun(X0,fun(X1,bool)),fun(X0,bool)),'_dst_x00x11_2'(s(fun(fun(X1,bool),bool),'?_0'),s(fun(fun(X0,fun(X1,bool)),fun(X0,fun(X1,bool))),'_dst_x000x111_2'(s(fun(fun(X0,bool),bool),'!_0'),s(fun(fun(X0,fun(X1,bool)),fun(X0,fun(X1,fun(X0,bool)))),'_dst_x0000x1111_2'(s(fun(fun(X1,bool),bool),'!_0'),s(fun(fun(X0,fun(X1,bool)),fun(X0,fun(X1,fun(X0,fun(X1,bool))))),'_dst_x00000x10011x01111_3'(s(fun(bool,fun(bool,bool)),'=_0'),s(fun(fun(X0,fun(X1,bool)),fun(X0,fun(X1,bool))),'_i_0_0'),s(fun(X0,fun(X1,fun(X0,fun(X1,bool)))),'_dst_x0000x1010x0101_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(X0,fun(X0,bool)),'_dst_x01x10_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(fun(X0,X0),'_i_0_0'))),s(fun(X1,fun(X1,bool)),'_dst_x01x10_2'(s(fun(X1,fun(X1,bool)),'=_0'),s(fun(X1,X1),'_i_0_0'))))))))))))))),s(fun('type.pair.prod'(X0,X1),fun(X0,fun(X1,bool))),sK91(X1,X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK91])],[f625])).
fof(f625,plain,(
  ! [X0,X1] : ? [X2] : p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun(fun(X0,fun(X1,bool)),bool),'_dst_x0x1_2'(s(fun(fun(X0,bool),bool),'?_0'),s(fun(fun(X0,fun(X1,bool)),fun(X0,bool)),'_dst_x00x11_2'(s(fun(fun(X1,bool),bool),'?_0'),s(fun(fun(X0,fun(X1,bool)),fun(X0,fun(X1,bool))),'_dst_x000x111_2'(s(fun(fun(X0,bool),bool),'!_0'),s(fun(fun(X0,fun(X1,bool)),fun(X0,fun(X1,fun(X0,bool)))),'_dst_x0000x1111_2'(s(fun(fun(X1,bool),bool),'!_0'),s(fun(fun(X0,fun(X1,bool)),fun(X0,fun(X1,fun(X0,fun(X1,bool))))),'_dst_x00000x10011x01111_3'(s(fun(bool,fun(bool,bool)),'=_0'),s(fun(fun(X0,fun(X1,bool)),fun(X0,fun(X1,bool))),'_i_0_0'),s(fun(X0,fun(X1,fun(X0,fun(X1,bool)))),'_dst_x0000x1010x0101_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(X0,fun(X0,bool)),'_dst_x01x10_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(fun(X0,X0),'_i_0_0'))),s(fun(X1,fun(X1,bool)),'_dst_x01x10_2'(s(fun(X1,fun(X1,bool)),'=_0'),s(fun(X1,X1),'_i_0_0'))))))))))))))),s(fun('type.pair.prod'(X0,X1),fun(X0,fun(X1,bool))),X2))))),
  inference(rectify,[],[f116])).
fof(f116,axiom,(
  ! [X6,X20] : ? [X53] : p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun(fun(X6,fun(X20,bool)),bool),'_dst_x0x1_2'(s(fun(fun(X6,bool),bool),'?_0'),s(fun(fun(X6,fun(X20,bool)),fun(X6,bool)),'_dst_x00x11_2'(s(fun(fun(X20,bool),bool),'?_0'),s(fun(fun(X6,fun(X20,bool)),fun(X6,fun(X20,bool))),'_dst_x000x111_2'(s(fun(fun(X6,bool),bool),'!_0'),s(fun(fun(X6,fun(X20,bool)),fun(X6,fun(X20,fun(X6,bool)))),'_dst_x0000x1111_2'(s(fun(fun(X20,bool),bool),'!_0'),s(fun(fun(X6,fun(X20,bool)),fun(X6,fun(X20,fun(X6,fun(X20,bool))))),'_dst_x00000x10011x01111_3'(s(fun(bool,fun(bool,bool)),'=_0'),s(fun(fun(X6,fun(X20,bool)),fun(X6,fun(X20,bool))),'_i_0_0'),s(fun(X6,fun(X20,fun(X6,fun(X20,bool)))),'_dst_x0000x1010x0101_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(X6,fun(X6,bool)),'_dst_x01x10_2'(s(fun(X6,fun(X6,bool)),'=_0'),s(fun(X6,X6),'_i_0_0'))),s(fun(X20,fun(X20,bool)),'_dst_x01x10_2'(s(fun(X20,fun(X20,bool)),'=_0'),s(fun(X20,X20),'_i_0_0'))))))))))))))),s(fun('type.pair.prod'(X6,X20),fun(X6,fun(X20,bool))),X53))))),
  file('thm.pred_set.EXTENSION','thm.pair.prod_TY_DEF')).
fof(f1518,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun(fun(bool,fun(X0,fun(X1,bool))),bool),'const.sum.IS_SUM_REP_0'),s(fun('type.sum.sum'(X0,X1),fun(bool,fun(X0,fun(X1,bool)))),sK90(X1,X0)))))) )),
  inference(cnf_transformation,[],[f1056])).
fof(f1056,plain,(
  ! [X0,X1] : p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun(fun(bool,fun(X0,fun(X1,bool))),bool),'const.sum.IS_SUM_REP_0'),s(fun('type.sum.sum'(X0,X1),fun(bool,fun(X0,fun(X1,bool)))),sK90(X1,X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK90])],[f624])).
fof(f624,plain,(
  ! [X0,X1] : ? [X2] : p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun(fun(bool,fun(X0,fun(X1,bool))),bool),'const.sum.IS_SUM_REP_0'),s(fun('type.sum.sum'(X0,X1),fun(bool,fun(X0,fun(X1,bool)))),X2))))),
  inference(rectify,[],[f225])).
fof(f225,axiom,(
  ! [X6,X20] : ? [X53] : p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun(fun(bool,fun(X6,fun(X20,bool))),bool),'const.sum.IS_SUM_REP_0'),s(fun('type.sum.sum'(X6,X20),fun(bool,fun(X6,fun(X20,bool)))),X53))))),
  file('thm.pred_set.EXTENSION','thm.sum.sum_TY_DEF')).
fof(f1899,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X3)))) | p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))),s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X5),s(X1,X7)))))) | ~p(s(bool,'const.relation.INDUCTIVE_INVARIANT_ON_4'(s(fun(X1,fun(X1,bool)),X3),s(fun(X1,bool),X4),s(fun(X1,fun(X0,bool)),X5),s(fun(fun(X1,X0),fun(X1,X0)),X6)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X7),s(fun(X1,bool),X4)))) | s(fun(X1,X0),X2) != s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X6)))) )),
  inference(forward_demodulation,[],[f1898,f1257])).
fof(f1898,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))),s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X5),s(X1,X7)))))) | ~p(s(bool,happ(s(fun(X1,bool),X4),s(X1,X7)))) | ~p(s(bool,'const.relation.INDUCTIVE_INVARIANT_ON_4'(s(fun(X1,fun(X1,bool)),X3),s(fun(X1,bool),X4),s(fun(X1,fun(X0,bool)),X5),s(fun(fun(X1,X0),fun(X1,X0)),X6)))) | ~p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X3)))) | s(fun(X1,X0),X2) != s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X6)))) )),
  inference(forward_demodulation,[],[f1517,f1257])).
fof(f1517,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X5),s(X1,X7))),s(X0,happ(s(fun(X1,X0),X2),s(X1,X7)))))) | ~p(s(bool,happ(s(fun(X1,bool),X4),s(X1,X7)))) | ~p(s(bool,'const.relation.INDUCTIVE_INVARIANT_ON_4'(s(fun(X1,fun(X1,bool)),X3),s(fun(X1,bool),X4),s(fun(X1,fun(X0,bool)),X5),s(fun(fun(X1,X0),fun(X1,X0)),X6)))) | ~p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X3)))) | s(fun(X1,X0),X2) != s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X6)))) )),
  inference(cnf_transformation,[],[f815])).
fof(f815,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (s(fun(X1,X0),X2) != s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X6))) | ~p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X3)))) | ~p(s(bool,'const.relation.INDUCTIVE_INVARIANT_ON_4'(s(fun(X1,fun(X1,bool)),X3),s(fun(X1,bool),X4),s(fun(X1,fun(X0,bool)),X5),s(fun(fun(X1,X0),fun(X1,X0)),X6)))) | ~p(s(bool,happ(s(fun(X1,bool),X4),s(X1,X7)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X5),s(X1,X7))),s(X0,happ(s(fun(X1,X0),X2),s(X1,X7)))))))),
  inference(flattening,[],[f814])).
fof(f814,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X6))) | ~p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X3)))) | ~p(s(bool,'const.relation.INDUCTIVE_INVARIANT_ON_4'(s(fun(X1,fun(X1,bool)),X3),s(fun(X1,bool),X4),s(fun(X1,fun(X0,bool)),X5),s(fun(fun(X1,X0),fun(X1,X0)),X6)))) | ~p(s(bool,happ(s(fun(X1,bool),X4),s(X1,X7))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X5),s(X1,X7))),s(X0,happ(s(fun(X1,X0),X2),s(X1,X7)))))))),
  inference(ennf_transformation,[],[f623])).
fof(f623,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((s(fun(X1,X0),X2) = s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X6))) & p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X3)))) & p(s(bool,'const.relation.INDUCTIVE_INVARIANT_ON_4'(s(fun(X1,fun(X1,bool)),X3),s(fun(X1,bool),X4),s(fun(X1,fun(X0,bool)),X5),s(fun(fun(X1,X0),fun(X1,X0)),X6)))) & p(s(bool,happ(s(fun(X1,bool),X4),s(X1,X7))))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X5),s(X1,X7))),s(X0,happ(s(fun(X1,X0),X2),s(X1,X7)))))))),
  inference(flattening,[],[f622])).
fof(f622,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6,X7] : ((s(fun(X1,X0),X2) = s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X6))) & p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X3)))) & p(s(bool,'const.relation.INDUCTIVE_INVARIANT_ON_4'(s(fun(X1,fun(X1,bool)),X3),s(fun(X1,bool),X4),s(fun(X1,fun(X0,bool)),X5),s(fun(fun(X1,X0),fun(X1,X0)),X6)))) & p(s(bool,happ(s(fun(X1,bool),X4),s(X1,X7))))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X5),s(X1,X7))),s(X0,happ(s(fun(X1,X0),X2),s(X1,X7)))))))),
  inference(rectify,[],[f137])).
fof(f137,axiom,(
  ! [X20,X6] : ! [X15,X18,X62,X17,X61,X9] : ((s(fun(X6,X20),X15) = s(fun(X6,X20),'const.relation.WFREC_2'(s(fun(X6,fun(X6,bool)),X18),s(fun(fun(X6,X20),fun(X6,X20)),X61))) & p(s(bool,'const.relation.WF_1'(s(fun(X6,fun(X6,bool)),X18)))) & p(s(bool,'const.relation.INDUCTIVE_INVARIANT_ON_4'(s(fun(X6,fun(X6,bool)),X18),s(fun(X6,bool),X62),s(fun(X6,fun(X20,bool)),X17),s(fun(fun(X6,X20),fun(X6,X20)),X61)))) & p(s(bool,happ(s(fun(X6,bool),X62),s(X6,X9))))) => p(s(bool,happ(s(fun(X20,bool),happ(s(fun(X6,fun(X20,bool)),X17),s(X6,X9))),s(X20,happ(s(fun(X6,X20),X15),s(X6,X9)))))))),
  file('thm.pred_set.EXTENSION','thm.relation.TFL_INDUCTIVE_INVARIANT_ON_WFREC')).
fof(f1897,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X8)))))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X7),s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X8)))))) | s(X0,happ(s(fun(X1,X0),X4),s(X1,X8))) = s(X0,happ(s(fun(X1,X0),X5),s(X1,X8))) | ~p(s(bool,'const.relation.approx_4'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X3),s(X1,X7),s(fun(X1,X0),X5)))) | ~p(s(bool,'const.relation.approx_4'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X3),s(X1,X6),s(fun(X1,X0),X4)))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(X1,fun(X1,bool)),X2)))) | ~p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X2))))) )),
  inference(forward_demodulation,[],[f1896,f1257])).
fof(f1896,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X1,X7),s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X8)))))) | s(X0,happ(s(fun(X1,X0),X4),s(X1,X8))) = s(X0,happ(s(fun(X1,X0),X5),s(X1,X8))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X8))),s(X1,X6)))) | ~p(s(bool,'const.relation.approx_4'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X3),s(X1,X7),s(fun(X1,X0),X5)))) | ~p(s(bool,'const.relation.approx_4'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X3),s(X1,X6),s(fun(X1,X0),X4)))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(X1,fun(X1,bool)),X2)))) | ~p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X2))))) )),
  inference(forward_demodulation,[],[f1516,f1257])).
fof(f1516,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),X4),s(X1,X8))) = s(X0,happ(s(fun(X1,X0),X5),s(X1,X8))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X8))),s(X1,X7)))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X8))),s(X1,X6)))) | ~p(s(bool,'const.relation.approx_4'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X3),s(X1,X7),s(fun(X1,X0),X5)))) | ~p(s(bool,'const.relation.approx_4'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X3),s(X1,X6),s(fun(X1,X0),X4)))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(X1,fun(X1,bool)),X2)))) | ~p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X2))))) )),
  inference(cnf_transformation,[],[f813])).
fof(f813,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (~p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X2)))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(X1,fun(X1,bool)),X2)))) | ~p(s(bool,'const.relation.approx_4'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X3),s(X1,X6),s(fun(X1,X0),X4)))) | ~p(s(bool,'const.relation.approx_4'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X3),s(X1,X7),s(fun(X1,X0),X5)))) | ! [X8] : (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X8))),s(X1,X6)))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X8))),s(X1,X7)))) | s(X0,happ(s(fun(X1,X0),X4),s(X1,X8))) = s(X0,happ(s(fun(X1,X0),X5),s(X1,X8)))))),
  inference(flattening,[],[f812])).
fof(f812,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((~p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X2)))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(X1,fun(X1,bool)),X2)))) | ~p(s(bool,'const.relation.approx_4'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X3),s(X1,X6),s(fun(X1,X0),X4)))) | ~p(s(bool,'const.relation.approx_4'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X3),s(X1,X7),s(fun(X1,X0),X5))))) | ! [X8] : (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X8))),s(X1,X6)))) | (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X8))),s(X1,X7)))) | s(X0,happ(s(fun(X1,X0),X4),s(X1,X8))) = s(X0,happ(s(fun(X1,X0),X5),s(X1,X8))))))),
  inference(ennf_transformation,[],[f621])).
fof(f621,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X2)))) & p(s(bool,'const.relation.transitive_1'(s(fun(X1,fun(X1,bool)),X2)))) & p(s(bool,'const.relation.approx_4'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X3),s(X1,X6),s(fun(X1,X0),X4)))) & p(s(bool,'const.relation.approx_4'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X3),s(X1,X7),s(fun(X1,X0),X5))))) => ! [X8] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X8))),s(X1,X6)))) => (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X8))),s(X1,X7)))) => s(X0,happ(s(fun(X1,X0),X4),s(X1,X8))) = s(X0,happ(s(fun(X1,X0),X5),s(X1,X8))))))),
  inference(flattening,[],[f620])).
fof(f620,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6,X7] : ((p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X2)))) & p(s(bool,'const.relation.transitive_1'(s(fun(X1,fun(X1,bool)),X2)))) & p(s(bool,'const.relation.approx_4'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X3),s(X1,X6),s(fun(X1,X0),X4)))) & p(s(bool,'const.relation.approx_4'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X3),s(X1,X7),s(fun(X1,X0),X5))))) => ! [X8] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X8))),s(X1,X6)))) => (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X8))),s(X1,X7)))) => s(X0,happ(s(fun(X1,X0),X4),s(X1,X8))) = s(X0,happ(s(fun(X1,X0),X5),s(X1,X8))))))),
  inference(rectify,[],[f138])).
fof(f138,axiom,(
  ! [X20,X6] : ! [X18,X61,X15,X14,X51,X52] : ((p(s(bool,'const.relation.WF_1'(s(fun(X6,fun(X6,bool)),X18)))) & p(s(bool,'const.relation.transitive_1'(s(fun(X6,fun(X6,bool)),X18)))) & p(s(bool,'const.relation.approx_4'(s(fun(X6,fun(X6,bool)),X18),s(fun(fun(X6,X20),fun(X6,X20)),X61),s(X6,X51),s(fun(X6,X20),X15)))) & p(s(bool,'const.relation.approx_4'(s(fun(X6,fun(X6,bool)),X18),s(fun(fun(X6,X20),fun(X6,X20)),X61),s(X6,X52),s(fun(X6,X20),X14))))) => ! [X9] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X18),s(X6,X9))),s(X6,X51)))) => (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X18),s(X6,X9))),s(X6,X52)))) => s(X20,happ(s(fun(X6,X20),X14),s(X6,X9))) = s(X20,happ(s(fun(X6,X20),X15),s(X6,X9))))))),
  file('thm.pred_set.EXTENSION','thm.relation.relation_|49|34_prove_ltb')).
fof(f1895,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X7)))))) | s(fun(X0,X1),X5) = s(fun(X0,X1),'const.relation.RESTRICT_3'(s(fun(X0,X1),X4),s(fun(X0,fun(X0,bool)),X2),s(X0,X7))) | ~p(s(bool,'const.relation.approx_4'(s(fun(X0,fun(X0,bool)),X2),s(fun(fun(X0,X1),fun(X0,X1)),X3),s(X0,X7),s(fun(X0,X1),X5)))) | ~p(s(bool,'const.relation.approx_4'(s(fun(X0,fun(X0,bool)),X2),s(fun(fun(X0,X1),fun(X0,X1)),X3),s(X0,X6),s(fun(X0,X1),X4)))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X2)))) | ~p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),X2))))) )),
  inference(forward_demodulation,[],[f1515,f1257])).
fof(f1515,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(fun(X0,X1),X5) = s(fun(X0,X1),'const.relation.RESTRICT_3'(s(fun(X0,X1),X4),s(fun(X0,fun(X0,bool)),X2),s(X0,X7))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X7))),s(X0,X6)))) | ~p(s(bool,'const.relation.approx_4'(s(fun(X0,fun(X0,bool)),X2),s(fun(fun(X0,X1),fun(X0,X1)),X3),s(X0,X7),s(fun(X0,X1),X5)))) | ~p(s(bool,'const.relation.approx_4'(s(fun(X0,fun(X0,bool)),X2),s(fun(fun(X0,X1),fun(X0,X1)),X3),s(X0,X6),s(fun(X0,X1),X4)))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X2)))) | ~p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),X2))))) )),
  inference(cnf_transformation,[],[f811])).
fof(f811,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (~p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),X2)))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X2)))) | ~p(s(bool,'const.relation.approx_4'(s(fun(X0,fun(X0,bool)),X2),s(fun(fun(X0,X1),fun(X0,X1)),X3),s(X0,X6),s(fun(X0,X1),X4)))) | ~p(s(bool,'const.relation.approx_4'(s(fun(X0,fun(X0,bool)),X2),s(fun(fun(X0,X1),fun(X0,X1)),X3),s(X0,X7),s(fun(X0,X1),X5)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X7))),s(X0,X6)))) | s(fun(X0,X1),X5) = s(fun(X0,X1),'const.relation.RESTRICT_3'(s(fun(X0,X1),X4),s(fun(X0,fun(X0,bool)),X2),s(X0,X7))))),
  inference(flattening,[],[f810])).
fof(f810,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((~p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),X2)))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X2)))) | ~p(s(bool,'const.relation.approx_4'(s(fun(X0,fun(X0,bool)),X2),s(fun(fun(X0,X1),fun(X0,X1)),X3),s(X0,X6),s(fun(X0,X1),X4)))) | ~p(s(bool,'const.relation.approx_4'(s(fun(X0,fun(X0,bool)),X2),s(fun(fun(X0,X1),fun(X0,X1)),X3),s(X0,X7),s(fun(X0,X1),X5)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X7))),s(X0,X6))))) | s(fun(X0,X1),X5) = s(fun(X0,X1),'const.relation.RESTRICT_3'(s(fun(X0,X1),X4),s(fun(X0,fun(X0,bool)),X2),s(X0,X7))))),
  inference(ennf_transformation,[],[f619])).
fof(f619,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),X2)))) & p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X2)))) & p(s(bool,'const.relation.approx_4'(s(fun(X0,fun(X0,bool)),X2),s(fun(fun(X0,X1),fun(X0,X1)),X3),s(X0,X6),s(fun(X0,X1),X4)))) & p(s(bool,'const.relation.approx_4'(s(fun(X0,fun(X0,bool)),X2),s(fun(fun(X0,X1),fun(X0,X1)),X3),s(X0,X7),s(fun(X0,X1),X5)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X7))),s(X0,X6))))) => s(fun(X0,X1),X5) = s(fun(X0,X1),'const.relation.RESTRICT_3'(s(fun(X0,X1),X4),s(fun(X0,fun(X0,bool)),X2),s(X0,X7))))),
  inference(flattening,[],[f618])).
fof(f618,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6,X7] : ((p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),X2)))) & p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X2)))) & p(s(bool,'const.relation.approx_4'(s(fun(X0,fun(X0,bool)),X2),s(fun(fun(X0,X1),fun(X0,X1)),X3),s(X0,X6),s(fun(X0,X1),X4)))) & p(s(bool,'const.relation.approx_4'(s(fun(X0,fun(X0,bool)),X2),s(fun(fun(X0,X1),fun(X0,X1)),X3),s(X0,X7),s(fun(X0,X1),X5)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X7))),s(X0,X6))))) => s(fun(X0,X1),X5) = s(fun(X0,X1),'const.relation.RESTRICT_3'(s(fun(X0,X1),X4),s(fun(X0,fun(X0,bool)),X2),s(X0,X7))))),
  inference(rectify,[],[f126])).
fof(f126,axiom,(
  ! [X6,X20] : ! [X18,X61,X15,X14,X51,X52] : ((p(s(bool,'const.relation.WF_1'(s(fun(X6,fun(X6,bool)),X18)))) & p(s(bool,'const.relation.transitive_1'(s(fun(X6,fun(X6,bool)),X18)))) & p(s(bool,'const.relation.approx_4'(s(fun(X6,fun(X6,bool)),X18),s(fun(fun(X6,X20),fun(X6,X20)),X61),s(X6,X51),s(fun(X6,X20),X15)))) & p(s(bool,'const.relation.approx_4'(s(fun(X6,fun(X6,bool)),X18),s(fun(fun(X6,X20),fun(X6,X20)),X61),s(X6,X52),s(fun(X6,X20),X14)))) & p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X18),s(X6,X52))),s(X6,X51))))) => s(fun(X6,X20),X14) = s(fun(X6,X20),'const.relation.RESTRICT_3'(s(fun(X6,X20),X15),s(fun(X6,fun(X6,bool)),X18),s(X6,X52))))),
  file('thm.pred_set.EXTENSION','thm.relation.relation_|49|35_prove_ltb')).
fof(f1894,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X3)))) | ~p(s(bool,'const.relation.INDUCTIVE_INVARIANT_3'(s(fun(X1,fun(X1,bool)),X3),s(fun(X1,fun(X0,bool)),X4),s(fun(fun(X1,X0),fun(X1,X0)),X5)))) | p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))),s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X4),s(X1,X6)))))) | s(fun(X1,X0),X2) != s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X5)))) )),
  inference(forward_demodulation,[],[f1514,f1257])).
fof(f1514,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X4),s(X1,X6))),s(X0,happ(s(fun(X1,X0),X2),s(X1,X6)))))) | ~p(s(bool,'const.relation.INDUCTIVE_INVARIANT_3'(s(fun(X1,fun(X1,bool)),X3),s(fun(X1,fun(X0,bool)),X4),s(fun(fun(X1,X0),fun(X1,X0)),X5)))) | ~p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X3)))) | s(fun(X1,X0),X2) != s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X5)))) )),
  inference(cnf_transformation,[],[f809])).
fof(f809,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : (s(fun(X1,X0),X2) != s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X5))) | ~p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X3)))) | ~p(s(bool,'const.relation.INDUCTIVE_INVARIANT_3'(s(fun(X1,fun(X1,bool)),X3),s(fun(X1,fun(X0,bool)),X4),s(fun(fun(X1,X0),fun(X1,X0)),X5)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X4),s(X1,X6))),s(X0,happ(s(fun(X1,X0),X2),s(X1,X6)))))))),
  inference(flattening,[],[f808])).
fof(f808,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X5))) | ~p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X3)))) | ~p(s(bool,'const.relation.INDUCTIVE_INVARIANT_3'(s(fun(X1,fun(X1,bool)),X3),s(fun(X1,fun(X0,bool)),X4),s(fun(fun(X1,X0),fun(X1,X0)),X5))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X4),s(X1,X6))),s(X0,happ(s(fun(X1,X0),X2),s(X1,X6)))))))),
  inference(ennf_transformation,[],[f617])).
fof(f617,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ((s(fun(X1,X0),X2) = s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X5))) & p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X3)))) & p(s(bool,'const.relation.INDUCTIVE_INVARIANT_3'(s(fun(X1,fun(X1,bool)),X3),s(fun(X1,fun(X0,bool)),X4),s(fun(fun(X1,X0),fun(X1,X0)),X5))))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X4),s(X1,X6))),s(X0,happ(s(fun(X1,X0),X2),s(X1,X6)))))))),
  inference(flattening,[],[f616])).
fof(f616,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6] : ((s(fun(X1,X0),X2) = s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X5))) & p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X3)))) & p(s(bool,'const.relation.INDUCTIVE_INVARIANT_3'(s(fun(X1,fun(X1,bool)),X3),s(fun(X1,fun(X0,bool)),X4),s(fun(fun(X1,X0),fun(X1,X0)),X5))))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X4),s(X1,X6))),s(X0,happ(s(fun(X1,X0),X2),s(X1,X6)))))))),
  inference(rectify,[],[f139])).
fof(f139,axiom,(
  ! [X20,X6] : ! [X15,X18,X17,X61,X9] : ((s(fun(X6,X20),X15) = s(fun(X6,X20),'const.relation.WFREC_2'(s(fun(X6,fun(X6,bool)),X18),s(fun(fun(X6,X20),fun(X6,X20)),X61))) & p(s(bool,'const.relation.WF_1'(s(fun(X6,fun(X6,bool)),X18)))) & p(s(bool,'const.relation.INDUCTIVE_INVARIANT_3'(s(fun(X6,fun(X6,bool)),X18),s(fun(X6,fun(X20,bool)),X17),s(fun(fun(X6,X20),fun(X6,X20)),X61))))) => p(s(bool,happ(s(fun(X20,bool),happ(s(fun(X6,fun(X20,bool)),X17),s(X6,X9))),s(X20,happ(s(fun(X6,X20),X15),s(X6,X9)))))))),
  file('thm.pred_set.EXTENSION','thm.relation.TFL_INDUCTIVE_INVARIANT_WFREC')).
fof(f1893,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,'const.relation.WFREC_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X4),s(X1,X6))),s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X3),s(X1,X6)))))) | ~p(s(bool,'const.relation.INDUCTIVE_INVARIANT_ON_4'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,bool),X5),s(fun(X1,fun(X0,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X4)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X5))))) )),
  inference(forward_demodulation,[],[f1892,f1257])).
fof(f1892,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,'const.relation.WFREC_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X4),s(X1,X6))),s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X3),s(X1,X6)))))) | ~p(s(bool,happ(s(fun(X1,bool),X5),s(X1,X6)))) | ~p(s(bool,'const.relation.INDUCTIVE_INVARIANT_ON_4'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,bool),X5),s(fun(X1,fun(X0,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X4)))) | ~p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X2))))) )),
  inference(forward_demodulation,[],[f1513,f1257])).
fof(f1513,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X3),s(X1,X6))),s(X0,'const.relation.WFREC_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X4),s(X1,X6)))))) | ~p(s(bool,happ(s(fun(X1,bool),X5),s(X1,X6)))) | ~p(s(bool,'const.relation.INDUCTIVE_INVARIANT_ON_4'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,bool),X5),s(fun(X1,fun(X0,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X4)))) | ~p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X2))))) )),
  inference(cnf_transformation,[],[f807])).
fof(f807,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : (~p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X2)))) | ~p(s(bool,'const.relation.INDUCTIVE_INVARIANT_ON_4'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,bool),X5),s(fun(X1,fun(X0,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X4)))) | ~p(s(bool,happ(s(fun(X1,bool),X5),s(X1,X6)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X3),s(X1,X6))),s(X0,'const.relation.WFREC_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X4),s(X1,X6)))))))),
  inference(flattening,[],[f806])).
fof(f806,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ((~p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X2)))) | ~p(s(bool,'const.relation.INDUCTIVE_INVARIANT_ON_4'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,bool),X5),s(fun(X1,fun(X0,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X4)))) | ~p(s(bool,happ(s(fun(X1,bool),X5),s(X1,X6))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X3),s(X1,X6))),s(X0,'const.relation.WFREC_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X4),s(X1,X6)))))))),
  inference(ennf_transformation,[],[f615])).
fof(f615,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ((p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X2)))) & p(s(bool,'const.relation.INDUCTIVE_INVARIANT_ON_4'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,bool),X5),s(fun(X1,fun(X0,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X4)))) & p(s(bool,happ(s(fun(X1,bool),X5),s(X1,X6))))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X3),s(X1,X6))),s(X0,'const.relation.WFREC_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X4),s(X1,X6)))))))),
  inference(flattening,[],[f614])).
fof(f614,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6] : ((p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X2)))) & p(s(bool,'const.relation.INDUCTIVE_INVARIANT_ON_4'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,bool),X5),s(fun(X1,fun(X0,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X4)))) & p(s(bool,happ(s(fun(X1,bool),X5),s(X1,X6))))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X3),s(X1,X6))),s(X0,'const.relation.WFREC_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X4),s(X1,X6)))))))),
  inference(rectify,[],[f151])).
fof(f151,axiom,(
  ! [X20,X6] : ! [X18,X17,X61,X62,X9] : ((p(s(bool,'const.relation.WF_1'(s(fun(X6,fun(X6,bool)),X18)))) & p(s(bool,'const.relation.INDUCTIVE_INVARIANT_ON_4'(s(fun(X6,fun(X6,bool)),X18),s(fun(X6,bool),X62),s(fun(X6,fun(X20,bool)),X17),s(fun(fun(X6,X20),fun(X6,X20)),X61)))) & p(s(bool,happ(s(fun(X6,bool),X62),s(X6,X9))))) => p(s(bool,happ(s(fun(X20,bool),happ(s(fun(X6,fun(X20,bool)),X17),s(X6,X9))),s(X20,'const.relation.WFREC_3'(s(fun(X6,fun(X6,bool)),X18),s(fun(fun(X6,X20),fun(X6,X20)),X61),s(X6,X9)))))))),
  file('thm.pred_set.EXTENSION','thm.relation.INDUCTIVE_INVARIANT_ON_WFREC')).
fof(f1891,plain,(
  ( ! [X4,X2,X0,X7,X5,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X7)))))) | s(X0,happ(s(fun(X1,X0),X3),s(X1,X7))) = s(X0,happ(s(fun(X1,X0),X4),s(X1,X7))) | s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),X3),s(fun(X1,fun(X1,bool)),X2),s(X1,X5))) != s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),X4),s(fun(X1,fun(X1,bool)),X2),s(X1,X5)))) )),
  inference(forward_demodulation,[],[f1510,f1257])).
fof(f1510,plain,(
  ( ! [X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),X3),s(X1,X7))) = s(X0,happ(s(fun(X1,X0),X4),s(X1,X7))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X7))),s(X1,X5)))) | s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),X3),s(fun(X1,fun(X1,bool)),X2),s(X1,X5))) != s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),X4),s(fun(X1,fun(X1,bool)),X2),s(X1,X5)))) )),
  inference(cnf_transformation,[],[f1055])).
fof(f1055,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),X3),s(fun(X1,fun(X1,bool)),X2),s(X1,X5))) != s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),X4),s(fun(X1,fun(X1,bool)),X2),s(X1,X5))) | ! [X7] : (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X7))),s(X1,X5)))) | s(X0,happ(s(fun(X1,X0),X3),s(X1,X7))) = s(X0,happ(s(fun(X1,X0),X4),s(X1,X7))))) & ((p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,sK89(X5,X4,X3,X2,X1,X0)))),s(X1,X5)))) & s(X0,happ(s(fun(X1,X0),X3),s(X1,sK89(X5,X4,X3,X2,X1,X0)))) != s(X0,happ(s(fun(X1,X0),X4),s(X1,sK89(X5,X4,X3,X2,X1,X0))))) | s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),X3),s(fun(X1,fun(X1,bool)),X2),s(X1,X5))) = s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),X4),s(fun(X1,fun(X1,bool)),X2),s(X1,X5)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK89])],[f1054])).
fof(f1054,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),X3),s(fun(X1,fun(X1,bool)),X2),s(X1,X5))) != s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),X4),s(fun(X1,fun(X1,bool)),X2),s(X1,X5))) | ! [X7] : (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X7))),s(X1,X5)))) | s(X0,happ(s(fun(X1,X0),X3),s(X1,X7))) = s(X0,happ(s(fun(X1,X0),X4),s(X1,X7))))) & (? [X6] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X6))),s(X1,X5)))) & s(X0,happ(s(fun(X1,X0),X3),s(X1,X6))) != s(X0,happ(s(fun(X1,X0),X4),s(X1,X6)))) | s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),X3),s(fun(X1,fun(X1,bool)),X2),s(X1,X5))) = s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),X4),s(fun(X1,fun(X1,bool)),X2),s(X1,X5)))))),
  inference(rectify,[],[f1053])).
fof(f1053,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),X3),s(fun(X1,fun(X1,bool)),X2),s(X1,X5))) != s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),X4),s(fun(X1,fun(X1,bool)),X2),s(X1,X5))) | ! [X6] : (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X6))),s(X1,X5)))) | s(X0,happ(s(fun(X1,X0),X3),s(X1,X6))) = s(X0,happ(s(fun(X1,X0),X4),s(X1,X6))))) & (? [X6] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X6))),s(X1,X5)))) & s(X0,happ(s(fun(X1,X0),X3),s(X1,X6))) != s(X0,happ(s(fun(X1,X0),X4),s(X1,X6)))) | s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),X3),s(fun(X1,fun(X1,bool)),X2),s(X1,X5))) = s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),X4),s(fun(X1,fun(X1,bool)),X2),s(X1,X5)))))),
  inference(nnf_transformation,[],[f805])).
fof(f805,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),X3),s(fun(X1,fun(X1,bool)),X2),s(X1,X5))) = s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),X4),s(fun(X1,fun(X1,bool)),X2),s(X1,X5))) <=> ! [X6] : (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X6))),s(X1,X5)))) | s(X0,happ(s(fun(X1,X0),X3),s(X1,X6))) = s(X0,happ(s(fun(X1,X0),X4),s(X1,X6)))))),
  inference(ennf_transformation,[],[f613])).
fof(f613,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),X3),s(fun(X1,fun(X1,bool)),X2),s(X1,X5))) = s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),X4),s(fun(X1,fun(X1,bool)),X2),s(X1,X5))) <=> ! [X6] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X6))),s(X1,X5)))) => s(X0,happ(s(fun(X1,X0),X3),s(X1,X6))) = s(X0,happ(s(fun(X1,X0),X4),s(X1,X6)))))),
  inference(flattening,[],[f612])).
fof(f612,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5] : (s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),X3),s(fun(X1,fun(X1,bool)),X2),s(X1,X5))) = s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),X4),s(fun(X1,fun(X1,bool)),X2),s(X1,X5))) <=> ! [X6] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X6))),s(X1,X5)))) => s(X0,happ(s(fun(X1,X0),X3),s(X1,X6))) = s(X0,happ(s(fun(X1,X0),X4),s(X1,X6)))))),
  inference(rectify,[],[f108])).
fof(f108,axiom,(
  ! [X20,X6] : ! [X18,X15,X14,X9] : (s(fun(X6,X20),'const.relation.RESTRICT_3'(s(fun(X6,X20),X14),s(fun(X6,fun(X6,bool)),X18),s(X6,X9))) = s(fun(X6,X20),'const.relation.RESTRICT_3'(s(fun(X6,X20),X15),s(fun(X6,fun(X6,bool)),X18),s(X6,X9))) <=> ! [X19] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X18),s(X6,X19))),s(X6,X9)))) => s(X20,happ(s(fun(X6,X20),X14),s(X6,X19))) = s(X20,happ(s(fun(X6,X20),X15),s(X6,X19)))))),
  file('thm.pred_set.EXTENSION','thm.relation.relation_|49|28_prove_ltb')).
fof(f1890,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,sK89(X5,X4,X3,X2,X1,X0))))))) | s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),X3),s(fun(X1,fun(X1,bool)),X2),s(X1,X5))) = s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),X4),s(fun(X1,fun(X1,bool)),X2),s(X1,X5)))) )),
  inference(forward_demodulation,[],[f1511,f1257])).
fof(f1511,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),X3),s(fun(X1,fun(X1,bool)),X2),s(X1,X5))) = s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),X4),s(fun(X1,fun(X1,bool)),X2),s(X1,X5))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,sK89(X5,X4,X3,X2,X1,X0)))),s(X1,X5))))) )),
  inference(cnf_transformation,[],[f1055])).
fof(f1512,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),X3),s(X1,sK89(X5,X4,X3,X2,X1,X0)))) != s(X0,happ(s(fun(X1,X0),X4),s(X1,sK89(X5,X4,X3,X2,X1,X0)))) | s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),X3),s(fun(X1,fun(X1,bool)),X2),s(X1,X5))) = s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),X4),s(fun(X1,fun(X1,bool)),X2),s(X1,X5)))) )),
  inference(cnf_transformation,[],[f1055])).
fof(f1889,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.relation.approx_4'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X3),s(X1,X4),s(fun(X1,X0),X5)))) | s(fun(X1,X0),X5) = s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),'_dst_x0x1x1_3'(s(fun(fun(X1,X0),fun(X1,X0)),X3),s(fun(X1,fun(X1,X0)),'const.relation.RESTRICT_2'(s(fun(X1,X0),X5),s(fun(X1,fun(X1,bool)),X2))),s(fun(X1,X1),'_i_0_0'))),s(fun(X1,fun(X1,bool)),X2),s(X1,X4)))) )),
  inference(forward_demodulation,[],[f1508,f1476])).
fof(f1508,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(fun(X1,X0),X5) = s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),'const.relation.RESTRICT_2'(s(fun(X1,X0),'_dst_x0x1x1_3'(s(fun(fun(X1,X0),fun(X1,X0)),X3),s(fun(X1,fun(X1,X0)),'const.relation.RESTRICT_2'(s(fun(X1,X0),X5),s(fun(X1,fun(X1,bool)),X2))),s(fun(X1,X1),'_i_0_0'))),s(fun(X1,fun(X1,bool)),X2))),s(X1,X4))) | ~p(s(bool,'const.relation.approx_4'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X3),s(X1,X4),s(fun(X1,X0),X5))))) )),
  inference(cnf_transformation,[],[f1052])).
fof(f1052,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((~p(s(bool,'const.relation.approx_4'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X3),s(X1,X4),s(fun(X1,X0),X5)))) | s(fun(X1,X0),X5) = s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),'const.relation.RESTRICT_2'(s(fun(X1,X0),'_dst_x0x1x1_3'(s(fun(fun(X1,X0),fun(X1,X0)),X3),s(fun(X1,fun(X1,X0)),'const.relation.RESTRICT_2'(s(fun(X1,X0),X5),s(fun(X1,fun(X1,bool)),X2))),s(fun(X1,X1),'_i_0_0'))),s(fun(X1,fun(X1,bool)),X2))),s(X1,X4)))) & (s(fun(X1,X0),X5) != s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),'const.relation.RESTRICT_2'(s(fun(X1,X0),'_dst_x0x1x1_3'(s(fun(fun(X1,X0),fun(X1,X0)),X3),s(fun(X1,fun(X1,X0)),'const.relation.RESTRICT_2'(s(fun(X1,X0),X5),s(fun(X1,fun(X1,bool)),X2))),s(fun(X1,X1),'_i_0_0'))),s(fun(X1,fun(X1,bool)),X2))),s(X1,X4))) | p(s(bool,'const.relation.approx_4'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X3),s(X1,X4),s(fun(X1,X0),X5))))))),
  inference(nnf_transformation,[],[f611])).
fof(f611,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (p(s(bool,'const.relation.approx_4'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X3),s(X1,X4),s(fun(X1,X0),X5)))) <=> s(fun(X1,X0),X5) = s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),'const.relation.RESTRICT_2'(s(fun(X1,X0),'_dst_x0x1x1_3'(s(fun(fun(X1,X0),fun(X1,X0)),X3),s(fun(X1,fun(X1,X0)),'const.relation.RESTRICT_2'(s(fun(X1,X0),X5),s(fun(X1,fun(X1,bool)),X2))),s(fun(X1,X1),'_i_0_0'))),s(fun(X1,fun(X1,bool)),X2))),s(X1,X4))))),
  inference(flattening,[],[f610])).
fof(f610,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5] : (p(s(bool,'const.relation.approx_4'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X3),s(X1,X4),s(fun(X1,X0),X5)))) <=> s(fun(X1,X0),X5) = s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),'const.relation.RESTRICT_2'(s(fun(X1,X0),'_dst_x0x1x1_3'(s(fun(fun(X1,X0),fun(X1,X0)),X3),s(fun(X1,fun(X1,X0)),'const.relation.RESTRICT_2'(s(fun(X1,X0),X5),s(fun(X1,fun(X1,bool)),X2))),s(fun(X1,X1),'_i_0_0'))),s(fun(X1,fun(X1,bool)),X2))),s(X1,X4))))),
  inference(rectify,[],[f125])).
fof(f125,axiom,(
  ! [X20,X6] : ! [X18,X61,X9,X15] : (p(s(bool,'const.relation.approx_4'(s(fun(X6,fun(X6,bool)),X18),s(fun(fun(X6,X20),fun(X6,X20)),X61),s(X6,X9),s(fun(X6,X20),X15)))) <=> s(fun(X6,X20),X15) = s(fun(X6,X20),happ(s(fun(X6,fun(X6,X20)),'const.relation.RESTRICT_2'(s(fun(X6,X20),'_dst_x0x1x1_3'(s(fun(fun(X6,X20),fun(X6,X20)),X61),s(fun(X6,fun(X6,X20)),'const.relation.RESTRICT_2'(s(fun(X6,X20),X15),s(fun(X6,fun(X6,bool)),X18))),s(fun(X6,X6),'_i_0_0'))),s(fun(X6,fun(X6,bool)),X18))),s(X6,X9))))),
  file('thm.pred_set.EXTENSION','thm.relation.approx_def')).
fof(f1888,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(fun(X1,X0),X5) != s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),'_dst_x0x1x1_3'(s(fun(fun(X1,X0),fun(X1,X0)),X3),s(fun(X1,fun(X1,X0)),'const.relation.RESTRICT_2'(s(fun(X1,X0),X5),s(fun(X1,fun(X1,bool)),X2))),s(fun(X1,X1),'_i_0_0'))),s(fun(X1,fun(X1,bool)),X2),s(X1,X4))) | p(s(bool,'const.relation.approx_4'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X3),s(X1,X4),s(fun(X1,X0),X5))))) )),
  inference(forward_demodulation,[],[f1509,f1476])).
fof(f1509,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.relation.approx_4'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X3),s(X1,X4),s(fun(X1,X0),X5)))) | s(fun(X1,X0),X5) != s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),'const.relation.RESTRICT_2'(s(fun(X1,X0),'_dst_x0x1x1_3'(s(fun(fun(X1,X0),fun(X1,X0)),X3),s(fun(X1,fun(X1,X0)),'const.relation.RESTRICT_2'(s(fun(X1,X0),X5),s(fun(X1,fun(X1,bool)),X2))),s(fun(X1,X1),'_i_0_0'))),s(fun(X1,fun(X1,bool)),X2))),s(X1,X4)))) )),
  inference(cnf_transformation,[],[f1052])).
fof(f1887,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.relation.RUNION_4'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3),s(X0,X4),s(X1,X5)))) | p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X4)))))) | p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))))))) )),
  inference(forward_demodulation,[],[f1886,f1257])).
fof(f1886,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X4)))))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))) | ~p(s(bool,'const.relation.RUNION_4'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3),s(X0,X4),s(X1,X5))))) )),
  inference(forward_demodulation,[],[f1505,f1257])).
fof(f1505,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X4))),s(X1,X5)))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))) | ~p(s(bool,'const.relation.RUNION_4'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3),s(X0,X4),s(X1,X5))))) )),
  inference(cnf_transformation,[],[f1051])).
fof(f1051,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((~p(s(bool,'const.relation.RUNION_4'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3),s(X0,X4),s(X1,X5)))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X4))),s(X1,X5))))) & ((~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))) & ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X4))),s(X1,X5))))) | p(s(bool,'const.relation.RUNION_4'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3),s(X0,X4),s(X1,X5))))))),
  inference(flattening,[],[f1050])).
fof(f1050,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((~p(s(bool,'const.relation.RUNION_4'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3),s(X0,X4),s(X1,X5)))) | (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X4))),s(X1,X5)))))) & ((~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))) & ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X4))),s(X1,X5))))) | p(s(bool,'const.relation.RUNION_4'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3),s(X0,X4),s(X1,X5))))))),
  inference(nnf_transformation,[],[f609])).
fof(f609,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (p(s(bool,'const.relation.RUNION_4'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3),s(X0,X4),s(X1,X5)))) <=> (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X4))),s(X1,X5))))))),
  inference(flattening,[],[f608])).
fof(f608,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5] : (p(s(bool,'const.relation.RUNION_4'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3),s(X0,X4),s(X1,X5)))) <=> (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X4))),s(X1,X5))))))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X6,X20] : ! [X30,X29,X9,X19] : (p(s(bool,'const.relation.RUNION_4'(s(fun(X6,fun(X20,bool)),X30),s(fun(X6,fun(X20,bool)),X29),s(X6,X9),s(X20,X19)))) <=> (p(s(bool,happ(s(fun(X20,bool),happ(s(fun(X6,fun(X20,bool)),X30),s(X6,X9))),s(X20,X19)))) | p(s(bool,happ(s(fun(X20,bool),happ(s(fun(X6,fun(X20,bool)),X29),s(X6,X9))),s(X20,X19))))))),
  file('thm.pred_set.EXTENSION','thm.relation.RUNION')).
fof(f1885,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4)))))) | p(s(bool,'const.relation.RUNION_4'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3),s(X0,X4),s(X1,X5))))) )),
  inference(forward_demodulation,[],[f1506,f1257])).
fof(f1506,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.relation.RUNION_4'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3),s(X0,X4),s(X1,X5)))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5))))) )),
  inference(cnf_transformation,[],[f1051])).
fof(f1884,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X4)))))) | p(s(bool,'const.relation.RUNION_4'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3),s(X0,X4),s(X1,X5))))) )),
  inference(forward_demodulation,[],[f1507,f1257])).
fof(f1507,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.relation.RUNION_4'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3),s(X0,X4),s(X1,X5)))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X4))),s(X1,X5))))) )),
  inference(cnf_transformation,[],[f1051])).
fof(f1883,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X1,fun(X0,bool)),X2),s(fun(X1,bool),X3),s(X1,X4),s(X0,X5)))) | p(s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X4))))))) )),
  inference(forward_demodulation,[],[f1502,f1257])).
fof(f1502,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X4))),s(X0,X5)))) | ~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X1,fun(X0,bool)),X2),s(fun(X1,bool),X3),s(X1,X4),s(X0,X5))))) )),
  inference(cnf_transformation,[],[f1049])).
fof(f1049,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X1,fun(X0,bool)),X2),s(fun(X1,bool),X3),s(X1,X4),s(X0,X5)))) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X4))),s(X0,X5)))) & p(s(bool,'const.bool.IN_2'(s(X1,X4),s(fun(X1,bool),X3)))))) & (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X4))),s(X0,X5)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X4),s(fun(X1,bool),X3)))) | p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X1,fun(X0,bool)),X2),s(fun(X1,bool),X3),s(X1,X4),s(X0,X5))))))),
  inference(flattening,[],[f1048])).
fof(f1048,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X1,fun(X0,bool)),X2),s(fun(X1,bool),X3),s(X1,X4),s(X0,X5)))) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X4))),s(X0,X5)))) & p(s(bool,'const.bool.IN_2'(s(X1,X4),s(fun(X1,bool),X3)))))) & ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X4))),s(X0,X5)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X4),s(fun(X1,bool),X3))))) | p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X1,fun(X0,bool)),X2),s(fun(X1,bool),X3),s(X1,X4),s(X0,X5))))))),
  inference(nnf_transformation,[],[f607])).
fof(f607,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X1,fun(X0,bool)),X2),s(fun(X1,bool),X3),s(X1,X4),s(X0,X5)))) <=> (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X4))),s(X0,X5)))) & p(s(bool,'const.bool.IN_2'(s(X1,X4),s(fun(X1,bool),X3))))))),
  inference(flattening,[],[f606])).
fof(f606,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5] : (p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X1,fun(X0,bool)),X2),s(fun(X1,bool),X3),s(X1,X4),s(X0,X5)))) <=> (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X4))),s(X0,X5)))) & p(s(bool,'const.bool.IN_2'(s(X1,X4),s(fun(X1,bool),X3))))))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X20,X6] : ! [X18,X7,X9,X19] : (p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X6,fun(X20,bool)),X18),s(fun(X6,bool),X7),s(X6,X9),s(X20,X19)))) <=> (p(s(bool,happ(s(fun(X20,bool),happ(s(fun(X6,fun(X20,bool)),X18),s(X6,X9))),s(X20,X19)))) & p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X7))))))),
  file('thm.pred_set.EXTENSION','thm.relation.RRESTRICT_DEF')).
fof(f1503,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X1,fun(X0,bool)),X2),s(fun(X1,bool),X3),s(X1,X4),s(X0,X5)))) | p(s(bool,'const.bool.IN_2'(s(X1,X4),s(fun(X1,bool),X3))))) )),
  inference(cnf_transformation,[],[f1049])).
fof(f1882,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X1,X4),s(fun(X1,bool),X3)))) | p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X1,fun(X0,bool)),X2),s(fun(X1,bool),X3),s(X1,X4),s(X0,X5)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X4))))))) )),
  inference(forward_demodulation,[],[f1504,f1257])).
fof(f1504,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X1,fun(X0,bool)),X2),s(fun(X1,bool),X3),s(X1,X4),s(X0,X5)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X4),s(fun(X1,bool),X3)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X4))),s(X0,X5))))) )),
  inference(cnf_transformation,[],[f1049])).
fof(f1881,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X1,fun(X0,bool)),X2),s(X1,X3),s(X1,X4),s(X0,X5)))) | p(s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X4))))))) )),
  inference(forward_demodulation,[],[f1499,f1257])).
fof(f1499,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X4))),s(X0,X5)))) | ~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X1,fun(X0,bool)),X2),s(X1,X3),s(X1,X4),s(X0,X5))))) )),
  inference(cnf_transformation,[],[f1047])).
fof(f1047,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X1,fun(X0,bool)),X2),s(X1,X3),s(X1,X4),s(X0,X5)))) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X4))),s(X0,X5)))) & s(X1,X3) != s(X1,X4))) & (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X4))),s(X0,X5)))) | s(X1,X3) = s(X1,X4) | p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X1,fun(X0,bool)),X2),s(X1,X3),s(X1,X4),s(X0,X5))))))),
  inference(flattening,[],[f1046])).
fof(f1046,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X1,fun(X0,bool)),X2),s(X1,X3),s(X1,X4),s(X0,X5)))) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X4))),s(X0,X5)))) & s(X1,X3) != s(X1,X4))) & ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X4))),s(X0,X5)))) | s(X1,X3) = s(X1,X4)) | p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X1,fun(X0,bool)),X2),s(X1,X3),s(X1,X4),s(X0,X5))))))),
  inference(nnf_transformation,[],[f605])).
fof(f605,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X1,fun(X0,bool)),X2),s(X1,X3),s(X1,X4),s(X0,X5)))) <=> (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X4))),s(X0,X5)))) & s(X1,X3) != s(X1,X4)))),
  inference(flattening,[],[f604])).
fof(f604,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5] : (p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X1,fun(X0,bool)),X2),s(X1,X3),s(X1,X4),s(X0,X5)))) <=> (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X4))),s(X0,X5)))) & ~s(X1,X3) = s(X1,X4)))),
  inference(rectify,[],[f103])).
fof(f103,axiom,(
  ! [X20,X6] : ! [X18,X9,X51,X52] : (p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X6,fun(X20,bool)),X18),s(X6,X9),s(X6,X51),s(X20,X52)))) <=> (p(s(bool,happ(s(fun(X20,bool),happ(s(fun(X6,fun(X20,bool)),X18),s(X6,X51))),s(X20,X52)))) & ~s(X6,X9) = s(X6,X51)))),
  file('thm.pred_set.EXTENSION','thm.relation.RDOM_DELETE_DEF')).
fof(f1500,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X1,fun(X0,bool)),X2),s(X1,X3),s(X1,X4),s(X0,X5)))) | s(X1,X3) != s(X1,X4)) )),
  inference(cnf_transformation,[],[f1047])).
fof(f1880,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X4)))))) | p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X1,fun(X0,bool)),X2),s(X1,X3),s(X1,X4),s(X0,X5)))) | s(X1,X3) = s(X1,X4)) )),
  inference(forward_demodulation,[],[f1501,f1257])).
fof(f1501,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X1,fun(X0,bool)),X2),s(X1,X3),s(X1,X4),s(X0,X5)))) | s(X1,X3) = s(X1,X4) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X4))),s(X0,X5))))) )),
  inference(cnf_transformation,[],[f1047])).
fof(f1498,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'const.relation.approx_4'(s(fun(X0,fun(X0,bool)),X2),s(fun(fun(X0,X1),fun(X0,X1)),X3),s(X0,X4),s(fun(X0,X1),X5))) = s(bool,happ(s(fun(fun(X0,X1),bool),'const.relation.approx_3'(s(fun(X0,fun(X0,bool)),X2),s(fun(fun(X0,X1),fun(X0,X1)),X3),s(X0,X4))),s(fun(X0,X1),X5)))) )),
  inference(cnf_transformation,[],[f603])).
fof(f603,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(bool,'const.relation.approx_4'(s(fun(X0,fun(X0,bool)),X2),s(fun(fun(X0,X1),fun(X0,X1)),X3),s(X0,X4),s(fun(X0,X1),X5))) = s(bool,happ(s(fun(fun(X0,X1),bool),'const.relation.approx_3'(s(fun(X0,fun(X0,bool)),X2),s(fun(fun(X0,X1),fun(X0,X1)),X3),s(X0,X4))),s(fun(X0,X1),X5)))),
  inference(flattening,[],[f602])).
fof(f602,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5] : s(bool,'const.relation.approx_4'(s(fun(X0,fun(X0,bool)),X2),s(fun(fun(X0,X1),fun(X0,X1)),X3),s(X0,X4),s(fun(X0,X1),X5))) = s(bool,happ(s(fun(fun(X0,X1),bool),'const.relation.approx_3'(s(fun(X0,fun(X0,bool)),X2),s(fun(fun(X0,X1),fun(X0,X1)),X3),s(X0,X4))),s(fun(X0,X1),X5)))),
  inference(rectify,[],[f115])).
fof(f115,axiom,(
  ! [X6,X20] : ! [X21,X22,X31,X32] : s(bool,'const.relation.approx_4'(s(fun(X6,fun(X6,bool)),X21),s(fun(fun(X6,X20),fun(X6,X20)),X22),s(X6,X31),s(fun(X6,X20),X32))) = s(bool,happ(s(fun(fun(X6,X20),bool),'const.relation.approx_3'(s(fun(X6,fun(X6,bool)),X21),s(fun(fun(X6,X20),fun(X6,X20)),X22),s(X6,X31))),s(fun(X6,X20),X32)))),
  file('thm.pred_set.EXTENSION','ARITY_3_4_const.relation.approx')).
fof(f1496,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'const.relation.RUNION_4'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3),s(X0,X4),s(X1,X5))) = s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),'const.relation.RUNION_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3))),s(X0,X4))),s(X1,X5)))) )),
  inference(cnf_transformation,[],[f599])).
fof(f599,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(bool,'const.relation.RUNION_4'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3),s(X0,X4),s(X1,X5))) = s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),'const.relation.RUNION_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3))),s(X0,X4))),s(X1,X5)))),
  inference(flattening,[],[f598])).
fof(f598,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5] : s(bool,'const.relation.RUNION_4'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3),s(X0,X4),s(X1,X5))) = s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),'const.relation.RUNION_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3))),s(X0,X4))),s(X1,X5)))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X6,X20] : ! [X21,X22,X31,X32] : s(bool,'const.relation.RUNION_4'(s(fun(X6,fun(X20,bool)),X21),s(fun(X6,fun(X20,bool)),X22),s(X6,X31),s(X20,X32))) = s(bool,happ(s(fun(X20,bool),happ(s(fun(X6,fun(X20,bool)),'const.relation.RUNION_2'(s(fun(X6,fun(X20,bool)),X21),s(fun(X6,fun(X20,bool)),X22))),s(X6,X31))),s(X20,X32)))),
  file('thm.pred_set.EXTENSION','ARITY_2_4_const.relation.RUNION')).
fof(f1495,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'const.relation.RRESTRICT_4'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,bool),X3),s(X0,X4),s(X1,X5))) = s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),'const.relation.RRESTRICT_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,bool),X3))),s(X0,X4))),s(X1,X5)))) )),
  inference(cnf_transformation,[],[f597])).
fof(f597,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(bool,'const.relation.RRESTRICT_4'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,bool),X3),s(X0,X4),s(X1,X5))) = s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),'const.relation.RRESTRICT_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,bool),X3))),s(X0,X4))),s(X1,X5)))),
  inference(flattening,[],[f596])).
fof(f596,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5] : s(bool,'const.relation.RRESTRICT_4'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,bool),X3),s(X0,X4),s(X1,X5))) = s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),'const.relation.RRESTRICT_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,bool),X3))),s(X0,X4))),s(X1,X5)))),
  inference(rectify,[],[f104])).
fof(f104,axiom,(
  ! [X6,X20] : ! [X21,X22,X31,X32] : s(bool,'const.relation.RRESTRICT_4'(s(fun(X6,fun(X20,bool)),X21),s(fun(X6,bool),X22),s(X6,X31),s(X20,X32))) = s(bool,happ(s(fun(X20,bool),happ(s(fun(X6,fun(X20,bool)),'const.relation.RRESTRICT_2'(s(fun(X6,fun(X20,bool)),X21),s(fun(X6,bool),X22))),s(X6,X31))),s(X20,X32)))),
  file('thm.pred_set.EXTENSION','ARITY_2_4_const.relation.RRESTRICT')).
fof(f1494,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3),s(X0,X4),s(X1,X5))) = s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))),s(X0,X4))),s(X1,X5)))) )),
  inference(cnf_transformation,[],[f595])).
fof(f595,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3),s(X0,X4),s(X1,X5))) = s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))),s(X0,X4))),s(X1,X5)))),
  inference(flattening,[],[f594])).
fof(f594,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5] : s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3),s(X0,X4),s(X1,X5))) = s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))),s(X0,X4))),s(X1,X5)))),
  inference(rectify,[],[f102])).
fof(f102,axiom,(
  ! [X6,X20] : ! [X21,X22,X31,X32] : s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X6,fun(X20,bool)),X21),s(X6,X22),s(X6,X31),s(X20,X32))) = s(bool,happ(s(fun(X20,bool),happ(s(fun(X6,fun(X20,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X6,fun(X20,bool)),X21),s(X6,X22))),s(X6,X31))),s(X20,X32)))),
  file('thm.pred_set.EXTENSION','ARITY_2_4_const.relation.RDOM_DELETE')).
fof(f1493,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.relation.RESTRICT_4'(s(fun(X1,X0),X2),s(fun(X1,fun(X1,bool)),X3),s(X1,X4),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),X2),s(fun(X1,fun(X1,bool)),X3),s(X1,X4))),s(X1,X5)))) )),
  inference(cnf_transformation,[],[f593])).
fof(f593,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,'const.relation.RESTRICT_4'(s(fun(X1,X0),X2),s(fun(X1,fun(X1,bool)),X3),s(X1,X4),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),X2),s(fun(X1,fun(X1,bool)),X3),s(X1,X4))),s(X1,X5)))),
  inference(flattening,[],[f592])).
fof(f592,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5] : s(X0,'const.relation.RESTRICT_4'(s(fun(X1,X0),X2),s(fun(X1,fun(X1,bool)),X3),s(X1,X4),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),X2),s(fun(X1,fun(X1,bool)),X3),s(X1,X4))),s(X1,X5)))),
  inference(rectify,[],[f130])).
fof(f130,axiom,(
  ! [X20,X6] : ! [X21,X22,X31,X32] : s(X20,'const.relation.RESTRICT_4'(s(fun(X6,X20),X21),s(fun(X6,fun(X6,bool)),X22),s(X6,X31),s(X6,X32))) = s(X20,happ(s(fun(X6,X20),'const.relation.RESTRICT_3'(s(fun(X6,X20),X21),s(fun(X6,fun(X6,bool)),X22),s(X6,X31))),s(X6,X32)))),
  file('thm.pred_set.EXTENSION','ARITY_3_4_const.relation.RESTRICT')).
fof(f1490,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.relation.RDOM_1'(s(fun(X0,fun(X1,bool)),X3)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.relation.RDOM_1'(s(fun(X0,fun(X1,bool)),X4)))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.relation.RDOM_1'(s(fun(X0,fun(X1,bool)),'const.relation.RUNION_2'(s(fun(X0,fun(X1,bool)),X4),s(fun(X0,fun(X1,bool)),X3))))))))) )),
  inference(cnf_transformation,[],[f1045])).
fof(f1045,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.relation.RDOM_1'(s(fun(X0,fun(X1,bool)),'const.relation.RUNION_2'(s(fun(X0,fun(X1,bool)),X4),s(fun(X0,fun(X1,bool)),X3)))))))) | p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.relation.RDOM_1'(s(fun(X0,fun(X1,bool)),X4)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.relation.RDOM_1'(s(fun(X0,fun(X1,bool)),X3))))))) & ((~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.relation.RDOM_1'(s(fun(X0,fun(X1,bool)),X4)))))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.relation.RDOM_1'(s(fun(X0,fun(X1,bool)),X3))))))) | p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.relation.RDOM_1'(s(fun(X0,fun(X1,bool)),'const.relation.RUNION_2'(s(fun(X0,fun(X1,bool)),X4),s(fun(X0,fun(X1,bool)),X3))))))))))),
  inference(flattening,[],[f1044])).
fof(f1044,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.relation.RDOM_1'(s(fun(X0,fun(X1,bool)),'const.relation.RUNION_2'(s(fun(X0,fun(X1,bool)),X4),s(fun(X0,fun(X1,bool)),X3)))))))) | (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.relation.RDOM_1'(s(fun(X0,fun(X1,bool)),X4)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.relation.RDOM_1'(s(fun(X0,fun(X1,bool)),X3)))))))) & ((~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.relation.RDOM_1'(s(fun(X0,fun(X1,bool)),X4)))))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.relation.RDOM_1'(s(fun(X0,fun(X1,bool)),X3))))))) | p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.relation.RDOM_1'(s(fun(X0,fun(X1,bool)),'const.relation.RUNION_2'(s(fun(X0,fun(X1,bool)),X4),s(fun(X0,fun(X1,bool)),X3))))))))))),
  inference(nnf_transformation,[],[f591])).
fof(f591,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.relation.RDOM_1'(s(fun(X0,fun(X1,bool)),'const.relation.RUNION_2'(s(fun(X0,fun(X1,bool)),X4),s(fun(X0,fun(X1,bool)),X3)))))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.relation.RDOM_1'(s(fun(X0,fun(X1,bool)),X4)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.relation.RDOM_1'(s(fun(X0,fun(X1,bool)),X3))))))))),
  inference(flattening,[],[f590])).
fof(f590,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.relation.RDOM_1'(s(fun(X0,fun(X1,bool)),'const.relation.RUNION_2'(s(fun(X0,fun(X1,bool)),X4),s(fun(X0,fun(X1,bool)),X3)))))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.relation.RDOM_1'(s(fun(X0,fun(X1,bool)),X4)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.relation.RDOM_1'(s(fun(X0,fun(X1,bool)),X3))))))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X6,X20] : ! [X9,X29,X30] : (p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),'const.relation.RDOM_1'(s(fun(X6,fun(X20,bool)),'const.relation.RUNION_2'(s(fun(X6,fun(X20,bool)),X30),s(fun(X6,fun(X20,bool)),X29)))))))) <=> (p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),'const.relation.RDOM_1'(s(fun(X6,fun(X20,bool)),X30)))))) | p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),'const.relation.RDOM_1'(s(fun(X6,fun(X20,bool)),X29))))))))),
  file('thm.pred_set.EXTENSION','thm.relation.IN_RDOM_RUNION')).
fof(f1491,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.relation.RDOM_1'(s(fun(X0,fun(X1,bool)),'const.relation.RUNION_2'(s(fun(X0,fun(X1,bool)),X4),s(fun(X0,fun(X1,bool)),X3)))))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.relation.RDOM_1'(s(fun(X0,fun(X1,bool)),X4))))))) )),
  inference(cnf_transformation,[],[f1045])).
fof(f1492,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.relation.RDOM_1'(s(fun(X0,fun(X1,bool)),'const.relation.RUNION_2'(s(fun(X0,fun(X1,bool)),X4),s(fun(X0,fun(X1,bool)),X3)))))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.relation.RDOM_1'(s(fun(X0,fun(X1,bool)),X3))))))) )),
  inference(cnf_transformation,[],[f1045])).
fof(f1487,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),X4)))))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),'const.relation.RRESTRICT_2'(s(fun(X1,fun(X0,bool)),X4),s(fun(X1,bool),X3))))))))) )),
  inference(cnf_transformation,[],[f1043])).
fof(f1043,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),'const.relation.RRESTRICT_2'(s(fun(X1,fun(X0,bool)),X4),s(fun(X1,bool),X3)))))))) | (p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),X4)))))) & p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),X3)))))) & (~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),X4)))))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),'const.relation.RRESTRICT_2'(s(fun(X1,fun(X0,bool)),X4),s(fun(X1,bool),X3))))))))))),
  inference(flattening,[],[f1042])).
fof(f1042,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),'const.relation.RRESTRICT_2'(s(fun(X1,fun(X0,bool)),X4),s(fun(X1,bool),X3)))))))) | (p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),X4)))))) & p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),X3)))))) & ((~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),X4)))))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),X3))))) | p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),'const.relation.RRESTRICT_2'(s(fun(X1,fun(X0,bool)),X4),s(fun(X1,bool),X3))))))))))),
  inference(nnf_transformation,[],[f589])).
fof(f589,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),'const.relation.RRESTRICT_2'(s(fun(X1,fun(X0,bool)),X4),s(fun(X1,bool),X3)))))))) <=> (p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),X4)))))) & p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),X3))))))),
  inference(flattening,[],[f588])).
fof(f588,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),'const.relation.RRESTRICT_2'(s(fun(X1,fun(X0,bool)),X4),s(fun(X1,bool),X3)))))))) <=> (p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),X4)))))) & p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),X3))))))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X20,X6] : ! [X9,X7,X18] : (p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),'const.relation.RDOM_1'(s(fun(X6,fun(X20,bool)),'const.relation.RRESTRICT_2'(s(fun(X6,fun(X20,bool)),X18),s(fun(X6,bool),X7)))))))) <=> (p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),'const.relation.RDOM_1'(s(fun(X6,fun(X20,bool)),X18)))))) & p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X7))))))),
  file('thm.pred_set.EXTENSION','thm.relation.IN_RDOM_RRESTRICT')).
fof(f1488,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),'const.relation.RRESTRICT_2'(s(fun(X1,fun(X0,bool)),X4),s(fun(X1,bool),X3))))))))) )),
  inference(cnf_transformation,[],[f1043])).
fof(f1489,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),'const.relation.RRESTRICT_2'(s(fun(X1,fun(X0,bool)),X4),s(fun(X1,bool),X3)))))))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),X4))))))) )),
  inference(cnf_transformation,[],[f1043])).
fof(f1484,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),X4)))))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X1,fun(X0,bool)),X4),s(X1,X3))))))))) )),
  inference(cnf_transformation,[],[f1041])).
fof(f1041,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X1,fun(X0,bool)),X4),s(X1,X3)))))))) | (p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),X4)))))) & s(X1,X2) != s(X1,X3))) & (~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),X4)))))) | s(X1,X2) = s(X1,X3) | p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X1,fun(X0,bool)),X4),s(X1,X3))))))))))),
  inference(flattening,[],[f1040])).
fof(f1040,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X1,fun(X0,bool)),X4),s(X1,X3)))))))) | (p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),X4)))))) & s(X1,X2) != s(X1,X3))) & ((~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),X4)))))) | s(X1,X2) = s(X1,X3)) | p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X1,fun(X0,bool)),X4),s(X1,X3))))))))))),
  inference(nnf_transformation,[],[f587])).
fof(f587,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X1,fun(X0,bool)),X4),s(X1,X3)))))))) <=> (p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),X4)))))) & s(X1,X2) != s(X1,X3)))),
  inference(flattening,[],[f586])).
fof(f586,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X1,fun(X0,bool)),X4),s(X1,X3)))))))) <=> (p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),X4)))))) & ~s(X1,X2) = s(X1,X3)))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X20,X6] : ! [X9,X28,X18] : (p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),'const.relation.RDOM_1'(s(fun(X6,fun(X20,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X6,fun(X20,bool)),X18),s(X6,X28)))))))) <=> (p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),'const.relation.RDOM_1'(s(fun(X6,fun(X20,bool)),X18)))))) & ~s(X6,X9) = s(X6,X28)))),
  file('thm.pred_set.EXTENSION','thm.relation.IN_RDOM_DELETE')).
fof(f1485,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X1,X2) != s(X1,X3) | ~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X1,fun(X0,bool)),X4),s(X1,X3))))))))) )),
  inference(cnf_transformation,[],[f1041])).
fof(f1486,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X1,fun(X0,bool)),X4),s(X1,X3)))))))) | s(X1,X2) = s(X1,X3) | ~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),X4))))))) )),
  inference(cnf_transformation,[],[f1041])).
fof(f1482,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK88(X4,X3,X2,X1,X0)),s(fun(X0,bool),X2)))) | s(fun(X0,X1),'const.bool.RES_ABSTRACT_2'(s(fun(X0,bool),X2),s(fun(X0,X1),X3))) = s(fun(X0,X1),'const.bool.RES_ABSTRACT_2'(s(fun(X0,bool),X2),s(fun(X0,X1),X4)))) )),
  inference(cnf_transformation,[],[f1039])).
fof(f1039,plain,(
  ! [X0,X1,X2,X3,X4] : ((p(s(bool,'const.bool.IN_2'(s(X0,sK88(X4,X3,X2,X1,X0)),s(fun(X0,bool),X2)))) & s(X1,happ(s(fun(X0,X1),X3),s(X0,sK88(X4,X3,X2,X1,X0)))) != s(X1,happ(s(fun(X0,X1),X4),s(X0,sK88(X4,X3,X2,X1,X0))))) | s(fun(X0,X1),'const.bool.RES_ABSTRACT_2'(s(fun(X0,bool),X2),s(fun(X0,X1),X3))) = s(fun(X0,X1),'const.bool.RES_ABSTRACT_2'(s(fun(X0,bool),X2),s(fun(X0,X1),X4))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK88])],[f804])).
fof(f804,plain,(
  ! [X0,X1,X2,X3,X4] : (? [X5] : (p(s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),X2)))) & s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))) != s(X1,happ(s(fun(X0,X1),X4),s(X0,X5)))) | s(fun(X0,X1),'const.bool.RES_ABSTRACT_2'(s(fun(X0,bool),X2),s(fun(X0,X1),X3))) = s(fun(X0,X1),'const.bool.RES_ABSTRACT_2'(s(fun(X0,bool),X2),s(fun(X0,X1),X4))))),
  inference(ennf_transformation,[],[f585])).
fof(f585,plain,(
  ! [X0,X1,X2,X3,X4] : (! [X5] : (p(s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),X2)))) => s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))) = s(X1,happ(s(fun(X0,X1),X4),s(X0,X5)))) => s(fun(X0,X1),'const.bool.RES_ABSTRACT_2'(s(fun(X0,bool),X2),s(fun(X0,X1),X3))) = s(fun(X0,X1),'const.bool.RES_ABSTRACT_2'(s(fun(X0,bool),X2),s(fun(X0,X1),X4))))),
  inference(flattening,[],[f584])).
fof(f584,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (! [X5] : (p(s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),X2)))) => s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))) = s(X1,happ(s(fun(X0,X1),X4),s(X0,X5)))) => s(fun(X0,X1),'const.bool.RES_ABSTRACT_2'(s(fun(X0,bool),X2),s(fun(X0,X1),X3))) = s(fun(X0,X1),'const.bool.RES_ABSTRACT_2'(s(fun(X0,bool),X2),s(fun(X0,X1),X4))))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X6,X20] : ! [X23,X24,X25] : (! [X9] : (p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X23)))) => s(X20,happ(s(fun(X6,X20),X24),s(X6,X9))) = s(X20,happ(s(fun(X6,X20),X25),s(X6,X9)))) => s(fun(X6,X20),'const.bool.RES_ABSTRACT_2'(s(fun(X6,bool),X23),s(fun(X6,X20),X24))) = s(fun(X6,X20),'const.bool.RES_ABSTRACT_2'(s(fun(X6,bool),X23),s(fun(X6,X20),X25))))),
  file('thm.pred_set.EXTENSION','thm.bool.RES_ABSTRACT_DEF|split|1')).
fof(f1483,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X1,happ(s(fun(X0,X1),X3),s(X0,sK88(X4,X3,X2,X1,X0)))) != s(X1,happ(s(fun(X0,X1),X4),s(X0,sK88(X4,X3,X2,X1,X0)))) | s(fun(X0,X1),'const.bool.RES_ABSTRACT_2'(s(fun(X0,bool),X2),s(fun(X0,X1),X3))) = s(fun(X0,X1),'const.bool.RES_ABSTRACT_2'(s(fun(X0,bool),X2),s(fun(X0,X1),X4)))) )),
  inference(cnf_transformation,[],[f1039])).
fof(f1878,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X2)))) | ~p(s(bool,'const.relation.INDUCTIVE_INVARIANT_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,fun(X0,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X4)))) | p(s(bool,'const.bool.IN_2'(s(X0,'const.relation.WFREC_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X4),s(X1,X5))),s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X3),s(X1,X5))))))) )),
  inference(forward_demodulation,[],[f1481,f1257])).
fof(f1481,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X3),s(X1,X5))),s(X0,'const.relation.WFREC_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X4),s(X1,X5)))))) | ~p(s(bool,'const.relation.INDUCTIVE_INVARIANT_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,fun(X0,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X4)))) | ~p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X2))))) )),
  inference(cnf_transformation,[],[f803])).
fof(f803,plain,(
  ! [X0,X1,X2,X3,X4] : (~p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X2)))) | ~p(s(bool,'const.relation.INDUCTIVE_INVARIANT_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,fun(X0,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X4)))) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X3),s(X1,X5))),s(X0,'const.relation.WFREC_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X4),s(X1,X5)))))))),
  inference(flattening,[],[f802])).
fof(f802,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X2)))) | ~p(s(bool,'const.relation.INDUCTIVE_INVARIANT_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,fun(X0,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X4))))) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X3),s(X1,X5))),s(X0,'const.relation.WFREC_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X4),s(X1,X5)))))))),
  inference(ennf_transformation,[],[f583])).
fof(f583,plain,(
  ! [X0,X1,X2,X3,X4] : ((p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X2)))) & p(s(bool,'const.relation.INDUCTIVE_INVARIANT_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,fun(X0,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X4))))) => ! [X5] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X3),s(X1,X5))),s(X0,'const.relation.WFREC_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X4),s(X1,X5)))))))),
  inference(flattening,[],[f582])).
fof(f582,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : ((p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X2)))) & p(s(bool,'const.relation.INDUCTIVE_INVARIANT_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,fun(X0,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X4))))) => ! [X5] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X3),s(X1,X5))),s(X0,'const.relation.WFREC_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X4),s(X1,X5)))))))),
  inference(rectify,[],[f153])).
fof(f153,axiom,(
  ! [X20,X6] : ! [X18,X17,X61] : ((p(s(bool,'const.relation.WF_1'(s(fun(X6,fun(X6,bool)),X18)))) & p(s(bool,'const.relation.INDUCTIVE_INVARIANT_3'(s(fun(X6,fun(X6,bool)),X18),s(fun(X6,fun(X20,bool)),X17),s(fun(fun(X6,X20),fun(X6,X20)),X61))))) => ! [X9] : p(s(bool,happ(s(fun(X20,bool),happ(s(fun(X6,fun(X20,bool)),X17),s(X6,X9))),s(X20,'const.relation.WFREC_3'(s(fun(X6,fun(X6,bool)),X18),s(fun(fun(X6,X20),fun(X6,X20)),X61),s(X6,X9)))))))),
  file('thm.pred_set.EXTENSION','thm.relation.INDUCTIVE_INVARIANT_WFREC')).
fof(f1480,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X3)))) | s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(fun(X1,X0),fun(X1,X0)),X2),s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),X4),s(fun(X1,fun(X1,bool)),X3),s(X1,X5))))),s(X1,X5))) | s(fun(X1,X0),X4) != s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X2)))) )),
  inference(cnf_transformation,[],[f801])).
fof(f801,plain,(
  ! [X0,X1,X2,X3,X4] : (s(fun(X1,X0),X4) != s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X2))) | ~p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X3)))) | ! [X5] : s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(fun(X1,X0),fun(X1,X0)),X2),s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),X4),s(fun(X1,fun(X1,bool)),X3),s(X1,X5))))),s(X1,X5))))),
  inference(flattening,[],[f800])).
fof(f800,plain,(
  ! [X0,X1,X2,X3,X4] : (s(fun(X1,X0),X4) != s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X2))) | (~p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X3)))) | ! [X5] : s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(fun(X1,X0),fun(X1,X0)),X2),s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),X4),s(fun(X1,fun(X1,bool)),X3),s(X1,X5))))),s(X1,X5)))))),
  inference(ennf_transformation,[],[f581])).
fof(f581,plain,(
  ! [X0,X1,X2,X3,X4] : (s(fun(X1,X0),X4) = s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X2))) => (p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X3)))) => ! [X5] : s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(fun(X1,X0),fun(X1,X0)),X2),s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),X4),s(fun(X1,fun(X1,bool)),X3),s(X1,X5))))),s(X1,X5)))))),
  inference(flattening,[],[f580])).
fof(f580,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (s(fun(X1,X0),X4) = s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X2))) => (p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X3)))) => ! [X5] : s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(fun(X1,X0),fun(X1,X0)),X2),s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),X4),s(fun(X1,fun(X1,bool)),X3),s(X1,X5))))),s(X1,X5)))))),
  inference(rectify,[],[f228])).
fof(f228,axiom,(
  ! [X20,X6] : ! [X61,X18,X15] : (s(fun(X6,X20),X15) = s(fun(X6,X20),'const.relation.WFREC_2'(s(fun(X6,fun(X6,bool)),X18),s(fun(fun(X6,X20),fun(X6,X20)),X61))) => (p(s(bool,'const.relation.WF_1'(s(fun(X6,fun(X6,bool)),X18)))) => ! [X9] : s(X20,happ(s(fun(X6,X20),X15),s(X6,X9))) = s(X20,happ(s(fun(X6,X20),happ(s(fun(fun(X6,X20),fun(X6,X20)),X61),s(fun(X6,X20),'const.relation.RESTRICT_3'(s(fun(X6,X20),X15),s(fun(X6,fun(X6,bool)),X18),s(X6,X9))))),s(X6,X9)))))),
  file('thm.pred_set.EXTENSION','thm.relation.WFREC_COROLLARY')).
fof(f1479,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X1,X4),s(fun(X1,bool),X2)))) | s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))) = s(X0,'const.bool.RES_ABSTRACT_3'(s(fun(X1,bool),X2),s(fun(X1,X0),X3),s(X1,X4)))) )),
  inference(cnf_transformation,[],[f799])).
fof(f799,plain,(
  ! [X0,X1,X2,X3,X4] : (~p(s(bool,'const.bool.IN_2'(s(X1,X4),s(fun(X1,bool),X2)))) | s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))) = s(X0,'const.bool.RES_ABSTRACT_3'(s(fun(X1,bool),X2),s(fun(X1,X0),X3),s(X1,X4))))),
  inference(ennf_transformation,[],[f579])).
fof(f579,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.bool.IN_2'(s(X1,X4),s(fun(X1,bool),X2)))) => s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))) = s(X0,'const.bool.RES_ABSTRACT_3'(s(fun(X1,bool),X2),s(fun(X1,X0),X3),s(X1,X4))))),
  inference(flattening,[],[f578])).
fof(f578,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.bool.IN_2'(s(X1,X4),s(fun(X1,bool),X2)))) => s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))) = s(X0,'const.bool.RES_ABSTRACT_3'(s(fun(X1,bool),X2),s(fun(X1,X0),X3),s(X1,X4))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X20,X6] : ! [X23,X26,X9] : (p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X23)))) => s(X20,happ(s(fun(X6,X20),X26),s(X6,X9))) = s(X20,'const.bool.RES_ABSTRACT_3'(s(fun(X6,bool),X23),s(fun(X6,X20),X26),s(X6,X9))))),
  file('thm.pred_set.EXTENSION','thm.bool.RES_ABSTRACT_DEF|split|0')).
fof(f1478,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun(X1,X0),'const.relation.the_fun_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X3),s(X1,X4))) = s(fun(X1,X0),'const.min.@_1'(s(fun(fun(X1,X0),bool),'const.relation.approx_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X3),s(X1,X4)))))) )),
  inference(cnf_transformation,[],[f577])).
fof(f577,plain,(
  ! [X0,X1,X2,X3,X4] : s(fun(X1,X0),'const.relation.the_fun_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X3),s(X1,X4))) = s(fun(X1,X0),'const.min.@_1'(s(fun(fun(X1,X0),bool),'const.relation.approx_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X3),s(X1,X4)))))),
  inference(flattening,[],[f576])).
fof(f576,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : s(fun(X1,X0),'const.relation.the_fun_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X3),s(X1,X4))) = s(fun(X1,X0),'const.min.@_1'(s(fun(fun(X1,X0),bool),'const.relation.approx_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X3),s(X1,X4)))))),
  inference(rectify,[],[f152])).
fof(f152,axiom,(
  ! [X20,X6] : ! [X18,X61,X9] : s(fun(X6,X20),'const.relation.the_fun_3'(s(fun(X6,fun(X6,bool)),X18),s(fun(fun(X6,X20),fun(X6,X20)),X61),s(X6,X9))) = s(fun(X6,X20),'const.min.@_1'(s(fun(fun(X6,X20),bool),'const.relation.approx_3'(s(fun(X6,fun(X6,bool)),X18),s(fun(fun(X6,X20),fun(X6,X20)),X61),s(X6,X9)))))),
  file('thm.pred_set.EXTENSION','thm.relation.the_fun_def')).
fof(f1877,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X2)))) | s(X0,'const.relation.WFREC_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X3),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(fun(X1,X0),fun(X1,X0)),X3),s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X3))),s(fun(X1,fun(X1,bool)),X2),s(X1,X4))))),s(X1,X4)))) )),
  inference(backward_demodulation,[],[f1477,f1456])).
fof(f1477,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.relation.WFREC_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X3),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X3))),s(X1,X4)))) )),
  inference(cnf_transformation,[],[f575])).
fof(f575,plain,(
  ! [X0,X1,X2,X3,X4] : s(X0,'const.relation.WFREC_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X3),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X3))),s(X1,X4)))),
  inference(flattening,[],[f574])).
fof(f574,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : s(X0,'const.relation.WFREC_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X3),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X3))),s(X1,X4)))),
  inference(rectify,[],[f150])).
fof(f150,axiom,(
  ! [X20,X6] : ! [X21,X22,X31] : s(X20,'const.relation.WFREC_3'(s(fun(X6,fun(X6,bool)),X21),s(fun(fun(X6,X20),fun(X6,X20)),X22),s(X6,X31))) = s(X20,happ(s(fun(X6,X20),'const.relation.WFREC_2'(s(fun(X6,fun(X6,bool)),X21),s(fun(fun(X6,X20),fun(X6,X20)),X22))),s(X6,X31)))),
  file('thm.pred_set.EXTENSION','ARITY_2_3_const.relation.WFREC')).
fof(f1476,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),X2),s(fun(X1,fun(X1,bool)),X3),s(X1,X4))) = s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),'const.relation.RESTRICT_2'(s(fun(X1,X0),X2),s(fun(X1,fun(X1,bool)),X3))),s(X1,X4)))) )),
  inference(cnf_transformation,[],[f573])).
fof(f573,plain,(
  ! [X0,X1,X2,X3,X4] : s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),X2),s(fun(X1,fun(X1,bool)),X3),s(X1,X4))) = s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),'const.relation.RESTRICT_2'(s(fun(X1,X0),X2),s(fun(X1,fun(X1,bool)),X3))),s(X1,X4)))),
  inference(flattening,[],[f572])).
fof(f572,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),X2),s(fun(X1,fun(X1,bool)),X3),s(X1,X4))) = s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),'const.relation.RESTRICT_2'(s(fun(X1,X0),X2),s(fun(X1,fun(X1,bool)),X3))),s(X1,X4)))),
  inference(rectify,[],[f127])).
fof(f127,axiom,(
  ! [X20,X6] : ! [X21,X22,X31] : s(fun(X6,X20),'const.relation.RESTRICT_3'(s(fun(X6,X20),X21),s(fun(X6,fun(X6,bool)),X22),s(X6,X31))) = s(fun(X6,X20),happ(s(fun(X6,fun(X6,X20)),'const.relation.RESTRICT_2'(s(fun(X6,X20),X21),s(fun(X6,fun(X6,bool)),X22))),s(X6,X31)))),
  file('thm.pred_set.EXTENSION','ARITY_2_3_const.relation.RESTRICT')).
fof(f1475,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.bool.RES_ABSTRACT_3'(s(fun(X1,bool),X2),s(fun(X1,X0),X3),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),'const.bool.RES_ABSTRACT_2'(s(fun(X1,bool),X2),s(fun(X1,X0),X3))),s(X1,X4)))) )),
  inference(cnf_transformation,[],[f571])).
fof(f571,plain,(
  ! [X0,X1,X2,X3,X4] : s(X0,'const.bool.RES_ABSTRACT_3'(s(fun(X1,bool),X2),s(fun(X1,X0),X3),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),'const.bool.RES_ABSTRACT_2'(s(fun(X1,bool),X2),s(fun(X1,X0),X3))),s(X1,X4)))),
  inference(flattening,[],[f570])).
fof(f570,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : s(X0,'const.bool.RES_ABSTRACT_3'(s(fun(X1,bool),X2),s(fun(X1,X0),X3),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),'const.bool.RES_ABSTRACT_2'(s(fun(X1,bool),X2),s(fun(X1,X0),X3))),s(X1,X4)))),
  inference(rectify,[],[f101])).
fof(f101,axiom,(
  ! [X20,X6] : ! [X21,X22,X31] : s(X20,'const.bool.RES_ABSTRACT_3'(s(fun(X6,bool),X21),s(fun(X6,X20),X22),s(X6,X31))) = s(X20,happ(s(fun(X6,X20),'const.bool.RES_ABSTRACT_2'(s(fun(X6,bool),X21),s(fun(X6,X20),X22))),s(X6,X31)))),
  file('thm.pred_set.EXTENSION','ARITY_2_3_const.bool.RES_ABSTRACT')).
fof(f1876,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X1,sK87(X3,X2,X1,X0)),s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X2)))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.relation.RDOM_1'(s(fun(X0,fun(X1,bool)),X3))))))) )),
  inference(forward_demodulation,[],[f1473,f1257])).
fof(f1473,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X2))),s(X1,sK87(X3,X2,X1,X0))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.relation.RDOM_1'(s(fun(X0,fun(X1,bool)),X3))))))) )),
  inference(cnf_transformation,[],[f1038])).
fof(f1038,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.relation.RDOM_1'(s(fun(X0,fun(X1,bool)),X3)))))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X2))),s(X1,sK87(X3,X2,X1,X0)))))) & (! [X4] : ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X2))),s(X1,X4)))) | p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.relation.RDOM_1'(s(fun(X0,fun(X1,bool)),X3))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK87])],[f1037])).
fof(f1037,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.relation.RDOM_1'(s(fun(X0,fun(X1,bool)),X3)))))) | ? [X5] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X2))),s(X1,X5))))) & (! [X4] : ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X2))),s(X1,X4)))) | p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.relation.RDOM_1'(s(fun(X0,fun(X1,bool)),X3))))))))),
  inference(rectify,[],[f1036])).
fof(f1036,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.relation.RDOM_1'(s(fun(X0,fun(X1,bool)),X3)))))) | ? [X4] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X2))),s(X1,X4))))) & (! [X4] : ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X2))),s(X1,X4)))) | p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.relation.RDOM_1'(s(fun(X0,fun(X1,bool)),X3))))))))),
  inference(nnf_transformation,[],[f569])).
fof(f569,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.relation.RDOM_1'(s(fun(X0,fun(X1,bool)),X3)))))) <=> ? [X4] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X2))),s(X1,X4)))))),
  inference(flattening,[],[f568])).
fof(f568,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.relation.RDOM_1'(s(fun(X0,fun(X1,bool)),X3)))))) <=> ? [X4] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X2))),s(X1,X4)))))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X6,X20] : ! [X9,X18] : (p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),'const.relation.RDOM_1'(s(fun(X6,fun(X20,bool)),X18)))))) <=> ? [X19] : p(s(bool,happ(s(fun(X20,bool),happ(s(fun(X6,fun(X20,bool)),X18),s(X6,X9))),s(X20,X19)))))),
  file('thm.pred_set.EXTENSION','thm.relation.IN_RDOM')).
fof(f1875,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X1,X4),s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X2)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.relation.RDOM_1'(s(fun(X0,fun(X1,bool)),X3))))))) )),
  inference(forward_demodulation,[],[f1474,f1257])).
fof(f1474,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.relation.RDOM_1'(s(fun(X0,fun(X1,bool)),X3)))))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X2))),s(X1,X4))))) )),
  inference(cnf_transformation,[],[f1038])).
fof(f1874,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,sK86(X3,X2,X1,X0))))))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RRANGE_1'(s(fun(X0,fun(X1,bool)),X3))))))) )),
  inference(forward_demodulation,[],[f1471,f1257])).
fof(f1471,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,sK86(X3,X2,X1,X0)))),s(X1,X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RRANGE_1'(s(fun(X0,fun(X1,bool)),X3))))))) )),
  inference(cnf_transformation,[],[f1035])).
fof(f1035,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RRANGE_1'(s(fun(X0,fun(X1,bool)),X3)))))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,sK86(X3,X2,X1,X0)))),s(X1,X2))))) & (! [X4] : ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X4))),s(X1,X2)))) | p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RRANGE_1'(s(fun(X0,fun(X1,bool)),X3))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK86])],[f1034])).
fof(f1034,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RRANGE_1'(s(fun(X0,fun(X1,bool)),X3)))))) | ? [X5] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X5))),s(X1,X2))))) & (! [X4] : ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X4))),s(X1,X2)))) | p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RRANGE_1'(s(fun(X0,fun(X1,bool)),X3))))))))),
  inference(rectify,[],[f1033])).
fof(f1033,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RRANGE_1'(s(fun(X0,fun(X1,bool)),X3)))))) | ? [X4] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X4))),s(X1,X2))))) & (! [X4] : ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X4))),s(X1,X2)))) | p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RRANGE_1'(s(fun(X0,fun(X1,bool)),X3))))))))),
  inference(nnf_transformation,[],[f567])).
fof(f567,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RRANGE_1'(s(fun(X0,fun(X1,bool)),X3)))))) <=> ? [X4] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X4))),s(X1,X2)))))),
  inference(flattening,[],[f566])).
fof(f566,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RRANGE_1'(s(fun(X0,fun(X1,bool)),X3)))))) <=> ? [X4] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X4))),s(X1,X2)))))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X20,X6] : ! [X19,X18] : (p(s(bool,'const.bool.IN_2'(s(X6,X19),s(fun(X6,bool),'const.relation.RRANGE_1'(s(fun(X20,fun(X6,bool)),X18)))))) <=> ? [X9] : p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X20,fun(X6,bool)),X18),s(X20,X9))),s(X6,X19)))))),
  file('thm.pred_set.EXTENSION','thm.relation.IN_RRANGE')).
fof(f1873,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X4)))))) | p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RRANGE_1'(s(fun(X0,fun(X1,bool)),X3))))))) )),
  inference(forward_demodulation,[],[f1472,f1257])).
fof(f1472,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RRANGE_1'(s(fun(X0,fun(X1,bool)),X3)))))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X4))),s(X1,X2))))) )),
  inference(cnf_transformation,[],[f1035])).
fof(f1872,plain,(
  ( ! [X6,X2,X0,X7,X3,X1] : (~p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3)))) | p(s(bool,'const.bool.IN_2'(s(X1,X7),s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X6)))))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X7),s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X6))))))) )),
  inference(forward_demodulation,[],[f1871,f1257])).
fof(f1871,plain,(
  ( ! [X6,X2,X0,X7,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X1,X7),s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X6)))))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X6))),s(X1,X7)))) | ~p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3))))) )),
  inference(forward_demodulation,[],[f1468,f1257])).
fof(f1468,plain,(
  ( ! [X6,X2,X0,X7,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X6))),s(X1,X7)))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X6))),s(X1,X7)))) | ~p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3))))) )),
  inference(cnf_transformation,[],[f1032])).
fof(f1032,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3)))) | ! [X6,X7] : (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X6))),s(X1,X7)))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X6))),s(X1,X7)))))) & ((p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,sK84(X3,X2,X1,X0)))),s(X1,sK85(X3,X2,X1,X0))))) & ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,sK84(X3,X2,X1,X0)))),s(X1,sK85(X3,X2,X1,X0)))))) | p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK84,sK85])],[f1031])).
fof(f1031,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3)))) | ! [X6,X7] : (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X6))),s(X1,X7)))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X6))),s(X1,X7)))))) & (? [X4,X5] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))) & ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X4))),s(X1,X5))))) | p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3))))))),
  inference(rectify,[],[f1030])).
fof(f1030,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3)))) | ! [X4,X5] : (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X4))),s(X1,X5)))))) & (? [X4,X5] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))) & ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X4))),s(X1,X5))))) | p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3))))))),
  inference(nnf_transformation,[],[f798])).
fof(f798,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3)))) <=> ! [X4,X5] : (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X4))),s(X1,X5))))))),
  inference(ennf_transformation,[],[f565])).
fof(f565,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3)))) <=> ! [X4,X5] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))) => p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X4))),s(X1,X5))))))),
  inference(flattening,[],[f564])).
fof(f564,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3)))) <=> ! [X4,X5] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))) => p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X4))),s(X1,X5))))))),
  inference(rectify,[],[f203])).
fof(f203,axiom,(
  ! [X6,X20] : ! [X30,X29] : (p(s(bool,'const.relation.RSUBSET_2'(s(fun(X6,fun(X20,bool)),X30),s(fun(X6,fun(X20,bool)),X29)))) <=> ! [X9,X19] : (p(s(bool,happ(s(fun(X20,bool),happ(s(fun(X6,fun(X20,bool)),X30),s(X6,X9))),s(X20,X19)))) => p(s(bool,happ(s(fun(X20,bool),happ(s(fun(X6,fun(X20,bool)),X29),s(X6,X9))),s(X20,X19))))))),
  file('thm.pred_set.EXTENSION','thm.relation.RSUBSET')).
fof(f1870,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X1,sK85(X3,X2,X1,X0)),s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,sK84(X3,X2,X1,X0))))))) | p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3))))) )),
  inference(forward_demodulation,[],[f1469,f1257])).
fof(f1469,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3)))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,sK84(X3,X2,X1,X0)))),s(X1,sK85(X3,X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f1032])).
fof(f1869,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X1,sK85(X3,X2,X1,X0)),s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,sK84(X3,X2,X1,X0))))))) | p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3))))) )),
  inference(forward_demodulation,[],[f1470,f1257])).
fof(f1470,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3)))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,sK84(X3,X2,X1,X0)))),s(X1,sK85(X3,X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f1032])).
fof(f1466,plain,(
  ( ! [X2,X0,X3,X1] : (s(fun(X0,fun(X1,bool)),X2) != s(fun(X0,fun(X1,bool)),X3) | p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X3),s(fun(X0,fun(X1,bool)),X2))))) )),
  inference(cnf_transformation,[],[f1029])).
fof(f1029,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X0,fun(X1,bool)),X2) != s(fun(X0,fun(X1,bool)),X3) | (p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3)))) & p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X3),s(fun(X0,fun(X1,bool)),X2)))))) & (~p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3)))) | ~p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X3),s(fun(X0,fun(X1,bool)),X2)))) | s(fun(X0,fun(X1,bool)),X2) = s(fun(X0,fun(X1,bool)),X3)))),
  inference(flattening,[],[f1028])).
fof(f1028,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X0,fun(X1,bool)),X2) != s(fun(X0,fun(X1,bool)),X3) | (p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3)))) & p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X3),s(fun(X0,fun(X1,bool)),X2)))))) & ((~p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3)))) | ~p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X3),s(fun(X0,fun(X1,bool)),X2))))) | s(fun(X0,fun(X1,bool)),X2) = s(fun(X0,fun(X1,bool)),X3)))),
  inference(nnf_transformation,[],[f563])).
fof(f563,plain,(
  ! [X0,X1,X2,X3] : (s(fun(X0,fun(X1,bool)),X2) = s(fun(X0,fun(X1,bool)),X3) <=> (p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3)))) & p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X3),s(fun(X0,fun(X1,bool)),X2))))))),
  inference(flattening,[],[f562])).
fof(f562,plain,(
  ! [X0,X1] : ! [X2,X3] : (s(fun(X0,fun(X1,bool)),X2) = s(fun(X0,fun(X1,bool)),X3) <=> (p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3)))) & p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X3),s(fun(X0,fun(X1,bool)),X2))))))),
  inference(rectify,[],[f218])).
fof(f218,axiom,(
  ! [X6,X20] : ! [X19,X49] : (s(fun(X6,fun(X20,bool)),X19) = s(fun(X6,fun(X20,bool)),X49) <=> (p(s(bool,'const.relation.RSUBSET_2'(s(fun(X6,fun(X20,bool)),X19),s(fun(X6,fun(X20,bool)),X49)))) & p(s(bool,'const.relation.RSUBSET_2'(s(fun(X6,fun(X20,bool)),X49),s(fun(X6,fun(X20,bool)),X19))))))),
  file('thm.pred_set.EXTENSION','thm.relation.EqIsBothRSUBSET')).
fof(f1868,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.relation.RRANGE_2'(s(fun(X0,fun(X1,bool)),X2),s(X1,X3)))) | p(s(bool,'const.bool.IN_2'(s(X1,X3),s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,sK83(X3,X2,X1,X0)))))))) )),
  inference(forward_demodulation,[],[f1463,f1257])).
fof(f1463,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,sK83(X3,X2,X1,X0)))),s(X1,X3)))) | ~p(s(bool,'const.relation.RRANGE_2'(s(fun(X0,fun(X1,bool)),X2),s(X1,X3))))) )),
  inference(cnf_transformation,[],[f1027])).
fof(f1027,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.RRANGE_2'(s(fun(X0,fun(X1,bool)),X2),s(X1,X3)))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,sK83(X3,X2,X1,X0)))),s(X1,X3))))) & (! [X4] : ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X3)))) | p(s(bool,'const.relation.RRANGE_2'(s(fun(X0,fun(X1,bool)),X2),s(X1,X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK83])],[f1026])).
fof(f1026,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.RRANGE_2'(s(fun(X0,fun(X1,bool)),X2),s(X1,X3)))) | ? [X5] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X3))))) & (! [X4] : ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X3)))) | p(s(bool,'const.relation.RRANGE_2'(s(fun(X0,fun(X1,bool)),X2),s(X1,X3))))))),
  inference(rectify,[],[f1025])).
fof(f1025,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.RRANGE_2'(s(fun(X0,fun(X1,bool)),X2),s(X1,X3)))) | ? [X4] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X3))))) & (! [X4] : ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X3)))) | p(s(bool,'const.relation.RRANGE_2'(s(fun(X0,fun(X1,bool)),X2),s(X1,X3))))))),
  inference(nnf_transformation,[],[f561])).
fof(f561,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.relation.RRANGE_2'(s(fun(X0,fun(X1,bool)),X2),s(X1,X3)))) <=> ? [X4] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X3)))))),
  inference(flattening,[],[f560])).
fof(f560,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.relation.RRANGE_2'(s(fun(X0,fun(X1,bool)),X2),s(X1,X3)))) <=> ? [X4] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X3)))))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X6,X20] : ! [X18,X19] : (p(s(bool,'const.relation.RRANGE_2'(s(fun(X6,fun(X20,bool)),X18),s(X20,X19)))) <=> ? [X9] : p(s(bool,happ(s(fun(X20,bool),happ(s(fun(X6,fun(X20,bool)),X18),s(X6,X9))),s(X20,X19)))))),
  file('thm.pred_set.EXTENSION','thm.relation.RRANGE')).
fof(f1867,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X1,X3),s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4)))))) | p(s(bool,'const.relation.RRANGE_2'(s(fun(X0,fun(X1,bool)),X2),s(X1,X3))))) )),
  inference(forward_demodulation,[],[f1464,f1257])).
fof(f1464,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.relation.RRANGE_2'(s(fun(X0,fun(X1,bool)),X2),s(X1,X3)))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X3))))) )),
  inference(cnf_transformation,[],[f1027])).
fof(f1866,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.relation.RDOM_2'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3)))) | p(s(bool,'const.bool.IN_2'(s(X1,sK82(X3,X2,X1,X0)),s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))))))) )),
  inference(forward_demodulation,[],[f1461,f1257])).
fof(f1461,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))),s(X1,sK82(X3,X2,X1,X0))))) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f1024])).
fof(f1024,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.RDOM_2'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3)))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))),s(X1,sK82(X3,X2,X1,X0)))))) & (! [X4] : ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))),s(X1,X4)))) | p(s(bool,'const.relation.RDOM_2'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK82])],[f1023])).
fof(f1023,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.RDOM_2'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3)))) | ? [X5] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))),s(X1,X5))))) & (! [X4] : ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))),s(X1,X4)))) | p(s(bool,'const.relation.RDOM_2'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))))))),
  inference(rectify,[],[f1022])).
fof(f1022,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.RDOM_2'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3)))) | ? [X4] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))),s(X1,X4))))) & (! [X4] : ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))),s(X1,X4)))) | p(s(bool,'const.relation.RDOM_2'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))))))),
  inference(nnf_transformation,[],[f559])).
fof(f559,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.relation.RDOM_2'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3)))) <=> ? [X4] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))),s(X1,X4)))))),
  inference(flattening,[],[f558])).
fof(f558,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.relation.RDOM_2'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3)))) <=> ? [X4] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))),s(X1,X4)))))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X6,X20] : ! [X18,X9] : (p(s(bool,'const.relation.RDOM_2'(s(fun(X6,fun(X20,bool)),X18),s(X6,X9)))) <=> ? [X19] : p(s(bool,happ(s(fun(X20,bool),happ(s(fun(X6,fun(X20,bool)),X18),s(X6,X9))),s(X20,X19)))))),
  file('thm.pred_set.EXTENSION','thm.relation.RDOM_DEF')).
fof(f1865,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X1,X4),s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X3)))))) | p(s(bool,'const.relation.RDOM_2'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))))) )),
  inference(forward_demodulation,[],[f1462,f1257])).
fof(f1462,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.relation.RDOM_2'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))),s(X1,X4))))) )),
  inference(cnf_transformation,[],[f1024])).
fof(f1459,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))) )),
  inference(cnf_transformation,[],[f1021])).
fof(f1457,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3)))) | ~p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X3),s(fun(X0,fun(X1,bool)),X2)))) | s(fun(X0,fun(X1,bool)),X2) = s(fun(X0,fun(X1,bool)),X3)) )),
  inference(cnf_transformation,[],[f796])).
fof(f796,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3)))) | ~p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X3),s(fun(X0,fun(X1,bool)),X2)))) | s(fun(X0,fun(X1,bool)),X2) = s(fun(X0,fun(X1,bool)),X3))),
  inference(flattening,[],[f795])).
fof(f795,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3)))) | ~p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X3),s(fun(X0,fun(X1,bool)),X2))))) | s(fun(X0,fun(X1,bool)),X2) = s(fun(X0,fun(X1,bool)),X3))),
  inference(ennf_transformation,[],[f553])).
fof(f553,plain,(
  ! [X0,X1,X2,X3] : ((p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3)))) & p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X3),s(fun(X0,fun(X1,bool)),X2))))) => s(fun(X0,fun(X1,bool)),X2) = s(fun(X0,fun(X1,bool)),X3))),
  inference(flattening,[],[f552])).
fof(f552,plain,(
  ! [X0,X1] : ! [X2,X3] : ((p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3)))) & p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X3),s(fun(X0,fun(X1,bool)),X2))))) => s(fun(X0,fun(X1,bool)),X2) = s(fun(X0,fun(X1,bool)),X3))),
  inference(rectify,[],[f202])).
fof(f202,axiom,(
  ! [X6,X20] : ! [X30,X29] : ((p(s(bool,'const.relation.RSUBSET_2'(s(fun(X6,fun(X20,bool)),X30),s(fun(X6,fun(X20,bool)),X29)))) & p(s(bool,'const.relation.RSUBSET_2'(s(fun(X6,fun(X20,bool)),X29),s(fun(X6,fun(X20,bool)),X30))))) => s(fun(X6,fun(X20,bool)),X29) = s(fun(X6,fun(X20,bool)),X30))),
  file('thm.pred_set.EXTENSION','thm.relation.RSUBSET_ANTISYM')).
fof(f1456,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,happ(s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X3))),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(fun(X1,X0),fun(X1,X0)),X3),s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X3))),s(fun(X1,fun(X1,bool)),X2),s(X1,X4))))),s(X1,X4))) | ~p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X2))))) )),
  inference(cnf_transformation,[],[f794])).
fof(f794,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X2)))) | ! [X4] : s(X0,happ(s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X3))),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(fun(X1,X0),fun(X1,X0)),X3),s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X3))),s(fun(X1,fun(X1,bool)),X2),s(X1,X4))))),s(X1,X4))))),
  inference(ennf_transformation,[],[f551])).
fof(f551,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X2)))) => ! [X4] : s(X0,happ(s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X3))),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(fun(X1,X0),fun(X1,X0)),X3),s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X3))),s(fun(X1,fun(X1,bool)),X2),s(X1,X4))))),s(X1,X4))))),
  inference(flattening,[],[f550])).
fof(f550,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X2)))) => ! [X4] : s(X0,happ(s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X3))),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(fun(X1,X0),fun(X1,X0)),X3),s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X3))),s(fun(X1,fun(X1,bool)),X2),s(X1,X4))))),s(X1,X4))))),
  inference(rectify,[],[f234])).
fof(f234,axiom,(
  ! [X20,X6] : ! [X18,X61] : (p(s(bool,'const.relation.WF_1'(s(fun(X6,fun(X6,bool)),X18)))) => ! [X9] : s(X20,happ(s(fun(X6,X20),'const.relation.WFREC_2'(s(fun(X6,fun(X6,bool)),X18),s(fun(fun(X6,X20),fun(X6,X20)),X61))),s(X6,X9))) = s(X20,happ(s(fun(X6,X20),happ(s(fun(fun(X6,X20),fun(X6,X20)),X61),s(fun(X6,X20),'const.relation.RESTRICT_3'(s(fun(X6,X20),'const.relation.WFREC_2'(s(fun(X6,fun(X6,bool)),X18),s(fun(fun(X6,X20),fun(X6,X20)),X61))),s(fun(X6,fun(X6,bool)),X18),s(X6,X9))))),s(X6,X9))))),
  file('thm.pred_set.EXTENSION','thm.relation.WFREC_THM')).
fof(f1455,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3))) = s(bool,happ(s(fun(fun(X0,fun(X1,bool)),bool),happ(s(fun(fun(X0,fun(X1,bool)),fun(fun(X0,fun(X1,bool)),bool)),'const.relation.RSUBSET_0'),s(fun(X0,fun(X1,bool)),X2))),s(fun(X0,fun(X1,bool)),X3)))) )),
  inference(cnf_transformation,[],[f549])).
fof(f549,plain,(
  ! [X0,X1,X2,X3] : s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3))) = s(bool,happ(s(fun(fun(X0,fun(X1,bool)),bool),happ(s(fun(fun(X0,fun(X1,bool)),fun(fun(X0,fun(X1,bool)),bool)),'const.relation.RSUBSET_0'),s(fun(X0,fun(X1,bool)),X2))),s(fun(X0,fun(X1,bool)),X3)))),
  inference(flattening,[],[f548])).
fof(f548,plain,(
  ! [X0,X1] : ! [X2,X3] : s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3))) = s(bool,happ(s(fun(fun(X0,fun(X1,bool)),bool),happ(s(fun(fun(X0,fun(X1,bool)),fun(fun(X0,fun(X1,bool)),bool)),'const.relation.RSUBSET_0'),s(fun(X0,fun(X1,bool)),X2))),s(fun(X0,fun(X1,bool)),X3)))),
  inference(rectify,[],[f220])).
fof(f220,axiom,(
  ! [X6,X20] : ! [X21,X22] : s(bool,'const.relation.RSUBSET_2'(s(fun(X6,fun(X20,bool)),X21),s(fun(X6,fun(X20,bool)),X22))) = s(bool,happ(s(fun(fun(X6,fun(X20,bool)),bool),happ(s(fun(fun(X6,fun(X20,bool)),fun(fun(X6,fun(X20,bool)),bool)),'const.relation.RSUBSET_0'),s(fun(X6,fun(X20,bool)),X21))),s(fun(X6,fun(X20,bool)),X22)))),
  file('thm.pred_set.EXTENSION','ARITY_0_2_const.relation.RSUBSET')).
fof(f1454,plain,(
  ( ! [X2,X0,X3,X1] : (s(fun(X0,fun(X0,bool)),'const.relation.inv_image_2'(s(fun(X1,fun(X1,bool)),X2),s(fun(X0,X1),X3))) = s(fun(X0,fun(X0,bool)),happ(s(fun(fun(X0,X1),fun(X0,fun(X0,bool))),'const.relation.inv_image_1'(s(fun(X1,fun(X1,bool)),X2))),s(fun(X0,X1),X3)))) )),
  inference(cnf_transformation,[],[f547])).
fof(f547,plain,(
  ! [X0,X1,X2,X3] : s(fun(X0,fun(X0,bool)),'const.relation.inv_image_2'(s(fun(X1,fun(X1,bool)),X2),s(fun(X0,X1),X3))) = s(fun(X0,fun(X0,bool)),happ(s(fun(fun(X0,X1),fun(X0,fun(X0,bool))),'const.relation.inv_image_1'(s(fun(X1,fun(X1,bool)),X2))),s(fun(X0,X1),X3)))),
  inference(flattening,[],[f546])).
fof(f546,plain,(
  ! [X0,X1] : ! [X2,X3] : s(fun(X0,fun(X0,bool)),'const.relation.inv_image_2'(s(fun(X1,fun(X1,bool)),X2),s(fun(X0,X1),X3))) = s(fun(X0,fun(X0,bool)),happ(s(fun(fun(X0,X1),fun(X0,fun(X0,bool))),'const.relation.inv_image_1'(s(fun(X1,fun(X1,bool)),X2))),s(fun(X0,X1),X3)))),
  inference(rectify,[],[f241])).
fof(f241,axiom,(
  ! [X6,X20] : ! [X21,X22] : s(fun(X6,fun(X6,bool)),'const.relation.inv_image_2'(s(fun(X20,fun(X20,bool)),X21),s(fun(X6,X20),X22))) = s(fun(X6,fun(X6,bool)),happ(s(fun(fun(X6,X20),fun(X6,fun(X6,bool))),'const.relation.inv_image_1'(s(fun(X20,fun(X20,bool)),X21))),s(fun(X6,X20),X22)))),
  file('thm.pred_set.EXTENSION','ARITY_1_2_const.relation.inv_image')).
fof(f1453,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.bool.LET_2'(s(fun(X1,X0),X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(fun(X1,X0),fun(X1,X0)),'const.bool.LET_0'),s(fun(X1,X0),X2))),s(X1,X3)))) )),
  inference(cnf_transformation,[],[f545])).
fof(f545,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.bool.LET_2'(s(fun(X1,X0),X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(fun(X1,X0),fun(X1,X0)),'const.bool.LET_0'),s(fun(X1,X0),X2))),s(X1,X3)))),
  inference(flattening,[],[f544])).
fof(f544,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X0,'const.bool.LET_2'(s(fun(X1,X0),X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(fun(X1,X0),fun(X1,X0)),'const.bool.LET_0'),s(fun(X1,X0),X2))),s(X1,X3)))),
  inference(rectify,[],[f165])).
fof(f165,axiom,(
  ! [X20,X6] : ! [X21,X22] : s(X20,'const.bool.LET_2'(s(fun(X6,X20),X21),s(X6,X22))) = s(X20,happ(s(fun(X6,X20),happ(s(fun(fun(X6,X20),fun(X6,X20)),'const.bool.LET_0'),s(fun(X6,X20),X21))),s(X6,X22)))),
  file('thm.pred_set.EXTENSION','ARITY_0_2_const.bool.LET')).
fof(f1452,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'const.relation.RDOM_2'(s(fun(X1,fun(X0,bool)),X2),s(X1,X3))) = s(bool,happ(s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),X2))),s(X1,X3)))) )),
  inference(cnf_transformation,[],[f543])).
fof(f543,plain,(
  ! [X0,X1,X2,X3] : s(bool,'const.relation.RDOM_2'(s(fun(X1,fun(X0,bool)),X2),s(X1,X3))) = s(bool,happ(s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),X2))),s(X1,X3)))),
  inference(flattening,[],[f542])).
fof(f542,plain,(
  ! [X0,X1] : ! [X2,X3] : s(bool,'const.relation.RDOM_2'(s(fun(X1,fun(X0,bool)),X2),s(X1,X3))) = s(bool,happ(s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),X2))),s(X1,X3)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X20,X6] : ! [X21,X22] : s(bool,'const.relation.RDOM_2'(s(fun(X6,fun(X20,bool)),X21),s(X6,X22))) = s(bool,happ(s(fun(X6,bool),'const.relation.RDOM_1'(s(fun(X6,fun(X20,bool)),X21))),s(X6,X22)))),
  file('thm.pred_set.EXTENSION','ARITY_1_2_const.relation.RDOM')).
fof(f1451,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'const.relation.RRANGE_2'(s(fun(X0,fun(X1,bool)),X2),s(X1,X3))) = s(bool,happ(s(fun(X1,bool),'const.relation.RRANGE_1'(s(fun(X0,fun(X1,bool)),X2))),s(X1,X3)))) )),
  inference(cnf_transformation,[],[f541])).
fof(f541,plain,(
  ! [X0,X1,X2,X3] : s(bool,'const.relation.RRANGE_2'(s(fun(X0,fun(X1,bool)),X2),s(X1,X3))) = s(bool,happ(s(fun(X1,bool),'const.relation.RRANGE_1'(s(fun(X0,fun(X1,bool)),X2))),s(X1,X3)))),
  inference(flattening,[],[f540])).
fof(f540,plain,(
  ! [X0,X1] : ! [X2,X3] : s(bool,'const.relation.RRANGE_2'(s(fun(X0,fun(X1,bool)),X2),s(X1,X3))) = s(bool,happ(s(fun(X1,bool),'const.relation.RRANGE_1'(s(fun(X0,fun(X1,bool)),X2))),s(X1,X3)))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X6,X20] : ! [X21,X22] : s(bool,'const.relation.RRANGE_2'(s(fun(X6,fun(X20,bool)),X21),s(X20,X22))) = s(bool,happ(s(fun(X20,bool),'const.relation.RRANGE_1'(s(fun(X6,fun(X20,bool)),X21))),s(X20,X22)))),
  file('thm.pred_set.EXTENSION','ARITY_1_2_const.relation.RRANGE')).
fof(f1864,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),'const.relation.RUNIV_0'),s(X0,X2))),s(X1,X3)))) )),
  inference(forward_demodulation,[],[f1450,f1447])).
fof(f1450,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'const.relation.RUNIV_2'(s(X0,X2),s(X1,X3))) = s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),'const.relation.RUNIV_0'),s(X0,X2))),s(X1,X3)))) )),
  inference(cnf_transformation,[],[f539])).
fof(f539,plain,(
  ! [X0,X1,X2,X3] : s(bool,'const.relation.RUNIV_2'(s(X0,X2),s(X1,X3))) = s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),'const.relation.RUNIV_0'),s(X0,X2))),s(X1,X3)))),
  inference(flattening,[],[f538])).
fof(f538,plain,(
  ! [X0,X1] : ! [X2,X3] : s(bool,'const.relation.RUNIV_2'(s(X0,X2),s(X1,X3))) = s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),'const.relation.RUNIV_0'),s(X0,X2))),s(X1,X3)))),
  inference(rectify,[],[f205])).
fof(f205,axiom,(
  ! [X6,X20] : ! [X21,X22] : s(bool,'const.relation.RUNIV_2'(s(X6,X21),s(X20,X22))) = s(bool,happ(s(fun(X20,bool),happ(s(fun(X6,fun(X20,bool)),'const.relation.RUNIV_0'),s(X6,X21))),s(X20,X22)))),
  file('thm.pred_set.EXTENSION','ARITY_0_2_const.relation.RUNIV')).
fof(f1448,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f535])).
fof(f535,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))),
  inference(flattening,[],[f534])).
fof(f534,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))),
  inference(rectify,[],[f179])).
fof(f179,axiom,(
  ! [X38,X35] : ! [X39,X42] : s(X35,X39) = s(X35,happ(s(fun(X38,X35),'_dst_x0_1'(s(X35,X39))),s(X38,X42)))),
  file('thm.pred_set.EXTENSION','_dst_x0_ax')).
fof(f1447,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,'const.relation.RUNIV_2'(s(X0,X2),s(X1,X3)))) )),
  inference(cnf_transformation,[],[f533])).
fof(f533,plain,(
  ! [X0,X1,X2,X3] : s(bool,'T_0') = s(bool,'const.relation.RUNIV_2'(s(X0,X2),s(X1,X3)))),
  inference(flattening,[],[f532])).
fof(f532,plain,(
  ! [X0,X1] : ! [X2,X3] : s(bool,'T_0') = s(bool,'const.relation.RUNIV_2'(s(X0,X2),s(X1,X3)))),
  inference(rectify,[],[f204])).
fof(f204,axiom,(
  ! [X6,X20] : ! [X9,X19] : s(bool,'T_0') = s(bool,'const.relation.RUNIV_2'(s(X6,X9),s(X20,X19)))),
  file('thm.pred_set.EXTENSION','thm.relation.RUNIV')).
fof(f1863,plain,(
  ( ! [X2,X0,X8,X7,X1] : (p(s(bool,'const.bool.IN_2'(s(X1,sK77(X8,X7,X2,X1,X0)),s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,sK76(X8,X7,X2,X1,X0))))))) | s(fun(X0,fun(X1,bool)),X2) = s(fun(X0,fun(X1,bool)),'const.pair.REP_prod_1'(s('type.pair.prod'(X0,X1),'const.pair.ABS_prod_1'(s(fun(X0,fun(X1,bool)),X2))))) | s(X0,X7) = s(X0,sK76(X8,X7,X2,X1,X0))) )),
  inference(forward_demodulation,[],[f1441,f1257])).
fof(f1441,plain,(
  ( ! [X2,X0,X8,X7,X1] : (s(fun(X0,fun(X1,bool)),X2) = s(fun(X0,fun(X1,bool)),'const.pair.REP_prod_1'(s('type.pair.prod'(X0,X1),'const.pair.ABS_prod_1'(s(fun(X0,fun(X1,bool)),X2))))) | s(X0,X7) = s(X0,sK76(X8,X7,X2,X1,X0)) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,sK76(X8,X7,X2,X1,X0)))),s(X1,sK77(X8,X7,X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f1017])).
fof(f1017,plain,(
  ! [X0,X1,X2] : ((! [X7,X8] : ((p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,sK76(X8,X7,X2,X1,X0)))),s(X1,sK77(X8,X7,X2,X1,X0))))) | (s(X0,X7) = s(X0,sK76(X8,X7,X2,X1,X0)) & s(X1,X8) = s(X1,sK77(X8,X7,X2,X1,X0)))) & (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,sK76(X8,X7,X2,X1,X0)))),s(X1,sK77(X8,X7,X2,X1,X0))))) | s(X0,X7) != s(X0,sK76(X8,X7,X2,X1,X0)) | s(X1,X8) != s(X1,sK77(X8,X7,X2,X1,X0)))) | s(fun(X0,fun(X1,bool)),X2) = s(fun(X0,fun(X1,bool)),'const.pair.REP_prod_1'(s('type.pair.prod'(X0,X1),'const.pair.ABS_prod_1'(s(fun(X0,fun(X1,bool)),X2)))))) & (s(fun(X0,fun(X1,bool)),X2) != s(fun(X0,fun(X1,bool)),'const.pair.REP_prod_1'(s('type.pair.prod'(X0,X1),'const.pair.ABS_prod_1'(s(fun(X0,fun(X1,bool)),X2))))) | ! [X5,X6] : ((~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X6)))) | (s(X0,X5) = s(X0,sK78(X2,X1,X0)) & s(X1,X6) = s(X1,sK79(X2,X1,X0)))) & (s(X0,X5) != s(X0,sK78(X2,X1,X0)) | s(X1,X6) != s(X1,sK79(X2,X1,X0)) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X6))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK76,sK77,sK78,sK79])],[f1016])).
fof(f1016,plain,(
  ! [X0,X1,X2] : ((! [X7,X8] : ? [X9,X10] : ((p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X9))),s(X1,X10)))) | (s(X0,X7) = s(X0,X9) & s(X1,X8) = s(X1,X10))) & (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X9))),s(X1,X10)))) | s(X0,X7) != s(X0,X9) | s(X1,X8) != s(X1,X10))) | s(fun(X0,fun(X1,bool)),X2) = s(fun(X0,fun(X1,bool)),'const.pair.REP_prod_1'(s('type.pair.prod'(X0,X1),'const.pair.ABS_prod_1'(s(fun(X0,fun(X1,bool)),X2)))))) & (s(fun(X0,fun(X1,bool)),X2) != s(fun(X0,fun(X1,bool)),'const.pair.REP_prod_1'(s('type.pair.prod'(X0,X1),'const.pair.ABS_prod_1'(s(fun(X0,fun(X1,bool)),X2))))) | ? [X3,X4] : ! [X5,X6] : ((~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X6)))) | (s(X0,X3) = s(X0,X5) & s(X1,X4) = s(X1,X6))) & (s(X0,X3) != s(X0,X5) | s(X1,X4) != s(X1,X6) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X6))))))))),
  inference(rectify,[],[f1015])).
fof(f1015,plain,(
  ! [X0,X1,X2] : ((! [X3,X4] : ? [X5,X6] : ((p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X6)))) | (s(X0,X3) = s(X0,X5) & s(X1,X4) = s(X1,X6))) & (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X6)))) | s(X0,X3) != s(X0,X5) | s(X1,X4) != s(X1,X6))) | s(fun(X0,fun(X1,bool)),X2) = s(fun(X0,fun(X1,bool)),'const.pair.REP_prod_1'(s('type.pair.prod'(X0,X1),'const.pair.ABS_prod_1'(s(fun(X0,fun(X1,bool)),X2)))))) & (s(fun(X0,fun(X1,bool)),X2) != s(fun(X0,fun(X1,bool)),'const.pair.REP_prod_1'(s('type.pair.prod'(X0,X1),'const.pair.ABS_prod_1'(s(fun(X0,fun(X1,bool)),X2))))) | ? [X3,X4] : ! [X5,X6] : ((~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X6)))) | (s(X0,X3) = s(X0,X5) & s(X1,X4) = s(X1,X6))) & (s(X0,X3) != s(X0,X5) | s(X1,X4) != s(X1,X6) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X6))))))))),
  inference(flattening,[],[f1014])).
fof(f1014,plain,(
  ! [X0,X1,X2] : ((! [X3,X4] : ? [X5,X6] : ((p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X6)))) | (s(X0,X3) = s(X0,X5) & s(X1,X4) = s(X1,X6))) & (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X6)))) | (s(X0,X3) != s(X0,X5) | s(X1,X4) != s(X1,X6)))) | s(fun(X0,fun(X1,bool)),X2) = s(fun(X0,fun(X1,bool)),'const.pair.REP_prod_1'(s('type.pair.prod'(X0,X1),'const.pair.ABS_prod_1'(s(fun(X0,fun(X1,bool)),X2)))))) & (s(fun(X0,fun(X1,bool)),X2) != s(fun(X0,fun(X1,bool)),'const.pair.REP_prod_1'(s('type.pair.prod'(X0,X1),'const.pair.ABS_prod_1'(s(fun(X0,fun(X1,bool)),X2))))) | ? [X3,X4] : ! [X5,X6] : ((~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X6)))) | (s(X0,X3) = s(X0,X5) & s(X1,X4) = s(X1,X6))) & ((s(X0,X3) != s(X0,X5) | s(X1,X4) != s(X1,X6)) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X6))))))))),
  inference(nnf_transformation,[],[f531])).
fof(f531,plain,(
  ! [X0,X1,X2] : (? [X3,X4] : ! [X5,X6] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X6)))) <=> (s(X0,X3) = s(X0,X5) & s(X1,X4) = s(X1,X6))) <=> s(fun(X0,fun(X1,bool)),X2) = s(fun(X0,fun(X1,bool)),'const.pair.REP_prod_1'(s('type.pair.prod'(X0,X1),'const.pair.ABS_prod_1'(s(fun(X0,fun(X1,bool)),X2))))))),
  inference(flattening,[],[f530])).
fof(f530,plain,(
  ! [X0,X1] : ! [X2] : (? [X3,X4] : ! [X5,X6] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X6)))) <=> (s(X0,X3) = s(X0,X5) & s(X1,X4) = s(X1,X6))) <=> s(fun(X0,fun(X1,bool)),X2) = s(fun(X0,fun(X1,bool)),'const.pair.REP_prod_1'(s('type.pair.prod'(X0,X1),'const.pair.ABS_prod_1'(s(fun(X0,fun(X1,bool)),X2))))))),
  inference(rectify,[],[f106])).
fof(f106,axiom,(
  ! [X6,X20] : ! [X33] : (? [X9,X19] : ! [X5,X13] : (p(s(bool,happ(s(fun(X20,bool),happ(s(fun(X6,fun(X20,bool)),X33),s(X6,X5))),s(X20,X13)))) <=> (s(X6,X5) = s(X6,X9) & s(X20,X13) = s(X20,X19))) <=> s(fun(X6,fun(X20,bool)),X33) = s(fun(X6,fun(X20,bool)),'const.pair.REP_prod_1'(s('type.pair.prod'(X6,X20),'const.pair.ABS_prod_1'(s(fun(X6,fun(X20,bool)),X33))))))),
  file('thm.pred_set.EXTENSION','thm.pair.ABS_REP_prod|split|1')).
fof(f1862,plain,(
  ( ! [X2,X0,X8,X7,X1] : (p(s(bool,'const.bool.IN_2'(s(X1,sK77(X8,X7,X2,X1,X0)),s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,sK76(X8,X7,X2,X1,X0))))))) | s(fun(X0,fun(X1,bool)),X2) = s(fun(X0,fun(X1,bool)),'const.pair.REP_prod_1'(s('type.pair.prod'(X0,X1),'const.pair.ABS_prod_1'(s(fun(X0,fun(X1,bool)),X2))))) | s(X1,X8) = s(X1,sK77(X8,X7,X2,X1,X0))) )),
  inference(forward_demodulation,[],[f1442,f1257])).
fof(f1442,plain,(
  ( ! [X2,X0,X8,X7,X1] : (s(fun(X0,fun(X1,bool)),X2) = s(fun(X0,fun(X1,bool)),'const.pair.REP_prod_1'(s('type.pair.prod'(X0,X1),'const.pair.ABS_prod_1'(s(fun(X0,fun(X1,bool)),X2))))) | s(X1,X8) = s(X1,sK77(X8,X7,X2,X1,X0)) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,sK76(X8,X7,X2,X1,X0)))),s(X1,sK77(X8,X7,X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f1017])).
fof(f1861,plain,(
  ( ! [X2,X0,X8,X7,X1] : (~p(s(bool,'const.bool.IN_2'(s(X1,sK77(X8,X7,X2,X1,X0)),s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,sK76(X8,X7,X2,X1,X0))))))) | s(fun(X0,fun(X1,bool)),X2) = s(fun(X0,fun(X1,bool)),'const.pair.REP_prod_1'(s('type.pair.prod'(X0,X1),'const.pair.ABS_prod_1'(s(fun(X0,fun(X1,bool)),X2))))) | s(X1,X8) != s(X1,sK77(X8,X7,X2,X1,X0)) | s(X0,X7) != s(X0,sK76(X8,X7,X2,X1,X0))) )),
  inference(forward_demodulation,[],[f1443,f1257])).
fof(f1443,plain,(
  ( ! [X2,X0,X8,X7,X1] : (s(fun(X0,fun(X1,bool)),X2) = s(fun(X0,fun(X1,bool)),'const.pair.REP_prod_1'(s('type.pair.prod'(X0,X1),'const.pair.ABS_prod_1'(s(fun(X0,fun(X1,bool)),X2))))) | s(X1,X8) != s(X1,sK77(X8,X7,X2,X1,X0)) | s(X0,X7) != s(X0,sK76(X8,X7,X2,X1,X0)) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,sK76(X8,X7,X2,X1,X0)))),s(X1,sK77(X8,X7,X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f1017])).
fof(f1860,plain,(
  ( ! [X6,X2,X0,X5,X1] : (s(fun(X0,fun(X1,bool)),X2) != s(fun(X0,fun(X1,bool)),'const.pair.REP_prod_1'(s('type.pair.prod'(X0,X1),'const.pair.ABS_prod_1'(s(fun(X0,fun(X1,bool)),X2))))) | s(X0,X5) = s(X0,sK78(X2,X1,X0)) | ~p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))))))) )),
  inference(forward_demodulation,[],[f1444,f1257])).
fof(f1444,plain,(
  ( ! [X6,X2,X0,X5,X1] : (s(X0,X5) = s(X0,sK78(X2,X1,X0)) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X6)))) | s(fun(X0,fun(X1,bool)),X2) != s(fun(X0,fun(X1,bool)),'const.pair.REP_prod_1'(s('type.pair.prod'(X0,X1),'const.pair.ABS_prod_1'(s(fun(X0,fun(X1,bool)),X2)))))) )),
  inference(cnf_transformation,[],[f1017])).
fof(f1859,plain,(
  ( ! [X6,X2,X0,X5,X1] : (s(fun(X0,fun(X1,bool)),X2) != s(fun(X0,fun(X1,bool)),'const.pair.REP_prod_1'(s('type.pair.prod'(X0,X1),'const.pair.ABS_prod_1'(s(fun(X0,fun(X1,bool)),X2))))) | s(X1,X6) = s(X1,sK79(X2,X1,X0)) | ~p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))))))) )),
  inference(forward_demodulation,[],[f1445,f1257])).
fof(f1445,plain,(
  ( ! [X6,X2,X0,X5,X1] : (s(X1,X6) = s(X1,sK79(X2,X1,X0)) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X6)))) | s(fun(X0,fun(X1,bool)),X2) != s(fun(X0,fun(X1,bool)),'const.pair.REP_prod_1'(s('type.pair.prod'(X0,X1),'const.pair.ABS_prod_1'(s(fun(X0,fun(X1,bool)),X2)))))) )),
  inference(cnf_transformation,[],[f1017])).
fof(f1858,plain,(
  ( ! [X6,X2,X0,X5,X1] : (s(fun(X0,fun(X1,bool)),X2) != s(fun(X0,fun(X1,bool)),'const.pair.REP_prod_1'(s('type.pair.prod'(X0,X1),'const.pair.ABS_prod_1'(s(fun(X0,fun(X1,bool)),X2))))) | s(X1,X6) != s(X1,sK79(X2,X1,X0)) | s(X0,X5) != s(X0,sK78(X2,X1,X0)) | p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))))))) )),
  inference(forward_demodulation,[],[f1446,f1257])).
fof(f1446,plain,(
  ( ! [X6,X2,X0,X5,X1] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X6)))) | s(X1,X6) != s(X1,sK79(X2,X1,X0)) | s(X0,X5) != s(X0,sK78(X2,X1,X0)) | s(fun(X0,fun(X1,bool)),X2) != s(fun(X0,fun(X1,bool)),'const.pair.REP_prod_1'(s('type.pair.prod'(X0,X1),'const.pair.ABS_prod_1'(s(fun(X0,fun(X1,bool)),X2)))))) )),
  inference(cnf_transformation,[],[f1017])).
fof(f1857,plain,(
  ( ! [X2,X0,X10,X8,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X10),s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,sK72(X2,X1,X0))))))) | p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),sK73(X2,X1,X0)),s(X1,X8))),s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X8))))))) )),
  inference(forward_demodulation,[],[f1856,f1257])).
fof(f1856,plain,(
  ( ! [X2,X0,X10,X8,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),sK73(X2,X1,X0)),s(X1,X8))),s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X8)))))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,sK72(X2,X1,X0)))),s(X0,X10))))) )),
  inference(forward_demodulation,[],[f1439,f1257])).
fof(f1439,plain,(
  ( ! [X2,X0,X10,X8,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X8))),s(X0,happ(s(fun(X1,X0),sK73(X2,X1,X0)),s(X1,X8)))))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,sK72(X2,X1,X0)))),s(X0,X10))))) )),
  inference(cnf_transformation,[],[f1013])).
fof(f1013,plain,(
  ! [X0,X1,X2] : ((! [X10] : ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,sK72(X2,X1,X0)))),s(X0,X10)))) | ! [X8] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X8))),s(X0,happ(s(fun(X1,X0),sK73(X2,X1,X0)),s(X1,X8))))))) & (! [X5] : ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,sK74(X5,X2,X1,X0)))),s(X0,happ(s(fun(X1,X0),X5),s(X1,sK74(X5,X2,X1,X0))))))) | ! [X3] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X3))),s(X0,sK75(X3,X2,X1,X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK72,sK73,sK74,sK75])],[f1012])).
fof(f1012,plain,(
  ! [X0,X1,X2] : ((? [X9] : ! [X10] : ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X9))),s(X0,X10)))) | ? [X7] : ! [X8] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X8))),s(X0,happ(s(fun(X1,X0),X7),s(X1,X8))))))) & (! [X5] : ? [X6] : ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X6))),s(X0,happ(s(fun(X1,X0),X5),s(X1,X6)))))) | ! [X3] : ? [X4] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X3))),s(X0,X4))))))),
  inference(rectify,[],[f1011])).
fof(f1011,plain,(
  ! [X0,X1,X2] : ((? [X3] : ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X3))),s(X0,X4)))) | ? [X5] : ! [X6] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X6))),s(X0,happ(s(fun(X1,X0),X5),s(X1,X6))))))) & (! [X5] : ? [X6] : ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X6))),s(X0,happ(s(fun(X1,X0),X5),s(X1,X6)))))) | ! [X3] : ? [X4] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X3))),s(X0,X4))))))),
  inference(nnf_transformation,[],[f529])).
fof(f529,plain,(
  ! [X0,X1,X2] : (! [X3] : ? [X4] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X3))),s(X0,X4)))) <=> ? [X5] : ! [X6] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X6))),s(X0,happ(s(fun(X1,X0),X5),s(X1,X6)))))))),
  inference(flattening,[],[f528])).
fof(f528,plain,(
  ! [X0,X1] : ! [X2] : (! [X3] : ? [X4] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X3))),s(X0,X4)))) <=> ? [X5] : ! [X6] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X6))),s(X0,happ(s(fun(X1,X0),X5),s(X1,X6)))))))),
  inference(rectify,[],[f98])).
fof(f98,axiom,(
  ! [X20,X6] : ! [X17] : (! [X9] : ? [X19] : p(s(bool,happ(s(fun(X20,bool),happ(s(fun(X6,fun(X20,bool)),X17),s(X6,X9))),s(X20,X19)))) <=> ? [X15] : ! [X9] : p(s(bool,happ(s(fun(X20,bool),happ(s(fun(X6,fun(X20,bool)),X17),s(X6,X9))),s(X20,happ(s(fun(X6,X20),X15),s(X6,X9)))))))),
  file('thm.pred_set.EXTENSION','thm.bool.SKOLEM_THM')).
fof(f1855,plain,(
  ( ! [X2,X0,X5,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X5),s(X1,sK74(X5,X2,X1,X0)))),s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,sK74(X5,X2,X1,X0))))))) | p(s(bool,'const.bool.IN_2'(s(X0,sK75(X3,X2,X1,X0)),s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X3))))))) )),
  inference(forward_demodulation,[],[f1854,f1257])).
fof(f1854,plain,(
  ( ! [X2,X0,X5,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK75(X3,X2,X1,X0)),s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X3)))))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,sK74(X5,X2,X1,X0)))),s(X0,happ(s(fun(X1,X0),X5),s(X1,sK74(X5,X2,X1,X0)))))))) )),
  inference(forward_demodulation,[],[f1440,f1257])).
fof(f1440,plain,(
  ( ! [X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X3))),s(X0,sK75(X3,X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,sK74(X5,X2,X1,X0)))),s(X0,happ(s(fun(X1,X0),X5),s(X1,sK74(X5,X2,X1,X0)))))))) )),
  inference(cnf_transformation,[],[f1013])).
fof(f1437,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),'const.relation.RUNIV_0'),s(fun(X0,fun(X1,bool)),X2)))) | s(fun(X0,fun(X1,bool)),'const.relation.RUNIV_0') = s(fun(X0,fun(X1,bool)),X2)) )),
  inference(cnf_transformation,[],[f1010])).
fof(f1010,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),'const.relation.RUNIV_0'),s(fun(X0,fun(X1,bool)),X2)))) | s(fun(X0,fun(X1,bool)),'const.relation.RUNIV_0') = s(fun(X0,fun(X1,bool)),X2)) & (s(fun(X0,fun(X1,bool)),'const.relation.RUNIV_0') != s(fun(X0,fun(X1,bool)),X2) | p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),'const.relation.RUNIV_0'),s(fun(X0,fun(X1,bool)),X2))))))),
  inference(nnf_transformation,[],[f527])).
fof(f527,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),'const.relation.RUNIV_0'),s(fun(X0,fun(X1,bool)),X2)))) <=> s(fun(X0,fun(X1,bool)),'const.relation.RUNIV_0') = s(fun(X0,fun(X1,bool)),X2))),
  inference(flattening,[],[f526])).
fof(f526,plain,(
  ! [X0,X1] : ! [X2] : (p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),'const.relation.RUNIV_0'),s(fun(X0,fun(X1,bool)),X2)))) <=> s(fun(X0,fun(X1,bool)),'const.relation.RUNIV_0') = s(fun(X0,fun(X1,bool)),X2))),
  inference(rectify,[],[f182])).
fof(f182,axiom,(
  ! [X6,X20] : ! [X18] : (p(s(bool,'const.relation.RSUBSET_2'(s(fun(X6,fun(X20,bool)),'const.relation.RUNIV_0'),s(fun(X6,fun(X20,bool)),X18)))) <=> s(fun(X6,fun(X20,bool)),'const.relation.RUNIV_0') = s(fun(X6,fun(X20,bool)),X18))),
  file('thm.pred_set.EXTENSION','thm.relation.RUNIV_SUBSET|split|0')).
fof(f1438,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),'const.relation.RUNIV_0'),s(fun(X0,fun(X1,bool)),X2)))) | s(fun(X0,fun(X1,bool)),'const.relation.RUNIV_0') != s(fun(X0,fun(X1,bool)),X2)) )),
  inference(cnf_transformation,[],[f1010])).
fof(f1435,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.sum.IS_SUM_REP_1'(s(fun(bool,fun(X0,fun(X1,bool))),X2)))) | s(fun(bool,fun(X0,fun(X1,bool))),X2) = s(fun(bool,fun(X0,fun(X1,bool))),'const.sum.REP_sum_1'(s('type.sum.sum'(X0,X1),'const.sum.ABS_sum_1'(s(fun(bool,fun(X0,fun(X1,bool))),X2)))))) )),
  inference(cnf_transformation,[],[f1009])).
fof(f1009,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.sum.IS_SUM_REP_1'(s(fun(bool,fun(X0,fun(X1,bool))),X2)))) | s(fun(bool,fun(X0,fun(X1,bool))),X2) = s(fun(bool,fun(X0,fun(X1,bool))),'const.sum.REP_sum_1'(s('type.sum.sum'(X0,X1),'const.sum.ABS_sum_1'(s(fun(bool,fun(X0,fun(X1,bool))),X2)))))) & (s(fun(bool,fun(X0,fun(X1,bool))),X2) != s(fun(bool,fun(X0,fun(X1,bool))),'const.sum.REP_sum_1'(s('type.sum.sum'(X0,X1),'const.sum.ABS_sum_1'(s(fun(bool,fun(X0,fun(X1,bool))),X2))))) | p(s(bool,'const.sum.IS_SUM_REP_1'(s(fun(bool,fun(X0,fun(X1,bool))),X2))))))),
  inference(nnf_transformation,[],[f525])).
fof(f525,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.sum.IS_SUM_REP_1'(s(fun(bool,fun(X0,fun(X1,bool))),X2)))) <=> s(fun(bool,fun(X0,fun(X1,bool))),X2) = s(fun(bool,fun(X0,fun(X1,bool))),'const.sum.REP_sum_1'(s('type.sum.sum'(X0,X1),'const.sum.ABS_sum_1'(s(fun(bool,fun(X0,fun(X1,bool))),X2))))))),
  inference(flattening,[],[f524])).
fof(f524,plain,(
  ! [X0,X1] : ! [X2] : (p(s(bool,'const.sum.IS_SUM_REP_1'(s(fun(bool,fun(X0,fun(X1,bool))),X2)))) <=> s(fun(bool,fun(X0,fun(X1,bool))),X2) = s(fun(bool,fun(X0,fun(X1,bool))),'const.sum.REP_sum_1'(s('type.sum.sum'(X0,X1),'const.sum.ABS_sum_1'(s(fun(bool,fun(X0,fun(X1,bool))),X2))))))),
  inference(rectify,[],[f224])).
fof(f224,axiom,(
  ! [X6,X20] : ! [X33] : (p(s(bool,'const.sum.IS_SUM_REP_1'(s(fun(bool,fun(X6,fun(X20,bool))),X33)))) <=> s(fun(bool,fun(X6,fun(X20,bool))),X33) = s(fun(bool,fun(X6,fun(X20,bool))),'const.sum.REP_sum_1'(s('type.sum.sum'(X6,X20),'const.sum.ABS_sum_1'(s(fun(bool,fun(X6,fun(X20,bool))),X33))))))),
  file('thm.pred_set.EXTENSION','thm.sum.sum_ISO_DEF|split|1')).
fof(f1436,plain,(
  ( ! [X2,X0,X1] : (s(fun(bool,fun(X0,fun(X1,bool))),X2) != s(fun(bool,fun(X0,fun(X1,bool))),'const.sum.REP_sum_1'(s('type.sum.sum'(X0,X1),'const.sum.ABS_sum_1'(s(fun(bool,fun(X0,fun(X1,bool))),X2))))) | p(s(bool,'const.sum.IS_SUM_REP_1'(s(fun(bool,fun(X0,fun(X1,bool))),X2))))) )),
  inference(cnf_transformation,[],[f1009])).
fof(f1432,plain,(
  ( ! [X6,X2,X0,X7,X1] : (~p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun(X1,bool),X2),s(fun(X0,X1),X7)))) | s(X0,X6) = s(X0,happ(s(fun(X1,X0),sK71(X2,X1,X0)),s(X1,happ(s(fun(X0,X1),sK70(X2,X1,X0)),s(X0,X6)))))) )),
  inference(cnf_transformation,[],[f1008])).
fof(f1008,plain,(
  ! [X0,X1,X2] : (! [X7] : ~p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun(X1,bool),X2),s(fun(X0,X1),X7)))) | (! [X6] : s(X0,X6) = s(X0,happ(s(fun(X1,X0),sK71(X2,X1,X0)),s(X1,happ(s(fun(X0,X1),sK70(X2,X1,X0)),s(X0,X6))))) & ! [X5] : ((~p(s(bool,happ(s(fun(X1,bool),X2),s(X1,X5)))) | s(X1,X5) = s(X1,happ(s(fun(X0,X1),sK70(X2,X1,X0)),s(X0,happ(s(fun(X1,X0),sK71(X2,X1,X0)),s(X1,X5)))))) & (s(X1,X5) != s(X1,happ(s(fun(X0,X1),sK70(X2,X1,X0)),s(X0,happ(s(fun(X1,X0),sK71(X2,X1,X0)),s(X1,X5))))) | p(s(bool,happ(s(fun(X1,bool),X2),s(X1,X5))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK70,sK71])],[f1007])).
fof(f1007,plain,(
  ! [X0,X1,X2] : (! [X7] : ~p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun(X1,bool),X2),s(fun(X0,X1),X7)))) | ? [X3,X4] : (! [X6] : s(X0,X6) = s(X0,happ(s(fun(X1,X0),X4),s(X1,happ(s(fun(X0,X1),X3),s(X0,X6))))) & ! [X5] : ((~p(s(bool,happ(s(fun(X1,bool),X2),s(X1,X5)))) | s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X4),s(X1,X5)))))) & (s(X1,X5) != s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))))) | p(s(bool,happ(s(fun(X1,bool),X2),s(X1,X5))))))))),
  inference(rectify,[],[f1006])).
fof(f1006,plain,(
  ! [X0,X1,X2] : (! [X3] : ~p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun(X1,bool),X2),s(fun(X0,X1),X3)))) | ? [X4,X5] : (! [X7] : s(X0,X7) = s(X0,happ(s(fun(X1,X0),X5),s(X1,happ(s(fun(X0,X1),X4),s(X0,X7))))) & ! [X6] : ((~p(s(bool,happ(s(fun(X1,bool),X2),s(X1,X6)))) | s(X1,X6) = s(X1,happ(s(fun(X0,X1),X4),s(X0,happ(s(fun(X1,X0),X5),s(X1,X6)))))) & (s(X1,X6) != s(X1,happ(s(fun(X0,X1),X4),s(X0,happ(s(fun(X1,X0),X5),s(X1,X6))))) | p(s(bool,happ(s(fun(X1,bool),X2),s(X1,X6))))))))),
  inference(nnf_transformation,[],[f793])).
fof(f793,plain,(
  ! [X0,X1,X2] : (! [X3] : ~p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun(X1,bool),X2),s(fun(X0,X1),X3)))) | ? [X4,X5] : (! [X7] : s(X0,X7) = s(X0,happ(s(fun(X1,X0),X5),s(X1,happ(s(fun(X0,X1),X4),s(X0,X7))))) & ! [X6] : (p(s(bool,happ(s(fun(X1,bool),X2),s(X1,X6)))) <=> s(X1,X6) = s(X1,happ(s(fun(X0,X1),X4),s(X0,happ(s(fun(X1,X0),X5),s(X1,X6))))))))),
  inference(ennf_transformation,[],[f523])).
fof(f523,plain,(
  ! [X0,X1,X2] : (? [X3] : p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun(X1,bool),X2),s(fun(X0,X1),X3)))) => ? [X4,X5] : (! [X7] : s(X0,X7) = s(X0,happ(s(fun(X1,X0),X5),s(X1,happ(s(fun(X0,X1),X4),s(X0,X7))))) & ! [X6] : (p(s(bool,happ(s(fun(X1,bool),X2),s(X1,X6)))) <=> s(X1,X6) = s(X1,happ(s(fun(X0,X1),X4),s(X0,happ(s(fun(X1,X0),X5),s(X1,X6))))))))),
  inference(flattening,[],[f522])).
fof(f522,plain,(
  ! [X0,X1] : ! [X2] : (? [X3] : p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun(X1,bool),X2),s(fun(X0,X1),X3)))) => ? [X4,X5] : (! [X7] : s(X0,X7) = s(X0,happ(s(fun(X1,X0),X5),s(X1,happ(s(fun(X0,X1),X4),s(X0,X7))))) & ! [X6] : (p(s(bool,happ(s(fun(X1,bool),X2),s(X1,X6)))) <=> s(X1,X6) = s(X1,happ(s(fun(X0,X1),X4),s(X0,happ(s(fun(X1,X0),X5),s(X1,X6))))))))),
  inference(rectify,[],[f107])).
fof(f107,axiom,(
  ! [X20,X6] : ! [X17] : (? [X53] : p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun(X6,bool),X17),s(fun(X20,X6),X53)))) => ? [X53,X54] : (! [X55] : s(X20,X55) = s(X20,happ(s(fun(X6,X20),X54),s(X6,happ(s(fun(X20,X6),X53),s(X20,X55))))) & ! [X33] : (p(s(bool,happ(s(fun(X6,bool),X17),s(X6,X33)))) <=> s(X6,X33) = s(X6,happ(s(fun(X20,X6),X53),s(X20,happ(s(fun(X6,X20),X54),s(X6,X33))))))))),
  file('thm.pred_set.EXTENSION','thm.bool.ABS_REP_THM')).
fof(f1853,plain,(
  ( ! [X2,X0,X7,X5,X1] : (~p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun(X1,bool),X2),s(fun(X0,X1),X7)))) | s(X1,X5) = s(X1,happ(s(fun(X0,X1),sK70(X2,X1,X0)),s(X0,happ(s(fun(X1,X0),sK71(X2,X1,X0)),s(X1,X5))))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X2))))) )),
  inference(forward_demodulation,[],[f1433,f1257])).
fof(f1433,plain,(
  ( ! [X2,X0,X7,X5,X1] : (s(X1,X5) = s(X1,happ(s(fun(X0,X1),sK70(X2,X1,X0)),s(X0,happ(s(fun(X1,X0),sK71(X2,X1,X0)),s(X1,X5))))) | ~p(s(bool,happ(s(fun(X1,bool),X2),s(X1,X5)))) | ~p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun(X1,bool),X2),s(fun(X0,X1),X7))))) )),
  inference(cnf_transformation,[],[f1008])).
fof(f1852,plain,(
  ( ! [X2,X0,X7,X5,X1] : (~p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun(X1,bool),X2),s(fun(X0,X1),X7)))) | s(X1,X5) != s(X1,happ(s(fun(X0,X1),sK70(X2,X1,X0)),s(X0,happ(s(fun(X1,X0),sK71(X2,X1,X0)),s(X1,X5))))) | p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X2))))) )),
  inference(forward_demodulation,[],[f1434,f1257])).
fof(f1434,plain,(
  ( ! [X2,X0,X7,X5,X1] : (p(s(bool,happ(s(fun(X1,bool),X2),s(X1,X5)))) | s(X1,X5) != s(X1,happ(s(fun(X0,X1),sK70(X2,X1,X0)),s(X0,happ(s(fun(X1,X0),sK71(X2,X1,X0)),s(X1,X5))))) | ~p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun(X1,bool),X2),s(fun(X0,X1),X7))))) )),
  inference(cnf_transformation,[],[f1008])).
fof(f1431,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.sum.IS_SUM_REP_1'(s(fun(bool,fun(X0,fun(X1,bool))),X2))) = s(bool,happ(s(fun(fun(bool,fun(X0,fun(X1,bool))),bool),'const.sum.IS_SUM_REP_0'),s(fun(bool,fun(X0,fun(X1,bool))),X2)))) )),
  inference(cnf_transformation,[],[f521])).
fof(f521,plain,(
  ! [X0,X1,X2] : s(bool,'const.sum.IS_SUM_REP_1'(s(fun(bool,fun(X0,fun(X1,bool))),X2))) = s(bool,happ(s(fun(fun(bool,fun(X0,fun(X1,bool))),bool),'const.sum.IS_SUM_REP_0'),s(fun(bool,fun(X0,fun(X1,bool))),X2)))),
  inference(flattening,[],[f520])).
fof(f520,plain,(
  ! [X0,X1] : ! [X2] : s(bool,'const.sum.IS_SUM_REP_1'(s(fun(bool,fun(X0,fun(X1,bool))),X2))) = s(bool,happ(s(fun(fun(bool,fun(X0,fun(X1,bool))),bool),'const.sum.IS_SUM_REP_0'),s(fun(bool,fun(X0,fun(X1,bool))),X2)))),
  inference(rectify,[],[f226])).
fof(f226,axiom,(
  ! [X6,X20] : ! [X21] : s(bool,'const.sum.IS_SUM_REP_1'(s(fun(bool,fun(X6,fun(X20,bool))),X21))) = s(bool,happ(s(fun(fun(bool,fun(X6,fun(X20,bool))),bool),'const.sum.IS_SUM_REP_0'),s(fun(bool,fun(X6,fun(X20,bool))),X21)))),
  file('thm.pred_set.EXTENSION','ARITY_0_1_const.sum.IS_SUM_REP')).
fof(f1430,plain,(
  ( ! [X2,X0,X1] : (s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) = s(X0,happ(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.pair.prod'(X0,X1),X2)))) )),
  inference(cnf_transformation,[],[f519])).
fof(f519,plain,(
  ! [X0,X1,X2] : s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) = s(X0,happ(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.pair.prod'(X0,X1),X2)))),
  inference(flattening,[],[f518])).
fof(f518,plain,(
  ! [X0,X1] : ! [X2] : s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) = s(X0,happ(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.pair.prod'(X0,X1),X2)))),
  inference(rectify,[],[f199])).
fof(f199,axiom,(
  ! [X6,X20] : ! [X21] : s(X6,'const.pair.FST_1'(s('type.pair.prod'(X6,X20),X21))) = s(X6,happ(s(fun('type.pair.prod'(X6,X20),X6),'const.pair.FST_0'),s('type.pair.prod'(X6,X20),X21)))),
  file('thm.pred_set.EXTENSION','ARITY_0_1_const.pair.FST')).
fof(f1429,plain,(
  ( ! [X2,X0,X1] : (s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X2))) = s(X1,happ(s(fun('type.pair.prod'(X0,X1),X1),'const.pair.SND_0'),s('type.pair.prod'(X0,X1),X2)))) )),
  inference(cnf_transformation,[],[f517])).
fof(f517,plain,(
  ! [X0,X1,X2] : s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X2))) = s(X1,happ(s(fun('type.pair.prod'(X0,X1),X1),'const.pair.SND_0'),s('type.pair.prod'(X0,X1),X2)))),
  inference(flattening,[],[f516])).
fof(f516,plain,(
  ! [X0,X1] : ! [X2] : s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X2))) = s(X1,happ(s(fun('type.pair.prod'(X0,X1),X1),'const.pair.SND_0'),s('type.pair.prod'(X0,X1),X2)))),
  inference(rectify,[],[f200])).
fof(f200,axiom,(
  ! [X6,X20] : ! [X21] : s(X20,'const.pair.SND_1'(s('type.pair.prod'(X6,X20),X21))) = s(X20,happ(s(fun('type.pair.prod'(X6,X20),X20),'const.pair.SND_0'),s('type.pair.prod'(X6,X20),X21)))),
  file('thm.pred_set.EXTENSION','ARITY_0_1_const.pair.SND')).
fof(f1850,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,X1),s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X0)))))) | p(s(bool,X0))) )),
  inference(forward_demodulation,[],[f1389,f1257])).
fof(f1389,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X0))),s(bool,X1))))) )),
  inference(cnf_transformation,[],[f989])).
fof(f989,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X0))),s(bool,X1)))) | (p(s(bool,X0)) & p(s(bool,X1)))) & (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X0))),s(bool,X1))))))),
  inference(flattening,[],[f988])).
fof(f988,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X0))),s(bool,X1)))) | (p(s(bool,X0)) & p(s(bool,X1)))) & ((~p(s(bool,X0)) | ~p(s(bool,X1))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X0))),s(bool,X1))))))),
  inference(nnf_transformation,[],[f504])).
fof(f504,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X0))),s(bool,X1)))) <=> (p(s(bool,X0)) & p(s(bool,X1))))),
  inference(rectify,[],[f69])).
fof(f69,axiom,(
  ! [X44,X45] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X44))),s(bool,X45)))) <=> (p(s(bool,X44)) & p(s(bool,X45))))),
  file('thm.pred_set.EXTENSION','|slash||bslash|_HL_0')).
fof(f1849,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,X1),s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X0)))))) | p(s(bool,X1))) )),
  inference(forward_demodulation,[],[f1390,f1257])).
fof(f1390,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X0))),s(bool,X1))))) )),
  inference(cnf_transformation,[],[f989])).
fof(f1848,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,'const.bool.IN_2'(s(bool,X1),s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X0))))))) )),
  inference(forward_demodulation,[],[f1391,f1257])).
fof(f1391,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X0))),s(bool,X1)))) | ~p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f989])).
fof(f1847,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X1),s(fun(fun(X0,bool),bool),'?_1')))) | p(s(bool,'const.bool.IN_2'(s(X0,sK69(X1,X0)),s(fun(X0,bool),X1))))) )),
  inference(forward_demodulation,[],[f1846,f1257])).
fof(f1846,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK69(X1,X0)),s(fun(X0,bool),X1)))) | ~p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1))))) )),
  inference(forward_demodulation,[],[f1387,f1257])).
fof(f1387,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK69(X1,X0))))) | ~p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f987])).
fof(f1843,plain,(
  ( ! [X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X1),s(fun(fun(X0,bool),bool),'!_1')))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))))) )),
  inference(forward_demodulation,[],[f1842,f1257])).
fof(f1842,plain,(
  ( ! [X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1))))) )),
  inference(forward_demodulation,[],[f1385,f1257])).
fof(f1385,plain,(
  ( ! [X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f984])).
fof(f1839,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,'const.bool.IN_2'(s(bool,X1),s(fun(bool,bool),'==>_1'(s(bool,X0))))))) )),
  inference(forward_demodulation,[],[f1382,f1257])).
fof(f1382,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,happ(s(fun(bool,bool),'==>_1'(s(bool,X0))),s(bool,X1))))) )),
  inference(cnf_transformation,[],[f981])).
fof(f981,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(bool,bool),'==>_1'(s(bool,X0))),s(bool,X1)))) | ~p(s(bool,X0)) | p(s(bool,X1))) & ((p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,happ(s(fun(bool,bool),'==>_1'(s(bool,X0))),s(bool,X1))))))),
  inference(flattening,[],[f980])).
fof(f980,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(bool,bool),'==>_1'(s(bool,X0))),s(bool,X1)))) | (~p(s(bool,X0)) | p(s(bool,X1)))) & ((p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,happ(s(fun(bool,bool),'==>_1'(s(bool,X0))),s(bool,X1))))))),
  inference(nnf_transformation,[],[f786])).
fof(f786,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),'==>_1'(s(bool,X0))),s(bool,X1)))) <=> (~p(s(bool,X0)) | p(s(bool,X1))))),
  inference(ennf_transformation,[],[f501])).
fof(f501,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),'==>_1'(s(bool,X0))),s(bool,X1)))) <=> (p(s(bool,X0)) => p(s(bool,X1))))),
  inference(rectify,[],[f75])).
fof(f75,axiom,(
  ! [X44,X45] : (p(s(bool,happ(s(fun(bool,bool),'==>_1'(s(bool,X44))),s(bool,X45)))) <=> (p(s(bool,X44)) => p(s(bool,X45))))),
  file('thm.pred_set.EXTENSION','==>_HL_1')).
fof(f1838,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(bool,X1),s(fun(bool,bool),'==>_1'(s(bool,X0)))))) | p(s(bool,X0))) )),
  inference(forward_demodulation,[],[f1383,f1257])).
fof(f1383,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),'==>_1'(s(bool,X0))),s(bool,X1)))) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f981])).
fof(f1837,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,'const.bool.IN_2'(s(bool,X1),s(fun(bool,bool),'==>_1'(s(bool,X0))))))) )),
  inference(forward_demodulation,[],[f1384,f1257])).
fof(f1384,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),'==>_1'(s(bool,X0))),s(bool,X1)))) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f981])).
fof(f1679,plain,(
  ( ! [X0,X1] : (~p(s(bool,'|dollar|or_2'(s(bool,X1),s(bool,X0)))) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(condensation,[],[f1375])).
fof(f1375,plain,(
  ( ! [X0,X3,X1] : (p(s(bool,X3)) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,'|dollar|or_2'(s(bool,X0),s(bool,X1))))) )),
  inference(cnf_transformation,[],[f979])).
fof(f979,plain,(
  ! [X0,X1] : ((~p(s(bool,'|dollar|or_2'(s(bool,X0),s(bool,X1)))) | ! [X3] : ((p(s(bool,X0)) & ~p(s(bool,X3))) | (p(s(bool,X1)) & ~p(s(bool,X3))) | p(s(bool,X3)))) & (((~p(s(bool,X0)) | p(s(bool,sK67(X1,X0)))) & (~p(s(bool,X1)) | p(s(bool,sK67(X1,X0)))) & ~p(s(bool,sK67(X1,X0)))) | p(s(bool,'|dollar|or_2'(s(bool,X0),s(bool,X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK67])],[f978])).
fof(f978,plain,(
  ! [X0,X1] : ((~p(s(bool,'|dollar|or_2'(s(bool,X0),s(bool,X1)))) | ! [X3] : ((p(s(bool,X0)) & ~p(s(bool,X3))) | (p(s(bool,X1)) & ~p(s(bool,X3))) | p(s(bool,X3)))) & (? [X2] : ((~p(s(bool,X0)) | p(s(bool,X2))) & (~p(s(bool,X1)) | p(s(bool,X2))) & ~p(s(bool,X2))) | p(s(bool,'|dollar|or_2'(s(bool,X0),s(bool,X1))))))),
  inference(rectify,[],[f977])).
fof(f977,plain,(
  ! [X0,X1] : ((~p(s(bool,'|dollar|or_2'(s(bool,X0),s(bool,X1)))) | ! [X2] : ((p(s(bool,X0)) & ~p(s(bool,X2))) | (p(s(bool,X1)) & ~p(s(bool,X2))) | p(s(bool,X2)))) & (? [X2] : ((~p(s(bool,X0)) | p(s(bool,X2))) & (~p(s(bool,X1)) | p(s(bool,X2))) & ~p(s(bool,X2))) | p(s(bool,'|dollar|or_2'(s(bool,X0),s(bool,X1))))))),
  inference(nnf_transformation,[],[f785])).
fof(f785,plain,(
  ! [X0,X1] : (p(s(bool,'|dollar|or_2'(s(bool,X0),s(bool,X1)))) <=> ! [X2] : ((p(s(bool,X0)) & ~p(s(bool,X2))) | (p(s(bool,X1)) & ~p(s(bool,X2))) | p(s(bool,X2))))),
  inference(flattening,[],[f784])).
fof(f784,plain,(
  ! [X0,X1] : (p(s(bool,'|dollar|or_2'(s(bool,X0),s(bool,X1)))) <=> ! [X2] : ((p(s(bool,X0)) & ~p(s(bool,X2))) | ((p(s(bool,X1)) & ~p(s(bool,X2))) | p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f500])).
fof(f500,plain,(
  ! [X0,X1] : (p(s(bool,'|dollar|or_2'(s(bool,X0),s(bool,X1)))) <=> ! [X2] : ((p(s(bool,X0)) => p(s(bool,X2))) => ((p(s(bool,X1)) => p(s(bool,X2))) => p(s(bool,X2)))))),
  inference(rectify,[],[f71])).
fof(f71,axiom,(
  ! [X5,X12] : (p(s(bool,'|dollar|or_2'(s(bool,X5),s(bool,X12)))) <=> ! [X8] : ((p(s(bool,X5)) => p(s(bool,X8))) => ((p(s(bool,X12)) => p(s(bool,X8))) => p(s(bool,X8)))))),
  file('thm.pred_set.EXTENSION','thm.bool.OR_DEF')).
fof(f1836,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,'|dollar|or_2'(s(bool,X0),s(bool,X1))))) )),
  inference(subsumption_resolution,[],[f1379,f1381])).
fof(f1379,plain,(
  ( ! [X0,X1] : (p(s(bool,'|dollar|or_2'(s(bool,X0),s(bool,X1)))) | p(s(bool,sK67(X1,X0))) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f979])).
fof(f1835,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,'|dollar|or_2'(s(bool,X0),s(bool,X1))))) )),
  inference(subsumption_resolution,[],[f1380,f1381])).
fof(f1380,plain,(
  ( ! [X0,X1] : (p(s(bool,'|dollar|or_2'(s(bool,X0),s(bool,X1)))) | p(s(bool,sK67(X1,X0))) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f979])).
fof(f1381,plain,(
  ( ! [X0,X1] : (~p(s(bool,sK67(X1,X0))) | p(s(bool,'|dollar|or_2'(s(bool,X0),s(bool,X1))))) )),
  inference(cnf_transformation,[],[f979])).
fof(f1370,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f976])).
fof(f976,plain,(
  ! [X0,X1] : ((s(bool,X0) != s(bool,X1) | ((~p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | p(s(bool,X0))))) & ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(flattening,[],[f975])).
fof(f975,plain,(
  ! [X0,X1] : ((s(bool,X0) != s(bool,X1) | ((~p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | p(s(bool,X0))))) & (((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0)))) | s(bool,X0) = s(bool,X1)))),
  inference(nnf_transformation,[],[f783])).
fof(f783,plain,(
  ! [X0,X1] : (s(bool,X0) = s(bool,X1) <=> ((~p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | p(s(bool,X0)))))),
  inference(ennf_transformation,[],[f499])).
fof(f499,plain,(
  ! [X0,X1] : (s(bool,X0) = s(bool,X1) <=> ((p(s(bool,X0)) => p(s(bool,X1))) & (p(s(bool,X1)) => p(s(bool,X0)))))),
  inference(rectify,[],[f95])).
fof(f95,axiom,(
  ! [X10,X11] : (s(bool,X10) = s(bool,X11) <=> ((p(s(bool,X10)) => p(s(bool,X11))) & (p(s(bool,X11)) => p(s(bool,X10)))))),
  file('thm.pred_set.EXTENSION','thm.bool.EQ_IMP_THM')).
fof(f1360,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),happ(s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.+_0'),s('type.num.num',X0))),s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f490])).
fof(f490,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),happ(s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.+_0'),s('type.num.num',X0))),s('type.num.num',X1)))),
  inference(rectify,[],[f193])).
fof(f193,axiom,(
  ! [X21,X22] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',X22))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),happ(s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.+_0'),s('type.num.num',X21))),s('type.num.num',X22)))),
  file('thm.pred_set.EXTENSION','ARITY_0_2_const.arithmetic.+')).
fof(f1359,plain,(
  ( ! [X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0'),s('type.num.num',X0))),s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f489])).
fof(f489,plain,(
  ! [X0,X1] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0'),s('type.num.num',X0))),s('type.num.num',X1)))),
  inference(rectify,[],[f168])).
fof(f168,axiom,(
  ! [X21,X22] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X21),s('type.num.num',X22))) = s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0'),s('type.num.num',X21))),s('type.num.num',X22)))),
  file('thm.pred_set.EXTENSION','ARITY_0_2_const.prim_rec.<')).
fof(f1358,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),happ(s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.-_0'),s('type.num.num',X0))),s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f488])).
fof(f488,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),happ(s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.-_0'),s('type.num.num',X0))),s('type.num.num',X1)))),
  inference(rectify,[],[f166])).
fof(f166,axiom,(
  ! [X21,X22] : s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X21),s('type.num.num',X22))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),happ(s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.-_0'),s('type.num.num',X21))),s('type.num.num',X22)))),
  file('thm.pred_set.EXTENSION','ARITY_0_2_const.arithmetic.-')).
fof(f1357,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),happ(s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.*_0'),s('type.num.num',X0))),s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f487])).
fof(f487,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),happ(s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.*_0'),s('type.num.num',X0))),s('type.num.num',X1)))),
  inference(rectify,[],[f192])).
fof(f192,axiom,(
  ! [X21,X22] : s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X21),s('type.num.num',X22))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),happ(s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.*_0'),s('type.num.num',X21))),s('type.num.num',X22)))),
  file('thm.pred_set.EXTENSION','ARITY_0_2_const.arithmetic.*')).
fof(f1356,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.*_1'(s('type.num.num',X0))),s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f486])).
fof(f486,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.*_1'(s('type.num.num',X0))),s('type.num.num',X1)))),
  inference(rectify,[],[f176])).
fof(f176,axiom,(
  ! [X21,X22] : s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X21),s('type.num.num',X22))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.*_1'(s('type.num.num',X21))),s('type.num.num',X22)))),
  file('thm.pred_set.EXTENSION','ARITY_1_2_const.arithmetic.*')).
fof(f1355,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.relation.WeakOrder_1'(s(fun(fun(X0,fun(X1,bool)),fun(fun(X0,fun(X1,bool)),bool)),'const.relation.RSUBSET_0'))))) )),
  inference(cnf_transformation,[],[f485])).
fof(f485,plain,(
  ! [X0,X1] : p(s(bool,'const.relation.WeakOrder_1'(s(fun(fun(X0,fun(X1,bool)),fun(fun(X0,fun(X1,bool)),bool)),'const.relation.RSUBSET_0'))))),
  inference(rectify,[],[f219])).
fof(f219,axiom,(
  ! [X6,X20] : p(s(bool,'const.relation.WeakOrder_1'(s(fun(fun(X6,fun(X20,bool)),fun(fun(X6,fun(X20,bool)),bool)),'const.relation.RSUBSET_0'))))),
  file('thm.pred_set.EXTENSION','thm.relation.RSUBSET_WeakOrder')).
fof(f1834,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2)))))) | p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4))))) )),
  inference(forward_demodulation,[],[f1354,f1257])).
fof(f1354,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f772])).
fof(f772,plain,(
  ! [X0,X1,X2,X3,X4] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4)))))),
  inference(flattening,[],[f771])).
fof(f771,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4))))) | p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4)))))),
  inference(ennf_transformation,[],[f484])).
fof(f484,plain,(
  ! [X0,X1,X2,X3,X4] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4))))) => p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4)))))),
  inference(flattening,[],[f483])).
fof(f483,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4))))) => p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4)))))),
  inference(rectify,[],[f257])).
fof(f257,axiom,(
  ! [X6] : ! [X18,X9,X19,X49] : ((p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X18),s(X6,X9))),s(X6,X19)))) & p(s(bool,'const.relation.RTC_3'(s(fun(X6,fun(X6,bool)),X18),s(X6,X19),s(X6,X49))))) => p(s(bool,'const.relation.RTC_3'(s(fun(X6,fun(X6,bool)),X18),s(X6,X9),s(X6,X49)))))),
  file('thm.pred_set.EXTENSION','thm.relation.RTC_RULES|split|1')).
fof(f1353,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4))))) )),
  inference(cnf_transformation,[],[f770])).
fof(f770,plain,(
  ! [X0,X1,X2,X3,X4] : (~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4)))))),
  inference(flattening,[],[f769])).
fof(f769,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4))))) | p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4)))))),
  inference(ennf_transformation,[],[f482])).
fof(f482,plain,(
  ! [X0,X1,X2,X3,X4] : ((p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4))))) => p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4)))))),
  inference(flattening,[],[f481])).
fof(f481,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : ((p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4))))) => p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4)))))),
  inference(rectify,[],[f214])).
fof(f214,axiom,(
  ! [X6] : ! [X18,X9,X19,X49] : ((p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X18),s(X6,X9),s(X6,X19)))) & p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X18),s(X6,X19),s(X6,X49))))) => p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X18),s(X6,X9),s(X6,X49)))))),
  file('thm.pred_set.EXTENSION','thm.relation.TC_RULES|split|1')).
fof(f1351,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK66(X4,X3,X2,X1,X0)),s(fun(X0,bool),X3)))) | s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X4),s(fun(X0,bool),X2))) | s(fun(X0,bool),X3) != s(fun(X0,bool),X4)) )),
  inference(cnf_transformation,[],[f974])).
fof(f974,plain,(
  ! [X0,X1,X2,X3,X4] : (s(fun(X0,bool),X3) != s(fun(X0,bool),X4) | (p(s(bool,'const.bool.IN_2'(s(X0,sK66(X4,X3,X2,X1,X0)),s(fun(X0,bool),X3)))) & s(bool,happ(s(fun(X0,bool),X1),s(X0,sK66(X4,X3,X2,X1,X0)))) != s(bool,happ(s(fun(X0,bool),X2),s(X0,sK66(X4,X3,X2,X1,X0))))) | s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X4),s(fun(X0,bool),X2))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK66])],[f768])).
fof(f768,plain,(
  ! [X0,X1,X2,X3,X4] : (s(fun(X0,bool),X3) != s(fun(X0,bool),X4) | ? [X5] : (p(s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),X3)))) & s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))) != s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) | s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X4),s(fun(X0,bool),X2))))),
  inference(flattening,[],[f767])).
fof(f767,plain,(
  ! [X0,X1,X2,X3,X4] : (s(fun(X0,bool),X3) != s(fun(X0,bool),X4) | (? [X5] : (p(s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),X3)))) & s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))) != s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) | s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X4),s(fun(X0,bool),X2)))))),
  inference(ennf_transformation,[],[f480])).
fof(f480,plain,(
  ! [X0,X1,X2,X3,X4] : (s(fun(X0,bool),X3) = s(fun(X0,bool),X4) => (! [X5] : (p(s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),X3)))) => s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))) = s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) => s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X4),s(fun(X0,bool),X2)))))),
  inference(flattening,[],[f479])).
fof(f479,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : (s(fun(X0,bool),X3) = s(fun(X0,bool),X4) => (! [X5] : (p(s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),X3)))) => s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))) = s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) => s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X4),s(fun(X0,bool),X2)))))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X6] : ! [X14,X15,X16,X17] : (s(fun(X6,bool),X16) = s(fun(X6,bool),X17) => (! [X9] : (p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X16)))) => s(bool,happ(s(fun(X6,bool),X14),s(X6,X9))) = s(bool,happ(s(fun(X6,bool),X15),s(X6,X9)))) => s(bool,'const.bool.RES_FORALL_2'(s(fun(X6,bool),X16),s(fun(X6,bool),X14))) = s(bool,'const.bool.RES_FORALL_2'(s(fun(X6,bool),X17),s(fun(X6,bool),X15)))))),
  file('thm.pred_set.EXTENSION','thm.bool.RES_FORALL_CONG')).
fof(f1352,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun(X0,bool),X3) != s(fun(X0,bool),X4) | s(bool,happ(s(fun(X0,bool),X1),s(X0,sK66(X4,X3,X2,X1,X0)))) != s(bool,happ(s(fun(X0,bool),X2),s(X0,sK66(X4,X3,X2,X1,X0)))) | s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X4),s(fun(X0,bool),X2)))) )),
  inference(cnf_transformation,[],[f974])).
fof(f1349,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK65(X4,X3,X2,X1,X0)),s(fun(X0,bool),X3)))) | s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X4),s(fun(X0,bool),X2))) | s(fun(X0,bool),X3) != s(fun(X0,bool),X4)) )),
  inference(cnf_transformation,[],[f973])).
fof(f973,plain,(
  ! [X0,X1,X2,X3,X4] : (s(fun(X0,bool),X3) != s(fun(X0,bool),X4) | (p(s(bool,'const.bool.IN_2'(s(X0,sK65(X4,X3,X2,X1,X0)),s(fun(X0,bool),X3)))) & s(bool,happ(s(fun(X0,bool),X1),s(X0,sK65(X4,X3,X2,X1,X0)))) != s(bool,happ(s(fun(X0,bool),X2),s(X0,sK65(X4,X3,X2,X1,X0))))) | s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X4),s(fun(X0,bool),X2))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK65])],[f766])).
fof(f766,plain,(
  ! [X0,X1,X2,X3,X4] : (s(fun(X0,bool),X3) != s(fun(X0,bool),X4) | ? [X5] : (p(s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),X3)))) & s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))) != s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) | s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X4),s(fun(X0,bool),X2))))),
  inference(flattening,[],[f765])).
fof(f765,plain,(
  ! [X0,X1,X2,X3,X4] : (s(fun(X0,bool),X3) != s(fun(X0,bool),X4) | (? [X5] : (p(s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),X3)))) & s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))) != s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) | s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X4),s(fun(X0,bool),X2)))))),
  inference(ennf_transformation,[],[f478])).
fof(f478,plain,(
  ! [X0,X1,X2,X3,X4] : (s(fun(X0,bool),X3) = s(fun(X0,bool),X4) => (! [X5] : (p(s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),X3)))) => s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))) = s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) => s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X4),s(fun(X0,bool),X2)))))),
  inference(flattening,[],[f477])).
fof(f477,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : (s(fun(X0,bool),X3) = s(fun(X0,bool),X4) => (! [X5] : (p(s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),X3)))) => s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))) = s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) => s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X4),s(fun(X0,bool),X2)))))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X6] : ! [X14,X15,X16,X17] : (s(fun(X6,bool),X16) = s(fun(X6,bool),X17) => (! [X9] : (p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X16)))) => s(bool,happ(s(fun(X6,bool),X14),s(X6,X9))) = s(bool,happ(s(fun(X6,bool),X15),s(X6,X9)))) => s(bool,'const.bool.RES_EXISTS_2'(s(fun(X6,bool),X16),s(fun(X6,bool),X14))) = s(bool,'const.bool.RES_EXISTS_2'(s(fun(X6,bool),X17),s(fun(X6,bool),X15)))))),
  file('thm.pred_set.EXTENSION','thm.bool.RES_EXISTS_CONG')).
fof(f1350,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun(X0,bool),X3) != s(fun(X0,bool),X4) | s(bool,happ(s(fun(X0,bool),X1),s(X0,sK65(X4,X3,X2,X1,X0)))) != s(bool,happ(s(fun(X0,bool),X2),s(X0,sK65(X4,X3,X2,X1,X0)))) | s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X4),s(fun(X0,bool),X2)))) )),
  inference(cnf_transformation,[],[f973])).
fof(f1345,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,sK63(X3,X2,X1,X0))) | s(bool,'F_0') != s(bool,X1)) )),
  inference(cnf_transformation,[],[f972])).
fof(f1347,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,X1) | p(s(bool,sK64(X3,X2,X1,X0)))) )),
  inference(cnf_transformation,[],[f972])).
fof(f1348,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.min.@_1'(s(fun(X0,bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(X0,bool),'_dst_x0x1_2'(s(fun(bool,bool),'==>_1'(s(bool,sK64(X3,X2,X1,X0)))),s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X2))))),s(fun(X0,bool),'_dst_x0x1_2'(s(fun(bool,bool),'==>_1'(s(bool,sK63(X3,X2,X1,X0)))),s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X3)))))))))) )),
  inference(cnf_transformation,[],[f972])).
fof(f1833,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3)))))) | ~p(s(bool,'const.relation.SC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))) )),
  inference(forward_demodulation,[],[f1832,f1257])).
fof(f1832,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3)))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | ~p(s(bool,'const.relation.SC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))) )),
  inference(forward_demodulation,[],[f1341,f1257])).
fof(f1341,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | ~p(s(bool,'const.relation.SC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f969])).
fof(f969,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.SC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2))))) & ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2))))) | p(s(bool,'const.relation.SC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))))),
  inference(flattening,[],[f968])).
fof(f968,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.SC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2)))))) & ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2))))) | p(s(bool,'const.relation.SC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))))),
  inference(nnf_transformation,[],[f474])).
fof(f474,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.relation.SC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) <=> (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2))))))),
  inference(flattening,[],[f473])).
fof(f473,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.relation.SC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) <=> (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2))))))),
  inference(rectify,[],[f216])).
fof(f216,axiom,(
  ! [X6] : ! [X18,X9,X19] : (p(s(bool,'const.relation.SC_3'(s(fun(X6,fun(X6,bool)),X18),s(X6,X9),s(X6,X19)))) <=> (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X18),s(X6,X9))),s(X6,X19)))) | p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X18),s(X6,X19))),s(X6,X9))))))),
  file('thm.pred_set.EXTENSION','thm.relation.SC_DEF')).
fof(f1831,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2)))))) | p(s(bool,'const.relation.SC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))) )),
  inference(forward_demodulation,[],[f1342,f1257])).
fof(f1342,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.relation.SC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f969])).
fof(f1830,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3)))))) | p(s(bool,'const.relation.SC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))) )),
  inference(forward_demodulation,[],[f1343,f1257])).
fof(f1343,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.relation.SC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2))))) )),
  inference(cnf_transformation,[],[f969])).
fof(f1829,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2)))))) | s(X0,X2) = s(X0,X3) | ~p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))) )),
  inference(forward_demodulation,[],[f1338,f1257])).
fof(f1338,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | s(X0,X2) = s(X0,X3) | ~p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f967])).
fof(f967,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | s(X0,X2) = s(X0,X3) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))))) & ((s(X0,X2) != s(X0,X3) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))))) | p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))))),
  inference(flattening,[],[f966])).
fof(f966,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | (s(X0,X2) = s(X0,X3) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))))) & ((s(X0,X2) != s(X0,X3) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))))) | p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))))),
  inference(nnf_transformation,[],[f472])).
fof(f472,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) <=> (s(X0,X2) = s(X0,X3) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))))))),
  inference(flattening,[],[f471])).
fof(f471,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) <=> (s(X0,X2) = s(X0,X3) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))))))),
  inference(rectify,[],[f209])).
fof(f209,axiom,(
  ! [X6] : ! [X18,X9,X19] : (p(s(bool,'const.relation.RC_3'(s(fun(X6,fun(X6,bool)),X18),s(X6,X9),s(X6,X19)))) <=> (s(X6,X9) = s(X6,X19) | p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X18),s(X6,X9))),s(X6,X19))))))),
  file('thm.pred_set.EXTENSION','thm.relation.RC_DEF')).
fof(f1339,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | s(X0,X2) != s(X0,X3)) )),
  inference(cnf_transformation,[],[f967])).
fof(f1828,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2)))))) | p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))) )),
  inference(forward_demodulation,[],[f1340,f1257])).
fof(f1340,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f967])).
fof(f1827,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2)))))) | ~p(s(bool,'const.relation.STRORD_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))) )),
  inference(forward_demodulation,[],[f1335,f1257])).
fof(f1335,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | ~p(s(bool,'const.relation.STRORD_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f965])).
fof(f965,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.STRORD_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & s(X0,X2) != s(X0,X3))) & (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | s(X0,X2) = s(X0,X3) | p(s(bool,'const.relation.STRORD_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))))),
  inference(flattening,[],[f964])).
fof(f964,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.STRORD_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & s(X0,X2) != s(X0,X3))) & ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | s(X0,X2) = s(X0,X3)) | p(s(bool,'const.relation.STRORD_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))))),
  inference(nnf_transformation,[],[f470])).
fof(f470,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.relation.STRORD_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) <=> (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & s(X0,X2) != s(X0,X3)))),
  inference(flattening,[],[f469])).
fof(f469,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.relation.STRORD_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) <=> (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & ~s(X0,X2) = s(X0,X3)))),
  inference(rectify,[],[f229])).
fof(f229,axiom,(
  ! [X6] : ! [X18,X55,X66] : (p(s(bool,'const.relation.STRORD_3'(s(fun(X6,fun(X6,bool)),X18),s(X6,X55),s(X6,X66)))) <=> (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X18),s(X6,X55))),s(X6,X66)))) & ~s(X6,X55) = s(X6,X66)))),
  file('thm.pred_set.EXTENSION','thm.relation.STRORD')).
fof(f1336,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.relation.STRORD_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | s(X0,X2) != s(X0,X3)) )),
  inference(cnf_transformation,[],[f965])).
fof(f1826,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2)))))) | p(s(bool,'const.relation.STRORD_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | s(X0,X2) = s(X0,X3)) )),
  inference(forward_demodulation,[],[f1337,f1257])).
fof(f1337,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.relation.STRORD_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | s(X0,X2) = s(X0,X3) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f965])).
fof(f1332,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.relation.diag_3'(s(fun(X0,bool),X1),s(X0,X2),s(X0,X3)))) | s(X0,X2) = s(X0,X3)) )),
  inference(cnf_transformation,[],[f963])).
fof(f963,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.diag_3'(s(fun(X0,bool),X1),s(X0,X2),s(X0,X3)))) | (s(X0,X2) = s(X0,X3) & p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))))) & (s(X0,X2) != s(X0,X3) | ~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) | p(s(bool,'const.relation.diag_3'(s(fun(X0,bool),X1),s(X0,X2),s(X0,X3))))))),
  inference(flattening,[],[f962])).
fof(f962,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.diag_3'(s(fun(X0,bool),X1),s(X0,X2),s(X0,X3)))) | (s(X0,X2) = s(X0,X3) & p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))))) & ((s(X0,X2) != s(X0,X3) | ~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1))))) | p(s(bool,'const.relation.diag_3'(s(fun(X0,bool),X1),s(X0,X2),s(X0,X3))))))),
  inference(nnf_transformation,[],[f468])).
fof(f468,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.relation.diag_3'(s(fun(X0,bool),X1),s(X0,X2),s(X0,X3)))) <=> (s(X0,X2) = s(X0,X3) & p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1))))))),
  inference(flattening,[],[f467])).
fof(f467,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.relation.diag_3'(s(fun(X0,bool),X1),s(X0,X2),s(X0,X3)))) <=> (s(X0,X2) = s(X0,X3) & p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1))))))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X6] : ! [X27,X9,X19] : (p(s(bool,'const.relation.diag_3'(s(fun(X6,bool),X27),s(X6,X9),s(X6,X19)))) <=> (s(X6,X9) = s(X6,X19) & p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X27))))))),
  file('thm.pred_set.EXTENSION','thm.relation.diag_def')).
fof(f1333,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.relation.diag_3'(s(fun(X0,bool),X1),s(X0,X2),s(X0,X3)))) | p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f963])).
fof(f1334,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) | p(s(bool,'const.relation.diag_3'(s(fun(X0,bool),X1),s(X0,X2),s(X0,X3)))) | s(X0,X2) != s(X0,X3)) )),
  inference(cnf_transformation,[],[f963])).
fof(f1825,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2)))))) | p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))) )),
  inference(forward_demodulation,[],[f1331,f1257])).
fof(f1331,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f764])).
fof(f764,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))))),
  inference(ennf_transformation,[],[f466])).
fof(f466,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) => p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))))),
  inference(flattening,[],[f465])).
fof(f465,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) => p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))))),
  inference(rectify,[],[f213])).
fof(f213,axiom,(
  ! [X6] : ! [X18,X9,X19] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X18),s(X6,X9))),s(X6,X19)))) => p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X18),s(X6,X9),s(X6,X19)))))),
  file('thm.pred_set.EXTENSION','thm.relation.TC_RULES|split|0')).
fof(f1330,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4))))) )),
  inference(cnf_transformation,[],[f763])).
fof(f763,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | ! [X4] : (~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4))))))),
  inference(ennf_transformation,[],[f464])).
fof(f464,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) => ! [X4] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) => p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4))))))),
  inference(flattening,[],[f463])).
fof(f463,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) => ! [X4] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) => p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4))))))),
  inference(rectify,[],[f255])).
fof(f255,axiom,(
  ! [X6] : ! [X18,X9,X19] : (p(s(bool,'const.relation.RTC_3'(s(fun(X6,fun(X6,bool)),X18),s(X6,X9),s(X6,X19)))) => ! [X49] : (p(s(bool,'const.relation.RTC_3'(s(fun(X6,fun(X6,bool)),X18),s(X6,X19),s(X6,X49)))) => p(s(bool,'const.relation.RTC_3'(s(fun(X6,fun(X6,bool)),X18),s(X6,X9),s(X6,X49))))))),
  file('thm.pred_set.EXTENSION','thm.relation.RTC_RTC')).
fof(f1329,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f762])).
fof(f762,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))))),
  inference(ennf_transformation,[],[f462])).
fof(f462,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) => p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))))),
  inference(flattening,[],[f461])).
fof(f461,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) => p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))))),
  inference(rectify,[],[f254])).
fof(f254,axiom,(
  ! [X6] : ! [X18,X9,X19] : (p(s(bool,'const.relation.RC_3'(s(fun(X6,fun(X6,bool)),X18),s(X6,X9),s(X6,X19)))) => p(s(bool,'const.relation.RTC_3'(s(fun(X6,fun(X6,bool)),X18),s(X6,X9),s(X6,X19)))))),
  file('thm.pred_set.EXTENSION','thm.relation.RC_RTC')).
fof(f1328,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f761])).
fof(f761,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))))),
  inference(ennf_transformation,[],[f460])).
fof(f460,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) => p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))))),
  inference(flattening,[],[f459])).
fof(f459,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) => p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))))),
  inference(rectify,[],[f253])).
fof(f253,axiom,(
  ! [X6] : ! [X18,X9,X19] : (p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X18),s(X6,X9),s(X6,X19)))) => p(s(bool,'const.relation.RTC_3'(s(fun(X6,fun(X6,bool)),X18),s(X6,X9),s(X6,X19)))))),
  file('thm.pred_set.EXTENSION','thm.relation.TC_RTC')).
fof(f1327,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),happ(s(fun(bool,fun(X0,fun(X0,X0))),'const.bool.COND_0'),s(bool,X1))),s(X0,X2))),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f458])).
fof(f458,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),happ(s(fun(bool,fun(X0,fun(X0,X0))),'const.bool.COND_0'),s(bool,X1))),s(X0,X2))),s(X0,X3)))),
  inference(flattening,[],[f457])).
fof(f457,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),happ(s(fun(bool,fun(X0,fun(X0,X0))),'const.bool.COND_0'),s(bool,X1))),s(X0,X2))),s(X0,X3)))),
  inference(rectify,[],[f169])).
fof(f169,axiom,(
  ! [X6] : ! [X21,X22,X31] : s(X6,'const.bool.COND_3'(s(bool,X21),s(X6,X22),s(X6,X31))) = s(X6,happ(s(fun(X6,X6),happ(s(fun(X6,fun(X6,X6)),happ(s(fun(bool,fun(X6,fun(X6,X6))),'const.bool.COND_0'),s(bool,X21))),s(X6,X22))),s(X6,X31)))),
  file('thm.pred_set.EXTENSION','ARITY_0_3_const.bool.COND')).
fof(f1326,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),X1))),s(X0,X2))),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f456])).
fof(f456,plain,(
  ! [X0,X1,X2,X3] : s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),X1))),s(X0,X2))),s(X0,X3)))),
  inference(flattening,[],[f455])).
fof(f455,plain,(
  ! [X0] : ! [X1,X2,X3] : s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),X1))),s(X0,X2))),s(X0,X3)))),
  inference(rectify,[],[f210])).
fof(f210,axiom,(
  ! [X6] : ! [X21,X22,X31] : s(bool,'const.relation.RC_3'(s(fun(X6,fun(X6,bool)),X21),s(X6,X22),s(X6,X31))) = s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),'const.relation.RC_1'(s(fun(X6,fun(X6,bool)),X21))),s(X6,X22))),s(X6,X31)))),
  file('thm.pred_set.EXTENSION','ARITY_1_3_const.relation.RC')).
fof(f1325,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'const.relation.RTC_1'(s(fun(X0,fun(X0,bool)),X1))),s(X0,X2))),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f454])).
fof(f454,plain,(
  ! [X0,X1,X2,X3] : s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'const.relation.RTC_1'(s(fun(X0,fun(X0,bool)),X1))),s(X0,X2))),s(X0,X3)))),
  inference(flattening,[],[f453])).
fof(f453,plain,(
  ! [X0] : ! [X1,X2,X3] : s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'const.relation.RTC_1'(s(fun(X0,fun(X0,bool)),X1))),s(X0,X2))),s(X0,X3)))),
  inference(rectify,[],[f243])).
fof(f243,axiom,(
  ! [X6] : ! [X21,X22,X31] : s(bool,'const.relation.RTC_3'(s(fun(X6,fun(X6,bool)),X21),s(X6,X22),s(X6,X31))) = s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),'const.relation.RTC_1'(s(fun(X6,fun(X6,bool)),X21))),s(X6,X22))),s(X6,X31)))),
  file('thm.pred_set.EXTENSION','ARITY_1_3_const.relation.RTC')).
fof(f1324,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),X1))),s(X0,X2))),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f452])).
fof(f452,plain,(
  ! [X0,X1,X2,X3] : s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),X1))),s(X0,X2))),s(X0,X3)))),
  inference(flattening,[],[f451])).
fof(f451,plain,(
  ! [X0] : ! [X1,X2,X3] : s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),X1))),s(X0,X2))),s(X0,X3)))),
  inference(rectify,[],[f211])).
fof(f211,axiom,(
  ! [X6] : ! [X21,X22,X31] : s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X21),s(X6,X22),s(X6,X31))) = s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),'const.relation.TC_1'(s(fun(X6,fun(X6,bool)),X21))),s(X6,X22))),s(X6,X31)))),
  file('thm.pred_set.EXTENSION','ARITY_1_3_const.relation.TC')).
fof(f1323,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'const.relation.STRORD_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'const.relation.STRORD_1'(s(fun(X0,fun(X0,bool)),X1))),s(X0,X2))),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f450])).
fof(f450,plain,(
  ! [X0,X1,X2,X3] : s(bool,'const.relation.STRORD_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'const.relation.STRORD_1'(s(fun(X0,fun(X0,bool)),X1))),s(X0,X2))),s(X0,X3)))),
  inference(flattening,[],[f449])).
fof(f449,plain,(
  ! [X0] : ! [X1,X2,X3] : s(bool,'const.relation.STRORD_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'const.relation.STRORD_1'(s(fun(X0,fun(X0,bool)),X1))),s(X0,X2))),s(X0,X3)))),
  inference(rectify,[],[f230])).
fof(f230,axiom,(
  ! [X6] : ! [X21,X22,X31] : s(bool,'const.relation.STRORD_3'(s(fun(X6,fun(X6,bool)),X21),s(X6,X22),s(X6,X31))) = s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),'const.relation.STRORD_1'(s(fun(X6,fun(X6,bool)),X21))),s(X6,X22))),s(X6,X31)))),
  file('thm.pred_set.EXTENSION','ARITY_1_3_const.relation.STRORD')).
fof(f1322,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'const.relation.inv_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),X1))),s(X0,X2))),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f448])).
fof(f448,plain,(
  ! [X0,X1,X2,X3] : s(bool,'const.relation.inv_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),X1))),s(X0,X2))),s(X0,X3)))),
  inference(flattening,[],[f447])).
fof(f447,plain,(
  ! [X0] : ! [X1,X2,X3] : s(bool,'const.relation.inv_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),X1))),s(X0,X2))),s(X0,X3)))),
  inference(rectify,[],[f157])).
fof(f157,axiom,(
  ! [X6] : ! [X21,X22,X31] : s(bool,'const.relation.inv_3'(s(fun(X6,fun(X6,bool)),X21),s(X6,X22),s(X6,X31))) = s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),'const.relation.inv_1'(s(fun(X6,fun(X6,bool)),X21))),s(X6,X22))),s(X6,X31)))),
  file('thm.pred_set.EXTENSION','ARITY_1_3_const.relation.inv')).
fof(f1321,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'const.relation.SC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1))),s(X0,X2))),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f446])).
fof(f446,plain,(
  ! [X0,X1,X2,X3] : s(bool,'const.relation.SC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1))),s(X0,X2))),s(X0,X3)))),
  inference(flattening,[],[f445])).
fof(f445,plain,(
  ! [X0] : ! [X1,X2,X3] : s(bool,'const.relation.SC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1))),s(X0,X2))),s(X0,X3)))),
  inference(rectify,[],[f217])).
fof(f217,axiom,(
  ! [X6] : ! [X21,X22,X31] : s(bool,'const.relation.SC_3'(s(fun(X6,fun(X6,bool)),X21),s(X6,X22),s(X6,X31))) = s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),'const.relation.SC_1'(s(fun(X6,fun(X6,bool)),X21))),s(X6,X22))),s(X6,X31)))),
  file('thm.pred_set.EXTENSION','ARITY_1_3_const.relation.SC')).
fof(f1320,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'const.relation.diag_3'(s(fun(X0,bool),X1),s(X0,X2),s(X0,X3))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'const.relation.diag_1'(s(fun(X0,bool),X1))),s(X0,X2))),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f444])).
fof(f444,plain,(
  ! [X0,X1,X2,X3] : s(bool,'const.relation.diag_3'(s(fun(X0,bool),X1),s(X0,X2),s(X0,X3))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'const.relation.diag_1'(s(fun(X0,bool),X1))),s(X0,X2))),s(X0,X3)))),
  inference(flattening,[],[f443])).
fof(f443,plain,(
  ! [X0] : ! [X1,X2,X3] : s(bool,'const.relation.diag_3'(s(fun(X0,bool),X1),s(X0,X2),s(X0,X3))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'const.relation.diag_1'(s(fun(X0,bool),X1))),s(X0,X2))),s(X0,X3)))),
  inference(rectify,[],[f158])).
fof(f158,axiom,(
  ! [X6] : ! [X21,X22,X31] : s(bool,'const.relation.diag_3'(s(fun(X6,bool),X21),s(X6,X22),s(X6,X31))) = s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),'const.relation.diag_1'(s(fun(X6,bool),X21))),s(X6,X22))),s(X6,X31)))),
  file('thm.pred_set.EXTENSION','ARITY_1_3_const.relation.diag')).
fof(f1812,plain,(
  ( ! [X6,X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,'const.bool.IN_2'(s(X0,sK58(X2,X1,X0)),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X1))))) )),
  inference(forward_demodulation,[],[f1811,f1257])).
fof(f1811,plain,(
  ( ! [X6,X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK58(X2,X1,X0)),s(fun(X0,bool),X1)))) | ~p(s(bool,X2)) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6))))) )),
  inference(forward_demodulation,[],[f1311,f1257])).
fof(f1311,plain,(
  ( ! [X6,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK58(X2,X1,X0))))) | ~p(s(bool,X2)) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6))))) )),
  inference(cnf_transformation,[],[f957])).
fof(f957,plain,(
  ! [X0,X1,X2] : ((! [X6] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) | ~p(s(bool,X2))) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK58(X2,X1,X0))))) & p(s(bool,X2)))) & (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,X2)) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK59(X2,X1,X0))))) & p(s(bool,X2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK58,sK59])],[f956])).
fof(f956,plain,(
  ! [X0,X1,X2] : ((! [X6] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) | ~p(s(bool,X2))) | (? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & p(s(bool,X2)))) & (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,X2)) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2)))))),
  inference(rectify,[],[f955])).
fof(f955,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,X2))) | (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & p(s(bool,X2)))) & (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,X2)) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2)))))),
  inference(flattening,[],[f954])).
fof(f954,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,X2))) | (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & p(s(bool,X2)))) & ((! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,X2))) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2)))))),
  inference(nnf_transformation,[],[f438])).
fof(f438,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2))) <=> (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & p(s(bool,X2))))),
  inference(flattening,[],[f437])).
fof(f437,plain,(
  ! [X0] : ! [X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2))) <=> (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & p(s(bool,X2))))),
  inference(rectify,[],[f100])).
fof(f100,axiom,(
  ! [X6] : ! [X17,X16] : (? [X9] : (p(s(bool,happ(s(fun(X6,bool),X17),s(X6,X9)))) & p(s(bool,X16))) <=> (? [X9] : p(s(bool,happ(s(fun(X6,bool),X17),s(X6,X9)))) & p(s(bool,X16))))),
  file('thm.pred_set.EXTENSION','thm.bool.LEFT_EXISTS_AND_THM')).
fof(f1810,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,'const.bool.IN_2'(s(X0,sK59(X2,X1,X0)),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))))) )),
  inference(forward_demodulation,[],[f1809,f1257])).
fof(f1809,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK59(X2,X1,X0)),s(fun(X0,bool),X1)))) | ~p(s(bool,X2)) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) )),
  inference(forward_demodulation,[],[f1313,f1257])).
fof(f1313,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK59(X2,X1,X0))))) | ~p(s(bool,X2)) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) )),
  inference(cnf_transformation,[],[f957])).
fof(f1808,plain,(
  ( ! [X6,X2,X0,X1] : (~p(s(bool,X1)) | p(s(bool,'const.bool.IN_2'(s(X0,sK56(X2,X1,X0)),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X2))))) )),
  inference(forward_demodulation,[],[f1807,f1257])).
fof(f1807,plain,(
  ( ! [X6,X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK56(X2,X1,X0)),s(fun(X0,bool),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))) | ~p(s(bool,X1))) )),
  inference(forward_demodulation,[],[f1308,f1257])).
fof(f1308,plain,(
  ( ! [X6,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK56(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f953])).
fof(f953,plain,(
  ! [X0,X1,X2] : ((! [X6] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK56(X2,X1,X0))))))) & (~p(s(bool,X1)) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK57(X2,X1,X0))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK56,sK57])],[f952])).
fof(f952,plain,(
  ! [X0,X1,X2] : ((! [X6] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | (p(s(bool,X1)) & ? [X5] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))))) & (~p(s(bool,X1)) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X3] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(rectify,[],[f951])).
fof(f951,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (p(s(bool,X1)) & ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (~p(s(bool,X1)) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X3] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(flattening,[],[f950])).
fof(f950,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (p(s(bool,X1)) & ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & ((~p(s(bool,X1)) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ? [X3] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f436])).
fof(f436,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (p(s(bool,X1)) & ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f435])).
fof(f435,plain,(
  ! [X0] : ! [X1,X2] : (? [X3] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (p(s(bool,X1)) & ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f96])).
fof(f96,axiom,(
  ! [X6] : ! [X17,X16] : (? [X9] : (p(s(bool,X17)) & p(s(bool,happ(s(fun(X6,bool),X16),s(X6,X9))))) <=> (p(s(bool,X17)) & ? [X9] : p(s(bool,happ(s(fun(X6,bool),X16),s(X6,X9))))))),
  file('thm.pred_set.EXTENSION','thm.bool.RIGHT_EXISTS_AND_THM')).
fof(f1806,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,X1)) | p(s(bool,'const.bool.IN_2'(s(X0,sK57(X2,X1,X0)),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2))))) )),
  inference(forward_demodulation,[],[f1805,f1257])).
fof(f1805,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK57(X2,X1,X0)),s(fun(X0,bool),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ~p(s(bool,X1))) )),
  inference(forward_demodulation,[],[f1310,f1257])).
fof(f1310,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK57(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f953])).
fof(f1800,plain,(
  ( ! [X2,X0,X7,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,sK51(X2,X1,X0)),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,sK51(X2,X1,X0)),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X7),s(fun(X0,bool),X1))))) )),
  inference(forward_demodulation,[],[f1799,f1257])).
fof(f1799,plain,(
  ( ! [X2,X0,X7,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,sK51(X2,X1,X0)),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,X7),s(fun(X0,bool),X1)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK51(X2,X1,X0)))))) )),
  inference(forward_demodulation,[],[f1798,f1257])).
fof(f1798,plain,(
  ( ! [X2,X0,X7,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X7),s(fun(X0,bool),X1)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK51(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK51(X2,X1,X0)))))) )),
  inference(forward_demodulation,[],[f1299,f1257])).
fof(f1299,plain,(
  ( ! [X2,X0,X7,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X7)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK51(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK51(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f945])).
fof(f1794,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,sK53(X2,X1,X0)),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,sK52(X2,X1,X0)),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))))) )),
  inference(forward_demodulation,[],[f1793,f1257])).
fof(f1793,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,sK53(X2,X1,X0)),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK52(X2,X1,X0)))))) )),
  inference(forward_demodulation,[],[f1792,f1257])).
fof(f1792,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK53(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK52(X2,X1,X0)))))) )),
  inference(forward_demodulation,[],[f1301,f1257])).
fof(f1301,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK53(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK52(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f945])).
fof(f1288,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK46(X2,X1,X0)),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f933])).
fof(f933,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | (p(s(bool,'const.bool.IN_2'(s(X0,sK46(X2,X1,X0)),s(fun(X0,bool),X1)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK46(X2,X1,X0))))))) & (! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK46])],[f932])).
fof(f932,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ? [X4] : (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f931])).
fof(f931,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))) & (! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(nnf_transformation,[],[f426])).
fof(f426,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(flattening,[],[f425])).
fof(f425,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6] : ! [X5,X12] : (p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X6,bool),X5),s(fun(X6,bool),X12)))) <=> ? [X9] : (p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X5)))) & p(s(bool,happ(s(fun(X6,bool),X12),s(X6,X9))))))),
  file('thm.pred_set.EXTENSION','thm.bool.RES_EXISTS_DEF')).
fof(f1780,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK46(X2,X1,X0)),s(fun(X0,bool),X2))))) )),
  inference(forward_demodulation,[],[f1289,f1257])).
fof(f1289,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK46(X2,X1,X0))))) | ~p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f933])).
fof(f1285,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK45(X2,X1,X0)),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f930])).
fof(f930,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | (p(s(bool,'const.bool.IN_2'(s(X0,sK45(X2,X1,X0)),s(fun(X0,bool),X1)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK45(X2,X1,X0))))))) & (! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK45])],[f929])).
fof(f929,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ? [X4] : (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f928])).
fof(f928,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))) & (! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(nnf_transformation,[],[f424])).
fof(f424,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(flattening,[],[f423])).
fof(f423,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X6] : ! [X17,X15] : (p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X6,bool),X17),s(fun(X6,bool),X15)))) <=> ? [X9] : (p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X17)))) & p(s(bool,happ(s(fun(X6,bool),X15),s(X6,X9))))))),
  file('thm.pred_set.EXTENSION','thm.bool.RES_EXISTS_THM')).
fof(f1778,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK45(X2,X1,X0)),s(fun(X0,bool),X2))))) )),
  inference(forward_demodulation,[],[f1286,f1257])).
fof(f1286,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK45(X2,X1,X0))))) | ~p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f930])).
fof(f1777,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) )),
  inference(forward_demodulation,[],[f1287,f1257])).
fof(f1287,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f930])).
fof(f1283,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK44(X2,X1,X0)),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f927])).
fof(f927,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & ((p(s(bool,'const.bool.IN_2'(s(X0,sK44(X2,X1,X0)),s(fun(X0,bool),X1)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK44(X2,X1,X0)))))) | p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK44])],[f926])).
fof(f926,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f925])).
fof(f925,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))) & (? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(nnf_transformation,[],[f760])).
fof(f760,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(ennf_transformation,[],[f422])).
fof(f422,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) => p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(flattening,[],[f421])).
fof(f421,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) => p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : ! [X5,X12] : (p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X6,bool),X5),s(fun(X6,bool),X12)))) <=> ! [X9] : (p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X5)))) => p(s(bool,happ(s(fun(X6,bool),X12),s(X6,X9))))))),
  file('thm.pred_set.EXTENSION','thm.bool.RES_FORALL_DEF')).
fof(f1775,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,sK44(X2,X1,X0)),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) )),
  inference(forward_demodulation,[],[f1284,f1257])).
fof(f1284,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK44(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f927])).
fof(f1774,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2))))) )),
  inference(forward_demodulation,[],[f1279,f1257])).
fof(f1279,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f924])).
fof(f924,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & ((p(s(bool,'const.bool.IN_2'(s(X0,sK43(X2,X1,X0)),s(fun(X0,bool),X1)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK43(X2,X1,X0)))))) | p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK43])],[f923])).
fof(f923,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f922])).
fof(f922,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))) & (? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(nnf_transformation,[],[f759])).
fof(f759,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(ennf_transformation,[],[f420])).
fof(f420,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) => p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(flattening,[],[f419])).
fof(f419,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) => p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X6] : ! [X17,X15] : (p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X6,bool),X17),s(fun(X6,bool),X15)))) <=> ! [X9] : (p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X17)))) => p(s(bool,happ(s(fun(X6,bool),X15),s(X6,X9))))))),
  file('thm.pred_set.EXTENSION','thm.bool.RES_FORALL_THM')).
fof(f1280,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK43(X2,X1,X0)),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f924])).
fof(f1773,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,sK43(X2,X1,X0)),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) )),
  inference(forward_demodulation,[],[f1281,f1257])).
fof(f1281,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK43(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f924])).
fof(f1772,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK39(X2,X1,X0)),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK38(X2,X1,X0))))))) | p(s(bool,'const.bool.IN_2'(s(X0,sK41(X2,X1,X0)),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK40(X2,X1,X0))))))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3)))))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4))))) )),
  inference(forward_demodulation,[],[f1771,f1257])).
fof(f1771,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK41(X2,X1,X0)),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK40(X2,X1,X0))))))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3)))))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK38(X2,X1,X0)))),s(X0,sK39(X2,X1,X0)))))) )),
  inference(forward_demodulation,[],[f1770,f1257])).
fof(f1770,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3)))))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK40(X2,X1,X0)))),s(X0,sK41(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK38(X2,X1,X0)))),s(X0,sK39(X2,X1,X0)))))) )),
  inference(forward_demodulation,[],[f1271,f1257])).
fof(f1271,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK40(X2,X1,X0)))),s(X0,sK41(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK38(X2,X1,X0)))),s(X0,sK39(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f920])).
fof(f920,plain,(
  ! [X0,X1,X2] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK38(X2,X1,X0)))),s(X0,sK39(X2,X1,X0))))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK38(X2,X1,X0)))),s(X0,sK39(X2,X1,X0)))))) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK40(X2,X1,X0)))),s(X0,sK41(X2,X1,X0))))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK41(X2,X1,X0)))),s(X0,sK42(X2,X1,X0))))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK40(X2,X1,X0)))),s(X0,sK42(X2,X1,X0)))))) | ! [X3,X4] : (~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK38,sK39,sK40,sK41,sK42])],[f919])).
fof(f919,plain,(
  ! [X0,X1,X2] : (? [X8,X9] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X8))),s(X0,X9)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X8))),s(X0,X9))))) | ? [X5,X6,X7] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X5))),s(X0,X6)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X6))),s(X0,X7)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X5))),s(X0,X7))))) | ! [X3,X4] : (~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4))))))),
  inference(rectify,[],[f758])).
fof(f758,plain,(
  ! [X0,X1,X2] : (? [X6,X7] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X6))),s(X0,X7)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X6))),s(X0,X7))))) | ? [X3,X4,X5] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X4))),s(X0,X5)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X5))))) | ! [X8,X9] : (~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X8),s(X0,X9)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X8))),s(X0,X9))))))),
  inference(flattening,[],[f757])).
fof(f757,plain,(
  ! [X0,X1,X2] : ((? [X6,X7] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X6))),s(X0,X7)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X6))),s(X0,X7))))) | ? [X3,X4,X5] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X4))),s(X0,X5))))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X5)))))) | ! [X8,X9] : (~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X8),s(X0,X9)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X8))),s(X0,X9))))))),
  inference(ennf_transformation,[],[f416])).
fof(f416,plain,(
  ! [X0,X1,X2] : ((! [X6,X7] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X6))),s(X0,X7)))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X6))),s(X0,X7))))) & ! [X3,X4,X5] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X4))),s(X0,X5))))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X5)))))) => ! [X8,X9] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X8),s(X0,X9)))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X8))),s(X0,X9))))))),
  inference(flattening,[],[f415])).
fof(f415,plain,(
  ! [X0] : ! [X1,X2] : ((! [X6,X7] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X6))),s(X0,X7)))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X6))),s(X0,X7))))) & ! [X3,X4,X5] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X4))),s(X0,X5))))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X5)))))) => ! [X8,X9] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X8),s(X0,X9)))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X8))),s(X0,X9))))))),
  inference(rectify,[],[f212])).
fof(f212,axiom,(
  ! [X6] : ! [X18,X17] : ((! [X9,X19] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X18),s(X6,X9))),s(X6,X19)))) => p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X17),s(X6,X9))),s(X6,X19))))) & ! [X9,X19,X49] : ((p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X17),s(X6,X9))),s(X6,X19)))) & p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X17),s(X6,X19))),s(X6,X49))))) => p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X17),s(X6,X9))),s(X6,X49)))))) => ! [X51,X52] : (p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X18),s(X6,X51),s(X6,X52)))) => p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X17),s(X6,X51))),s(X6,X52))))))),
  file('thm.pred_set.EXTENSION','thm.relation.TC_INDUCT')).
fof(f1769,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,sK39(X2,X1,X0)),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK38(X2,X1,X0))))))) | p(s(bool,'const.bool.IN_2'(s(X0,sK41(X2,X1,X0)),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK40(X2,X1,X0))))))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3)))))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4))))) )),
  inference(forward_demodulation,[],[f1768,f1257])).
fof(f1768,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK41(X2,X1,X0)),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK40(X2,X1,X0))))))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3)))))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK38(X2,X1,X0)))),s(X0,sK39(X2,X1,X0)))))) )),
  inference(forward_demodulation,[],[f1767,f1257])).
fof(f1767,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3)))))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK40(X2,X1,X0)))),s(X0,sK41(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK38(X2,X1,X0)))),s(X0,sK39(X2,X1,X0)))))) )),
  inference(forward_demodulation,[],[f1272,f1257])).
fof(f1272,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK40(X2,X1,X0)))),s(X0,sK41(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK38(X2,X1,X0)))),s(X0,sK39(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f920])).
fof(f1766,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK39(X2,X1,X0)),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK38(X2,X1,X0))))))) | p(s(bool,'const.bool.IN_2'(s(X0,sK42(X2,X1,X0)),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK41(X2,X1,X0))))))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3)))))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4))))) )),
  inference(forward_demodulation,[],[f1765,f1257])).
fof(f1765,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK42(X2,X1,X0)),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK41(X2,X1,X0))))))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3)))))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK38(X2,X1,X0)))),s(X0,sK39(X2,X1,X0)))))) )),
  inference(forward_demodulation,[],[f1764,f1257])).
fof(f1764,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3)))))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK41(X2,X1,X0)))),s(X0,sK42(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK38(X2,X1,X0)))),s(X0,sK39(X2,X1,X0)))))) )),
  inference(forward_demodulation,[],[f1273,f1257])).
fof(f1273,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK41(X2,X1,X0)))),s(X0,sK42(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK38(X2,X1,X0)))),s(X0,sK39(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f920])).
fof(f1763,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,sK39(X2,X1,X0)),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK38(X2,X1,X0))))))) | p(s(bool,'const.bool.IN_2'(s(X0,sK42(X2,X1,X0)),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK41(X2,X1,X0))))))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3)))))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4))))) )),
  inference(forward_demodulation,[],[f1762,f1257])).
fof(f1762,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK42(X2,X1,X0)),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK41(X2,X1,X0))))))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3)))))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK38(X2,X1,X0)))),s(X0,sK39(X2,X1,X0)))))) )),
  inference(forward_demodulation,[],[f1761,f1257])).
fof(f1761,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3)))))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK41(X2,X1,X0)))),s(X0,sK42(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK38(X2,X1,X0)))),s(X0,sK39(X2,X1,X0)))))) )),
  inference(forward_demodulation,[],[f1274,f1257])).
fof(f1274,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK41(X2,X1,X0)))),s(X0,sK42(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK38(X2,X1,X0)))),s(X0,sK39(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f920])).
fof(f1760,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK39(X2,X1,X0)),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK38(X2,X1,X0))))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,sK42(X2,X1,X0)),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK40(X2,X1,X0))))))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3)))))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4))))) )),
  inference(forward_demodulation,[],[f1759,f1257])).
fof(f1759,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,sK42(X2,X1,X0)),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK40(X2,X1,X0))))))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3)))))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK38(X2,X1,X0)))),s(X0,sK39(X2,X1,X0)))))) )),
  inference(forward_demodulation,[],[f1758,f1257])).
fof(f1758,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3)))))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK40(X2,X1,X0)))),s(X0,sK42(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK38(X2,X1,X0)))),s(X0,sK39(X2,X1,X0)))))) )),
  inference(forward_demodulation,[],[f1275,f1257])).
fof(f1275,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK40(X2,X1,X0)))),s(X0,sK42(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK38(X2,X1,X0)))),s(X0,sK39(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f920])).
fof(f1757,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,sK39(X2,X1,X0)),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK38(X2,X1,X0))))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,sK42(X2,X1,X0)),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK40(X2,X1,X0))))))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3)))))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4))))) )),
  inference(forward_demodulation,[],[f1756,f1257])).
fof(f1756,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,sK42(X2,X1,X0)),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK40(X2,X1,X0))))))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3)))))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK38(X2,X1,X0)))),s(X0,sK39(X2,X1,X0)))))) )),
  inference(forward_demodulation,[],[f1755,f1257])).
fof(f1755,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3)))))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK40(X2,X1,X0)))),s(X0,sK42(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK38(X2,X1,X0)))),s(X0,sK39(X2,X1,X0)))))) )),
  inference(forward_demodulation,[],[f1276,f1257])).
fof(f1276,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK40(X2,X1,X0)))),s(X0,sK42(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK38(X2,X1,X0)))),s(X0,sK39(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f920])).
fof(f1754,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,sK34(X2,X1,X0)),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK34(X2,X1,X0))))))) | p(s(bool,'const.bool.IN_2'(s(X0,sK36(X2,X1,X0)),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK35(X2,X1,X0))))))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3)))))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4))))) )),
  inference(forward_demodulation,[],[f1753,f1257])).
fof(f1753,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK36(X2,X1,X0)),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK35(X2,X1,X0))))))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3)))))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK34(X2,X1,X0)))),s(X0,sK34(X2,X1,X0)))))) )),
  inference(forward_demodulation,[],[f1752,f1257])).
fof(f1752,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3)))))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK35(X2,X1,X0)))),s(X0,sK36(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK34(X2,X1,X0)))),s(X0,sK34(X2,X1,X0)))))) )),
  inference(forward_demodulation,[],[f1268,f1257])).
fof(f1268,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK35(X2,X1,X0)))),s(X0,sK36(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK34(X2,X1,X0)))),s(X0,sK34(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f918])).
fof(f918,plain,(
  ! [X0,X1,X2] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK34(X2,X1,X0)))),s(X0,sK34(X2,X1,X0))))) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK35(X2,X1,X0)))),s(X0,sK36(X2,X1,X0))))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK36(X2,X1,X0)))),s(X0,sK37(X2,X1,X0))))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK35(X2,X1,X0)))),s(X0,sK37(X2,X1,X0)))))) | ! [X3,X4] : (~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK34,sK35,sK36,sK37])],[f917])).
fof(f917,plain,(
  ! [X0,X1,X2] : (? [X8] : ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X8))),s(X0,X8)))) | ? [X5,X6,X7] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X6)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X6))),s(X0,X7)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X5))),s(X0,X7))))) | ! [X3,X4] : (~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4))))))),
  inference(rectify,[],[f756])).
fof(f756,plain,(
  ! [X0,X1,X2] : (? [X6] : ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X6))),s(X0,X6)))) | ? [X3,X4,X5] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X4)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X4))),s(X0,X5)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X5))))) | ! [X7,X8] : (~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X7),s(X0,X8)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X7))),s(X0,X8))))))),
  inference(flattening,[],[f755])).
fof(f755,plain,(
  ! [X0,X1,X2] : ((? [X6] : ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X6))),s(X0,X6)))) | ? [X3,X4,X5] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X4)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X4))),s(X0,X5))))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X5)))))) | ! [X7,X8] : (~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X7),s(X0,X8)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X7))),s(X0,X8))))))),
  inference(ennf_transformation,[],[f414])).
fof(f414,plain,(
  ! [X0,X1,X2] : ((! [X6] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X6))),s(X0,X6)))) & ! [X3,X4,X5] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X4)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X4))),s(X0,X5))))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X5)))))) => ! [X7,X8] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X7),s(X0,X8)))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X7))),s(X0,X8))))))),
  inference(flattening,[],[f413])).
fof(f413,plain,(
  ! [X0] : ! [X1,X2] : ((! [X6] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X6))),s(X0,X6)))) & ! [X3,X4,X5] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X4)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X4))),s(X0,X5))))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X5)))))) => ! [X7,X8] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X7),s(X0,X8)))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X7))),s(X0,X8))))))),
  inference(rectify,[],[f259])).
fof(f259,axiom,(
  ! [X6] : ! [X18,X17] : ((! [X9] : p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X17),s(X6,X9))),s(X6,X9)))) & ! [X9,X19,X49] : ((p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X18),s(X6,X9))),s(X6,X19)))) & p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X17),s(X6,X19))),s(X6,X49))))) => p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X17),s(X6,X9))),s(X6,X49)))))) => ! [X9,X19] : (p(s(bool,'const.relation.RTC_3'(s(fun(X6,fun(X6,bool)),X18),s(X6,X9),s(X6,X19)))) => p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X17),s(X6,X9))),s(X6,X19))))))),
  file('thm.pred_set.EXTENSION','thm.relation.RTC_INDUCT')).
fof(f1751,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,sK34(X2,X1,X0)),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK34(X2,X1,X0))))))) | p(s(bool,'const.bool.IN_2'(s(X0,sK37(X2,X1,X0)),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK36(X2,X1,X0))))))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3)))))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4))))) )),
  inference(forward_demodulation,[],[f1750,f1257])).
fof(f1750,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK37(X2,X1,X0)),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK36(X2,X1,X0))))))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3)))))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK34(X2,X1,X0)))),s(X0,sK34(X2,X1,X0)))))) )),
  inference(forward_demodulation,[],[f1749,f1257])).
fof(f1749,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3)))))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK36(X2,X1,X0)))),s(X0,sK37(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK34(X2,X1,X0)))),s(X0,sK34(X2,X1,X0)))))) )),
  inference(forward_demodulation,[],[f1269,f1257])).
fof(f1269,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK36(X2,X1,X0)))),s(X0,sK37(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK34(X2,X1,X0)))),s(X0,sK34(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f918])).
fof(f1748,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,sK34(X2,X1,X0)),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK34(X2,X1,X0))))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,sK37(X2,X1,X0)),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK35(X2,X1,X0))))))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3)))))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4))))) )),
  inference(forward_demodulation,[],[f1747,f1257])).
fof(f1747,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,sK37(X2,X1,X0)),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK35(X2,X1,X0))))))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3)))))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK34(X2,X1,X0)))),s(X0,sK34(X2,X1,X0)))))) )),
  inference(forward_demodulation,[],[f1746,f1257])).
fof(f1746,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3)))))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK35(X2,X1,X0)))),s(X0,sK37(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK34(X2,X1,X0)))),s(X0,sK34(X2,X1,X0)))))) )),
  inference(forward_demodulation,[],[f1270,f1257])).
fof(f1270,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK35(X2,X1,X0)))),s(X0,sK37(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK34(X2,X1,X0)))),s(X0,sK34(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f918])).
fof(f1745,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,sK30(X2,X1,X0)),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK30(X2,X1,X0))))))) | p(s(bool,'const.bool.IN_2'(s(X0,sK32(X2,X1,X0)),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK31(X2,X1,X0))))))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3)))))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4))))) )),
  inference(forward_demodulation,[],[f1744,f1257])).
fof(f1744,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK32(X2,X1,X0)),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK31(X2,X1,X0))))))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3)))))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK30(X2,X1,X0)))),s(X0,sK30(X2,X1,X0)))))) )),
  inference(forward_demodulation,[],[f1743,f1257])).
fof(f1743,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3)))))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK31(X2,X1,X0)))),s(X0,sK32(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK30(X2,X1,X0)))),s(X0,sK30(X2,X1,X0)))))) )),
  inference(forward_demodulation,[],[f1264,f1257])).
fof(f1264,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK31(X2,X1,X0)))),s(X0,sK32(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK30(X2,X1,X0)))),s(X0,sK30(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f916])).
fof(f916,plain,(
  ! [X0,X1,X2] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK30(X2,X1,X0)))),s(X0,sK30(X2,X1,X0))))) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK31(X2,X1,X0)))),s(X0,sK32(X2,X1,X0))))) & p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,sK32(X2,X1,X0)),s(X0,sK33(X2,X1,X0))))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK32(X2,X1,X0)))),s(X0,sK33(X2,X1,X0))))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK31(X2,X1,X0)))),s(X0,sK33(X2,X1,X0)))))) | ! [X3,X4] : (~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK30,sK31,sK32,sK33])],[f915])).
fof(f915,plain,(
  ! [X0,X1,X2] : (? [X8] : ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X8))),s(X0,X8)))) | ? [X5,X6,X7] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X6)))) & p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X6),s(X0,X7)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X6))),s(X0,X7)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X5))),s(X0,X7))))) | ! [X3,X4] : (~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4))))))),
  inference(rectify,[],[f754])).
fof(f754,plain,(
  ! [X0,X1,X2] : (? [X6] : ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X6))),s(X0,X6)))) | ? [X3,X4,X5] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X4)))) & p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X5)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X4))),s(X0,X5)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X5))))) | ! [X7,X8] : (~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X7),s(X0,X8)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X7))),s(X0,X8))))))),
  inference(flattening,[],[f753])).
fof(f753,plain,(
  ! [X0,X1,X2] : ((? [X6] : ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X6))),s(X0,X6)))) | ? [X3,X4,X5] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X4)))) & p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X5)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X4))),s(X0,X5))))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X5)))))) | ! [X7,X8] : (~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X7),s(X0,X8)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X7))),s(X0,X8))))))),
  inference(ennf_transformation,[],[f412])).
fof(f412,plain,(
  ! [X0,X1,X2] : ((! [X6] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X6))),s(X0,X6)))) & ! [X3,X4,X5] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X4)))) & p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X5)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X4))),s(X0,X5))))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X5)))))) => ! [X7,X8] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X7),s(X0,X8)))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X7))),s(X0,X8))))))),
  inference(flattening,[],[f411])).
fof(f411,plain,(
  ! [X0] : ! [X1,X2] : ((! [X6] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X6))),s(X0,X6)))) & ! [X3,X4,X5] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X4)))) & p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X5)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X4))),s(X0,X5))))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X5)))))) => ! [X7,X8] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X7),s(X0,X8)))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X7))),s(X0,X8))))))),
  inference(rectify,[],[f256])).
fof(f256,axiom,(
  ! [X6] : ! [X18,X17] : ((! [X9] : p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X17),s(X6,X9))),s(X6,X9)))) & ! [X9,X19,X49] : ((p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X18),s(X6,X9))),s(X6,X19)))) & p(s(bool,'const.relation.RTC_3'(s(fun(X6,fun(X6,bool)),X18),s(X6,X19),s(X6,X49)))) & p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X17),s(X6,X19))),s(X6,X49))))) => p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X17),s(X6,X9))),s(X6,X49)))))) => ! [X9,X19] : (p(s(bool,'const.relation.RTC_3'(s(fun(X6,fun(X6,bool)),X18),s(X6,X9),s(X6,X19)))) => p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X17),s(X6,X9))),s(X6,X19))))))),
  file('thm.pred_set.EXTENSION','thm.relation.RTC_STRONG_INDUCT')).
fof(f1742,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,sK30(X2,X1,X0)),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK30(X2,X1,X0))))))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3)))))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,sK32(X2,X1,X0)),s(X0,sK33(X2,X1,X0)))))) )),
  inference(forward_demodulation,[],[f1741,f1257])).
fof(f1741,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3)))))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,sK32(X2,X1,X0)),s(X0,sK33(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK30(X2,X1,X0)))),s(X0,sK30(X2,X1,X0)))))) )),
  inference(forward_demodulation,[],[f1265,f1257])).
fof(f1265,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,sK32(X2,X1,X0)),s(X0,sK33(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK30(X2,X1,X0)))),s(X0,sK30(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f916])).
fof(f1740,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,sK30(X2,X1,X0)),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK30(X2,X1,X0))))))) | p(s(bool,'const.bool.IN_2'(s(X0,sK33(X2,X1,X0)),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK32(X2,X1,X0))))))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3)))))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4))))) )),
  inference(forward_demodulation,[],[f1739,f1257])).
fof(f1739,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK33(X2,X1,X0)),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK32(X2,X1,X0))))))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3)))))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK30(X2,X1,X0)))),s(X0,sK30(X2,X1,X0)))))) )),
  inference(forward_demodulation,[],[f1738,f1257])).
fof(f1738,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3)))))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK32(X2,X1,X0)))),s(X0,sK33(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK30(X2,X1,X0)))),s(X0,sK30(X2,X1,X0)))))) )),
  inference(forward_demodulation,[],[f1266,f1257])).
fof(f1266,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK32(X2,X1,X0)))),s(X0,sK33(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK30(X2,X1,X0)))),s(X0,sK30(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f916])).
fof(f1737,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,sK30(X2,X1,X0)),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK30(X2,X1,X0))))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,sK33(X2,X1,X0)),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK31(X2,X1,X0))))))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3)))))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4))))) )),
  inference(forward_demodulation,[],[f1736,f1257])).
fof(f1736,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,sK33(X2,X1,X0)),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK31(X2,X1,X0))))))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3)))))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK30(X2,X1,X0)))),s(X0,sK30(X2,X1,X0)))))) )),
  inference(forward_demodulation,[],[f1735,f1257])).
fof(f1735,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3)))))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK31(X2,X1,X0)))),s(X0,sK33(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK30(X2,X1,X0)))),s(X0,sK30(X2,X1,X0)))))) )),
  inference(forward_demodulation,[],[f1267,f1257])).
fof(f1267,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK31(X2,X1,X0)))),s(X0,sK33(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK30(X2,X1,X0)))),s(X0,sK30(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f916])).
fof(f1262,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2))) = s(bool,happ(s(fun(fun(X0,bool),bool),happ(s(fun(X0,fun(fun(X0,bool),bool)),'const.bool.IN_0'),s(X0,X1))),s(fun(X0,bool),X2)))) )),
  inference(cnf_transformation,[],[f408])).
fof(f408,plain,(
  ! [X0,X1,X2] : s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2))) = s(bool,happ(s(fun(fun(X0,bool),bool),happ(s(fun(X0,fun(fun(X0,bool),bool)),'const.bool.IN_0'),s(X0,X1))),s(fun(X0,bool),X2)))),
  inference(flattening,[],[f407])).
fof(f407,plain,(
  ! [X0] : ! [X1,X2] : s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2))) = s(bool,happ(s(fun(fun(X0,bool),bool),happ(s(fun(X0,fun(fun(X0,bool),bool)),'const.bool.IN_0'),s(X0,X1))),s(fun(X0,bool),X2)))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X6] : ! [X21,X22] : s(bool,'const.bool.IN_2'(s(X6,X21),s(fun(X6,bool),X22))) = s(bool,happ(s(fun(fun(X6,bool),bool),happ(s(fun(X6,fun(fun(X6,bool),bool)),'const.bool.IN_0'),s(X6,X21))),s(fun(X6,bool),X22)))),
  file('thm.pred_set.EXTENSION','ARITY_0_2_const.bool.IN')).
fof(f1732,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'),s(X0,X1))),s(X0,X2)))) )),
  inference(forward_demodulation,[],[f1261,f1253])).
fof(f1261,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.relation.EMPTY_REL_2'(s(X0,X1),s(X0,X2))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'),s(X0,X1))),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f406])).
fof(f406,plain,(
  ! [X0,X1,X2] : s(bool,'const.relation.EMPTY_REL_2'(s(X0,X1),s(X0,X2))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'),s(X0,X1))),s(X0,X2)))),
  inference(flattening,[],[f405])).
fof(f405,plain,(
  ! [X0] : ! [X1,X2] : s(bool,'const.relation.EMPTY_REL_2'(s(X0,X1),s(X0,X2))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'),s(X0,X1))),s(X0,X2)))),
  inference(rectify,[],[f223])).
fof(f223,axiom,(
  ! [X6] : ! [X21,X22] : s(bool,'const.relation.EMPTY_REL_2'(s(X6,X21),s(X6,X22))) = s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),'const.relation.EMPTY_REL_0'),s(X6,X21))),s(X6,X22)))),
  file('thm.pred_set.EXTENSION','ARITY_0_2_const.relation.EMPTY_REL')).
fof(f1721,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK27(X1,X0)),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK26(X1,X0))))))) | p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(backward_demodulation,[],[f1257,f1245])).
fof(f1720,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK26(X1,X0)),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK25(X1,X0))))))) | p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(backward_demodulation,[],[f1257,f1244])).
fof(f1719,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,sK27(X1,X0)),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK25(X1,X0))))))) | p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(backward_demodulation,[],[f1257,f1246])).
fof(f1718,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,sK24(X1,X0)),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK23(X1,X0))))))) | p(s(bool,'const.relation.equivalence_1'(s(fun(X0,fun(X0,bool)),X1)))) | s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK23(X1,X0)))) != s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK24(X1,X0))))) )),
  inference(backward_demodulation,[],[f1257,f1242])).
fof(f1717,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK24(X1,X0)),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK23(X1,X0))))))) | p(s(bool,'const.relation.equivalence_1'(s(fun(X0,fun(X0,bool)),X1)))) | s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK23(X1,X0)))) = s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK24(X1,X0))))) )),
  inference(backward_demodulation,[],[f1257,f1241])).
fof(f1716,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK21(X1,X0)),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK22(X1,X0))))))) | p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(backward_demodulation,[],[f1257,f1237])).
fof(f1715,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK22(X1,X0)),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK21(X1,X0))))))) | p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(backward_demodulation,[],[f1257,f1236])).
fof(f1729,plain,(
  ( ! [X0,X1] : (s(bool,'const.bool.IN_2'(s(X0,sK19(X1,X0)),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK20(X1,X0)))))) != s(bool,'const.bool.IN_2'(s(X0,sK20(X1,X0)),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK19(X1,X0)))))) | p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(forward_demodulation,[],[f1714,f1257])).
fof(f1714,plain,(
  ( ! [X0,X1] : (s(bool,'const.bool.IN_2'(s(X0,sK20(X1,X0)),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK19(X1,X0)))))) != s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK20(X1,X0)))),s(X0,sK19(X1,X0)))) | p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(backward_demodulation,[],[f1257,f1234])).
fof(f1713,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK18(X1,X0)),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK18(X1,X0))))))) | p(s(bool,'const.relation.irreflexive_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(backward_demodulation,[],[f1257,f1232])).
fof(f1712,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,sK17(X1,X0)),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK17(X1,X0))))))) | p(s(bool,'const.relation.reflexive_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(backward_demodulation,[],[f1257,f1230])).
fof(f1711,plain,(
  ( ! [X0,X3,X1] : (~p(s(bool,'const.relation.reflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))))))) )),
  inference(backward_demodulation,[],[f1257,f1229])).
fof(f1710,plain,(
  ( ! [X0,X3,X1] : (~p(s(bool,'const.relation.irreflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))))))) )),
  inference(backward_demodulation,[],[f1257,f1231])).
fof(f1728,plain,(
  ( ! [X6,X0,X7,X5,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5)))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X7),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X6)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X7),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5)))))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(forward_demodulation,[],[f1727,f1257])).
fof(f1727,plain,(
  ( ! [X6,X0,X7,X5,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X7),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X6)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X7),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5)))))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X6)))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(forward_demodulation,[],[f1709,f1257])).
fof(f1709,plain,(
  ( ! [X6,X0,X7,X5,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X7),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5)))))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X6))),s(X0,X7)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X6)))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(backward_demodulation,[],[f1257,f1243])).
fof(f1708,plain,(
  ( ! [X4,X0,X5,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4)))))) | s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))) != s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))) | ~p(s(bool,'const.relation.equivalence_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(backward_demodulation,[],[f1257,f1240])).
fof(f1707,plain,(
  ( ! [X4,X0,X5,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4)))))) | s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))) = s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))) | ~p(s(bool,'const.relation.equivalence_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(backward_demodulation,[],[f1257,f1239])).
fof(f1726,plain,(
  ( ! [X4,X0,X5,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4)))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5)))))) | s(X0,X4) = s(X0,X5) | ~p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(forward_demodulation,[],[f1706,f1257])).
fof(f1706,plain,(
  ( ! [X4,X0,X5,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5)))))) | s(X0,X4) = s(X0,X5) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X5)))) | ~p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(backward_demodulation,[],[f1257,f1235])).
fof(f1725,plain,(
  ( ! [X4,X0,X5,X1] : (s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))))) = s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))))) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(forward_demodulation,[],[f1705,f1257])).
fof(f1705,plain,(
  ( ! [X4,X0,X5,X1] : (s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X4))) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(backward_demodulation,[],[f1257,f1233])).
fof(f1724,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3)))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2)))))) | s(X0,X2) = s(X0,X3) | ~p(s(bool,'const.relation.WeakOrder_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(forward_demodulation,[],[f1704,f1257])).
fof(f1704,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2)))))) | s(X0,X2) = s(X0,X3) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2)))) | ~p(s(bool,'const.relation.WeakOrder_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(backward_demodulation,[],[f1257,f1215])).
fof(f1703,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3)))))) | s(X0,X2) != s(X0,X3) | ~p(s(bool,'const.relation.WeakOrder_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(backward_demodulation,[],[f1257,f1214])).
fof(f1256,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2))))) )),
  inference(cnf_transformation,[],[f396])).
fof(f396,plain,(
  ! [X0,X1,X2] : p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2))))),
  inference(flattening,[],[f395])).
fof(f395,plain,(
  ! [X0] : ! [X1,X2] : p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2))))),
  inference(rectify,[],[f258])).
fof(f258,axiom,(
  ! [X6] : ! [X18,X9] : p(s(bool,'const.relation.RTC_3'(s(fun(X6,fun(X6,bool)),X18),s(X6,X9),s(X6,X9))))),
  file('thm.pred_set.EXTENSION','thm.relation.RTC_RULES|split|0')).
fof(f1255,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f394])).
fof(f394,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f393])).
fof(f393,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f128])).
fof(f128,axiom,(
  ! [X6] : ! [X10,X11] : s(X6,X10) = s(X6,'const.bool.COND_3'(s(bool,'T_0'),s(X6,X10),s(X6,X11)))),
  file('thm.pred_set.EXTENSION','thm.bool.COND_CLAUSES|split|0')).
fof(f1254,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f392])).
fof(f392,plain,(
  ! [X0,X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f391])).
fof(f391,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f129])).
fof(f129,axiom,(
  ! [X6] : ! [X10,X11] : s(X6,X11) = s(X6,'const.bool.COND_3'(s(bool,'F_0'),s(X6,X10),s(X6,X11)))),
  file('thm.pred_set.EXTENSION','thm.bool.COND_CLAUSES|split|1')).
fof(f1253,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,'const.relation.EMPTY_REL_2'(s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f390])).
fof(f390,plain,(
  ! [X0,X1,X2] : s(bool,'F_0') = s(bool,'const.relation.EMPTY_REL_2'(s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f389])).
fof(f389,plain,(
  ! [X0] : ! [X1,X2] : s(bool,'F_0') = s(bool,'const.relation.EMPTY_REL_2'(s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f222])).
fof(f222,axiom,(
  ! [X6] : ! [X9,X19] : s(bool,'F_0') = s(bool,'const.relation.EMPTY_REL_2'(s(X6,X9),s(X6,X19)))),
  file('thm.pred_set.EXTENSION','thm.relation.EMPTY_REL_DEF')).
fof(f1247,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X0,bool)),X1),s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0')))) | s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0') = s(fun(X0,fun(X0,bool)),X1)) )),
  inference(cnf_transformation,[],[f909])).
fof(f909,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X0,bool)),X1),s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0')))) | s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0') = s(fun(X0,fun(X0,bool)),X1)) & (s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0') != s(fun(X0,fun(X0,bool)),X1) | p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X0,bool)),X1),s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'))))))),
  inference(nnf_transformation,[],[f385])).
fof(f385,plain,(
  ! [X0,X1] : (p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X0,bool)),X1),s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0')))) <=> s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0') = s(fun(X0,fun(X0,bool)),X1))),
  inference(flattening,[],[f384])).
fof(f384,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X0,bool)),X1),s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0')))) <=> s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0') = s(fun(X0,fun(X0,bool)),X1))),
  inference(rectify,[],[f215])).
fof(f215,axiom,(
  ! [X6] : ! [X18] : (p(s(bool,'const.relation.RSUBSET_2'(s(fun(X6,fun(X6,bool)),X18),s(fun(X6,fun(X6,bool)),'const.relation.EMPTY_REL_0')))) <=> s(fun(X6,fun(X6,bool)),'const.relation.EMPTY_REL_0') = s(fun(X6,fun(X6,bool)),X18))),
  file('thm.pred_set.EXTENSION','thm.relation.REMPTY_SUBSET|split|1')).
fof(f1248,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X0,bool)),X1),s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0')))) | s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0') != s(fun(X0,fun(X0,bool)),X1)) )),
  inference(cnf_transformation,[],[f909])).
fof(f1243,plain,(
  ( ! [X6,X0,X7,X5,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X7)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X6))),s(X0,X7)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X6)))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(cnf_transformation,[],[f908])).
fof(f908,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1)))) | ! [X5,X6,X7] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X6)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X6))),s(X0,X7)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X7)))))) & ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK25(X1,X0)))),s(X0,sK26(X1,X0))))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK26(X1,X0)))),s(X0,sK27(X1,X0))))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK25(X1,X0)))),s(X0,sK27(X1,X0)))))) | p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK25,sK26,sK27])],[f907])).
fof(f907,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1)))) | ! [X5,X6,X7] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X6)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X6))),s(X0,X7)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X7)))))) & (? [X2,X3,X4] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X4)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X4))))) | p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(rectify,[],[f906])).
fof(f906,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1)))) | ! [X2,X3,X4] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X4)))))) & (? [X2,X3,X4] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X4)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X4))))) | p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(nnf_transformation,[],[f750])).
fof(f750,plain,(
  ! [X0,X1] : (p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> ! [X2,X3,X4] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X4))))))),
  inference(flattening,[],[f749])).
fof(f749,plain,(
  ! [X0,X1] : (p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> ! [X2,X3,X4] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X4))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X4))))))),
  inference(ennf_transformation,[],[f383])).
fof(f383,plain,(
  ! [X0,X1] : (p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> ! [X2,X3,X4] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X4))))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X4))))))),
  inference(flattening,[],[f382])).
fof(f382,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> ! [X2,X3,X4] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X4))))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X4))))))),
  inference(rectify,[],[f89])).
fof(f89,axiom,(
  ! [X6] : ! [X18] : (p(s(bool,'const.relation.transitive_1'(s(fun(X6,fun(X6,bool)),X18)))) <=> ! [X9,X19,X49] : ((p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X18),s(X6,X9))),s(X6,X19)))) & p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X18),s(X6,X19))),s(X6,X49))))) => p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X18),s(X6,X9))),s(X6,X49))))))),
  file('thm.pred_set.EXTENSION','thm.relation.transitive_def')).
fof(f1244,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK25(X1,X0)))),s(X0,sK26(X1,X0)))))) )),
  inference(cnf_transformation,[],[f908])).
fof(f1245,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK26(X1,X0)))),s(X0,sK27(X1,X0)))))) )),
  inference(cnf_transformation,[],[f908])).
fof(f1246,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK25(X1,X0)))),s(X0,sK27(X1,X0)))))) )),
  inference(cnf_transformation,[],[f908])).
fof(f1239,plain,(
  ( ! [X4,X0,X5,X1] : (s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))) = s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X5)))) | ~p(s(bool,'const.relation.equivalence_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(cnf_transformation,[],[f905])).
fof(f905,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.relation.equivalence_1'(s(fun(X0,fun(X0,bool)),X1)))) | ! [X4,X5] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X5)))) | s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))) = s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5)))) & (s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))) != s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X5))))))) & (((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK23(X1,X0)))),s(X0,sK24(X1,X0))))) | s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK23(X1,X0)))) = s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK24(X1,X0))))) & (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK23(X1,X0)))),s(X0,sK24(X1,X0))))) | s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK23(X1,X0)))) != s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK24(X1,X0)))))) | p(s(bool,'const.relation.equivalence_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK23,sK24])],[f904])).
fof(f904,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.relation.equivalence_1'(s(fun(X0,fun(X0,bool)),X1)))) | ! [X4,X5] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X5)))) | s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))) = s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5)))) & (s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))) != s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X5))))))) & (? [X2,X3] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))) = s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3)))) & (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))) != s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))))) | p(s(bool,'const.relation.equivalence_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(rectify,[],[f903])).
fof(f903,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.relation.equivalence_1'(s(fun(X0,fun(X0,bool)),X1)))) | ! [X2,X3] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))) = s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3)))) & (s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))) != s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))))))) & (? [X2,X3] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))) = s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3)))) & (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))) != s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))))) | p(s(bool,'const.relation.equivalence_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(nnf_transformation,[],[f381])).
fof(f381,plain,(
  ! [X0,X1] : (p(s(bool,'const.relation.equivalence_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> ! [X2,X3] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) <=> s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))) = s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3)))))),
  inference(flattening,[],[f380])).
fof(f380,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.relation.equivalence_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> ! [X2,X3] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) <=> s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))) = s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3)))))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X6] : ! [X18] : (p(s(bool,'const.relation.equivalence_1'(s(fun(X6,fun(X6,bool)),X18)))) <=> ! [X9,X19] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X18),s(X6,X9))),s(X6,X19)))) <=> s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X18),s(X6,X9))) = s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X18),s(X6,X19)))))),
  file('thm.pred_set.EXTENSION','thm.relation.ALT_equivalence')).
fof(f1240,plain,(
  ( ! [X4,X0,X5,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X5)))) | s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))) != s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))) | ~p(s(bool,'const.relation.equivalence_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(cnf_transformation,[],[f905])).
fof(f1241,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.relation.equivalence_1'(s(fun(X0,fun(X0,bool)),X1)))) | s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK23(X1,X0)))) = s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK24(X1,X0)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK23(X1,X0)))),s(X0,sK24(X1,X0)))))) )),
  inference(cnf_transformation,[],[f905])).
fof(f1242,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.relation.equivalence_1'(s(fun(X0,fun(X0,bool)),X1)))) | s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK23(X1,X0)))) != s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK24(X1,X0)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK23(X1,X0)))),s(X0,sK24(X1,X0)))))) )),
  inference(cnf_transformation,[],[f905])).
fof(f1235,plain,(
  ( ! [X4,X0,X5,X1] : (s(X0,X4) = s(X0,X5) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X4)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X5)))) | ~p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(cnf_transformation,[],[f902])).
fof(f902,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) | ! [X4,X5] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X5)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X4)))) | s(X0,X4) = s(X0,X5))) & ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK21(X1,X0)))),s(X0,sK22(X1,X0))))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK22(X1,X0)))),s(X0,sK21(X1,X0))))) & s(X0,sK21(X1,X0)) != s(X0,sK22(X1,X0))) | p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK21,sK22])],[f901])).
fof(f901,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) | ! [X4,X5] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X5)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X4)))) | s(X0,X4) = s(X0,X5))) & (? [X2,X3] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2)))) & s(X0,X2) != s(X0,X3)) | p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(rectify,[],[f900])).
fof(f900,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) | ! [X2,X3] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2)))) | s(X0,X2) = s(X0,X3))) & (? [X2,X3] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2)))) & s(X0,X2) != s(X0,X3)) | p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(nnf_transformation,[],[f748])).
fof(f748,plain,(
  ! [X0,X1] : (p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> ! [X2,X3] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2)))) | s(X0,X2) = s(X0,X3)))),
  inference(flattening,[],[f747])).
fof(f747,plain,(
  ! [X0,X1] : (p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> ! [X2,X3] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2))))) | s(X0,X2) = s(X0,X3)))),
  inference(ennf_transformation,[],[f379])).
fof(f379,plain,(
  ! [X0,X1] : (p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> ! [X2,X3] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2))))) => s(X0,X2) = s(X0,X3)))),
  inference(flattening,[],[f378])).
fof(f378,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> ! [X2,X3] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2))))) => s(X0,X2) = s(X0,X3)))),
  inference(rectify,[],[f238])).
fof(f238,axiom,(
  ! [X6] : ! [X18] : (p(s(bool,'const.relation.antisymmetric_1'(s(fun(X6,fun(X6,bool)),X18)))) <=> ! [X9,X19] : ((p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X18),s(X6,X9))),s(X6,X19)))) & p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X18),s(X6,X19))),s(X6,X9))))) => s(X6,X9) = s(X6,X19)))),
  file('thm.pred_set.EXTENSION','thm.relation.antisymmetric_def')).
fof(f1236,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK21(X1,X0)))),s(X0,sK22(X1,X0)))))) )),
  inference(cnf_transformation,[],[f902])).
fof(f1237,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK22(X1,X0)))),s(X0,sK21(X1,X0)))))) )),
  inference(cnf_transformation,[],[f902])).
fof(f1238,plain,(
  ( ! [X0,X1] : (s(X0,sK21(X1,X0)) != s(X0,sK22(X1,X0)) | p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(cnf_transformation,[],[f902])).
fof(f1233,plain,(
  ( ! [X4,X0,X5,X1] : (s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X5))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X4))) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(cnf_transformation,[],[f899])).
fof(f899,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) | ! [X4,X5] : s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X5))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X4)))) & (s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK19(X1,X0)))),s(X0,sK20(X1,X0)))) != s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK20(X1,X0)))),s(X0,sK19(X1,X0)))) | p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK19,sK20])],[f898])).
fof(f898,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) | ! [X4,X5] : s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X5))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X4)))) & (? [X2,X3] : s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))) != s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2))) | p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(rectify,[],[f897])).
fof(f897,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) | ! [X2,X3] : s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2)))) & (? [X2,X3] : s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))) != s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2))) | p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(nnf_transformation,[],[f377])).
fof(f377,plain,(
  ! [X0,X1] : (p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> ! [X2,X3] : s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2))))),
  inference(flattening,[],[f376])).
fof(f376,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> ! [X2,X3] : s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2))))),
  inference(rectify,[],[f87])).
fof(f87,axiom,(
  ! [X6] : ! [X18] : (p(s(bool,'const.relation.symmetric_1'(s(fun(X6,fun(X6,bool)),X18)))) <=> ! [X9,X19] : s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X18),s(X6,X9))),s(X6,X19))) = s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X18),s(X6,X19))),s(X6,X9))))),
  file('thm.pred_set.EXTENSION','thm.relation.symmetric_def')).
fof(f1234,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) | s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK19(X1,X0)))),s(X0,sK20(X1,X0)))) != s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK20(X1,X0)))),s(X0,sK19(X1,X0))))) )),
  inference(cnf_transformation,[],[f899])).
fof(f1231,plain,(
  ( ! [X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X3)))) | ~p(s(bool,'const.relation.irreflexive_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(cnf_transformation,[],[f896])).
fof(f896,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.relation.irreflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) | ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X3))))) & (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK18(X1,X0)))),s(X0,sK18(X1,X0))))) | p(s(bool,'const.relation.irreflexive_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK18])],[f895])).
fof(f895,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.relation.irreflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) | ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X3))))) & (? [X2] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X2)))) | p(s(bool,'const.relation.irreflexive_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(rectify,[],[f894])).
fof(f894,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.relation.irreflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) | ! [X2] : ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X2))))) & (? [X2] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X2)))) | p(s(bool,'const.relation.irreflexive_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(nnf_transformation,[],[f375])).
fof(f375,plain,(
  ! [X0,X1] : (p(s(bool,'const.relation.irreflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> ! [X2] : ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X2)))))),
  inference(flattening,[],[f374])).
fof(f374,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.relation.irreflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> ! [X2] : ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X2)))))),
  inference(rectify,[],[f232])).
fof(f232,axiom,(
  ! [X6] : ! [X18] : (p(s(bool,'const.relation.irreflexive_1'(s(fun(X6,fun(X6,bool)),X18)))) <=> ! [X9] : ~p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X18),s(X6,X9))),s(X6,X9)))))),
  file('thm.pred_set.EXTENSION','thm.relation.irreflexive_def')).
fof(f1232,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.relation.irreflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK18(X1,X0)))),s(X0,sK18(X1,X0)))))) )),
  inference(cnf_transformation,[],[f896])).
fof(f1229,plain,(
  ( ! [X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X3)))) | ~p(s(bool,'const.relation.reflexive_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(cnf_transformation,[],[f893])).
fof(f893,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.relation.reflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) | ! [X3] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X3))))) & (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK17(X1,X0)))),s(X0,sK17(X1,X0))))) | p(s(bool,'const.relation.reflexive_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK17])],[f892])).
fof(f892,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.relation.reflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) | ! [X3] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X3))))) & (? [X2] : ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X2)))) | p(s(bool,'const.relation.reflexive_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(rectify,[],[f891])).
fof(f891,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.relation.reflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) | ! [X2] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X2))))) & (? [X2] : ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X2)))) | p(s(bool,'const.relation.reflexive_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(nnf_transformation,[],[f373])).
fof(f373,plain,(
  ! [X0,X1] : (p(s(bool,'const.relation.reflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> ! [X2] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X2)))))),
  inference(flattening,[],[f372])).
fof(f372,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.relation.reflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> ! [X2] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X2)))))),
  inference(rectify,[],[f91])).
fof(f91,axiom,(
  ! [X6] : ! [X18] : (p(s(bool,'const.relation.reflexive_1'(s(fun(X6,fun(X6,bool)),X18)))) <=> ! [X9] : p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X18),s(X6,X9))),s(X6,X9)))))),
  file('thm.pred_set.EXTENSION','thm.relation.reflexive_def')).
fof(f1230,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.relation.reflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK17(X1,X0)))),s(X0,sK17(X1,X0)))))) )),
  inference(cnf_transformation,[],[f893])).
fof(f1225,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.relation.equivalence_1'(s(fun(X0,fun(X0,bool)),X1)))) | p(s(bool,'const.relation.reflexive_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(cnf_transformation,[],[f890])).
fof(f890,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.relation.equivalence_1'(s(fun(X0,fun(X0,bool)),X1)))) | (p(s(bool,'const.relation.reflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) & p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) & p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1)))))) & (~p(s(bool,'const.relation.reflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1)))) | p(s(bool,'const.relation.equivalence_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(flattening,[],[f889])).
fof(f889,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.relation.equivalence_1'(s(fun(X0,fun(X0,bool)),X1)))) | (p(s(bool,'const.relation.reflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) & p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) & p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1)))))) & ((~p(s(bool,'const.relation.reflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1))))) | p(s(bool,'const.relation.equivalence_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(nnf_transformation,[],[f371])).
fof(f371,plain,(
  ! [X0,X1] : (p(s(bool,'const.relation.equivalence_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> (p(s(bool,'const.relation.reflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) & p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) & p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(flattening,[],[f370])).
fof(f370,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.relation.equivalence_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> (p(s(bool,'const.relation.reflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) & p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) & p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(rectify,[],[f90])).
fof(f90,axiom,(
  ! [X6] : ! [X18] : (p(s(bool,'const.relation.equivalence_1'(s(fun(X6,fun(X6,bool)),X18)))) <=> (p(s(bool,'const.relation.reflexive_1'(s(fun(X6,fun(X6,bool)),X18)))) & p(s(bool,'const.relation.symmetric_1'(s(fun(X6,fun(X6,bool)),X18)))) & p(s(bool,'const.relation.transitive_1'(s(fun(X6,fun(X6,bool)),X18))))))),
  file('thm.pred_set.EXTENSION','thm.relation.equivalence_def')).
fof(f1226,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.relation.equivalence_1'(s(fun(X0,fun(X0,bool)),X1)))) | p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(cnf_transformation,[],[f890])).
fof(f1227,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.relation.equivalence_1'(s(fun(X0,fun(X0,bool)),X1)))) | p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(cnf_transformation,[],[f890])).
fof(f1228,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.relation.reflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1)))) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) | p(s(bool,'const.relation.equivalence_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(cnf_transformation,[],[f890])).
fof(f1221,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.relation.WeakOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) | p(s(bool,'const.relation.reflexive_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(cnf_transformation,[],[f888])).
fof(f888,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.relation.WeakOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) | (p(s(bool,'const.relation.reflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) & p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) & p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1)))))) & (~p(s(bool,'const.relation.reflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) | ~p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1)))) | p(s(bool,'const.relation.WeakOrder_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(flattening,[],[f887])).
fof(f887,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.relation.WeakOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) | (p(s(bool,'const.relation.reflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) & p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) & p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1)))))) & ((~p(s(bool,'const.relation.reflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) | ~p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1))))) | p(s(bool,'const.relation.WeakOrder_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(nnf_transformation,[],[f369])).
fof(f369,plain,(
  ! [X0,X1] : (p(s(bool,'const.relation.WeakOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> (p(s(bool,'const.relation.reflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) & p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) & p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(flattening,[],[f368])).
fof(f368,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.relation.WeakOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> (p(s(bool,'const.relation.reflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) & p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) & p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(rectify,[],[f237])).
fof(f237,axiom,(
  ! [X67] : ! [X68] : (p(s(bool,'const.relation.WeakOrder_1'(s(fun(X67,fun(X67,bool)),X68)))) <=> (p(s(bool,'const.relation.reflexive_1'(s(fun(X67,fun(X67,bool)),X68)))) & p(s(bool,'const.relation.antisymmetric_1'(s(fun(X67,fun(X67,bool)),X68)))) & p(s(bool,'const.relation.transitive_1'(s(fun(X67,fun(X67,bool)),X68))))))),
  file('thm.pred_set.EXTENSION','thm.relation.WeakOrder')).
fof(f1222,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.relation.WeakOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) | p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(cnf_transformation,[],[f888])).
fof(f1223,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.relation.WeakOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) | p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(cnf_transformation,[],[f888])).
fof(f1224,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1)))) | p(s(bool,'const.relation.WeakOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) | ~p(s(bool,'const.relation.reflexive_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(cnf_transformation,[],[f888])).
fof(f1218,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.relation.StrongOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) | p(s(bool,'const.relation.irreflexive_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(cnf_transformation,[],[f886])).
fof(f886,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.relation.StrongOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) | (p(s(bool,'const.relation.irreflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) & p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1)))))) & (~p(s(bool,'const.relation.irreflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1)))) | p(s(bool,'const.relation.StrongOrder_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(flattening,[],[f885])).
fof(f885,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.relation.StrongOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) | (p(s(bool,'const.relation.irreflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) & p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1)))))) & ((~p(s(bool,'const.relation.irreflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1))))) | p(s(bool,'const.relation.StrongOrder_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(nnf_transformation,[],[f367])).
fof(f367,plain,(
  ! [X0,X1] : (p(s(bool,'const.relation.StrongOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> (p(s(bool,'const.relation.irreflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) & p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(flattening,[],[f366])).
fof(f366,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.relation.StrongOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> (p(s(bool,'const.relation.irreflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) & p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(rectify,[],[f231])).
fof(f231,axiom,(
  ! [X67] : ! [X68] : (p(s(bool,'const.relation.StrongOrder_1'(s(fun(X67,fun(X67,bool)),X68)))) <=> (p(s(bool,'const.relation.irreflexive_1'(s(fun(X67,fun(X67,bool)),X68)))) & p(s(bool,'const.relation.transitive_1'(s(fun(X67,fun(X67,bool)),X68))))))),
  file('thm.pred_set.EXTENSION','thm.relation.StrongOrder')).
fof(f1219,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.relation.StrongOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) | p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(cnf_transformation,[],[f886])).
fof(f1220,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.relation.irreflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1)))) | p(s(bool,'const.relation.StrongOrder_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(cnf_transformation,[],[f886])).
fof(f1214,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2)))) | s(X0,X2) != s(X0,X3) | ~p(s(bool,'const.relation.WeakOrder_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(cnf_transformation,[],[f883])).
fof(f883,plain,(
  ! [X0,X1] : (~p(s(bool,'const.relation.WeakOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) | ! [X2,X3] : ((s(X0,X2) != s(X0,X3) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2)))))) & (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2)))) | s(X0,X2) = s(X0,X3))))),
  inference(flattening,[],[f882])).
fof(f882,plain,(
  ! [X0,X1] : (~p(s(bool,'const.relation.WeakOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) | ! [X2,X3] : ((s(X0,X2) != s(X0,X3) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2)))))) & ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2))))) | s(X0,X2) = s(X0,X3))))),
  inference(nnf_transformation,[],[f746])).
fof(f746,plain,(
  ! [X0,X1] : (~p(s(bool,'const.relation.WeakOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) | ! [X2,X3] : (s(X0,X2) = s(X0,X3) <=> (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2)))))))),
  inference(ennf_transformation,[],[f363])).
fof(f363,plain,(
  ! [X0,X1] : (p(s(bool,'const.relation.WeakOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) => ! [X2,X3] : (s(X0,X2) = s(X0,X3) <=> (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2)))))))),
  inference(flattening,[],[f362])).
fof(f362,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.relation.WeakOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) => ! [X2,X3] : (s(X0,X2) = s(X0,X3) <=> (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2)))))))),
  inference(rectify,[],[f221])).
fof(f221,axiom,(
  ! [X6] : ! [X18] : (p(s(bool,'const.relation.WeakOrder_1'(s(fun(X6,fun(X6,bool)),X18)))) => ! [X19,X49] : (s(X6,X19) = s(X6,X49) <=> (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X18),s(X6,X19))),s(X6,X49)))) & p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X18),s(X6,X49))),s(X6,X19)))))))),
  file('thm.pred_set.EXTENSION','thm.relation.WeakOrder_EQ')).
fof(f1215,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X2) = s(X0,X3) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | ~p(s(bool,'const.relation.WeakOrder_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(cnf_transformation,[],[f883])).
fof(f1212,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.relation.WeakOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) | s(fun(X0,fun(X0,bool)),X1) = s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),'const.relation.STRORD_1'(s(fun(X0,fun(X0,bool)),X1)))))) )),
  inference(cnf_transformation,[],[f745])).
fof(f745,plain,(
  ! [X0,X1] : (~p(s(bool,'const.relation.WeakOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) | s(fun(X0,fun(X0,bool)),X1) = s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),'const.relation.STRORD_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(ennf_transformation,[],[f361])).
fof(f361,plain,(
  ! [X0,X1] : (p(s(bool,'const.relation.WeakOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) => s(fun(X0,fun(X0,bool)),X1) = s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),'const.relation.STRORD_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(flattening,[],[f360])).
fof(f360,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.relation.WeakOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) => s(fun(X0,fun(X0,bool)),X1) = s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),'const.relation.STRORD_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(rectify,[],[f233])).
fof(f233,axiom,(
  ! [X6] : ! [X18] : (p(s(bool,'const.relation.WeakOrder_1'(s(fun(X6,fun(X6,bool)),X18)))) => s(fun(X6,fun(X6,bool)),X18) = s(fun(X6,fun(X6,bool)),'const.relation.RC_1'(s(fun(X6,fun(X6,bool)),'const.relation.STRORD_1'(s(fun(X6,fun(X6,bool)),X18))))))),
  file('thm.pred_set.EXTENSION','thm.relation.RC_STRORD')).
fof(f1211,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.relation.StrongOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) | s(fun(X0,fun(X0,bool)),X1) = s(fun(X0,fun(X0,bool)),'const.relation.STRORD_1'(s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),X1)))))) )),
  inference(cnf_transformation,[],[f744])).
fof(f744,plain,(
  ! [X0,X1] : (~p(s(bool,'const.relation.StrongOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) | s(fun(X0,fun(X0,bool)),X1) = s(fun(X0,fun(X0,bool)),'const.relation.STRORD_1'(s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(ennf_transformation,[],[f359])).
fof(f359,plain,(
  ! [X0,X1] : (p(s(bool,'const.relation.StrongOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) => s(fun(X0,fun(X0,bool)),X1) = s(fun(X0,fun(X0,bool)),'const.relation.STRORD_1'(s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(flattening,[],[f358])).
fof(f358,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.relation.StrongOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) => s(fun(X0,fun(X0,bool)),X1) = s(fun(X0,fun(X0,bool)),'const.relation.STRORD_1'(s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(rectify,[],[f227])).
fof(f227,axiom,(
  ! [X6] : ! [X18] : (p(s(bool,'const.relation.StrongOrder_1'(s(fun(X6,fun(X6,bool)),X18)))) => s(fun(X6,fun(X6,bool)),X18) = s(fun(X6,fun(X6,bool)),'const.relation.STRORD_1'(s(fun(X6,fun(X6,bool)),'const.relation.RC_1'(s(fun(X6,fun(X6,bool)),X18))))))),
  file('thm.pred_set.EXTENSION','thm.relation.STRORD_RC')).
fof(f1210,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1)))) | s(fun(X0,fun(X0,bool)),X1) = s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),X1)))) )),
  inference(cnf_transformation,[],[f743])).
fof(f743,plain,(
  ! [X0,X1] : (~p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1)))) | s(fun(X0,fun(X0,bool)),X1) = s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),X1))))),
  inference(ennf_transformation,[],[f357])).
fof(f357,plain,(
  ! [X0,X1] : (p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1)))) => s(fun(X0,fun(X0,bool)),X1) = s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),X1))))),
  inference(flattening,[],[f356])).
fof(f356,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1)))) => s(fun(X0,fun(X0,bool)),X1) = s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),X1))))),
  inference(rectify,[],[f207])).
fof(f207,axiom,(
  ! [X6] : ! [X18] : (p(s(bool,'const.relation.transitive_1'(s(fun(X6,fun(X6,bool)),X18)))) => s(fun(X6,fun(X6,bool)),X18) = s(fun(X6,fun(X6,bool)),'const.relation.TC_1'(s(fun(X6,fun(X6,bool)),X18))))),
  file('thm.pred_set.EXTENSION','thm.relation.transitive_TC_identity')).
fof(f1209,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) | s(fun(X0,fun(X0,bool)),X1) = s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))) )),
  inference(cnf_transformation,[],[f742])).
fof(f742,plain,(
  ! [X0,X1] : (~p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) | s(fun(X0,fun(X0,bool)),X1) = s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1))))),
  inference(ennf_transformation,[],[f355])).
fof(f355,plain,(
  ! [X0,X1] : (p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) => s(fun(X0,fun(X0,bool)),X1) = s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1))))),
  inference(flattening,[],[f354])).
fof(f354,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) => s(fun(X0,fun(X0,bool)),X1) = s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1))))),
  inference(rectify,[],[f208])).
fof(f208,axiom,(
  ! [X6] : ! [X18] : (p(s(bool,'const.relation.symmetric_1'(s(fun(X6,fun(X6,bool)),X18)))) => s(fun(X6,fun(X6,bool)),X18) = s(fun(X6,fun(X6,bool)),'const.relation.SC_1'(s(fun(X6,fun(X6,bool)),X18))))),
  file('thm.pred_set.EXTENSION','thm.relation.symmetric_SC_identity')).
fof(f1208,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) | s(fun(X0,fun(X0,bool)),X1) = s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),X1)))) )),
  inference(cnf_transformation,[],[f741])).
fof(f741,plain,(
  ! [X0,X1] : (~p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) | s(fun(X0,fun(X0,bool)),X1) = s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),X1))))),
  inference(ennf_transformation,[],[f353])).
fof(f353,plain,(
  ! [X0,X1] : (p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) => s(fun(X0,fun(X0,bool)),X1) = s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),X1))))),
  inference(flattening,[],[f352])).
fof(f352,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) => s(fun(X0,fun(X0,bool)),X1) = s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),X1))))),
  inference(rectify,[],[f235])).
fof(f235,axiom,(
  ! [X6] : ! [X18] : (p(s(bool,'const.relation.symmetric_1'(s(fun(X6,fun(X6,bool)),X18)))) => s(fun(X6,fun(X6,bool)),X18) = s(fun(X6,fun(X6,bool)),'const.relation.inv_1'(s(fun(X6,fun(X6,bool)),X18))))),
  file('thm.pred_set.EXTENSION','thm.relation.symmetric_inv_identity')).
fof(f1207,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.relation.reflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) | s(fun(X0,fun(X0,bool)),X1) = s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),X1)))) )),
  inference(cnf_transformation,[],[f740])).
fof(f740,plain,(
  ! [X0,X1] : (~p(s(bool,'const.relation.reflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) | s(fun(X0,fun(X0,bool)),X1) = s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),X1))))),
  inference(ennf_transformation,[],[f351])).
fof(f351,plain,(
  ! [X0,X1] : (p(s(bool,'const.relation.reflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) => s(fun(X0,fun(X0,bool)),X1) = s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),X1))))),
  inference(flattening,[],[f350])).
fof(f350,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.relation.reflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) => s(fun(X0,fun(X0,bool)),X1) = s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),X1))))),
  inference(rectify,[],[f206])).
fof(f206,axiom,(
  ! [X6] : ! [X18] : (p(s(bool,'const.relation.reflexive_1'(s(fun(X6,fun(X6,bool)),X18)))) => s(fun(X6,fun(X6,bool)),X18) = s(fun(X6,fun(X6,bool)),'const.relation.RC_1'(s(fun(X6,fun(X6,bool)),X18))))),
  file('thm.pred_set.EXTENSION','thm.relation.reflexive_RC_identity')).
fof(f1206,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.relation.equivalence_1'(s(fun(X0,fun(X0,bool)),X1)))) | s(fun(X0,fun(X0,bool)),X1) = s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),X1)))) )),
  inference(cnf_transformation,[],[f739])).
fof(f739,plain,(
  ! [X0,X1] : (~p(s(bool,'const.relation.equivalence_1'(s(fun(X0,fun(X0,bool)),X1)))) | s(fun(X0,fun(X0,bool)),X1) = s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),X1))))),
  inference(ennf_transformation,[],[f349])).
fof(f349,plain,(
  ! [X0,X1] : (p(s(bool,'const.relation.equivalence_1'(s(fun(X0,fun(X0,bool)),X1)))) => s(fun(X0,fun(X0,bool)),X1) = s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),X1))))),
  inference(flattening,[],[f348])).
fof(f348,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.relation.equivalence_1'(s(fun(X0,fun(X0,bool)),X1)))) => s(fun(X0,fun(X0,bool)),X1) = s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),X1))))),
  inference(rectify,[],[f236])).
fof(f236,axiom,(
  ! [X6] : ! [X18] : (p(s(bool,'const.relation.equivalence_1'(s(fun(X6,fun(X6,bool)),X18)))) => s(fun(X6,fun(X6,bool)),X18) = s(fun(X6,fun(X6,bool)),'const.relation.inv_1'(s(fun(X6,fun(X6,bool)),X18))))),
  file('thm.pred_set.EXTENSION','thm.relation.equivalence_inv_identity')).
fof(f1204,plain,(
  ( ! [X0,X1] : (s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),X1))))) = s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),X1)))))) )),
  inference(cnf_transformation,[],[f345])).
fof(f345,plain,(
  ! [X0,X1] : s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),X1))))) = s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),X1)))))),
  inference(flattening,[],[f344])).
fof(f344,plain,(
  ! [X0] : ! [X1] : s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),X1))))) = s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),X1)))))),
  inference(rectify,[],[f247])).
fof(f247,axiom,(
  ! [X6] : ! [X18] : s(fun(X6,fun(X6,bool)),'const.relation.inv_1'(s(fun(X6,fun(X6,bool)),'const.relation.RC_1'(s(fun(X6,fun(X6,bool)),X18))))) = s(fun(X6,fun(X6,bool)),'const.relation.RC_1'(s(fun(X6,fun(X6,bool)),'const.relation.inv_1'(s(fun(X6,fun(X6,bool)),X18)))))),
  file('thm.pred_set.EXTENSION','thm.relation.inv_RC')).
fof(f1203,plain,(
  ( ! [X0,X1] : (s(fun(X0,fun(X0,bool)),'const.relation.EQC_1'(s(fun(X0,fun(X0,bool)),X1))) = s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))))))) )),
  inference(cnf_transformation,[],[f343])).
fof(f343,plain,(
  ! [X0,X1] : s(fun(X0,fun(X0,bool)),'const.relation.EQC_1'(s(fun(X0,fun(X0,bool)),X1))) = s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))))))),
  inference(flattening,[],[f342])).
fof(f342,plain,(
  ! [X0] : ! [X1] : s(fun(X0,fun(X0,bool)),'const.relation.EQC_1'(s(fun(X0,fun(X0,bool)),X1))) = s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))))))),
  inference(rectify,[],[f242])).
fof(f242,axiom,(
  ! [X6] : ! [X18] : s(fun(X6,fun(X6,bool)),'const.relation.EQC_1'(s(fun(X6,fun(X6,bool)),X18))) = s(fun(X6,fun(X6,bool)),'const.relation.RC_1'(s(fun(X6,fun(X6,bool)),'const.relation.TC_1'(s(fun(X6,fun(X6,bool)),'const.relation.SC_1'(s(fun(X6,fun(X6,bool)),X18)))))))),
  file('thm.pred_set.EXTENSION','thm.relation.EQC_DEF')).
fof(f1202,plain,(
  ( ! [X0,X1] : (s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),X1))))) = s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),X1)))))) )),
  inference(cnf_transformation,[],[f341])).
fof(f341,plain,(
  ! [X0,X1] : s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),X1))))) = s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),X1)))))),
  inference(flattening,[],[f340])).
fof(f340,plain,(
  ! [X0] : ! [X1] : s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),X1))))) = s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),X1)))))),
  inference(rectify,[],[f240])).
fof(f240,axiom,(
  ! [X6] : ! [X18] : s(fun(X6,fun(X6,bool)),'const.relation.inv_1'(s(fun(X6,fun(X6,bool)),'const.relation.TC_1'(s(fun(X6,fun(X6,bool)),X18))))) = s(fun(X6,fun(X6,bool)),'const.relation.TC_1'(s(fun(X6,fun(X6,bool)),'const.relation.inv_1'(s(fun(X6,fun(X6,bool)),X18)))))),
  file('thm.pred_set.EXTENSION','thm.relation.inv_MOVES_OUT|split|3')).
fof(f1200,plain,(
  ( ! [X0,X1] : (s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1))))) = s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),X1)))))) )),
  inference(cnf_transformation,[],[f337])).
fof(f337,plain,(
  ! [X0,X1] : s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1))))) = s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),X1)))))),
  inference(flattening,[],[f336])).
fof(f336,plain,(
  ! [X0] : ! [X1] : s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1))))) = s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),X1)))))),
  inference(rectify,[],[f261])).
fof(f261,axiom,(
  ! [X6] : ! [X18] : s(fun(X6,fun(X6,bool)),'const.relation.RC_1'(s(fun(X6,fun(X6,bool)),'const.relation.SC_1'(s(fun(X6,fun(X6,bool)),X18))))) = s(fun(X6,fun(X6,bool)),'const.relation.SC_1'(s(fun(X6,fun(X6,bool)),'const.relation.RC_1'(s(fun(X6,fun(X6,bool)),X18)))))),
  file('thm.pred_set.EXTENSION','thm.relation.RC_MOVES_OUT|split|0')).
fof(f1199,plain,(
  ( ! [X0,X1] : (s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1))) = s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))))) )),
  inference(cnf_transformation,[],[f335])).
fof(f335,plain,(
  ! [X0,X1] : s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1))) = s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))))),
  inference(flattening,[],[f334])).
fof(f334,plain,(
  ! [X0] : ! [X1] : s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1))) = s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))))),
  inference(rectify,[],[f250])).
fof(f250,axiom,(
  ! [X6] : ! [X18] : s(fun(X6,fun(X6,bool)),'const.relation.SC_1'(s(fun(X6,fun(X6,bool)),X18))) = s(fun(X6,fun(X6,bool)),'const.relation.SC_1'(s(fun(X6,fun(X6,bool)),'const.relation.SC_1'(s(fun(X6,fun(X6,bool)),X18)))))),
  file('thm.pred_set.EXTENSION','thm.relation.SC_IDEM')).
fof(f1198,plain,(
  ( ! [X0,X1] : (s(fun(X0,fun(X0,bool)),'const.relation.EQC_1'(s(fun(X0,fun(X0,bool)),X1))) = s(fun(X0,fun(X0,bool)),'const.relation.EQC_1'(s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),X1)))))) )),
  inference(cnf_transformation,[],[f333])).
fof(f333,plain,(
  ! [X0,X1] : s(fun(X0,fun(X0,bool)),'const.relation.EQC_1'(s(fun(X0,fun(X0,bool)),X1))) = s(fun(X0,fun(X0,bool)),'const.relation.EQC_1'(s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),X1)))))),
  inference(flattening,[],[f332])).
fof(f332,plain,(
  ! [X0] : ! [X1] : s(fun(X0,fun(X0,bool)),'const.relation.EQC_1'(s(fun(X0,fun(X0,bool)),X1))) = s(fun(X0,fun(X0,bool)),'const.relation.EQC_1'(s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),X1)))))),
  inference(rectify,[],[f252])).
fof(f252,axiom,(
  ! [X6] : ! [X18] : s(fun(X6,fun(X6,bool)),'const.relation.EQC_1'(s(fun(X6,fun(X6,bool)),X18))) = s(fun(X6,fun(X6,bool)),'const.relation.EQC_1'(s(fun(X6,fun(X6,bool)),'const.relation.TC_1'(s(fun(X6,fun(X6,bool)),X18)))))),
  file('thm.pred_set.EXTENSION','thm.relation.EQC_MOVES_IN|split|2')).
fof(f1197,plain,(
  ( ! [X0,X1] : (s(fun(X0,fun(X0,bool)),'const.relation.RTC_1'(s(fun(X0,fun(X0,bool)),X1))) = s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),X1)))))) )),
  inference(cnf_transformation,[],[f331])).
fof(f331,plain,(
  ! [X0,X1] : s(fun(X0,fun(X0,bool)),'const.relation.RTC_1'(s(fun(X0,fun(X0,bool)),X1))) = s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),X1)))))),
  inference(flattening,[],[f330])).
fof(f330,plain,(
  ! [X0] : ! [X1] : s(fun(X0,fun(X0,bool)),'const.relation.RTC_1'(s(fun(X0,fun(X0,bool)),X1))) = s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),X1)))))),
  inference(rectify,[],[f249])).
fof(f249,axiom,(
  ! [X6] : ! [X18] : s(fun(X6,fun(X6,bool)),'const.relation.RTC_1'(s(fun(X6,fun(X6,bool)),X18))) = s(fun(X6,fun(X6,bool)),'const.relation.TC_1'(s(fun(X6,fun(X6,bool)),'const.relation.RC_1'(s(fun(X6,fun(X6,bool)),X18)))))),
  file('thm.pred_set.EXTENSION','thm.relation.TC_RC_EQNS|split|1')).
fof(f1196,plain,(
  ( ! [X0,X1] : (s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1))) = s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))))) )),
  inference(cnf_transformation,[],[f329])).
fof(f329,plain,(
  ! [X0,X1] : s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1))) = s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))))),
  inference(flattening,[],[f328])).
fof(f328,plain,(
  ! [X0] : ! [X1] : s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1))) = s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))))),
  inference(rectify,[],[f246])).
fof(f246,axiom,(
  ! [X6] : ! [X18] : s(fun(X6,fun(X6,bool)),'const.relation.SC_1'(s(fun(X6,fun(X6,bool)),X18))) = s(fun(X6,fun(X6,bool)),'const.relation.inv_1'(s(fun(X6,fun(X6,bool)),'const.relation.SC_1'(s(fun(X6,fun(X6,bool)),X18)))))),
  file('thm.pred_set.EXTENSION','thm.relation.inv_SC|split|0')).
fof(f1195,plain,(
  ( ! [X0,X1] : (s(fun(X0,fun(X0,bool)),'const.relation.RTC_1'(s(fun(X0,fun(X0,bool)),X1))) = s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),X1)))))) )),
  inference(cnf_transformation,[],[f327])).
fof(f327,plain,(
  ! [X0,X1] : s(fun(X0,fun(X0,bool)),'const.relation.RTC_1'(s(fun(X0,fun(X0,bool)),X1))) = s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),X1)))))),
  inference(flattening,[],[f326])).
fof(f326,plain,(
  ! [X0] : ! [X1] : s(fun(X0,fun(X0,bool)),'const.relation.RTC_1'(s(fun(X0,fun(X0,bool)),X1))) = s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),X1)))))),
  inference(rectify,[],[f239])).
fof(f239,axiom,(
  ! [X6] : ! [X18] : s(fun(X6,fun(X6,bool)),'const.relation.RTC_1'(s(fun(X6,fun(X6,bool)),X18))) = s(fun(X6,fun(X6,bool)),'const.relation.RC_1'(s(fun(X6,fun(X6,bool)),'const.relation.TC_1'(s(fun(X6,fun(X6,bool)),X18)))))),
  file('thm.pred_set.EXTENSION','thm.relation.TC_RC_EQNS|split|0')).
fof(f1194,plain,(
  ( ! [X0,X1] : (s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1))) = s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),X1)))))) )),
  inference(cnf_transformation,[],[f325])).
fof(f325,plain,(
  ! [X0,X1] : s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1))) = s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),X1)))))),
  inference(flattening,[],[f324])).
fof(f324,plain,(
  ! [X0] : ! [X1] : s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1))) = s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),X1)))))),
  inference(rectify,[],[f245])).
fof(f245,axiom,(
  ! [X6] : ! [X18] : s(fun(X6,fun(X6,bool)),'const.relation.SC_1'(s(fun(X6,fun(X6,bool)),X18))) = s(fun(X6,fun(X6,bool)),'const.relation.SC_1'(s(fun(X6,fun(X6,bool)),'const.relation.inv_1'(s(fun(X6,fun(X6,bool)),X18)))))),
  file('thm.pred_set.EXTENSION','thm.relation.inv_SC|split|1')).
fof(f1193,plain,(
  ( ! [X0,X1] : (s(fun(X0,fun(X0,bool)),'const.relation.diag_1'(s(fun(X0,bool),X1))) = s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),'const.relation.diag_1'(s(fun(X0,bool),X1)))))) )),
  inference(cnf_transformation,[],[f323])).
fof(f323,plain,(
  ! [X0,X1] : s(fun(X0,fun(X0,bool)),'const.relation.diag_1'(s(fun(X0,bool),X1))) = s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),'const.relation.diag_1'(s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f322])).
fof(f322,plain,(
  ! [X0] : ! [X1] : s(fun(X0,fun(X0,bool)),'const.relation.diag_1'(s(fun(X0,bool),X1))) = s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),'const.relation.diag_1'(s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f154])).
fof(f154,axiom,(
  ! [X6] : ! [X27] : s(fun(X6,fun(X6,bool)),'const.relation.diag_1'(s(fun(X6,bool),X27))) = s(fun(X6,fun(X6,bool)),'const.relation.inv_1'(s(fun(X6,fun(X6,bool)),'const.relation.diag_1'(s(fun(X6,bool),X27)))))),
  file('thm.pred_set.EXTENSION','thm.relation.inv_diag')).
fof(f1192,plain,(
  ( ! [X0,X1] : (s(fun(X0,fun(X0,bool)),X1) = s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),X1)))))) )),
  inference(cnf_transformation,[],[f321])).
fof(f321,plain,(
  ! [X0,X1] : s(fun(X0,fun(X0,bool)),X1) = s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),X1)))))),
  inference(flattening,[],[f320])).
fof(f320,plain,(
  ! [X0] : ! [X1] : s(fun(X0,fun(X0,bool)),X1) = s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),X1)))))),
  inference(rectify,[],[f248])).
fof(f248,axiom,(
  ! [X6] : ! [X18] : s(fun(X6,fun(X6,bool)),X18) = s(fun(X6,fun(X6,bool)),'const.relation.inv_1'(s(fun(X6,fun(X6,bool)),'const.relation.inv_1'(s(fun(X6,fun(X6,bool)),X18)))))),
  file('thm.pred_set.EXTENSION','thm.relation.inv_inv')).
fof(f1191,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1))))))) )),
  inference(cnf_transformation,[],[f319])).
fof(f319,plain,(
  ! [X0,X1] : p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(flattening,[],[f318])).
fof(f318,plain,(
  ! [X0] : ! [X1] : p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(rectify,[],[f251])).
fof(f251,axiom,(
  ! [X6] : ! [X18] : p(s(bool,'const.relation.symmetric_1'(s(fun(X6,fun(X6,bool)),'const.relation.SC_1'(s(fun(X6,fun(X6,bool)),X18))))))),
  file('thm.pred_set.EXTENSION','thm.relation.SC_SYMMETRIC')).
fof(f1701,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))) )),
  inference(forward_demodulation,[],[f1190,f1188])).
fof(f1190,plain,(
  ( ! [X0,X1] : (s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f317])).
fof(f317,plain,(
  ! [X0,X1] : s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))),
  inference(flattening,[],[f316])).
fof(f316,plain,(
  ! [X0] : ! [X1] : s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))),
  inference(rectify,[],[f149])).
fof(f149,axiom,(
  ! [X6] : ! [X21] : s(X6,'const.combin.I_1'(s(X6,X21))) = s(X6,happ(s(fun(X6,X6),'const.combin.I_0'),s(X6,X21)))),
  file('thm.pred_set.EXTENSION','ARITY_0_1_const.combin.I')).
fof(f1133,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f847])).
fof(f847,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f288])).
fof(f288,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f50])).
fof(f50,axiom,(
  ! [X8] : (s(bool,'T_0') = s(bool,X8) <=> p(s(bool,X8)))),
  file('thm.pred_set.EXTENSION','thm.bool.EQ_CLAUSES|split|0')).
fof(f1134,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f847])).
fof(f1118,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.num.num','type.num.num'),'const.arithmetic.DIVMOD_1'(s('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),X0))) = s('type.pair.prod'('type.num.num','type.num.num'),happ(s(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.pair.prod'('type.num.num','type.num.num')),'const.arithmetic.DIVMOD_0'),s('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),X0)))) )),
  inference(cnf_transformation,[],[f271])).
fof(f271,plain,(
  ! [X0] : s('type.pair.prod'('type.num.num','type.num.num'),'const.arithmetic.DIVMOD_1'(s('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),X0))) = s('type.pair.prod'('type.num.num','type.num.num'),happ(s(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.pair.prod'('type.num.num','type.num.num')),'const.arithmetic.DIVMOD_0'),s('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),X0)))),
  inference(rectify,[],[f196])).
fof(f196,axiom,(
  ! [X21] : s('type.pair.prod'('type.num.num','type.num.num'),'const.arithmetic.DIVMOD_1'(s('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),X21))) = s('type.pair.prod'('type.num.num','type.num.num'),happ(s(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.pair.prod'('type.num.num','type.num.num')),'const.arithmetic.DIVMOD_0'),s('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),X21)))),
  file('thm.pred_set.EXTENSION','ARITY_0_1_const.arithmetic.DIVMOD')).
fof(f1117,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.findq_1'(s('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),X0))) = s('type.num.num',happ(s(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.num.num'),'const.arithmetic.findq_0'),s('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),X0)))) )),
  inference(cnf_transformation,[],[f270])).
fof(f270,plain,(
  ! [X0] : s('type.num.num','const.arithmetic.findq_1'(s('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),X0))) = s('type.num.num',happ(s(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.num.num'),'const.arithmetic.findq_0'),s('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),X0)))),
  inference(rectify,[],[f164])).
fof(f164,axiom,(
  ! [X21] : s('type.num.num','const.arithmetic.findq_1'(s('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),X21))) = s('type.num.num',happ(s(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.num.num'),'const.arithmetic.findq_0'),s('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),X21)))),
  file('thm.pred_set.EXTENSION','ARITY_0_1_const.arithmetic.findq')).
fof(f1107,plain,(
  s(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.pair.prod'('type.num.num','type.num.num')),'const.arithmetic.DIVMOD_0') = s(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.pair.prod'('type.num.num','type.num.num')),'const.relation.WFREC_2'(s(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),bool)),'const.prim_rec.measure_1'(s(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.num.num'),'const.combin.o_2'(s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.num.num'),'const.pair.FST_0'),s(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.pair.prod'('type.num.num','type.num.num')),'const.pair.SND_0'))))),s(fun(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.pair.prod'('type.num.num','type.num.num')),fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.pair.prod'('type.num.num','type.num.num'))),'_dst_x0x1_2'(s(fun(fun('type.num.num',fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.pair.prod'('type.num.num','type.num.num'))),'const.pair.UNCURRY_0'),s(fun(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.pair.prod'('type.num.num','type.num.num')),fun('type.num.num',fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num')))),'_dst_x00x11_2'(s(fun(fun('type.num.num',fun('type.num.num','type.pair.prod'('type.num.num','type.num.num'))),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.pair.UNCURRY_0'),s(fun(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.pair.prod'('type.num.num','type.num.num')),fun('type.num.num',fun('type.num.num',fun('type.num.num','type.pair.prod'('type.num.num','type.num.num'))))),'_dst_x0000x0001x0000x1111_4'(s(fun(bool,fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num')))),'const.bool.COND_0'),s(fun('type.num.num',bool),'_dst_x1x0_2'(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.num.|48|_0'))),s('type.pair.prod'('type.num.num','type.num.num'),happ(s(fun('type.num.num','type.pair.prod'('type.num.num','type.num.num')),happ(s(fun('type.num.num',fun('type.num.num','type.pair.prod'('type.num.num','type.num.num'))),'const.pair.,_0'),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0'))),s(fun(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.pair.prod'('type.num.num','type.num.num')),fun('type.num.num',fun('type.num.num',fun('type.num.num','type.pair.prod'('type.num.num','type.num.num'))))),'_dst_x0000x0011x0110x1111_4'(s(fun(bool,fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num')))),'const.bool.COND_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0'),s(fun('type.num.num',fun('type.num.num','type.pair.prod'('type.num.num','type.num.num'))),'const.pair.,_0'),s(fun(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.pair.prod'('type.num.num','type.num.num')),fun('type.num.num',fun('type.num.num',fun('type.num.num','type.pair.prod'('type.num.num','type.num.num'))))),'_dst_x0000x1111x0011_3'(s(fun(fun('type.num.num','type.pair.prod'('type.num.num','type.num.num')),fun('type.num.num','type.pair.prod'('type.num.num','type.num.num'))),'const.bool.LET_0'),s(fun(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.pair.prod'('type.num.num','type.num.num')),fun('type.num.num',fun('type.num.num',fun('type.num.num',fun('type.num.num','type.pair.prod'('type.num.num','type.num.num')))))),'_dstsel_0_x01111_1'(s(fun('type.num.num',fun('type.num.num',fun('type.num.num',fun('type.num.num','type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')))))),'_dst_x0000x1001x0111_3'(s(fun('type.num.num',fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')))),'const.pair.,_0'),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.+_0'),s(fun('type.num.num',fun('type.num.num',fun('type.num.num','type.pair.prod'('type.num.num','type.num.num')))),'_dst_x000x111x010_3'(s(fun('type.num.num',fun('type.num.num','type.pair.prod'('type.num.num','type.num.num'))),'const.pair.,_0'),s(fun('type.num.num',fun('type.num.num',fun('type.num.num','type.num.num'))),'_dst_x000x100x011_3'(s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.-_0'),s(fun('type.num.num','type.num.num'),'_i_0_0'),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.*_0'))),s(fun('type.num.num','type.num.num'),'_i_0_0'))))))),s(fun('type.num.num',fun('type.num.num','type.num.num')),'_dst_x00x11_2'(s(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.num.num'),'const.arithmetic.findq_0'),s(fun('type.num.num',fun('type.num.num','type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')))),'_dst_x00x11_2'(s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num'))),happ(s(fun('type.num.num',fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')))),'const.pair.,_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s(fun('type.num.num',fun('type.num.num','type.pair.prod'('type.num.num','type.num.num'))),'const.pair.,_0')))))))))))))))))),
  inference(cnf_transformation,[],[f159])).
fof(f159,axiom,(
  s(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.pair.prod'('type.num.num','type.num.num')),'const.arithmetic.DIVMOD_0') = s(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.pair.prod'('type.num.num','type.num.num')),'const.relation.WFREC_2'(s(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),bool)),'const.prim_rec.measure_1'(s(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.num.num'),'const.combin.o_2'(s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.num.num'),'const.pair.FST_0'),s(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.pair.prod'('type.num.num','type.num.num')),'const.pair.SND_0'))))),s(fun(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.pair.prod'('type.num.num','type.num.num')),fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.pair.prod'('type.num.num','type.num.num'))),'_dst_x0x1_2'(s(fun(fun('type.num.num',fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.pair.prod'('type.num.num','type.num.num'))),'const.pair.UNCURRY_0'),s(fun(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.pair.prod'('type.num.num','type.num.num')),fun('type.num.num',fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num')))),'_dst_x00x11_2'(s(fun(fun('type.num.num',fun('type.num.num','type.pair.prod'('type.num.num','type.num.num'))),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.pair.UNCURRY_0'),s(fun(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.pair.prod'('type.num.num','type.num.num')),fun('type.num.num',fun('type.num.num',fun('type.num.num','type.pair.prod'('type.num.num','type.num.num'))))),'_dst_x0000x0001x0000x1111_4'(s(fun(bool,fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num')))),'const.bool.COND_0'),s(fun('type.num.num',bool),'_dst_x1x0_2'(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.num.|48|_0'))),s('type.pair.prod'('type.num.num','type.num.num'),happ(s(fun('type.num.num','type.pair.prod'('type.num.num','type.num.num')),happ(s(fun('type.num.num',fun('type.num.num','type.pair.prod'('type.num.num','type.num.num'))),'const.pair.,_0'),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0'))),s(fun(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.pair.prod'('type.num.num','type.num.num')),fun('type.num.num',fun('type.num.num',fun('type.num.num','type.pair.prod'('type.num.num','type.num.num'))))),'_dst_x0000x0011x0110x1111_4'(s(fun(bool,fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num')))),'const.bool.COND_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0'),s(fun('type.num.num',fun('type.num.num','type.pair.prod'('type.num.num','type.num.num'))),'const.pair.,_0'),s(fun(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.pair.prod'('type.num.num','type.num.num')),fun('type.num.num',fun('type.num.num',fun('type.num.num','type.pair.prod'('type.num.num','type.num.num'))))),'_dst_x0000x1111x0011_3'(s(fun(fun('type.num.num','type.pair.prod'('type.num.num','type.num.num')),fun('type.num.num','type.pair.prod'('type.num.num','type.num.num'))),'const.bool.LET_0'),s(fun(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.pair.prod'('type.num.num','type.num.num')),fun('type.num.num',fun('type.num.num',fun('type.num.num',fun('type.num.num','type.pair.prod'('type.num.num','type.num.num')))))),'_dstsel_0_x01111_1'(s(fun('type.num.num',fun('type.num.num',fun('type.num.num',fun('type.num.num','type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')))))),'_dst_x0000x1001x0111_3'(s(fun('type.num.num',fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')))),'const.pair.,_0'),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.+_0'),s(fun('type.num.num',fun('type.num.num',fun('type.num.num','type.pair.prod'('type.num.num','type.num.num')))),'_dst_x000x111x010_3'(s(fun('type.num.num',fun('type.num.num','type.pair.prod'('type.num.num','type.num.num'))),'const.pair.,_0'),s(fun('type.num.num',fun('type.num.num',fun('type.num.num','type.num.num'))),'_dst_x000x100x011_3'(s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.-_0'),s(fun('type.num.num','type.num.num'),'_i_0_0'),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.*_0'))),s(fun('type.num.num','type.num.num'),'_i_0_0'))))))),s(fun('type.num.num',fun('type.num.num','type.num.num')),'_dst_x00x11_2'(s(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.num.num'),'const.arithmetic.findq_0'),s(fun('type.num.num',fun('type.num.num','type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')))),'_dst_x00x11_2'(s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num'))),happ(s(fun('type.num.num',fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')))),'const.pair.,_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s(fun('type.num.num',fun('type.num.num','type.pair.prod'('type.num.num','type.num.num'))),'const.pair.,_0')))))))))))))))))),
  file('thm.pred_set.EXTENSION','thm.arithmetic.DIVMOD_DEF')).
fof(f1106,plain,(
  s(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.num.num'),'const.arithmetic.findq_0') = s(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.num.num'),'const.relation.WFREC_2'(s(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),bool)),'const.prim_rec.measure_1'(s(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.num.num'),happ(s(fun(fun('type.num.num',fun('type.pair.prod'('type.num.num','type.num.num'),'type.num.num')),fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.num.num')),'const.pair.UNCURRY_0'),s(fun('type.num.num',fun('type.pair.prod'('type.num.num','type.num.num'),'type.num.num')),'_dst_x0_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.num.num'),happ(s(fun(fun('type.num.num',fun('type.num.num','type.num.num')),fun('type.pair.prod'('type.num.num','type.num.num'),'type.num.num')),'const.pair.UNCURRY_0'),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.-_0'))))))))),s(fun(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.num.num'),fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.num.num')),'_dst_x0x1_2'(s(fun(fun('type.num.num',fun('type.pair.prod'('type.num.num','type.num.num'),'type.num.num')),fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.num.num')),'const.pair.UNCURRY_0'),s(fun(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.num.num'),fun('type.num.num',fun('type.pair.prod'('type.num.num','type.num.num'),'type.num.num'))),'_dst_x00x11_2'(s(fun(fun('type.num.num',fun('type.num.num','type.num.num')),fun('type.pair.prod'('type.num.num','type.num.num'),'type.num.num')),'const.pair.UNCURRY_0'),s(fun(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.num.num'),fun('type.num.num',fun('type.num.num',fun('type.num.num','type.num.num')))),'_dst_x0000x0001x0100x1111_4'(s(fun(bool,fun('type.num.num',fun('type.num.num','type.num.num'))),'const.bool.COND_0'),s(fun('type.num.num',bool),'_dst_x1x0_2'(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.num.|48|_0'))),s(fun('type.num.num','type.num.num'),'_i_0_0'),s(fun(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.num.num'),fun('type.num.num',fun('type.num.num',fun('type.num.num','type.num.num')))),'_dst_x0000x1110x0001_3'(s(fun(fun('type.num.num','type.num.num'),fun('type.num.num','type.num.num')),'const.bool.LET_0'),s(fun(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.num.num'),fun('type.num.num',fun('type.num.num',fun('type.num.num','type.num.num')))),'_dst_x0000x0011x0100x1111_4'(s(fun(bool,fun('type.num.num',fun('type.num.num','type.num.num'))),'const.bool.COND_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0'),s(fun('type.num.num','type.num.num'),'_i_0_0'),s(fun(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.num.num'),fun('type.num.num',fun('type.num.num',fun('type.num.num','type.num.num')))),'_dstsel_0_x0111_1'(s(fun('type.num.num',fun('type.num.num',fun('type.num.num','type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num'))))),'_dst_x000x100x011_3'(s(fun('type.num.num',fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')))),'const.pair.,_0'),s(fun('type.num.num','type.num.num'),'const.arithmetic.*_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s(fun('type.num.num',fun('type.num.num','type.pair.prod'('type.num.num','type.num.num'))),'const.pair.,_0'))))))),s(fun('type.num.num','type.num.num'),'const.arithmetic.*_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))))))))),
  inference(cnf_transformation,[],[f155])).
fof(f155,axiom,(
  s(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.num.num'),'const.arithmetic.findq_0') = s(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.num.num'),'const.relation.WFREC_2'(s(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),bool)),'const.prim_rec.measure_1'(s(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.num.num'),happ(s(fun(fun('type.num.num',fun('type.pair.prod'('type.num.num','type.num.num'),'type.num.num')),fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.num.num')),'const.pair.UNCURRY_0'),s(fun('type.num.num',fun('type.pair.prod'('type.num.num','type.num.num'),'type.num.num')),'_dst_x0_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.num.num'),happ(s(fun(fun('type.num.num',fun('type.num.num','type.num.num')),fun('type.pair.prod'('type.num.num','type.num.num'),'type.num.num')),'const.pair.UNCURRY_0'),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.-_0'))))))))),s(fun(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.num.num'),fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.num.num')),'_dst_x0x1_2'(s(fun(fun('type.num.num',fun('type.pair.prod'('type.num.num','type.num.num'),'type.num.num')),fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.num.num')),'const.pair.UNCURRY_0'),s(fun(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.num.num'),fun('type.num.num',fun('type.pair.prod'('type.num.num','type.num.num'),'type.num.num'))),'_dst_x00x11_2'(s(fun(fun('type.num.num',fun('type.num.num','type.num.num')),fun('type.pair.prod'('type.num.num','type.num.num'),'type.num.num')),'const.pair.UNCURRY_0'),s(fun(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.num.num'),fun('type.num.num',fun('type.num.num',fun('type.num.num','type.num.num')))),'_dst_x0000x0001x0100x1111_4'(s(fun(bool,fun('type.num.num',fun('type.num.num','type.num.num'))),'const.bool.COND_0'),s(fun('type.num.num',bool),'_dst_x1x0_2'(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.num.|48|_0'))),s(fun('type.num.num','type.num.num'),'_i_0_0'),s(fun(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.num.num'),fun('type.num.num',fun('type.num.num',fun('type.num.num','type.num.num')))),'_dst_x0000x1110x0001_3'(s(fun(fun('type.num.num','type.num.num'),fun('type.num.num','type.num.num')),'const.bool.LET_0'),s(fun(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.num.num'),fun('type.num.num',fun('type.num.num',fun('type.num.num','type.num.num')))),'_dst_x0000x0011x0100x1111_4'(s(fun(bool,fun('type.num.num',fun('type.num.num','type.num.num'))),'const.bool.COND_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0'),s(fun('type.num.num','type.num.num'),'_i_0_0'),s(fun(fun('type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')),'type.num.num'),fun('type.num.num',fun('type.num.num',fun('type.num.num','type.num.num')))),'_dstsel_0_x0111_1'(s(fun('type.num.num',fun('type.num.num',fun('type.num.num','type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num'))))),'_dst_x000x100x011_3'(s(fun('type.num.num',fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.num.num','type.num.num')))),'const.pair.,_0'),s(fun('type.num.num','type.num.num'),'const.arithmetic.*_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s(fun('type.num.num',fun('type.num.num','type.pair.prod'('type.num.num','type.num.num'))),'const.pair.,_0'))))))),s(fun('type.num.num','type.num.num'),'const.arithmetic.*_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))))))))),
  file('thm.pred_set.EXTENSION','thm.arithmetic.findq_def')).
fof(f4664,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2)))))) | s(X0,X2) = s(X0,X3)) ) | $spl12),
  inference(subsumption_resolution,[],[f4663,f4647])).
fof(f4663,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2)))))) | s(X0,X2) = s(X0,X3) | ~p(s(bool,'const.relation.WeakOrder_1'(s(fun(X0,fun(X0,bool)),X1))))) ) | $spl12),
  inference(subsumption_resolution,[],[f4427,f1104])).
fof(f4427,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2)))))) | s(X0,X2) = s(X0,X3) | ~p(s(bool,'const.relation.WeakOrder_1'(s(fun(X0,fun(X0,bool)),X1))))) ) | $spl12),
  inference(backward_demodulation,[],[f4285,f1724])).
fof(f4678,plain,(
  ( ! [X2,X0,X3] : (s(X0,X2) != s(X0,X3)) ) | $spl12),
  inference(subsumption_resolution,[],[f4608,f1104])).
fof(f4608,plain,(
  ( ! [X2,X0,X3] : (~p(s(bool,'T_0')) | s(X0,X2) != s(X0,X3)) ) | $spl12),
  inference(backward_demodulation,[],[f4285,f1336])).
fof(f4677,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f4676,f4665])).
fof(f4676,plain,(
  ( ! [X4,X3,X1] : (s(X1,X3) != s(X1,X4)) ) | $spl12),
  inference(subsumption_resolution,[],[f4583,f1104])).
fof(f4583,plain,(
  ( ! [X4,X3,X1] : (~p(s(bool,'T_0')) | s(X1,X3) != s(X1,X4)) ) | $spl12),
  inference(backward_demodulation,[],[f4285,f1500])).
fof(f4669,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f4668,f4665])).
fof(f4668,plain,(
  ( ! [X2,X3,X1] : (s(X1,X2) != s(X1,X3)) ) | $spl12),
  inference(subsumption_resolution,[],[f4486,f1104])).
fof(f4486,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,'T_0')) | s(X1,X2) != s(X1,X3)) ) | $spl12),
  inference(backward_demodulation,[],[f4285,f1485])).
fof(f4667,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f4666,f4647])).
fof(f4666,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.relation.irreflexive_1'(s(fun(X0,fun(X0,bool)),X1))))) ) | $spl12),
  inference(subsumption_resolution,[],[f4458,f1104])).
fof(f4458,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.relation.irreflexive_1'(s(fun(X0,fun(X0,bool)),X1))))) ) | $spl12),
  inference(backward_demodulation,[],[f4285,f1710])).
fof(f4656,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f4385,f4647])).
fof(f4385,plain,(
  ~p(s(bool,'|dollar|or_2'(s(bool,'T_0'),s(bool,'T_0')))) | $spl12),
  inference(backward_demodulation,[],[f4285,f4258])).
fof(f4636,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f4322,f1104])).
fof(f4322,plain,(
  ~p(s(bool,'T_0')) | $spl12),
  inference(backward_demodulation,[],[f4285,f1700])).
fof(f4633,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f4632,f4285])).
fof(f4632,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0)) ) | $spl12),
  inference(subsumption_resolution,[],[f4317,f1133])).
fof(f4317,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl12),
  inference(backward_demodulation,[],[f4285,f1137])).
fof(f4285,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl12),
  inference(cnf_transformation,[],[f4285_D])).
fof(f4285_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f4301,plain,(
  ( ! [X1] : (~p(s(bool,'|dollar|or_2'(s(bool,X1),s(bool,'F_0')))) | p(s(bool,X1))) ) | $spl16),
  inference(cnf_transformation,[],[f4301_D])).
fof(f4301_D,plain,(
  ( ! [X1] : (~p(s(bool,'|dollar|or_2'(s(bool,X1),s(bool,'F_0')))) | p(s(bool,X1))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f4300,plain,(
  ( ! [X0,X1] : (~p(s(bool,'|dollar|or_2'(s(bool,X1),s(bool,'F_0')))) | p(s(bool,X1)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f4271,f1134])).
fof(f4271,plain,(
  ( ! [X0,X1] : (~p(s(bool,'|dollar|or_2'(s(bool,X1),s(bool,'F_0')))) | p(s(bool,X1)) | p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f1679,f1120])).
fof(f4286,plain,(
  ( ! [X1] : (~p(s(bool,'|dollar|or_2'(s(bool,'F_0'),s(bool,X1)))) | p(s(bool,X1))) ) | $spl14),
  inference(cnf_transformation,[],[f4286_D])).
fof(f4286_D,plain,(
  ( ! [X1] : (~p(s(bool,'|dollar|or_2'(s(bool,'F_0'),s(bool,X1)))) | p(s(bool,X1))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f4284,plain,(
  ( ! [X0,X1] : (~p(s(bool,'|dollar|or_2'(s(bool,'F_0'),s(bool,X1)))) | p(s(bool,X1)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f4259,f1134])).
fof(f4259,plain,(
  ( ! [X0,X1] : (~p(s(bool,'|dollar|or_2'(s(bool,'F_0'),s(bool,X1)))) | p(s(bool,X0)) | p(s(bool,X1)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f1679,f1120])).
fof(f1697,plain,(
  ~p(s(bool,'F_0')) | $spl9),
  inference(cnf_transformation,[],[f1697_D])).
fof(f1697_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f1108,plain,(
  ( ! [X1] : (p(s(bool,X1)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f844])).
fof(f844,plain,(
  (~p(s(bool,'F_0')) | ! [X1] : p(s(bool,X1))) & (~p(s(bool,sK16)) | p(s(bool,'F_0')))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK16])],[f843])).
fof(f843,plain,(
  (~p(s(bool,'F_0')) | ! [X1] : p(s(bool,X1))) & (? [X0] : ~p(s(bool,X0)) | p(s(bool,'F_0')))),
  inference(rectify,[],[f842])).
fof(f842,plain,(
  (~p(s(bool,'F_0')) | ! [X0] : p(s(bool,X0))) & (? [X0] : ~p(s(bool,X0)) | p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f264])).
fof(f264,plain,(
  p(s(bool,'F_0')) <=> ! [X0] : p(s(bool,X0))),
  inference(rectify,[],[f77])).
fof(f77,axiom,(
  p(s(bool,'F_0')) <=> ! [X8] : p(s(bool,X8))),
  file('thm.pred_set.EXTENSION','thm.bool.F_DEF')).
fof(f1696,plain,(
  ~p(s(bool,sK16)) | $spl7),
  inference(cnf_transformation,[],[f1696_D])).
fof(f1696_D,plain,(
  ~p(s(bool,sK16)) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f1698,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(cnf_transformation,[],[f1698_D])).
fof(f1698_D,plain,(
  p(s(bool,'F_0')) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f1109,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,sK16))),
  inference(cnf_transformation,[],[f844])).
fof(f1690,plain,(
  s(fun(sK12,bool),sK13) = s(fun(sK12,bool),sK14) | $spl0),
  inference(cnf_transformation,[],[f1690_D])).
fof(f1690_D,plain,(
  s(fun(sK12,bool),sK13) = s(fun(sK12,bool),sK14) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f1694,plain,(
  ( ! [X4] : (s(bool,'const.bool.IN_2'(s(sK12,X4),s(fun(sK12,bool),sK13))) = s(bool,'const.bool.IN_2'(s(sK12,X4),s(fun(sK12,bool),sK14)))) ) | $spl4),
  inference(cnf_transformation,[],[f1694_D])).
fof(f1694_D,plain,(
  ( ! [X4] : (s(bool,'const.bool.IN_2'(s(sK12,X4),s(fun(sK12,bool),sK13))) = s(bool,'const.bool.IN_2'(s(sK12,X4),s(fun(sK12,bool),sK14)))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f1102,plain,(
  ( ! [X4] : (s(bool,'const.bool.IN_2'(s(sK12,X4),s(fun(sK12,bool),sK13))) = s(bool,'const.bool.IN_2'(s(sK12,X4),s(fun(sK12,bool),sK14))) | s(fun(sK12,bool),sK13) = s(fun(sK12,bool),sK14)) )),
  inference(cnf_transformation,[],[f841])).
fof(f841,plain,(
  (s(fun(sK12,bool),sK13) = s(fun(sK12,bool),sK14) | ! [X4] : s(bool,'const.bool.IN_2'(s(sK12,X4),s(fun(sK12,bool),sK13))) = s(bool,'const.bool.IN_2'(s(sK12,X4),s(fun(sK12,bool),sK14)))) & (s(fun(sK12,bool),sK13) != s(fun(sK12,bool),sK14) | s(bool,'const.bool.IN_2'(s(sK12,sK15),s(fun(sK12,bool),sK13))) != s(bool,'const.bool.IN_2'(s(sK12,sK15),s(fun(sK12,bool),sK14))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13,sK14,sK15])],[f840])).
fof(f840,plain,(
  ? [X0,X1,X2] : ((s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))),
  inference(rectify,[],[f839])).
fof(f839,plain,(
  ? [X0,X1,X2] : ((s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) & (s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))),
  inference(nnf_transformation,[],[f723])).
fof(f723,plain,(
  ? [X0,X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <~> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(ennf_transformation,[],[f263])).
fof(f263,plain,(
  ~! [X0,X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(flattening,[],[f262])).
fof(f262,plain,(
  ~! [X0] : ! [X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : (s(fun(X6,bool),X7) = s(fun(X6,bool),X8) <=> ! [X9] : s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X7))) = s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X8))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : (s(fun(X6,bool),X7) = s(fun(X6,bool),X8) <=> ! [X9] : s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X7))) = s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X8))))),
  file('thm.pred_set.EXTENSION','thm.pred_set.EXTENSION')).
fof(f1691,plain,(
  s(fun(sK12,bool),sK13) != s(fun(sK12,bool),sK14) | $spl1),
  inference(cnf_transformation,[],[f1691_D])).
fof(f1691_D,plain,(
  s(fun(sK12,bool),sK13) != s(fun(sK12,bool),sK14) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f1693,plain,(
  s(bool,'const.bool.IN_2'(s(sK12,sK15),s(fun(sK12,bool),sK13))) != s(bool,'const.bool.IN_2'(s(sK12,sK15),s(fun(sK12,bool),sK14))) | $spl3),
  inference(cnf_transformation,[],[f1693_D])).
fof(f1693_D,plain,(
  s(bool,'const.bool.IN_2'(s(sK12,sK15),s(fun(sK12,bool),sK13))) != s(bool,'const.bool.IN_2'(s(sK12,sK15),s(fun(sK12,bool),sK14))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f1103,plain,(
  s(bool,'const.bool.IN_2'(s(sK12,sK15),s(fun(sK12,bool),sK13))) != s(bool,'const.bool.IN_2'(s(sK12,sK15),s(fun(sK12,bool),sK14))) | s(fun(sK12,bool),sK13) != s(fun(sK12,bool),sK14)),
  inference(cnf_transformation,[],[f841])).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 1466
% Passive clauses: 16676
% Generated clauses: 89944
% Final active clauses: 1208
% Final passive clauses: 3914
% Input formulas: 260
% Initial clauses: 562
% Discarded non-redundant clauses: 60251
% 
% Duplicate literals: 11
% Trivial inequalities: 23
% Fw subsumption resolutions: 601
% Fw demodulations: 23332
% Bw demodulations: 1518
% Condensations: 41
% Global subsumptions: 4
% 
% Simple tautologies: 115
% Equational tautologies: 58
% Forward subsumptions: 44582
% Fw demodulations to eq. taut.: 60
% Bw demodulations to eq. taut.: 2
% 
% Unit resulting resolution: 28161
% Factoring: 3
% Forward superposition: 18572
% Backward superposition: 16468
% Self superposition: 160
% Equality factoring: 2
% Equality resolution: 10
% 
% Split clauses: 37
% Split components: 74
% SAT solver clauses: 17299
% SAT solver unit clauses: 3684
% SAT solver binary clauses: 8850
% 
% Sat splits: 37
% Sat splitting refutations: 543
% 
% Memory used [KB]: 137012
% Time elapsed: 8.842 s
% ------------------------------
% ------------------------------
% Success in time 10.18 s
