Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.rat
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.rat
% SZS output start Proof for thm.rat
fof(f23715,plain,(
  $false),
  inference(subsumption_resolution,[],[f23714,f1480])).
fof(f1480,plain,(
  s('type.rat.rat','const.rat.rat_|49|_0') = s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac','const.frac.frac_|49|_0')))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  s('type.rat.rat','const.rat.rat_|49|_0') = s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac','const.frac.frac_|49|_0')))),
  file('thm.rat.rat_|49|','thm.rat.rat_|49|_def')).
fof(f23714,plain,(
  s('type.rat.rat','const.rat.rat_|49|_0') != s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac','const.frac.frac_|49|_0')))),
  inference(forward_demodulation,[],[f23713,f4190])).
fof(f4190,plain,(
  s('type.rat.rat','const.rat.rat_|49|_0') = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(backward_demodulation,[],[f4143,f1488])).
fof(f1488,plain,(
  s('type.rat.rat','const.rat.rat_|49|_0') = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(cnf_transformation,[],[f15])).
fof(f15,axiom,(
  s('type.rat.rat','const.rat.rat_|49|_0') = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.rat.rat_|49|','thm.rat.rat_of_num_def|split|1')).
fof(f4143,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.num.|48|_0')),
  inference(superposition,[],[f1538,f1535])).
fof(f1535,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f535])).
fof(f535,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f130])).
fof(f130,axiom,(
  ! [X29] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X29)))),
  file('thm.rat.rat_|49|','thm.arithmetic.MULT_CLAUSES|split|0')).
fof(f1538,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(cnf_transformation,[],[f538])).
fof(f538,plain,(
  ! [X0] : s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(rectify,[],[f499])).
fof(f499,axiom,(
  ! [X106] : s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X106),s('type.num.num','const.arithmetic.ZERO_0')))),
  file('thm.rat.rat_|49|','thm.numRing.num_rewrites|split|69')).
fof(f23713,plain,(
  s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac','const.frac.frac_|49|_0'))) != s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(forward_demodulation,[],[f1475,f9011])).
fof(f9011,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(forward_demodulation,[],[f1493,f4143])).
fof(f1493,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(cnf_transformation,[],[f55])).
fof(f55,axiom,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  file('thm.rat.rat_|49|','thm.numeral.numeral_distrib|split|13')).
fof(f1475,plain,(
  s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac','const.frac.frac_|49|_0'))) != s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(cnf_transformation,[],[f519])).
fof(f519,plain,(
  s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac','const.frac.frac_|49|_0'))) != s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(flattening,[],[f6])).
fof(f6,negated_conjecture,(
  ~s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac','const.frac.frac_|49|_0'))) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac','const.frac.frac_|49|_0'))) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  file('thm.rat.rat_|49|','thm.rat.rat_|49|')).
% SZS output end Proof for thm.rat
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 1123
% Passive clauses: 4460
% Generated clauses: 22204
% Final active clauses: 1014
% Final passive clauses: 2969
% Input formulas: 517
% Initial clauses: 851
% 
% Duplicate literals: 173
% Trivial inequalities: 97
% Fw subsumption resolutions: 1383
% Fw demodulations: 2056
% Bw demodulations: 104
% 
% Simple tautologies: 283
% Equational tautologies: 23
% Forward subsumptions: 13526
% 
% Binary resolution: 4585
% Factoring: 382
% Forward superposition: 4798
% Backward superposition: 7532
% Self superposition: 52
% Equality factoring: 2
% Equality resolution: 11
% 
% Split clauses: 359
% Split components: 720
% SAT solver clauses: 552
% SAT solver unit clauses: 21
% SAT solver binary clauses: 135
% 
% Sat splits: 359
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 552
% TWLsolver calls for satisfiability: 412
% 
% Memory used [KB]: 14967
% Time elapsed: 0.355 s
% ------------------------------
% ------------------------------
% Success in time 0.382 s
