Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.integer
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.integer
% SZS output start Proof for thm.integer
fof(f24317,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f24315,f23618,f366,f636,f377,f641,f365,f634,f632,f24161,f367,f637,f635,f633,f23772,f23769])).
fof(f23769,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11)))) | $spl238),
  inference(cnf_transformation,[],[f23769_D])).
fof(f23769_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11)))) <=> ~$spl238),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl238])])).
fof(f23772,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11)))) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f23771,f381])).
fof(f381,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f114])).
fof(f114,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X8] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X8)))),
  file('thm.integer.integer_326_prove_ltb','thm.arithmetic.MULT_CLAUSES|split|0')).
fof(f23771,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))),s('type.num.num',sK11)))) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f632,f634])).
fof(f633,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK10),s('type.num.num',sK11))),s('type.num.num',sK11)))) | $spl1),
  inference(cnf_transformation,[],[f633_D])).
fof(f633_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK10),s('type.num.num',sK11))),s('type.num.num',sK11)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f635,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK10) | $spl3),
  inference(cnf_transformation,[],[f635_D])).
fof(f635_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK10) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f637,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK11) | $spl4),
  inference(cnf_transformation,[],[f637_D])).
fof(f637_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK11) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f367,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK11) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK10) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK10),s('type.num.num',sK11))),s('type.num.num',sK11))))),
  inference(cnf_transformation,[],[f265])).
fof(f265,plain,(
  (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK10),s('type.num.num',sK11))),s('type.num.num',sK11)))) | (s('type.num.num','const.num.|48|_0') = s('type.num.num',sK10) & s('type.num.num','const.num.|48|_0') != s('type.num.num',sK11))) & (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK10),s('type.num.num',sK11))),s('type.num.num',sK11)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK10) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK11))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK10,sK11])],[f264])).
fof(f264,plain,(
  ? [X0,X1] : ((p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X1)))) | (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') != s('type.num.num',X1))) & (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X1)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)))),
  inference(flattening,[],[f263])).
fof(f263,plain,(
  ? [X0,X1] : ((p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X1)))) | (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') != s('type.num.num',X1))) & (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X1)))) | (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1))))),
  inference(nnf_transformation,[],[f226])).
fof(f226,plain,(
  ? [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X1)))) <~> (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)))),
  inference(ennf_transformation,[],[f107])).
fof(f107,plain,(
  ~! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X1)))) <=> (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)))),
  inference(flattening,[],[f106])).
fof(f106,plain,(
  ~! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X1)))) <=> (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) & ~s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X6),s('type.num.num',X7))),s('type.num.num',X7)))) <=> (s('type.num.num','const.num.|48|_0') = s('type.num.num',X6) & ~s('type.num.num','const.num.|48|_0') = s('type.num.num',X7)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X6),s('type.num.num',X7))),s('type.num.num',X7)))) <=> (s('type.num.num','const.num.|48|_0') = s('type.num.num',X6) & ~s('type.num.num','const.num.|48|_0') = s('type.num.num',X7)))),
  file('thm.integer.integer_326_prove_ltb','thm.integer.integer_326_prove_ltb')).
fof(f24161,plain,(
  $false | ($spl4 | $spl9 | $spl238)),
  inference(subsumption_resolution,[],[f24160,f641])).
fof(f24160,plain,(
  p(s(bool,'F_0')) | ($spl4 | $spl238)),
  inference(forward_demodulation,[],[f24159,f391])).
fof(f391,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f124])).
fof(f124,plain,(
  ! [X0] : s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f74])).
fof(f74,axiom,(
  ! [X9] : s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num','const.num.|48|_0')))),
  file('thm.integer.integer_326_prove_ltb','thm.numeral.numeral_distrib|split|20')).
fof(f24159,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | ($spl4 | $spl238)),
  inference(forward_demodulation,[],[f23769,f637])).
fof(f632,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK10),s('type.num.num',sK11))),s('type.num.num',sK11)))) | $spl0),
  inference(cnf_transformation,[],[f632_D])).
fof(f632_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK10),s('type.num.num',sK11))),s('type.num.num',sK11)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f634,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK10) | $spl2),
  inference(cnf_transformation,[],[f634_D])).
fof(f634_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK10) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f365,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK10) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK10),s('type.num.num',sK11))),s('type.num.num',sK11))))),
  inference(cnf_transformation,[],[f265])).
fof(f641,plain,(
  ~p(s(bool,'F_0')) | $spl9),
  inference(cnf_transformation,[],[f641_D])).
fof(f641_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f377,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f110])).
fof(f110,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.integer.integer_326_prove_ltb',aHL_FALSITY)).
fof(f636,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK11) | $spl5),
  inference(cnf_transformation,[],[f636_D])).
fof(f636_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK11) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f366,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK11) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK10),s('type.num.num',sK11))),s('type.num.num',sK11))))),
  inference(cnf_transformation,[],[f265])).
fof(f23618,plain,(
  $false | ($spl1 | $spl2 | $spl5 | $spl9)),
  inference(subsumption_resolution,[],[f23617,f641])).
fof(f23617,plain,(
  p(s(bool,'F_0')) | ($spl1 | $spl2 | $spl5 | $spl9)),
  inference(forward_demodulation,[],[f23616,f391])).
fof(f23616,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK11),s('type.num.num','const.num.|48|_0')))) | ($spl1 | $spl2 | $spl5 | $spl9)),
  inference(subsumption_resolution,[],[f23615,f636])).
fof(f23615,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK11),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK11) | ($spl1 | $spl2 | $spl9)),
  inference(subsumption_resolution,[],[f23414,f641])).
fof(f23414,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK11),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK11) | ($spl1 | $spl2)),
  inference(superposition,[],[f23127,f3544])).
fof(f3544,plain,(
  s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))) | ($spl1 | $spl2)),
  inference(resolution,[],[f3543,f422])).
fof(f422,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f272])).
fof(f272,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f150])).
fof(f150,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f149])).
fof(f149,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X13] : (s(bool,'F_0') = s(bool,X13) <=> ~p(s(bool,X13)))),
  file('thm.integer.integer_326_prove_ltb','thm.bool.EQ_CLAUSES|split|2')).
fof(f3543,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11)))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f3542,f381])).
fof(f3542,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))),s('type.num.num',sK11)))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f633,f634])).
fof(f23127,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))) | s('type.num.num',X0) = s('type.num.num',X1)) )),
  inference(forward_demodulation,[],[f23126,f472])).
fof(f472,plain,(
  ( ! [X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f174])).
fof(f174,plain,(
  ! [X0,X1] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X8,X9] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X8),s('type.num.num',X9))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.num.num',X9)))),
  file('thm.integer.integer_326_prove_ltb','thm.arithmetic.LESS_EQ')).
fof(f23126,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))) | s('type.num.num',X0) = s('type.num.num',X1)) )),
  inference(forward_demodulation,[],[f534,f472])).
fof(f534,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))) | s('type.num.num',X0) = s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f325])).
fof(f325,plain,(
  ! [X0,X1] : ((s('type.num.num',X0) = s('type.num.num',X1) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0))))) & ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))) & ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0))))) | s('type.num.num',X0) != s('type.num.num',X1)))),
  inference(flattening,[],[f324])).
fof(f324,plain,(
  ! [X0,X1] : ((s('type.num.num',X0) = s('type.num.num',X1) | (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0)))))) & ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))) & ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0))))) | s('type.num.num',X0) != s('type.num.num',X1)))),
  inference(nnf_transformation,[],[f204])).
fof(f204,plain,(
  ! [X0,X1] : (s('type.num.num',X0) != s('type.num.num',X1) <=> (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0))))))),
  inference(flattening,[],[f203])).
fof(f203,plain,(
  ! [X0,X1] : (~s('type.num.num',X0) = s('type.num.num',X1) <=> (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0))))))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X8,X9] : (~s('type.num.num',X8) = s('type.num.num',X9) <=> (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.num.num',X9)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))),s('type.num.num',X8))))))),
  file('thm.integer.integer_326_prove_ltb','thm.arithmetic.NOT_NUM_EQ')).
fof(f24315,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f24262,f635])).
fof(f24262,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK10) | $spl0),
  inference(resolution,[],[f632,f505])).
fof(f505,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.num.num',X0)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f312])).
fof(f312,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.num.num',X0)))) | (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X1))) & (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.num.num',X0))))))),
  inference(flattening,[],[f311])).
fof(f311,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.num.num',X0)))) | (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X1))) & ((~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.num.num',X0))))))),
  inference(nnf_transformation,[],[f194])).
fof(f194,plain,(
  ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.num.num',X0)))) <=> (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X9,X8] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X8),s('type.num.num',X9))),s('type.num.num',X9)))) <=> (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X9)))) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X8)))),
  file('thm.integer.integer_326_prove_ltb','thm.arithmetic.LT_MULT_CANCEL_RBARE|split|1')).
% SZS output end Proof for thm.integer
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 852
% Passive clauses: 4681
% Generated clauses: 23958
% Final active clauses: 795
% Final passive clauses: 3475
% Input formulas: 104
% Initial clauses: 261
% 
% Duplicate literals: 215
% Trivial inequalities: 58
% Fw subsumption resolutions: 1546
% Fw demodulations: 2060
% Bw demodulations: 10
% 
% Simple tautologies: 207
% Equational tautologies: 25
% Forward subsumptions: 15024
% Fw demodulations to eq. taut.: 2
% 
% Binary resolution: 5494
% Factoring: 1181
% Forward superposition: 6379
% Backward superposition: 6561
% Self superposition: 31
% Equality factoring: 2
% Equality resolution: 12
% 
% Split clauses: 204
% Split components: 409
% SAT solver clauses: 369
% SAT solver unit clauses: 15
% SAT solver binary clauses: 50
% SAT solver learnt clauses: 3
% SAT solver learnt literals: 2
% 
% Sat splits: 204
% Sat splitting refutations: 10
% 
% TWLsolver clauses: 366
% TWLsolver calls for satisfiability: 328
% 
% Memory used [KB]: 12409
% Time elapsed: 0.414 s
% ------------------------------
% ------------------------------
% Success in time 0.419 s
