Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.rat
% (59408)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 160
% Passive clauses: 2579
% Generated clauses: 25074
% Final active clauses: 131
% Final passive clauses: 2302
% Input formulas: 11
% Initial clauses: 15
% 
% Duplicate literals: 1591
% Trivial inequalities: 298
% Fw subsumption resolutions: 2894
% Fw demodulations: 3087
% Bw demodulations: 25
% 
% Simple tautologies: 383
% Equational tautologies: 69
% Forward subsumptions: 12064
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 80
% Forward superposition: 9347
% Backward superposition: 6990
% Self superposition: 244
% Equality factoring: 180
% Equality resolution: 17
% 
% Split clauses: 1528
% Split components: 3112
% SAT solver clauses: 2221
% SAT solver unit clauses: 20
% SAT solver binary clauses: 260
% 
% Sat splits: 1528
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 2210
% TWLsolver calls for satisfiability: 148
% 
% Memory used [KB]: 13816
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.rat
% (59429)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 113
% Passive clauses: 673
% Generated clauses: 10885
% Final active clauses: 79
% Final passive clauses: 328
% Input formulas: 11
% Initial clauses: 15
% 
% Duplicate literals: 1321
% Trivial inequalities: 92
% Fw subsumption resolutions: 604
% Fw demodulations: 540
% Bw demodulations: 235
% Condensations: 1563
% 
% Simple tautologies: 138
% Equational tautologies: 49
% Forward subsumptions: 4442
% Bw demodulations to eq. taut.: 18
% 
% Binary resolution: 14
% Forward superposition: 2975
% Backward superposition: 2539
% Self superposition: 207
% Equality factoring: 142
% Equality resolution: 8
% 
% Split clauses: 204
% Split components: 412
% SAT solver clauses: 1248
% SAT solver unit clauses: 23
% SAT solver binary clauses: 82
% 
% Sat splits: 204
% Sat splitting refutations: 6
% 
% TWLsolver clauses: 1248
% TWLsolver calls for satisfiability: 118
% 
% Memory used [KB]: 8571
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.rat
% (59456)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 678
% Passive clauses: 10430
% Generated clauses: 136364
% Final active clauses: 49
% Final passive clauses: 394
% Input formulas: 11
% Initial clauses: 15
% 
% Duplicate literals: 6219
% Trivial inequalities: 709
% Fw subsumption resolutions: 5413
% Fw demodulations: 27621
% Bw demodulations: 7223
% Condensations: 7132
% Global subsumptions: 1583
% 
% Simple tautologies: 670
% Equational tautologies: 4662
% Forward subsumptions: 69018
% Fw demodulations to eq. taut.: 28
% Bw demodulations to eq. taut.: 403
% 
% Unit resulting resolution: 19341
% Forward superposition: 24824
% Backward superposition: 15494
% Self superposition: 493
% Equality factoring: 259
% Equality resolution: 50
% 
% Split clauses: 784
% Split components: 1575
% SAT solver clauses: 15477
% SAT solver unit clauses: 166
% SAT solver binary clauses: 1097
% 
% Sat splits: 784
% Sat splitting refutations: 1894
% 
% Memory used [KB]: 68570
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.rat
% (59703)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 2075
% Passive clauses: 29274
% Generated clauses: 132314
% Final active clauses: 299
% Final passive clauses: 13186
% Input formulas: 11
% Initial clauses: 13
% 
% Selected by SInE selection: 9
% SInE iterations: 1
% 
% Duplicate literals: 1931
% Trivial inequalities: 3648
% Fw subsumption resolutions: 10331
% Fw demodulations: 17795
% Bw demodulations: 847
% 
% Simple tautologies: 135
% Equational tautologies: 18
% Forward subsumptions: 67238
% Fw demodulations to eq. taut.: 17
% Bw demodulations to eq. taut.: 314
% 
% Binary resolution: 4900
% Forward superposition: 50693
% Backward superposition: 32976
% Self superposition: 100
% Equality factoring: 2
% Equality resolution: 277
% 
% Split clauses: 4750
% Split components: 9575
% SAT solver clauses: 4922
% SAT solver unit clauses: 3
% SAT solver binary clauses: 479
% 
% Sat splits: 4750
% Sat splitting refutations: 159
% 
% Memory used [KB]: 51299
% Time elapsed: 4.0000 s
% ------------------------------
% ------------------------------
% remaining time: 142 next slice time: 6
lrs+1002_2:3_br=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:sas=minisat:stl=30:sac=on:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none:sp=reverse_arity:urr=on_4 on thm.rat
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.rat
% SZS output start Proof for thm.rat
fof(f11286,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f53,f70,f68,f606,f648,f698,f709,f296,f753,f860,f870,f944,f962,f967,f969,f300,f752,f1267,f96,f1631,f1629,f1543,f2170,f2198,f2547,f3370,f3627,f3626,f3647,f3653,f3655,f3662,f3665,f3671,f3673,f3675,f3677,f3681,f3683,f3685,f3687,f3689,f3691,f3694,f3962,f3963,f3964,f3965,f5335,f5499,f647,f5573,f7519,f7518,f443,f7522,f7521,f440,f8173,f8172,f9790,f10628,f10824,f10826])).
fof(f10826,plain,(
  $false | ($spl10 | $spl14 | $spl17 | $spl26)),
  inference(subsumption_resolution,[],[f10825,f9805])).
fof(f9805,plain,(
  s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl26),
  inference(unit_resulting_resolution,[],[f7517,f43])).
fof(f43,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,sK1(X0)))) )),
  inference(cnf_transformation,[],[f30])).
fof(f30,plain,(
  ! [X0] : ((~p(s(bool,sK1(X0))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,sK1(X0)))) & s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X0))) = s('type.integer.int','const.bool.COND_3'(s(bool,sK1(X0)),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK1])],[f29])).
fof(f29,plain,(
  ! [X0] : ? [X1] : ((~p(s(bool,X1)) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,X1))) & s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X0))) = s('type.integer.int','const.bool.COND_3'(s(bool,X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))))),
  inference(flattening,[],[f28])).
fof(f28,plain,(
  ! [X0] : ? [X1] : (((~p(s(bool,X1)) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,X1)))) & s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X0))) = s('type.integer.int','const.bool.COND_3'(s(bool,X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))))),
  inference(nnf_transformation,[],[f20])).
fof(f20,plain,(
  ! [X0] : ? [X1] : ((p(s(bool,X1)) <=> s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) & s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X0))) = s('type.integer.int','const.bool.COND_3'(s(bool,X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X8] : ? [X9] : ((p(s(bool,X9)) <=> s('type.integer.int',X8) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) & s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X8))) = s('type.integer.int','const.bool.COND_3'(s(bool,X9),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X8),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))))),
  file('thm.rat.RAT_SGN_CONG','thm.intExtension.SGN_def')).
fof(f7517,plain,(
  p(s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac',sK0))))) | $spl26),
  inference(cnf_transformation,[],[f7517_D])).
fof(f7517_D,plain,(
  p(s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac',sK0))))) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f10825,plain,(
  s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ($spl10 | $spl14 | $spl17 | $spl26)),
  inference(subsumption_resolution,[],[f10148,f5572])).
fof(f5572,plain,(
  s(bool,'T_0') != s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0)))))) | ($spl10 | $spl17)),
  inference(unit_resulting_resolution,[],[f33,f5506,f2685])).
fof(f2685,plain,(
  ( ! [X6,X7] : (s(bool,X6) != s(bool,X7) | p(s(bool,X7)) | ~p(s(bool,X6))) ) | $spl10),
  inference(superposition,[],[f753,f50])).
fof(f50,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(superposition,[],[f35,f38])).
fof(f38,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.rat.RAT_SGN_CONG',aHL_BOOL_CASES)).
fof(f35,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f14])).
fof(f14,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.rat.RAT_SGN_CONG',aHL_FALSITY)).
fof(f5506,plain,(
  ~p(s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))))))) | $spl17),
  inference(unit_resulting_resolution,[],[f1631,f85])).
fof(f85,plain,(
  ( ! [X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X1))) | ~p(s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X1)))))))) )),
  inference(superposition,[],[f36,f43])).
fof(f36,plain,(
  ( ! [X0] : (s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X0))) = s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0)))))) )),
  inference(cnf_transformation,[],[f15])).
fof(f15,plain,(
  ! [X0] : s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X0))) = s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0)))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6] : s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X6))) = s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X6)))))),
  file('thm.rat.RAT_SGN_CONG','thm.frac.frac_sgn_def')).
fof(f33,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.rat.RAT_SGN_CONG',aHL_TRUTH)).
fof(f10148,plain,(
  s(bool,'T_0') = s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0)))))) | s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ($spl14 | $spl26)),
  inference(backward_demodulation,[],[f9805,f7289])).
fof(f7289,plain,(
  s(bool,'T_0') = s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl14),
  inference(superposition,[],[f5616,f40])).
fof(f40,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X0))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0)))) )),
  inference(cnf_transformation,[],[f26])).
fof(f26,plain,(
  ! [X0] : ((s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X0))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0)))) & (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X0)))))))))),
  inference(nnf_transformation,[],[f17])).
fof(f17,plain,(
  ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X0))))))) <=> s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X6] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X6))))))) <=> s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X6))))),
  file('thm.rat.RAT_SGN_CONG','thm.rat.RAT_NMREQ|48|_CONG')).
fof(f5616,plain,(
  s(bool,'T_0') = s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0)))))))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1628,f1626])).
fof(f1626,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0)))))) | ~p(s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac',X0)))))) )),
  inference(subsumption_resolution,[],[f1611,f35])).
fof(f1611,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | ~p(s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac',X0))))) | s(bool,'T_0') = s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0))))))) )),
  inference(superposition,[],[f1605,f38])).
fof(f1605,plain,(
  ( ! [X6] : (p(s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X6))))))) | ~p(s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac',X6)))))) )),
  inference(trivial_inequality_removal,[],[f1604])).
fof(f1604,plain,(
  ( ! [X6] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X6))))))) | ~p(s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac',X6)))))) )),
  inference(superposition,[],[f86,f43])).
fof(f86,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X0))) | p(s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0)))))))) )),
  inference(superposition,[],[f44,f36])).
fof(f44,plain,(
  ( ! [X0] : (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,sK1(X0)))) )),
  inference(cnf_transformation,[],[f30])).
fof(f1628,plain,(
  p(s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))))))))) | $spl14),
  inference(cnf_transformation,[],[f1628_D])).
fof(f1628_D,plain,(
  p(s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))))))))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f10824,plain,(
  $false | ($spl14 | $spl17 | $spl26)),
  inference(subsumption_resolution,[],[f10823,f9812])).
fof(f9812,plain,(
  s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',sK0))) != s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))) | ($spl17 | $spl26)),
  inference(unit_resulting_resolution,[],[f7517,f5508])).
fof(f5508,plain,(
  ( ! [X1] : (s('type.integer.int',X1) != s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',sK0))) | ~p(s(bool,sK1(X1)))) ) | $spl17),
  inference(superposition,[],[f1631,f43])).
fof(f10823,plain,(
  s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',sK0))) = s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))) | ($spl14 | $spl26)),
  inference(forward_demodulation,[],[f10822,f36])).
fof(f10822,plain,(
  s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))) = s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))))) | ($spl14 | $spl26)),
  inference(subsumption_resolution,[],[f10146,f9805])).
fof(f10146,plain,(
  s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))) = s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))))) | s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ($spl14 | $spl26)),
  inference(backward_demodulation,[],[f9805,f7248])).
fof(f7248,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl14),
  inference(superposition,[],[f5624,f197])).
fof(f197,plain,(
  ( ! [X1] : (s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X1))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X1)))) )),
  inference(superposition,[],[f36,f40])).
fof(f5624,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1628,f43])).
fof(f10628,plain,(
  $false | $spl26),
  inference(subsumption_resolution,[],[f10627,f9805])).
fof(f10627,plain,(
  s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl26),
  inference(subsumption_resolution,[],[f9949,f36])).
fof(f9949,plain,(
  s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',sK0))) != s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))))) | s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl26),
  inference(backward_demodulation,[],[f9805,f1875])).
fof(f1875,plain,(
  s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',sK0))) != s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f32,f197])).
fof(f32,plain,(
  s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',sK0))) != s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0)))))))),
  inference(cnf_transformation,[],[f25])).
fof(f25,plain,(
  s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',sK0))) != s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f23])).
fof(f23,plain,(
  ? [X0] : s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X0))) != s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X0)))))))),
  inference(ennf_transformation,[],[f13])).
fof(f13,plain,(
  ~! [X0] : s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X0))) = s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X0)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X6))) = s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X6)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X6))) = s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X6)))))))),
  file('thm.rat.RAT_SGN_CONG','thm.rat.RAT_SGN_CONG')).
fof(f9790,plain,(
  $false | ($spl10 | $spl27)),
  inference(subsumption_resolution,[],[f9789,f32])).
fof(f9789,plain,(
  s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',sK0))) = s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))))))) | ($spl10 | $spl27)),
  inference(forward_demodulation,[],[f9701,f7666])).
fof(f7666,plain,(
  s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',sK0))) = s('type.integer.int','const.bool.COND_3'(s(bool,'F_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | ($spl10 | $spl27)),
  inference(forward_demodulation,[],[f7621,f36])).
fof(f7621,plain,(
  s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))))) = s('type.integer.int','const.bool.COND_3'(s(bool,'F_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | ($spl10 | $spl27)),
  inference(superposition,[],[f45,f7542])).
fof(f7542,plain,(
  s(bool,'F_0') = s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac',sK0)))) | ($spl10 | $spl27)),
  inference(unit_resulting_resolution,[],[f7518,f5500])).
fof(f5500,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,X1) | p(s(bool,X1))) ) | $spl10),
  inference(global_subsumption,[],[f32,f33,f36,f38,f40,f39,f46,f48,f50,f53,f57,f76,f75,f89,f208,f209,f219,f246,f265,f266,f267,f279,f296,f676,f197,f668,f712,f96,f154,f300,f250,f35,f1563,f1546,f1605,f1729,f1803,f1875,f1950,f1956,f1534,f753,f1606,f2683,f2684,f2685,f2174,f1977,f1531,f1499,f514,f291,f249,f247,f253,f2150,f1595,f649,f86,f85,f1870,f1876,f1601,f669,f661,f641,f631,f628,f84,f632,f1802,f1730,f204,f203,f141,f2549,f2137,f2104,f1626,f1625,f1610,f1423,f1422,f1409,f971,f634,f623,f2101,f1879,f1873,f1607,f1406,f261,f2169,f2167,f2102,f1975,f1884,f1883,f1871,f1608,f1407,f646,f645,f639,f638,f629,f621,f259,f251,f187,f182,f155,f146,f73,f54,f44,f43,f671,f422,f421,f411,f635,f1880,f457,f352,f351,f346,f343,f37,f423,f642,f1529,f1527,f1497,f1495,f513,f511,f455,f454,f451,f449,f447,f446,f444,f443,f441,f440,f438,f436,f434,f433,f432,f431,f425,f417,f415,f414,f413,f412,f344,f45,f2059,f2060,f2066,f657,f2063,f660,f715,f426,f2075,f2074,f2073,f2072,f2070,f2069,f2068,f2067,f2052,f347,f456,f2051,f2054,f1532,f1500,f515,f292,f258,f257,f255,f254,f1971,f1970,f1969,f1968,f1966,f1965,f1964,f1963,f2163,f2162,f2161,f2160,f2158,f2157,f2156,f2155,f5367,f5339,f5334,f5253,f5254,f5332,f5265,f5266,f5267,f5424,f5309,f5310,f5321,f5322,f5323,f5464,f5465,f5480,f5489,f5490,f5331,f5407,f5381])).
fof(f5381,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0) | p(s(bool,X2)) | p(s(bool,X1))) ) | $spl10),
  inference(subsumption_resolution,[],[f5212,f219])).
fof(f5212,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0) | p(s(bool,X2)) | s(bool,X0) != s(bool,X2) | p(s(bool,X1))) ) | $spl10),
  inference(superposition,[],[f38,f2683])).
fof(f5407,plain,(
  ( ! [X171,X169,X170] : (s(bool,'T_0') = s(bool,X170) | ~p(s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac',X169))))) | p(s(bool,X171)) | s(bool,X171) != s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X169))))))) ) | $spl10),
  inference(subsumption_resolution,[],[f5259,f50])).
fof(f5259,plain,(
  ( ! [X171,X169,X170] : (s(bool,'T_0') = s(bool,X170) | ~p(s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac',X169))))) | p(s(bool,X171)) | s(bool,X171) != s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X169)))))) | p(s(bool,X170))) ) | $spl10),
  inference(superposition,[],[f1626,f2683])).
fof(f5331,plain,(
  ( ! [X173,X174,X172] : (p(s(bool,X173)) | ~p(s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac',X172))))) | p(s(bool,X174)) | s(bool,X174) != s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X172))))))) ) | $spl10),
  inference(duplicate_literal_removal,[],[f5260])).
fof(f5260,plain,(
  ( ! [X173,X174,X172] : (p(s(bool,X173)) | ~p(s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac',X172))))) | p(s(bool,X174)) | s(bool,X174) != s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X172)))))) | p(s(bool,X173))) ) | $spl10),
  inference(superposition,[],[f2549,f2683])).
fof(f5490,plain,(
  ( ! [X68,X66,X67] : (s(bool,X68) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X66))),s('type.integer.int',X67))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X66))))))),s('type.integer.int',X67)))) | ~p(s(bool,X68)) | ~p(s(bool,sK1(X67)))) ) | $spl10),
  inference(superposition,[],[f2685,f347])).
fof(f5489,plain,(
  ( ! [X64,X65,X63] : (s(bool,X65) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X63))),s('type.integer.int',X64))) | p(s(bool,X65)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X63))))))),s('type.integer.int',X64)))) | ~p(s(bool,sK1(X64)))) ) | $spl10),
  inference(superposition,[],[f2685,f347])).
fof(f5480,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X34))),s('type.integer.int',X35))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X34))))))),s('type.integer.int',X35)))) | ~p(s(bool,sK1(X35)))) ) | $spl10),
  inference(superposition,[],[f753,f347])).
fof(f5465,plain,(
  ( ! [X59,X57,X58,X56] : (s(bool,X58) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X56))),s('type.integer.int',X57))) | ~p(s(bool,sK1(X57))) | p(s(bool,X59)) | s(bool,X58) != s(bool,X59) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X56))))))),s('type.integer.int',X57))))) ) | $spl10),
  inference(superposition,[],[f347,f2683])).
fof(f5464,plain,(
  ( ! [X54,X52,X55,X53] : (s(bool,X54) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X52))),s('type.integer.int',X53))) | ~p(s(bool,sK1(X53))) | p(s(bool,X55)) | s(bool,X55) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X52))))))),s('type.integer.int',X53))) | p(s(bool,X54))) ) | $spl10),
  inference(superposition,[],[f347,f2683])).
fof(f5323,plain,(
  ( ! [X202,X200,X199,X201] : (s(bool,X199) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X200) = s(fun(bool,bool),X201) | s(bool,'T_0') = s(bool,sK2(X200,X201,bool,bool)) | p(s(bool,X202)) | s(bool,X199) != s(bool,X202) | p(s(bool,g(s(bool,'F_0'))))) ) | $spl10),
  inference(superposition,[],[f250,f2683])).
fof(f5322,plain,(
  ( ! [X198,X196,X194,X197,X195] : (s(bool,X197) != s(bool,f(s(X194,sK2(X195,X196,bool,X194)))) | s(fun(X194,bool),X195) = s(fun(X194,bool),X196) | p(s(bool,X198)) | s(bool,X197) != s(bool,X198) | p(s(bool,g(s(X194,sK2(X195,X196,bool,X194)))))) ) | $spl10),
  inference(superposition,[],[f46,f2683])).
fof(f5321,plain,(
  ( ! [X191,X189,X192,X190,X193] : (s(X191,f(s(bool,X192))) != s(X191,g(s(bool,X192))) | s(fun(bool,X191),X189) = s(fun(bool,X191),X190) | p(s(bool,X193)) | s(bool,X192) != s(bool,X193) | p(s(bool,sK2(X189,X190,X191,bool)))) ) | $spl10),
  inference(superposition,[],[f46,f2683])).
fof(f5310,plain,(
  ( ! [X158,X156,X157] : (s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X156))) = s('type.integer.int','const.bool.COND_3'(s(bool,X157),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X156),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | p(s(bool,X158)) | s(bool,X157) != s(bool,X158) | p(s(bool,sK1(X156)))) ) | $spl10),
  inference(superposition,[],[f45,f2683])).
fof(f5309,plain,(
  ( ! [X154,X155,X153] : (s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X153))) = s('type.integer.int','const.bool.COND_3'(s(bool,sK1(X153)),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,X154),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | p(s(bool,X155)) | s(bool,X154) != s(bool,X155) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X153),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl10),
  inference(superposition,[],[f45,f2683])).
fof(f5424,plain,(
  ( ! [X134,X132,X136,X135,X133] : (s(X133,f(s(bool,X132))) != s(X133,g(s(bool,X132))) | s(fun(bool,X133),X134) = s(fun(bool,X133),X135) | s(bool,'T_0') = s(bool,sK2(X134,X135,X133,bool)) | p(s(bool,X136)) | s(bool,X132) != s(bool,X136)) ) | $spl10),
  inference(subsumption_resolution,[],[f5303,f35])).
fof(f5303,plain,(
  ( ! [X134,X132,X136,X135,X133] : (s(X133,f(s(bool,X132))) != s(X133,g(s(bool,X132))) | s(fun(bool,X133),X134) = s(fun(bool,X133),X135) | s(bool,'T_0') = s(bool,sK2(X134,X135,X133,bool)) | p(s(bool,X136)) | s(bool,X132) != s(bool,X136) | p(s(bool,'F_0'))) ) | $spl10),
  inference(superposition,[],[f250,f2683])).
fof(f5267,plain,(
  ( ! [X202,X200,X199,X201] : (s(bool,X199) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X200) = s(fun(bool,bool),X201) | s(bool,'T_0') = s(bool,sK2(X200,X201,bool,bool)) | p(s(bool,X202)) | s(bool,X202) != s(bool,g(s(bool,'F_0'))) | p(s(bool,X199))) ) | $spl10),
  inference(superposition,[],[f250,f2683])).
fof(f5266,plain,(
  ( ! [X198,X196,X194,X197,X195] : (s(bool,X197) != s(bool,f(s(X194,sK2(X195,X196,bool,X194)))) | s(fun(X194,bool),X195) = s(fun(X194,bool),X196) | p(s(bool,X198)) | s(bool,X198) != s(bool,g(s(X194,sK2(X195,X196,bool,X194)))) | p(s(bool,X197))) ) | $spl10),
  inference(superposition,[],[f46,f2683])).
fof(f5265,plain,(
  ( ! [X191,X189,X192,X190,X193] : (s(X191,f(s(bool,X192))) != s(X191,g(s(bool,X192))) | s(fun(bool,X191),X189) = s(fun(bool,X191),X190) | p(s(bool,X193)) | s(bool,X193) != s(bool,sK2(X189,X190,X191,bool)) | p(s(bool,X192))) ) | $spl10),
  inference(superposition,[],[f46,f2683])).
fof(f5332,plain,(
  ( ! [X165,X163,X164] : (p(s(bool,X164)) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X163))) | p(s(bool,X165)) | s(bool,X165) != s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X163))))))))) ) | $spl10),
  inference(duplicate_literal_removal,[],[f5257])).
fof(f5257,plain,(
  ( ! [X165,X163,X164] : (p(s(bool,X164)) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X163))) | p(s(bool,X165)) | s(bool,X165) != s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X163)))))))) | p(s(bool,X164))) ) | $spl10),
  inference(superposition,[],[f203,f2683])).
fof(f5254,plain,(
  ( ! [X158,X156,X157] : (s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X156))) = s('type.integer.int','const.bool.COND_3'(s(bool,X157),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X156),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | p(s(bool,X158)) | s(bool,X158) != s(bool,sK1(X156)) | p(s(bool,X157))) ) | $spl10),
  inference(superposition,[],[f45,f2683])).
fof(f5253,plain,(
  ( ! [X154,X155,X153] : (s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X153))) = s('type.integer.int','const.bool.COND_3'(s(bool,sK1(X153)),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,X154),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | p(s(bool,X155)) | s(bool,X155) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',X153),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,X154))) ) | $spl10),
  inference(superposition,[],[f45,f2683])).
fof(f5334,plain,(
  ( ! [X70,X74,X72,X71,X73] : (s(bool,X71) = s(bool,X72) | p(s(bool,X73)) | s(bool,X72) != s(bool,X73) | p(s(bool,X70)) | p(s(bool,X74)) | s(bool,X71) != s(bool,X74)) ) | $spl10),
  inference(duplicate_literal_removal,[],[f5196])).
fof(f5196,plain,(
  ( ! [X70,X74,X72,X71,X73] : (s(bool,X71) = s(bool,X72) | p(s(bool,X73)) | s(bool,X72) != s(bool,X73) | p(s(bool,X70)) | p(s(bool,X74)) | s(bool,X71) != s(bool,X74) | p(s(bool,X70))) ) | $spl10),
  inference(superposition,[],[f2683,f2683])).
fof(f5339,plain,(
  ( ! [X30,X28,X31,X29,X27] : (s(bool,X28) = s(bool,X29) | p(s(bool,X30)) | s(bool,X29) != s(bool,X30) | p(s(bool,X27)) | s(bool,X28) = s(bool,X31) | ~p(s(bool,X31))) ) | $spl10),
  inference(duplicate_literal_removal,[],[f5186])).
fof(f5186,plain,(
  ( ! [X30,X28,X31,X29,X27] : (s(bool,X28) = s(bool,X29) | p(s(bool,X30)) | s(bool,X29) != s(bool,X30) | p(s(bool,X27)) | s(bool,X28) = s(bool,X31) | p(s(bool,X27)) | ~p(s(bool,X31))) ) | $spl10),
  inference(superposition,[],[f2683,f267])).
fof(f5367,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X1) | p(s(bool,X2)) | s(bool,X1) != s(bool,X2) | s(bool,'T_0') = s(bool,X0)) ) | $spl10),
  inference(subsumption_resolution,[],[f5179,f50])).
fof(f5179,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X1) | p(s(bool,X2)) | s(bool,X1) != s(bool,X2) | p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl10),
  inference(superposition,[],[f2683,f38])).
fof(f2155,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,f(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(fun('type.integer.int',bool),X0) = s(fun('type.integer.int',bool),X1) | ~p(s(bool,sK1(sK2(X0,X1,bool,'type.integer.int')))) | s(bool,'T_0') = s(bool,g(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f253,f38])).
fof(f2156,plain,(
  ( ! [X4,X2,X3] : (s(bool,'T_0') != s(bool,f(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(fun('type.integer.int',bool),X2) = s(fun('type.integer.int',bool),X3) | ~p(s(bool,sK1(sK2(X2,X3,bool,'type.integer.int')))) | s(bool,X4) = s(bool,g(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f253,f48])).
fof(f2157,plain,(
  ( ! [X6,X5] : (s(bool,'T_0') != s(bool,f(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(fun('type.integer.int',bool),X5) = s(fun('type.integer.int',bool),X6) | ~p(s(bool,sK1(sK2(X5,X6,bool,'type.integer.int')))) | ~p(s(bool,g(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f253,f50])).
fof(f2158,plain,(
  ( ! [X8,X7,X9] : (s(bool,X7) != s(bool,f(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(fun('type.integer.int',bool),X8) = s(fun('type.integer.int',bool),X9) | ~p(s(bool,sK1(sK2(X8,X9,bool,'type.integer.int')))) | ~p(s(bool,X7)) | ~p(s(bool,g(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f253,f57])).
fof(f2160,plain,(
  ( ! [X14,X15,X13] : (s(bool,'T_0') != s(bool,f(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(fun('type.integer.int',bool),X13) = s(fun('type.integer.int',bool),X14) | ~p(s(bool,sK1(sK2(X13,X14,bool,'type.integer.int')))) | s(bool,X15) = s(bool,g(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,X15))) )),
  inference(superposition,[],[f253,f219])).
fof(f2161,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,X16) != s(bool,f(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(fun('type.integer.int',bool),X17) = s(fun('type.integer.int',bool),X18) | ~p(s(bool,sK1(sK2(X17,X18,bool,'type.integer.int')))) | s(bool,X19) = s(bool,g(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,X16)) | ~p(s(bool,X19))) )),
  inference(superposition,[],[f253,f267])).
fof(f2162,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,X20) != s(bool,f(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(fun('type.integer.int',bool),X21) = s(fun('type.integer.int',bool),X22) | ~p(s(bool,sK1(sK2(X21,X22,bool,'type.integer.int')))) | s(bool,X20) = s(bool,X23) | p(s(bool,g(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,X23))) )),
  inference(superposition,[],[f253,f267])).
fof(f2163,plain,(
  ( ! [X26,X24,X25] : (s(bool,X24) != s(bool,f(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(fun('type.integer.int',bool),X25) = s(fun('type.integer.int',bool),X26) | ~p(s(bool,sK1(sK2(X25,X26,bool,'type.integer.int')))) | p(s(bool,X24)) | p(s(bool,g(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f253,f279])).
fof(f1963,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK2(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f250,f38])).
fof(f1964,plain,(
  ( ! [X4,X2,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,sK2(X2,X3,bool,bool)) | s(bool,X4) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f250,f48])).
fof(f1965,plain,(
  ( ! [X6,X5] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X5) = s(fun(bool,bool),X6) | s(bool,'T_0') = s(bool,sK2(X5,X6,bool,bool)) | ~p(s(bool,g(s(bool,'F_0'))))) )),
  inference(superposition,[],[f250,f50])).
fof(f1966,plain,(
  ( ! [X8,X7,X9] : (s(bool,X7) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X8) = s(fun(bool,bool),X9) | s(bool,'T_0') = s(bool,sK2(X8,X9,bool,bool)) | ~p(s(bool,X7)) | ~p(s(bool,g(s(bool,'F_0'))))) )),
  inference(superposition,[],[f250,f57])).
fof(f1968,plain,(
  ( ! [X14,X15,X13] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X13) = s(fun(bool,bool),X14) | s(bool,'T_0') = s(bool,sK2(X13,X14,bool,bool)) | s(bool,X15) = s(bool,g(s(bool,'F_0'))) | p(s(bool,X15))) )),
  inference(superposition,[],[f250,f219])).
fof(f1969,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,X16) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X17) = s(fun(bool,bool),X18) | s(bool,'T_0') = s(bool,sK2(X17,X18,bool,bool)) | s(bool,X19) = s(bool,g(s(bool,'F_0'))) | p(s(bool,X16)) | ~p(s(bool,X19))) )),
  inference(superposition,[],[f250,f267])).
fof(f1970,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,X20) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X21) = s(fun(bool,bool),X22) | s(bool,'T_0') = s(bool,sK2(X21,X22,bool,bool)) | s(bool,X20) = s(bool,X23) | p(s(bool,g(s(bool,'F_0')))) | ~p(s(bool,X23))) )),
  inference(superposition,[],[f250,f267])).
fof(f1971,plain,(
  ( ! [X26,X24,X25] : (s(bool,X24) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X25) = s(fun(bool,bool),X26) | s(bool,'T_0') = s(bool,sK2(X25,X26,bool,bool)) | p(s(bool,X24)) | p(s(bool,g(s(bool,'F_0'))))) )),
  inference(superposition,[],[f250,f279])).
fof(f254,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,f(s(X0,sK2(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,X3) = s(bool,g(s(X0,sK2(X1,X2,bool,X0)))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f46,f48])).
fof(f255,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X7) != s(bool,f(s(X4,sK2(X5,X6,bool,X4)))) | s(fun(X4,bool),X5) = s(fun(X4,bool),X6) | ~p(s(bool,g(s(X4,sK2(X5,X6,bool,X4))))) | ~p(s(bool,X7))) )),
  inference(superposition,[],[f46,f57])).
fof(f257,plain,(
  ( ! [X14,X12,X13] : (s(bool,'T_0') != s(bool,f(s(X12,sK2(X13,X14,bool,X12)))) | s(fun(X12,bool),X13) = s(fun(X12,bool),X14) | ~p(s(bool,g(s(X12,sK2(X13,X14,bool,X12)))))) )),
  inference(superposition,[],[f46,f50])).
fof(f258,plain,(
  ( ! [X17,X15,X16] : (s(bool,'T_0') = s(bool,g(s(X15,sK2(X16,X17,bool,X15)))) | s(fun(X15,bool),X16) = s(fun(X15,bool),X17) | s(bool,'F_0') != s(bool,f(s(X15,sK2(X16,X17,bool,X15))))) )),
  inference(superposition,[],[f46,f38])).
fof(f292,plain,(
  ( ! [X30,X28,X29,X27] : (s(bool,'T_0') != s(bool,f(s(X27,sK2(X28,X29,bool,X27)))) | s(fun(X27,bool),X28) = s(fun(X27,bool),X29) | s(bool,X30) = s(bool,g(s(X27,sK2(X28,X29,bool,X27)))) | p(s(bool,X30))) )),
  inference(superposition,[],[f46,f219])).
fof(f515,plain,(
  ( ! [X54,X56,X55,X53] : (s(bool,X56) != s(bool,f(s(X53,sK2(X54,X55,bool,X53)))) | s(fun(X53,bool),X54) = s(fun(X53,bool),X55) | p(s(bool,g(s(X53,sK2(X54,X55,bool,X53))))) | p(s(bool,X56))) )),
  inference(superposition,[],[f46,f279])).
fof(f1500,plain,(
  ( ! [X111,X109,X112,X110,X108] : (s(bool,X111) != s(bool,f(s(X108,sK2(X109,X110,bool,X108)))) | s(fun(X108,bool),X109) = s(fun(X108,bool),X110) | s(bool,X112) = s(bool,g(s(X108,sK2(X109,X110,bool,X108)))) | p(s(bool,X111)) | ~p(s(bool,X112))) )),
  inference(superposition,[],[f46,f267])).
fof(f1532,plain,(
  ( ! [X111,X109,X112,X110,X108] : (s(bool,X111) != s(bool,f(s(X108,sK2(X109,X110,bool,X108)))) | s(fun(X108,bool),X109) = s(fun(X108,bool),X110) | s(bool,X111) = s(bool,X112) | p(s(bool,g(s(X108,sK2(X109,X110,bool,X108))))) | ~p(s(bool,X112))) )),
  inference(superposition,[],[f46,f267])).
fof(f2054,plain,(
  ( ! [X8,X9] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X9))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X8))),s('type.integer.int',X9))) | ~p(s(bool,sK1(X9))) | ~p(s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X8)))))))))) )),
  inference(superposition,[],[f347,f43])).
fof(f2051,plain,(
  ( ! [X0,X1] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0))),s('type.integer.int',X1))) | ~p(s(bool,sK1(X1))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0)))) )),
  inference(superposition,[],[f347,f40])).
fof(f456,plain,(
  ( ! [X7] : (s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X7))) = s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',X7),s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X7),s('type.integer.int',X7))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | ~p(s(bool,sK1(X7)))) )),
  inference(forward_demodulation,[],[f424,f89])).
fof(f424,plain,(
  ( ! [X7] : (s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X7))) = s('type.integer.int','const.bool.COND_3'(s(bool,sK1('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),s('type.integer.int',X7),s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X7),s('type.integer.int',X7))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | ~p(s(bool,sK1(X7)))) )),
  inference(superposition,[],[f45,f43])).
fof(f347,plain,(
  ( ! [X0,X1] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X1))),s('type.integer.int',X0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X1))))))),s('type.integer.int',X0))) | ~p(s(bool,sK1(X0)))) )),
  inference(superposition,[],[f37,f43])).
fof(f2052,plain,(
  ( ! [X4,X2,X3] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',X3),s('type.integer.int',X4))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X2))),s('type.integer.int',X4))) | ~p(s(bool,sK1(X4))) | ~p(s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X2))))))))) | ~p(s(bool,sK1(X3)))) )),
  inference(superposition,[],[f347,f54])).
fof(f2067,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X3))),s('type.integer.int',X4))) | ~p(s(bool,sK1(X4))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X3))))))),s('type.integer.int',X4)))) )),
  inference(superposition,[],[f347,f38])).
fof(f2068,plain,(
  ( ! [X6,X7,X5] : (s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X5))),s('type.integer.int',X6))) | ~p(s(bool,sK1(X6))) | s(bool,X7) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X5))))))),s('type.integer.int',X6))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f347,f48])).
fof(f2069,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X8))),s('type.integer.int',X9))) | ~p(s(bool,sK1(X9))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X8))))))),s('type.integer.int',X9))))) )),
  inference(superposition,[],[f347,f50])).
fof(f2070,plain,(
  ( ! [X12,X10,X11] : (s(bool,X12) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X10))),s('type.integer.int',X11))) | ~p(s(bool,sK1(X11))) | ~p(s(bool,X12)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X10))))))),s('type.integer.int',X11))))) )),
  inference(superposition,[],[f347,f57])).
fof(f2072,plain,(
  ( ! [X17,X18,X16] : (s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X16))),s('type.integer.int',X17))) | ~p(s(bool,sK1(X17))) | s(bool,X18) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X16))))))),s('type.integer.int',X17))) | p(s(bool,X18))) )),
  inference(superposition,[],[f347,f219])).
fof(f2073,plain,(
  ( ! [X21,X19,X22,X20] : (s(bool,X21) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X19))),s('type.integer.int',X20))) | ~p(s(bool,sK1(X20))) | s(bool,X22) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X19))))))),s('type.integer.int',X20))) | p(s(bool,X21)) | ~p(s(bool,X22))) )),
  inference(superposition,[],[f347,f267])).
fof(f2074,plain,(
  ( ! [X26,X24,X23,X25] : (s(bool,X25) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X23))),s('type.integer.int',X24))) | ~p(s(bool,sK1(X24))) | s(bool,X25) = s(bool,X26) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X23))))))),s('type.integer.int',X24)))) | ~p(s(bool,X26))) )),
  inference(superposition,[],[f347,f267])).
fof(f2075,plain,(
  ( ! [X28,X29,X27] : (s(bool,X29) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X27))),s('type.integer.int',X28))) | ~p(s(bool,sK1(X28))) | p(s(bool,X29)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X27))))))),s('type.integer.int',X28))))) )),
  inference(superposition,[],[f347,f279])).
fof(f426,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X1))) = s('type.integer.int','const.bool.COND_3'(s(bool,sK1(X1)),s('type.integer.int',X0),s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | ~p(s(bool,sK1(X0)))) )),
  inference(superposition,[],[f45,f43])).
fof(f715,plain,(
  ( ! [X14] : (s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X14))))) = s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X14))),s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X14))),s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X14))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | ~p(s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X14)))))))) )),
  inference(forward_demodulation,[],[f663,f89])).
fof(f663,plain,(
  ( ! [X14] : (s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X14))))) = s('type.integer.int','const.bool.COND_3'(s(bool,sK1('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X14))),s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X14))),s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X14))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | ~p(s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X14)))))))) )),
  inference(superposition,[],[f45,f85])).
fof(f660,plain,(
  ( ! [X8,X9] : (s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X9))) = s('type.integer.int','const.bool.COND_3'(s(bool,sK1(X9)),s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X8))),s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X9),s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X8))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | ~p(s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X8)))))))) )),
  inference(superposition,[],[f45,f85])).
fof(f2063,plain,(
  ( ! [X19,X18] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X19))),s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X18))))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X19))))))),s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X18))))) | ~p(s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | ~p(s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac',X18)))))) )),
  inference(superposition,[],[f347,f84])).
fof(f657,plain,(
  ( ! [X0,X1] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X1))),s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X0))))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X1))))))),s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X0))))) | ~p(s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0)))))))) )),
  inference(superposition,[],[f37,f85])).
fof(f2066,plain,(
  ( ! [X24,X25] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X25))),s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X24))))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X25))))))),s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X24))))) | ~p(s(bool,sK1('const.bool.COND_3'(s(bool,sK1(X24)),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X24),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))))) )),
  inference(superposition,[],[f347,f45])).
fof(f2060,plain,(
  ( ! [X12,X13] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X13))),s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X13))))))),s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | ~p(s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X12))))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X12)))) )),
  inference(superposition,[],[f347,f197])).
fof(f2059,plain,(
  ( ! [X10,X11] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X11))),s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X11))))))),s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | ~p(s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac',X10))))) | ~p(s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac',X10)))))) )),
  inference(superposition,[],[f347,f84])).
fof(f344,plain,(
  ( ! [X2,X1] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X1))))))))) | ~p(s(bool,sK1(X2)))) )),
  inference(superposition,[],[f37,f54])).
fof(f412,plain,(
  ( ! [X0] : (s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X0))) = s('type.integer.int','const.bool.COND_3'(s(bool,'F_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | s(bool,'T_0') = s(bool,sK1(X0))) )),
  inference(superposition,[],[f45,f38])).
fof(f413,plain,(
  ( ! [X2,X1] : (s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X1))) = s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | s(bool,X2) = s(bool,sK1(X1)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f45,f48])).
fof(f414,plain,(
  ( ! [X3] : (s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X3))) = s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | ~p(s(bool,sK1(X3)))) )),
  inference(superposition,[],[f45,f50])).
fof(f415,plain,(
  ( ! [X4,X5] : (s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X4))) = s('type.integer.int','const.bool.COND_3'(s(bool,X5),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X4),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | ~p(s(bool,X5)) | ~p(s(bool,sK1(X4)))) )),
  inference(superposition,[],[f45,f57])).
fof(f417,plain,(
  ( ! [X8,X9] : (s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X8))) = s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X8),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | s(bool,X9) = s(bool,sK1(X8)) | p(s(bool,X9))) )),
  inference(superposition,[],[f45,f219])).
fof(f425,plain,(
  ( ! [X8] : (s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X8))))) = s('type.integer.int','const.bool.COND_3'(s(bool,sK1('const.bool.COND_3'(s(bool,sK1(X8)),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X8),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X8))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) )),
  inference(superposition,[],[f45,f45])).
fof(f431,plain,(
  ( ! [X1] : (s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X1))) = s('type.integer.int','const.bool.COND_3'(s(bool,sK1(X1)),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'F_0'),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f45,f38])).
fof(f432,plain,(
  ( ! [X2,X3] : (s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X2))) = s('type.integer.int','const.bool.COND_3'(s(bool,sK1(X2)),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | s(bool,X3) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f45,f48])).
fof(f433,plain,(
  ( ! [X4] : (s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X4))) = s('type.integer.int','const.bool.COND_3'(s(bool,sK1(X4)),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X4),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f45,f50])).
fof(f434,plain,(
  ( ! [X6,X5] : (s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X5))) = s('type.integer.int','const.bool.COND_3'(s(bool,sK1(X5)),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,X6),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | ~p(s(bool,X6)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X5),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f45,f57])).
fof(f436,plain,(
  ( ! [X10,X9] : (s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X9))) = s('type.integer.int','const.bool.COND_3'(s(bool,sK1(X9)),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | s(bool,X10) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X9),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,X10))) )),
  inference(superposition,[],[f45,f219])).
fof(f438,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X1))) = s('type.integer.int','const.bool.COND_3'(s(bool,sK1(X1)),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | ~p(s(bool,sK1('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | ~p(s(bool,sK1(X0)))) )),
  inference(superposition,[],[f45,f54])).
fof(f441,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X1))) = s('type.integer.int','const.bool.COND_3'(s(bool,sK1(X1)),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int',X0))))) | ~p(s(bool,sK1('const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | ~p(s(bool,sK1(X0)))) )),
  inference(superposition,[],[f45,f54])).
fof(f444,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X0))) = s('type.integer.int','const.bool.COND_3'(s(bool,sK1(X0)),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1))) | ~p(s(bool,sK1('const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | ~p(s(bool,sK1(X1)))) )),
  inference(superposition,[],[f45,f54])).
fof(f446,plain,(
  ( ! [X4] : (~p(s(bool,sK1('const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X4),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X4))) = s('type.integer.int','const.bool.COND_3'(s(bool,sK1(X4)),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f45,f43])).
fof(f447,plain,(
  ( ! [X2,X1] : (s('type.integer.int',X2) = s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X1))) | ~p(s(bool,sK1('const.bool.COND_3'(s(bool,sK1(X1)),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))))) | ~p(s(bool,sK1(X2)))) )),
  inference(superposition,[],[f45,f54])).
fof(f449,plain,(
  ( ! [X5] : (~p(s(bool,sK1('const.bool.COND_3'(s(bool,sK1(X5)),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X5),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X5)))) )),
  inference(superposition,[],[f45,f43])).
fof(f451,plain,(
  ( ! [X1] : (p(s(bool,sK1('const.bool.COND_3'(s(bool,sK1(X1)),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X1)))) )),
  inference(superposition,[],[f44,f45])).
fof(f454,plain,(
  ( ! [X6,X7] : (s('type.integer.int',X7) != s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X6))) | p(s(bool,sK1(X7))) | ~p(s(bool,sK1('const.bool.COND_3'(s(bool,sK1(X6)),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X6),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))))) )),
  inference(superposition,[],[f73,f45])).
fof(f455,plain,(
  ( ! [X8,X9] : (s('type.integer.int',X9) != s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X8))) | p(s(bool,sK1('const.bool.COND_3'(s(bool,sK1(X8)),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X8),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))))) | ~p(s(bool,sK1(X9)))) )),
  inference(superposition,[],[f73,f45])).
fof(f511,plain,(
  ( ! [X41,X40] : (s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X40))) = s('type.integer.int','const.bool.COND_3'(s(bool,sK1(X40)),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,X41),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X40),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X41))) )),
  inference(superposition,[],[f45,f279])).
fof(f513,plain,(
  ( ! [X43,X44] : (s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X43))) = s('type.integer.int','const.bool.COND_3'(s(bool,X44),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X43),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | p(s(bool,sK1(X43))) | p(s(bool,X44))) )),
  inference(superposition,[],[f45,f279])).
fof(f1495,plain,(
  ( ! [X90,X88,X89] : (s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X88))) = s('type.integer.int','const.bool.COND_3'(s(bool,sK1(X88)),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,X89),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | s(bool,X90) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X88),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,X89)) | ~p(s(bool,X90))) )),
  inference(superposition,[],[f45,f267])).
fof(f1497,plain,(
  ( ! [X94,X95,X93] : (s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X93))) = s('type.integer.int','const.bool.COND_3'(s(bool,X94),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X93),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | s(bool,X95) = s(bool,sK1(X93)) | p(s(bool,X94)) | ~p(s(bool,X95))) )),
  inference(superposition,[],[f45,f267])).
fof(f1527,plain,(
  ( ! [X90,X88,X89] : (s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X88))) = s('type.integer.int','const.bool.COND_3'(s(bool,sK1(X88)),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,X89),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | s(bool,X89) = s(bool,X90) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X88),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,X90))) )),
  inference(superposition,[],[f45,f267])).
fof(f1529,plain,(
  ( ! [X94,X95,X93] : (s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X93))) = s('type.integer.int','const.bool.COND_3'(s(bool,X94),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X93),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | s(bool,X94) = s(bool,X95) | p(s(bool,sK1(X93))) | ~p(s(bool,X95))) )),
  inference(superposition,[],[f45,f267])).
fof(f642,plain,(
  ( ! [X2] : (s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X2))))) = s('type.integer.int','const.bool.COND_3'(s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X2))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | ~p(s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac',X2)))))) )),
  inference(superposition,[],[f45,f84])).
fof(f423,plain,(
  ( ! [X6] : (s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X6))))) = s('type.integer.int','const.bool.COND_3'(s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X6)))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X6))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) )),
  inference(superposition,[],[f45,f36])).
fof(f37,plain,(
  ( ! [X0] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X0))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(cnf_transformation,[],[f16])).
fof(f16,plain,(
  ! [X0] : s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X0))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6] : s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X6))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X6))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.rat.RAT_SGN_CONG','thm.rat.RAT_NMRLT|48|_CONG')).
fof(f343,plain,(
  ( ! [X0] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0)))) )),
  inference(superposition,[],[f37,f40])).
fof(f346,plain,(
  ( ! [X5] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X5))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X5)))))))))) )),
  inference(superposition,[],[f37,f43])).
fof(f351,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X1))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f37,f38])).
fof(f352,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X2))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,X3) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X2))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f37,f48])).
fof(f457,plain,(
  ( ! [X0] : (s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X0))))))) = s('type.integer.int','const.bool.COND_3'(s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X0)))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) )),
  inference(forward_demodulation,[],[f430,f36])).
fof(f430,plain,(
  ( ! [X0] : (s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X0))))))))) = s('type.integer.int','const.bool.COND_3'(s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X0)))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) )),
  inference(superposition,[],[f45,f37])).
fof(f1880,plain,(
  ( ! [X5] : (s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) = s('type.integer.int','const.bool.COND_3'(s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X5)))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X5)))) )),
  inference(superposition,[],[f45,f197])).
fof(f635,plain,(
  ( ! [X2] : (s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) = s('type.integer.int','const.bool.COND_3'(s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac',X2)))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | ~p(s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac',X2)))))) )),
  inference(superposition,[],[f45,f84])).
fof(f411,plain,(
  s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))),
  inference(superposition,[],[f45,f89])).
fof(f421,plain,(
  ( ! [X4] : (s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.bool.COND_3'(s(bool,sK1(X4)),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | ~p(s(bool,sK1(X4)))) )),
  inference(superposition,[],[f45,f43])).
fof(f422,plain,(
  ( ! [X5] : (s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.bool.COND_3'(s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X5)))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X5)))) )),
  inference(superposition,[],[f45,f40])).
fof(f671,plain,(
  ( ! [X3] : (s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.bool.COND_3'(s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac',X3)))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | ~p(s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X3)))))))) )),
  inference(superposition,[],[f45,f85])).
fof(f54,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X0) = s('type.integer.int',X1) | ~p(s(bool,sK1(X1))) | ~p(s(bool,sK1(X0)))) )),
  inference(superposition,[],[f43,f43])).
fof(f73,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X0) != s('type.integer.int',X1) | p(s(bool,sK1(X1))) | ~p(s(bool,sK1(X0)))) )),
  inference(superposition,[],[f44,f43])).
fof(f146,plain,(
  ( ! [X10,X11] : (s('type.integer.int',X11) = s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X10))) | ~p(s(bool,sK1(X11))) | ~p(s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X10)))))))) )),
  inference(superposition,[],[f54,f36])).
fof(f155,plain,(
  ( ! [X6,X7] : (s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X6))) = s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X7))) | ~p(s(bool,sK1(X7))) | ~p(s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac',X6)))))) )),
  inference(superposition,[],[f36,f54])).
fof(f182,plain,(
  ( ! [X10,X11] : (s('type.integer.int',X11) != s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X10))) | p(s(bool,sK1(X11))) | ~p(s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X10)))))))) )),
  inference(superposition,[],[f73,f36])).
fof(f187,plain,(
  ( ! [X10,X11] : (s('type.integer.int',X11) != s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X10))) | p(s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X10))))))) | ~p(s(bool,sK1(X11)))) )),
  inference(superposition,[],[f73,f36])).
fof(f251,plain,(
  ( ! [X21,X19,X20,X18] : (s(X20,f(s('type.integer.int',X21))) != s(X20,g(s('type.integer.int',X21))) | s(fun('type.integer.int',X20),X18) = s(fun('type.integer.int',X20),X19) | ~p(s(bool,sK1(sK2(X18,X19,X20,'type.integer.int')))) | ~p(s(bool,sK1(X21)))) )),
  inference(superposition,[],[f46,f54])).
fof(f259,plain,(
  ( ! [X21,X19,X20,X18] : (s('type.integer.int',X21) != s('type.integer.int',f(s(X18,sK2(X19,X20,'type.integer.int',X18)))) | s(fun(X18,'type.integer.int'),X19) = s(fun(X18,'type.integer.int'),X20) | ~p(s(bool,sK1(g(s(X18,sK2(X19,X20,'type.integer.int',X18)))))) | ~p(s(bool,sK1(X21)))) )),
  inference(superposition,[],[f46,f54])).
fof(f621,plain,(
  ( ! [X2,X1] : (s('type.integer.int',X2) = s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac',X1))))) | ~p(s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac',X1))))) | ~p(s(bool,sK1(X2)))) )),
  inference(superposition,[],[f84,f54])).
fof(f629,plain,(
  ( ! [X2,X3] : (s('type.integer.int',X2) = s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X3))) | ~p(s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac',X3))))) | ~p(s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | ~p(s(bool,sK1(X2)))) )),
  inference(superposition,[],[f84,f54])).
fof(f638,plain,(
  ( ! [X8,X7] : (s('type.integer.int',X8) != s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,sK1(X8))) | ~p(s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac',X7))))) | ~p(s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac',X7)))))) )),
  inference(superposition,[],[f73,f84])).
fof(f639,plain,(
  ( ! [X10,X9] : (s('type.integer.int',X10) != s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac',X9))))) | ~p(s(bool,sK1(X10))) | ~p(s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac',X9)))))) )),
  inference(superposition,[],[f73,f84])).
fof(f645,plain,(
  ( ! [X8,X7] : (s('type.integer.int',X8) != s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X7))) | p(s(bool,sK1(X8))) | ~p(s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | ~p(s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac',X7)))))) )),
  inference(superposition,[],[f73,f84])).
fof(f646,plain,(
  ( ! [X10,X9] : (s('type.integer.int',X10) != s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X9))) | p(s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | ~p(s(bool,sK1(X10))) | ~p(s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac',X9)))))) )),
  inference(superposition,[],[f73,f84])).
fof(f1407,plain,(
  ( ! [X2,X1] : (~p(s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int',X2))))) | p(s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac',X1))))) | ~p(s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac',X1))))) | ~p(s(bool,sK1(X2)))) )),
  inference(superposition,[],[f971,f54])).
fof(f1608,plain,(
  ( ! [X2,X1] : (p(s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int',X2))))) | ~p(s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac',X1))))) | ~p(s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac',X1))))) | ~p(s(bool,sK1(X2)))) )),
  inference(superposition,[],[f1605,f54])).
fof(f1871,plain,(
  ( ! [X4,X3] : (s('type.integer.int',X4) = s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X3))) | ~p(s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X3))))))))) | ~p(s(bool,sK1(X4)))) )),
  inference(superposition,[],[f197,f54])).
fof(f1883,plain,(
  ( ! [X10,X11] : (s('type.integer.int',X11) != s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,sK1(X11))) | ~p(s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X10))))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X10)))) )),
  inference(superposition,[],[f73,f197])).
fof(f1884,plain,(
  ( ! [X12,X13] : (s('type.integer.int',X13) != s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X12))))))))) | ~p(s(bool,sK1(X13))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X12)))) )),
  inference(superposition,[],[f73,f197])).
fof(f1975,plain,(
  ( ! [X37,X35,X36] : (s('type.integer.int',X35) != s('type.integer.int',f(s(bool,'F_0'))) | s(fun(bool,'type.integer.int'),X36) = s(fun(bool,'type.integer.int'),X37) | s(bool,'T_0') = s(bool,sK2(X36,X37,'type.integer.int',bool)) | ~p(s(bool,sK1(g(s(bool,'F_0'))))) | ~p(s(bool,sK1(X35)))) )),
  inference(superposition,[],[f250,f54])).
fof(f2102,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') = s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int',X2)))) | ~p(s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac',X1))))) | ~p(s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac',X1))))) | ~p(s(bool,sK1(X2)))) )),
  inference(superposition,[],[f1626,f54])).
fof(f2167,plain,(
  ( ! [X37,X35,X36] : (s('type.integer.int',X35) != s('type.integer.int',f(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(fun('type.integer.int','type.integer.int'),X36) = s(fun('type.integer.int','type.integer.int'),X37) | ~p(s(bool,sK1(sK2(X36,X37,'type.integer.int','type.integer.int')))) | ~p(s(bool,sK1(g(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | ~p(s(bool,sK1(X35)))) )),
  inference(superposition,[],[f253,f54])).
fof(f2169,plain,(
  ( ! [X41,X42] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int',f(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(fun('type.integer.int','type.integer.int'),X41) = s(fun('type.integer.int','type.integer.int'),X42) | ~p(s(bool,sK1(sK2(X41,X42,'type.integer.int','type.integer.int')))) | ~p(s(bool,sK1(g(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f253,f43])).
fof(f261,plain,(
  ( ! [X28,X26,X27] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int',f(s(X26,sK2(X27,X28,'type.integer.int',X26)))) | s(fun(X26,'type.integer.int'),X27) = s(fun(X26,'type.integer.int'),X28) | ~p(s(bool,sK1(g(s(X26,sK2(X27,X28,'type.integer.int',X26))))))) )),
  inference(superposition,[],[f46,f43])).
fof(f1406,plain,(
  ( ! [X0] : (~p(s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | p(s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X0))))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0)))) )),
  inference(superposition,[],[f971,f40])).
fof(f1607,plain,(
  ( ! [X0] : (p(s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | ~p(s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X0))))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0)))) )),
  inference(superposition,[],[f1605,f40])).
fof(f1873,plain,(
  ( ! [X7] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X7))) | ~p(s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X7)))))))))) )),
  inference(superposition,[],[f197,f43])).
fof(f1879,plain,(
  ( ! [X4] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X4))))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X4)))) )),
  inference(superposition,[],[f44,f197])).
fof(f2101,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X0))))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0)))) )),
  inference(superposition,[],[f1626,f40])).
fof(f623,plain,(
  ( ! [X5] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac',X5))))) | ~p(s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac',X5)))))) )),
  inference(superposition,[],[f84,f43])).
fof(f634,plain,(
  ( ! [X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac',X1))))) | ~p(s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac',X1)))))) )),
  inference(superposition,[],[f44,f84])).
fof(f971,plain,(
  ( ! [X10] : (~p(s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X10))))))) | p(s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac',X10)))))) )),
  inference(global_subsumption,[],[f32,f33,f35,f36,f37,f38,f40,f39,f45,f44,f43,f46,f48,f50,f53,f54,f57,f76,f73,f75,f89,f96,f146,f154,f155,f182,f187,f208,f209,f219,f246,f250,f251,f254,f255,f257,f258,f259,f261,f265,f266,f267,f300,f279,f292,f296,f343,f344,f346,f347,f351,f352,f412,f413,f414,f415,f417,f423,f456,f425,f426,f431,f432,f433,f434,f436,f438,f440,f441,f443,f444,f446,f447,f449,f451,f454,f455,f511,f513,f515,f629,f642,f645,f646,f671,f676,f715,f660,f657,f457,f639,f638,f635,f634,f623,f621,f422,f421,f411,f197,f514,f291,f249,f247,f253,f668,f649,f86,f85,f632,f204,f203,f141,f669,f661,f641,f631,f628,f84,f712])).
fof(f1409,plain,(
  ( ! [X5] : (~p(s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | p(s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac',X5))))) | ~p(s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac',X5)))))) )),
  inference(superposition,[],[f971,f43])).
fof(f1422,plain,(
  ( ! [X2,X1] : (p(s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac',X1))))) | s(bool,X2) = s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X1)))))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f1411,f33])).
fof(f1411,plain,(
  ( ! [X2,X1] : (~p(s(bool,'T_0')) | p(s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac',X1))))) | s(bool,X2) = s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X1)))))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f971,f48])).
fof(f1423,plain,(
  ( ! [X8,X9] : (p(s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac',X8))))) | s(bool,X9) = s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X8)))))) | p(s(bool,X9))) )),
  inference(subsumption_resolution,[],[f1415,f33])).
fof(f1415,plain,(
  ( ! [X8,X9] : (~p(s(bool,'T_0')) | p(s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac',X8))))) | s(bool,X9) = s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X8)))))) | p(s(bool,X9))) )),
  inference(superposition,[],[f971,f219])).
fof(f1610,plain,(
  ( ! [X5] : (p(s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | ~p(s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac',X5))))) | ~p(s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac',X5)))))) )),
  inference(superposition,[],[f1605,f43])).
fof(f1625,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X11)) | ~p(s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac',X10))))) | s(bool,X12) = s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X10)))))) | ~p(s(bool,X12))) )),
  inference(duplicate_literal_removal,[],[f1617])).
fof(f1617,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X11)) | ~p(s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac',X10))))) | s(bool,X12) = s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X10)))))) | p(s(bool,X11)) | ~p(s(bool,X12))) )),
  inference(superposition,[],[f1605,f267])).
fof(f2104,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac',X5))))) | ~p(s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac',X5)))))) )),
  inference(superposition,[],[f1626,f43])).
fof(f2137,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') = s(bool,X10) | ~p(s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac',X9))))) | s(bool,X11) = s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X9)))))) | ~p(s(bool,X11))) )),
  inference(subsumption_resolution,[],[f2109,f50])).
fof(f2109,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') = s(bool,X10) | ~p(s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac',X9))))) | s(bool,X11) = s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X9)))))) | p(s(bool,X10)) | ~p(s(bool,X11))) )),
  inference(superposition,[],[f1626,f267])).
fof(f2549,plain,(
  ( ! [X2] : (p(s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X2))))))) | ~p(s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac',X2)))))) )),
  inference(global_subsumption,[],[f32,f33,f36,f38,f40,f39,f46,f48,f50,f53,f57,f76,f75,f89,f208,f209,f219,f246,f265,f266,f267,f279,f296,f676,f197,f668,f712,f971,f634,f623,f96,f261,f646,f645,f639,f638,f629,f621,f259,f251,f187,f182,f155,f154,f146,f73,f54,f44,f43,f671,f422,f421,f411,f635,f457,f352,f346,f343,f37,f423,f642,f513,f511,f455,f454,f451,f449,f447,f446,f444,f443,f441,f440,f438,f436,f434,f433,f432,f425,f417,f415,f414,f413,f344,f45,f657,f660,f715,f426,f347,f456,f515,f292,f257,f255,f254,f431,f412,f351,f300,f258,f250,f35,f1407,f1409,f1422,f1423,f1563,f1546,f1495,f1497,f1500,f1527,f1529,f1532,f1605,f1608,f1729,f1803,f1871,f1875,f1880,f1883,f1884,f1950,f1963,f1964,f1965,f1966,f1968,f1969,f1970,f1971,f1975,f2051,f2052,f2054,f2059,f2060,f2063,f2066,f2067,f2068,f2069,f2070,f2072,f2073,f2074,f2075,f2102,f2155,f2156,f2157,f2158,f2160,f2161,f2162,f2163,f2167,f2169,f1956,f1534,f2174,f1977,f1531,f1499,f514,f291,f249,f247,f253,f2150,f1595,f649,f86,f85,f1870,f1876,f1601,f669,f661,f641,f631,f628,f84,f1802,f1730,f204,f203,f141,f632,f2101,f1879,f1873,f1607,f1406,f2137,f2104,f1626,f1625,f1610,f1606])).
fof(f141,plain,(
  ( ! [X0] : (~p(s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X0))))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0)))) )),
  inference(trivial_inequality_removal,[],[f140])).
fof(f140,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0))) | ~p(s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X0)))))))))) )),
  inference(superposition,[],[f39,f43])).
fof(f203,plain,(
  ( ! [X10] : (p(s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X10))))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X10)))) )),
  inference(trivial_inequality_removal,[],[f202])).
fof(f202,plain,(
  ( ! [X10] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X10))))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X10)))) )),
  inference(superposition,[],[f44,f40])).
fof(f204,plain,(
  ( ! [X2,X3] : (s('type.integer.int',X3) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X2))))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X2)))) )),
  inference(subsumption_resolution,[],[f198,f44])).
fof(f198,plain,(
  ( ! [X2,X3] : (s('type.integer.int',X3) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X2))))))))) | ~p(s(bool,sK1(X3))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X2)))) )),
  inference(superposition,[],[f73,f40])).
fof(f1730,plain,(
  ( ! [X8,X9] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X8))) | s(bool,X9) = s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X8)))))))) | p(s(bool,X9))) )),
  inference(subsumption_resolution,[],[f1720,f33])).
fof(f1720,plain,(
  ( ! [X8,X9] : (~p(s(bool,'T_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X8))) | s(bool,X9) = s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X8)))))))) | p(s(bool,X9))) )),
  inference(superposition,[],[f141,f219])).
fof(f1802,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X11)) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X10))) | s(bool,X12) = s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X10)))))))) | ~p(s(bool,X12))) )),
  inference(duplicate_literal_removal,[],[f1794])).
fof(f1794,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X11)) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X10))) | s(bool,X12) = s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X10)))))))) | p(s(bool,X11)) | ~p(s(bool,X12))) )),
  inference(superposition,[],[f203,f267])).
fof(f632,plain,(
  s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',sK0))) != s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0)))))))))),
  inference(superposition,[],[f32,f84])).
fof(f84,plain,(
  ( ! [X0] : (s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X0))) = s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac',X0)))))) )),
  inference(superposition,[],[f36,f43])).
fof(f628,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X0))) = s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X1))) | ~p(s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac',X1))))) | ~p(s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac',X0)))))) )),
  inference(superposition,[],[f84,f84])).
fof(f631,plain,(
  ( ! [X6] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X6))) | ~p(s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac',X6))))) | ~p(s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f84,f43])).
fof(f641,plain,(
  ( ! [X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X1))) | p(s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | ~p(s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac',X1)))))) )),
  inference(superposition,[],[f44,f84])).
fof(f661,plain,(
  ( ! [X10,X11] : (s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X11))) = s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X10))))) | ~p(s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac',X11))))) | ~p(s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X10)))))))) )),
  inference(superposition,[],[f84,f85])).
fof(f669,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac',X0))))) | ~p(s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0)))))))) )),
  inference(superposition,[],[f84,f85])).
fof(f1601,plain,(
  ( ! [X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X1))))))) | ~p(s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac',X1)))))) )),
  inference(superposition,[],[f86,f84])).
fof(f1876,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X0))))))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0)))) )),
  inference(superposition,[],[f86,f197])).
fof(f1870,plain,(
  ( ! [X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X1))) | ~p(s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X1)))))))))))) )),
  inference(superposition,[],[f197,f85])).
fof(f649,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X0))) = s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X1))) | ~p(s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X1))))))) | ~p(s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0)))))))) )),
  inference(superposition,[],[f85,f85])).
fof(f1595,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X0))) != s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X1))) | p(s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X1))))))) | ~p(s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0)))))))) )),
  inference(superposition,[],[f86,f85])).
fof(f2150,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X0))))) != s(X1,g(s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X0))))) | s(fun('type.integer.int',X1),X2) = s(fun('type.integer.int',X1),X3) | ~p(s(bool,sK1(sK2(X2,X3,X1,'type.integer.int')))) | ~p(s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0)))))))) )),
  inference(superposition,[],[f253,f85])).
fof(f253,plain,(
  ( ! [X28,X26,X27] : (s(X28,f(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s(X28,g(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(fun('type.integer.int',X28),X26) = s(fun('type.integer.int',X28),X27) | ~p(s(bool,sK1(sK2(X26,X27,X28,'type.integer.int'))))) )),
  inference(superposition,[],[f46,f43])).
fof(f247,plain,(
  ( ! [X6,X4,X7,X5] : (s(X6,f(s(bool,X7))) != s(X6,g(s(bool,X7))) | s(fun(bool,X6),X4) = s(fun(bool,X6),X5) | ~p(s(bool,sK2(X4,X5,X6,bool))) | ~p(s(bool,X7))) )),
  inference(superposition,[],[f46,f57])).
fof(f249,plain,(
  ( ! [X14,X12,X13] : (s(X14,f(s(bool,'T_0'))) != s(X14,g(s(bool,'T_0'))) | s(fun(bool,X14),X12) = s(fun(bool,X14),X13) | ~p(s(bool,sK2(X12,X13,X14,bool)))) )),
  inference(superposition,[],[f46,f50])).
fof(f291,plain,(
  ( ! [X26,X24,X23,X25] : (s(X25,f(s(bool,'T_0'))) != s(X25,g(s(bool,'T_0'))) | s(fun(bool,X25),X23) = s(fun(bool,X25),X24) | s(bool,X26) = s(bool,sK2(X23,X24,X25,bool)) | p(s(bool,X26))) )),
  inference(superposition,[],[f46,f219])).
fof(f514,plain,(
  ( ! [X52,X50,X51,X49] : (s(X51,f(s(bool,X52))) != s(X51,g(s(bool,X52))) | s(fun(bool,X51),X49) = s(fun(bool,X51),X50) | p(s(bool,sK2(X49,X50,X51,bool))) | p(s(bool,X52))) )),
  inference(superposition,[],[f46,f279])).
fof(f1499,plain,(
  ( ! [X103,X107,X105,X106,X104] : (s(X105,f(s(bool,X106))) != s(X105,g(s(bool,X106))) | s(fun(bool,X105),X103) = s(fun(bool,X105),X104) | s(bool,X107) = s(bool,sK2(X103,X104,X105,bool)) | p(s(bool,X106)) | ~p(s(bool,X107))) )),
  inference(superposition,[],[f46,f267])).
fof(f1531,plain,(
  ( ! [X103,X107,X105,X106,X104] : (s(X105,f(s(bool,X106))) != s(X105,g(s(bool,X106))) | s(fun(bool,X105),X103) = s(fun(bool,X105),X104) | s(bool,X106) = s(bool,X107) | p(s(bool,sK2(X103,X104,X105,bool))) | ~p(s(bool,X107))) )),
  inference(superposition,[],[f46,f267])).
fof(f1977,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s(X36,f(s(bool,X35))) != s(X36,g(s(bool,X35))) | s(fun(bool,X36),X37) = s(fun(bool,X36),X38) | s(bool,'T_0') = s(bool,sK2(X37,X38,X36,bool)) | s(bool,X35) = s(bool,X39) | ~p(s(bool,X39))) )),
  inference(subsumption_resolution,[],[f1958,f35])).
fof(f1958,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s(X36,f(s(bool,X35))) != s(X36,g(s(bool,X35))) | s(fun(bool,X36),X37) = s(fun(bool,X36),X38) | s(bool,'T_0') = s(bool,sK2(X37,X38,X36,bool)) | s(bool,X35) = s(bool,X39) | p(s(bool,'F_0')) | ~p(s(bool,X39))) )),
  inference(superposition,[],[f250,f267])).
fof(f2174,plain,(
  ( ! [X47,X45,X46,X44] : (s(X45,f(s(bool,X44))) != s(X45,g(s(bool,X44))) | s(fun(bool,X45),X46) = s(fun(bool,X45),X47) | s(bool,'T_0') = s(bool,sK2(X46,X47,X45,bool)) | p(s(bool,X44))) )),
  inference(subsumption_resolution,[],[f1960,f35])).
fof(f1960,plain,(
  ( ! [X47,X45,X46,X44] : (s(X45,f(s(bool,X44))) != s(X45,g(s(bool,X44))) | s(fun(bool,X45),X46) = s(fun(bool,X45),X47) | s(bool,'T_0') = s(bool,sK2(X46,X47,X45,bool)) | p(s(bool,'F_0')) | p(s(bool,X44))) )),
  inference(superposition,[],[f250,f279])).
fof(f2684,plain,(
  ( ! [X4,X5,X3] : (s(bool,X3) != s(bool,X4) | p(s(bool,X4)) | s(bool,X3) = s(bool,X5) | s(bool,'T_0') = s(bool,X5)) ) | $spl10),
  inference(superposition,[],[f753,f48])).
fof(f2683,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X2) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | p(s(bool,X2))) ) | $spl10),
  inference(superposition,[],[f753,f219])).
fof(f1606,plain,(
  ( ! [X2,X3] : (s('type.integer.int',X3) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X2))))))) | ~p(s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac',X2)))))) )),
  inference(subsumption_resolution,[],[f1602,f44])).
fof(f1602,plain,(
  ( ! [X2,X3] : (s('type.integer.int',X3) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X2))))))) | ~p(s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac',X2))))) | ~p(s(bool,sK1(X3)))) )),
  inference(superposition,[],[f86,f54])).
fof(f1534,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) != s(bool,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(equality_factoring,[],[f267])).
fof(f1956,plain,(
  ( ! [X28,X26,X29,X27] : (s(X26,f(s(bool,'T_0'))) != s(X26,g(s(bool,'T_0'))) | s(fun(bool,X26),X27) = s(fun(bool,X26),X28) | s(bool,'T_0') = s(bool,sK2(X27,X28,X26,bool)) | s(bool,'F_0') = s(bool,X29) | p(s(bool,X29))) )),
  inference(superposition,[],[f250,f219])).
fof(f1950,plain,(
  ( ! [X6,X4,X7,X5] : (s(X5,f(s(bool,X4))) != s(X5,g(s(bool,X4))) | s(fun(bool,X5),X6) = s(fun(bool,X5),X7) | s(bool,'T_0') = s(bool,sK2(X6,X7,X5,bool)) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f250,f38])).
fof(f1803,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0))) | s(bool,'T_0') = s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X0))))))))) )),
  inference(subsumption_resolution,[],[f1788,f35])).
fof(f1788,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0))) | s(bool,'T_0') = s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X0))))))))) )),
  inference(superposition,[],[f203,f38])).
fof(f1729,plain,(
  ( ! [X2,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X1))) | s(bool,X2) = s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X1)))))))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f1716,f33])).
fof(f1716,plain,(
  ( ! [X2,X1] : (~p(s(bool,'T_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X1))) | s(bool,X2) = s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X1)))))))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f141,f48])).
fof(f1546,plain,(
  ( ! [X30,X28,X31,X29,X27] : (s(bool,X28) = s(bool,X29) | s(bool,X29) = s(bool,X30) | p(s(bool,X27)) | ~p(s(bool,X30)) | s(bool,X28) = s(bool,X31) | ~p(s(bool,X31))) )),
  inference(duplicate_literal_removal,[],[f1456])).
fof(f1456,plain,(
  ( ! [X30,X28,X31,X29,X27] : (s(bool,X28) = s(bool,X29) | s(bool,X29) = s(bool,X30) | p(s(bool,X27)) | ~p(s(bool,X30)) | s(bool,X28) = s(bool,X31) | p(s(bool,X27)) | ~p(s(bool,X31))) )),
  inference(superposition,[],[f267,f267])).
fof(f1563,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,X2) | ~p(s(bool,X2)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f1449,f50])).
fof(f1449,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,X2) | p(s(bool,X0)) | ~p(s(bool,X2)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f267,f38])).
fof(f250,plain,(
  ( ! [X17,X15,X16] : (s(X17,f(s(bool,'F_0'))) != s(X17,g(s(bool,'F_0'))) | s(fun(bool,X17),X15) = s(fun(bool,X17),X16) | s(bool,'T_0') = s(bool,sK2(X15,X16,X17,bool))) )),
  inference(superposition,[],[f46,f38])).
fof(f154,plain,(
  ( ! [X5] : (s('type.integer.int',X5) != s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',sK0))) | ~p(s(bool,sK1(X5))) | ~p(s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0)))))))))) )),
  inference(superposition,[],[f32,f54])).
fof(f712,plain,(
  ( ! [X10,X11] : (s('type.integer.int',X11) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac',X10))))) | ~p(s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X10)))))))) )),
  inference(subsumption_resolution,[],[f675,f44])).
fof(f675,plain,(
  ( ! [X10,X11] : (s('type.integer.int',X11) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac',X10))))) | ~p(s(bool,sK1(X11))) | ~p(s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X10)))))))) )),
  inference(superposition,[],[f73,f85])).
fof(f668,plain,(
  s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',sK0))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0)))))))))))),
  inference(superposition,[],[f32,f85])).
fof(f676,plain,(
  ( ! [X2] : (p(s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac',X2))))) | ~p(s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X2)))))))) )),
  inference(trivial_inequality_removal,[],[f670])).
fof(f670,plain,(
  ( ! [X2] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac',X2))))) | ~p(s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X2)))))))) )),
  inference(superposition,[],[f44,f85])).
fof(f279,plain,(
  ( ! [X8,X9] : (s(bool,X8) = s(bool,X9) | p(s(bool,X8)) | p(s(bool,X9))) )),
  inference(superposition,[],[f33,f219])).
fof(f267,plain,(
  ( ! [X10,X8,X9] : (s(bool,X9) = s(bool,X10) | s(bool,X8) = s(bool,X9) | p(s(bool,X10)) | ~p(s(bool,X8))) )),
  inference(superposition,[],[f219,f50])).
fof(f266,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X4) = s(bool,X5) | s(bool,X5) = s(bool,X6) | p(s(bool,X6)) | s(bool,X4) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f219,f48])).
fof(f265,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | p(s(bool,X2)) | s(bool,X0) = s(bool,X3) | p(s(bool,X3))) )),
  inference(superposition,[],[f219,f219])).
fof(f246,plain,(
  ( ! [X2,X0,X3,X1] : (s(X2,f(s(bool,'T_0'))) != s(X2,g(s(bool,'T_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,X3) = s(bool,sK2(X0,X1,X2,bool)) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f46,f48])).
fof(f219,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') = s(bool,X5) | s(bool,X4) = s(bool,X5) | p(s(bool,X4))) )),
  inference(superposition,[],[f33,f48])).
fof(f209,plain,(
  ( ! [X6,X4,X5] : (s(bool,X4) = s(bool,X5) | s(bool,X5) = s(bool,X6) | s(bool,'T_0') = s(bool,X6) | ~p(s(bool,X4))) )),
  inference(superposition,[],[f48,f50])).
fof(f208,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f48,f48])).
fof(f89,plain,(
  s(bool,'T_0') = s(bool,sK1('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(unit_resulting_resolution,[],[f75,f50])).
fof(f75,plain,(
  p(s(bool,sK1('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(equality_resolution,[],[f44])).
fof(f76,plain,(
  ( ! [X1] : (p(s(bool,sK1('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,sK1(X1)))) )),
  inference(trivial_inequality_removal,[],[f72])).
fof(f72,plain,(
  ( ! [X1] : (s('type.integer.int',X1) != s('type.integer.int',X1) | p(s(bool,sK1('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,sK1(X1)))) )),
  inference(superposition,[],[f44,f43])).
fof(f57,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(superposition,[],[f50,f50])).
fof(f48,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X1) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f38,f38])).
fof(f46,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK2(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK2(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f31])).
fof(f31,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK2(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK2(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2])],[f24])).
fof(f24,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f22])).
fof(f22,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f21])).
fof(f21,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.rat.RAT_SGN_CONG',aHL_EXT)).
fof(f39,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X0))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0)))) )),
  inference(cnf_transformation,[],[f26])).
fof(f45,plain,(
  ( ! [X0] : (s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X0))) = s('type.integer.int','const.bool.COND_3'(s(bool,sK1(X0)),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) )),
  inference(cnf_transformation,[],[f30])).
fof(f9701,plain,(
  s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))))))) = s('type.integer.int','const.bool.COND_3'(s(bool,'F_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | ($spl10 | $spl27)),
  inference(superposition,[],[f457,f7987])).
fof(f7987,plain,(
  s(bool,'F_0') = s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0)))))))) | ($spl10 | $spl27)),
  inference(unit_resulting_resolution,[],[f7956,f5500])).
fof(f7956,plain,(
  ~p(s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))))))))) | $spl27),
  inference(unit_resulting_resolution,[],[f7523,f141])).
fof(f7523,plain,(
  s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl27),
  inference(unit_resulting_resolution,[],[f7518,f44])).
fof(f8172,plain,(
  ~p(s(bool,sK1('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | $spl35),
  inference(cnf_transformation,[],[f8172_D])).
fof(f8172_D,plain,(
  ~p(s(bool,sK1('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) <=> ~$spl35),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl35])])).
fof(f8173,plain,(
  ( ! [X4] : (s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X4))) = s('type.integer.int','const.bool.COND_3'(s(bool,sK1(X4)),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X4),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) ) | $spl36),
  inference(cnf_transformation,[],[f8173_D])).
fof(f8173_D,plain,(
  ( ! [X4] : (s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X4))) = s('type.integer.int','const.bool.COND_3'(s(bool,sK1(X4)),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X4),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f440,plain,(
  ( ! [X4] : (s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X4))) = s('type.integer.int','const.bool.COND_3'(s(bool,sK1(X4)),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X4),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | ~p(s(bool,sK1('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) )),
  inference(superposition,[],[f45,f43])).
fof(f7521,plain,(
  ~p(s(bool,sK1('const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | $spl31),
  inference(cnf_transformation,[],[f7521_D])).
fof(f7521_D,plain,(
  ~p(s(bool,sK1('const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) <=> ~$spl31),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl31])])).
fof(f7522,plain,(
  ( ! [X4] : (s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X4))) = s('type.integer.int','const.bool.COND_3'(s(bool,sK1(X4)),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X4),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl32),
  inference(cnf_transformation,[],[f7522_D])).
fof(f7522_D,plain,(
  ( ! [X4] : (s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X4))) = s('type.integer.int','const.bool.COND_3'(s(bool,sK1(X4)),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X4),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f443,plain,(
  ( ! [X4] : (s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X4))) = s('type.integer.int','const.bool.COND_3'(s(bool,sK1(X4)),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X4),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | ~p(s(bool,sK1('const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) )),
  inference(superposition,[],[f45,f43])).
fof(f7518,plain,(
  ~p(s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac',sK0))))) | $spl27),
  inference(cnf_transformation,[],[f7518_D])).
fof(f7518_D,plain,(
  ~p(s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac',sK0))))) <=> ~$spl27),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl27])])).
fof(f7519,plain,(
  ( ! [X0] : (~p(s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int',X0))))) | ~p(s(bool,sK1(X0)))) ) | $spl28),
  inference(cnf_transformation,[],[f7519_D])).
fof(f7519_D,plain,(
  ( ! [X0] : (~p(s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int',X0))))) | ~p(s(bool,sK1(X0)))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f5573,plain,(
  ( ! [X0] : (~p(s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int',X0))))) | ~p(s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac',sK0))))) | ~p(s(bool,sK1(X0)))) ) | $spl17),
  inference(superposition,[],[f5506,f54])).
fof(f647,plain,(
  ( ! [X22] : (p(s(bool,X22))) ) | $spl4),
  inference(cnf_transformation,[],[f647_D])).
fof(f647_D,plain,(
  ( ! [X22] : (p(s(bool,X22))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f5499,plain,(
  ( ! [X56,X55] : (s(bool,'F_0') = s(bool,X55) | s(bool,X55) != s(bool,X56) | p(s(bool,X56))) ) | $spl24),
  inference(cnf_transformation,[],[f5499_D])).
fof(f5499_D,plain,(
  ( ! [X56,X55] : (s(bool,'F_0') = s(bool,X55) | s(bool,X55) != s(bool,X56) | p(s(bool,X56))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f5335,plain,(
  ( ! [X54,X56,X55] : (s(bool,'F_0') = s(bool,X55) | p(s(bool,X56)) | s(bool,X55) != s(bool,X56) | p(s(bool,X54))) ) | ($spl10 | $spl15)),
  inference(duplicate_literal_removal,[],[f5192])).
fof(f5192,plain,(
  ( ! [X54,X56,X55] : (s(bool,'F_0') = s(bool,X55) | p(s(bool,X56)) | s(bool,X55) != s(bool,X56) | p(s(bool,X54)) | p(s(bool,X54))) ) | ($spl10 | $spl15)),
  inference(superposition,[],[f2683,f2175])).
fof(f2175,plain,(
  ( ! [X3] : (s(bool,'F_0') = s(bool,X3) | p(s(bool,X3))) ) | $spl15),
  inference(backward_demodulation,[],[f1676,f1711])).
fof(f1711,plain,(
  ( ! [X3] : (s(bool,X3) = s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0)))))))))) | p(s(bool,X3))) ) | $spl15),
  inference(subsumption_resolution,[],[f1692,f33])).
fof(f1692,plain,(
  ( ! [X3] : (~p(s(bool,'T_0')) | s(bool,X3) = s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0)))))))))) | p(s(bool,X3))) ) | $spl15),
  inference(superposition,[],[f1633,f219])).
fof(f1633,plain,(
  ~p(s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))))))))))) | $spl15),
  inference(unit_resulting_resolution,[],[f36,f1629,f73])).
fof(f1676,plain,(
  s(bool,'F_0') = s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0)))))))))) | $spl15),
  inference(unit_resulting_resolution,[],[f35,f1633,f279])).
fof(f3965,plain,(
  $false | ($spl3 | $spl20)),
  inference(subsumption_resolution,[],[f3674,f35])).
fof(f3674,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl20)),
  inference(unit_resulting_resolution,[],[f70,f3626,f279])).
fof(f3964,plain,(
  $false | ($spl3 | $spl15 | $spl20)),
  inference(subsumption_resolution,[],[f3676,f3626])).
fof(f3676,plain,(
  p(s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))))))))) | ($spl3 | $spl15 | $spl20)),
  inference(unit_resulting_resolution,[],[f3394,f3626,f279])).
fof(f3394,plain,(
  s(bool,'T_0') != s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0)))))))) | ($spl3 | $spl15)),
  inference(global_subsumption,[],[f32,f33,f36,f38,f40,f39,f46,f48,f50,f53,f57,f76,f75,f89,f208,f209,f219,f246,f265,f266,f267,f279,f296,f320,f377,f378,f379,f614,f606,f617,f676,f197,f668,f712,f96,f154,f437,f418,f338,f337,f300,f299,f250,f35,f1296,f1297,f1298,f1299,f1321,f1350,f1351,f1352,f1353,f1424,f1563,f1546,f1543,f1538,f1605,f1629,f1643,f1641,f1634,f1633,f1660,f1682,f1680,f1673,f1706,f1729,f1732,f1746,f1803,f1874,f1875,f1888,f1889,f1891,f1949,f1950,f1973,f1974,f2001,f2002,f2004,f2077,f2078,f2091,f2092,f2093,f2165,f2166,f2044,f1956,f1860,f1711,f1676,f2175,f1665,f1637,f1534,f2174,f1977,f1531,f1499,f514,f291,f249,f247,f253,f1683,f2150,f1595,f649,f86,f85,f1870,f1876,f1601,f669,f661,f641,f631,f628,f84,f1802,f1730,f204,f203,f141,f1856,f1686,f632,f2101,f1879,f1873,f1607,f1406,f2137,f2104,f1626,f1625,f1610,f1606,f2549,f1423,f1422,f1409,f971,f634,f623,f2169,f261,f2167,f2102,f2037,f1975,f1884,f1883,f1871,f1684,f1608,f1407,f646,f645,f639,f638,f629,f621,f259,f251,f187,f182,f155,f146,f73,f54,f44,f43,f671,f422,f421,f411,f635,f1880,f1703,f457,f352,f351,f346,f343,f37,f423,f642,f1529,f1527,f1497,f1495,f513,f511,f455,f454,f451,f449,f447,f446,f444,f443,f441,f440,f438,f436,f434,f433,f432,f431,f425,f417,f415,f414,f413,f412,f344,f45,f2066,f2059,f2060,f715,f660,f657,f2063,f2075,f2074,f2073,f2072,f2070,f2069,f2068,f2067,f347,f426,f2052,f2051,f2054,f456,f2163,f2162,f2161,f2160,f2158,f2157,f2156,f2155,f1971,f1970,f1969,f1968,f1966,f1965,f1964,f1963,f1532,f1500,f515,f292,f258,f257,f255,f254,f70,f2587,f2664,f2665,f2666,f2963,f2987,f3067,f3393])).
fof(f3393,plain,(
  ( ! [X59] : (s(bool,'F_0') = s(bool,X59) | s(bool,'T_0') != s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))))))))) ) | ($spl3 | $spl15)),
  inference(subsumption_resolution,[],[f3271,f2175])).
fof(f3271,plain,(
  ( ! [X59] : (s(bool,'F_0') = s(bool,X59) | s(bool,'T_0') != s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0)))))))) | ~p(s(bool,X59))) ) | ($spl3 | $spl15)),
  inference(superposition,[],[f2587,f1637])).
fof(f3067,plain,(
  s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))))))))) = s('type.integer.int','const.bool.COND_3'(s(bool,'F_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | $spl15),
  inference(superposition,[],[f45,f1637])).
fof(f2987,plain,(
  ( ! [X21] : (s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X21))) = s('type.integer.int','const.bool.COND_3'(s(bool,sK1(X21)),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'F_0'),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X21),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl15),
  inference(superposition,[],[f45,f2175])).
fof(f2963,plain,(
  ( ! [X20] : (s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X20))) = s('type.integer.int','const.bool.COND_3'(s(bool,'F_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X20),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | p(s(bool,sK1(X20)))) ) | $spl15),
  inference(superposition,[],[f45,f2175])).
fof(f2666,plain,(
  ( ! [X33,X34] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X33) = s(fun(bool,bool),X34) | s(bool,'T_0') = s(bool,sK2(X33,X34,bool,bool)) | p(s(bool,g(s(bool,'F_0'))))) ) | $spl15),
  inference(superposition,[],[f250,f2175])).
fof(f2665,plain,(
  ( ! [X30,X31,X32] : (s(bool,'F_0') != s(bool,f(s(X30,sK2(X31,X32,bool,X30)))) | s(fun(X30,bool),X31) = s(fun(X30,bool),X32) | p(s(bool,g(s(X30,sK2(X31,X32,bool,X30)))))) ) | $spl15),
  inference(superposition,[],[f46,f2175])).
fof(f2664,plain,(
  ( ! [X28,X29,X27] : (s(X29,f(s(bool,'F_0'))) != s(X29,g(s(bool,'F_0'))) | s(fun(bool,X29),X27) = s(fun(bool,X29),X28) | p(s(bool,sK2(X27,X28,X29,bool)))) ) | $spl15),
  inference(superposition,[],[f46,f2175])).
fof(f2587,plain,(
  ( ! [X8,X7] : (s(bool,X7) = s(bool,X8) | s(bool,'T_0') != s(bool,X7) | ~p(s(bool,X8))) ) | $spl3),
  inference(subsumption_resolution,[],[f2581,f35])).
fof(f2581,plain,(
  ( ! [X8,X7] : (s(bool,'T_0') != s(bool,X7) | s(bool,X7) = s(bool,X8) | p(s(bool,'F_0')) | ~p(s(bool,X8))) ) | $spl3),
  inference(superposition,[],[f70,f267])).
fof(f1703,plain,(
  ~p(s(bool,sK1('const.bool.COND_3'(s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0)))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))))) | $spl15),
  inference(forward_demodulation,[],[f1672,f37])).
fof(f1672,plain,(
  ~p(s(bool,sK1('const.bool.COND_3'(s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0)))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))))) | $spl15),
  inference(unit_resulting_resolution,[],[f45,f1633,f73])).
fof(f1684,plain,(
  ( ! [X0] : (~p(s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int',X0))))) | ~p(s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))))))))) | ~p(s(bool,sK1(X0)))) ) | $spl15),
  inference(superposition,[],[f1633,f54])).
fof(f2037,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X0))) | ~p(s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))))))))) | ~p(s(bool,sK1(X0)))) ) | $spl15),
  inference(superposition,[],[f1673,f54])).
fof(f1686,plain,(
  ~p(s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | ~p(s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))))))))) | $spl15),
  inference(superposition,[],[f1633,f43])).
fof(f1856,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))))))))) | $spl15),
  inference(superposition,[],[f1634,f84])).
fof(f1683,plain,(
  ~p(s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl15),
  inference(superposition,[],[f1633,f40])).
fof(f1637,plain,(
  s(bool,'F_0') = s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0)))))))) | $spl15),
  inference(unit_resulting_resolution,[],[f35,f1629,f279])).
fof(f1665,plain,(
  ( ! [X3] : (s(bool,X3) = s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0)))))))) | p(s(bool,X3))) ) | $spl15),
  inference(subsumption_resolution,[],[f1649,f33])).
fof(f1649,plain,(
  ( ! [X3] : (~p(s(bool,'T_0')) | s(bool,X3) = s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0)))))))) | p(s(bool,X3))) ) | $spl15),
  inference(superposition,[],[f1629,f219])).
fof(f1860,plain,(
  ~p(s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))))))))) | $spl15),
  inference(trivial_inequality_removal,[],[f1859])).
fof(f1859,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))))))))) | $spl15),
  inference(superposition,[],[f1634,f43])).
fof(f2044,plain,(
  ~p(s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))))))))))) | $spl15),
  inference(trivial_inequality_removal,[],[f2043])).
fof(f2043,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))))))))))) | $spl15),
  inference(superposition,[],[f1673,f43])).
fof(f2166,plain,(
  ( ! [X33,X34,X32] : (s(bool,'T_0') != s(bool,f(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(fun('type.integer.int',bool),X32) = s(fun('type.integer.int',bool),X33) | ~p(s(bool,sK1(sK2(X32,X33,bool,'type.integer.int')))) | ~p(s(bool,X34)) | s(bool,X34) != s(bool,g(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl3),
  inference(superposition,[],[f253,f617])).
fof(f2165,plain,(
  ( ! [X30,X31] : (s(bool,'F_0') != s(bool,f(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(fun('type.integer.int',bool),X30) = s(fun('type.integer.int',bool),X31) | ~p(s(bool,sK1(sK2(X30,X31,bool,'type.integer.int')))) | p(s(bool,g(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl3),
  inference(superposition,[],[f253,f299])).
fof(f2093,plain,(
  ( ! [X39,X41,X40] : (s(bool,X41) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X39))),s('type.integer.int',X40))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X39))))))),s('type.integer.int',X40)))) | ~p(s(bool,X41)) | ~p(s(bool,sK1(X40)))) ) | $spl3),
  inference(superposition,[],[f379,f347])).
fof(f2092,plain,(
  ( ! [X37,X38,X36] : (s(bool,X38) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X36))),s('type.integer.int',X37))) | p(s(bool,X38)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X36))))))),s('type.integer.int',X37)))) | ~p(s(bool,sK1(X37)))) ) | $spl3),
  inference(superposition,[],[f379,f347])).
fof(f2091,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X34))),s('type.integer.int',X35))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X34))))))),s('type.integer.int',X35)))) | ~p(s(bool,sK1(X35)))) ) | $spl3),
  inference(superposition,[],[f320,f347])).
fof(f2078,plain,(
  ( ! [X37,X35,X36] : (s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X35))),s('type.integer.int',X36))) | ~p(s(bool,sK1(X36))) | ~p(s(bool,X37)) | s(bool,X37) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X35))))))),s('type.integer.int',X36)))) ) | $spl3),
  inference(superposition,[],[f347,f617])).
fof(f2077,plain,(
  ( ! [X33,X34] : (s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X33))),s('type.integer.int',X34))) | ~p(s(bool,sK1(X34))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X33))))))),s('type.integer.int',X34))))) ) | $spl3),
  inference(superposition,[],[f347,f299])).
fof(f2004,plain,(
  s(bool,'T_0') != s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))))))))) | ($spl3 | $spl15)),
  inference(superposition,[],[f1682,f43])).
fof(f2002,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int',X0)))) | ~p(s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))))))))) | ~p(s(bool,sK1(X0)))) ) | ($spl3 | $spl15)),
  inference(superposition,[],[f1682,f54])).
fof(f2001,plain,(
  s(bool,'T_0') != s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ($spl3 | $spl15)),
  inference(superposition,[],[f1682,f40])).
fof(f1974,plain,(
  ( ! [X33,X34,X32] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X32) = s(fun(bool,bool),X33) | s(bool,'T_0') = s(bool,sK2(X32,X33,bool,bool)) | ~p(s(bool,X34)) | s(bool,X34) != s(bool,g(s(bool,'F_0')))) ) | $spl3),
  inference(superposition,[],[f250,f617])).
fof(f1973,plain,(
  ( ! [X30,X31] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X30) = s(fun(bool,bool),X31) | s(bool,'T_0') = s(bool,sK2(X30,X31,bool,bool)) | p(s(bool,g(s(bool,'F_0'))))) ) | $spl3),
  inference(superposition,[],[f250,f299])).
fof(f1949,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(bool,X0))) != s(X1,g(s(bool,X0))) | s(fun(bool,X1),X2) = s(fun(bool,X1),X3) | s(bool,'T_0') = s(bool,sK2(X2,X3,X1,bool)) | p(s(bool,X0))) ) | $spl3),
  inference(superposition,[],[f250,f299])).
fof(f1891,plain,(
  s(bool,'F_0') = s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))))))))) | ($spl3 | $spl15)),
  inference(superposition,[],[f1680,f43])).
fof(f1889,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int',X0)))) | ~p(s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))))))))) | ~p(s(bool,sK1(X0)))) ) | ($spl3 | $spl15)),
  inference(superposition,[],[f1680,f54])).
fof(f1888,plain,(
  s(bool,'F_0') = s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ($spl3 | $spl15)),
  inference(superposition,[],[f1680,f40])).
fof(f1874,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl15),
  inference(superposition,[],[f1634,f197])).
fof(f1746,plain,(
  s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))))))))) = s('type.integer.int','const.bool.COND_3'(s(bool,'F_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | ($spl3 | $spl15)),
  inference(superposition,[],[f45,f1641])).
fof(f1732,plain,(
  ( ! [X21,X22] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X21))) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X21))))))))) ) | $spl3),
  inference(subsumption_resolution,[],[f1726,f33])).
fof(f1726,plain,(
  ( ! [X21,X22] : (~p(s(bool,'T_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X21))) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X21))))))))) ) | $spl3),
  inference(superposition,[],[f141,f617])).
fof(f1706,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0)))))))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl15),
  inference(subsumption_resolution,[],[f1688,f33])).
fof(f1688,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0)))))))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl15),
  inference(superposition,[],[f1633,f48])).
fof(f1673,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))))))))) | $spl15),
  inference(unit_resulting_resolution,[],[f1633,f44])).
fof(f1680,plain,(
  s(bool,'F_0') = s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0)))))))))) | ($spl3 | $spl15)),
  inference(unit_resulting_resolution,[],[f1633,f299])).
fof(f1682,plain,(
  s(bool,'T_0') != s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0)))))))))) | ($spl3 | $spl15)),
  inference(unit_resulting_resolution,[],[f33,f1633,f379])).
fof(f1660,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0)))))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl15),
  inference(subsumption_resolution,[],[f1645,f33])).
fof(f1645,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0)))))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl15),
  inference(superposition,[],[f1629,f48])).
fof(f1634,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))))))) | $spl15),
  inference(unit_resulting_resolution,[],[f1629,f44])).
fof(f1641,plain,(
  s(bool,'F_0') = s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0)))))))) | ($spl3 | $spl15)),
  inference(unit_resulting_resolution,[],[f1629,f299])).
fof(f1643,plain,(
  s(bool,'T_0') != s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0)))))))) | ($spl3 | $spl15)),
  inference(unit_resulting_resolution,[],[f33,f1629,f379])).
fof(f1538,plain,(
  ( ! [X35,X33,X34] : (s(bool,'T_0') != s(bool,X34) | p(s(bool,X33)) | s(bool,X34) = s(bool,X35) | ~p(s(bool,X35))) ) | $spl3),
  inference(duplicate_literal_removal,[],[f1510])).
fof(f1510,plain,(
  ( ! [X35,X33,X34] : (s(bool,'T_0') != s(bool,X34) | p(s(bool,X33)) | s(bool,X34) = s(bool,X35) | p(s(bool,X33)) | ~p(s(bool,X35))) ) | $spl3),
  inference(superposition,[],[f320,f267])).
fof(f1424,plain,(
  ( ! [X15,X16] : (p(s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac',X15))))) | ~p(s(bool,X16)) | s(bool,X16) != s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X15))))))) ) | $spl3),
  inference(subsumption_resolution,[],[f1419,f33])).
fof(f1419,plain,(
  ( ! [X15,X16] : (~p(s(bool,'T_0')) | p(s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac',X15))))) | ~p(s(bool,X16)) | s(bool,X16) != s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X15))))))) ) | $spl3),
  inference(superposition,[],[f971,f617])).
fof(f1353,plain,(
  ( ! [X61,X64,X62,X63] : (s(bool,'T_0') != s(bool,f(s(X61,sK2(X62,X63,bool,X61)))) | s(fun(X61,bool),X62) = s(fun(X61,bool),X63) | ~p(s(bool,X64)) | s(bool,X64) != s(bool,g(s(X61,sK2(X62,X63,bool,X61))))) ) | $spl3),
  inference(superposition,[],[f46,f617])).
fof(f1352,plain,(
  ( ! [X59,X57,X60,X58] : (s(X59,f(s(bool,'T_0'))) != s(X59,g(s(bool,'T_0'))) | s(fun(bool,X59),X57) = s(fun(bool,X59),X58) | ~p(s(bool,X60)) | s(bool,X60) != s(bool,sK2(X57,X58,X59,bool))) ) | $spl3),
  inference(superposition,[],[f46,f617])).
fof(f1351,plain,(
  ( ! [X52,X53] : (s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X52))) = s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X52),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | ~p(s(bool,X53)) | s(bool,X53) != s(bool,sK1(X52))) ) | $spl3),
  inference(superposition,[],[f45,f617])).
fof(f1350,plain,(
  ( ! [X50,X51] : (s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X50))) = s('type.integer.int','const.bool.COND_3'(s(bool,sK1(X50)),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | ~p(s(bool,X51)) | s(bool,X51) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',X50),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl3),
  inference(superposition,[],[f45,f617])).
fof(f1321,plain,(
  ( ! [X14,X15,X16] : (s(bool,X14) != s(bool,X15) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X14) != s(bool,X16)) ) | $spl3),
  inference(superposition,[],[f320,f617])).
fof(f1299,plain,(
  ( ! [X14,X12,X13] : (s(bool,X12) = s(bool,X13) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X14) | ~p(s(bool,X12))) ) | $spl3),
  inference(superposition,[],[f617,f50])).
fof(f1298,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,X9) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X10) | s(bool,X8) = s(bool,X11) | s(bool,'T_0') = s(bool,X11)) ) | $spl3),
  inference(superposition,[],[f617,f48])).
fof(f1297,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X4) = s(bool,X5) | ~p(s(bool,X6)) | s(bool,X5) != s(bool,X6) | s(bool,X4) = s(bool,X7) | p(s(bool,X7))) ) | $spl3),
  inference(superposition,[],[f617,f219])).
fof(f1296,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X2) | ~p(s(bool,X3)) | s(bool,X0) != s(bool,X3)) ) | $spl3),
  inference(superposition,[],[f617,f617])).
fof(f299,plain,(
  ( ! [X16] : (s(bool,'F_0') = s(bool,X16) | p(s(bool,X16))) ) | $spl3),
  inference(trivial_inequality_removal,[],[f288])).
fof(f288,plain,(
  ( ! [X16] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'F_0') = s(bool,X16) | p(s(bool,X16))) ) | $spl3),
  inference(superposition,[],[f70,f219])).
fof(f337,plain,(
  ( ! [X14,X15,X16] : (s(X16,f(s(bool,'F_0'))) != s(X16,g(s(bool,'F_0'))) | s(fun(bool,X16),X14) = s(fun(bool,X16),X15) | p(s(bool,sK2(X14,X15,X16,bool)))) ) | $spl3),
  inference(superposition,[],[f46,f299])).
fof(f338,plain,(
  ( ! [X19,X17,X18] : (s(bool,'F_0') != s(bool,f(s(X17,sK2(X18,X19,bool,X17)))) | s(fun(X17,bool),X18) = s(fun(X17,bool),X19) | p(s(bool,g(s(X17,sK2(X18,X19,bool,X17)))))) ) | $spl3),
  inference(superposition,[],[f46,f299])).
fof(f418,plain,(
  ( ! [X10] : (s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X10))) = s('type.integer.int','const.bool.COND_3'(s(bool,'F_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X10),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | p(s(bool,sK1(X10)))) ) | $spl3),
  inference(superposition,[],[f45,f299])).
fof(f437,plain,(
  ( ! [X11] : (s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X11))) = s('type.integer.int','const.bool.COND_3'(s(bool,sK1(X11)),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'F_0'),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X11),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl3),
  inference(superposition,[],[f45,f299])).
fof(f617,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') = s(bool,X34) | ~p(s(bool,X35)) | s(bool,X34) != s(bool,X35)) ) | $spl3),
  inference(subsumption_resolution,[],[f602,f35])).
fof(f602,plain,(
  ( ! [X35,X34] : (s(bool,X34) != s(bool,X35) | p(s(bool,'F_0')) | ~p(s(bool,X35)) | s(bool,'T_0') = s(bool,X34)) ) | $spl3),
  inference(superposition,[],[f379,f38])).
fof(f614,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,X1) | ~p(s(bool,X1)) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(subsumption_resolution,[],[f589,f50])).
fof(f589,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(superposition,[],[f379,f38])).
fof(f379,plain,(
  ( ! [X6,X7] : (s(bool,X6) != s(bool,X7) | p(s(bool,X7)) | ~p(s(bool,X6))) ) | $spl3),
  inference(superposition,[],[f320,f50])).
fof(f378,plain,(
  ( ! [X4,X5,X3] : (s(bool,X3) != s(bool,X4) | p(s(bool,X4)) | s(bool,X3) = s(bool,X5) | s(bool,'T_0') = s(bool,X5)) ) | $spl3),
  inference(superposition,[],[f320,f48])).
fof(f377,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | p(s(bool,X1)) | s(bool,X0) = s(bool,X2) | p(s(bool,X2))) ) | $spl3),
  inference(superposition,[],[f320,f219])).
fof(f320,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) ) | $spl3),
  inference(superposition,[],[f70,f299])).
fof(f3963,plain,(
  $false | ($spl3 | $spl20)),
  inference(subsumption_resolution,[],[f3680,f35])).
fof(f3680,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl20)),
  inference(unit_resulting_resolution,[],[f70,f3626,f279])).
fof(f3962,plain,(
  $false | ($spl3 | $spl15 | $spl20)),
  inference(subsumption_resolution,[],[f3682,f3626])).
fof(f3682,plain,(
  p(s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))))))))) | ($spl3 | $spl15 | $spl20)),
  inference(unit_resulting_resolution,[],[f3394,f3626,f279])).
fof(f3694,plain,(
  $false | ($spl10 | $spl20)),
  inference(unit_resulting_resolution,[],[f33,f89,f3626,f2685])).
fof(f3691,plain,(
  $false | ($spl3 | $spl15 | $spl20)),
  inference(unit_resulting_resolution,[],[f70,f3626,f2175])).
fof(f3689,plain,(
  $false | ($spl10 | $spl20)),
  inference(unit_resulting_resolution,[],[f89,f3626,f753])).
fof(f3687,plain,(
  $false | ($spl3 | $spl15 | $spl20)),
  inference(unit_resulting_resolution,[],[f3626,f3394,f3626,f279])).
fof(f3685,plain,(
  $false | ($spl3 | $spl20)),
  inference(unit_resulting_resolution,[],[f3626,f70,f3626,f279])).
fof(f3683,plain,(
  $false | ($spl3 | $spl15 | $spl20)),
  inference(unit_resulting_resolution,[],[f3626,f3394,f3626,f279])).
fof(f3681,plain,(
  $false | ($spl3 | $spl20)),
  inference(unit_resulting_resolution,[],[f35,f70,f3626,f279])).
fof(f3677,plain,(
  $false | ($spl3 | $spl15 | $spl20)),
  inference(unit_resulting_resolution,[],[f3626,f3394,f3626,f279])).
fof(f3675,plain,(
  $false | ($spl3 | $spl20)),
  inference(unit_resulting_resolution,[],[f35,f70,f3626,f279])).
fof(f3673,plain,(
  $false | ($spl3 | $spl15 | $spl20)),
  inference(unit_resulting_resolution,[],[f3626,f3394,f3626,f279])).
fof(f3671,plain,(
  $false | ($spl3 | $spl20)),
  inference(unit_resulting_resolution,[],[f3626,f70,f3626,f279])).
fof(f3665,plain,(
  $false | ($spl3 | $spl15 | $spl20)),
  inference(unit_resulting_resolution,[],[f33,f3394,f3394,f3626,f267])).
fof(f3662,plain,(
  $false | ($spl3 | $spl20)),
  inference(unit_resulting_resolution,[],[f33,f70,f70,f3626,f267])).
fof(f3655,plain,(
  $false | ($spl3 | $spl15 | $spl20)),
  inference(unit_resulting_resolution,[],[f3394,f3394,f3626,f219])).
fof(f3653,plain,(
  $false | ($spl3 | $spl20)),
  inference(unit_resulting_resolution,[],[f70,f70,f3626,f219])).
fof(f3647,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f33,f3626])).
fof(f3626,plain,(
  ( ! [X30] : (~p(s(bool,X30))) ) | $spl20),
  inference(cnf_transformation,[],[f3626_D])).
fof(f3626_D,plain,(
  ( ! [X30] : (~p(s(bool,X30))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f3627,plain,(
  ( ! [X31] : (s(bool,'F_0') != s(bool,X31) | s(bool,'T_0') != s(bool,X31)) ) | $spl22),
  inference(cnf_transformation,[],[f3627_D])).
fof(f3627_D,plain,(
  ( ! [X31] : (s(bool,'F_0') != s(bool,X31) | s(bool,'T_0') != s(bool,X31)) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f3370,plain,(
  ( ! [X30,X31] : (s(bool,'F_0') != s(bool,X31) | ~p(s(bool,X30)) | s(bool,'T_0') != s(bool,X31)) ) | ($spl3 | $spl6)),
  inference(duplicate_literal_removal,[],[f3341])).
fof(f3341,plain,(
  ( ! [X30,X31] : (s(bool,'F_0') != s(bool,X31) | ~p(s(bool,X30)) | s(bool,'T_0') != s(bool,X31) | ~p(s(bool,X30))) ) | ($spl3 | $spl6)),
  inference(superposition,[],[f648,f2587])).
fof(f2547,plain,(
  $false | $spl0),
  inference(equality_resolution,[],[f2434])).
fof(f2434,plain,(
  ( ! [X0] : (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl0),
  inference(subsumption_resolution,[],[f2290,f35])).
fof(f2290,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl0),
  inference(backward_demodulation,[],[f68,f44])).
fof(f2198,plain,(
  $false | ($spl2 | $spl6)),
  inference(subsumption_resolution,[],[f69,f1061])).
fof(f1061,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl6),
  inference(unit_resulting_resolution,[],[f33,f648])).
fof(f69,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl2),
  inference(cnf_transformation,[],[f69_D])).
fof(f69_D,plain,(
  s(bool,'T_0') = s(bool,'F_0') <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f2170,plain,(
  ( ! [X41,X42] : (s(bool,'F_0') = s(bool,X41) | ~p(s(bool,X42)) | s(bool,X41) = s(bool,X42)) ) | $spl18),
  inference(cnf_transformation,[],[f2170_D])).
fof(f2170_D,plain,(
  ( ! [X41,X42] : (s(bool,'F_0') = s(bool,X41) | ~p(s(bool,X42)) | s(bool,X41) = s(bool,X42)) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f1543,plain,(
  ( ! [X41,X42,X40] : (s(bool,'F_0') = s(bool,X41) | s(bool,X41) = s(bool,X42) | p(s(bool,X40)) | ~p(s(bool,X42))) ) | $spl3),
  inference(duplicate_literal_removal,[],[f1459])).
fof(f1459,plain,(
  ( ! [X41,X42,X40] : (s(bool,'F_0') = s(bool,X41) | s(bool,X41) = s(bool,X42) | p(s(bool,X40)) | ~p(s(bool,X42)) | p(s(bool,X40))) ) | $spl3),
  inference(superposition,[],[f267,f299])).
fof(f1629,plain,(
  ~p(s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))))))))) | $spl15),
  inference(cnf_transformation,[],[f1629_D])).
fof(f1629_D,plain,(
  ~p(s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))))))))) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f1631,plain,(
  s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',sK0))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl17),
  inference(cnf_transformation,[],[f1631_D])).
fof(f1631_D,plain,(
  s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',sK0))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f96,plain,(
  s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',sK0))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0)))))))))),
  inference(superposition,[],[f32,f43])).
fof(f1267,plain,(
  ( ! [X15] : (s(bool,'F_0') = s(bool,X15) | p(s(bool,X15))) ) | $spl12),
  inference(cnf_transformation,[],[f1267_D])).
fof(f1267_D,plain,(
  ( ! [X15] : (s(bool,'F_0') = s(bool,X15) | p(s(bool,X15))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f752,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl8),
  inference(cnf_transformation,[],[f752_D])).
fof(f752_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f300,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') = s(bool,X14) | s(bool,'F_0') = s(bool,X15) | p(s(bool,X15))) )),
  inference(duplicate_literal_removal,[],[f276])).
fof(f276,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') = s(bool,X14) | s(bool,'F_0') = s(bool,X15) | p(s(bool,X15)) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f219,f38])).
fof(f969,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f968,f927])).
fof(f927,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl8),
  inference(subsumption_resolution,[],[f814,f33])).
fof(f814,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl8),
  inference(backward_demodulation,[],[f752,f43])).
fof(f968,plain,(
  s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',sK0))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl8),
  inference(subsumption_resolution,[],[f850,f33])).
fof(f850,plain,(
  ~p(s(bool,'T_0')) | s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',sK0))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl8),
  inference(backward_demodulation,[],[f752,f668])).
fof(f967,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f966,f915])).
fof(f915,plain,(
  ( ! [X2,X1] : (s('type.integer.int',X2) = s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X1)))) ) | $spl8),
  inference(subsumption_resolution,[],[f914,f33])).
fof(f914,plain,(
  ( ! [X2,X1] : (~p(s(bool,'T_0')) | s('type.integer.int',X2) = s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X1)))) ) | $spl8),
  inference(forward_demodulation,[],[f913,f752])).
fof(f913,plain,(
  ( ! [X2,X1] : (s('type.integer.int',X2) = s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X1))) | ~p(s(bool,sK1(X2)))) ) | $spl8),
  inference(subsumption_resolution,[],[f912,f33])).
fof(f912,plain,(
  ( ! [X2,X1] : (~p(s(bool,'T_0')) | s('type.integer.int',X2) = s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X1))) | ~p(s(bool,sK1(X2)))) ) | $spl8),
  inference(forward_demodulation,[],[f796,f752])).
fof(f796,plain,(
  ( ! [X2,X1] : (~p(s(bool,sK1('const.bool.COND_3'(s(bool,sK1(X1)),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))))) | s('type.integer.int',X2) = s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X1))) | ~p(s(bool,sK1(X2)))) ) | $spl8),
  inference(backward_demodulation,[],[f752,f447])).
fof(f966,plain,(
  s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',sK0))) != s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl8),
  inference(subsumption_resolution,[],[f846,f33])).
fof(f846,plain,(
  ~p(s(bool,'T_0')) | s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',sK0))) != s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl8),
  inference(backward_demodulation,[],[f752,f632])).
fof(f962,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f961,f927])).
fof(f961,plain,(
  s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',sK0))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl8),
  inference(subsumption_resolution,[],[f832,f33])).
fof(f832,plain,(
  ~p(s(bool,'T_0')) | s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',sK0))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl8),
  inference(backward_demodulation,[],[f752,f96])).
fof(f944,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f943,f33])).
fof(f943,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(forward_demodulation,[],[f942,f752])).
fof(f942,plain,(
  ~p(s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))))))))) | $spl8),
  inference(subsumption_resolution,[],[f941,f937])).
fof(f937,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X0) = s('type.integer.int',X1)) ) | $spl8),
  inference(subsumption_resolution,[],[f936,f33])).
fof(f936,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s('type.integer.int',X0) = s('type.integer.int',X1)) ) | $spl8),
  inference(forward_demodulation,[],[f935,f752])).
fof(f935,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X0) = s('type.integer.int',X1) | ~p(s(bool,sK1(X1)))) ) | $spl8),
  inference(subsumption_resolution,[],[f816,f33])).
fof(f816,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s('type.integer.int',X0) = s('type.integer.int',X1) | ~p(s(bool,sK1(X1)))) ) | $spl8),
  inference(backward_demodulation,[],[f752,f54])).
fof(f941,plain,(
  ( ! [X5] : (s('type.integer.int',X5) != s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',sK0))) | ~p(s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0)))))))))) ) | $spl8),
  inference(subsumption_resolution,[],[f819,f33])).
fof(f819,plain,(
  ( ! [X5] : (~p(s(bool,'T_0')) | s('type.integer.int',X5) != s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',sK0))) | ~p(s(bool,sK1('const.frac.frac_sgn_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0)))))))))) ) | $spl8),
  inference(backward_demodulation,[],[f752,f154])).
fof(f870,plain,(
  $false | ($spl3 | $spl6 | $spl8)),
  inference(subsumption_resolution,[],[f869,f752])).
fof(f869,plain,(
  ( ! [X23] : (s(bool,'T_0') != s(bool,X23)) ) | ($spl3 | $spl6 | $spl8)),
  inference(subsumption_resolution,[],[f767,f320])).
fof(f767,plain,(
  ( ! [X23] : (s(bool,'T_0') != s(bool,X23) | ~p(s(bool,X23))) ) | ($spl6 | $spl8)),
  inference(backward_demodulation,[],[f752,f648])).
fof(f860,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f754,f33])).
fof(f754,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(backward_demodulation,[],[f752,f35])).
fof(f753,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,X1) | p(s(bool,X1))) ) | $spl10),
  inference(cnf_transformation,[],[f753_D])).
fof(f753_D,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,X1) | p(s(bool,X1))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f296,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,'T_0') = s(bool,X0) | p(s(bool,X1))) )),
  inference(equality_factoring,[],[f219])).
fof(f709,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f708,f647])).
fof(f708,plain,(
  ~p(s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))))))))))) | $spl4),
  inference(subsumption_resolution,[],[f668,f694])).
fof(f694,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X0)))) ) | $spl4),
  inference(subsumption_resolution,[],[f682,f647])).
fof(f682,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X0))) | ~p(s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac',X0)))))) ) | $spl4),
  inference(backward_demodulation,[],[f681,f84])).
fof(f681,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl4),
  inference(subsumption_resolution,[],[f680,f647])).
fof(f680,plain,(
  ( ! [X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X1)))))))) ) | $spl4),
  inference(subsumption_resolution,[],[f654,f647])).
fof(f654,plain,(
  ( ! [X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,sK1('const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X1))))))) | ~p(s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac',X1)))))) )),
  inference(superposition,[],[f85,f84])).
fof(f698,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f697,f647])).
fof(f697,plain,(
  ~p(s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))))))))) | $spl4),
  inference(subsumption_resolution,[],[f688,f694])).
fof(f688,plain,(
  s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',sK0))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,sK1('const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))))))))) | $spl4),
  inference(backward_demodulation,[],[f681,f632])).
fof(f648,plain,(
  ( ! [X23] : (s(bool,'F_0') != s(bool,X23) | ~p(s(bool,X23))) ) | $spl6),
  inference(cnf_transformation,[],[f648_D])).
fof(f648_D,plain,(
  ( ! [X23] : (s(bool,'F_0') != s(bool,X23) | ~p(s(bool,X23))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f606,plain,(
  ( ! [X23,X22] : (s(bool,'F_0') != s(bool,X23) | p(s(bool,X22)) | ~p(s(bool,X23))) ) | $spl3),
  inference(duplicate_literal_removal,[],[f597])).
fof(f597,plain,(
  ( ! [X23,X22] : (s(bool,'F_0') != s(bool,X23) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X22))) ) | $spl3),
  inference(superposition,[],[f379,f299])).
fof(f68,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl0),
  inference(cnf_transformation,[],[f68_D])).
fof(f68_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f70,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl3),
  inference(cnf_transformation,[],[f70_D])).
fof(f70_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f53,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f38])).
% SZS output end Proof for thm.rat
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 154
% Passive clauses: 2367
% Generated clauses: 11913
% Final active clauses: 48
% Final passive clauses: 893
% Input formulas: 11
% Initial clauses: 15
% 
% Duplicate literals: 364
% Trivial inequalities: 216
% Fw subsumption resolutions: 1047
% Fw demodulations: 1934
% Bw demodulations: 1226
% Global subsumptions: 4
% 
% Simple tautologies: 322
% Equational tautologies: 137
% Forward subsumptions: 5488
% Bw demodulations to eq. taut.: 5
% 
% Unit resulting resolution: 742
% Forward superposition: 3548
% Backward superposition: 1988
% Self superposition: 96
% Equality factoring: 26
% Equality resolution: 23
% 
% Split clauses: 11
% Split components: 22
% SAT solver clauses: 2418
% SAT solver unit clauses: 46
% SAT solver binary clauses: 333
% 
% Sat splits: 11
% Sat splitting refutations: 34
% 
% Memory used [KB]: 8699
% Time elapsed: 0.558 s
% ------------------------------
% ------------------------------
% Success in time 16.289 s
