Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.list
% (85059)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 883
% Passive clauses: 8090
% Generated clauses: 23939
% Final active clauses: 873
% Final passive clauses: 6888
% Input formulas: 60
% Initial clauses: 236
% 
% Duplicate literals: 608
% Trivial inequalities: 93
% Fw subsumption resolutions: 2763
% Fw demodulations: 1515
% Bw demodulations: 10
% 
% Simple tautologies: 321
% Equational tautologies: 41
% Forward subsumptions: 9626
% 
% Binary resolution: 9093
% Factoring: 1528
% Forward superposition: 3480
% Backward superposition: 4253
% Self superposition: 36
% Equality factoring: 28
% Equality resolution: 20
% 
% Split clauses: 703
% Split components: 1413
% SAT solver clauses: 1117
% SAT solver unit clauses: 17
% SAT solver binary clauses: 198
% 
% Sat splits: 703
% 
% TWLsolver clauses: 1117
% TWLsolver calls for satisfiability: 573
% 
% Memory used [KB]: 12153
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.list
% (85083)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 418
% Passive clauses: 2489
% Generated clauses: 14025
% Final active clauses: 413
% Final passive clauses: 2066
% Input formulas: 60
% Initial clauses: 236
% 
% Duplicate literals: 306
% Trivial inequalities: 94
% Fw subsumption resolutions: 1160
% Fw demodulations: 1383
% Bw demodulations: 9
% Condensations: 369
% Global subsumptions: 22
% 
% Simple tautologies: 338
% Equational tautologies: 19
% Forward subsumptions: 4786
% 
% Binary resolution: 2796
% Factoring: 119
% Forward superposition: 2657
% Backward superposition: 2280
% Self superposition: 43
% Equality factoring: 30
% Equality resolution: 15
% 
% Split clauses: 87
% Split components: 179
% SAT solver clauses: 3309
% SAT solver unit clauses: 225
% SAT solver binary clauses: 93
% SAT solver learnt clauses: 44
% SAT solver learnt literals: 10
% 
% Sat splits: 87
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 3221
% TWLsolver calls for satisfiability: 1895
% 
% Memory used [KB]: 13560
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.list
% (85100)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 5114
% Passive clauses: 21062
% Generated clauses: 129472
% Final active clauses: 1319
% Final passive clauses: 2700
% Input formulas: 60
% Initial clauses: 236
% 
% Duplicate literals: 3724
% Trivial inequalities: 212
% Fw subsumption resolutions: 3122
% Fw demodulations: 23212
% Bw demodulations: 5994
% Condensations: 1625
% Global subsumptions: 148
% 
% Simple tautologies: 189
% Equational tautologies: 350
% Forward subsumptions: 73936
% Bw demodulations to eq. taut.: 41
% 
% Unit resulting resolution: 56163
% Factoring: 1
% Forward superposition: 23357
% Backward superposition: 10613
% Self superposition: 79
% Equality factoring: 36
% Equality resolution: 22
% 
% Split clauses: 121
% Split components: 247
% SAT solver clauses: 22956
% SAT solver unit clauses: 638
% SAT solver binary clauses: 8131
% 
% Sat splits: 121
% Sat splitting refutations: 1501
% 
% Memory used [KB]: 35052
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.list
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.list
% SZS output start Proof for thm.list
fof(f34072,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f304,f546,f545,f303,f548,f302,f544,f552,f554,f555,f556,f559,f649,f676,f711,f713,f714,f716,f718,f809,f882,f884,f885,f886,f889,f979,f1091,f1093,f1094,f1096,f1098,f1189,f1655,f1879,f1880,f1886,f1888,f1889,f1891,f1893,f1984,f2137,f2207,f2208,f2209,f2210,f2279,f2281,f2282,f2283,f2286,f2376,f1007,f2410,f2523,f2524,f2527,f2528,f2537,f2539,f2540,f2542,f2544,f2635,f2011,f4156,f4158,f4159,f4160,f4163,f4253,f4280,f6154,f6157,f6207,f6221,f6223,f6282,f6285,f6286,f6287,f3574,f3575,f3576,f3577,f6289,f6291,f6290,f6292,f6293,f6357,f7169,f7167,f7159,f7158,f7157,f7248,f7249,f7250,f7251,f7252,f7253,f6603,f6601,f6600,f6599,f6284,f6612,f6611,f6610,f6609,f7303,f7317,f7425,f7424,f7430,f7431,f2534,f7433,f7434,f7435,f7436,f7437,f7438,f7439,f7440,f2518,f11296,f6155,f11298,f6158,f11347,f6208,f11360,f6222,f11361,f6224,f11535,f11537,f11538,f11540,f11542,f11633,f11703,f11636,f11637,f11638,f11640,f11641,f11642,f11643,f11644,f11645,f11654,f11656,f11714,f11720,f11721,f11698,f11697,f11709,f6584,f6565,f6564,f6571,f6570,f6583,f2512,f2521,f2520,f2533,f2532,f2531,f2536,f2513,f11808,f549,f11887,f11886,f11892,f11893,f2519,f11896,f11897,f11900,f11901,f11904,f11905,f550,f11906,f11907,f2535,f12699,f12690,f12689,f12688,f12779,f12780,f12781,f12782,f12783,f12784,f13434,f13442,f14129,f14131,f14185,f14187,f14188,f14190,f14192,f14283,f837,f14586,f14706,f14705,f14704,f14793,f14794,f14795,f14796,f14797,f14798,f14807,f14810,f14590,f16000,f16008,f16014,f16021,f16028,f16293,f16331,f16332,f16336,f16342,f16350,f16356,f16368,f16375,f2413,f16429,f16430,f16457,f16458,f16459,f16460,f16494,f16615,f16617,f16618,f16619,f16622,f16712,f16783,f16718,f16719,f16720,f16788,f16722,f16723,f16724,f16725,f16726,f16727,f16796,f16797,f16773,f16772,f16787,f16840,f16842,f16843,f16844,f16847,f16937,f17008,f16943,f16944,f16945,f17012,f16947,f16948,f16949,f16950,f16951,f16952,f17020,f17021,f16998,f16997,f17011,f17148,f17312,f17149,f18521,f18714,f18550,f18552,f18553,f18554,f18557,f18647,f18727,f18653,f18654,f18655,f18731,f18657,f18658,f18659,f18660,f18661,f18662,f18739,f18740,f18709,f18708,f18730,f18766,f18768,f18765,f6288,f18769,f18767,f18770,f18771,f18791,f18804,f18806,f18807,f18809,f18811,f18902,f18974,f18908,f18909,f18910,f18912,f18913,f18914,f18915,f18916,f18917,f18926,f18928,f18986,f18989,f18990,f18969,f18968,f18981,f18991,f19015,f19016,f19059,f19066,f19063,f19078,f19088,f19099,f19108,f19118,f19119,f19125,f19126,f19132,f19142,f18773,f19231,f19233,f19301,f19303,f19304,f19306,f19308,f19399,f19471,f19405,f19406,f19407,f19409,f19410,f19411,f19412,f19413,f19414,f19423,f19425,f19484,f19485,f19466,f19465,f19477,f19486,f19528,f19530,f19539,f19541,f19542,f19543,f19546,f19636,f19704,f19715,f19722,f19848,f19854,f19856,f19857,f19859,f19861,f19952,f20035,f19959,f19960,f19961,f19963,f19964,f19965,f19966,f19967,f19968,f20031,f20029,f20051,f20052,f20024,f20023,f20041,f20053,f20143,f20145,f20146,f20147,f20150,f20240,f20309,f20320,f20327,f20495,f20566,f20619,f20631,f20654,f20657,f20659,f20660,f20662,f20664,f20755,f20829,f20839,f20849,f20856,f20853,f20868,f20878,f20889,f20898,f20937,f20938,f20944,f20945,f20951,f20961,f20841,f21019,f21042,f21029,f21086,f21099,f21090,f21148,f21139,f21184,f21237,f21229,f21291,f21335,f21336,f21376,f21377,f21417,f21461,f19487,f21521,f21534,f21525,f21527,f21535,f21686,f21677,f21678,f21687,f21775,f21777,f21785,f21868,f21860,f21903,f21906,f21913,f21993,f21994,f21995,f22001,f22043,f22044,f22045,f22051,f22122,f22124,f22132,f22194,f20054,f22195,f22202,f22359,f22496,f22497,f22480,f22702,f22534,f22536,f22537,f22538,f22541,f22631,f22715,f22733,f22737,f22772,f22774,f22762,f22792,f22793,f22795,f22797,f22876,f22917,f22907,f23037,f23111,f23112,f23113,f23114,f23122,f23174,f23177,f23388,f23213,f23215,f23216,f23217,f23220,f23310,f23402,f23420,f23461,f23450,f23452,f23463,f23526,f23567,f23557,f23558,f23568,f23628,f20653,f23680,f23681,f23841,f24004,f23988,f24068,f22769,f24073,f24056,f24074,f24075,f24060,f24126,f22913,f24130,f24115,f24131,f24132,f24005,f24119,f24148,f24157,f24207,f23457,f24211,f24195,f24213,f24214,f24199,f24265,f23563,f24269,f16786,f24254,f24270,f24271,f24212,f24258,f24336,f23170,f24340,f24341,f24409,f23676,f24413,f24414,f24632,f24457,f24459,f24460,f24461,f24464,f24554,f24646,f24861,f24693,f24695,f24696,f24697,f24700,f24790,f24876,f24944,f24945,f24950,f24953,f25023,f25030,f25015,f25083,f25090,f25142,f25151,f25159,f25160,f25165,f25168,f25171,f25234,f25243,f25294,f25303,f25356,f25363,f25415,f25422,f25475,f25482,f25492,f25495,f25498,f25501,f25504,f25564,f25571,f25628,f25635,f25620,f25623,f25637,f25698,f25705,f25690,f25693,f25707,f25727,f25730,f25733,f25736,f25739,f25800,f25809,f25795,f25811,f25902,f25911,f25894,f25897,f25913,f25928,f25931,f25934,f25937,f25940,f25943,f26007,f26016,f26001,f26002,f26019,f26099,f26106,f26094,f26108,f26123,f26124,f26129,f26132,f26135,f26136,f26211,f26218,f26206,f26221,f26320,f26327,f26315,f26329,f26424,f26431,f26419,f23176,f26432,f26537,f25024,f26541,f26528,f26595,f25084,f26599,f26586,f26654,f25143,f26661,f26642,f26716,f25235,f26723,f26707,f26778,f25295,f26785,f26771,f26795,f26797,f26801,f26804,f26807,f26809,f26871,f25357,f26875,f26862,f26928,f25416,f26932,f26986,f25476,f26990,f26977,f27042,f25565,f27046,f27103,f25629,f27107,f27094,f27115,f27117,f27121,f27124,f27127,f27129,f27131,f27194,f25699,f27198,f27185,f27253,f25801,f27260,f27241,f27315,f25903,f27322,f27306,f27377,f26008,f27384,f25152,f27370,f27439,f26100,f27443,f27430,f27526,f26212,f27530,f27516,f27585,f26321,f27589,f27576,f27641,f26425,f24069,f27645,f27661,f27670,f27700,f27702,f27703,f27705,f27707,f27798,f27876,f27828,f27830,f27882,f27887,f27890,f27892,f27893,f27895,f27897,f27988,f28066,f28018,f28020,f28072,f28077,f28098,f28100,f28101,f28103,f28105,f28196,f28274,f28226,f28228,f28280,f28285,f28332,f28334,f28335,f28337,f28339,f28430,f28508,f28460,f28462,f28514,f28519,f28571,f18525,f28572,f19243,f28577,f19718,f28580,f20323,f28583,f22775,f28584,f23462,f28585,f24076,f28586,f24215,f29161,f29224,f29226,f29227,f29229,f29231,f29322,f29410,f29406,f29404,f29424,f29427,f29429,f29430,f29432,f29434,f29525,f29603,f29555,f29557,f19481,f11713,f29617,f29637,f29639,f29670,f29672,f29659,f29693,f29695,f29696,f29698,f29700,f29791,f29832,f29856,f29858,f29859,f29861,f29863,f678,f29954,f29995,f30184,f30185,f30186,f30229,f30230,f30190,f30220,f30233,f30225,f30234,f30235,f30236,f30279,f30280,f30240,f30276,f30285,f30286,f30287,f30326,f30327,f30291,f30322,f30331,f30332,f30333,f30372,f30373,f30337,f30369,f30378,f30379,f30380,f30419,f30420,f30384,f30413,f30423,f30424,f30425,f30464,f30465,f30429,f30460,f30461,f30467,f30468,f30469,f30508,f30509,f30473,f30504,f30505,f30511,f30512,f30513,f30551,f30552,f30517,f30547,f30556,f30557,f30558,f30596,f30597,f30562,f30593,f27888,f30600,f30228,f30601,f30603,f30643,f30644,f30607,f30231,f30636,f30646,f30641,f30647,f30278,f30648,f30650,f30690,f30691,f30654,f30281,f30689,f30695,f30325,f30696,f30698,f30734,f30735,f30702,f30328,f30732,f30739,f30371,f30740,f30742,f30778,f30779,f30746,f30374,f30777,f30783,f30418,f30784,f30786,f30822,f30823,f30790,f30421,f30818,f30825,f30463,f30826,f30828,f30864,f30865,f30832,f30466,f30862,f30863,f30866,f30507,f30867,f30869,f30905,f30906,f30873,f30510,f30903,f30904,f30907,f30550,f30908,f30910,f30945,f30946,f30914,f30553,f30943,f30991,f30595,f30992,f30994,f31029,f31030,f30998,f30598,f31028,f28078,f31032,f31033,f31034,f31077,f31078,f31038,f31068,f31081,f31073,f31082,f31083,f31084,f31127,f31128,f31088,f31124,f31133,f31134,f31135,f31174,f31175,f31139,f31170,f31179,f31180,f31181,f31220,f31221,f31185,f31217,f31227,f31228,f31229,f31268,f31269,f31233,f31262,f31272,f31273,f31274,f31313,f31314,f31278,f31309,f31310,f31316,f31317,f31318,f31357,f31358,f31322,f31353,f31354,f31360,f31361,f31362,f31400,f31401,f31366,f31396,f31405,f31406,f31407,f31445,f31446,f31411,f31442,f28286,f31508,f31076,f31509,f31511,f31551,f31552,f31515,f31079,f31544,f31554,f31549,f31555,f31126,f31556,f31558,f31598,f31599,f31562,f31129,f31597,f31603,f31173,f31604,f31606,f31642,f31643,f31610,f31176,f31640,f31646,f31219,f31647,f31649,f31685,f31686,f31653,f31222,f31684,f31690,f31267,f31691,f31693,f31729,f31730,f31697,f31270,f31725,f31769,f31312,f31770,f31772,f31808,f31809,f31776,f31315,f31806,f31807,f31810,f31356,f31811,f31813,f31849,f31850,f31817,f31359,f31847,f31848,f31851,f31399,f31852,f31854,f31889,f31890,f31858,f31402,f31887,f31893,f31444,f31894,f31896,f19849,f31931,f31932,f31900,f31447,f31930,f28520,f32079,f32065,f32081,f32070,f32143,f32146,f32135,f32204,f32195,f32264,f32266,f32256,f32325,f32314,f32384,f32375,f32376,f32442,f32433,f32434,f32499,f32490,f32591,f32583,f29425,f32654,f32640,f32656,f32645,f32718,f25031,f32721,f21149,f32710,f32779,f32770,f32839,f25091,f32841,f23178,f21238,f32831,f32941,f32930,f33000,f25364,f32991,f7334,f32992,f33058,f25423,f33049,f7342,f33050,f33115,f33106,f33174,f23175,f25572,f33166,f29618,f33242,f33284,f33344,f33409,f33451,f33452,f33515,f33520,f33521,f33460,f33553,f33554,f33617,f33618,f33621,f33620,f33622,f33623,f677,f33619,f33624,f33625,f33562,f16001,f33522,f33680,f33681,f33682,f33683,f33684,f33685,f33686,f33687,f33688,f33689,f30227,f33690,f30277,f33691,f30324,f33692,f30370,f33693,f30417,f33694,f30462,f33695,f30506,f33696,f30549,f33697,f30594,f33698,f33699,f33700,f33701,f33702,f33703,f33704,f33705,f33706,f33707,f31075,f33708,f31125,f33709,f31172,f33710,f31218,f33711,f31266,f33712,f31311,f33713,f31355,f33714,f31398,f33715,f31443,f33766,f33774,f33784,f680,f33785,f33966,f33243,f679,f34071])).
fof(f34071,plain,(
  $false | ($spl1 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f34070,f534])).
fof(f534,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f8])).
fof(f8,axiom,(
  p(s(bool,'T_0'))),
  file('thm.list.MEM_GENLIST','thm.bool.TRUTH')).
fof(f34070,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl4 | $spl46)),
  inference(forward_demodulation,[],[f34067,f2535])).
fof(f34067,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | ($spl1 | $spl4)),
  inference(trivial_inequality_removal,[],[f34066])).
fof(f34066,plain,(
  s(sK12,sK13) != s(sK12,sK13) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | ($spl1 | $spl4)),
  inference(superposition,[],[f29626,f548])).
fof(f29626,plain,(
  ( ! [X0] : (s(sK12,sK13) != s(sK12,happ(s(fun('type.num.num',sK12),sK15),s('type.num.num',X0))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',sK14))))) ) | $spl1),
  inference(duplicate_literal_removal,[],[f29624])).
fof(f29624,plain,(
  ( ! [X0] : (s(sK12,sK13) != s(sK12,happ(s(fun('type.num.num',sK12),sK15),s('type.num.num',X0))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',sK14)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',sK14))))) ) | $spl1),
  inference(superposition,[],[f14310,f360])).
fof(f360,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun('type.num.num',X0),X1),s('type.num.num',X3))) = s(X0,'const.list.EL_2'(s('type.num.num',X3),s('type.list.list'(X0),'const.list.GENLIST_2'(s(fun('type.num.num',X0),X1),s('type.num.num',X2))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num',X2))))) )),
  inference(cnf_transformation,[],[f150])).
fof(f150,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num',X2)))) | s(X0,happ(s(fun('type.num.num',X0),X1),s('type.num.num',X3))) = s(X0,'const.list.EL_2'(s('type.num.num',X3),s('type.list.list'(X0),'const.list.GENLIST_2'(s(fun('type.num.num',X0),X1),s('type.num.num',X2))))))),
  inference(ennf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num',X2)))) => s(X0,happ(s(fun('type.num.num',X0),X1),s('type.num.num',X3))) = s(X0,'const.list.EL_2'(s('type.num.num',X3),s('type.list.list'(X0),'const.list.GENLIST_2'(s(fun('type.num.num',X0),X1),s('type.num.num',X2))))))),
  inference(flattening,[],[f80])).
fof(f80,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num',X2)))) => s(X0,happ(s(fun('type.num.num',X0),X1),s('type.num.num',X3))) = s(X0,'const.list.EL_2'(s('type.num.num',X3),s('type.list.list'(X0),'const.list.GENLIST_2'(s(fun('type.num.num',X0),X1),s('type.num.num',X2))))))),
  inference(rectify,[],[f50])).
fof(f50,axiom,(
  ! [X6] : ! [X9,X8,X7] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num',X8)))) => s(X6,happ(s(fun('type.num.num',X6),X9),s('type.num.num',X7))) = s(X6,'const.list.EL_2'(s('type.num.num',X7),s('type.list.list'(X6),'const.list.GENLIST_2'(s(fun('type.num.num',X6),X9),s('type.num.num',X8))))))),
  file('thm.list.MEM_GENLIST','thm.list.EL_GENLIST')).
fof(f14310,plain,(
  ( ! [X0] : (s(sK12,sK13) != s(sK12,'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',sK14))))) ) | $spl1),
  inference(forward_demodulation,[],[f14184,f535])).
fof(f535,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num',X2) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.GENLIST_2'(s(fun('type.num.num',X0),X1),s('type.num.num',X2)))))) )),
  inference(cnf_transformation,[],[f143])).
fof(f143,plain,(
  ! [X0,X1,X2] : s('type.num.num',X2) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.GENLIST_2'(s(fun('type.num.num',X0),X1),s('type.num.num',X2)))))),
  inference(flattening,[],[f142])).
fof(f142,plain,(
  ! [X0] : ! [X1,X2] : s('type.num.num',X2) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.GENLIST_2'(s(fun('type.num.num',X0),X1),s('type.num.num',X2)))))),
  inference(rectify,[],[f49])).
fof(f49,axiom,(
  ! [X6] : ! [X9,X8] : s('type.num.num',X8) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),'const.list.GENLIST_2'(s(fun('type.num.num',X6),X9),s('type.num.num',X8)))))),
  file('thm.list.MEM_GENLIST','thm.list.LENGTH_GENLIST')).
fof(f14184,plain,(
  ( ! [X0] : (s(sK12,sK13) != s(sK12,'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))))) ) | $spl1),
  inference(resolution,[],[f545,f366])).
fof(f366,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1)))))) | s(X0,X2) != s(X0,'const.list.EL_2'(s('type.num.num',X3),s('type.list.list'(X0),X1))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))))))) )),
  inference(cnf_transformation,[],[f222])).
fof(f222,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1)))))) | (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK19(X2,X1,X0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1)))))) & s(X0,X2) = s(X0,'const.list.EL_2'(s('type.num.num',sK19(X2,X1,X0)),s('type.list.list'(X0),X1))))) & (! [X3] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1)))))) | s(X0,X2) != s(X0,'const.list.EL_2'(s('type.num.num',X3),s('type.list.list'(X0),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK19])],[f221])).
fof(f221,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1)))))) | ? [X4] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1)))))) & s(X0,X2) = s(X0,'const.list.EL_2'(s('type.num.num',X4),s('type.list.list'(X0),X1))))) & (! [X3] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1)))))) | s(X0,X2) != s(X0,'const.list.EL_2'(s('type.num.num',X3),s('type.list.list'(X0),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1))))))))),
  inference(rectify,[],[f220])).
fof(f220,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1)))))) | ? [X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1)))))) & s(X0,X2) = s(X0,'const.list.EL_2'(s('type.num.num',X3),s('type.list.list'(X0),X1))))) & (! [X3] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1)))))) | s(X0,X2) != s(X0,'const.list.EL_2'(s('type.num.num',X3),s('type.list.list'(X0),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1))))))))),
  inference(nnf_transformation,[],[f87])).
fof(f87,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1)))))) <=> ? [X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1)))))) & s(X0,X2) = s(X0,'const.list.EL_2'(s('type.num.num',X3),s('type.list.list'(X0),X1)))))),
  inference(flattening,[],[f86])).
fof(f86,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1)))))) <=> ? [X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1)))))) & s(X0,X2) = s(X0,'const.list.EL_2'(s('type.num.num',X3),s('type.list.list'(X0),X1)))))),
  inference(rectify,[],[f48])).
fof(f48,axiom,(
  ! [X6] : ! [X21,X7] : (p(s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X6),X21)))))) <=> ? [X8] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X8),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X21)))))) & s(X6,X7) = s(X6,'const.list.EL_2'(s('type.num.num',X8),s('type.list.list'(X6),X21)))))),
  file('thm.list.MEM_GENLIST','thm.list.MEM_EL')).
fof(f679,plain,(
  ( ! [X15] : (~p(s(bool,X15))) ) | $spl14),
  inference(cnf_transformation,[],[f679_D])).
fof(f679_D,plain,(
  ( ! [X15] : (~p(s(bool,X15))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f33243,plain,(
  ( ! [X4,X3] : (p(s(bool,X3)) | s(bool,X3) != s(bool,X4) | s(bool,'T_0') != s(bool,X4)) ) | $spl470),
  inference(cnf_transformation,[],[f33243_D])).
fof(f33243_D,plain,(
  ( ! [X4,X3] : (p(s(bool,X3)) | s(bool,X3) != s(bool,X4) | s(bool,'T_0') != s(bool,X4)) ) <=> ~$spl470),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl470])])).
fof(f33966,plain,(
  ( ! [X28,X29,X27] : (s(bool,'T_0') != s(bool,X29) | p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,X29)) ) | ($spl8 | $spl46)),
  inference(forward_demodulation,[],[f33959,f2535])).
fof(f33959,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,X29) | s(bool,X29) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | $spl8),
  inference(duplicate_literal_removal,[],[f33939])).
fof(f33939,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,X29) | s(bool,X29) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | ~p(s(bool,X28))) ) | $spl8),
  inference(resolution,[],[f11868,f333])).
fof(f333,plain,(
  ( ! [X2,X0,X1] : (sP3(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f208])).
fof(f208,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP3(X2,X1,X0)) & (~sP3(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f178])).
fof(f178,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP3(X2,X1,X0))),
  inference(definition_folding,[],[f66,f177,f176,f175,f174])).
fof(f174,plain,(
  ! [X0,X1,X2] : (sP0(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f175,plain,(
  ! [X1,X0,X2] : (sP1(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f176,plain,(
  ! [X2,X0,X1] : (sP2(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f177,plain,(
  ! [X2,X1,X0] : (sP3(X2,X1,X0) <=> (sP0(X0,X1,X2) & sP1(X1,X0,X2) & sP2(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f66,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f65])).
fof(f65,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f56])).
fof(f56,axiom,(
  ! [X22,X23,X24] : ((p(s(bool,X24)) <=> s(bool,X22) = s(bool,X23)) <=> ((p(s(bool,X24)) | p(s(bool,X23)) | p(s(bool,X22))) & (p(s(bool,X24)) | ~p(s(bool,X22)) | ~p(s(bool,X23))) & (p(s(bool,X23)) | ~p(s(bool,X22)) | ~p(s(bool,X24))) & (p(s(bool,X22)) | ~p(s(bool,X23)) | ~p(s(bool,X24)))))),
  file('thm.list.MEM_GENLIST','thm.sat.sat_2|49|_prove_ltb')).
fof(f11868,plain,(
  ( ! [X10,X11,X9] : (~sP3(X11,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X9) | p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X9) != s(bool,X10)) ) | $spl8),
  inference(superposition,[],[f676,f334])).
fof(f334,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X2)) | ~sP3(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f208])).
fof(f33785,plain,(
  ( ! [X56] : (~p(s(bool,X56)) | ~sP2(X56,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) ) | $spl486),
  inference(cnf_transformation,[],[f33785_D])).
fof(f33785_D,plain,(
  ( ! [X56] : (~p(s(bool,X56)) | ~sP2(X56,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) ) <=> ~$spl486),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl486])])).
fof(f680,plain,(
  ( ! [X14] : (p(s(bool,X14))) ) | $spl16),
  inference(cnf_transformation,[],[f680_D])).
fof(f680_D,plain,(
  ( ! [X14] : (p(s(bool,X14))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f33784,plain,(
  ( ! [X56,X55] : (p(s(bool,X55)) | ~p(s(bool,X56)) | ~sP2(X56,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) ) | ($spl1 | $spl8)),
  inference(subsumption_resolution,[],[f33783,f351])).
fof(f351,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f213])).
fof(f213,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f73])).
fof(f73,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X11] : (s(bool,'F_0') = s(bool,X11) <=> ~p(s(bool,X11)))),
  file('thm.list.MEM_GENLIST','thm.bool.EQ_CLAUSES|split|2')).
fof(f33783,plain,(
  ( ! [X56,X55] : (s(bool,'F_0') != s(bool,X55) | p(s(bool,X55)) | ~p(s(bool,X56)) | ~sP2(X56,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) ) | ($spl1 | $spl8)),
  inference(forward_demodulation,[],[f33782,f14213])).
fof(f14213,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | $spl1),
  inference(resolution,[],[f545,f351])).
fof(f33782,plain,(
  ( ! [X56,X55] : (p(s(bool,X55)) | ~p(s(bool,X56)) | s(bool,X55) != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | ~sP2(X56,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) ) | ($spl1 | $spl8)),
  inference(subsumption_resolution,[],[f33746,f329])).
fof(f329,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f207])).
fof(f207,plain,(
  ! [X0,X1,X2] : ((~sP0(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP0(X0,X1,X2)))),
  inference(flattening,[],[f206])).
fof(f206,plain,(
  ! [X0,X1,X2] : ((~sP0(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP0(X0,X1,X2)))),
  inference(nnf_transformation,[],[f174])).
fof(f33746,plain,(
  ( ! [X56,X55] : (p(s(bool,X55)) | ~p(s(bool,X56)) | s(bool,X55) != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | ~sP2(X56,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | ~sP0('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X56)) ) | ($spl1 | $spl8)),
  inference(resolution,[],[f11867,f30738])).
fof(f30738,plain,(
  ( ! [X28,X27] : (sP3(X27,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X28) | ~sP2(X27,X28,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | ~sP0(X28,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X27)) ) | $spl1),
  inference(subsumption_resolution,[],[f14194,f727])).
fof(f727,plain,(
  ( ! [X41,X42] : (sP1('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X41,X42)) ) | $spl1),
  inference(resolution,[],[f545,f327])).
fof(f327,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f205])).
fof(f205,plain,(
  ! [X0,X1,X2] : ((~sP1(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1,X2)))),
  inference(rectify,[],[f204])).
fof(f204,plain,(
  ! [X1,X0,X2] : ((~sP1(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0,X2)))),
  inference(flattening,[],[f203])).
fof(f203,plain,(
  ! [X1,X0,X2] : ((~sP1(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0,X2)))),
  inference(nnf_transformation,[],[f175])).
fof(f14194,plain,(
  ( ! [X28,X27] : (sP3(X27,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X28) | ~sP2(X27,X28,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | ~sP1('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X28,X27) | ~sP0(X28,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X27)) ) | $spl1),
  inference(resolution,[],[f545,f318])).
fof(f318,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X1)) | ~sP2(X0,X2,X1) | ~sP1(X1,X2,X0) | ~sP0(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f199])).
fof(f199,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | (sP0(X2,X1,X0) & sP1(X1,X2,X0) & sP2(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP0(X2,X1,X0) | ~sP1(X1,X2,X0) | ~sP2(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f198])).
fof(f198,plain,(
  ! [X2,X1,X0] : ((~sP3(X2,X1,X0) | (sP0(X0,X1,X2) & sP1(X1,X0,X2) & sP2(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP0(X0,X1,X2) | ~sP1(X1,X0,X2) | ~sP2(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP3(X2,X1,X0)))),
  inference(flattening,[],[f197])).
fof(f197,plain,(
  ! [X2,X1,X0] : ((~sP3(X2,X1,X0) | (sP0(X0,X1,X2) & sP1(X1,X0,X2) & sP2(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP0(X0,X1,X2) | ~sP1(X1,X0,X2) | ~sP2(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP3(X2,X1,X0)))),
  inference(nnf_transformation,[],[f177])).
fof(f11867,plain,(
  ( ! [X6,X8,X7] : (~sP3(X8,X6,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X6) != s(bool,X7)) ) | $spl8),
  inference(superposition,[],[f676,f334])).
fof(f33774,plain,(
  ( ! [X30,X31,X32] : (s(bool,'T_0') != s(bool,X32) | p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X30) != s(bool,X32)) ) | ($spl8 | $spl46)),
  inference(forward_demodulation,[],[f33765,f2535])).
fof(f33765,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X30) != s(bool,X32) | s(bool,X32) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | $spl8),
  inference(duplicate_literal_removal,[],[f33736])).
fof(f33736,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X30) != s(bool,X32) | s(bool,X32) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | ~p(s(bool,X31))) ) | $spl8),
  inference(resolution,[],[f11867,f333])).
fof(f33766,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,'T_0') != s(bool,X8)) ) | ($spl8 | $spl46)),
  inference(duplicate_literal_removal,[],[f33728])).
fof(f33728,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,'T_0') != s(bool,X8) | ~p(s(bool,X7))) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f11867,f12162])).
fof(f12162,plain,(
  ( ! [X158,X159] : (sP3(X159,X158,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | s(bool,'T_0') != s(bool,X158) | ~p(s(bool,X159))) ) | $spl46),
  inference(superposition,[],[f333,f2535])).
fof(f31443,plain,(
  ( ! [X2,X0,X3,X1] : (sP10(X0,X1) | p(s(bool,X0)) | sP8(X1,X2,X3)) ) | $spl458),
  inference(cnf_transformation,[],[f31443_D])).
fof(f31443_D,plain,(
  ( ! [X2,X0,X3,X1] : (sP10(X0,X1) | p(s(bool,X0)) | sP8(X1,X2,X3)) ) <=> ~$spl458),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl458])])).
fof(f33715,plain,(
  ( ! [X239,X237,X235,X238,X236] : (p(s(bool,X235)) | p(s(bool,X236)) | sP10(X235,X237) | sP8(X237,X238,X239)) ) | ($spl84 | $spl122)),
  inference(resolution,[],[f11713,f28401])).
fof(f28401,plain,(
  ( ! [X341,X343,X339,X340,X342] : (sP11(X339,X340) | sP10(X339,X341) | sP8(X341,X342,X343)) ) | $spl122),
  inference(resolution,[],[f23560,f411])).
fof(f411,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f243])).
fof(f243,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(flattening,[],[f242])).
fof(f242,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(nnf_transformation,[],[f185])).
fof(f185,plain,(
  ! [X0,X1,X2] : (sP8(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f23560,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,X6)) | sP11(X5,X7) | sP10(X5,X6)) ) | $spl122),
  inference(subsumption_resolution,[],[f23521,f537])).
fof(f537,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f348])).
fof(f348,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f212])).
fof(f212,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f72])).
fof(f72,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X6] : ! [X7] : (s(X6,X7) = s(X6,X7) <=> p(s(bool,'T_0')))),
  file('thm.list.MEM_GENLIST','thm.bool.REFL_CLAUSE')).
fof(f23521,plain,(
  ( ! [X6,X7,X5] : (sP10(X5,X6) | sP11(X5,X7) | ~p(s(bool,'T_0')) | ~p(s(bool,X6))) ) | $spl122),
  inference(resolution,[],[f23299,f439])).
fof(f439,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | ~sP10(X0,X1)) )),
  inference(cnf_transformation,[],[f256])).
fof(f256,plain,(
  ! [X0,X1] : ((~sP10(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP10(X0,X1)))),
  inference(rectify,[],[f255])).
fof(f255,plain,(
  ! [X1,X0] : ((~sP10(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP10(X1,X0)))),
  inference(flattening,[],[f254])).
fof(f254,plain,(
  ! [X1,X0] : ((~sP10(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP10(X1,X0)))),
  inference(nnf_transformation,[],[f188])).
fof(f188,plain,(
  ! [X1,X0] : (sP10(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f23299,plain,(
  ( ! [X329,X331,X330] : (sP10('T_0',X329) | sP10(X330,X329) | sP11(X330,X331)) ) | $spl122),
  inference(resolution,[],[f19619,f456])).
fof(f456,plain,(
  ( ! [X0,X1] : (sP11(X0,X1) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f261])).
fof(f261,plain,(
  ! [X0,X1] : ((~sP11(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP11(X0,X1)))),
  inference(flattening,[],[f260])).
fof(f260,plain,(
  ! [X0,X1] : ((~sP11(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP11(X0,X1)))),
  inference(nnf_transformation,[],[f190])).
fof(f190,plain,(
  ! [X0,X1] : (sP11(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f19619,plain,(
  ( ! [X231,X232] : (sP10('T_0',X231) | sP10(X232,X231) | p(s(bool,X232))) ) | $spl122),
  inference(resolution,[],[f19515,f443])).
fof(f443,plain,(
  ( ! [X0,X1] : (sP10(X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f256])).
fof(f19515,plain,(
  ( ! [X0] : (p(s(bool,X0)) | sP10('T_0',X0)) ) | $spl122),
  inference(subsumption_resolution,[],[f19501,f492])).
fof(f492,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f283])).
fof(f283,plain,(
  ! [X0] : (((~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | p(s(bool,X0))) & (~p(s(bool,X0)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f282])).
fof(f282,plain,(
  ! [X0] : (((~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0] : ((p(s(bool,X0)) | p(s(bool,'F_0'))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X11] : ((p(s(bool,X11)) | p(s(bool,'F_0'))) <=> p(s(bool,X11)))),
  file('thm.list.MEM_GENLIST','thm.bool.OR_CLAUSES|split|3')).
fof(f19501,plain,(
  ( ! [X0] : (sP10('T_0',X0) | p(s(bool,'F_0')) | p(s(bool,X0))) ) | $spl122),
  inference(resolution,[],[f19417,f438])).
fof(f438,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | ~sP10(X0,X1)) )),
  inference(cnf_transformation,[],[f256])).
fof(f19417,plain,(
  ( ! [X399] : (sP10('T_0',X399) | sP10(X399,'F_0')) ) | $spl122),
  inference(resolution,[],[f18883,f533])).
fof(f533,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.list.MEM_GENLIST',aHL_TRUTH)).
fof(f18883,plain,(
  ( ! [X229,X230] : (~p(s(bool,X230)) | sP10(X230,X229) | sP10(X229,'F_0')) ) | $spl122),
  inference(resolution,[],[f18767,f440])).
fof(f440,plain,(
  ( ! [X0,X1] : (sP10(X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f256])).
fof(f31398,plain,(
  ( ! [X2,X0,X3,X1] : (sP10(X0,X1) | p(s(bool,X0)) | sP8(X2,X1,X3)) ) | $spl452),
  inference(cnf_transformation,[],[f31398_D])).
fof(f31398_D,plain,(
  ( ! [X2,X0,X3,X1] : (sP10(X0,X1) | p(s(bool,X0)) | sP8(X2,X1,X3)) ) <=> ~$spl452),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl452])])).
fof(f33714,plain,(
  ( ! [X231,X233,X230,X234,X232] : (p(s(bool,X230)) | p(s(bool,X231)) | sP10(X230,X232) | sP8(X233,X232,X234)) ) | ($spl84 | $spl122)),
  inference(resolution,[],[f11713,f28400])).
fof(f28400,plain,(
  ( ! [X335,X337,X334,X336,X338] : (sP11(X334,X335) | sP10(X334,X336) | sP8(X337,X336,X338)) ) | $spl122),
  inference(resolution,[],[f23560,f410])).
fof(f410,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f243])).
fof(f31355,plain,(
  ( ! [X2,X0,X3,X1] : (sP10(X0,X1) | p(s(bool,X0)) | sP6(X1,X2,X3)) ) | $spl446),
  inference(cnf_transformation,[],[f31355_D])).
fof(f31355_D,plain,(
  ( ! [X2,X0,X3,X1] : (sP10(X0,X1) | p(s(bool,X0)) | sP6(X1,X2,X3)) ) <=> ~$spl446),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl446])])).
fof(f33713,plain,(
  ( ! [X229,X227,X225,X228,X226] : (p(s(bool,X225)) | p(s(bool,X226)) | sP10(X225,X227) | sP6(X227,X228,X229)) ) | ($spl84 | $spl122)),
  inference(resolution,[],[f11713,f28387])).
fof(f28387,plain,(
  ( ! [X269,X271,X273,X270,X272] : (sP11(X269,X270) | sP10(X269,X271) | sP6(X271,X272,X273)) ) | $spl122),
  inference(resolution,[],[f23560,f394])).
fof(f394,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f236])).
fof(f236,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(rectify,[],[f235])).
fof(f235,plain,(
  ! [X2,X0,X1] : ((~sP6(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP6(X2,X0,X1)))),
  inference(flattening,[],[f234])).
fof(f234,plain,(
  ! [X2,X0,X1] : ((~sP6(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP6(X2,X0,X1)))),
  inference(nnf_transformation,[],[f182])).
fof(f182,plain,(
  ! [X2,X0,X1] : (sP6(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f31311,plain,(
  ( ! [X2,X0,X3,X1] : (sP10(X0,X1) | p(s(bool,X0)) | sP4(X1,X2,X3)) ) | $spl440),
  inference(cnf_transformation,[],[f31311_D])).
fof(f31311_D,plain,(
  ( ! [X2,X0,X3,X1] : (sP10(X0,X1) | p(s(bool,X0)) | sP4(X1,X2,X3)) ) <=> ~$spl440),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl440])])).
fof(f33712,plain,(
  ( ! [X222,X220,X223,X221,X224] : (p(s(bool,X220)) | p(s(bool,X221)) | sP10(X220,X222) | sP4(X222,X223,X224)) ) | ($spl84 | $spl122)),
  inference(resolution,[],[f11713,f28371])).
fof(f28371,plain,(
  ( ! [X191,X189,X192,X190,X193] : (sP11(X189,X190) | sP10(X189,X191) | sP4(X191,X192,X193)) ) | $spl122),
  inference(resolution,[],[f23560,f377])).
fof(f377,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f228])).
fof(f228,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f227])).
fof(f227,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(flattening,[],[f226])).
fof(f226,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(nnf_transformation,[],[f179])).
fof(f179,plain,(
  ! [X2,X0,X1] : (sP4(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f31266,plain,(
  ( ! [X2,X0,X3,X1] : (sP10(X0,X1) | p(s(bool,X0)) | sP4(X2,X3,X1)) ) | $spl434),
  inference(cnf_transformation,[],[f31266_D])).
fof(f31266_D,plain,(
  ( ! [X2,X0,X3,X1] : (sP10(X0,X1) | p(s(bool,X0)) | sP4(X2,X3,X1)) ) <=> ~$spl434),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl434])])).
fof(f33711,plain,(
  ( ! [X218,X216,X215,X219,X217] : (p(s(bool,X215)) | p(s(bool,X216)) | sP10(X215,X217) | sP4(X218,X219,X217)) ) | ($spl84 | $spl122)),
  inference(resolution,[],[f11713,f28370])).
fof(f28370,plain,(
  ( ! [X187,X185,X188,X186,X184] : (sP11(X184,X185) | sP10(X184,X186) | sP4(X187,X188,X186)) ) | $spl122),
  inference(resolution,[],[f23560,f375])).
fof(f375,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f228])).
fof(f31218,plain,(
  ( ! [X2,X0,X3,X1] : (sP10(X0,X1) | p(s(bool,X0)) | sP1(X1,X2,X3)) ) | $spl428),
  inference(cnf_transformation,[],[f31218_D])).
fof(f31218_D,plain,(
  ( ! [X2,X0,X3,X1] : (sP10(X0,X1) | p(s(bool,X0)) | sP1(X1,X2,X3)) ) <=> ~$spl428),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl428])])).
fof(f33710,plain,(
  ( ! [X212,X210,X213,X211,X214] : (p(s(bool,X210)) | p(s(bool,X211)) | sP10(X210,X212) | sP1(X212,X213,X214)) ) | ($spl84 | $spl122)),
  inference(resolution,[],[f11713,f28348])).
fof(f28348,plain,(
  ( ! [X92,X90,X88,X91,X89] : (sP11(X88,X89) | sP10(X88,X90) | sP1(X90,X91,X92)) ) | $spl122),
  inference(resolution,[],[f23560,f327])).
fof(f31172,plain,(
  ( ! [X2,X0,X3,X1] : (sP10(X0,X1) | p(s(bool,X0)) | sP1(X2,X1,X3)) ) | $spl422),
  inference(cnf_transformation,[],[f31172_D])).
fof(f31172_D,plain,(
  ( ! [X2,X0,X3,X1] : (sP10(X0,X1) | p(s(bool,X0)) | sP1(X2,X1,X3)) ) <=> ~$spl422),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl422])])).
fof(f33709,plain,(
  ( ! [X208,X206,X209,X207,X205] : (p(s(bool,X205)) | p(s(bool,X206)) | sP10(X205,X207) | sP1(X208,X207,X209)) ) | ($spl84 | $spl122)),
  inference(resolution,[],[f11713,f28347])).
fof(f28347,plain,(
  ( ! [X87,X85,X83,X86,X84] : (sP11(X83,X84) | sP10(X83,X85) | sP1(X86,X85,X87)) ) | $spl122),
  inference(resolution,[],[f23560,f326])).
fof(f326,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f205])).
fof(f31125,plain,(
  ( ! [X2,X0,X3,X1] : (sP10(X0,X1) | p(s(bool,X0)) | sP2(X1,X2,X3)) ) | $spl416),
  inference(cnf_transformation,[],[f31125_D])).
fof(f31125_D,plain,(
  ( ! [X2,X0,X3,X1] : (sP10(X0,X1) | p(s(bool,X0)) | sP2(X1,X2,X3)) ) <=> ~$spl416),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl416])])).
fof(f33708,plain,(
  ( ! [X204,X202,X200,X203,X201] : (p(s(bool,X200)) | p(s(bool,X201)) | sP10(X200,X202) | sP2(X202,X203,X204)) ) | ($spl84 | $spl122)),
  inference(resolution,[],[f11713,f28345])).
fof(f28345,plain,(
  ( ! [X76,X74,X77,X75,X73] : (sP11(X73,X74) | sP10(X73,X75) | sP2(X75,X76,X77)) ) | $spl122),
  inference(resolution,[],[f23560,f323])).
fof(f323,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f202])).
fof(f202,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(rectify,[],[f201])).
fof(f201,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(flattening,[],[f200])).
fof(f200,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(nnf_transformation,[],[f176])).
fof(f31075,plain,(
  ( ! [X2,X0,X3,X1] : (sP10(X0,X1) | p(s(bool,X0)) | sP2(X2,X1,X3)) ) | $spl408),
  inference(cnf_transformation,[],[f31075_D])).
fof(f31075_D,plain,(
  ( ! [X2,X0,X3,X1] : (sP10(X0,X1) | p(s(bool,X0)) | sP2(X2,X1,X3)) ) <=> ~$spl408),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl408])])).
fof(f33707,plain,(
  ( ! [X198,X196,X199,X197,X195] : (p(s(bool,X195)) | p(s(bool,X196)) | sP10(X195,X197) | sP2(X198,X197,X199)) ) | ($spl84 | $spl122)),
  inference(resolution,[],[f11713,f28344])).
fof(f28344,plain,(
  ( ! [X70,X68,X72,X71,X69] : (sP11(X68,X69) | sP10(X68,X70) | sP2(X71,X70,X72)) ) | $spl122),
  inference(resolution,[],[f23560,f322])).
fof(f322,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f202])).
fof(f33706,plain,(
  ( ! [X191,X194,X192,X190,X193] : (p(s(bool,X190)) | p(s(bool,X191)) | sP10(X191,X192) | sP8(X192,X193,X194)) ) | ($spl84 | $spl122)),
  inference(resolution,[],[f11713,f28167])).
fof(f28167,plain,(
  ( ! [X341,X343,X339,X340,X342] : (sP11(X339,X340) | sP10(X340,X341) | sP8(X341,X342,X343)) ) | $spl122),
  inference(resolution,[],[f23454,f411])).
fof(f23454,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,X6)) | sP11(X7,X5) | sP10(X5,X6)) ) | $spl122),
  inference(subsumption_resolution,[],[f23415,f537])).
fof(f23415,plain,(
  ( ! [X6,X7,X5] : (sP10(X5,X6) | sP11(X7,X5) | ~p(s(bool,'T_0')) | ~p(s(bool,X6))) ) | $spl122),
  inference(resolution,[],[f23298,f439])).
fof(f23298,plain,(
  ( ! [X327,X326,X328] : (sP10('T_0',X326) | sP10(X327,X326) | sP11(X328,X327)) ) | $spl122),
  inference(resolution,[],[f19619,f455])).
fof(f455,plain,(
  ( ! [X0,X1] : (sP11(X0,X1) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f261])).
fof(f33705,plain,(
  ( ! [X189,X187,X185,X188,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | sP10(X186,X187) | sP8(X188,X187,X189)) ) | ($spl84 | $spl122)),
  inference(resolution,[],[f11713,f28166])).
fof(f28166,plain,(
  ( ! [X335,X337,X334,X336,X338] : (sP11(X334,X335) | sP10(X335,X336) | sP8(X337,X336,X338)) ) | $spl122),
  inference(resolution,[],[f23454,f410])).
fof(f33704,plain,(
  ( ! [X182,X180,X184,X183,X181] : (p(s(bool,X180)) | p(s(bool,X181)) | sP10(X181,X182) | sP6(X182,X183,X184)) ) | ($spl84 | $spl122)),
  inference(resolution,[],[f11713,f28153])).
fof(f28153,plain,(
  ( ! [X269,X271,X273,X270,X272] : (sP11(X269,X270) | sP10(X270,X271) | sP6(X271,X272,X273)) ) | $spl122),
  inference(resolution,[],[f23454,f394])).
fof(f33703,plain,(
  ( ! [X177,X175,X178,X176,X179] : (p(s(bool,X175)) | p(s(bool,X176)) | sP10(X176,X177) | sP4(X177,X178,X179)) ) | ($spl84 | $spl122)),
  inference(resolution,[],[f11713,f28137])).
fof(f28137,plain,(
  ( ! [X191,X189,X192,X190,X193] : (sP11(X189,X190) | sP10(X190,X191) | sP4(X191,X192,X193)) ) | $spl122),
  inference(resolution,[],[f23454,f377])).
fof(f33702,plain,(
  ( ! [X173,X171,X174,X172,X170] : (p(s(bool,X170)) | p(s(bool,X171)) | sP10(X171,X172) | sP4(X173,X174,X172)) ) | ($spl84 | $spl122)),
  inference(resolution,[],[f11713,f28136])).
fof(f28136,plain,(
  ( ! [X187,X185,X188,X186,X184] : (sP11(X184,X185) | sP10(X185,X186) | sP4(X187,X188,X186)) ) | $spl122),
  inference(resolution,[],[f23454,f375])).
fof(f33701,plain,(
  ( ! [X167,X165,X169,X166,X168] : (p(s(bool,X165)) | p(s(bool,X166)) | sP10(X166,X167) | sP1(X167,X168,X169)) ) | ($spl84 | $spl122)),
  inference(resolution,[],[f11713,f28114])).
fof(f28114,plain,(
  ( ! [X92,X90,X88,X91,X89] : (sP11(X88,X89) | sP10(X89,X90) | sP1(X90,X91,X92)) ) | $spl122),
  inference(resolution,[],[f23454,f327])).
fof(f33700,plain,(
  ( ! [X163,X161,X164,X162,X160] : (p(s(bool,X160)) | p(s(bool,X161)) | sP10(X161,X162) | sP1(X163,X162,X164)) ) | ($spl84 | $spl122)),
  inference(resolution,[],[f11713,f28113])).
fof(f28113,plain,(
  ( ! [X87,X85,X83,X86,X84] : (sP11(X83,X84) | sP10(X84,X85) | sP1(X86,X85,X87)) ) | $spl122),
  inference(resolution,[],[f23454,f326])).
fof(f33699,plain,(
  ( ! [X158,X156,X159,X157,X155] : (p(s(bool,X155)) | p(s(bool,X156)) | sP10(X156,X157) | sP2(X157,X158,X159)) ) | ($spl84 | $spl122)),
  inference(resolution,[],[f11713,f28111])).
fof(f28111,plain,(
  ( ! [X76,X74,X77,X75,X73] : (sP11(X73,X74) | sP10(X74,X75) | sP2(X75,X76,X77)) ) | $spl122),
  inference(resolution,[],[f23454,f323])).
fof(f33698,plain,(
  ( ! [X154,X152,X151,X153,X150] : (p(s(bool,X150)) | p(s(bool,X151)) | sP10(X151,X152) | sP2(X153,X152,X154)) ) | ($spl84 | $spl122)),
  inference(resolution,[],[f11713,f28110])).
fof(f28110,plain,(
  ( ! [X70,X68,X72,X71,X69] : (sP11(X68,X69) | sP10(X69,X70) | sP2(X71,X70,X72)) ) | $spl122),
  inference(resolution,[],[f23454,f322])).
fof(f30594,plain,(
  ( ! [X2,X0,X3,X1] : (sP10(X0,X1) | p(s(bool,X1)) | sP8(X0,X2,X3)) ) | $spl400),
  inference(cnf_transformation,[],[f30594_D])).
fof(f30594_D,plain,(
  ( ! [X2,X0,X3,X1] : (sP10(X0,X1) | p(s(bool,X1)) | sP8(X0,X2,X3)) ) <=> ~$spl400),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl400])])).
fof(f33697,plain,(
  ( ! [X146,X149,X147,X145,X148] : (p(s(bool,X145)) | p(s(bool,X146)) | sP10(X147,X145) | sP8(X147,X148,X149)) ) | ($spl84 | $spl122)),
  inference(resolution,[],[f11713,f27959])).
fof(f27959,plain,(
  ( ! [X341,X343,X339,X340,X342] : (sP11(X339,X340) | sP10(X341,X339) | sP8(X341,X342,X343)) ) | $spl122),
  inference(resolution,[],[f22910,f411])).
fof(f22910,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,X5)) | sP11(X6,X7) | sP10(X5,X6)) ) | $spl122),
  inference(subsumption_resolution,[],[f22871,f537])).
fof(f22871,plain,(
  ( ! [X6,X7,X5] : (sP10(X5,X6) | sP11(X6,X7) | ~p(s(bool,'T_0')) | ~p(s(bool,X5))) ) | $spl122),
  inference(resolution,[],[f22620,f439])).
fof(f22620,plain,(
  ( ! [X329,X331,X330] : (sP10('T_0',X329) | sP10(X329,X330) | sP11(X330,X331)) ) | $spl122),
  inference(resolution,[],[f19618,f456])).
fof(f19618,plain,(
  ( ! [X229,X230] : (sP10('T_0',X229) | sP10(X229,X230) | p(s(bool,X230))) ) | $spl122),
  inference(resolution,[],[f19515,f440])).
fof(f30549,plain,(
  ( ! [X2,X0,X3,X1] : (sP10(X0,X1) | p(s(bool,X1)) | sP8(X2,X0,X3)) ) | $spl394),
  inference(cnf_transformation,[],[f30549_D])).
fof(f30549_D,plain,(
  ( ! [X2,X0,X3,X1] : (sP10(X0,X1) | p(s(bool,X1)) | sP8(X2,X0,X3)) ) <=> ~$spl394),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl394])])).
fof(f33696,plain,(
  ( ! [X144,X142,X140,X143,X141] : (p(s(bool,X140)) | p(s(bool,X141)) | sP10(X142,X140) | sP8(X143,X142,X144)) ) | ($spl84 | $spl122)),
  inference(resolution,[],[f11713,f27958])).
fof(f27958,plain,(
  ( ! [X335,X337,X334,X336,X338] : (sP11(X334,X335) | sP10(X336,X334) | sP8(X337,X336,X338)) ) | $spl122),
  inference(resolution,[],[f22910,f410])).
fof(f30506,plain,(
  ( ! [X2,X0,X3,X1] : (sP10(X0,X1) | p(s(bool,X1)) | sP6(X0,X2,X3)) ) | $spl388),
  inference(cnf_transformation,[],[f30506_D])).
fof(f30506_D,plain,(
  ( ! [X2,X0,X3,X1] : (sP10(X0,X1) | p(s(bool,X1)) | sP6(X0,X2,X3)) ) <=> ~$spl388),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl388])])).
fof(f33695,plain,(
  ( ! [X138,X136,X135,X139,X137] : (p(s(bool,X135)) | p(s(bool,X136)) | sP10(X137,X135) | sP6(X137,X138,X139)) ) | ($spl84 | $spl122)),
  inference(resolution,[],[f11713,f27945])).
fof(f27945,plain,(
  ( ! [X269,X271,X273,X270,X272] : (sP11(X269,X270) | sP10(X271,X269) | sP6(X271,X272,X273)) ) | $spl122),
  inference(resolution,[],[f22910,f394])).
fof(f30462,plain,(
  ( ! [X2,X0,X3,X1] : (sP10(X0,X1) | p(s(bool,X1)) | sP4(X0,X2,X3)) ) | $spl382),
  inference(cnf_transformation,[],[f30462_D])).
fof(f30462_D,plain,(
  ( ! [X2,X0,X3,X1] : (sP10(X0,X1) | p(s(bool,X1)) | sP4(X0,X2,X3)) ) <=> ~$spl382),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl382])])).
fof(f33694,plain,(
  ( ! [X134,X132,X130,X133,X131] : (p(s(bool,X130)) | p(s(bool,X131)) | sP10(X132,X130) | sP4(X132,X133,X134)) ) | ($spl84 | $spl122)),
  inference(resolution,[],[f11713,f27929])).
fof(f27929,plain,(
  ( ! [X191,X189,X192,X190,X193] : (sP11(X189,X190) | sP10(X191,X189) | sP4(X191,X192,X193)) ) | $spl122),
  inference(resolution,[],[f22910,f377])).
fof(f30417,plain,(
  ( ! [X2,X0,X3,X1] : (sP10(X0,X1) | p(s(bool,X1)) | sP4(X2,X3,X0)) ) | $spl376),
  inference(cnf_transformation,[],[f30417_D])).
fof(f30417_D,plain,(
  ( ! [X2,X0,X3,X1] : (sP10(X0,X1) | p(s(bool,X1)) | sP4(X2,X3,X0)) ) <=> ~$spl376),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl376])])).
fof(f33693,plain,(
  ( ! [X127,X125,X128,X126,X129] : (p(s(bool,X125)) | p(s(bool,X126)) | sP10(X127,X125) | sP4(X128,X129,X127)) ) | ($spl84 | $spl122)),
  inference(resolution,[],[f11713,f27928])).
fof(f27928,plain,(
  ( ! [X187,X185,X188,X186,X184] : (sP11(X184,X185) | sP10(X186,X184) | sP4(X187,X188,X186)) ) | $spl122),
  inference(resolution,[],[f22910,f375])).
fof(f30370,plain,(
  ( ! [X2,X0,X3,X1] : (sP10(X0,X1) | p(s(bool,X1)) | sP1(X0,X2,X3)) ) | $spl370),
  inference(cnf_transformation,[],[f30370_D])).
fof(f30370_D,plain,(
  ( ! [X2,X0,X3,X1] : (sP10(X0,X1) | p(s(bool,X1)) | sP1(X0,X2,X3)) ) <=> ~$spl370),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl370])])).
fof(f33692,plain,(
  ( ! [X123,X121,X124,X122,X120] : (p(s(bool,X120)) | p(s(bool,X121)) | sP10(X122,X120) | sP1(X122,X123,X124)) ) | ($spl84 | $spl122)),
  inference(resolution,[],[f11713,f27906])).
fof(f27906,plain,(
  ( ! [X92,X90,X88,X91,X89] : (sP11(X88,X89) | sP10(X90,X88) | sP1(X90,X91,X92)) ) | $spl122),
  inference(resolution,[],[f22910,f327])).
fof(f30324,plain,(
  ( ! [X2,X0,X3,X1] : (sP10(X0,X1) | p(s(bool,X1)) | sP1(X2,X0,X3)) ) | $spl364),
  inference(cnf_transformation,[],[f30324_D])).
fof(f30324_D,plain,(
  ( ! [X2,X0,X3,X1] : (sP10(X0,X1) | p(s(bool,X1)) | sP1(X2,X0,X3)) ) <=> ~$spl364),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl364])])).
fof(f33691,plain,(
  ( ! [X118,X116,X119,X117,X115] : (p(s(bool,X115)) | p(s(bool,X116)) | sP10(X117,X115) | sP1(X118,X117,X119)) ) | ($spl84 | $spl122)),
  inference(resolution,[],[f11713,f27905])).
fof(f27905,plain,(
  ( ! [X87,X85,X83,X86,X84] : (sP11(X83,X84) | sP10(X85,X83) | sP1(X86,X85,X87)) ) | $spl122),
  inference(resolution,[],[f22910,f326])).
fof(f30277,plain,(
  ( ! [X2,X0,X3,X1] : (sP10(X0,X1) | p(s(bool,X1)) | sP2(X0,X2,X3)) ) | $spl358),
  inference(cnf_transformation,[],[f30277_D])).
fof(f30277_D,plain,(
  ( ! [X2,X0,X3,X1] : (sP10(X0,X1) | p(s(bool,X1)) | sP2(X0,X2,X3)) ) <=> ~$spl358),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl358])])).
fof(f33690,plain,(
  ( ! [X111,X114,X112,X110,X113] : (p(s(bool,X110)) | p(s(bool,X111)) | sP10(X112,X110) | sP2(X112,X113,X114)) ) | ($spl84 | $spl122)),
  inference(resolution,[],[f11713,f27903])).
fof(f27903,plain,(
  ( ! [X76,X74,X77,X75,X73] : (sP11(X73,X74) | sP10(X75,X73) | sP2(X75,X76,X77)) ) | $spl122),
  inference(resolution,[],[f22910,f323])).
fof(f30227,plain,(
  ( ! [X2,X0,X3,X1] : (sP10(X0,X1) | p(s(bool,X1)) | sP2(X2,X0,X3)) ) | $spl350),
  inference(cnf_transformation,[],[f30227_D])).
fof(f30227_D,plain,(
  ( ! [X2,X0,X3,X1] : (sP10(X0,X1) | p(s(bool,X1)) | sP2(X2,X0,X3)) ) <=> ~$spl350),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl350])])).
fof(f33689,plain,(
  ( ! [X109,X107,X105,X108,X106] : (p(s(bool,X105)) | p(s(bool,X106)) | sP10(X107,X105) | sP2(X108,X107,X109)) ) | ($spl84 | $spl122)),
  inference(resolution,[],[f11713,f27902])).
fof(f27902,plain,(
  ( ! [X70,X68,X72,X71,X69] : (sP11(X68,X69) | sP10(X70,X68) | sP2(X71,X70,X72)) ) | $spl122),
  inference(resolution,[],[f22910,f322])).
fof(f33688,plain,(
  ( ! [X103,X101,X102,X100,X104] : (p(s(bool,X100)) | p(s(bool,X101)) | sP10(X102,X101) | sP8(X102,X103,X104)) ) | ($spl84 | $spl122)),
  inference(resolution,[],[f11713,f27769])).
fof(f27769,plain,(
  ( ! [X341,X343,X339,X340,X342] : (sP11(X339,X340) | sP10(X341,X340) | sP8(X341,X342,X343)) ) | $spl122),
  inference(resolution,[],[f22766,f411])).
fof(f22766,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | sP11(X5,X4) | sP10(X3,X4)) ) | $spl122),
  inference(subsumption_resolution,[],[f22728,f537])).
fof(f22728,plain,(
  ( ! [X4,X5,X3] : (sP10(X3,X4) | sP11(X5,X4) | ~p(s(bool,'T_0')) | ~p(s(bool,X3))) ) | $spl122),
  inference(resolution,[],[f22619,f439])).
fof(f22619,plain,(
  ( ! [X327,X326,X328] : (sP10('T_0',X326) | sP10(X326,X327) | sP11(X328,X327)) ) | $spl122),
  inference(resolution,[],[f19618,f455])).
fof(f33687,plain,(
  ( ! [X99,X97,X95,X98,X96] : (p(s(bool,X95)) | p(s(bool,X96)) | sP10(X97,X96) | sP8(X98,X97,X99)) ) | ($spl84 | $spl122)),
  inference(resolution,[],[f11713,f27768])).
fof(f27768,plain,(
  ( ! [X335,X337,X334,X336,X338] : (sP11(X334,X335) | sP10(X336,X335) | sP8(X337,X336,X338)) ) | $spl122),
  inference(resolution,[],[f22766,f410])).
fof(f33686,plain,(
  ( ! [X94,X92,X90,X93,X91] : (p(s(bool,X90)) | p(s(bool,X91)) | sP10(X92,X91) | sP6(X92,X93,X94)) ) | ($spl84 | $spl122)),
  inference(resolution,[],[f11713,f27755])).
fof(f27755,plain,(
  ( ! [X269,X271,X273,X270,X272] : (sP11(X269,X270) | sP10(X271,X270) | sP6(X271,X272,X273)) ) | $spl122),
  inference(resolution,[],[f22766,f394])).
fof(f33685,plain,(
  ( ! [X88,X87,X85,X89,X86] : (p(s(bool,X85)) | p(s(bool,X86)) | sP10(X87,X86) | sP4(X87,X88,X89)) ) | ($spl84 | $spl122)),
  inference(resolution,[],[f11713,f27739])).
fof(f27739,plain,(
  ( ! [X191,X189,X192,X190,X193] : (sP11(X189,X190) | sP10(X191,X190) | sP4(X191,X192,X193)) ) | $spl122),
  inference(resolution,[],[f22766,f377])).
fof(f33684,plain,(
  ( ! [X80,X83,X81,X84,X82] : (p(s(bool,X80)) | p(s(bool,X81)) | sP10(X82,X81) | sP4(X83,X84,X82)) ) | ($spl84 | $spl122)),
  inference(resolution,[],[f11713,f27738])).
fof(f27738,plain,(
  ( ! [X187,X185,X188,X186,X184] : (sP11(X184,X185) | sP10(X186,X185) | sP4(X187,X188,X186)) ) | $spl122),
  inference(resolution,[],[f22766,f375])).
fof(f33683,plain,(
  ( ! [X78,X76,X79,X77,X75] : (p(s(bool,X75)) | p(s(bool,X76)) | sP10(X77,X76) | sP1(X77,X78,X79)) ) | ($spl84 | $spl122)),
  inference(resolution,[],[f11713,f27716])).
fof(f27716,plain,(
  ( ! [X92,X90,X88,X91,X89] : (sP11(X88,X89) | sP10(X90,X89) | sP1(X90,X91,X92)) ) | $spl122),
  inference(resolution,[],[f22766,f327])).
fof(f33682,plain,(
  ( ! [X70,X74,X72,X71,X73] : (p(s(bool,X70)) | p(s(bool,X71)) | sP10(X72,X71) | sP1(X73,X72,X74)) ) | ($spl84 | $spl122)),
  inference(resolution,[],[f11713,f27715])).
fof(f27715,plain,(
  ( ! [X87,X85,X83,X86,X84] : (sP11(X83,X84) | sP10(X85,X84) | sP1(X86,X85,X87)) ) | $spl122),
  inference(resolution,[],[f22766,f326])).
fof(f33681,plain,(
  ( ! [X68,X66,X69,X67,X65] : (p(s(bool,X65)) | p(s(bool,X66)) | sP10(X67,X66) | sP2(X67,X68,X69)) ) | ($spl84 | $spl122)),
  inference(resolution,[],[f11713,f27713])).
fof(f27713,plain,(
  ( ! [X76,X74,X77,X75,X73] : (sP11(X73,X74) | sP10(X75,X74) | sP2(X75,X76,X77)) ) | $spl122),
  inference(resolution,[],[f22766,f323])).
fof(f33680,plain,(
  ( ! [X61,X64,X62,X60,X63] : (p(s(bool,X60)) | p(s(bool,X61)) | sP10(X62,X61) | sP2(X63,X62,X64)) ) | ($spl84 | $spl122)),
  inference(resolution,[],[f11713,f27712])).
fof(f27712,plain,(
  ( ! [X70,X68,X72,X71,X69] : (sP11(X68,X69) | sP10(X70,X69) | sP2(X71,X70,X72)) ) | $spl122),
  inference(resolution,[],[f22766,f322])).
fof(f33522,plain,(
  ( ! [X107] : (sP3(X107,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X107))) ) | $spl476),
  inference(cnf_transformation,[],[f33522_D])).
fof(f33522_D,plain,(
  ( ! [X107] : (sP3(X107,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X107))) ) <=> ~$spl476),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl476])])).
fof(f16001,plain,(
  ( ! [X26,X24] : (p(s(bool,X26)) | s(bool,'F_0') = s(bool,X24) | s(bool,X24) != s(bool,X26)) ) | $spl94),
  inference(cnf_transformation,[],[f16001_D])).
fof(f16001_D,plain,(
  ( ! [X26,X24] : (p(s(bool,X26)) | s(bool,'F_0') = s(bool,X24) | s(bool,X24) != s(bool,X26)) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f33562,plain,(
  ( ! [X107,X105,X106] : (s(bool,X105) != s(bool,X106) | p(s(bool,X106)) | sP3(X107,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X107)) | s(bool,'F_0') = s(bool,X105)) ) | ($spl8 | $spl46)),
  inference(superposition,[],[f11866,f15318])).
fof(f15318,plain,(
  ( ! [X34] : (s(bool,X34) = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | s(bool,'F_0') = s(bool,X34)) ) | ($spl8 | $spl46)),
  inference(trivial_inequality_removal,[],[f15299])).
fof(f15299,plain,(
  ( ! [X34] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X34) = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | s(bool,'F_0') = s(bool,X34)) ) | ($spl8 | $spl46)),
  inference(superposition,[],[f15244,f359])).
fof(f359,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X11] : (s(bool,'T_0') = s(bool,X11) | s(bool,'F_0') = s(bool,X11))),
  file('thm.list.MEM_GENLIST','thm.bool.BOOL_CASES_AX')).
fof(f15244,plain,(
  ( ! [X28] : (s(bool,'T_0') != s(bool,X28) | s(bool,X28) = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | ($spl8 | $spl46)),
  inference(subsumption_resolution,[],[f15243,f537])).
fof(f15243,plain,(
  ( ! [X28] : (~p(s(bool,'T_0')) | s(bool,'T_0') != s(bool,X28) | s(bool,X28) = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | ($spl8 | $spl46)),
  inference(forward_demodulation,[],[f15227,f2535])).
fof(f15227,plain,(
  ( ! [X28] : (s(bool,'T_0') != s(bool,X28) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | s(bool,X28) = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f12162,f6412])).
fof(f6412,plain,(
  ( ! [X52,X53] : (~sP3('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X53,X52) | s(bool,X52) = s(bool,X53)) ) | $spl8),
  inference(resolution,[],[f2511,f334])).
fof(f2511,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | $spl8),
  inference(equality_resolution,[],[f676])).
fof(f11866,plain,(
  ( ! [X4,X5,X3] : (s(bool,X3) != s(bool,X4) | p(s(bool,X4)) | sP3(X5,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X3) | p(s(bool,X5))) ) | $spl8),
  inference(superposition,[],[f676,f332])).
fof(f332,plain,(
  ( ! [X2,X0,X1] : (sP3(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f208])).
fof(f33625,plain,(
  ( ! [X104] : (sP3(X104,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | p(s(bool,X104))) ) | $spl484),
  inference(cnf_transformation,[],[f33625_D])).
fof(f33625_D,plain,(
  ( ! [X104] : (sP3(X104,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | p(s(bool,X104))) ) <=> ~$spl484),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl484])])).
fof(f33624,plain,(
  ( ! [X103,X104] : (p(s(bool,X103)) | sP3(X104,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | p(s(bool,X104))) ) | ($spl1 | $spl8)),
  inference(subsumption_resolution,[],[f33561,f351])).
fof(f33561,plain,(
  ( ! [X103,X104] : (s(bool,'F_0') != s(bool,X103) | p(s(bool,X103)) | sP3(X104,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | p(s(bool,X104))) ) | ($spl1 | $spl8)),
  inference(superposition,[],[f11866,f14213])).
fof(f33619,plain,(
  ( ! [X93] : (sP3(X93,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),'F_0') | p(s(bool,X93))) ) | $spl480),
  inference(cnf_transformation,[],[f33619_D])).
fof(f33619_D,plain,(
  ( ! [X93] : (sP3(X93,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),'F_0') | p(s(bool,X93))) ) <=> ~$spl480),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl480])])).
fof(f677,plain,(
  ( ! [X4,X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,X4)) ) | $spl10),
  inference(cnf_transformation,[],[f677_D])).
fof(f677_D,plain,(
  ( ! [X4,X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,X4)) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f33623,plain,(
  ( ! [X101,X102,X100] : (s(bool,X100) != s(bool,X101) | p(s(bool,X101)) | sP3(X102,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),'F_0') | p(s(bool,X102))) ) | $spl8),
  inference(subsumption_resolution,[],[f33560,f337])).
fof(f337,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f210])).
fof(f210,plain,(
  ! [X0,X1] : ((s(bool,X0) != s(bool,X1) | ((~p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | p(s(bool,X0))))) & ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(flattening,[],[f209])).
fof(f209,plain,(
  ! [X0,X1] : ((s(bool,X0) != s(bool,X1) | ((~p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | p(s(bool,X0))))) & (((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0)))) | s(bool,X0) = s(bool,X1)))),
  inference(nnf_transformation,[],[f147])).
fof(f147,plain,(
  ! [X0,X1] : (s(bool,X0) = s(bool,X1) <=> ((~p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | p(s(bool,X0)))))),
  inference(ennf_transformation,[],[f67])).
fof(f67,plain,(
  ! [X0,X1] : (s(bool,X0) = s(bool,X1) <=> ((p(s(bool,X0)) => p(s(bool,X1))) & (p(s(bool,X1)) => p(s(bool,X0)))))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X12,X13] : (s(bool,X12) = s(bool,X13) <=> ((p(s(bool,X12)) => p(s(bool,X13))) & (p(s(bool,X13)) => p(s(bool,X12)))))),
  file('thm.list.MEM_GENLIST','thm.bool.EQ_IMP_THM')).
fof(f33560,plain,(
  ( ! [X101,X102,X100] : (s(bool,X100) != s(bool,X101) | p(s(bool,X101)) | sP3(X102,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),'F_0') | p(s(bool,X102)) | p(s(bool,X100))) ) | $spl8),
  inference(superposition,[],[f11866,f351])).
fof(f33622,plain,(
  ( ! [X99,X97,X98] : (s(bool,X97) != s(bool,X98) | p(s(bool,X98)) | sP3(X99,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),'F_0') | p(s(bool,X99))) ) | $spl8),
  inference(subsumption_resolution,[],[f33559,f337])).
fof(f33559,plain,(
  ( ! [X99,X97,X98] : (s(bool,X97) != s(bool,X98) | p(s(bool,X98)) | sP3(X99,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),'F_0') | p(s(bool,X99)) | p(s(bool,X97))) ) | $spl8),
  inference(superposition,[],[f11866,f353])).
fof(f353,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f214])).
fof(f214,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f75])).
fof(f75,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X11] : (s(bool,'F_0') = s(bool,X11) <=> ~p(s(bool,X11)))),
  file('thm.list.MEM_GENLIST','thm.bool.EQ_CLAUSES|split|3')).
fof(f33620,plain,(
  ( ! [X92,X91] : (s(bool,X91) != s(bool,X92) | s(bool,'T_0') = s(bool,X91)) ) | $spl482),
  inference(cnf_transformation,[],[f33620_D])).
fof(f33620_D,plain,(
  ( ! [X92,X91] : (s(bool,X91) != s(bool,X92) | s(bool,'T_0') = s(bool,X91)) ) <=> ~$spl482),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl482])])).
fof(f33621,plain,(
  ( ! [X94,X95,X96] : (s(bool,X94) != s(bool,X95) | sP3(X96,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),'F_0') | p(s(bool,X96)) | s(bool,'T_0') = s(bool,X94)) ) | ($spl8 | $spl22)),
  inference(subsumption_resolution,[],[f33558,f4447])).
fof(f4447,plain,(
  ( ! [X54,X55] : (s(bool,X54) != s(bool,X55) | ~p(s(bool,X55)) | s(bool,'T_0') = s(bool,X54)) ) | $spl22),
  inference(trivial_inequality_removal,[],[f4432])).
fof(f4432,plain,(
  ( ! [X54,X55] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,X54) != s(bool,X55) | ~p(s(bool,X55)) | s(bool,'T_0') = s(bool,X54)) ) | $spl22),
  inference(superposition,[],[f3521,f359])).
fof(f3521,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,X5) | s(bool,X4) != s(bool,X5) | ~p(s(bool,X4))) ) | $spl22),
  inference(superposition,[],[f1880,f357])).
fof(f357,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f216])).
fof(f216,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X11] : (s(bool,'T_0') = s(bool,X11) <=> p(s(bool,X11)))),
  file('thm.list.MEM_GENLIST','thm.bool.EQ_CLAUSES|split|1')).
fof(f33558,plain,(
  ( ! [X94,X95,X96] : (s(bool,X94) != s(bool,X95) | p(s(bool,X95)) | sP3(X96,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),'F_0') | p(s(bool,X96)) | s(bool,'T_0') = s(bool,X94)) ) | $spl8),
  inference(superposition,[],[f11866,f358])).
fof(f358,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.list.MEM_GENLIST',aHL_BOOL_CASES)).
fof(f33618,plain,(
  ( ! [X92,X93,X91] : (s(bool,X91) != s(bool,X92) | sP3(X93,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),'F_0') | p(s(bool,X93)) | s(bool,'T_0') = s(bool,X91)) ) | ($spl8 | $spl22)),
  inference(subsumption_resolution,[],[f33557,f4447])).
fof(f33557,plain,(
  ( ! [X92,X93,X91] : (s(bool,X91) != s(bool,X92) | p(s(bool,X92)) | sP3(X93,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),'F_0') | p(s(bool,X93)) | s(bool,'T_0') = s(bool,X91)) ) | $spl8),
  inference(superposition,[],[f11866,f359])).
fof(f33617,plain,(
  ( ! [X81] : (sP3(X81,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),'T_0') | p(s(bool,X81))) ) | $spl478),
  inference(cnf_transformation,[],[f33617_D])).
fof(f33617_D,plain,(
  ( ! [X81] : (sP3(X81,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),'T_0') | p(s(bool,X81))) ) <=> ~$spl478),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl478])])).
fof(f33554,plain,(
  ( ! [X83,X84,X82] : (s(bool,X82) != s(bool,X83) | p(s(bool,X83)) | sP3(X84,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),'T_0') | p(s(bool,X84)) | s(bool,'F_0') = s(bool,X82)) ) | $spl8),
  inference(superposition,[],[f11866,f358])).
fof(f33553,plain,(
  ( ! [X80,X81,X79] : (s(bool,X79) != s(bool,X80) | p(s(bool,X80)) | sP3(X81,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),'T_0') | p(s(bool,X81)) | s(bool,'F_0') = s(bool,X79)) ) | $spl8),
  inference(superposition,[],[f11866,f359])).
fof(f33460,plain,(
  ( ! [X107,X105,X106] : (s(bool,X105) != s(bool,X106) | p(s(bool,X106)) | sP3(X107,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X107)) | s(bool,'F_0') = s(bool,X105)) ) | ($spl8 | $spl46)),
  inference(superposition,[],[f11865,f15318])).
fof(f11865,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | p(s(bool,X1)) | sP3(X2,X0,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X2))) ) | $spl8),
  inference(superposition,[],[f676,f332])).
fof(f33521,plain,(
  ( ! [X104] : (sP3(X104,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X104))) ) | $spl474),
  inference(cnf_transformation,[],[f33521_D])).
fof(f33521_D,plain,(
  ( ! [X104] : (sP3(X104,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X104))) ) <=> ~$spl474),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl474])])).
fof(f33520,plain,(
  ( ! [X103,X104] : (p(s(bool,X103)) | sP3(X104,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X104))) ) | ($spl1 | $spl8)),
  inference(subsumption_resolution,[],[f33459,f351])).
fof(f33459,plain,(
  ( ! [X103,X104] : (s(bool,'F_0') != s(bool,X103) | p(s(bool,X103)) | sP3(X104,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X104))) ) | ($spl1 | $spl8)),
  inference(superposition,[],[f11865,f14213])).
fof(f33515,plain,(
  ( ! [X81] : (sP3(X81,'T_0','const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X81))) ) | $spl472),
  inference(cnf_transformation,[],[f33515_D])).
fof(f33515_D,plain,(
  ( ! [X81] : (sP3(X81,'T_0','const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X81))) ) <=> ~$spl472),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl472])])).
fof(f33452,plain,(
  ( ! [X83,X84,X82] : (s(bool,X82) != s(bool,X83) | p(s(bool,X83)) | sP3(X84,'T_0','const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X84)) | s(bool,'F_0') = s(bool,X82)) ) | $spl8),
  inference(superposition,[],[f11865,f358])).
fof(f33451,plain,(
  ( ! [X80,X81,X79] : (s(bool,X79) != s(bool,X80) | p(s(bool,X80)) | sP3(X81,'T_0','const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X81)) | s(bool,'F_0') = s(bool,X79)) ) | $spl8),
  inference(superposition,[],[f11865,f359])).
fof(f33409,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,'T_0') != s(bool,X14) | s(bool,X12) != s(bool,X14)) ) | ($spl8 | $spl46)),
  inference(duplicate_literal_removal,[],[f33364])).
fof(f33364,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,'T_0') != s(bool,X14) | s(bool,X12) != s(bool,X14) | ~p(s(bool,X13))) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f33359,f333])).
fof(f33359,plain,(
  ( ! [X10,X11,X9] : (~sP3(X11,X9,X10) | p(s(bool,X9)) | ~p(s(bool,X11)) | s(bool,'T_0') != s(bool,X10)) ) | ($spl8 | $spl46)),
  inference(forward_demodulation,[],[f11844,f2535])).
fof(f11844,plain,(
  ( ! [X10,X11,X9] : (s(bool,X10) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X9)) | ~p(s(bool,X11)) | ~sP3(X11,X9,X10)) ) | $spl8),
  inference(superposition,[],[f676,f334])).
fof(f33344,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,'T_0') != s(bool,X14) | s(bool,X12) != s(bool,X14)) ) | ($spl8 | $spl46)),
  inference(duplicate_literal_removal,[],[f33299])).
fof(f33299,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,'T_0') != s(bool,X14) | s(bool,X12) != s(bool,X14) | ~p(s(bool,X13))) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f33294,f333])).
fof(f33294,plain,(
  ( ! [X6,X8,X7] : (~sP3(X8,X7,X6) | p(s(bool,X6)) | ~p(s(bool,X8)) | s(bool,'T_0') != s(bool,X7)) ) | ($spl8 | $spl46)),
  inference(forward_demodulation,[],[f11843,f2535])).
fof(f11843,plain,(
  ( ! [X6,X8,X7] : (s(bool,X7) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X6)) | ~p(s(bool,X8)) | ~sP3(X8,X7,X6)) ) | $spl8),
  inference(superposition,[],[f676,f334])).
fof(f33284,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | s(bool,'T_0') != s(bool,X4) | p(s(bool,X5)) | s(bool,X3) != s(bool,X4)) ) | ($spl8 | $spl46)),
  inference(duplicate_literal_removal,[],[f33255])).
fof(f33255,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | s(bool,'T_0') != s(bool,X4) | p(s(bool,X5)) | p(s(bool,X5)) | s(bool,X3) != s(bool,X4)) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f33253,f335])).
fof(f335,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | s(bool,X0) != s(bool,X1) | ~sP3(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f208])).
fof(f33253,plain,(
  ( ! [X4,X5,X3] : (sP3(X5,X3,X4) | p(s(bool,X3)) | s(bool,'T_0') != s(bool,X4) | p(s(bool,X5))) ) | ($spl8 | $spl46)),
  inference(forward_demodulation,[],[f11842,f2535])).
fof(f11842,plain,(
  ( ! [X4,X5,X3] : (s(bool,X4) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X3)) | sP3(X5,X3,X4) | p(s(bool,X5))) ) | $spl8),
  inference(superposition,[],[f676,f332])).
fof(f33242,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | s(bool,'T_0') != s(bool,X4) | p(s(bool,X5)) | s(bool,X3) != s(bool,X4)) ) | ($spl8 | $spl46)),
  inference(duplicate_literal_removal,[],[f33212])).
fof(f33212,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | s(bool,'T_0') != s(bool,X4) | p(s(bool,X5)) | p(s(bool,X5)) | s(bool,X3) != s(bool,X4)) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f33210,f335])).
fof(f33210,plain,(
  ( ! [X2,X0,X1] : (sP3(X2,X1,X0) | p(s(bool,X0)) | s(bool,'T_0') != s(bool,X1) | p(s(bool,X2))) ) | ($spl8 | $spl46)),
  inference(forward_demodulation,[],[f11841,f2535])).
fof(f11841,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X0)) | sP3(X2,X1,X0) | p(s(bool,X2))) ) | $spl8),
  inference(superposition,[],[f676,f332])).
fof(f29618,plain,(
  ( ! [X76] : (sP10(X76,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | sP10('F_0',X76)) ) | $spl342),
  inference(cnf_transformation,[],[f29618_D])).
fof(f29618_D,plain,(
  ( ! [X76] : (sP10(X76,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | sP10('F_0',X76)) ) <=> ~$spl342),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl342])])).
fof(f33166,plain,(
  ( ! [X28,X26,X27] : (sP10('F_0',X26) | sP10(X26,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X27)) | ~p(s(bool,X28))) ) | ($spl8 | $spl122 | $spl133)),
  inference(resolution,[],[f29191,f11976])).
fof(f11976,plain,(
  ( ! [X130,X131] : (~sP8('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X130,X131) | p(s(bool,X131)) | ~p(s(bool,X130))) ) | $spl8),
  inference(resolution,[],[f2511,f408])).
fof(f408,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f243])).
fof(f29191,plain,(
  ( ! [X546,X549,X547,X548] : (sP10('F_0',X546) | sP10(X546,X547) | sP8(X547,X548,X549)) ) | ($spl122 | $spl133)),
  inference(subsumption_resolution,[],[f29103,f528])).
fof(f528,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f139])).
fof(f139,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.list.MEM_GENLIST',aHL_FALSITY)).
fof(f29103,plain,(
  ( ! [X546,X549,X547,X548] : (p(s(bool,'F_0')) | sP10('F_0',X546) | sP10(X546,X547) | sP8(X547,X548,X549)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f20735,f19923])).
fof(f19923,plain,(
  ( ! [X271,X273,X270,X272] : (sP10(X270,X271) | sP10(X270,'F_0') | sP8(X271,X272,X273)) ) | $spl122),
  inference(resolution,[],[f18884,f411])).
fof(f18884,plain,(
  ( ! [X231,X232] : (~p(s(bool,X232)) | sP10(X231,X232) | sP10(X231,'F_0')) ) | $spl122),
  inference(resolution,[],[f18767,f443])).
fof(f20735,plain,(
  ( ! [X227,X228] : (~sP10(X227,X228) | p(s(bool,X228)) | sP10('F_0',X227)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f20649,f438])).
fof(f20649,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | sP10('F_0',X1)) ) | ($spl122 | $spl133)),
  inference(subsumption_resolution,[],[f20625,f537])).
fof(f20625,plain,(
  ( ! [X1] : (sP10('F_0',X1) | ~p(s(bool,X1)) | ~p(s(bool,'T_0'))) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f20608,f439])).
fof(f20608,plain,(
  ( ! [X22] : (sP10(X22,'T_0') | sP10('F_0',X22)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f20245,f18985])).
fof(f18985,plain,(
  ~sP10('F_0','F_0') | $spl133),
  inference(cnf_transformation,[],[f18985_D])).
fof(f18985_D,plain,(
  ~sP10('F_0','F_0') <=> ~$spl133),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl133])])).
fof(f20245,plain,(
  ( ! [X269,X270] : (sP10(X269,'T_0') | sP10(X270,X269) | sP10(X270,'F_0')) ) | $spl122),
  inference(resolution,[],[f20069,f18884])).
fof(f20069,plain,(
  ( ! [X0] : (p(s(bool,X0)) | sP10(X0,'T_0')) ) | $spl122),
  inference(subsumption_resolution,[],[f20056,f492])).
fof(f20056,plain,(
  ( ! [X0] : (sP10(X0,'T_0') | p(s(bool,'F_0')) | p(s(bool,X0))) ) | $spl122),
  inference(resolution,[],[f19971,f438])).
fof(f19971,plain,(
  ( ! [X401] : (sP10(X401,'F_0') | sP10(X401,'T_0')) ) | $spl122),
  inference(resolution,[],[f18884,f533])).
fof(f25572,plain,(
  ( ! [X70,X68,X69] : (sP8(X68,X69,X70)) ) | $spl274),
  inference(cnf_transformation,[],[f25572_D])).
fof(f25572_D,plain,(
  ( ! [X70,X68,X69] : (sP8(X68,X69,X70)) ) <=> ~$spl274),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl274])])).
fof(f23175,plain,(
  ( ! [X17,X16] : (sP10('F_0',X16) | s(bool,'F_0') = s(bool,X17) | sP10(X17,X16)) ) | $spl212),
  inference(cnf_transformation,[],[f23175_D])).
fof(f23175_D,plain,(
  ( ! [X17,X16] : (sP10('F_0',X16) | s(bool,'F_0') = s(bool,X17) | sP10(X17,X16)) ) <=> ~$spl212),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl212])])).
fof(f33174,plain,(
  ( ! [X28,X26,X24,X27,X25] : (sP10('F_0',X24) | sP8(X25,X26,X27) | s(bool,'F_0') = s(bool,X28) | sP10(X28,X24)) ) | ($spl8 | $spl46 | $spl122 | $spl133)),
  inference(subsumption_resolution,[],[f33143,f411])).
fof(f33143,plain,(
  ( ! [X28,X26,X24,X27,X25] : (sP10('F_0',X24) | sP8(X25,X26,X27) | s(bool,'F_0') = s(bool,X28) | ~p(s(bool,X25)) | sP10(X28,X24)) ) | ($spl8 | $spl46 | $spl122 | $spl133)),
  inference(resolution,[],[f29191,f16692])).
fof(f16692,plain,(
  ( ! [X302,X304,X303] : (~sP10(X303,X304) | s(bool,'F_0') = s(bool,X302) | ~p(s(bool,X304)) | sP10(X302,X303)) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16255,f439])).
fof(f16255,plain,(
  ( ! [X592,X593] : (sP10(X592,X593) | p(s(bool,X593)) | s(bool,'F_0') = s(bool,X592)) ) | ($spl8 | $spl46)),
  inference(subsumption_resolution,[],[f16254,f537])).
fof(f16254,plain,(
  ( ! [X592,X593] : (~p(s(bool,'T_0')) | sP10(X592,X593) | p(s(bool,X593)) | s(bool,'F_0') = s(bool,X592)) ) | ($spl8 | $spl46)),
  inference(forward_demodulation,[],[f15610,f2535])).
fof(f15610,plain,(
  ( ! [X592,X593] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | sP10(X592,X593) | p(s(bool,X593)) | s(bool,'F_0') = s(bool,X592)) ) | ($spl8 | $spl46)),
  inference(superposition,[],[f440,f15318])).
fof(f33106,plain,(
  ( ! [X24,X23,X25] : (sP10('F_0',X23) | sP10(X23,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X24)) | ~p(s(bool,X25))) ) | ($spl8 | $spl122 | $spl133)),
  inference(resolution,[],[f29190,f11977])).
fof(f11977,plain,(
  ( ! [X132,X133] : (~sP8(X132,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X133) | p(s(bool,X133)) | ~p(s(bool,X132))) ) | $spl8),
  inference(resolution,[],[f2511,f408])).
fof(f29190,plain,(
  ( ! [X542,X545,X543,X544] : (sP10('F_0',X542) | sP10(X542,X543) | sP8(X544,X543,X545)) ) | ($spl122 | $spl133)),
  inference(subsumption_resolution,[],[f29102,f528])).
fof(f29102,plain,(
  ( ! [X542,X545,X543,X544] : (p(s(bool,'F_0')) | sP10('F_0',X542) | sP10(X542,X543) | sP8(X544,X543,X545)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f20735,f19922])).
fof(f19922,plain,(
  ( ! [X269,X267,X268,X266] : (sP10(X266,X267) | sP10(X266,'F_0') | sP8(X268,X267,X269)) ) | $spl122),
  inference(resolution,[],[f18884,f410])).
fof(f33115,plain,(
  ( ! [X28,X26,X24,X27,X25] : (sP10('F_0',X24) | sP8(X25,X26,X27) | s(bool,'F_0') = s(bool,X28) | sP10(X28,X24)) ) | ($spl8 | $spl46 | $spl122 | $spl133)),
  inference(subsumption_resolution,[],[f33084,f410])).
fof(f33084,plain,(
  ( ! [X28,X26,X24,X27,X25] : (sP10('F_0',X24) | sP8(X25,X26,X27) | s(bool,'F_0') = s(bool,X28) | ~p(s(bool,X26)) | sP10(X28,X24)) ) | ($spl8 | $spl46 | $spl122 | $spl133)),
  inference(resolution,[],[f29190,f16692])).
fof(f33050,plain,(
  ( ! [X30,X31,X29] : (sP10('F_0',X29) | sP10(X29,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X30)) | p(s(bool,X31))) ) | ($spl8 | $spl122 | $spl133)),
  inference(resolution,[],[f29189,f11960])).
fof(f11960,plain,(
  ( ! [X99,X98] : (~sP6('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X98,X99) | p(s(bool,X99)) | p(s(bool,X98))) ) | $spl8),
  inference(resolution,[],[f2511,f391])).
fof(f391,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f236])).
fof(f29189,plain,(
  ( ! [X540,X538,X541,X539] : (sP10('F_0',X538) | sP10(X538,X539) | sP6(X539,X540,X541)) ) | ($spl122 | $spl133)),
  inference(subsumption_resolution,[],[f29101,f528])).
fof(f29101,plain,(
  ( ! [X540,X538,X541,X539] : (p(s(bool,'F_0')) | sP10('F_0',X538) | sP10(X538,X539) | sP6(X539,X540,X541)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f20735,f19909])).
fof(f19909,plain,(
  ( ! [X216,X215,X217,X214] : (sP10(X214,X215) | sP10(X214,'F_0') | sP6(X215,X216,X217)) ) | $spl122),
  inference(resolution,[],[f18884,f394])).
fof(f7342,plain,(
  ( ! [X2,X3] : (sP7('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X2,X3)) ) | $spl74),
  inference(cnf_transformation,[],[f7342_D])).
fof(f7342_D,plain,(
  ( ! [X2,X3] : (sP7('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X2,X3)) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f33049,plain,(
  ( ! [X28,X26,X27] : (sP10('F_0',X26) | sP10(X26,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | sP7('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X27,X28)) ) | ($spl8 | $spl122 | $spl133)),
  inference(resolution,[],[f29189,f12008])).
fof(f12008,plain,(
  ( ! [X167,X166] : (~sP6('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X166,X167) | sP7('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X166,X167)) ) | $spl8),
  inference(resolution,[],[f2511,f539])).
fof(f539,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~sP6(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(duplicate_literal_removal,[],[f387])).
fof(f387,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~sP6(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f233])).
fof(f233,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP6(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP6(X0,X1,X2) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f232])).
fof(f232,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP6(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP6(X2,X0,X1) | sP7(X2,X0,X1)))),
  inference(flattening,[],[f231])).
fof(f231,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP6(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP6(X2,X0,X1)) | sP7(X2,X0,X1)))),
  inference(nnf_transformation,[],[f183])).
fof(f183,plain,(
  ! [X2,X0,X1] : (sP7(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP6(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f25423,plain,(
  ( ! [X70,X68,X69] : (sP6(X68,X69,X70)) ) | $spl268),
  inference(cnf_transformation,[],[f25423_D])).
fof(f25423_D,plain,(
  ( ! [X70,X68,X69] : (sP6(X68,X69,X70)) ) <=> ~$spl268),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl268])])).
fof(f33058,plain,(
  ( ! [X28,X26,X24,X27,X25] : (sP10('F_0',X24) | sP6(X25,X26,X27) | s(bool,'F_0') = s(bool,X28) | sP10(X28,X24)) ) | ($spl8 | $spl46 | $spl122 | $spl133)),
  inference(subsumption_resolution,[],[f33026,f394])).
fof(f33026,plain,(
  ( ! [X28,X26,X24,X27,X25] : (sP10('F_0',X24) | sP6(X25,X26,X27) | s(bool,'F_0') = s(bool,X28) | ~p(s(bool,X25)) | sP10(X28,X24)) ) | ($spl8 | $spl46 | $spl122 | $spl133)),
  inference(resolution,[],[f29189,f16692])).
fof(f32992,plain,(
  ( ! [X30,X31,X29] : (sP10('F_0',X29) | sP10(X29,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | ~p(s(bool,X30)) | p(s(bool,X31))) ) | ($spl8 | $spl122 | $spl133)),
  inference(resolution,[],[f29188,f11946])).
fof(f11946,plain,(
  ( ! [X70,X71] : (~sP4('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X70,X71) | ~p(s(bool,X71)) | p(s(bool,X70))) ) | $spl8),
  inference(resolution,[],[f2511,f374])).
fof(f374,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f228])).
fof(f29188,plain,(
  ( ! [X536,X534,X537,X535] : (sP10('F_0',X534) | sP10(X534,X535) | sP4(X535,X536,X537)) ) | ($spl122 | $spl133)),
  inference(subsumption_resolution,[],[f29100,f528])).
fof(f29100,plain,(
  ( ! [X536,X534,X537,X535] : (p(s(bool,'F_0')) | sP10('F_0',X534) | sP10(X534,X535) | sP4(X535,X536,X537)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f20735,f19893])).
fof(f19893,plain,(
  ( ! [X152,X151,X153,X150] : (sP10(X150,X151) | sP10(X150,'F_0') | sP4(X151,X152,X153)) ) | $spl122),
  inference(resolution,[],[f18884,f377])).
fof(f7334,plain,(
  ( ! [X2,X3] : (sP5('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X2,X3)) ) | $spl72),
  inference(cnf_transformation,[],[f7334_D])).
fof(f7334_D,plain,(
  ( ! [X2,X3] : (sP5('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X2,X3)) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f32991,plain,(
  ( ! [X28,X26,X27] : (sP10('F_0',X26) | sP10(X26,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | sP5('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X27,X28)) ) | ($spl8 | $spl122 | $spl133)),
  inference(resolution,[],[f29188,f12007])).
fof(f12007,plain,(
  ( ! [X165,X164] : (~sP4('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X164,X165) | sP5('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X164,X165)) ) | $spl8),
  inference(resolution,[],[f2511,f538])).
fof(f538,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | ~sP4(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(duplicate_literal_removal,[],[f370])).
fof(f370,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | ~sP4(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f225])).
fof(f225,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP4(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP4(X0,X1,X2) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f224])).
fof(f224,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP4(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP4(X2,X0,X1) | sP5(X2,X0,X1)))),
  inference(flattening,[],[f223])).
fof(f223,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP4(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP4(X2,X0,X1)) | sP5(X2,X0,X1)))),
  inference(nnf_transformation,[],[f180])).
fof(f180,plain,(
  ! [X2,X0,X1] : (sP5(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP4(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f25364,plain,(
  ( ! [X70,X68,X69] : (sP4(X68,X69,X70)) ) | $spl264),
  inference(cnf_transformation,[],[f25364_D])).
fof(f25364_D,plain,(
  ( ! [X70,X68,X69] : (sP4(X68,X69,X70)) ) <=> ~$spl264),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl264])])).
fof(f33000,plain,(
  ( ! [X28,X26,X24,X27,X25] : (sP10('F_0',X24) | sP4(X25,X26,X27) | s(bool,'F_0') = s(bool,X28) | sP10(X28,X24)) ) | ($spl8 | $spl46 | $spl122 | $spl133)),
  inference(subsumption_resolution,[],[f32968,f377])).
fof(f32968,plain,(
  ( ! [X28,X26,X24,X27,X25] : (sP10('F_0',X24) | sP4(X25,X26,X27) | s(bool,'F_0') = s(bool,X28) | ~p(s(bool,X25)) | sP10(X28,X24)) ) | ($spl8 | $spl46 | $spl122 | $spl133)),
  inference(resolution,[],[f29188,f16692])).
fof(f32930,plain,(
  ( ! [X21,X22,X20] : (sP10('F_0',X20) | sP10(X20,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X21)) | ~p(s(bool,X22))) ) | ($spl8 | $spl122 | $spl133)),
  inference(resolution,[],[f29187,f11947])).
fof(f11947,plain,(
  ( ! [X72,X73] : (~sP4(X72,X73,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X73)) | ~p(s(bool,X72))) ) | $spl8),
  inference(resolution,[],[f2511,f374])).
fof(f29187,plain,(
  ( ! [X532,X530,X533,X531] : (sP10('F_0',X530) | sP10(X530,X531) | sP4(X532,X533,X531)) ) | ($spl122 | $spl133)),
  inference(subsumption_resolution,[],[f29099,f528])).
fof(f29099,plain,(
  ( ! [X532,X530,X533,X531] : (p(s(bool,'F_0')) | sP10('F_0',X530) | sP10(X530,X531) | sP4(X532,X533,X531)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f20735,f19892])).
fof(f19892,plain,(
  ( ! [X146,X149,X147,X148] : (sP10(X146,X147) | sP10(X146,'F_0') | sP4(X148,X149,X147)) ) | $spl122),
  inference(resolution,[],[f18884,f375])).
fof(f32941,plain,(
  ( ! [X28,X26,X24,X27,X25] : (sP10('F_0',X24) | sP4(X25,X26,X27) | s(bool,'F_0') = s(bool,X28) | sP10(X28,X24)) ) | ($spl8 | $spl46 | $spl122 | $spl133)),
  inference(subsumption_resolution,[],[f32909,f375])).
fof(f32909,plain,(
  ( ! [X28,X26,X24,X27,X25] : (sP10('F_0',X24) | sP4(X25,X26,X27) | s(bool,'F_0') = s(bool,X28) | ~p(s(bool,X27)) | sP10(X28,X24)) ) | ($spl8 | $spl46 | $spl122 | $spl133)),
  inference(resolution,[],[f29187,f16692])).
fof(f32831,plain,(
  ( ! [X30,X31,X32] : (sP10('F_0',X30) | sP10(X30,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X31)) | ~p(s(bool,X32))) ) | ($spl8 | $spl122 | $spl133)),
  inference(resolution,[],[f29186,f11923])).
fof(f11923,plain,(
  ( ! [X39,X38] : (~sP1('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X38,X39) | p(s(bool,X39)) | ~p(s(bool,X38))) ) | $spl8),
  inference(resolution,[],[f2511,f324])).
fof(f324,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP1(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f205])).
fof(f29186,plain,(
  ( ! [X528,X526,X529,X527] : (sP10('F_0',X526) | sP10(X526,X527) | sP1(X527,X528,X529)) ) | ($spl122 | $spl133)),
  inference(subsumption_resolution,[],[f29098,f528])).
fof(f29098,plain,(
  ( ! [X528,X526,X529,X527] : (p(s(bool,'F_0')) | sP10('F_0',X526) | sP10(X526,X527) | sP1(X527,X528,X529)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f20735,f19870])).
fof(f19870,plain,(
  ( ! [X74,X72,X75,X73] : (sP10(X72,X73) | sP10(X72,'F_0') | sP1(X73,X74,X75)) ) | $spl122),
  inference(resolution,[],[f18884,f327])).
fof(f21238,plain,(
  ( ! [X19,X18,X16] : (sP3(X18,X16,X19) | ~sP0(X19,X16,X18) | ~sP2(X18,X19,X16) | p(s(bool,X16))) ) | $spl178),
  inference(cnf_transformation,[],[f21238_D])).
fof(f21238_D,plain,(
  ( ! [X19,X18,X16] : (sP3(X18,X16,X19) | ~sP0(X19,X16,X18) | ~sP2(X18,X19,X16) | p(s(bool,X16))) ) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f23178,plain,(
  ( ! [X28] : (sP10('F_0',X28)) ) | $spl216),
  inference(cnf_transformation,[],[f23178_D])).
fof(f23178_D,plain,(
  ( ! [X28] : (sP10('F_0',X28)) ) <=> ~$spl216),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl216])])).
fof(f32841,plain,(
  ( ! [X19,X17,X18,X16] : (sP10('F_0',X16) | sP3(X18,X17,X19) | p(s(bool,X17)) | ~sP2(X18,X19,X17) | ~sP0(X19,X17,X18)) ) | ($spl122 | $spl133)),
  inference(subsumption_resolution,[],[f32827,f20735])).
fof(f32827,plain,(
  ( ! [X19,X17,X18,X16] : (sP10('F_0',X16) | sP10(X16,X17) | sP3(X18,X17,X19) | p(s(bool,X17)) | ~sP2(X18,X19,X17) | ~sP0(X19,X17,X18)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f29186,f318])).
fof(f25091,plain,(
  ( ! [X70,X68,X69] : (sP1(X68,X69,X70)) ) | $spl252),
  inference(cnf_transformation,[],[f25091_D])).
fof(f25091_D,plain,(
  ( ! [X70,X68,X69] : (sP1(X68,X69,X70)) ) <=> ~$spl252),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl252])])).
fof(f32839,plain,(
  ( ! [X28,X26,X24,X27,X25] : (sP10('F_0',X24) | sP1(X25,X26,X27) | s(bool,'F_0') = s(bool,X28) | sP10(X28,X24)) ) | ($spl8 | $spl46 | $spl122 | $spl133)),
  inference(subsumption_resolution,[],[f32807,f327])).
fof(f32807,plain,(
  ( ! [X28,X26,X24,X27,X25] : (sP10('F_0',X24) | sP1(X25,X26,X27) | s(bool,'F_0') = s(bool,X28) | ~p(s(bool,X25)) | sP10(X28,X24)) ) | ($spl8 | $spl46 | $spl122 | $spl133)),
  inference(resolution,[],[f29186,f16692])).
fof(f32770,plain,(
  ( ! [X28,X29,X27] : (sP10('F_0',X27) | sP10(X27,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X28)) | ~p(s(bool,X29))) ) | ($spl8 | $spl122 | $spl133)),
  inference(resolution,[],[f29185,f11924])).
fof(f11924,plain,(
  ( ! [X41,X40] : (~sP1(X40,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X41) | p(s(bool,X41)) | ~p(s(bool,X40))) ) | $spl8),
  inference(resolution,[],[f2511,f324])).
fof(f29185,plain,(
  ( ! [X524,X522,X525,X523] : (sP10('F_0',X522) | sP10(X522,X523) | sP1(X524,X523,X525)) ) | ($spl122 | $spl133)),
  inference(subsumption_resolution,[],[f29097,f528])).
fof(f29097,plain,(
  ( ! [X524,X522,X525,X523] : (p(s(bool,'F_0')) | sP10('F_0',X522) | sP10(X522,X523) | sP1(X524,X523,X525)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f20735,f19869])).
fof(f19869,plain,(
  ( ! [X70,X68,X71,X69] : (sP10(X68,X69) | sP10(X68,'F_0') | sP1(X70,X69,X71)) ) | $spl122),
  inference(resolution,[],[f18884,f326])).
fof(f32779,plain,(
  ( ! [X28,X26,X24,X27,X25] : (sP10('F_0',X24) | sP1(X25,X26,X27) | s(bool,'F_0') = s(bool,X28) | sP10(X28,X24)) ) | ($spl8 | $spl46 | $spl122 | $spl133)),
  inference(subsumption_resolution,[],[f32747,f326])).
fof(f32747,plain,(
  ( ! [X28,X26,X24,X27,X25] : (sP10('F_0',X24) | sP1(X25,X26,X27) | s(bool,'F_0') = s(bool,X28) | ~p(s(bool,X26)) | sP10(X28,X24)) ) | ($spl8 | $spl46 | $spl122 | $spl133)),
  inference(resolution,[],[f29185,f16692])).
fof(f32710,plain,(
  ( ! [X43,X41,X42] : (sP10('F_0',X41) | sP10(X41,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X42)) | ~p(s(bool,X43))) ) | ($spl8 | $spl122 | $spl133)),
  inference(resolution,[],[f29184,f11920])).
fof(f11920,plain,(
  ( ! [X33,X32] : (~sP2('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X32,X33) | p(s(bool,X33)) | ~p(s(bool,X32))) ) | $spl8),
  inference(resolution,[],[f2511,f320])).
fof(f320,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f202])).
fof(f29184,plain,(
  ( ! [X520,X518,X521,X519] : (sP10('F_0',X518) | sP10(X518,X519) | sP2(X519,X520,X521)) ) | ($spl122 | $spl133)),
  inference(subsumption_resolution,[],[f29096,f528])).
fof(f29096,plain,(
  ( ! [X520,X518,X521,X519] : (p(s(bool,'F_0')) | sP10('F_0',X518) | sP10(X518,X519) | sP2(X519,X520,X521)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f20735,f19867])).
fof(f19867,plain,(
  ( ! [X61,X62,X60,X63] : (sP10(X60,X61) | sP10(X60,'F_0') | sP2(X61,X62,X63)) ) | $spl122),
  inference(resolution,[],[f18884,f323])).
fof(f21149,plain,(
  ( ! [X19,X18,X16] : (sP3(X16,X18,X19) | ~sP0(X19,X18,X16) | ~sP1(X18,X19,X16) | p(s(bool,X16))) ) | $spl176),
  inference(cnf_transformation,[],[f21149_D])).
fof(f21149_D,plain,(
  ( ! [X19,X18,X16] : (sP3(X16,X18,X19) | ~sP0(X19,X18,X16) | ~sP1(X18,X19,X16) | p(s(bool,X16))) ) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f32721,plain,(
  ( ! [X19,X17,X18,X16] : (sP10('F_0',X16) | sP3(X17,X18,X19) | p(s(bool,X17)) | ~sP1(X18,X19,X17) | ~sP0(X19,X18,X17)) ) | ($spl122 | $spl133)),
  inference(subsumption_resolution,[],[f32702,f20735])).
fof(f32702,plain,(
  ( ! [X19,X17,X18,X16] : (sP10('F_0',X16) | sP10(X16,X17) | sP3(X17,X18,X19) | p(s(bool,X17)) | ~sP1(X18,X19,X17) | ~sP0(X19,X18,X17)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f29184,f319])).
fof(f319,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X0)) | ~sP2(X0,X2,X1) | ~sP1(X1,X2,X0) | ~sP0(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f199])).
fof(f25031,plain,(
  ( ! [X70,X68,X69] : (sP2(X68,X69,X70)) ) | $spl248),
  inference(cnf_transformation,[],[f25031_D])).
fof(f25031_D,plain,(
  ( ! [X70,X68,X69] : (sP2(X68,X69,X70)) ) <=> ~$spl248),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl248])])).
fof(f32718,plain,(
  ( ! [X28,X26,X24,X27,X25] : (sP10('F_0',X24) | sP2(X25,X26,X27) | s(bool,'F_0') = s(bool,X28) | sP10(X28,X24)) ) | ($spl8 | $spl46 | $spl122 | $spl133)),
  inference(subsumption_resolution,[],[f32682,f323])).
fof(f32682,plain,(
  ( ! [X28,X26,X24,X27,X25] : (sP10('F_0',X24) | sP2(X25,X26,X27) | s(bool,'F_0') = s(bool,X28) | ~p(s(bool,X25)) | sP10(X28,X24)) ) | ($spl8 | $spl46 | $spl122 | $spl133)),
  inference(resolution,[],[f29184,f16692])).
fof(f32645,plain,(
  ( ! [X39,X38,X40] : (sP10('F_0',X38) | sP10(X38,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X39)) | ~p(s(bool,X40))) ) | ($spl8 | $spl122 | $spl133)),
  inference(resolution,[],[f29183,f11921])).
fof(f11921,plain,(
  ( ! [X35,X34] : (~sP2(X34,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X35) | p(s(bool,X35)) | ~p(s(bool,X34))) ) | $spl8),
  inference(resolution,[],[f2511,f320])).
fof(f29183,plain,(
  ( ! [X516,X514,X517,X515] : (sP10('F_0',X514) | sP10(X514,X515) | sP2(X516,X515,X517)) ) | ($spl122 | $spl133)),
  inference(subsumption_resolution,[],[f29095,f528])).
fof(f29095,plain,(
  ( ! [X516,X514,X517,X515] : (p(s(bool,'F_0')) | sP10('F_0',X514) | sP10(X514,X515) | sP2(X516,X515,X517)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f20735,f19866])).
fof(f19866,plain,(
  ( ! [X59,X57,X58,X56] : (sP10(X56,X57) | sP10(X56,'F_0') | sP2(X58,X57,X59)) ) | $spl122),
  inference(resolution,[],[f18884,f322])).
fof(f32656,plain,(
  ( ! [X24,X23] : (sP10('F_0',X23) | s(bool,'T_0') != s(bool,X24) | sP10(X23,X24)) ) | $spl468),
  inference(cnf_transformation,[],[f32656_D])).
fof(f32656_D,plain,(
  ( ! [X24,X23] : (sP10('F_0',X23) | s(bool,'T_0') != s(bool,X24) | sP10(X23,X24)) ) <=> ~$spl468),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl468])])).
fof(f32640,plain,(
  ( ! [X24,X23,X25] : (sP10('F_0',X23) | sP10(X23,X24) | s(bool,'T_0') != s(bool,X24) | ~p(s(bool,X25))) ) | ($spl6 | $spl122 | $spl133)),
  inference(resolution,[],[f29183,f14550])).
fof(f14550,plain,(
  ( ! [X14,X13] : (~sP2(X13,X14,'F_0') | s(bool,'T_0') != s(bool,X14) | ~p(s(bool,X13))) ) | $spl6),
  inference(subsumption_resolution,[],[f14549,f14372])).
fof(f14372,plain,(
  ( ! [X4,X5] : (sP0(X4,'F_0',X5) | s(bool,'T_0') != s(bool,X4)) ) | $spl6),
  inference(subsumption_resolution,[],[f14364,f329])).
fof(f14364,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | p(s(bool,X5)) | sP0(X4,'F_0',X5)) ) | $spl6),
  inference(resolution,[],[f14153,f313])).
fof(f313,plain,(
  ( ! [X2,X0,X1] : (sP0(X2,X1,X0) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f199])).
fof(f14153,plain,(
  ( ! [X6,X7] : (sP3(X7,'F_0',X6) | s(bool,'T_0') != s(bool,X6) | p(s(bool,X7))) ) | $spl6),
  inference(superposition,[],[f1060,f332])).
fof(f1060,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl6),
  inference(subsumption_resolution,[],[f1046,f550])).
fof(f1046,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | $spl6),
  inference(superposition,[],[f580,f355])).
fof(f355,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f215])).
fof(f215,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X11] : (s(bool,'T_0') = s(bool,X11) <=> p(s(bool,X11)))),
  file('thm.list.MEM_GENLIST','thm.bool.EQ_CLAUSES|split|0')).
fof(f580,plain,(
  s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | $spl6),
  inference(resolution,[],[f550,f350])).
fof(f350,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f213])).
fof(f14549,plain,(
  ( ! [X14,X13] : (~p(s(bool,X13)) | s(bool,'T_0') != s(bool,X14) | ~sP2(X13,X14,'F_0') | ~sP0(X14,'F_0',X13)) ) | $spl6),
  inference(subsumption_resolution,[],[f14548,f474])).
fof(f474,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f271])).
fof(f271,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f270])).
fof(f270,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f161])).
fof(f161,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f113])).
fof(f113,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f112])).
fof(f112,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X11] : ((p(s(bool,X11)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X11)))),
  file('thm.list.MEM_GENLIST','thm.bool.IMP_CLAUSES|split|4')).
fof(f14548,plain,(
  ( ! [X14,X13] : (~p(s(bool,X13)) | s(bool,'T_0') != s(bool,X14) | p(s(bool,'F_0')) | ~sP2(X13,X14,'F_0') | ~sP0(X14,'F_0',X13)) ) | $spl6),
  inference(subsumption_resolution,[],[f14535,f327])).
fof(f14535,plain,(
  ( ! [X14,X13] : (~p(s(bool,X13)) | s(bool,'T_0') != s(bool,X14) | p(s(bool,'F_0')) | ~sP2(X13,X14,'F_0') | ~sP1('F_0',X14,X13) | ~sP0(X14,'F_0',X13)) ) | $spl6),
  inference(resolution,[],[f14155,f318])).
fof(f14155,plain,(
  ( ! [X10,X11] : (~sP3(X11,'F_0',X10) | ~p(s(bool,X11)) | s(bool,'T_0') != s(bool,X10)) ) | $spl6),
  inference(superposition,[],[f1060,f334])).
fof(f32654,plain,(
  ( ! [X28,X26,X24,X27,X25] : (sP10('F_0',X24) | sP2(X25,X26,X27) | s(bool,'F_0') = s(bool,X28) | sP10(X28,X24)) ) | ($spl8 | $spl46 | $spl122 | $spl133)),
  inference(subsumption_resolution,[],[f32618,f322])).
fof(f32618,plain,(
  ( ! [X28,X26,X24,X27,X25] : (sP10('F_0',X24) | sP2(X25,X26,X27) | s(bool,'F_0') = s(bool,X28) | ~p(s(bool,X26)) | sP10(X28,X24)) ) | ($spl8 | $spl46 | $spl122 | $spl133)),
  inference(resolution,[],[f29183,f16692])).
fof(f29425,plain,(
  ( ! [X76] : (sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X76) | sP10('F_0',X76)) ) | $spl338),
  inference(cnf_transformation,[],[f29425_D])).
fof(f29425_D,plain,(
  ( ! [X76] : (sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X76) | sP10('F_0',X76)) ) <=> ~$spl338),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl338])])).
fof(f32583,plain,(
  ( ! [X28,X26,X27] : (sP10('F_0',X26) | sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X26) | p(s(bool,X27)) | ~p(s(bool,X28))) ) | ($spl8 | $spl122 | $spl133)),
  inference(resolution,[],[f29182,f11976])).
fof(f29182,plain,(
  ( ! [X509,X507,X508,X510] : (sP10('F_0',X507) | sP10(X508,X507) | sP8(X508,X509,X510)) ) | ($spl122 | $spl133)),
  inference(subsumption_resolution,[],[f29092,f528])).
fof(f29092,plain,(
  ( ! [X509,X507,X508,X510] : (p(s(bool,'F_0')) | sP10('F_0',X507) | sP10(X508,X507) | sP8(X508,X509,X510)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f20735,f19370])).
fof(f19370,plain,(
  ( ! [X271,X273,X270,X272] : (sP10(X270,X271) | sP10(X271,'F_0') | sP8(X270,X272,X273)) ) | $spl122),
  inference(resolution,[],[f18883,f411])).
fof(f32591,plain,(
  ( ! [X30,X33,X31,X29,X32] : (sP10('F_0',X29) | sP8(X30,X31,X32) | s(bool,'F_0') = s(bool,X33) | sP10(X33,X29)) ) | ($spl8 | $spl46 | $spl122 | $spl133)),
  inference(subsumption_resolution,[],[f32561,f411])).
fof(f32561,plain,(
  ( ! [X30,X33,X31,X29,X32] : (sP10('F_0',X29) | sP8(X30,X31,X32) | s(bool,'F_0') = s(bool,X33) | ~p(s(bool,X30)) | sP10(X33,X29)) ) | ($spl8 | $spl46 | $spl122 | $spl133)),
  inference(resolution,[],[f29182,f16693])).
fof(f16693,plain,(
  ( ! [X306,X305,X307] : (~sP10(X307,X306) | s(bool,'F_0') = s(bool,X305) | ~p(s(bool,X307)) | sP10(X305,X306)) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16255,f439])).
fof(f32490,plain,(
  ( ! [X24,X23,X25] : (sP10('F_0',X23) | sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X23) | p(s(bool,X24)) | ~p(s(bool,X25))) ) | ($spl8 | $spl122 | $spl133)),
  inference(resolution,[],[f29181,f11977])).
fof(f29181,plain,(
  ( ! [X505,X504,X506,X503] : (sP10('F_0',X503) | sP10(X504,X503) | sP8(X505,X504,X506)) ) | ($spl122 | $spl133)),
  inference(subsumption_resolution,[],[f29091,f528])).
fof(f29091,plain,(
  ( ! [X505,X504,X506,X503] : (p(s(bool,'F_0')) | sP10('F_0',X503) | sP10(X504,X503) | sP8(X505,X504,X506)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f20735,f19369])).
fof(f19369,plain,(
  ( ! [X269,X267,X268,X266] : (sP10(X266,X267) | sP10(X267,'F_0') | sP8(X268,X266,X269)) ) | $spl122),
  inference(resolution,[],[f18883,f410])).
fof(f32499,plain,(
  ( ! [X30,X33,X31,X29,X32] : (sP10('F_0',X29) | sP8(X30,X31,X32) | s(bool,'F_0') = s(bool,X33) | sP10(X33,X29)) ) | ($spl8 | $spl46 | $spl122 | $spl133)),
  inference(subsumption_resolution,[],[f32469,f410])).
fof(f32469,plain,(
  ( ! [X30,X33,X31,X29,X32] : (sP10('F_0',X29) | sP8(X30,X31,X32) | s(bool,'F_0') = s(bool,X33) | ~p(s(bool,X31)) | sP10(X33,X29)) ) | ($spl8 | $spl46 | $spl122 | $spl133)),
  inference(resolution,[],[f29181,f16693])).
fof(f32434,plain,(
  ( ! [X30,X31,X29] : (sP10('F_0',X29) | sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X29) | p(s(bool,X30)) | p(s(bool,X31))) ) | ($spl8 | $spl122 | $spl133)),
  inference(resolution,[],[f29180,f11960])).
fof(f29180,plain,(
  ( ! [X499,X500,X502,X501] : (sP10('F_0',X499) | sP10(X500,X499) | sP6(X500,X501,X502)) ) | ($spl122 | $spl133)),
  inference(subsumption_resolution,[],[f29090,f528])).
fof(f29090,plain,(
  ( ! [X499,X500,X502,X501] : (p(s(bool,'F_0')) | sP10('F_0',X499) | sP10(X500,X499) | sP6(X500,X501,X502)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f20735,f19356])).
fof(f19356,plain,(
  ( ! [X216,X215,X217,X214] : (sP10(X214,X215) | sP10(X215,'F_0') | sP6(X214,X216,X217)) ) | $spl122),
  inference(resolution,[],[f18883,f394])).
fof(f32433,plain,(
  ( ! [X28,X26,X27] : (sP10('F_0',X26) | sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X26) | sP7('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X27,X28)) ) | ($spl8 | $spl122 | $spl133)),
  inference(resolution,[],[f29180,f12008])).
fof(f32442,plain,(
  ( ! [X30,X33,X31,X29,X32] : (sP10('F_0',X29) | sP6(X30,X31,X32) | s(bool,'F_0') = s(bool,X33) | sP10(X33,X29)) ) | ($spl8 | $spl46 | $spl122 | $spl133)),
  inference(subsumption_resolution,[],[f32411,f394])).
fof(f32411,plain,(
  ( ! [X30,X33,X31,X29,X32] : (sP10('F_0',X29) | sP6(X30,X31,X32) | s(bool,'F_0') = s(bool,X33) | ~p(s(bool,X30)) | sP10(X33,X29)) ) | ($spl8 | $spl46 | $spl122 | $spl133)),
  inference(resolution,[],[f29180,f16693])).
fof(f32376,plain,(
  ( ! [X30,X31,X29] : (sP10('F_0',X29) | sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X29) | ~p(s(bool,X30)) | p(s(bool,X31))) ) | ($spl8 | $spl122 | $spl133)),
  inference(resolution,[],[f29179,f11946])).
fof(f29179,plain,(
  ( ! [X496,X498,X495,X497] : (sP10('F_0',X495) | sP10(X496,X495) | sP4(X496,X497,X498)) ) | ($spl122 | $spl133)),
  inference(subsumption_resolution,[],[f29089,f528])).
fof(f29089,plain,(
  ( ! [X496,X498,X495,X497] : (p(s(bool,'F_0')) | sP10('F_0',X495) | sP10(X496,X495) | sP4(X496,X497,X498)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f20735,f19340])).
fof(f19340,plain,(
  ( ! [X152,X151,X153,X150] : (sP10(X150,X151) | sP10(X151,'F_0') | sP4(X150,X152,X153)) ) | $spl122),
  inference(resolution,[],[f18883,f377])).
fof(f32375,plain,(
  ( ! [X28,X26,X27] : (sP10('F_0',X26) | sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X26) | sP5('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X27,X28)) ) | ($spl8 | $spl122 | $spl133)),
  inference(resolution,[],[f29179,f12007])).
fof(f32384,plain,(
  ( ! [X30,X33,X31,X29,X32] : (sP10('F_0',X29) | sP4(X30,X31,X32) | s(bool,'F_0') = s(bool,X33) | sP10(X33,X29)) ) | ($spl8 | $spl46 | $spl122 | $spl133)),
  inference(subsumption_resolution,[],[f32353,f377])).
fof(f32353,plain,(
  ( ! [X30,X33,X31,X29,X32] : (sP10('F_0',X29) | sP4(X30,X31,X32) | s(bool,'F_0') = s(bool,X33) | ~p(s(bool,X30)) | sP10(X33,X29)) ) | ($spl8 | $spl46 | $spl122 | $spl133)),
  inference(resolution,[],[f29179,f16693])).
fof(f32314,plain,(
  ( ! [X21,X22,X20] : (sP10('F_0',X20) | sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X20) | p(s(bool,X21)) | ~p(s(bool,X22))) ) | ($spl8 | $spl122 | $spl133)),
  inference(resolution,[],[f29178,f11947])).
fof(f29178,plain,(
  ( ! [X492,X494,X493,X491] : (sP10('F_0',X491) | sP10(X492,X491) | sP4(X493,X494,X492)) ) | ($spl122 | $spl133)),
  inference(subsumption_resolution,[],[f29088,f528])).
fof(f29088,plain,(
  ( ! [X492,X494,X493,X491] : (p(s(bool,'F_0')) | sP10('F_0',X491) | sP10(X492,X491) | sP4(X493,X494,X492)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f20735,f19339])).
fof(f19339,plain,(
  ( ! [X146,X149,X147,X148] : (sP10(X146,X147) | sP10(X147,'F_0') | sP4(X148,X149,X146)) ) | $spl122),
  inference(resolution,[],[f18883,f375])).
fof(f32325,plain,(
  ( ! [X30,X33,X31,X29,X32] : (sP10('F_0',X29) | sP4(X30,X31,X32) | s(bool,'F_0') = s(bool,X33) | sP10(X33,X29)) ) | ($spl8 | $spl46 | $spl122 | $spl133)),
  inference(subsumption_resolution,[],[f32294,f375])).
fof(f32294,plain,(
  ( ! [X30,X33,X31,X29,X32] : (sP10('F_0',X29) | sP4(X30,X31,X32) | s(bool,'F_0') = s(bool,X33) | ~p(s(bool,X32)) | sP10(X33,X29)) ) | ($spl8 | $spl46 | $spl122 | $spl133)),
  inference(resolution,[],[f29178,f16693])).
fof(f32256,plain,(
  ( ! [X30,X31,X32] : (sP10('F_0',X30) | sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X30) | p(s(bool,X31)) | ~p(s(bool,X32))) ) | ($spl8 | $spl122 | $spl133)),
  inference(resolution,[],[f29177,f11923])).
fof(f29177,plain,(
  ( ! [X488,X490,X487,X489] : (sP10('F_0',X487) | sP10(X488,X487) | sP1(X488,X489,X490)) ) | ($spl122 | $spl133)),
  inference(subsumption_resolution,[],[f29087,f528])).
fof(f29087,plain,(
  ( ! [X488,X490,X487,X489] : (p(s(bool,'F_0')) | sP10('F_0',X487) | sP10(X488,X487) | sP1(X488,X489,X490)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f20735,f19317])).
fof(f19317,plain,(
  ( ! [X74,X72,X75,X73] : (sP10(X72,X73) | sP10(X73,'F_0') | sP1(X72,X74,X75)) ) | $spl122),
  inference(resolution,[],[f18883,f327])).
fof(f32266,plain,(
  ( ! [X19,X17,X18,X16] : (sP10('F_0',X16) | sP3(X18,X17,X19) | p(s(bool,X17)) | ~sP2(X18,X19,X17) | ~sP0(X19,X17,X18)) ) | ($spl122 | $spl133)),
  inference(subsumption_resolution,[],[f32252,f20734])).
fof(f20734,plain,(
  ( ! [X225,X226] : (~sP10(X226,X225) | p(s(bool,X226)) | sP10('F_0',X225)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f20649,f438])).
fof(f32252,plain,(
  ( ! [X19,X17,X18,X16] : (sP10('F_0',X16) | sP10(X17,X16) | sP3(X18,X17,X19) | p(s(bool,X17)) | ~sP2(X18,X19,X17) | ~sP0(X19,X17,X18)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f29177,f318])).
fof(f32264,plain,(
  ( ! [X30,X33,X31,X29,X32] : (sP10('F_0',X29) | sP1(X30,X31,X32) | s(bool,'F_0') = s(bool,X33) | sP10(X33,X29)) ) | ($spl8 | $spl46 | $spl122 | $spl133)),
  inference(subsumption_resolution,[],[f32233,f327])).
fof(f32233,plain,(
  ( ! [X30,X33,X31,X29,X32] : (sP10('F_0',X29) | sP1(X30,X31,X32) | s(bool,'F_0') = s(bool,X33) | ~p(s(bool,X30)) | sP10(X33,X29)) ) | ($spl8 | $spl46 | $spl122 | $spl133)),
  inference(resolution,[],[f29177,f16693])).
fof(f32195,plain,(
  ( ! [X28,X29,X27] : (sP10('F_0',X27) | sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X27) | p(s(bool,X28)) | ~p(s(bool,X29))) ) | ($spl8 | $spl122 | $spl133)),
  inference(resolution,[],[f29176,f11924])).
fof(f29176,plain,(
  ( ! [X484,X486,X485,X483] : (sP10('F_0',X483) | sP10(X484,X483) | sP1(X485,X484,X486)) ) | ($spl122 | $spl133)),
  inference(subsumption_resolution,[],[f29086,f528])).
fof(f29086,plain,(
  ( ! [X484,X486,X485,X483] : (p(s(bool,'F_0')) | sP10('F_0',X483) | sP10(X484,X483) | sP1(X485,X484,X486)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f20735,f19316])).
fof(f19316,plain,(
  ( ! [X70,X68,X71,X69] : (sP10(X68,X69) | sP10(X69,'F_0') | sP1(X70,X68,X71)) ) | $spl122),
  inference(resolution,[],[f18883,f326])).
fof(f32204,plain,(
  ( ! [X30,X33,X31,X29,X32] : (sP10('F_0',X29) | sP1(X30,X31,X32) | s(bool,'F_0') = s(bool,X33) | sP10(X33,X29)) ) | ($spl8 | $spl46 | $spl122 | $spl133)),
  inference(subsumption_resolution,[],[f32173,f326])).
fof(f32173,plain,(
  ( ! [X30,X33,X31,X29,X32] : (sP10('F_0',X29) | sP1(X30,X31,X32) | s(bool,'F_0') = s(bool,X33) | ~p(s(bool,X31)) | sP10(X33,X29)) ) | ($spl8 | $spl46 | $spl122 | $spl133)),
  inference(resolution,[],[f29176,f16693])).
fof(f32135,plain,(
  ( ! [X43,X41,X42] : (sP10('F_0',X41) | sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X41) | p(s(bool,X42)) | ~p(s(bool,X43))) ) | ($spl8 | $spl122 | $spl133)),
  inference(resolution,[],[f29175,f11920])).
fof(f29175,plain,(
  ( ! [X480,X482,X479,X481] : (sP10('F_0',X479) | sP10(X480,X479) | sP2(X480,X481,X482)) ) | ($spl122 | $spl133)),
  inference(subsumption_resolution,[],[f29085,f528])).
fof(f29085,plain,(
  ( ! [X480,X482,X479,X481] : (p(s(bool,'F_0')) | sP10('F_0',X479) | sP10(X480,X479) | sP2(X480,X481,X482)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f20735,f19314])).
fof(f19314,plain,(
  ( ! [X61,X62,X60,X63] : (sP10(X60,X61) | sP10(X61,'F_0') | sP2(X60,X62,X63)) ) | $spl122),
  inference(resolution,[],[f18883,f323])).
fof(f32146,plain,(
  ( ! [X19,X17,X18,X16] : (sP10('F_0',X16) | sP3(X17,X18,X19) | p(s(bool,X17)) | ~sP1(X18,X19,X17) | ~sP0(X19,X18,X17)) ) | ($spl122 | $spl133)),
  inference(subsumption_resolution,[],[f32127,f20734])).
fof(f32127,plain,(
  ( ! [X19,X17,X18,X16] : (sP10('F_0',X16) | sP10(X17,X16) | sP3(X17,X18,X19) | p(s(bool,X17)) | ~sP1(X18,X19,X17) | ~sP0(X19,X18,X17)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f29175,f319])).
fof(f32143,plain,(
  ( ! [X30,X33,X31,X29,X32] : (sP10('F_0',X29) | sP2(X30,X31,X32) | s(bool,'F_0') = s(bool,X33) | sP10(X33,X29)) ) | ($spl8 | $spl46 | $spl122 | $spl133)),
  inference(subsumption_resolution,[],[f32108,f323])).
fof(f32108,plain,(
  ( ! [X30,X33,X31,X29,X32] : (sP10('F_0',X29) | sP2(X30,X31,X32) | s(bool,'F_0') = s(bool,X33) | ~p(s(bool,X30)) | sP10(X33,X29)) ) | ($spl8 | $spl46 | $spl122 | $spl133)),
  inference(resolution,[],[f29175,f16693])).
fof(f32070,plain,(
  ( ! [X39,X38,X40] : (sP10('F_0',X38) | sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X38) | p(s(bool,X39)) | ~p(s(bool,X40))) ) | ($spl8 | $spl122 | $spl133)),
  inference(resolution,[],[f29174,f11921])).
fof(f29174,plain,(
  ( ! [X476,X478,X475,X477] : (sP10('F_0',X475) | sP10(X476,X475) | sP2(X477,X476,X478)) ) | ($spl122 | $spl133)),
  inference(subsumption_resolution,[],[f29084,f528])).
fof(f29084,plain,(
  ( ! [X476,X478,X475,X477] : (p(s(bool,'F_0')) | sP10('F_0',X475) | sP10(X476,X475) | sP2(X477,X476,X478)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f20735,f19313])).
fof(f19313,plain,(
  ( ! [X59,X57,X58,X56] : (sP10(X56,X57) | sP10(X57,'F_0') | sP2(X58,X56,X59)) ) | $spl122),
  inference(resolution,[],[f18883,f322])).
fof(f32081,plain,(
  ( ! [X24,X23] : (sP10('F_0',X23) | s(bool,'T_0') != s(bool,X24) | sP10(X24,X23)) ) | $spl466),
  inference(cnf_transformation,[],[f32081_D])).
fof(f32081_D,plain,(
  ( ! [X24,X23] : (sP10('F_0',X23) | s(bool,'T_0') != s(bool,X24) | sP10(X24,X23)) ) <=> ~$spl466),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl466])])).
fof(f32065,plain,(
  ( ! [X24,X23,X25] : (sP10('F_0',X23) | sP10(X24,X23) | s(bool,'T_0') != s(bool,X24) | ~p(s(bool,X25))) ) | ($spl6 | $spl122 | $spl133)),
  inference(resolution,[],[f29174,f14550])).
fof(f32079,plain,(
  ( ! [X30,X33,X31,X29,X32] : (sP10('F_0',X29) | sP2(X30,X31,X32) | s(bool,'F_0') = s(bool,X33) | sP10(X33,X29)) ) | ($spl8 | $spl46 | $spl122 | $spl133)),
  inference(subsumption_resolution,[],[f32044,f322])).
fof(f32044,plain,(
  ( ! [X30,X33,X31,X29,X32] : (sP10('F_0',X29) | sP2(X30,X31,X32) | s(bool,'F_0') = s(bool,X33) | ~p(s(bool,X31)) | sP10(X33,X29)) ) | ($spl8 | $spl46 | $spl122 | $spl133)),
  inference(resolution,[],[f29174,f16693])).
fof(f28520,plain,(
  ( ! [X107,X106] : (sP11(X106,X107) | sP10(X106,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | $spl334),
  inference(cnf_transformation,[],[f28520_D])).
fof(f28520_D,plain,(
  ( ! [X107,X106] : (sP11(X106,X107) | sP10(X106,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) <=> ~$spl334),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl334])])).
fof(f31930,plain,(
  ( ! [X35,X33,X36,X34] : (sP11(X33,X34) | sP10(X33,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X35)) | ~p(s(bool,X36))) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f28401,f11976])).
fof(f31447,plain,(
  ( ! [X30,X31,X32] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X30) | sP8(X30,X31,X32)) ) | $spl462),
  inference(cnf_transformation,[],[f31447_D])).
fof(f31447_D,plain,(
  ( ! [X30,X31,X32] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X30) | sP8(X30,X31,X32)) ) <=> ~$spl462),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl462])])).
fof(f31900,plain,(
  ( ! [X37,X35,X36,X34] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X34) | sP8(X34,X35,X36) | p(s(bool,X37))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f28401,f14313])).
fof(f14313,plain,(
  ( ! [X158] : (~sP11('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X158) | p(s(bool,X158))) ) | $spl1),
  inference(subsumption_resolution,[],[f14271,f492])).
fof(f14271,plain,(
  ( ! [X158] : (p(s(bool,'F_0')) | p(s(bool,X158)) | ~sP11('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X158)) ) | $spl1),
  inference(resolution,[],[f545,f454])).
fof(f454,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP11(X0,X1)) )),
  inference(cnf_transformation,[],[f261])).
fof(f31932,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP8(X26,X27,X28) | p(s(bool,X25)) | p(s(bool,X29))) ) | $spl122),
  inference(subsumption_resolution,[],[f31898,f492])).
fof(f31898,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP8(X26,X27,X28) | p(s(bool,'F_0')) | p(s(bool,X25)) | p(s(bool,X29))) ) | $spl122),
  inference(resolution,[],[f28401,f454])).
fof(f31931,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP8(X21,X22,X23) | p(s(bool,X20)) | p(s(bool,X24))) ) | $spl122),
  inference(subsumption_resolution,[],[f31897,f492])).
fof(f31897,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP8(X21,X22,X23) | p(s(bool,'F_0')) | p(s(bool,X20)) | p(s(bool,X24))) ) | $spl122),
  inference(resolution,[],[f28401,f458])).
fof(f458,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP11(X0,X1)) )),
  inference(cnf_transformation,[],[f263])).
fof(f263,plain,(
  ! [X0,X1] : ((~sP11(X0,X1) | (p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP11(X0,X1)))),
  inference(flattening,[],[f262])).
fof(f262,plain,(
  ! [X0,X1] : ((~sP11(X0,X1) | ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0')))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP11(X0,X1)))),
  inference(nnf_transformation,[],[f191])).
fof(f191,plain,(
  ! [X0,X1] : (sP11(X0,X1) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(definition_folding,[],[f160,f190])).
fof(f160,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f159])).
fof(f159,plain,(
  ! [X0,X1] : (((p(s(bool,X1)) | p(s(bool,X0))) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | (p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(ennf_transformation,[],[f105])).
fof(f105,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(flattening,[],[f104])).
fof(f104,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(rectify,[],[f53])).
fof(f53,axiom,(
  ! [X16,X6] : ((~(p(s(bool,X6)) | p(s(bool,X16))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X6)) => p(s(bool,'F_0'))) => (~p(s(bool,X16)) => p(s(bool,'F_0')))))),
  file('thm.list.MEM_GENLIST','thm.sat.sat_9_prove_ltb')).
fof(f19849,plain,(
  ( ! [X24] : (sP10(X24,'F_0')) ) | $spl150),
  inference(cnf_transformation,[],[f19849_D])).
fof(f19849_D,plain,(
  ( ! [X24] : (sP10(X24,'F_0')) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f31896,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP10(X15,X16) | sP8(X16,X17,X18) | p(s(bool,X15)) | sP10(X19,'F_0')) ) | $spl122),
  inference(resolution,[],[f28401,f18976])).
fof(f18976,plain,(
  ( ! [X241,X242] : (~sP11(X242,X241) | p(s(bool,X242)) | sP10(X241,'F_0')) ) | $spl122),
  inference(subsumption_resolution,[],[f18889,f492])).
fof(f18889,plain,(
  ( ! [X241,X242] : (sP10(X241,'F_0') | p(s(bool,'F_0')) | p(s(bool,X242)) | ~sP11(X242,X241)) ) | $spl122),
  inference(resolution,[],[f18767,f454])).
fof(f31894,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP10(X5,X6) | sP8(X6,X7,X8) | p(s(bool,X5)) | sP10('F_0',X9)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f28401,f20831])).
fof(f20831,plain,(
  ( ! [X241,X242] : (~sP11(X242,X241) | p(s(bool,X242)) | sP10('F_0',X241)) ) | ($spl122 | $spl133)),
  inference(subsumption_resolution,[],[f20742,f492])).
fof(f20742,plain,(
  ( ! [X241,X242] : (sP10('F_0',X241) | p(s(bool,'F_0')) | p(s(bool,X242)) | ~sP11(X242,X241)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f20649,f454])).
fof(f31444,plain,(
  ( ! [X6,X8,X7,X5] : (sP10(X5,X6) | sP10('F_0',X5) | sP8(X6,X7,X8)) ) | $spl460),
  inference(cnf_transformation,[],[f31444_D])).
fof(f31444_D,plain,(
  ( ! [X6,X8,X7,X5] : (sP10(X5,X6) | sP10('F_0',X5) | sP8(X6,X7,X8)) ) <=> ~$spl460),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl460])])).
fof(f31893,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10(X0,X1) | sP8(X1,X2,X3) | p(s(bool,X4)) | sP10('F_0',X0)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f28401,f20832])).
fof(f20832,plain,(
  ( ! [X243,X244] : (~sP11(X243,X244) | p(s(bool,X244)) | sP10('F_0',X243)) ) | ($spl122 | $spl133)),
  inference(subsumption_resolution,[],[f20743,f492])).
fof(f20743,plain,(
  ( ! [X243,X244] : (sP10('F_0',X243) | p(s(bool,'F_0')) | p(s(bool,X244)) | ~sP11(X243,X244)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f20649,f454])).
fof(f31887,plain,(
  ( ! [X30,X31,X29,X32] : (sP11(X29,X30) | sP10(X29,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X31)) | ~p(s(bool,X32))) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f28400,f11977])).
fof(f31402,plain,(
  ( ! [X30,X31,X32] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X30) | sP8(X31,X30,X32)) ) | $spl456),
  inference(cnf_transformation,[],[f31402_D])).
fof(f31402_D,plain,(
  ( ! [X30,X31,X32] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X30) | sP8(X31,X30,X32)) ) <=> ~$spl456),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl456])])).
fof(f31858,plain,(
  ( ! [X37,X35,X36,X34] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X34) | sP8(X35,X34,X36) | p(s(bool,X37))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f28400,f14313])).
fof(f31890,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP8(X27,X26,X28) | p(s(bool,X25)) | p(s(bool,X29))) ) | $spl122),
  inference(subsumption_resolution,[],[f31856,f492])).
fof(f31856,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP8(X27,X26,X28) | p(s(bool,'F_0')) | p(s(bool,X25)) | p(s(bool,X29))) ) | $spl122),
  inference(resolution,[],[f28400,f454])).
fof(f31889,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP8(X22,X21,X23) | p(s(bool,X20)) | p(s(bool,X24))) ) | $spl122),
  inference(subsumption_resolution,[],[f31855,f492])).
fof(f31855,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP8(X22,X21,X23) | p(s(bool,'F_0')) | p(s(bool,X20)) | p(s(bool,X24))) ) | $spl122),
  inference(resolution,[],[f28400,f458])).
fof(f31854,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP10(X15,X16) | sP8(X17,X16,X18) | p(s(bool,X15)) | sP10(X19,'F_0')) ) | $spl122),
  inference(resolution,[],[f28400,f18976])).
fof(f31852,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP10(X5,X6) | sP8(X7,X6,X8) | p(s(bool,X5)) | sP10('F_0',X9)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f28400,f20831])).
fof(f31399,plain,(
  ( ! [X6,X8,X7,X5] : (sP10(X5,X6) | sP10('F_0',X5) | sP8(X7,X6,X8)) ) | $spl454),
  inference(cnf_transformation,[],[f31399_D])).
fof(f31399_D,plain,(
  ( ! [X6,X8,X7,X5] : (sP10(X5,X6) | sP10('F_0',X5) | sP8(X7,X6,X8)) ) <=> ~$spl454),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl454])])).
fof(f31851,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10(X0,X1) | sP8(X2,X1,X3) | p(s(bool,X4)) | sP10('F_0',X0)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f28400,f20832])).
fof(f31848,plain,(
  ( ! [X39,X37,X38,X40] : (sP11(X37,X38) | sP10(X37,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X39)) | p(s(bool,X40))) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f28387,f11960])).
fof(f31847,plain,(
  ( ! [X35,X33,X36,X34] : (sP11(X33,X34) | sP10(X33,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | sP7('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X35,X36)) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f28387,f12008])).
fof(f31359,plain,(
  ( ! [X30,X31,X32] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X30) | sP6(X30,X31,X32)) ) | $spl450),
  inference(cnf_transformation,[],[f31359_D])).
fof(f31359_D,plain,(
  ( ! [X30,X31,X32] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X30) | sP6(X30,X31,X32)) ) <=> ~$spl450),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl450])])).
fof(f31817,plain,(
  ( ! [X37,X35,X36,X34] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X34) | sP6(X34,X35,X36) | p(s(bool,X37))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f28387,f14313])).
fof(f31850,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP6(X26,X27,X28) | p(s(bool,X25)) | p(s(bool,X29))) ) | $spl122),
  inference(subsumption_resolution,[],[f31815,f492])).
fof(f31815,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP6(X26,X27,X28) | p(s(bool,'F_0')) | p(s(bool,X25)) | p(s(bool,X29))) ) | $spl122),
  inference(resolution,[],[f28387,f454])).
fof(f31849,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP6(X21,X22,X23) | p(s(bool,X20)) | p(s(bool,X24))) ) | $spl122),
  inference(subsumption_resolution,[],[f31814,f492])).
fof(f31814,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP6(X21,X22,X23) | p(s(bool,'F_0')) | p(s(bool,X20)) | p(s(bool,X24))) ) | $spl122),
  inference(resolution,[],[f28387,f458])).
fof(f31813,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP10(X15,X16) | sP6(X16,X17,X18) | p(s(bool,X15)) | sP10(X19,'F_0')) ) | $spl122),
  inference(resolution,[],[f28387,f18976])).
fof(f31811,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP10(X5,X6) | sP6(X6,X7,X8) | p(s(bool,X5)) | sP10('F_0',X9)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f28387,f20831])).
fof(f31356,plain,(
  ( ! [X6,X8,X7,X5] : (sP10(X5,X6) | sP10('F_0',X5) | sP6(X6,X7,X8)) ) | $spl448),
  inference(cnf_transformation,[],[f31356_D])).
fof(f31356_D,plain,(
  ( ! [X6,X8,X7,X5] : (sP10(X5,X6) | sP10('F_0',X5) | sP6(X6,X7,X8)) ) <=> ~$spl448),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl448])])).
fof(f31810,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10(X0,X1) | sP6(X1,X2,X3) | p(s(bool,X4)) | sP10('F_0',X0)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f28387,f20832])).
fof(f31807,plain,(
  ( ! [X39,X37,X38,X40] : (sP11(X37,X38) | sP10(X37,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | ~p(s(bool,X39)) | p(s(bool,X40))) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f28371,f11946])).
fof(f31806,plain,(
  ( ! [X35,X33,X36,X34] : (sP11(X33,X34) | sP10(X33,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | sP5('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X35,X36)) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f28371,f12007])).
fof(f31315,plain,(
  ( ! [X30,X31,X32] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X30) | sP4(X30,X31,X32)) ) | $spl444),
  inference(cnf_transformation,[],[f31315_D])).
fof(f31315_D,plain,(
  ( ! [X30,X31,X32] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X30) | sP4(X30,X31,X32)) ) <=> ~$spl444),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl444])])).
fof(f31776,plain,(
  ( ! [X37,X35,X36,X34] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X34) | sP4(X34,X35,X36) | p(s(bool,X37))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f28371,f14313])).
fof(f31809,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP4(X26,X27,X28) | p(s(bool,X25)) | p(s(bool,X29))) ) | $spl122),
  inference(subsumption_resolution,[],[f31774,f492])).
fof(f31774,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP4(X26,X27,X28) | p(s(bool,'F_0')) | p(s(bool,X25)) | p(s(bool,X29))) ) | $spl122),
  inference(resolution,[],[f28371,f454])).
fof(f31808,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP4(X21,X22,X23) | p(s(bool,X20)) | p(s(bool,X24))) ) | $spl122),
  inference(subsumption_resolution,[],[f31773,f492])).
fof(f31773,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP4(X21,X22,X23) | p(s(bool,'F_0')) | p(s(bool,X20)) | p(s(bool,X24))) ) | $spl122),
  inference(resolution,[],[f28371,f458])).
fof(f31772,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP10(X15,X16) | sP4(X16,X17,X18) | p(s(bool,X15)) | sP10(X19,'F_0')) ) | $spl122),
  inference(resolution,[],[f28371,f18976])).
fof(f31770,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP10(X5,X6) | sP4(X6,X7,X8) | p(s(bool,X5)) | sP10('F_0',X9)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f28371,f20831])).
fof(f31312,plain,(
  ( ! [X6,X8,X7,X5] : (sP10(X5,X6) | sP10('F_0',X5) | sP4(X6,X7,X8)) ) | $spl442),
  inference(cnf_transformation,[],[f31312_D])).
fof(f31312_D,plain,(
  ( ! [X6,X8,X7,X5] : (sP10(X5,X6) | sP10('F_0',X5) | sP4(X6,X7,X8)) ) <=> ~$spl442),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl442])])).
fof(f31769,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10(X0,X1) | sP4(X1,X2,X3) | p(s(bool,X4)) | sP10('F_0',X0)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f28371,f20832])).
fof(f31725,plain,(
  ( ! [X28,X26,X27,X25] : (sP11(X25,X26) | sP10(X25,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X27)) | ~p(s(bool,X28))) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f28370,f11947])).
fof(f31270,plain,(
  ( ! [X30,X31,X32] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X30) | sP4(X31,X32,X30)) ) | $spl438),
  inference(cnf_transformation,[],[f31270_D])).
fof(f31270_D,plain,(
  ( ! [X30,X31,X32] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X30) | sP4(X31,X32,X30)) ) <=> ~$spl438),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl438])])).
fof(f31697,plain,(
  ( ! [X37,X35,X36,X34] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X34) | sP4(X35,X36,X34) | p(s(bool,X37))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f28370,f14313])).
fof(f31730,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP4(X27,X28,X26) | p(s(bool,X25)) | p(s(bool,X29))) ) | $spl122),
  inference(subsumption_resolution,[],[f31695,f492])).
fof(f31695,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP4(X27,X28,X26) | p(s(bool,'F_0')) | p(s(bool,X25)) | p(s(bool,X29))) ) | $spl122),
  inference(resolution,[],[f28370,f454])).
fof(f31729,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP4(X22,X23,X21) | p(s(bool,X20)) | p(s(bool,X24))) ) | $spl122),
  inference(subsumption_resolution,[],[f31694,f492])).
fof(f31694,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP4(X22,X23,X21) | p(s(bool,'F_0')) | p(s(bool,X20)) | p(s(bool,X24))) ) | $spl122),
  inference(resolution,[],[f28370,f458])).
fof(f31693,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP10(X15,X16) | sP4(X17,X18,X16) | p(s(bool,X15)) | sP10(X19,'F_0')) ) | $spl122),
  inference(resolution,[],[f28370,f18976])).
fof(f31691,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP10(X5,X6) | sP4(X7,X8,X6) | p(s(bool,X5)) | sP10('F_0',X9)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f28370,f20831])).
fof(f31267,plain,(
  ( ! [X6,X8,X7,X5] : (sP10(X5,X6) | sP10('F_0',X5) | sP4(X7,X8,X6)) ) | $spl436),
  inference(cnf_transformation,[],[f31267_D])).
fof(f31267_D,plain,(
  ( ! [X6,X8,X7,X5] : (sP10(X5,X6) | sP10('F_0',X5) | sP4(X7,X8,X6)) ) <=> ~$spl436),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl436])])).
fof(f31690,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10(X0,X1) | sP4(X2,X3,X1) | p(s(bool,X4)) | sP10('F_0',X0)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f28370,f20832])).
fof(f31684,plain,(
  ( ! [X39,X41,X38,X40] : (sP11(X38,X39) | sP10(X38,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X40)) | ~p(s(bool,X41))) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f28348,f11923])).
fof(f31222,plain,(
  ( ! [X30,X31,X32] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X30) | sP1(X30,X31,X32)) ) | $spl432),
  inference(cnf_transformation,[],[f31222_D])).
fof(f31222_D,plain,(
  ( ! [X30,X31,X32] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X30) | sP1(X30,X31,X32)) ) <=> ~$spl432),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl432])])).
fof(f31653,plain,(
  ( ! [X37,X35,X36,X34] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X34) | sP1(X34,X35,X36) | p(s(bool,X37))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f28348,f14313])).
fof(f31686,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP1(X26,X27,X28) | p(s(bool,X25)) | p(s(bool,X29))) ) | $spl122),
  inference(subsumption_resolution,[],[f31651,f492])).
fof(f31651,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP1(X26,X27,X28) | p(s(bool,'F_0')) | p(s(bool,X25)) | p(s(bool,X29))) ) | $spl122),
  inference(resolution,[],[f28348,f454])).
fof(f31685,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP1(X21,X22,X23) | p(s(bool,X20)) | p(s(bool,X24))) ) | $spl122),
  inference(subsumption_resolution,[],[f31650,f492])).
fof(f31650,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP1(X21,X22,X23) | p(s(bool,'F_0')) | p(s(bool,X20)) | p(s(bool,X24))) ) | $spl122),
  inference(resolution,[],[f28348,f458])).
fof(f31649,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP10(X15,X16) | sP1(X16,X17,X18) | p(s(bool,X15)) | sP10(X19,'F_0')) ) | $spl122),
  inference(resolution,[],[f28348,f18976])).
fof(f31647,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP10(X5,X6) | sP1(X6,X7,X8) | p(s(bool,X5)) | sP10('F_0',X9)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f28348,f20831])).
fof(f31219,plain,(
  ( ! [X6,X8,X7,X5] : (sP10(X5,X6) | sP10('F_0',X5) | sP1(X6,X7,X8)) ) | $spl430),
  inference(cnf_transformation,[],[f31219_D])).
fof(f31219_D,plain,(
  ( ! [X6,X8,X7,X5] : (sP10(X5,X6) | sP10('F_0',X5) | sP1(X6,X7,X8)) ) <=> ~$spl430),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl430])])).
fof(f31646,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10(X0,X1) | sP1(X1,X2,X3) | p(s(bool,X4)) | sP10('F_0',X0)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f28348,f20832])).
fof(f31640,plain,(
  ( ! [X37,X35,X36,X34] : (sP11(X34,X35) | sP10(X34,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X36)) | ~p(s(bool,X37))) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f28347,f11924])).
fof(f31176,plain,(
  ( ! [X30,X31,X32] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X30) | sP1(X31,X30,X32)) ) | $spl426),
  inference(cnf_transformation,[],[f31176_D])).
fof(f31176_D,plain,(
  ( ! [X30,X31,X32] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X30) | sP1(X31,X30,X32)) ) <=> ~$spl426),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl426])])).
fof(f31610,plain,(
  ( ! [X37,X35,X36,X34] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X34) | sP1(X35,X34,X36) | p(s(bool,X37))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f28347,f14313])).
fof(f31643,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP1(X27,X26,X28) | p(s(bool,X25)) | p(s(bool,X29))) ) | $spl122),
  inference(subsumption_resolution,[],[f31608,f492])).
fof(f31608,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP1(X27,X26,X28) | p(s(bool,'F_0')) | p(s(bool,X25)) | p(s(bool,X29))) ) | $spl122),
  inference(resolution,[],[f28347,f454])).
fof(f31642,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP1(X22,X21,X23) | p(s(bool,X20)) | p(s(bool,X24))) ) | $spl122),
  inference(subsumption_resolution,[],[f31607,f492])).
fof(f31607,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP1(X22,X21,X23) | p(s(bool,'F_0')) | p(s(bool,X20)) | p(s(bool,X24))) ) | $spl122),
  inference(resolution,[],[f28347,f458])).
fof(f31606,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP10(X15,X16) | sP1(X17,X16,X18) | p(s(bool,X15)) | sP10(X19,'F_0')) ) | $spl122),
  inference(resolution,[],[f28347,f18976])).
fof(f31604,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP10(X5,X6) | sP1(X7,X6,X8) | p(s(bool,X5)) | sP10('F_0',X9)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f28347,f20831])).
fof(f31173,plain,(
  ( ! [X6,X8,X7,X5] : (sP10(X5,X6) | sP10('F_0',X5) | sP1(X7,X6,X8)) ) | $spl424),
  inference(cnf_transformation,[],[f31173_D])).
fof(f31173_D,plain,(
  ( ! [X6,X8,X7,X5] : (sP10(X5,X6) | sP10('F_0',X5) | sP1(X7,X6,X8)) ) <=> ~$spl424),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl424])])).
fof(f31603,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10(X0,X1) | sP1(X2,X1,X3) | p(s(bool,X4)) | sP10('F_0',X0)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f28347,f20832])).
fof(f31597,plain,(
  ( ! [X54,X56,X55,X53] : (sP11(X53,X54) | sP10(X53,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X55)) | ~p(s(bool,X56))) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f28345,f11920])).
fof(f31129,plain,(
  ( ! [X30,X31,X32] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X30) | sP2(X30,X31,X32)) ) | $spl420),
  inference(cnf_transformation,[],[f31129_D])).
fof(f31129_D,plain,(
  ( ! [X30,X31,X32] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X30) | sP2(X30,X31,X32)) ) <=> ~$spl420),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl420])])).
fof(f31562,plain,(
  ( ! [X37,X35,X36,X34] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X34) | sP2(X34,X35,X36) | p(s(bool,X37))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f28345,f14313])).
fof(f31599,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP2(X26,X27,X28) | p(s(bool,X25)) | p(s(bool,X29))) ) | $spl122),
  inference(subsumption_resolution,[],[f31560,f492])).
fof(f31560,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP2(X26,X27,X28) | p(s(bool,'F_0')) | p(s(bool,X25)) | p(s(bool,X29))) ) | $spl122),
  inference(resolution,[],[f28345,f454])).
fof(f31598,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP2(X21,X22,X23) | p(s(bool,X20)) | p(s(bool,X24))) ) | $spl122),
  inference(subsumption_resolution,[],[f31559,f492])).
fof(f31559,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP2(X21,X22,X23) | p(s(bool,'F_0')) | p(s(bool,X20)) | p(s(bool,X24))) ) | $spl122),
  inference(resolution,[],[f28345,f458])).
fof(f31558,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP10(X15,X16) | sP2(X16,X17,X18) | p(s(bool,X15)) | sP10(X19,'F_0')) ) | $spl122),
  inference(resolution,[],[f28345,f18976])).
fof(f31556,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP10(X5,X6) | sP2(X6,X7,X8) | p(s(bool,X5)) | sP10('F_0',X9)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f28345,f20831])).
fof(f31126,plain,(
  ( ! [X6,X8,X7,X5] : (sP10(X5,X6) | sP10('F_0',X5) | sP2(X6,X7,X8)) ) | $spl418),
  inference(cnf_transformation,[],[f31126_D])).
fof(f31126_D,plain,(
  ( ! [X6,X8,X7,X5] : (sP10(X5,X6) | sP10('F_0',X5) | sP2(X6,X7,X8)) ) <=> ~$spl418),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl418])])).
fof(f31555,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10(X0,X1) | sP2(X1,X2,X3) | p(s(bool,X4)) | sP10('F_0',X0)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f28345,f20832])).
fof(f31549,plain,(
  ( ! [X52,X50,X51,X49] : (sP11(X49,X50) | sP10(X49,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X51)) | ~p(s(bool,X52))) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f28344,f11921])).
fof(f31554,plain,(
  ( ! [X30,X31,X29] : (sP11(X29,X30) | s(bool,'T_0') != s(bool,X31) | sP10(X29,X31)) ) | $spl464),
  inference(cnf_transformation,[],[f31554_D])).
fof(f31554_D,plain,(
  ( ! [X30,X31,X29] : (sP11(X29,X30) | s(bool,'T_0') != s(bool,X31) | sP10(X29,X31)) ) <=> ~$spl464),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl464])])).
fof(f31544,plain,(
  ( ! [X30,X31,X29,X32] : (sP11(X29,X30) | sP10(X29,X31) | s(bool,'T_0') != s(bool,X31) | ~p(s(bool,X32))) ) | ($spl6 | $spl122)),
  inference(resolution,[],[f28344,f14550])).
fof(f31079,plain,(
  ( ! [X30,X31,X32] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X30) | sP2(X31,X30,X32)) ) | $spl412),
  inference(cnf_transformation,[],[f31079_D])).
fof(f31079_D,plain,(
  ( ! [X30,X31,X32] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X30) | sP2(X31,X30,X32)) ) <=> ~$spl412),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl412])])).
fof(f31515,plain,(
  ( ! [X37,X35,X36,X34] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X34) | sP2(X35,X34,X36) | p(s(bool,X37))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f28344,f14313])).
fof(f31552,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP2(X27,X26,X28) | p(s(bool,X25)) | p(s(bool,X29))) ) | $spl122),
  inference(subsumption_resolution,[],[f31513,f492])).
fof(f31513,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP2(X27,X26,X28) | p(s(bool,'F_0')) | p(s(bool,X25)) | p(s(bool,X29))) ) | $spl122),
  inference(resolution,[],[f28344,f454])).
fof(f31551,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP2(X22,X21,X23) | p(s(bool,X20)) | p(s(bool,X24))) ) | $spl122),
  inference(subsumption_resolution,[],[f31512,f492])).
fof(f31512,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP2(X22,X21,X23) | p(s(bool,'F_0')) | p(s(bool,X20)) | p(s(bool,X24))) ) | $spl122),
  inference(resolution,[],[f28344,f458])).
fof(f31511,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP10(X15,X16) | sP2(X17,X16,X18) | p(s(bool,X15)) | sP10(X19,'F_0')) ) | $spl122),
  inference(resolution,[],[f28344,f18976])).
fof(f31509,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP10(X5,X6) | sP2(X7,X6,X8) | p(s(bool,X5)) | sP10('F_0',X9)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f28344,f20831])).
fof(f31076,plain,(
  ( ! [X6,X8,X7,X5] : (sP10(X5,X6) | sP10('F_0',X5) | sP2(X7,X6,X8)) ) | $spl410),
  inference(cnf_transformation,[],[f31076_D])).
fof(f31076_D,plain,(
  ( ! [X6,X8,X7,X5] : (sP10(X5,X6) | sP10('F_0',X5) | sP2(X7,X6,X8)) ) <=> ~$spl410),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl410])])).
fof(f31508,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10(X0,X1) | sP2(X2,X1,X3) | p(s(bool,X4)) | sP10('F_0',X0)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f28344,f20832])).
fof(f28286,plain,(
  ( ! [X107,X106] : (sP11(X106,X107) | sP10(X107,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | $spl328),
  inference(cnf_transformation,[],[f28286_D])).
fof(f28286_D,plain,(
  ( ! [X107,X106] : (sP11(X106,X107) | sP10(X107,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) <=> ~$spl328),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl328])])).
fof(f31442,plain,(
  ( ! [X35,X33,X36,X34] : (sP11(X33,X34) | sP10(X34,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X35)) | ~p(s(bool,X36))) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f28167,f11976])).
fof(f31411,plain,(
  ( ! [X30,X33,X31,X32] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X30) | sP8(X30,X31,X32) | p(s(bool,X33))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f28167,f14312])).
fof(f14312,plain,(
  ( ! [X157] : (~sP11(X157,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | p(s(bool,X157))) ) | $spl1),
  inference(subsumption_resolution,[],[f14270,f492])).
fof(f14270,plain,(
  ( ! [X157] : (p(s(bool,'F_0')) | p(s(bool,X157)) | ~sP11(X157,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) ) | $spl1),
  inference(resolution,[],[f545,f454])).
fof(f31446,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP8(X26,X27,X28) | p(s(bool,X29)) | p(s(bool,X25))) ) | $spl122),
  inference(subsumption_resolution,[],[f31410,f492])).
fof(f31410,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP8(X26,X27,X28) | p(s(bool,'F_0')) | p(s(bool,X29)) | p(s(bool,X25))) ) | $spl122),
  inference(resolution,[],[f28167,f454])).
fof(f31445,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP8(X21,X22,X23) | p(s(bool,X24)) | p(s(bool,X20))) ) | $spl122),
  inference(subsumption_resolution,[],[f31409,f492])).
fof(f31409,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP8(X21,X22,X23) | p(s(bool,'F_0')) | p(s(bool,X24)) | p(s(bool,X20))) ) | $spl122),
  inference(resolution,[],[f28167,f458])).
fof(f31407,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP10(X10,X11) | sP8(X11,X12,X13) | p(s(bool,X10)) | sP10(X14,'F_0')) ) | $spl122),
  inference(resolution,[],[f28167,f18977])).
fof(f18977,plain,(
  ( ! [X243,X244] : (~sP11(X243,X244) | p(s(bool,X244)) | sP10(X243,'F_0')) ) | $spl122),
  inference(subsumption_resolution,[],[f18890,f492])).
fof(f18890,plain,(
  ( ! [X243,X244] : (sP10(X243,'F_0') | p(s(bool,'F_0')) | p(s(bool,X244)) | ~sP11(X243,X244)) ) | $spl122),
  inference(resolution,[],[f18767,f454])).
fof(f31406,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP10(X5,X6) | sP8(X6,X7,X8) | p(s(bool,X9)) | sP10('F_0',X5)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f28167,f20831])).
fof(f31405,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10(X0,X1) | sP8(X1,X2,X3) | p(s(bool,X0)) | sP10('F_0',X4)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f28167,f20832])).
fof(f31396,plain,(
  ( ! [X30,X31,X29,X32] : (sP11(X29,X30) | sP10(X30,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X31)) | ~p(s(bool,X32))) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f28166,f11977])).
fof(f31366,plain,(
  ( ! [X30,X33,X31,X32] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X30) | sP8(X31,X30,X32) | p(s(bool,X33))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f28166,f14312])).
fof(f31401,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP8(X27,X26,X28) | p(s(bool,X29)) | p(s(bool,X25))) ) | $spl122),
  inference(subsumption_resolution,[],[f31365,f492])).
fof(f31365,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP8(X27,X26,X28) | p(s(bool,'F_0')) | p(s(bool,X29)) | p(s(bool,X25))) ) | $spl122),
  inference(resolution,[],[f28166,f454])).
fof(f31400,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP8(X22,X21,X23) | p(s(bool,X24)) | p(s(bool,X20))) ) | $spl122),
  inference(subsumption_resolution,[],[f31364,f492])).
fof(f31364,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP8(X22,X21,X23) | p(s(bool,'F_0')) | p(s(bool,X24)) | p(s(bool,X20))) ) | $spl122),
  inference(resolution,[],[f28166,f458])).
fof(f31362,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP10(X10,X11) | sP8(X12,X11,X13) | p(s(bool,X10)) | sP10(X14,'F_0')) ) | $spl122),
  inference(resolution,[],[f28166,f18977])).
fof(f31361,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP10(X5,X6) | sP8(X7,X6,X8) | p(s(bool,X9)) | sP10('F_0',X5)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f28166,f20831])).
fof(f31360,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10(X0,X1) | sP8(X2,X1,X3) | p(s(bool,X0)) | sP10('F_0',X4)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f28166,f20832])).
fof(f31354,plain,(
  ( ! [X39,X37,X38,X40] : (sP11(X37,X38) | sP10(X38,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X39)) | p(s(bool,X40))) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f28153,f11960])).
fof(f31353,plain,(
  ( ! [X35,X33,X36,X34] : (sP11(X33,X34) | sP10(X34,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | sP7('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X35,X36)) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f28153,f12008])).
fof(f31322,plain,(
  ( ! [X30,X33,X31,X32] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X30) | sP6(X30,X31,X32) | p(s(bool,X33))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f28153,f14312])).
fof(f31358,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP6(X26,X27,X28) | p(s(bool,X29)) | p(s(bool,X25))) ) | $spl122),
  inference(subsumption_resolution,[],[f31321,f492])).
fof(f31321,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP6(X26,X27,X28) | p(s(bool,'F_0')) | p(s(bool,X29)) | p(s(bool,X25))) ) | $spl122),
  inference(resolution,[],[f28153,f454])).
fof(f31357,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP6(X21,X22,X23) | p(s(bool,X24)) | p(s(bool,X20))) ) | $spl122),
  inference(subsumption_resolution,[],[f31320,f492])).
fof(f31320,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP6(X21,X22,X23) | p(s(bool,'F_0')) | p(s(bool,X24)) | p(s(bool,X20))) ) | $spl122),
  inference(resolution,[],[f28153,f458])).
fof(f31318,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP10(X10,X11) | sP6(X11,X12,X13) | p(s(bool,X10)) | sP10(X14,'F_0')) ) | $spl122),
  inference(resolution,[],[f28153,f18977])).
fof(f31317,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP10(X5,X6) | sP6(X6,X7,X8) | p(s(bool,X9)) | sP10('F_0',X5)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f28153,f20831])).
fof(f31316,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10(X0,X1) | sP6(X1,X2,X3) | p(s(bool,X0)) | sP10('F_0',X4)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f28153,f20832])).
fof(f31310,plain,(
  ( ! [X39,X37,X38,X40] : (sP11(X37,X38) | sP10(X38,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | ~p(s(bool,X39)) | p(s(bool,X40))) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f28137,f11946])).
fof(f31309,plain,(
  ( ! [X35,X33,X36,X34] : (sP11(X33,X34) | sP10(X34,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | sP5('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X35,X36)) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f28137,f12007])).
fof(f31278,plain,(
  ( ! [X30,X33,X31,X32] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X30) | sP4(X30,X31,X32) | p(s(bool,X33))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f28137,f14312])).
fof(f31314,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP4(X26,X27,X28) | p(s(bool,X29)) | p(s(bool,X25))) ) | $spl122),
  inference(subsumption_resolution,[],[f31277,f492])).
fof(f31277,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP4(X26,X27,X28) | p(s(bool,'F_0')) | p(s(bool,X29)) | p(s(bool,X25))) ) | $spl122),
  inference(resolution,[],[f28137,f454])).
fof(f31313,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP4(X21,X22,X23) | p(s(bool,X24)) | p(s(bool,X20))) ) | $spl122),
  inference(subsumption_resolution,[],[f31276,f492])).
fof(f31276,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP4(X21,X22,X23) | p(s(bool,'F_0')) | p(s(bool,X24)) | p(s(bool,X20))) ) | $spl122),
  inference(resolution,[],[f28137,f458])).
fof(f31274,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP10(X10,X11) | sP4(X11,X12,X13) | p(s(bool,X10)) | sP10(X14,'F_0')) ) | $spl122),
  inference(resolution,[],[f28137,f18977])).
fof(f31273,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP10(X5,X6) | sP4(X6,X7,X8) | p(s(bool,X9)) | sP10('F_0',X5)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f28137,f20831])).
fof(f31272,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10(X0,X1) | sP4(X1,X2,X3) | p(s(bool,X0)) | sP10('F_0',X4)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f28137,f20832])).
fof(f31262,plain,(
  ( ! [X28,X26,X27,X25] : (sP11(X25,X26) | sP10(X26,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X27)) | ~p(s(bool,X28))) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f28136,f11947])).
fof(f31233,plain,(
  ( ! [X30,X33,X31,X32] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X30) | sP4(X31,X32,X30) | p(s(bool,X33))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f28136,f14312])).
fof(f31269,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP4(X27,X28,X26) | p(s(bool,X29)) | p(s(bool,X25))) ) | $spl122),
  inference(subsumption_resolution,[],[f31232,f492])).
fof(f31232,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP4(X27,X28,X26) | p(s(bool,'F_0')) | p(s(bool,X29)) | p(s(bool,X25))) ) | $spl122),
  inference(resolution,[],[f28136,f454])).
fof(f31268,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP4(X22,X23,X21) | p(s(bool,X24)) | p(s(bool,X20))) ) | $spl122),
  inference(subsumption_resolution,[],[f31231,f492])).
fof(f31231,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP4(X22,X23,X21) | p(s(bool,'F_0')) | p(s(bool,X24)) | p(s(bool,X20))) ) | $spl122),
  inference(resolution,[],[f28136,f458])).
fof(f31229,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP10(X10,X11) | sP4(X12,X13,X11) | p(s(bool,X10)) | sP10(X14,'F_0')) ) | $spl122),
  inference(resolution,[],[f28136,f18977])).
fof(f31228,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP10(X5,X6) | sP4(X7,X8,X6) | p(s(bool,X9)) | sP10('F_0',X5)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f28136,f20831])).
fof(f31227,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10(X0,X1) | sP4(X2,X3,X1) | p(s(bool,X0)) | sP10('F_0',X4)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f28136,f20832])).
fof(f31217,plain,(
  ( ! [X39,X41,X38,X40] : (sP11(X38,X39) | sP10(X39,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X40)) | ~p(s(bool,X41))) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f28114,f11923])).
fof(f31185,plain,(
  ( ! [X30,X33,X31,X32] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X30) | sP1(X30,X31,X32) | p(s(bool,X33))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f28114,f14312])).
fof(f31221,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP1(X26,X27,X28) | p(s(bool,X29)) | p(s(bool,X25))) ) | $spl122),
  inference(subsumption_resolution,[],[f31184,f492])).
fof(f31184,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP1(X26,X27,X28) | p(s(bool,'F_0')) | p(s(bool,X29)) | p(s(bool,X25))) ) | $spl122),
  inference(resolution,[],[f28114,f454])).
fof(f31220,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP1(X21,X22,X23) | p(s(bool,X24)) | p(s(bool,X20))) ) | $spl122),
  inference(subsumption_resolution,[],[f31183,f492])).
fof(f31183,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP1(X21,X22,X23) | p(s(bool,'F_0')) | p(s(bool,X24)) | p(s(bool,X20))) ) | $spl122),
  inference(resolution,[],[f28114,f458])).
fof(f31181,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP10(X10,X11) | sP1(X11,X12,X13) | p(s(bool,X10)) | sP10(X14,'F_0')) ) | $spl122),
  inference(resolution,[],[f28114,f18977])).
fof(f31180,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP10(X5,X6) | sP1(X6,X7,X8) | p(s(bool,X9)) | sP10('F_0',X5)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f28114,f20831])).
fof(f31179,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10(X0,X1) | sP1(X1,X2,X3) | p(s(bool,X0)) | sP10('F_0',X4)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f28114,f20832])).
fof(f31170,plain,(
  ( ! [X37,X35,X36,X34] : (sP11(X34,X35) | sP10(X35,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X36)) | ~p(s(bool,X37))) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f28113,f11924])).
fof(f31139,plain,(
  ( ! [X30,X33,X31,X32] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X30) | sP1(X31,X30,X32) | p(s(bool,X33))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f28113,f14312])).
fof(f31175,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP1(X27,X26,X28) | p(s(bool,X29)) | p(s(bool,X25))) ) | $spl122),
  inference(subsumption_resolution,[],[f31138,f492])).
fof(f31138,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP1(X27,X26,X28) | p(s(bool,'F_0')) | p(s(bool,X29)) | p(s(bool,X25))) ) | $spl122),
  inference(resolution,[],[f28113,f454])).
fof(f31174,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP1(X22,X21,X23) | p(s(bool,X24)) | p(s(bool,X20))) ) | $spl122),
  inference(subsumption_resolution,[],[f31137,f492])).
fof(f31137,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP1(X22,X21,X23) | p(s(bool,'F_0')) | p(s(bool,X24)) | p(s(bool,X20))) ) | $spl122),
  inference(resolution,[],[f28113,f458])).
fof(f31135,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP10(X10,X11) | sP1(X12,X11,X13) | p(s(bool,X10)) | sP10(X14,'F_0')) ) | $spl122),
  inference(resolution,[],[f28113,f18977])).
fof(f31134,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP10(X5,X6) | sP1(X7,X6,X8) | p(s(bool,X9)) | sP10('F_0',X5)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f28113,f20831])).
fof(f31133,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10(X0,X1) | sP1(X2,X1,X3) | p(s(bool,X0)) | sP10('F_0',X4)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f28113,f20832])).
fof(f31124,plain,(
  ( ! [X54,X56,X55,X53] : (sP11(X53,X54) | sP10(X54,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X55)) | ~p(s(bool,X56))) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f28111,f11920])).
fof(f31088,plain,(
  ( ! [X30,X33,X31,X32] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X30) | sP2(X30,X31,X32) | p(s(bool,X33))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f28111,f14312])).
fof(f31128,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP2(X26,X27,X28) | p(s(bool,X29)) | p(s(bool,X25))) ) | $spl122),
  inference(subsumption_resolution,[],[f31087,f492])).
fof(f31087,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP2(X26,X27,X28) | p(s(bool,'F_0')) | p(s(bool,X29)) | p(s(bool,X25))) ) | $spl122),
  inference(resolution,[],[f28111,f454])).
fof(f31127,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP2(X21,X22,X23) | p(s(bool,X24)) | p(s(bool,X20))) ) | $spl122),
  inference(subsumption_resolution,[],[f31086,f492])).
fof(f31086,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP2(X21,X22,X23) | p(s(bool,'F_0')) | p(s(bool,X24)) | p(s(bool,X20))) ) | $spl122),
  inference(resolution,[],[f28111,f458])).
fof(f31084,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP10(X10,X11) | sP2(X11,X12,X13) | p(s(bool,X10)) | sP10(X14,'F_0')) ) | $spl122),
  inference(resolution,[],[f28111,f18977])).
fof(f31083,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP10(X5,X6) | sP2(X6,X7,X8) | p(s(bool,X9)) | sP10('F_0',X5)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f28111,f20831])).
fof(f31082,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10(X0,X1) | sP2(X1,X2,X3) | p(s(bool,X0)) | sP10('F_0',X4)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f28111,f20832])).
fof(f31073,plain,(
  ( ! [X52,X50,X51,X49] : (sP11(X49,X50) | sP10(X50,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X51)) | ~p(s(bool,X52))) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f28110,f11921])).
fof(f31081,plain,(
  ( ! [X30,X31,X29] : (sP11(X29,X30) | s(bool,'T_0') != s(bool,X31) | sP10(X30,X31)) ) | $spl414),
  inference(cnf_transformation,[],[f31081_D])).
fof(f31081_D,plain,(
  ( ! [X30,X31,X29] : (sP11(X29,X30) | s(bool,'T_0') != s(bool,X31) | sP10(X30,X31)) ) <=> ~$spl414),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl414])])).
fof(f31068,plain,(
  ( ! [X30,X31,X29,X32] : (sP11(X29,X30) | sP10(X30,X31) | s(bool,'T_0') != s(bool,X31) | ~p(s(bool,X32))) ) | ($spl6 | $spl122)),
  inference(resolution,[],[f28110,f14550])).
fof(f31038,plain,(
  ( ! [X30,X33,X31,X32] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X30) | sP2(X31,X30,X32) | p(s(bool,X33))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f28110,f14312])).
fof(f31078,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP2(X27,X26,X28) | p(s(bool,X29)) | p(s(bool,X25))) ) | $spl122),
  inference(subsumption_resolution,[],[f31037,f492])).
fof(f31037,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP2(X27,X26,X28) | p(s(bool,'F_0')) | p(s(bool,X29)) | p(s(bool,X25))) ) | $spl122),
  inference(resolution,[],[f28110,f454])).
fof(f31077,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP2(X22,X21,X23) | p(s(bool,X24)) | p(s(bool,X20))) ) | $spl122),
  inference(subsumption_resolution,[],[f31036,f492])).
fof(f31036,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP2(X22,X21,X23) | p(s(bool,'F_0')) | p(s(bool,X24)) | p(s(bool,X20))) ) | $spl122),
  inference(resolution,[],[f28110,f458])).
fof(f31034,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP10(X10,X11) | sP2(X12,X11,X13) | p(s(bool,X10)) | sP10(X14,'F_0')) ) | $spl122),
  inference(resolution,[],[f28110,f18977])).
fof(f31033,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP10(X5,X6) | sP2(X7,X6,X8) | p(s(bool,X9)) | sP10('F_0',X5)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f28110,f20831])).
fof(f31032,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10(X0,X1) | sP2(X2,X1,X3) | p(s(bool,X0)) | sP10('F_0',X4)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f28110,f20832])).
fof(f28078,plain,(
  ( ! [X107,X106] : (sP11(X106,X107) | sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X106)) ) | $spl322),
  inference(cnf_transformation,[],[f28078_D])).
fof(f28078_D,plain,(
  ( ! [X107,X106] : (sP11(X106,X107) | sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X106)) ) <=> ~$spl322),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl322])])).
fof(f31028,plain,(
  ( ! [X35,X33,X36,X34] : (sP11(X33,X34) | sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X33) | p(s(bool,X35)) | ~p(s(bool,X36))) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f27959,f11976])).
fof(f30598,plain,(
  ( ! [X30,X31,X32] : (sP10(X30,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP8(X30,X31,X32)) ) | $spl404),
  inference(cnf_transformation,[],[f30598_D])).
fof(f30598_D,plain,(
  ( ! [X30,X31,X32] : (sP10(X30,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP8(X30,X31,X32)) ) <=> ~$spl404),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl404])])).
fof(f30998,plain,(
  ( ! [X37,X35,X36,X34] : (sP10(X34,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP8(X34,X35,X36) | p(s(bool,X37))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f27959,f14313])).
fof(f31030,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP8(X25,X27,X28) | p(s(bool,X26)) | p(s(bool,X29))) ) | $spl122),
  inference(subsumption_resolution,[],[f30996,f492])).
fof(f30996,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP8(X25,X27,X28) | p(s(bool,'F_0')) | p(s(bool,X26)) | p(s(bool,X29))) ) | $spl122),
  inference(resolution,[],[f27959,f454])).
fof(f31029,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP8(X20,X22,X23) | p(s(bool,X21)) | p(s(bool,X24))) ) | $spl122),
  inference(subsumption_resolution,[],[f30995,f492])).
fof(f30995,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP8(X20,X22,X23) | p(s(bool,'F_0')) | p(s(bool,X21)) | p(s(bool,X24))) ) | $spl122),
  inference(resolution,[],[f27959,f458])).
fof(f30994,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP10(X15,X16) | sP8(X15,X17,X18) | p(s(bool,X16)) | sP10(X19,'F_0')) ) | $spl122),
  inference(resolution,[],[f27959,f18976])).
fof(f30992,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP10(X5,X6) | sP8(X5,X7,X8) | p(s(bool,X6)) | sP10('F_0',X9)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f27959,f20831])).
fof(f30595,plain,(
  ( ! [X6,X8,X7,X5] : (sP10(X5,X6) | sP10('F_0',X6) | sP8(X5,X7,X8)) ) | $spl402),
  inference(cnf_transformation,[],[f30595_D])).
fof(f30595_D,plain,(
  ( ! [X6,X8,X7,X5] : (sP10(X5,X6) | sP10('F_0',X6) | sP8(X5,X7,X8)) ) <=> ~$spl402),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl402])])).
fof(f30991,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10(X0,X1) | sP8(X0,X2,X3) | p(s(bool,X4)) | sP10('F_0',X1)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f27959,f20832])).
fof(f30943,plain,(
  ( ! [X30,X31,X29,X32] : (sP11(X29,X30) | sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X29) | p(s(bool,X31)) | ~p(s(bool,X32))) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f27958,f11977])).
fof(f30553,plain,(
  ( ! [X30,X31,X32] : (sP10(X30,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP8(X31,X30,X32)) ) | $spl398),
  inference(cnf_transformation,[],[f30553_D])).
fof(f30553_D,plain,(
  ( ! [X30,X31,X32] : (sP10(X30,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP8(X31,X30,X32)) ) <=> ~$spl398),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl398])])).
fof(f30914,plain,(
  ( ! [X37,X35,X36,X34] : (sP10(X34,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP8(X35,X34,X36) | p(s(bool,X37))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f27958,f14313])).
fof(f30946,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP8(X27,X25,X28) | p(s(bool,X26)) | p(s(bool,X29))) ) | $spl122),
  inference(subsumption_resolution,[],[f30912,f492])).
fof(f30912,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP8(X27,X25,X28) | p(s(bool,'F_0')) | p(s(bool,X26)) | p(s(bool,X29))) ) | $spl122),
  inference(resolution,[],[f27958,f454])).
fof(f30945,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP8(X22,X20,X23) | p(s(bool,X21)) | p(s(bool,X24))) ) | $spl122),
  inference(subsumption_resolution,[],[f30911,f492])).
fof(f30911,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP8(X22,X20,X23) | p(s(bool,'F_0')) | p(s(bool,X21)) | p(s(bool,X24))) ) | $spl122),
  inference(resolution,[],[f27958,f458])).
fof(f30910,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP10(X15,X16) | sP8(X17,X15,X18) | p(s(bool,X16)) | sP10(X19,'F_0')) ) | $spl122),
  inference(resolution,[],[f27958,f18976])).
fof(f30908,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP10(X5,X6) | sP8(X7,X5,X8) | p(s(bool,X6)) | sP10('F_0',X9)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f27958,f20831])).
fof(f30550,plain,(
  ( ! [X6,X8,X7,X5] : (sP10(X5,X6) | sP10('F_0',X6) | sP8(X7,X5,X8)) ) | $spl396),
  inference(cnf_transformation,[],[f30550_D])).
fof(f30550_D,plain,(
  ( ! [X6,X8,X7,X5] : (sP10(X5,X6) | sP10('F_0',X6) | sP8(X7,X5,X8)) ) <=> ~$spl396),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl396])])).
fof(f30907,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10(X0,X1) | sP8(X2,X0,X3) | p(s(bool,X4)) | sP10('F_0',X1)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f27958,f20832])).
fof(f30904,plain,(
  ( ! [X39,X37,X38,X40] : (sP11(X37,X38) | sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X37) | p(s(bool,X39)) | p(s(bool,X40))) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f27945,f11960])).
fof(f30903,plain,(
  ( ! [X35,X33,X36,X34] : (sP11(X33,X34) | sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X33) | sP7('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X35,X36)) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f27945,f12008])).
fof(f30510,plain,(
  ( ! [X30,X31,X32] : (sP10(X30,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP6(X30,X31,X32)) ) | $spl392),
  inference(cnf_transformation,[],[f30510_D])).
fof(f30510_D,plain,(
  ( ! [X30,X31,X32] : (sP10(X30,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP6(X30,X31,X32)) ) <=> ~$spl392),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl392])])).
fof(f30873,plain,(
  ( ! [X37,X35,X36,X34] : (sP10(X34,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP6(X34,X35,X36) | p(s(bool,X37))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f27945,f14313])).
fof(f30906,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP6(X25,X27,X28) | p(s(bool,X26)) | p(s(bool,X29))) ) | $spl122),
  inference(subsumption_resolution,[],[f30871,f492])).
fof(f30871,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP6(X25,X27,X28) | p(s(bool,'F_0')) | p(s(bool,X26)) | p(s(bool,X29))) ) | $spl122),
  inference(resolution,[],[f27945,f454])).
fof(f30905,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP6(X20,X22,X23) | p(s(bool,X21)) | p(s(bool,X24))) ) | $spl122),
  inference(subsumption_resolution,[],[f30870,f492])).
fof(f30870,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP6(X20,X22,X23) | p(s(bool,'F_0')) | p(s(bool,X21)) | p(s(bool,X24))) ) | $spl122),
  inference(resolution,[],[f27945,f458])).
fof(f30869,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP10(X15,X16) | sP6(X15,X17,X18) | p(s(bool,X16)) | sP10(X19,'F_0')) ) | $spl122),
  inference(resolution,[],[f27945,f18976])).
fof(f30867,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP10(X5,X6) | sP6(X5,X7,X8) | p(s(bool,X6)) | sP10('F_0',X9)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f27945,f20831])).
fof(f30507,plain,(
  ( ! [X6,X8,X7,X5] : (sP10(X5,X6) | sP10('F_0',X6) | sP6(X5,X7,X8)) ) | $spl390),
  inference(cnf_transformation,[],[f30507_D])).
fof(f30507_D,plain,(
  ( ! [X6,X8,X7,X5] : (sP10(X5,X6) | sP10('F_0',X6) | sP6(X5,X7,X8)) ) <=> ~$spl390),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl390])])).
fof(f30866,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10(X0,X1) | sP6(X0,X2,X3) | p(s(bool,X4)) | sP10('F_0',X1)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f27945,f20832])).
fof(f30863,plain,(
  ( ! [X39,X37,X38,X40] : (sP11(X37,X38) | sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X37) | ~p(s(bool,X39)) | p(s(bool,X40))) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f27929,f11946])).
fof(f30862,plain,(
  ( ! [X35,X33,X36,X34] : (sP11(X33,X34) | sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X33) | sP5('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X35,X36)) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f27929,f12007])).
fof(f30466,plain,(
  ( ! [X30,X31,X32] : (sP10(X30,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP4(X30,X31,X32)) ) | $spl386),
  inference(cnf_transformation,[],[f30466_D])).
fof(f30466_D,plain,(
  ( ! [X30,X31,X32] : (sP10(X30,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP4(X30,X31,X32)) ) <=> ~$spl386),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl386])])).
fof(f30832,plain,(
  ( ! [X37,X35,X36,X34] : (sP10(X34,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP4(X34,X35,X36) | p(s(bool,X37))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f27929,f14313])).
fof(f30865,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP4(X25,X27,X28) | p(s(bool,X26)) | p(s(bool,X29))) ) | $spl122),
  inference(subsumption_resolution,[],[f30830,f492])).
fof(f30830,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP4(X25,X27,X28) | p(s(bool,'F_0')) | p(s(bool,X26)) | p(s(bool,X29))) ) | $spl122),
  inference(resolution,[],[f27929,f454])).
fof(f30864,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP4(X20,X22,X23) | p(s(bool,X21)) | p(s(bool,X24))) ) | $spl122),
  inference(subsumption_resolution,[],[f30829,f492])).
fof(f30829,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP4(X20,X22,X23) | p(s(bool,'F_0')) | p(s(bool,X21)) | p(s(bool,X24))) ) | $spl122),
  inference(resolution,[],[f27929,f458])).
fof(f30828,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP10(X15,X16) | sP4(X15,X17,X18) | p(s(bool,X16)) | sP10(X19,'F_0')) ) | $spl122),
  inference(resolution,[],[f27929,f18976])).
fof(f30826,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP10(X5,X6) | sP4(X5,X7,X8) | p(s(bool,X6)) | sP10('F_0',X9)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f27929,f20831])).
fof(f30463,plain,(
  ( ! [X6,X8,X7,X5] : (sP10(X5,X6) | sP10('F_0',X6) | sP4(X5,X7,X8)) ) | $spl384),
  inference(cnf_transformation,[],[f30463_D])).
fof(f30463_D,plain,(
  ( ! [X6,X8,X7,X5] : (sP10(X5,X6) | sP10('F_0',X6) | sP4(X5,X7,X8)) ) <=> ~$spl384),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl384])])).
fof(f30825,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10(X0,X1) | sP4(X0,X2,X3) | p(s(bool,X4)) | sP10('F_0',X1)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f27929,f20832])).
fof(f30818,plain,(
  ( ! [X28,X26,X27,X25] : (sP11(X25,X26) | sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X25) | p(s(bool,X27)) | ~p(s(bool,X28))) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f27928,f11947])).
fof(f30421,plain,(
  ( ! [X30,X31,X32] : (sP10(X30,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP4(X31,X32,X30)) ) | $spl380),
  inference(cnf_transformation,[],[f30421_D])).
fof(f30421_D,plain,(
  ( ! [X30,X31,X32] : (sP10(X30,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP4(X31,X32,X30)) ) <=> ~$spl380),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl380])])).
fof(f30790,plain,(
  ( ! [X37,X35,X36,X34] : (sP10(X34,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP4(X35,X36,X34) | p(s(bool,X37))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f27928,f14313])).
fof(f30823,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP4(X27,X28,X25) | p(s(bool,X26)) | p(s(bool,X29))) ) | $spl122),
  inference(subsumption_resolution,[],[f30788,f492])).
fof(f30788,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP4(X27,X28,X25) | p(s(bool,'F_0')) | p(s(bool,X26)) | p(s(bool,X29))) ) | $spl122),
  inference(resolution,[],[f27928,f454])).
fof(f30822,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP4(X22,X23,X20) | p(s(bool,X21)) | p(s(bool,X24))) ) | $spl122),
  inference(subsumption_resolution,[],[f30787,f492])).
fof(f30787,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP4(X22,X23,X20) | p(s(bool,'F_0')) | p(s(bool,X21)) | p(s(bool,X24))) ) | $spl122),
  inference(resolution,[],[f27928,f458])).
fof(f30786,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP10(X15,X16) | sP4(X17,X18,X15) | p(s(bool,X16)) | sP10(X19,'F_0')) ) | $spl122),
  inference(resolution,[],[f27928,f18976])).
fof(f30784,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP10(X5,X6) | sP4(X7,X8,X5) | p(s(bool,X6)) | sP10('F_0',X9)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f27928,f20831])).
fof(f30418,plain,(
  ( ! [X6,X8,X7,X5] : (sP10(X5,X6) | sP10('F_0',X6) | sP4(X7,X8,X5)) ) | $spl378),
  inference(cnf_transformation,[],[f30418_D])).
fof(f30418_D,plain,(
  ( ! [X6,X8,X7,X5] : (sP10(X5,X6) | sP10('F_0',X6) | sP4(X7,X8,X5)) ) <=> ~$spl378),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl378])])).
fof(f30783,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10(X0,X1) | sP4(X2,X3,X0) | p(s(bool,X4)) | sP10('F_0',X1)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f27928,f20832])).
fof(f30777,plain,(
  ( ! [X39,X41,X38,X40] : (sP11(X38,X39) | sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X38) | p(s(bool,X40)) | ~p(s(bool,X41))) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f27906,f11923])).
fof(f30374,plain,(
  ( ! [X30,X31,X32] : (sP10(X30,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP1(X30,X31,X32)) ) | $spl374),
  inference(cnf_transformation,[],[f30374_D])).
fof(f30374_D,plain,(
  ( ! [X30,X31,X32] : (sP10(X30,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP1(X30,X31,X32)) ) <=> ~$spl374),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl374])])).
fof(f30746,plain,(
  ( ! [X37,X35,X36,X34] : (sP10(X34,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP1(X34,X35,X36) | p(s(bool,X37))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f27906,f14313])).
fof(f30779,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP1(X25,X27,X28) | p(s(bool,X26)) | p(s(bool,X29))) ) | $spl122),
  inference(subsumption_resolution,[],[f30744,f492])).
fof(f30744,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP1(X25,X27,X28) | p(s(bool,'F_0')) | p(s(bool,X26)) | p(s(bool,X29))) ) | $spl122),
  inference(resolution,[],[f27906,f454])).
fof(f30778,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP1(X20,X22,X23) | p(s(bool,X21)) | p(s(bool,X24))) ) | $spl122),
  inference(subsumption_resolution,[],[f30743,f492])).
fof(f30743,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP1(X20,X22,X23) | p(s(bool,'F_0')) | p(s(bool,X21)) | p(s(bool,X24))) ) | $spl122),
  inference(resolution,[],[f27906,f458])).
fof(f30742,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP10(X15,X16) | sP1(X15,X17,X18) | p(s(bool,X16)) | sP10(X19,'F_0')) ) | $spl122),
  inference(resolution,[],[f27906,f18976])).
fof(f30740,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP10(X5,X6) | sP1(X5,X7,X8) | p(s(bool,X6)) | sP10('F_0',X9)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f27906,f20831])).
fof(f30371,plain,(
  ( ! [X6,X8,X7,X5] : (sP10(X5,X6) | sP10('F_0',X6) | sP1(X5,X7,X8)) ) | $spl372),
  inference(cnf_transformation,[],[f30371_D])).
fof(f30371_D,plain,(
  ( ! [X6,X8,X7,X5] : (sP10(X5,X6) | sP10('F_0',X6) | sP1(X5,X7,X8)) ) <=> ~$spl372),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl372])])).
fof(f30739,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10(X0,X1) | sP1(X0,X2,X3) | p(s(bool,X4)) | sP10('F_0',X1)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f27906,f20832])).
fof(f30732,plain,(
  ( ! [X37,X35,X36,X34] : (sP11(X34,X35) | sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X34) | p(s(bool,X36)) | ~p(s(bool,X37))) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f27905,f11924])).
fof(f30328,plain,(
  ( ! [X30,X31,X32] : (sP10(X30,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP1(X31,X30,X32)) ) | $spl368),
  inference(cnf_transformation,[],[f30328_D])).
fof(f30328_D,plain,(
  ( ! [X30,X31,X32] : (sP10(X30,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP1(X31,X30,X32)) ) <=> ~$spl368),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl368])])).
fof(f30702,plain,(
  ( ! [X37,X35,X36,X34] : (sP10(X34,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP1(X35,X34,X36) | p(s(bool,X37))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f27905,f14313])).
fof(f30735,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP1(X27,X25,X28) | p(s(bool,X26)) | p(s(bool,X29))) ) | $spl122),
  inference(subsumption_resolution,[],[f30700,f492])).
fof(f30700,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP1(X27,X25,X28) | p(s(bool,'F_0')) | p(s(bool,X26)) | p(s(bool,X29))) ) | $spl122),
  inference(resolution,[],[f27905,f454])).
fof(f30734,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP1(X22,X20,X23) | p(s(bool,X21)) | p(s(bool,X24))) ) | $spl122),
  inference(subsumption_resolution,[],[f30699,f492])).
fof(f30699,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP1(X22,X20,X23) | p(s(bool,'F_0')) | p(s(bool,X21)) | p(s(bool,X24))) ) | $spl122),
  inference(resolution,[],[f27905,f458])).
fof(f30698,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP10(X15,X16) | sP1(X17,X15,X18) | p(s(bool,X16)) | sP10(X19,'F_0')) ) | $spl122),
  inference(resolution,[],[f27905,f18976])).
fof(f30696,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP10(X5,X6) | sP1(X7,X5,X8) | p(s(bool,X6)) | sP10('F_0',X9)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f27905,f20831])).
fof(f30325,plain,(
  ( ! [X6,X8,X7,X5] : (sP10(X5,X6) | sP10('F_0',X6) | sP1(X7,X5,X8)) ) | $spl366),
  inference(cnf_transformation,[],[f30325_D])).
fof(f30325_D,plain,(
  ( ! [X6,X8,X7,X5] : (sP10(X5,X6) | sP10('F_0',X6) | sP1(X7,X5,X8)) ) <=> ~$spl366),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl366])])).
fof(f30695,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10(X0,X1) | sP1(X2,X0,X3) | p(s(bool,X4)) | sP10('F_0',X1)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f27905,f20832])).
fof(f30689,plain,(
  ( ! [X54,X56,X55,X53] : (sP11(X53,X54) | sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X53) | p(s(bool,X55)) | ~p(s(bool,X56))) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f27903,f11920])).
fof(f30281,plain,(
  ( ! [X30,X31,X32] : (sP10(X30,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP2(X30,X31,X32)) ) | $spl362),
  inference(cnf_transformation,[],[f30281_D])).
fof(f30281_D,plain,(
  ( ! [X30,X31,X32] : (sP10(X30,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP2(X30,X31,X32)) ) <=> ~$spl362),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl362])])).
fof(f30654,plain,(
  ( ! [X37,X35,X36,X34] : (sP10(X34,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP2(X34,X35,X36) | p(s(bool,X37))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f27903,f14313])).
fof(f30691,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP2(X25,X27,X28) | p(s(bool,X26)) | p(s(bool,X29))) ) | $spl122),
  inference(subsumption_resolution,[],[f30652,f492])).
fof(f30652,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP2(X25,X27,X28) | p(s(bool,'F_0')) | p(s(bool,X26)) | p(s(bool,X29))) ) | $spl122),
  inference(resolution,[],[f27903,f454])).
fof(f30690,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP2(X20,X22,X23) | p(s(bool,X21)) | p(s(bool,X24))) ) | $spl122),
  inference(subsumption_resolution,[],[f30651,f492])).
fof(f30651,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP2(X20,X22,X23) | p(s(bool,'F_0')) | p(s(bool,X21)) | p(s(bool,X24))) ) | $spl122),
  inference(resolution,[],[f27903,f458])).
fof(f30650,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP10(X15,X16) | sP2(X15,X17,X18) | p(s(bool,X16)) | sP10(X19,'F_0')) ) | $spl122),
  inference(resolution,[],[f27903,f18976])).
fof(f30648,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP10(X5,X6) | sP2(X5,X7,X8) | p(s(bool,X6)) | sP10('F_0',X9)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f27903,f20831])).
fof(f30278,plain,(
  ( ! [X6,X8,X7,X5] : (sP10(X5,X6) | sP10('F_0',X6) | sP2(X5,X7,X8)) ) | $spl360),
  inference(cnf_transformation,[],[f30278_D])).
fof(f30278_D,plain,(
  ( ! [X6,X8,X7,X5] : (sP10(X5,X6) | sP10('F_0',X6) | sP2(X5,X7,X8)) ) <=> ~$spl360),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl360])])).
fof(f30647,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10(X0,X1) | sP2(X0,X2,X3) | p(s(bool,X4)) | sP10('F_0',X1)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f27903,f20832])).
fof(f30641,plain,(
  ( ! [X52,X50,X51,X49] : (sP11(X49,X50) | sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X49) | p(s(bool,X51)) | ~p(s(bool,X52))) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f27902,f11921])).
fof(f30646,plain,(
  ( ! [X30,X31,X29] : (sP11(X29,X30) | s(bool,'T_0') != s(bool,X31) | sP10(X31,X29)) ) | $spl406),
  inference(cnf_transformation,[],[f30646_D])).
fof(f30646_D,plain,(
  ( ! [X30,X31,X29] : (sP11(X29,X30) | s(bool,'T_0') != s(bool,X31) | sP10(X31,X29)) ) <=> ~$spl406),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl406])])).
fof(f30636,plain,(
  ( ! [X30,X31,X29,X32] : (sP11(X29,X30) | sP10(X31,X29) | s(bool,'T_0') != s(bool,X31) | ~p(s(bool,X32))) ) | ($spl6 | $spl122)),
  inference(resolution,[],[f27902,f14550])).
fof(f30231,plain,(
  ( ! [X30,X31,X32] : (sP10(X30,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP2(X31,X30,X32)) ) | $spl354),
  inference(cnf_transformation,[],[f30231_D])).
fof(f30231_D,plain,(
  ( ! [X30,X31,X32] : (sP10(X30,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP2(X31,X30,X32)) ) <=> ~$spl354),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl354])])).
fof(f30607,plain,(
  ( ! [X37,X35,X36,X34] : (sP10(X34,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP2(X35,X34,X36) | p(s(bool,X37))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f27902,f14313])).
fof(f30644,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP2(X27,X25,X28) | p(s(bool,X26)) | p(s(bool,X29))) ) | $spl122),
  inference(subsumption_resolution,[],[f30605,f492])).
fof(f30605,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP2(X27,X25,X28) | p(s(bool,'F_0')) | p(s(bool,X26)) | p(s(bool,X29))) ) | $spl122),
  inference(resolution,[],[f27902,f454])).
fof(f30643,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP2(X22,X20,X23) | p(s(bool,X21)) | p(s(bool,X24))) ) | $spl122),
  inference(subsumption_resolution,[],[f30604,f492])).
fof(f30604,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP2(X22,X20,X23) | p(s(bool,'F_0')) | p(s(bool,X21)) | p(s(bool,X24))) ) | $spl122),
  inference(resolution,[],[f27902,f458])).
fof(f30603,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP10(X15,X16) | sP2(X17,X15,X18) | p(s(bool,X16)) | sP10(X19,'F_0')) ) | $spl122),
  inference(resolution,[],[f27902,f18976])).
fof(f30601,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP10(X5,X6) | sP2(X7,X5,X8) | p(s(bool,X6)) | sP10('F_0',X9)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f27902,f20831])).
fof(f30228,plain,(
  ( ! [X6,X8,X7,X5] : (sP10(X5,X6) | sP10('F_0',X6) | sP2(X7,X5,X8)) ) | $spl352),
  inference(cnf_transformation,[],[f30228_D])).
fof(f30228_D,plain,(
  ( ! [X6,X8,X7,X5] : (sP10(X5,X6) | sP10('F_0',X6) | sP2(X7,X5,X8)) ) <=> ~$spl352),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl352])])).
fof(f30600,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10(X0,X1) | sP2(X2,X0,X3) | p(s(bool,X4)) | sP10('F_0',X1)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f27902,f20832])).
fof(f27888,plain,(
  ( ! [X107,X106] : (sP11(X106,X107) | sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X107)) ) | $spl316),
  inference(cnf_transformation,[],[f27888_D])).
fof(f27888_D,plain,(
  ( ! [X107,X106] : (sP11(X106,X107) | sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X107)) ) <=> ~$spl316),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl316])])).
fof(f30593,plain,(
  ( ! [X35,X33,X36,X34] : (sP11(X33,X34) | sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X34) | p(s(bool,X35)) | ~p(s(bool,X36))) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f27769,f11976])).
fof(f30562,plain,(
  ( ! [X30,X33,X31,X32] : (sP10(X30,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP8(X30,X31,X32) | p(s(bool,X33))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f27769,f14312])).
fof(f30597,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP8(X25,X27,X28) | p(s(bool,X29)) | p(s(bool,X26))) ) | $spl122),
  inference(subsumption_resolution,[],[f30561,f492])).
fof(f30561,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP8(X25,X27,X28) | p(s(bool,'F_0')) | p(s(bool,X29)) | p(s(bool,X26))) ) | $spl122),
  inference(resolution,[],[f27769,f454])).
fof(f30596,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP8(X20,X22,X23) | p(s(bool,X24)) | p(s(bool,X21))) ) | $spl122),
  inference(subsumption_resolution,[],[f30560,f492])).
fof(f30560,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP8(X20,X22,X23) | p(s(bool,'F_0')) | p(s(bool,X24)) | p(s(bool,X21))) ) | $spl122),
  inference(resolution,[],[f27769,f458])).
fof(f30558,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP10(X10,X11) | sP8(X10,X12,X13) | p(s(bool,X11)) | sP10(X14,'F_0')) ) | $spl122),
  inference(resolution,[],[f27769,f18977])).
fof(f30557,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP10(X5,X6) | sP8(X5,X7,X8) | p(s(bool,X9)) | sP10('F_0',X6)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f27769,f20831])).
fof(f30556,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10(X0,X1) | sP8(X0,X2,X3) | p(s(bool,X1)) | sP10('F_0',X4)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f27769,f20832])).
fof(f30547,plain,(
  ( ! [X30,X31,X29,X32] : (sP11(X29,X30) | sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X30) | p(s(bool,X31)) | ~p(s(bool,X32))) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f27768,f11977])).
fof(f30517,plain,(
  ( ! [X30,X33,X31,X32] : (sP10(X30,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP8(X31,X30,X32) | p(s(bool,X33))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f27768,f14312])).
fof(f30552,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP8(X27,X25,X28) | p(s(bool,X29)) | p(s(bool,X26))) ) | $spl122),
  inference(subsumption_resolution,[],[f30516,f492])).
fof(f30516,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP8(X27,X25,X28) | p(s(bool,'F_0')) | p(s(bool,X29)) | p(s(bool,X26))) ) | $spl122),
  inference(resolution,[],[f27768,f454])).
fof(f30551,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP8(X22,X20,X23) | p(s(bool,X24)) | p(s(bool,X21))) ) | $spl122),
  inference(subsumption_resolution,[],[f30515,f492])).
fof(f30515,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP8(X22,X20,X23) | p(s(bool,'F_0')) | p(s(bool,X24)) | p(s(bool,X21))) ) | $spl122),
  inference(resolution,[],[f27768,f458])).
fof(f30513,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP10(X10,X11) | sP8(X12,X10,X13) | p(s(bool,X11)) | sP10(X14,'F_0')) ) | $spl122),
  inference(resolution,[],[f27768,f18977])).
fof(f30512,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP10(X5,X6) | sP8(X7,X5,X8) | p(s(bool,X9)) | sP10('F_0',X6)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f27768,f20831])).
fof(f30511,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10(X0,X1) | sP8(X2,X0,X3) | p(s(bool,X1)) | sP10('F_0',X4)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f27768,f20832])).
fof(f30505,plain,(
  ( ! [X39,X37,X38,X40] : (sP11(X37,X38) | sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X38) | p(s(bool,X39)) | p(s(bool,X40))) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f27755,f11960])).
fof(f30504,plain,(
  ( ! [X35,X33,X36,X34] : (sP11(X33,X34) | sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X34) | sP7('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X35,X36)) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f27755,f12008])).
fof(f30473,plain,(
  ( ! [X30,X33,X31,X32] : (sP10(X30,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP6(X30,X31,X32) | p(s(bool,X33))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f27755,f14312])).
fof(f30509,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP6(X25,X27,X28) | p(s(bool,X29)) | p(s(bool,X26))) ) | $spl122),
  inference(subsumption_resolution,[],[f30472,f492])).
fof(f30472,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP6(X25,X27,X28) | p(s(bool,'F_0')) | p(s(bool,X29)) | p(s(bool,X26))) ) | $spl122),
  inference(resolution,[],[f27755,f454])).
fof(f30508,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP6(X20,X22,X23) | p(s(bool,X24)) | p(s(bool,X21))) ) | $spl122),
  inference(subsumption_resolution,[],[f30471,f492])).
fof(f30471,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP6(X20,X22,X23) | p(s(bool,'F_0')) | p(s(bool,X24)) | p(s(bool,X21))) ) | $spl122),
  inference(resolution,[],[f27755,f458])).
fof(f30469,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP10(X10,X11) | sP6(X10,X12,X13) | p(s(bool,X11)) | sP10(X14,'F_0')) ) | $spl122),
  inference(resolution,[],[f27755,f18977])).
fof(f30468,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP10(X5,X6) | sP6(X5,X7,X8) | p(s(bool,X9)) | sP10('F_0',X6)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f27755,f20831])).
fof(f30467,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10(X0,X1) | sP6(X0,X2,X3) | p(s(bool,X1)) | sP10('F_0',X4)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f27755,f20832])).
fof(f30461,plain,(
  ( ! [X39,X37,X38,X40] : (sP11(X37,X38) | sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X38) | ~p(s(bool,X39)) | p(s(bool,X40))) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f27739,f11946])).
fof(f30460,plain,(
  ( ! [X35,X33,X36,X34] : (sP11(X33,X34) | sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X34) | sP5('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X35,X36)) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f27739,f12007])).
fof(f30429,plain,(
  ( ! [X30,X33,X31,X32] : (sP10(X30,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP4(X30,X31,X32) | p(s(bool,X33))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f27739,f14312])).
fof(f30465,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP4(X25,X27,X28) | p(s(bool,X29)) | p(s(bool,X26))) ) | $spl122),
  inference(subsumption_resolution,[],[f30428,f492])).
fof(f30428,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP4(X25,X27,X28) | p(s(bool,'F_0')) | p(s(bool,X29)) | p(s(bool,X26))) ) | $spl122),
  inference(resolution,[],[f27739,f454])).
fof(f30464,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP4(X20,X22,X23) | p(s(bool,X24)) | p(s(bool,X21))) ) | $spl122),
  inference(subsumption_resolution,[],[f30427,f492])).
fof(f30427,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP4(X20,X22,X23) | p(s(bool,'F_0')) | p(s(bool,X24)) | p(s(bool,X21))) ) | $spl122),
  inference(resolution,[],[f27739,f458])).
fof(f30425,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP10(X10,X11) | sP4(X10,X12,X13) | p(s(bool,X11)) | sP10(X14,'F_0')) ) | $spl122),
  inference(resolution,[],[f27739,f18977])).
fof(f30424,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP10(X5,X6) | sP4(X5,X7,X8) | p(s(bool,X9)) | sP10('F_0',X6)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f27739,f20831])).
fof(f30423,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10(X0,X1) | sP4(X0,X2,X3) | p(s(bool,X1)) | sP10('F_0',X4)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f27739,f20832])).
fof(f30413,plain,(
  ( ! [X28,X26,X27,X25] : (sP11(X25,X26) | sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X26) | p(s(bool,X27)) | ~p(s(bool,X28))) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f27738,f11947])).
fof(f30384,plain,(
  ( ! [X30,X33,X31,X32] : (sP10(X30,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP4(X31,X32,X30) | p(s(bool,X33))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f27738,f14312])).
fof(f30420,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP4(X27,X28,X25) | p(s(bool,X29)) | p(s(bool,X26))) ) | $spl122),
  inference(subsumption_resolution,[],[f30383,f492])).
fof(f30383,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP4(X27,X28,X25) | p(s(bool,'F_0')) | p(s(bool,X29)) | p(s(bool,X26))) ) | $spl122),
  inference(resolution,[],[f27738,f454])).
fof(f30419,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP4(X22,X23,X20) | p(s(bool,X24)) | p(s(bool,X21))) ) | $spl122),
  inference(subsumption_resolution,[],[f30382,f492])).
fof(f30382,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP4(X22,X23,X20) | p(s(bool,'F_0')) | p(s(bool,X24)) | p(s(bool,X21))) ) | $spl122),
  inference(resolution,[],[f27738,f458])).
fof(f30380,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP10(X10,X11) | sP4(X12,X13,X10) | p(s(bool,X11)) | sP10(X14,'F_0')) ) | $spl122),
  inference(resolution,[],[f27738,f18977])).
fof(f30379,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP10(X5,X6) | sP4(X7,X8,X5) | p(s(bool,X9)) | sP10('F_0',X6)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f27738,f20831])).
fof(f30378,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10(X0,X1) | sP4(X2,X3,X0) | p(s(bool,X1)) | sP10('F_0',X4)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f27738,f20832])).
fof(f30369,plain,(
  ( ! [X39,X41,X38,X40] : (sP11(X38,X39) | sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X39) | p(s(bool,X40)) | ~p(s(bool,X41))) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f27716,f11923])).
fof(f30337,plain,(
  ( ! [X30,X33,X31,X32] : (sP10(X30,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP1(X30,X31,X32) | p(s(bool,X33))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f27716,f14312])).
fof(f30373,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP1(X25,X27,X28) | p(s(bool,X29)) | p(s(bool,X26))) ) | $spl122),
  inference(subsumption_resolution,[],[f30336,f492])).
fof(f30336,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP1(X25,X27,X28) | p(s(bool,'F_0')) | p(s(bool,X29)) | p(s(bool,X26))) ) | $spl122),
  inference(resolution,[],[f27716,f454])).
fof(f30372,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP1(X20,X22,X23) | p(s(bool,X24)) | p(s(bool,X21))) ) | $spl122),
  inference(subsumption_resolution,[],[f30335,f492])).
fof(f30335,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP1(X20,X22,X23) | p(s(bool,'F_0')) | p(s(bool,X24)) | p(s(bool,X21))) ) | $spl122),
  inference(resolution,[],[f27716,f458])).
fof(f30333,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP10(X10,X11) | sP1(X10,X12,X13) | p(s(bool,X11)) | sP10(X14,'F_0')) ) | $spl122),
  inference(resolution,[],[f27716,f18977])).
fof(f30332,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP10(X5,X6) | sP1(X5,X7,X8) | p(s(bool,X9)) | sP10('F_0',X6)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f27716,f20831])).
fof(f30331,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10(X0,X1) | sP1(X0,X2,X3) | p(s(bool,X1)) | sP10('F_0',X4)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f27716,f20832])).
fof(f30322,plain,(
  ( ! [X37,X35,X36,X34] : (sP11(X34,X35) | sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X35) | p(s(bool,X36)) | ~p(s(bool,X37))) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f27715,f11924])).
fof(f30291,plain,(
  ( ! [X30,X33,X31,X32] : (sP10(X30,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP1(X31,X30,X32) | p(s(bool,X33))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f27715,f14312])).
fof(f30327,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP1(X27,X25,X28) | p(s(bool,X29)) | p(s(bool,X26))) ) | $spl122),
  inference(subsumption_resolution,[],[f30290,f492])).
fof(f30290,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP1(X27,X25,X28) | p(s(bool,'F_0')) | p(s(bool,X29)) | p(s(bool,X26))) ) | $spl122),
  inference(resolution,[],[f27715,f454])).
fof(f30326,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP1(X22,X20,X23) | p(s(bool,X24)) | p(s(bool,X21))) ) | $spl122),
  inference(subsumption_resolution,[],[f30289,f492])).
fof(f30289,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP1(X22,X20,X23) | p(s(bool,'F_0')) | p(s(bool,X24)) | p(s(bool,X21))) ) | $spl122),
  inference(resolution,[],[f27715,f458])).
fof(f30287,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP10(X10,X11) | sP1(X12,X10,X13) | p(s(bool,X11)) | sP10(X14,'F_0')) ) | $spl122),
  inference(resolution,[],[f27715,f18977])).
fof(f30286,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP10(X5,X6) | sP1(X7,X5,X8) | p(s(bool,X9)) | sP10('F_0',X6)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f27715,f20831])).
fof(f30285,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10(X0,X1) | sP1(X2,X0,X3) | p(s(bool,X1)) | sP10('F_0',X4)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f27715,f20832])).
fof(f30276,plain,(
  ( ! [X54,X56,X55,X53] : (sP11(X53,X54) | sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X54) | p(s(bool,X55)) | ~p(s(bool,X56))) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f27713,f11920])).
fof(f30240,plain,(
  ( ! [X30,X33,X31,X32] : (sP10(X30,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP2(X30,X31,X32) | p(s(bool,X33))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f27713,f14312])).
fof(f30280,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP2(X25,X27,X28) | p(s(bool,X29)) | p(s(bool,X26))) ) | $spl122),
  inference(subsumption_resolution,[],[f30239,f492])).
fof(f30239,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP2(X25,X27,X28) | p(s(bool,'F_0')) | p(s(bool,X29)) | p(s(bool,X26))) ) | $spl122),
  inference(resolution,[],[f27713,f454])).
fof(f30279,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP2(X20,X22,X23) | p(s(bool,X24)) | p(s(bool,X21))) ) | $spl122),
  inference(subsumption_resolution,[],[f30238,f492])).
fof(f30238,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP2(X20,X22,X23) | p(s(bool,'F_0')) | p(s(bool,X24)) | p(s(bool,X21))) ) | $spl122),
  inference(resolution,[],[f27713,f458])).
fof(f30236,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP10(X10,X11) | sP2(X10,X12,X13) | p(s(bool,X11)) | sP10(X14,'F_0')) ) | $spl122),
  inference(resolution,[],[f27713,f18977])).
fof(f30235,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP10(X5,X6) | sP2(X5,X7,X8) | p(s(bool,X9)) | sP10('F_0',X6)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f27713,f20831])).
fof(f30234,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10(X0,X1) | sP2(X0,X2,X3) | p(s(bool,X1)) | sP10('F_0',X4)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f27713,f20832])).
fof(f30225,plain,(
  ( ! [X52,X50,X51,X49] : (sP11(X49,X50) | sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X50) | p(s(bool,X51)) | ~p(s(bool,X52))) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f27712,f11921])).
fof(f30233,plain,(
  ( ! [X30,X31,X29] : (sP11(X29,X30) | s(bool,'T_0') != s(bool,X31) | sP10(X31,X30)) ) | $spl356),
  inference(cnf_transformation,[],[f30233_D])).
fof(f30233_D,plain,(
  ( ! [X30,X31,X29] : (sP11(X29,X30) | s(bool,'T_0') != s(bool,X31) | sP10(X31,X30)) ) <=> ~$spl356),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl356])])).
fof(f30220,plain,(
  ( ! [X30,X31,X29,X32] : (sP11(X29,X30) | sP10(X31,X30) | s(bool,'T_0') != s(bool,X31) | ~p(s(bool,X32))) ) | ($spl6 | $spl122)),
  inference(resolution,[],[f27712,f14550])).
fof(f30190,plain,(
  ( ! [X30,X33,X31,X32] : (sP10(X30,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP2(X31,X30,X32) | p(s(bool,X33))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f27712,f14312])).
fof(f30230,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP2(X27,X25,X28) | p(s(bool,X29)) | p(s(bool,X26))) ) | $spl122),
  inference(subsumption_resolution,[],[f30189,f492])).
fof(f30189,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP2(X27,X25,X28) | p(s(bool,'F_0')) | p(s(bool,X29)) | p(s(bool,X26))) ) | $spl122),
  inference(resolution,[],[f27712,f454])).
fof(f30229,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP2(X22,X20,X23) | p(s(bool,X24)) | p(s(bool,X21))) ) | $spl122),
  inference(subsumption_resolution,[],[f30188,f492])).
fof(f30188,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | sP2(X22,X20,X23) | p(s(bool,'F_0')) | p(s(bool,X24)) | p(s(bool,X21))) ) | $spl122),
  inference(resolution,[],[f27712,f458])).
fof(f30186,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP10(X10,X11) | sP2(X12,X10,X13) | p(s(bool,X11)) | sP10(X14,'F_0')) ) | $spl122),
  inference(resolution,[],[f27712,f18977])).
fof(f30185,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP10(X5,X6) | sP2(X7,X5,X8) | p(s(bool,X9)) | sP10('F_0',X6)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f27712,f20831])).
fof(f30184,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10(X0,X1) | sP2(X2,X0,X3) | p(s(bool,X1)) | sP10('F_0',X4)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f27712,f20832])).
fof(f29995,plain,(
  ( ! [X4,X2] : (s(sK12,sK13) = s(sK12,'const.list.EL_2'(s('type.num.num',sK19(happ(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK16)),X2,sK12)),s('type.list.list'(sK12),X2))) | ~p(s(bool,X4)) | s(bool,X4) != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),X2)))))) ) | $spl348),
  inference(cnf_transformation,[],[f29995_D])).
fof(f29995_D,plain,(
  ( ! [X4,X2] : (s(sK12,sK13) = s(sK12,'const.list.EL_2'(s('type.num.num',sK19(happ(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK16)),X2,sK12)),s('type.list.list'(sK12),X2))) | ~p(s(bool,X4)) | s(bool,X4) != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),X2)))))) ) <=> ~$spl348),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl348])])).
fof(f29954,plain,(
  ( ! [X261,X260,X262] : (s(sK12,sK13) = s(sK12,'const.list.EL_2'(s('type.num.num',sK19(happ(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK16)),X260,sK12)),s('type.list.list'(sK12),X260))) | p(s(bool,X261)) | ~p(s(bool,X262)) | s(bool,X262) != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),X260)))))) ) | $spl4),
  inference(resolution,[],[f14181,f536])).
fof(f536,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f309])).
fof(f309,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f196])).
fof(f196,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f195])).
fof(f195,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f145])).
fof(f145,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f64])).
fof(f64,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X7,X19,X14,X20] : ((s(bool,X7) = s(bool,X19) & (p(s(bool,X19)) => s(bool,X14) = s(bool,X20))) => ((p(s(bool,X7)) => p(s(bool,X14))) <=> (p(s(bool,X19)) => p(s(bool,X20)))))),
  file('thm.list.MEM_GENLIST','thm.bool.IMP_CONG')).
fof(f14181,plain,(
  ( ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),X3)))))) | s(sK12,sK13) = s(sK12,'const.list.EL_2'(s('type.num.num',sK19(happ(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK16)),X3,sK12)),s('type.list.list'(sK12),X3)))) ) | $spl4),
  inference(forward_demodulation,[],[f14178,f548])).
fof(f14178,plain,(
  ( ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),X3)))))) | s(sK12,happ(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK16))) = s(sK12,'const.list.EL_2'(s('type.num.num',sK19(happ(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK16)),X3,sK12)),s('type.list.list'(sK12),X3)))) ) | $spl4),
  inference(superposition,[],[f365,f548])).
fof(f365,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.list.EL_2'(s('type.num.num',sK19(X2,X1,X0)),s('type.list.list'(X0),X1))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1))))))) )),
  inference(cnf_transformation,[],[f222])).
fof(f678,plain,(
  ( ! [X10,X9] : (~p(s(bool,X9)) | s(bool,X9) != s(bool,X10)) ) | $spl12),
  inference(cnf_transformation,[],[f678_D])).
fof(f678_D,plain,(
  ( ! [X10,X9] : (~p(s(bool,X9)) | s(bool,X9) != s(bool,X10)) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f29863,plain,(
  ( ! [X30,X33,X31,X32] : (s(sK12,sK13) = s(sK12,'const.list.EL_2'(s('type.num.num',sK19(happ(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK16)),X30,sK12)),s('type.list.list'(sK12),X30))) | ~p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X32) != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),X30))))) | s(bool,X31) != s(bool,X33)) ) | $spl4),
  inference(resolution,[],[f14181,f312])).
fof(f312,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f196])).
fof(f29861,plain,(
  ( ! [X24,X23,X25,X22] : (s(sK12,sK13) = s(sK12,'const.list.EL_2'(s('type.num.num',sK19(happ(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK16)),X22,sK12)),s('type.list.list'(sK12),X22))) | p(s(bool,X23)) | ~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),X22)))))) ) | $spl4),
  inference(resolution,[],[f14181,f311])).
fof(f311,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f196])).
fof(f29859,plain,(
  ( ! [X14,X17,X15,X16] : (s(sK12,sK13) = s(sK12,'const.list.EL_2'(s('type.num.num',sK19(happ(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK16)),X14,sK12)),s('type.list.list'(sK12),X14))) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),X14)))))) ) | $spl4),
  inference(resolution,[],[f14181,f310])).
fof(f310,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f196])).
fof(f29858,plain,(
  ( ! [X12,X10,X13,X11] : (s(sK12,sK13) = s(sK12,'const.list.EL_2'(s('type.num.num',sK19(happ(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK16)),X10,sK12)),s('type.list.list'(sK12),X10))) | ~p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X12) != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),X10))))) | s(bool,X11) != s(bool,X13)) ) | $spl4),
  inference(resolution,[],[f14181,f308])).
fof(f308,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f196])).
fof(f29856,plain,(
  ( ! [X4,X2,X5,X3] : (s(sK12,sK13) = s(sK12,'const.list.EL_2'(s('type.num.num',sK19(happ(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK16)),X2,sK12)),s('type.list.list'(sK12),X2))) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),X2)))))) ) | $spl4),
  inference(resolution,[],[f14181,f306])).
fof(f306,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f196])).
fof(f29832,plain,(
  ( ! [X4,X2] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK19(happ(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK16)),X2,sK12)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),X2)))))) | ~p(s(bool,X4)) | s(bool,X4) != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),X2)))))) ) | $spl346),
  inference(cnf_transformation,[],[f29832_D])).
fof(f29832_D,plain,(
  ( ! [X4,X2] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK19(happ(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK16)),X2,sK12)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),X2)))))) | ~p(s(bool,X4)) | s(bool,X4) != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),X2)))))) ) <=> ~$spl346),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl346])])).
fof(f29791,plain,(
  ( ! [X261,X260,X262] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK19(happ(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK16)),X260,sK12)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),X260)))))) | p(s(bool,X261)) | ~p(s(bool,X262)) | s(bool,X262) != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),X260)))))) ) | $spl4),
  inference(resolution,[],[f14177,f536])).
fof(f14177,plain,(
  ( ! [X2] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),X2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK19(happ(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK16)),X2,sK12)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),X2))))))) ) | $spl4),
  inference(superposition,[],[f364,f548])).
fof(f364,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK19(X2,X1,X0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1))))))) )),
  inference(cnf_transformation,[],[f222])).
fof(f29700,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK19(happ(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK16)),X30,sK12)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),X30)))))) | ~p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X32) != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),X30))))) | s(bool,X31) != s(bool,X33)) ) | $spl4),
  inference(resolution,[],[f14177,f312])).
fof(f29698,plain,(
  ( ! [X24,X23,X25,X22] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK19(happ(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK16)),X22,sK12)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),X22)))))) | p(s(bool,X23)) | ~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),X22)))))) ) | $spl4),
  inference(resolution,[],[f14177,f311])).
fof(f29696,plain,(
  ( ! [X14,X17,X15,X16] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK19(happ(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK16)),X14,sK12)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),X14)))))) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),X14)))))) ) | $spl4),
  inference(resolution,[],[f14177,f310])).
fof(f29695,plain,(
  ( ! [X12,X10,X13,X11] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK19(happ(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK16)),X10,sK12)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),X10)))))) | ~p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X12) != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),X10))))) | s(bool,X11) != s(bool,X13)) ) | $spl4),
  inference(resolution,[],[f14177,f308])).
fof(f29693,plain,(
  ( ! [X4,X2,X5,X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK19(happ(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK16)),X2,sK12)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),X2)))))) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),X2)))))) ) | $spl4),
  inference(resolution,[],[f14177,f306])).
fof(f29659,plain,(
  ( ! [X24,X23] : (sP10('F_0',X23) | s(bool,'F_0') = s(bool,X24) | sP10(X23,X24)) ) | $spl344),
  inference(cnf_transformation,[],[f29659_D])).
fof(f29659_D,plain,(
  ( ! [X24,X23] : (sP10('F_0',X23) | s(bool,'F_0') = s(bool,X24) | sP10(X23,X24)) ) <=> ~$spl344),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl344])])).
fof(f29672,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | sP10('F_0',X26) | sP10(X26,X27) | s(bool,'F_0') = s(bool,X27)) ) | ($spl8 | $spl46 | $spl122 | $spl133)),
  inference(resolution,[],[f20832,f16926])).
fof(f16926,plain,(
  ( ! [X329,X331,X330] : (sP10(X329,X330) | sP11(X329,X331) | s(bool,'F_0') = s(bool,X330)) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16258,f456])).
fof(f16258,plain,(
  ( ! [X596,X597] : (sP10(X597,X596) | p(s(bool,X597)) | s(bool,'F_0') = s(bool,X596)) ) | ($spl8 | $spl46)),
  inference(subsumption_resolution,[],[f16257,f537])).
fof(f16257,plain,(
  ( ! [X596,X597] : (~p(s(bool,'T_0')) | sP10(X597,X596) | p(s(bool,X597)) | s(bool,'F_0') = s(bool,X596)) ) | ($spl8 | $spl46)),
  inference(forward_demodulation,[],[f15612,f2535])).
fof(f15612,plain,(
  ( ! [X596,X597] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | sP10(X597,X596) | p(s(bool,X597)) | s(bool,'F_0') = s(bool,X596)) ) | ($spl8 | $spl46)),
  inference(superposition,[],[f443,f15318])).
fof(f29670,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | sP10('F_0',X20) | sP10(X21,X20) | s(bool,'F_0') = s(bool,X21)) ) | ($spl8 | $spl46 | $spl122 | $spl133)),
  inference(resolution,[],[f20832,f16701])).
fof(f16701,plain,(
  ( ! [X329,X331,X330] : (sP10(X329,X330) | sP11(X330,X331) | s(bool,'F_0') = s(bool,X329)) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16255,f456])).
fof(f29639,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | sP10('F_0',X23) | sP10(X23,X24) | s(bool,'F_0') = s(bool,X24)) ) | ($spl8 | $spl46 | $spl122 | $spl133)),
  inference(resolution,[],[f20831,f16925])).
fof(f16925,plain,(
  ( ! [X327,X326,X328] : (sP10(X326,X327) | sP11(X328,X326) | s(bool,'F_0') = s(bool,X327)) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16258,f455])).
fof(f29637,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | sP10('F_0',X17) | sP10(X18,X17) | s(bool,'F_0') = s(bool,X18)) ) | ($spl8 | $spl46 | $spl122 | $spl133)),
  inference(resolution,[],[f20831,f16700])).
fof(f16700,plain,(
  ( ! [X327,X326,X328] : (sP10(X326,X327) | sP11(X328,X327) | s(bool,'F_0') = s(bool,X326)) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16255,f455])).
fof(f29617,plain,(
  ( ! [X76,X75] : (~p(s(bool,X75)) | sP10(X76,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | sP10('F_0',X76)) ) | ($spl8 | $spl46 | $spl122 | $spl133)),
  inference(subsumption_resolution,[],[f29592,f350])).
fof(f29592,plain,(
  ( ! [X76,X75] : (~p(s(bool,X75)) | sP10(X76,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | sP10('F_0',X76) | s(bool,'F_0') = s(bool,X75)) ) | ($spl8 | $spl46 | $spl122 | $spl133)),
  inference(superposition,[],[f20737,f15318])).
fof(f20737,plain,(
  ( ! [X231,X232] : (~p(s(bool,X232)) | sP10(X231,X232) | sP10('F_0',X231)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f20649,f443])).
fof(f11713,plain,(
  ( ! [X397,X396] : (~sP11(X396,X397) | p(s(bool,X396)) | p(s(bool,X397))) ) | $spl84),
  inference(cnf_transformation,[],[f11713_D])).
fof(f11713_D,plain,(
  ( ! [X397,X396] : (~sP11(X396,X397) | p(s(bool,X396)) | p(s(bool,X397))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f19481,plain,(
  ( ! [X405] : (sP10('F_0',X405) | sP10(X405,'F_0')) ) | $spl142),
  inference(cnf_transformation,[],[f19481_D])).
fof(f19481_D,plain,(
  ( ! [X405] : (sP10('F_0',X405) | sP10(X405,'F_0')) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f29557,plain,(
  ( ! [X428,X429,X427] : (sP10(X427,'F_0') | sP10('F_0',X427) | p(s(bool,X428)) | p(s(bool,X429)) | ~sP11(X428,X429)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f20737,f454])).
fof(f29555,plain,(
  ( ! [X422,X421,X423] : (sP10(X421,'F_0') | sP10('F_0',X421) | p(s(bool,X422)) | p(s(bool,X423)) | ~sP11(X422,X423)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f20737,f458])).
fof(f29603,plain,(
  ( ! [X0,X3,X1] : (sP10(X0,X1) | s(bool,X1) != s(bool,X3) | ~p(s(bool,X3)) | sP10('F_0',X0)) ) | $spl340),
  inference(cnf_transformation,[],[f29603_D])).
fof(f29603_D,plain,(
  ( ! [X0,X3,X1] : (sP10(X0,X1) | s(bool,X1) != s(bool,X3) | ~p(s(bool,X3)) | sP10('F_0',X0)) ) <=> ~$spl340),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl340])])).
fof(f29525,plain,(
  ( ! [X356,X358,X357,X359] : (sP10(X356,X357) | sP10('F_0',X356) | p(s(bool,X358)) | ~p(s(bool,X359)) | s(bool,X357) != s(bool,X359)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f20737,f536])).
fof(f29434,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP10(X35,X36) | sP10('F_0',X35) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X38) | s(bool,X37) != s(bool,X39)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f20737,f312])).
fof(f29432,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP10('F_0',X25) | p(s(bool,X27)) | ~p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X26) != s(bool,X28)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f20737,f311])).
fof(f29430,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP10(X15,X16) | sP10('F_0',X15) | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f20737,f310])).
fof(f29429,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP10(X10,X11) | sP10('F_0',X10) | ~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,X14)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f20737,f308])).
fof(f29427,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10(X0,X1) | sP10('F_0',X0) | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f20737,f306])).
fof(f29424,plain,(
  ( ! [X76,X75] : (~p(s(bool,X75)) | sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X76) | sP10('F_0',X76)) ) | ($spl8 | $spl46 | $spl122 | $spl133)),
  inference(subsumption_resolution,[],[f29389,f350])).
fof(f29389,plain,(
  ( ! [X76,X75] : (~p(s(bool,X75)) | sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X76) | sP10('F_0',X76) | s(bool,'F_0') = s(bool,X75)) ) | ($spl8 | $spl46 | $spl122 | $spl133)),
  inference(superposition,[],[f20736,f15318])).
fof(f20736,plain,(
  ( ! [X229,X230] : (~p(s(bool,X230)) | sP10(X230,X229) | sP10('F_0',X229)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f20649,f440])).
fof(f29404,plain,(
  ( ! [X428,X429,X427] : (sP10('F_0',X427) | p(s(bool,X428)) | p(s(bool,X429)) | ~sP11(X428,X429)) ) | ($spl122 | $spl133)),
  inference(duplicate_literal_removal,[],[f29354])).
fof(f29354,plain,(
  ( ! [X428,X429,X427] : (sP10('F_0',X427) | sP10('F_0',X427) | p(s(bool,X428)) | p(s(bool,X429)) | ~sP11(X428,X429)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f20736,f454])).
fof(f29406,plain,(
  ( ! [X422,X421,X423] : (sP10('F_0',X421) | p(s(bool,X422)) | p(s(bool,X423)) | ~sP11(X422,X423)) ) | ($spl122 | $spl133)),
  inference(duplicate_literal_removal,[],[f29352])).
fof(f29352,plain,(
  ( ! [X422,X421,X423] : (sP10('F_0',X421) | sP10('F_0',X421) | p(s(bool,X422)) | p(s(bool,X423)) | ~sP11(X422,X423)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f20736,f458])).
fof(f29410,plain,(
  ( ! [X0,X3,X1] : (sP10(X0,X1) | s(bool,X0) != s(bool,X3) | ~p(s(bool,X3)) | sP10('F_0',X1)) ) | $spl336),
  inference(cnf_transformation,[],[f29410_D])).
fof(f29410_D,plain,(
  ( ! [X0,X3,X1] : (sP10(X0,X1) | s(bool,X0) != s(bool,X3) | ~p(s(bool,X3)) | sP10('F_0',X1)) ) <=> ~$spl336),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl336])])).
fof(f29322,plain,(
  ( ! [X356,X358,X357,X359] : (sP10(X356,X357) | sP10('F_0',X357) | p(s(bool,X358)) | ~p(s(bool,X359)) | s(bool,X356) != s(bool,X359)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f20736,f536])).
fof(f29231,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP10(X35,X36) | sP10('F_0',X36) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X35) != s(bool,X38) | s(bool,X37) != s(bool,X39)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f20736,f312])).
fof(f29229,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP10('F_0',X26) | p(s(bool,X27)) | ~p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X25) != s(bool,X28)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f20736,f311])).
fof(f29227,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP10(X15,X16) | sP10('F_0',X16) | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X15) != s(bool,X18)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f20736,f310])).
fof(f29226,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP10(X10,X11) | sP10('F_0',X11) | ~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X10) != s(bool,X13) | s(bool,X12) != s(bool,X14)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f20736,f308])).
fof(f29224,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10(X0,X1) | sP10('F_0',X1) | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X0) != s(bool,X3)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f20736,f306])).
fof(f29161,plain,(
  ( ! [X54,X52,X53] : (p(s(bool,X52)) | sP10('F_0',X53) | sP10(X54,X53) | s(bool,'F_0') = s(bool,X54)) ) | ($spl8 | $spl46 | $spl122 | $spl133)),
  inference(duplicate_literal_removal,[],[f28960])).
fof(f28960,plain,(
  ( ! [X54,X52,X53] : (p(s(bool,X52)) | sP10('F_0',X53) | sP10(X54,X53) | p(s(bool,X52)) | s(bool,'F_0') = s(bool,X54)) ) | ($spl8 | $spl46 | $spl122 | $spl133)),
  inference(resolution,[],[f20735,f16694])).
fof(f16694,plain,(
  ( ! [X308,X310,X309] : (sP10(X308,X309) | sP10(X309,X310) | p(s(bool,X310)) | s(bool,'F_0') = s(bool,X308)) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16255,f440])).
fof(f24215,plain,(
  ( ! [X12] : (sP10(X12,'T_0') | sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X12)) ) | $spl240),
  inference(cnf_transformation,[],[f24215_D])).
fof(f24215_D,plain,(
  ( ! [X12] : (sP10(X12,'T_0') | sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X12)) ) <=> ~$spl240),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl240])])).
fof(f28586,plain,(
  ( ! [X50,X51,X49] : (~p(s(bool,X49)) | ~p(s(bool,X50)) | sP10(X51,'T_0') | sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X51)) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f14252,f24762])).
fof(f24762,plain,(
  ( ! [X271,X273,X270,X272] : (sP10(X270,'T_0') | sP10(X271,X270) | sP8(X272,X273,X271)) ) | $spl122),
  inference(resolution,[],[f20223,f409])).
fof(f409,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f243])).
fof(f20223,plain,(
  ( ! [X231,X232] : (sP10(X231,'T_0') | sP10(X232,X231) | p(s(bool,X232))) ) | $spl122),
  inference(resolution,[],[f20069,f443])).
fof(f14252,plain,(
  ( ! [X130,X129] : (~sP8(X129,X130,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | ~p(s(bool,X130)) | ~p(s(bool,X129))) ) | $spl1),
  inference(resolution,[],[f545,f408])).
fof(f24076,plain,(
  ( ! [X12] : (sP10(X12,'T_0') | sP10(X12,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) ) | $spl236),
  inference(cnf_transformation,[],[f24076_D])).
fof(f24076_D,plain,(
  ( ! [X12] : (sP10(X12,'T_0') | sP10(X12,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) ) <=> ~$spl236),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl236])])).
fof(f28585,plain,(
  ( ! [X47,X48,X46] : (~p(s(bool,X46)) | ~p(s(bool,X47)) | sP10(X48,'T_0') | sP10(X48,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f14252,f24526])).
fof(f24526,plain,(
  ( ! [X271,X273,X270,X272] : (sP10(X270,'T_0') | sP10(X270,X271) | sP8(X272,X273,X271)) ) | $spl122),
  inference(resolution,[],[f20222,f409])).
fof(f20222,plain,(
  ( ! [X229,X230] : (sP10(X229,'T_0') | sP10(X229,X230) | p(s(bool,X230))) ) | $spl122),
  inference(resolution,[],[f20069,f440])).
fof(f23462,plain,(
  ( ! [X12] : (sP10('T_0',X12) | sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X12)) ) | $spl222),
  inference(cnf_transformation,[],[f23462_D])).
fof(f23462_D,plain,(
  ( ! [X12] : (sP10('T_0',X12) | sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X12)) ) <=> ~$spl222),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl222])])).
fof(f28584,plain,(
  ( ! [X45,X43,X44] : (~p(s(bool,X43)) | ~p(s(bool,X44)) | sP10('T_0',X45) | sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X45)) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f14252,f23282])).
fof(f23282,plain,(
  ( ! [X271,X273,X270,X272] : (sP10('T_0',X270) | sP10(X271,X270) | sP8(X272,X273,X271)) ) | $spl122),
  inference(resolution,[],[f19619,f409])).
fof(f22775,plain,(
  ( ! [X12] : (sP10('T_0',X12) | sP10(X12,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) ) | $spl206),
  inference(cnf_transformation,[],[f22775_D])).
fof(f22775_D,plain,(
  ( ! [X12] : (sP10('T_0',X12) | sP10(X12,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) ) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f28583,plain,(
  ( ! [X41,X42,X40] : (~p(s(bool,X40)) | ~p(s(bool,X41)) | sP10('T_0',X42) | sP10(X42,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f14252,f22603])).
fof(f22603,plain,(
  ( ! [X271,X273,X270,X272] : (sP10('T_0',X270) | sP10(X270,X271) | sP8(X272,X273,X271)) ) | $spl122),
  inference(resolution,[],[f19618,f409])).
fof(f20323,plain,(
  sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0') | $spl160),
  inference(cnf_transformation,[],[f20323_D])).
fof(f20323_D,plain,(
  sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0') <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f28580,plain,(
  ( ! [X35,X34] : (~p(s(bool,X34)) | ~p(s(bool,X35)) | sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0')) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f14252,f20212])).
fof(f20212,plain,(
  ( ! [X202,X203,X201] : (sP8(X202,X203,X201) | sP10(X201,'T_0')) ) | $spl122),
  inference(resolution,[],[f20069,f409])).
fof(f19718,plain,(
  sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | $spl148),
  inference(cnf_transformation,[],[f19718_D])).
fof(f19718_D,plain,(
  sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f28577,plain,(
  ( ! [X26,X27] : (~p(s(bool,X26)) | ~p(s(bool,X27)) | sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f14252,f19608])).
fof(f19608,plain,(
  ( ! [X202,X203,X201] : (sP8(X202,X203,X201) | sP10('T_0',X201)) ) | $spl122),
  inference(resolution,[],[f19515,f409])).
fof(f19243,plain,(
  ( ! [X13] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X13) | s(bool,'F_0') = s(bool,X13)) ) | $spl136),
  inference(cnf_transformation,[],[f19243_D])).
fof(f19243_D,plain,(
  ( ! [X13] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X13) | s(bool,'F_0') = s(bool,X13)) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f28572,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,X13)) | ~p(s(bool,X14)) | sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X15) | s(bool,'F_0') = s(bool,X15)) ) | ($spl1 | $spl8 | $spl46)),
  inference(resolution,[],[f14252,f16909])).
fof(f16909,plain,(
  ( ! [X271,X273,X270,X272] : (sP10(X270,X271) | sP8(X272,X273,X270) | s(bool,'F_0') = s(bool,X271)) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16258,f409])).
fof(f18525,plain,(
  ( ! [X132] : (s(bool,'F_0') = s(bool,X132) | sP10(X132,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) ) | $spl116),
  inference(cnf_transformation,[],[f18525_D])).
fof(f18525_D,plain,(
  ( ! [X132] : (s(bool,'F_0') = s(bool,X132) | sP10(X132,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f28571,plain,(
  ( ! [X12,X10,X11] : (~p(s(bool,X10)) | ~p(s(bool,X11)) | sP10(X12,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | s(bool,'F_0') = s(bool,X12)) ) | ($spl1 | $spl8 | $spl46)),
  inference(resolution,[],[f14252,f16684])).
fof(f16684,plain,(
  ( ! [X271,X273,X270,X272] : (sP10(X270,X271) | sP8(X272,X273,X271) | s(bool,'F_0') = s(bool,X270)) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16255,f409])).
fof(f28519,plain,(
  ( ! [X107,X105,X106] : (~p(s(bool,X105)) | sP11(X106,X107) | sP10(X106,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | ($spl8 | $spl46 | $spl122)),
  inference(subsumption_resolution,[],[f28497,f350])).
fof(f28497,plain,(
  ( ! [X107,X105,X106] : (~p(s(bool,X105)) | sP11(X106,X107) | sP10(X106,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | s(bool,'F_0') = s(bool,X105)) ) | ($spl8 | $spl46 | $spl122)),
  inference(superposition,[],[f23560,f15318])).
fof(f28514,plain,(
  ( ! [X549,X550] : (sP11(X549,X550) | sP10(X549,'F_0')) ) | $spl332),
  inference(cnf_transformation,[],[f28514_D])).
fof(f28514_D,plain,(
  ( ! [X549,X550] : (sP11(X549,X550) | sP10(X549,'F_0')) ) <=> ~$spl332),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl332])])).
fof(f28462,plain,(
  ( ! [X557,X559,X560,X558] : (sP11(X557,X558) | sP10(X557,'F_0') | p(s(bool,X559)) | p(s(bool,X560)) | ~sP11(X559,X560)) ) | $spl122),
  inference(resolution,[],[f23560,f454])).
fof(f28460,plain,(
  ( ! [X549,X551,X552,X550] : (sP11(X549,X550) | sP10(X549,'F_0') | p(s(bool,X551)) | p(s(bool,X552)) | ~sP11(X551,X552)) ) | $spl122),
  inference(resolution,[],[f23560,f458])).
fof(f28508,plain,(
  ( ! [X4,X2,X0,X1] : (sP11(X0,X1) | s(bool,X2) != s(bool,X4) | ~p(s(bool,X4)) | sP10(X0,X2)) ) | $spl330),
  inference(cnf_transformation,[],[f28508_D])).
fof(f28508_D,plain,(
  ( ! [X4,X2,X0,X1] : (sP11(X0,X1) | s(bool,X2) != s(bool,X4) | ~p(s(bool,X4)) | sP10(X0,X2)) ) <=> ~$spl330),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl330])])).
fof(f28430,plain,(
  ( ! [X455,X457,X454,X456,X458] : (sP11(X454,X455) | sP10(X454,X456) | p(s(bool,X457)) | ~p(s(bool,X458)) | s(bool,X456) != s(bool,X458)) ) | $spl122),
  inference(resolution,[],[f23560,f536])).
fof(f28339,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (sP11(X42,X43) | sP10(X42,X44) | ~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X44) != s(bool,X46) | s(bool,X45) != s(bool,X47)) ) | $spl122),
  inference(resolution,[],[f23560,f312])).
fof(f28337,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (sP11(X30,X31) | sP10(X30,X32) | p(s(bool,X33)) | ~p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X32) != s(bool,X34)) ) | $spl122),
  inference(resolution,[],[f23560,f311])).
fof(f28335,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (sP11(X18,X19) | sP10(X18,X20) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X20) != s(bool,X22)) ) | $spl122),
  inference(resolution,[],[f23560,f310])).
fof(f28334,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (sP11(X12,X13) | sP10(X12,X14) | ~p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,X17)) ) | $spl122),
  inference(resolution,[],[f23560,f308])).
fof(f28332,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP11(X0,X1) | sP10(X0,X2) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X2) != s(bool,X4)) ) | $spl122),
  inference(resolution,[],[f23560,f306])).
fof(f28285,plain,(
  ( ! [X107,X105,X106] : (~p(s(bool,X105)) | sP11(X106,X107) | sP10(X107,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | ($spl8 | $spl46 | $spl122)),
  inference(subsumption_resolution,[],[f28263,f350])).
fof(f28263,plain,(
  ( ! [X107,X105,X106] : (~p(s(bool,X105)) | sP11(X106,X107) | sP10(X107,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | s(bool,'F_0') = s(bool,X105)) ) | ($spl8 | $spl46 | $spl122)),
  inference(superposition,[],[f23454,f15318])).
fof(f28280,plain,(
  ( ! [X549,X550] : (sP11(X549,X550) | sP10(X550,'F_0')) ) | $spl326),
  inference(cnf_transformation,[],[f28280_D])).
fof(f28280_D,plain,(
  ( ! [X549,X550] : (sP11(X549,X550) | sP10(X550,'F_0')) ) <=> ~$spl326),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl326])])).
fof(f28228,plain,(
  ( ! [X557,X559,X560,X558] : (sP11(X557,X558) | sP10(X558,'F_0') | p(s(bool,X559)) | p(s(bool,X560)) | ~sP11(X559,X560)) ) | $spl122),
  inference(resolution,[],[f23454,f454])).
fof(f28226,plain,(
  ( ! [X549,X551,X552,X550] : (sP11(X549,X550) | sP10(X550,'F_0') | p(s(bool,X551)) | p(s(bool,X552)) | ~sP11(X551,X552)) ) | $spl122),
  inference(resolution,[],[f23454,f458])).
fof(f28274,plain,(
  ( ! [X4,X2,X0,X1] : (sP11(X0,X1) | s(bool,X2) != s(bool,X4) | ~p(s(bool,X4)) | sP10(X1,X2)) ) | $spl324),
  inference(cnf_transformation,[],[f28274_D])).
fof(f28274_D,plain,(
  ( ! [X4,X2,X0,X1] : (sP11(X0,X1) | s(bool,X2) != s(bool,X4) | ~p(s(bool,X4)) | sP10(X1,X2)) ) <=> ~$spl324),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl324])])).
fof(f28196,plain,(
  ( ! [X455,X457,X454,X456,X458] : (sP11(X454,X455) | sP10(X455,X456) | p(s(bool,X457)) | ~p(s(bool,X458)) | s(bool,X456) != s(bool,X458)) ) | $spl122),
  inference(resolution,[],[f23454,f536])).
fof(f28105,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (sP11(X42,X43) | sP10(X43,X44) | ~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X44) != s(bool,X46) | s(bool,X45) != s(bool,X47)) ) | $spl122),
  inference(resolution,[],[f23454,f312])).
fof(f28103,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (sP11(X30,X31) | sP10(X31,X32) | p(s(bool,X33)) | ~p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X32) != s(bool,X34)) ) | $spl122),
  inference(resolution,[],[f23454,f311])).
fof(f28101,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (sP11(X18,X19) | sP10(X19,X20) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X20) != s(bool,X22)) ) | $spl122),
  inference(resolution,[],[f23454,f310])).
fof(f28100,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (sP11(X12,X13) | sP10(X13,X14) | ~p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,X17)) ) | $spl122),
  inference(resolution,[],[f23454,f308])).
fof(f28098,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP11(X0,X1) | sP10(X1,X2) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X2) != s(bool,X4)) ) | $spl122),
  inference(resolution,[],[f23454,f306])).
fof(f28077,plain,(
  ( ! [X107,X105,X106] : (~p(s(bool,X105)) | sP11(X106,X107) | sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X106)) ) | ($spl8 | $spl46 | $spl122)),
  inference(subsumption_resolution,[],[f28055,f350])).
fof(f28055,plain,(
  ( ! [X107,X105,X106] : (~p(s(bool,X105)) | sP11(X106,X107) | sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X106) | s(bool,'F_0') = s(bool,X105)) ) | ($spl8 | $spl46 | $spl122)),
  inference(superposition,[],[f22910,f15318])).
fof(f28072,plain,(
  ( ! [X549,X550] : (sP11(X549,X550) | sP10('F_0',X549)) ) | $spl320),
  inference(cnf_transformation,[],[f28072_D])).
fof(f28072_D,plain,(
  ( ! [X549,X550] : (sP11(X549,X550) | sP10('F_0',X549)) ) <=> ~$spl320),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl320])])).
fof(f28020,plain,(
  ( ! [X557,X559,X560,X558] : (sP11(X557,X558) | sP10('F_0',X557) | p(s(bool,X559)) | p(s(bool,X560)) | ~sP11(X559,X560)) ) | $spl122),
  inference(resolution,[],[f22910,f454])).
fof(f28018,plain,(
  ( ! [X549,X551,X552,X550] : (sP11(X549,X550) | sP10('F_0',X549) | p(s(bool,X551)) | p(s(bool,X552)) | ~sP11(X551,X552)) ) | $spl122),
  inference(resolution,[],[f22910,f458])).
fof(f28066,plain,(
  ( ! [X4,X2,X0,X1] : (sP11(X0,X1) | s(bool,X2) != s(bool,X4) | ~p(s(bool,X4)) | sP10(X2,X0)) ) | $spl318),
  inference(cnf_transformation,[],[f28066_D])).
fof(f28066_D,plain,(
  ( ! [X4,X2,X0,X1] : (sP11(X0,X1) | s(bool,X2) != s(bool,X4) | ~p(s(bool,X4)) | sP10(X2,X0)) ) <=> ~$spl318),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl318])])).
fof(f27988,plain,(
  ( ! [X455,X457,X454,X456,X458] : (sP11(X454,X455) | sP10(X456,X454) | p(s(bool,X457)) | ~p(s(bool,X458)) | s(bool,X456) != s(bool,X458)) ) | $spl122),
  inference(resolution,[],[f22910,f536])).
fof(f27897,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (sP11(X42,X43) | sP10(X44,X42) | ~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X44) != s(bool,X46) | s(bool,X45) != s(bool,X47)) ) | $spl122),
  inference(resolution,[],[f22910,f312])).
fof(f27895,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (sP11(X30,X31) | sP10(X32,X30) | p(s(bool,X33)) | ~p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X32) != s(bool,X34)) ) | $spl122),
  inference(resolution,[],[f22910,f311])).
fof(f27893,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (sP11(X18,X19) | sP10(X20,X18) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X20) != s(bool,X22)) ) | $spl122),
  inference(resolution,[],[f22910,f310])).
fof(f27892,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (sP11(X12,X13) | sP10(X14,X12) | ~p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,X17)) ) | $spl122),
  inference(resolution,[],[f22910,f308])).
fof(f27890,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP11(X0,X1) | sP10(X2,X0) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X2) != s(bool,X4)) ) | $spl122),
  inference(resolution,[],[f22910,f306])).
fof(f27887,plain,(
  ( ! [X107,X105,X106] : (~p(s(bool,X105)) | sP11(X106,X107) | sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X107)) ) | ($spl8 | $spl46 | $spl122)),
  inference(subsumption_resolution,[],[f27865,f350])).
fof(f27865,plain,(
  ( ! [X107,X105,X106] : (~p(s(bool,X105)) | sP11(X106,X107) | sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X107) | s(bool,'F_0') = s(bool,X105)) ) | ($spl8 | $spl46 | $spl122)),
  inference(superposition,[],[f22766,f15318])).
fof(f27882,plain,(
  ( ! [X549,X550] : (sP11(X549,X550) | sP10('F_0',X550)) ) | $spl314),
  inference(cnf_transformation,[],[f27882_D])).
fof(f27882_D,plain,(
  ( ! [X549,X550] : (sP11(X549,X550) | sP10('F_0',X550)) ) <=> ~$spl314),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl314])])).
fof(f27830,plain,(
  ( ! [X557,X559,X560,X558] : (sP11(X557,X558) | sP10('F_0',X558) | p(s(bool,X559)) | p(s(bool,X560)) | ~sP11(X559,X560)) ) | $spl122),
  inference(resolution,[],[f22766,f454])).
fof(f27828,plain,(
  ( ! [X549,X551,X552,X550] : (sP11(X549,X550) | sP10('F_0',X550) | p(s(bool,X551)) | p(s(bool,X552)) | ~sP11(X551,X552)) ) | $spl122),
  inference(resolution,[],[f22766,f458])).
fof(f27876,plain,(
  ( ! [X4,X2,X0,X1] : (sP11(X0,X1) | s(bool,X2) != s(bool,X4) | ~p(s(bool,X4)) | sP10(X2,X1)) ) | $spl312),
  inference(cnf_transformation,[],[f27876_D])).
fof(f27876_D,plain,(
  ( ! [X4,X2,X0,X1] : (sP11(X0,X1) | s(bool,X2) != s(bool,X4) | ~p(s(bool,X4)) | sP10(X2,X1)) ) <=> ~$spl312),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl312])])).
fof(f27798,plain,(
  ( ! [X455,X457,X454,X456,X458] : (sP11(X454,X455) | sP10(X456,X455) | p(s(bool,X457)) | ~p(s(bool,X458)) | s(bool,X456) != s(bool,X458)) ) | $spl122),
  inference(resolution,[],[f22766,f536])).
fof(f27707,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (sP11(X42,X43) | sP10(X44,X43) | ~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X44) != s(bool,X46) | s(bool,X45) != s(bool,X47)) ) | $spl122),
  inference(resolution,[],[f22766,f312])).
fof(f27705,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (sP11(X30,X31) | sP10(X32,X31) | p(s(bool,X33)) | ~p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X32) != s(bool,X34)) ) | $spl122),
  inference(resolution,[],[f22766,f311])).
fof(f27703,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (sP11(X18,X19) | sP10(X20,X19) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X20) != s(bool,X22)) ) | $spl122),
  inference(resolution,[],[f22766,f310])).
fof(f27702,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (sP11(X12,X13) | sP10(X14,X13) | ~p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,X17)) ) | $spl122),
  inference(resolution,[],[f22766,f308])).
fof(f27700,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP11(X0,X1) | sP10(X2,X1) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X2) != s(bool,X4)) ) | $spl122),
  inference(resolution,[],[f22766,f306])).
fof(f27670,plain,(
  $false | ($spl8 | $spl122 | $spl162)),
  inference(subsumption_resolution,[],[f27663,f534])).
fof(f27663,plain,(
  ~p(s(bool,'T_0')) | ($spl8 | $spl122 | $spl162)),
  inference(resolution,[],[f27646,f6468])).
fof(f6468,plain,(
  ( ! [X150] : (~sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X150) | ~p(s(bool,X150))) ) | $spl8),
  inference(resolution,[],[f2511,f439])).
fof(f27646,plain,(
  ( ! [X9] : (sP10(X9,'T_0')) ) | ($spl122 | $spl162)),
  inference(subsumption_resolution,[],[f20653,f20171])).
fof(f20171,plain,(
  ( ! [X86] : (sP10(X86,'T_0') | s(bool,'F_0') != s(bool,X86)) ) | $spl122),
  inference(resolution,[],[f20069,f350])).
fof(f27661,plain,(
  $false | ($spl122 | $spl162 | $spl205)),
  inference(resolution,[],[f27646,f22771])).
fof(f22771,plain,(
  ~sP10('T_0','T_0') | $spl205),
  inference(cnf_transformation,[],[f22771_D])).
fof(f22771_D,plain,(
  ~sP10('T_0','T_0') <=> ~$spl205),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl205])])).
fof(f27645,plain,(
  ( ! [X61,X62,X63] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0') | sP8(X61,X62,X63)) ) | ($spl1 | $spl122)),
  inference(subsumption_resolution,[],[f27623,f409])).
fof(f27623,plain,(
  ( ! [X61,X62,X63] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0') | sP8(X61,X62,X63) | p(s(bool,X63))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f24762,f14262])).
fof(f14262,plain,(
  ( ! [X149] : (~sP10(X149,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | p(s(bool,X149))) ) | $spl1),
  inference(resolution,[],[f545,f438])).
fof(f24069,plain,(
  ( ! [X25] : (sP10(X25,'T_0')) ) | $spl234),
  inference(cnf_transformation,[],[f24069_D])).
fof(f24069_D,plain,(
  ( ! [X25] : (sP10(X25,'T_0')) ) <=> ~$spl234),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl234])])).
fof(f26425,plain,(
  ( ! [X30,X28,X29,X27] : (sP10(X27,X28) | ~p(s(bool,X28)) | sP8(X29,X30,X27)) ) | $spl308),
  inference(cnf_transformation,[],[f26425_D])).
fof(f26425_D,plain,(
  ( ! [X30,X28,X29,X27] : (sP10(X27,X28) | ~p(s(bool,X28)) | sP8(X29,X30,X27)) ) <=> ~$spl308),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl308])])).
fof(f27641,plain,(
  ( ! [X30,X33,X31,X29,X32] : (sP10(X29,X30) | sP8(X31,X32,X29) | ~p(s(bool,X30)) | sP10(X33,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(subsumption_resolution,[],[f27598,f20171])).
fof(f27598,plain,(
  ( ! [X30,X33,X31,X29,X32] : (sP10(X29,X30) | sP8(X31,X32,X29) | s(bool,'F_0') = s(bool,X33) | ~p(s(bool,X30)) | sP10(X33,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(resolution,[],[f24762,f16693])).
fof(f27576,plain,(
  ( ! [X24,X23,X25] : (sP10(X23,'T_0') | sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X23) | p(s(bool,X24)) | ~p(s(bool,X25))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f24746,f14238])).
fof(f14238,plain,(
  ( ! [X101,X102] : (~sP6(X101,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X102) | p(s(bool,X102)) | ~p(s(bool,X101))) ) | $spl1),
  inference(resolution,[],[f545,f391])).
fof(f24746,plain,(
  ( ! [X208,X206,X209,X207] : (sP10(X206,'T_0') | sP10(X207,X206) | sP6(X208,X207,X209)) ) | $spl122),
  inference(resolution,[],[f20223,f393])).
fof(f393,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f236])).
fof(f27589,plain,(
  ( ! [X61,X62,X63] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0') | sP6(X61,X62,X63)) ) | ($spl1 | $spl122)),
  inference(subsumption_resolution,[],[f27565,f393])).
fof(f27565,plain,(
  ( ! [X61,X62,X63] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0') | sP6(X61,X62,X63) | p(s(bool,X62))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f24746,f14262])).
fof(f26321,plain,(
  ( ! [X30,X28,X29,X27] : (sP10(X27,X28) | ~p(s(bool,X28)) | sP6(X29,X27,X30)) ) | $spl304),
  inference(cnf_transformation,[],[f26321_D])).
fof(f26321_D,plain,(
  ( ! [X30,X28,X29,X27] : (sP10(X27,X28) | ~p(s(bool,X28)) | sP6(X29,X27,X30)) ) <=> ~$spl304),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl304])])).
fof(f27585,plain,(
  ( ! [X30,X33,X31,X29,X32] : (sP10(X29,X30) | sP6(X31,X29,X32) | ~p(s(bool,X30)) | sP10(X33,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(subsumption_resolution,[],[f27540,f20171])).
fof(f27540,plain,(
  ( ! [X30,X33,X31,X29,X32] : (sP10(X29,X30) | sP6(X31,X29,X32) | s(bool,'F_0') = s(bool,X33) | ~p(s(bool,X30)) | sP10(X33,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(resolution,[],[f24746,f16693])).
fof(f27516,plain,(
  ( ! [X21,X22,X20] : (sP10(X20,'T_0') | sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X20) | p(s(bool,X21)) | ~p(s(bool,X22))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f24745,f14239])).
fof(f14239,plain,(
  ( ! [X103,X104] : (~sP6(X103,X104,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | p(s(bool,X104)) | ~p(s(bool,X103))) ) | $spl1),
  inference(resolution,[],[f545,f391])).
fof(f24745,plain,(
  ( ! [X204,X202,X205,X203] : (sP10(X202,'T_0') | sP10(X203,X202) | sP6(X204,X205,X203)) ) | $spl122),
  inference(resolution,[],[f20223,f392])).
fof(f392,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f236])).
fof(f27530,plain,(
  ( ! [X61,X62,X63] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0') | sP6(X61,X62,X63)) ) | ($spl1 | $spl122)),
  inference(subsumption_resolution,[],[f27506,f392])).
fof(f27506,plain,(
  ( ! [X61,X62,X63] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0') | sP6(X61,X62,X63) | p(s(bool,X63))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f24745,f14262])).
fof(f26212,plain,(
  ( ! [X30,X28,X29,X27] : (sP10(X27,X28) | ~p(s(bool,X28)) | sP6(X29,X30,X27)) ) | $spl300),
  inference(cnf_transformation,[],[f26212_D])).
fof(f26212_D,plain,(
  ( ! [X30,X28,X29,X27] : (sP10(X27,X28) | ~p(s(bool,X28)) | sP6(X29,X30,X27)) ) <=> ~$spl300),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl300])])).
fof(f27526,plain,(
  ( ! [X30,X33,X31,X29,X32] : (sP10(X29,X30) | sP6(X31,X32,X29) | ~p(s(bool,X30)) | sP10(X33,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(subsumption_resolution,[],[f27481,f20171])).
fof(f27481,plain,(
  ( ! [X30,X33,X31,X29,X32] : (sP10(X29,X30) | sP6(X31,X32,X29) | s(bool,'F_0') = s(bool,X33) | ~p(s(bool,X30)) | sP10(X33,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(resolution,[],[f24745,f16693])).
fof(f27430,plain,(
  ( ! [X24,X23,X25] : (sP10(X23,'T_0') | sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X23) | ~p(s(bool,X24)) | ~p(s(bool,X25))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f24732,f14222])).
fof(f14222,plain,(
  ( ! [X70,X69] : (~sP4(X69,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X70) | ~p(s(bool,X70)) | ~p(s(bool,X69))) ) | $spl1),
  inference(resolution,[],[f545,f374])).
fof(f24732,plain,(
  ( ! [X152,X151,X153,X150] : (sP10(X150,'T_0') | sP10(X151,X150) | sP4(X152,X151,X153)) ) | $spl122),
  inference(resolution,[],[f20223,f376])).
fof(f376,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f228])).
fof(f27443,plain,(
  ( ! [X61,X62,X63] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0') | sP4(X61,X62,X63)) ) | ($spl1 | $spl122)),
  inference(subsumption_resolution,[],[f27419,f376])).
fof(f27419,plain,(
  ( ! [X61,X62,X63] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0') | sP4(X61,X62,X63) | p(s(bool,X62))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f24732,f14262])).
fof(f26100,plain,(
  ( ! [X30,X28,X29,X27] : (sP10(X27,X28) | ~p(s(bool,X28)) | sP4(X29,X27,X30)) ) | $spl296),
  inference(cnf_transformation,[],[f26100_D])).
fof(f26100_D,plain,(
  ( ! [X30,X28,X29,X27] : (sP10(X27,X28) | ~p(s(bool,X28)) | sP4(X29,X27,X30)) ) <=> ~$spl296),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl296])])).
fof(f27439,plain,(
  ( ! [X30,X33,X31,X29,X32] : (sP10(X29,X30) | sP4(X31,X29,X32) | ~p(s(bool,X30)) | sP10(X33,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(subsumption_resolution,[],[f27394,f20171])).
fof(f27394,plain,(
  ( ! [X30,X33,X31,X29,X32] : (sP10(X29,X30) | sP4(X31,X29,X32) | s(bool,'F_0') = s(bool,X33) | ~p(s(bool,X30)) | sP10(X33,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(resolution,[],[f24732,f16693])).
fof(f27370,plain,(
  ( ! [X30,X31,X32] : (sP10(X30,'T_0') | sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X30) | p(s(bool,X31)) | p(s(bool,X32))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f24712,f14203])).
fof(f14203,plain,(
  ( ! [X45,X46] : (~sP0('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X45,X46) | p(s(bool,X46)) | p(s(bool,X45))) ) | $spl1),
  inference(resolution,[],[f545,f328])).
fof(f328,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP0(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f207])).
fof(f24712,plain,(
  ( ! [X87,X85,X86,X84] : (sP10(X84,'T_0') | sP10(X85,X84) | sP0(X85,X86,X87)) ) | $spl122),
  inference(resolution,[],[f20223,f331])).
fof(f331,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f207])).
fof(f25152,plain,(
  ( ! [X70,X68,X69] : (sP0(X68,X69,X70)) ) | $spl256),
  inference(cnf_transformation,[],[f25152_D])).
fof(f25152_D,plain,(
  ( ! [X70,X68,X69] : (sP0(X68,X69,X70)) ) <=> ~$spl256),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl256])])).
fof(f27384,plain,(
  ( ! [X61,X62,X63] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0') | sP0(X61,X62,X63)) ) | ($spl1 | $spl122)),
  inference(subsumption_resolution,[],[f27356,f331])).
fof(f27356,plain,(
  ( ! [X61,X62,X63] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0') | sP0(X61,X62,X63) | p(s(bool,X61))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f24712,f14262])).
fof(f26008,plain,(
  ( ! [X30,X28,X29,X27] : (sP10(X27,X28) | ~p(s(bool,X28)) | sP0(X27,X29,X30)) ) | $spl292),
  inference(cnf_transformation,[],[f26008_D])).
fof(f26008_D,plain,(
  ( ! [X30,X28,X29,X27] : (sP10(X27,X28) | ~p(s(bool,X28)) | sP0(X27,X29,X30)) ) <=> ~$spl292),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl292])])).
fof(f27377,plain,(
  ( ! [X30,X33,X31,X29,X32] : (sP10(X29,X30) | sP0(X29,X31,X32) | ~p(s(bool,X30)) | sP10(X33,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(subsumption_resolution,[],[f27331,f20171])).
fof(f27331,plain,(
  ( ! [X30,X33,X31,X29,X32] : (sP10(X29,X30) | sP0(X29,X31,X32) | s(bool,'F_0') = s(bool,X33) | ~p(s(bool,X30)) | sP10(X33,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(resolution,[],[f24712,f16693])).
fof(f27306,plain,(
  ( ! [X26,X27,X25] : (sP10(X25,'T_0') | sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X25) | p(s(bool,X26)) | p(s(bool,X27))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f24711,f14202])).
fof(f14202,plain,(
  ( ! [X43,X44] : (~sP0(X43,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X44) | p(s(bool,X44)) | p(s(bool,X43))) ) | $spl1),
  inference(resolution,[],[f545,f328])).
fof(f24711,plain,(
  ( ! [X80,X83,X81,X82] : (sP10(X80,'T_0') | sP10(X81,X80) | sP0(X82,X81,X83)) ) | $spl122),
  inference(resolution,[],[f20223,f330])).
fof(f330,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f207])).
fof(f27322,plain,(
  ( ! [X61,X62,X63] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0') | sP0(X61,X62,X63)) ) | ($spl1 | $spl122)),
  inference(subsumption_resolution,[],[f27294,f330])).
fof(f27294,plain,(
  ( ! [X61,X62,X63] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0') | sP0(X61,X62,X63) | p(s(bool,X62))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f24711,f14262])).
fof(f25903,plain,(
  ( ! [X30,X28,X29,X27] : (sP10(X27,X28) | ~p(s(bool,X28)) | sP0(X29,X27,X30)) ) | $spl288),
  inference(cnf_transformation,[],[f25903_D])).
fof(f25903_D,plain,(
  ( ! [X30,X28,X29,X27] : (sP10(X27,X28) | ~p(s(bool,X28)) | sP0(X29,X27,X30)) ) <=> ~$spl288),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl288])])).
fof(f27315,plain,(
  ( ! [X30,X33,X31,X29,X32] : (sP10(X29,X30) | sP0(X31,X29,X32) | ~p(s(bool,X30)) | sP10(X33,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(subsumption_resolution,[],[f27269,f20171])).
fof(f27269,plain,(
  ( ! [X30,X33,X31,X29,X32] : (sP10(X29,X30) | sP0(X31,X29,X32) | s(bool,'F_0') = s(bool,X33) | ~p(s(bool,X30)) | sP10(X33,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(resolution,[],[f24711,f16693])).
fof(f27241,plain,(
  ( ! [X17,X18,X16] : (sP10(X16,'T_0') | sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X16) | p(s(bool,X17)) | p(s(bool,X18))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f24710,f14204])).
fof(f14204,plain,(
  ( ! [X47,X48] : (~sP0(X47,X48,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | p(s(bool,X48)) | p(s(bool,X47))) ) | $spl1),
  inference(resolution,[],[f545,f328])).
fof(f24710,plain,(
  ( ! [X78,X76,X79,X77] : (sP10(X76,'T_0') | sP10(X77,X76) | sP0(X78,X79,X77)) ) | $spl122),
  inference(resolution,[],[f20223,f329])).
fof(f27260,plain,(
  ( ! [X61,X62,X63] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0') | sP0(X61,X62,X63)) ) | ($spl1 | $spl122)),
  inference(subsumption_resolution,[],[f27232,f329])).
fof(f27232,plain,(
  ( ! [X61,X62,X63] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0') | sP0(X61,X62,X63) | p(s(bool,X63))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f24710,f14262])).
fof(f25801,plain,(
  ( ! [X30,X28,X29,X27] : (sP10(X27,X28) | ~p(s(bool,X28)) | sP0(X29,X30,X27)) ) | $spl284),
  inference(cnf_transformation,[],[f25801_D])).
fof(f25801_D,plain,(
  ( ! [X30,X28,X29,X27] : (sP10(X27,X28) | ~p(s(bool,X28)) | sP0(X29,X30,X27)) ) <=> ~$spl284),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl284])])).
fof(f27253,plain,(
  ( ! [X30,X33,X31,X29,X32] : (sP10(X29,X30) | sP0(X31,X32,X29) | ~p(s(bool,X30)) | sP10(X33,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(subsumption_resolution,[],[f27207,f20171])).
fof(f27207,plain,(
  ( ! [X30,X33,X31,X29,X32] : (sP10(X29,X30) | sP0(X31,X32,X29) | s(bool,'F_0') = s(bool,X33) | ~p(s(bool,X30)) | sP10(X33,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(resolution,[],[f24710,f16693])).
fof(f27185,plain,(
  ( ! [X26,X24,X25] : (sP10(X24,'T_0') | sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X24) | ~p(s(bool,X25)) | ~p(s(bool,X26))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f24709,f14199])).
fof(f14199,plain,(
  ( ! [X37,X38] : (~sP1(X37,X38,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | ~p(s(bool,X38)) | ~p(s(bool,X37))) ) | $spl1),
  inference(resolution,[],[f545,f324])).
fof(f24709,plain,(
  ( ! [X74,X72,X75,X73] : (sP10(X72,'T_0') | sP10(X73,X72) | sP1(X74,X75,X73)) ) | $spl122),
  inference(resolution,[],[f20223,f325])).
fof(f325,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f205])).
fof(f27198,plain,(
  ( ! [X61,X62,X63] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0') | sP1(X61,X62,X63)) ) | ($spl1 | $spl122)),
  inference(subsumption_resolution,[],[f27174,f325])).
fof(f27174,plain,(
  ( ! [X61,X62,X63] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0') | sP1(X61,X62,X63) | p(s(bool,X63))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f24709,f14262])).
fof(f25699,plain,(
  ( ! [X30,X28,X29,X27] : (sP10(X27,X28) | ~p(s(bool,X28)) | sP1(X29,X30,X27)) ) | $spl280),
  inference(cnf_transformation,[],[f25699_D])).
fof(f25699_D,plain,(
  ( ! [X30,X28,X29,X27] : (sP10(X27,X28) | ~p(s(bool,X28)) | sP1(X29,X30,X27)) ) <=> ~$spl280),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl280])])).
fof(f27194,plain,(
  ( ! [X30,X33,X31,X29,X32] : (sP10(X29,X30) | sP1(X31,X32,X29) | ~p(s(bool,X30)) | sP10(X33,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(subsumption_resolution,[],[f27149,f20171])).
fof(f27149,plain,(
  ( ! [X30,X33,X31,X29,X32] : (sP10(X29,X30) | sP1(X31,X32,X29) | s(bool,'F_0') = s(bool,X33) | ~p(s(bool,X30)) | sP10(X33,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(resolution,[],[f24709,f16693])).
fof(f27131,plain,(
  ( ! [X54,X56,X55] : (p(s(bool,X54)) | ~p(s(bool,X55)) | sP10(X56,'T_0') | sP10(X56,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f14239,f24509])).
fof(f24509,plain,(
  ( ! [X204,X202,X205,X203] : (sP10(X202,'T_0') | sP10(X202,X203) | sP6(X204,X205,X203)) ) | $spl122),
  inference(resolution,[],[f20222,f392])).
fof(f27129,plain,(
  ( ! [X50,X48,X49] : (p(s(bool,X48)) | ~p(s(bool,X49)) | sP10('T_0',X50) | sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X50)) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f14239,f23265])).
fof(f23265,plain,(
  ( ! [X204,X202,X205,X203] : (sP10('T_0',X202) | sP10(X203,X202) | sP6(X204,X205,X203)) ) | $spl122),
  inference(resolution,[],[f19619,f392])).
fof(f27127,plain,(
  ( ! [X43,X44,X42] : (p(s(bool,X42)) | ~p(s(bool,X43)) | sP10('T_0',X44) | sP10(X44,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f14239,f22586])).
fof(f22586,plain,(
  ( ! [X204,X202,X205,X203] : (sP10('T_0',X202) | sP10(X202,X203) | sP6(X204,X205,X203)) ) | $spl122),
  inference(resolution,[],[f19618,f392])).
fof(f27124,plain,(
  ( ! [X37,X36] : (p(s(bool,X36)) | ~p(s(bool,X37)) | sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0')) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f14239,f20195])).
fof(f20195,plain,(
  ( ! [X152,X151,X150] : (sP6(X151,X152,X150) | sP10(X150,'T_0')) ) | $spl122),
  inference(resolution,[],[f20069,f392])).
fof(f27121,plain,(
  ( ! [X30,X29] : (p(s(bool,X29)) | ~p(s(bool,X30)) | sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f14239,f19591])).
fof(f19591,plain,(
  ( ! [X152,X151,X150] : (sP6(X151,X152,X150) | sP10('T_0',X150)) ) | $spl122),
  inference(resolution,[],[f19515,f392])).
fof(f27117,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X20) | s(bool,'F_0') = s(bool,X20)) ) | ($spl1 | $spl8 | $spl46)),
  inference(resolution,[],[f14239,f16892])).
fof(f16892,plain,(
  ( ! [X204,X202,X205,X203] : (sP10(X202,X203) | sP6(X204,X205,X202) | s(bool,'F_0') = s(bool,X203)) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16258,f392])).
fof(f27115,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | sP10(X14,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | s(bool,'F_0') = s(bool,X14)) ) | ($spl1 | $spl8 | $spl46)),
  inference(resolution,[],[f14239,f16667])).
fof(f16667,plain,(
  ( ! [X204,X202,X205,X203] : (sP10(X202,X203) | sP6(X204,X205,X203) | s(bool,'F_0') = s(bool,X202)) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16255,f392])).
fof(f27094,plain,(
  ( ! [X37,X35,X36] : (sP10(X35,'T_0') | sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X35) | ~p(s(bool,X36)) | ~p(s(bool,X37))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f24706,f14196])).
fof(f14196,plain,(
  ( ! [X31,X32] : (~sP2(X31,X32,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | ~p(s(bool,X32)) | ~p(s(bool,X31))) ) | $spl1),
  inference(resolution,[],[f545,f320])).
fof(f24706,plain,(
  ( ! [X61,X62,X60,X63] : (sP10(X60,'T_0') | sP10(X61,X60) | sP2(X62,X63,X61)) ) | $spl122),
  inference(resolution,[],[f20223,f321])).
fof(f321,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f202])).
fof(f27107,plain,(
  ( ! [X61,X62,X63] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0') | sP2(X61,X62,X63)) ) | ($spl1 | $spl122)),
  inference(subsumption_resolution,[],[f27079,f321])).
fof(f27079,plain,(
  ( ! [X61,X62,X63] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0') | sP2(X61,X62,X63) | p(s(bool,X63))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f24706,f14262])).
fof(f25629,plain,(
  ( ! [X30,X28,X29,X27] : (sP10(X27,X28) | ~p(s(bool,X28)) | sP2(X29,X30,X27)) ) | $spl276),
  inference(cnf_transformation,[],[f25629_D])).
fof(f25629_D,plain,(
  ( ! [X30,X28,X29,X27] : (sP10(X27,X28) | ~p(s(bool,X28)) | sP2(X29,X30,X27)) ) <=> ~$spl276),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl276])])).
fof(f27103,plain,(
  ( ! [X30,X33,X31,X29,X32] : (sP10(X29,X30) | sP2(X31,X32,X29) | ~p(s(bool,X30)) | sP10(X33,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(subsumption_resolution,[],[f27054,f20171])).
fof(f27054,plain,(
  ( ! [X30,X33,X31,X29,X32] : (sP10(X29,X30) | sP2(X31,X32,X29) | s(bool,'F_0') = s(bool,X33) | ~p(s(bool,X30)) | sP10(X33,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(resolution,[],[f24706,f16693])).
fof(f27046,plain,(
  ( ! [X70,X68,X69] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0') | sP8(X68,X69,X70)) ) | ($spl1 | $spl122)),
  inference(subsumption_resolution,[],[f27027,f409])).
fof(f27027,plain,(
  ( ! [X70,X68,X69] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0') | sP8(X68,X69,X70) | p(s(bool,X70))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f24526,f14263])).
fof(f14263,plain,(
  ( ! [X150] : (~sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X150) | p(s(bool,X150))) ) | $spl1),
  inference(resolution,[],[f545,f438])).
fof(f25565,plain,(
  ( ! [X30,X28,X29,X27] : (sP10(X27,X28) | ~p(s(bool,X27)) | sP8(X29,X30,X28)) ) | $spl272),
  inference(cnf_transformation,[],[f25565_D])).
fof(f25565_D,plain,(
  ( ! [X30,X28,X29,X27] : (sP10(X27,X28) | ~p(s(bool,X27)) | sP8(X29,X30,X28)) ) <=> ~$spl272),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl272])])).
fof(f27042,plain,(
  ( ! [X30,X33,X31,X29,X32] : (sP10(X29,X30) | sP8(X31,X32,X30) | ~p(s(bool,X29)) | sP10(X33,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(subsumption_resolution,[],[f26999,f20171])).
fof(f26999,plain,(
  ( ! [X30,X33,X31,X29,X32] : (sP10(X29,X30) | sP8(X31,X32,X30) | s(bool,'F_0') = s(bool,X33) | ~p(s(bool,X29)) | sP10(X33,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(resolution,[],[f24526,f16693])).
fof(f26977,plain,(
  ( ! [X21,X22,X20] : (sP10(X20,'T_0') | sP10(X20,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | p(s(bool,X21)) | ~p(s(bool,X22))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f24510,f14238])).
fof(f24510,plain,(
  ( ! [X208,X206,X209,X207] : (sP10(X206,'T_0') | sP10(X206,X207) | sP6(X208,X207,X209)) ) | $spl122),
  inference(resolution,[],[f20222,f393])).
fof(f26990,plain,(
  ( ! [X70,X68,X69] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0') | sP6(X68,X69,X70)) ) | ($spl1 | $spl122)),
  inference(subsumption_resolution,[],[f26970,f393])).
fof(f26970,plain,(
  ( ! [X70,X68,X69] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0') | sP6(X68,X69,X70) | p(s(bool,X69))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f24510,f14263])).
fof(f25476,plain,(
  ( ! [X30,X28,X29,X27] : (sP10(X27,X28) | ~p(s(bool,X27)) | sP6(X29,X28,X30)) ) | $spl270),
  inference(cnf_transformation,[],[f25476_D])).
fof(f25476_D,plain,(
  ( ! [X30,X28,X29,X27] : (sP10(X27,X28) | ~p(s(bool,X27)) | sP6(X29,X28,X30)) ) <=> ~$spl270),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl270])])).
fof(f26986,plain,(
  ( ! [X30,X33,X31,X29,X32] : (sP10(X29,X30) | sP6(X31,X30,X32) | ~p(s(bool,X29)) | sP10(X33,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(subsumption_resolution,[],[f26942,f20171])).
fof(f26942,plain,(
  ( ! [X30,X33,X31,X29,X32] : (sP10(X29,X30) | sP6(X31,X30,X32) | s(bool,'F_0') = s(bool,X33) | ~p(s(bool,X29)) | sP10(X33,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(resolution,[],[f24510,f16693])).
fof(f26932,plain,(
  ( ! [X70,X68,X69] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0') | sP6(X68,X69,X70)) ) | ($spl1 | $spl122)),
  inference(subsumption_resolution,[],[f26912,f392])).
fof(f26912,plain,(
  ( ! [X70,X68,X69] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0') | sP6(X68,X69,X70) | p(s(bool,X70))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f24509,f14263])).
fof(f25416,plain,(
  ( ! [X30,X28,X29,X27] : (sP10(X27,X28) | ~p(s(bool,X27)) | sP6(X29,X30,X28)) ) | $spl266),
  inference(cnf_transformation,[],[f25416_D])).
fof(f25416_D,plain,(
  ( ! [X30,X28,X29,X27] : (sP10(X27,X28) | ~p(s(bool,X27)) | sP6(X29,X30,X28)) ) <=> ~$spl266),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl266])])).
fof(f26928,plain,(
  ( ! [X30,X33,X31,X29,X32] : (sP10(X29,X30) | sP6(X31,X32,X30) | ~p(s(bool,X29)) | sP10(X33,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(subsumption_resolution,[],[f26884,f20171])).
fof(f26884,plain,(
  ( ! [X30,X33,X31,X29,X32] : (sP10(X29,X30) | sP6(X31,X32,X30) | s(bool,'F_0') = s(bool,X33) | ~p(s(bool,X29)) | sP10(X33,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(resolution,[],[f24509,f16693])).
fof(f26862,plain,(
  ( ! [X24,X23,X25] : (sP10(X23,'T_0') | sP10(X23,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | ~p(s(bool,X24)) | ~p(s(bool,X25))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f24496,f14222])).
fof(f24496,plain,(
  ( ! [X152,X151,X153,X150] : (sP10(X150,'T_0') | sP10(X150,X151) | sP4(X152,X151,X153)) ) | $spl122),
  inference(resolution,[],[f20222,f376])).
fof(f26875,plain,(
  ( ! [X70,X68,X69] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0') | sP4(X68,X69,X70)) ) | ($spl1 | $spl122)),
  inference(subsumption_resolution,[],[f26854,f376])).
fof(f26854,plain,(
  ( ! [X70,X68,X69] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0') | sP4(X68,X69,X70) | p(s(bool,X69))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f24496,f14263])).
fof(f25357,plain,(
  ( ! [X30,X28,X29,X27] : (sP10(X27,X28) | ~p(s(bool,X27)) | sP4(X29,X28,X30)) ) | $spl262),
  inference(cnf_transformation,[],[f25357_D])).
fof(f25357_D,plain,(
  ( ! [X30,X28,X29,X27] : (sP10(X27,X28) | ~p(s(bool,X27)) | sP4(X29,X28,X30)) ) <=> ~$spl262),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl262])])).
fof(f26871,plain,(
  ( ! [X30,X33,X31,X29,X32] : (sP10(X29,X30) | sP4(X31,X30,X32) | ~p(s(bool,X29)) | sP10(X33,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(subsumption_resolution,[],[f26826,f20171])).
fof(f26826,plain,(
  ( ! [X30,X33,X31,X29,X32] : (sP10(X29,X30) | sP4(X31,X30,X32) | s(bool,'F_0') = s(bool,X33) | ~p(s(bool,X29)) | sP10(X33,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(resolution,[],[f24496,f16693])).
fof(f26809,plain,(
  ( ! [X52,X53,X51] : (p(s(bool,X51)) | ~p(s(bool,X52)) | sP10('T_0',X53) | sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X53)) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f14238,f23266])).
fof(f23266,plain,(
  ( ! [X208,X206,X209,X207] : (sP10('T_0',X206) | sP10(X207,X206) | sP6(X208,X207,X209)) ) | $spl122),
  inference(resolution,[],[f19619,f393])).
fof(f26807,plain,(
  ( ! [X47,X45,X46] : (p(s(bool,X45)) | ~p(s(bool,X46)) | sP10('T_0',X47) | sP10(X47,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f14238,f22587])).
fof(f22587,plain,(
  ( ! [X208,X206,X209,X207] : (sP10('T_0',X206) | sP10(X206,X207) | sP6(X208,X207,X209)) ) | $spl122),
  inference(resolution,[],[f19618,f393])).
fof(f26804,plain,(
  ( ! [X39,X38] : (p(s(bool,X38)) | ~p(s(bool,X39)) | sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0')) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f14238,f20196])).
fof(f20196,plain,(
  ( ! [X154,X155,X153] : (sP6(X154,X153,X155) | sP10(X153,'T_0')) ) | $spl122),
  inference(resolution,[],[f20069,f393])).
fof(f26801,plain,(
  ( ! [X31,X32] : (p(s(bool,X31)) | ~p(s(bool,X32)) | sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f14238,f19592])).
fof(f19592,plain,(
  ( ! [X154,X155,X153] : (sP6(X154,X153,X155) | sP10('T_0',X153)) ) | $spl122),
  inference(resolution,[],[f19515,f393])).
fof(f26797,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X23) | s(bool,'F_0') = s(bool,X23)) ) | ($spl1 | $spl8 | $spl46)),
  inference(resolution,[],[f14238,f16893])).
fof(f16893,plain,(
  ( ! [X208,X206,X209,X207] : (sP10(X206,X207) | sP6(X208,X206,X209) | s(bool,'F_0') = s(bool,X207)) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16258,f393])).
fof(f26795,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | sP10(X17,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | s(bool,'F_0') = s(bool,X17)) ) | ($spl1 | $spl8 | $spl46)),
  inference(resolution,[],[f14238,f16668])).
fof(f16668,plain,(
  ( ! [X208,X206,X209,X207] : (sP10(X206,X207) | sP6(X208,X207,X209) | s(bool,'F_0') = s(bool,X206)) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16255,f393])).
fof(f26771,plain,(
  ( ! [X30,X31,X32] : (sP10(X30,'T_0') | sP10(X30,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | p(s(bool,X31)) | p(s(bool,X32))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f24476,f14203])).
fof(f24476,plain,(
  ( ! [X87,X85,X86,X84] : (sP10(X84,'T_0') | sP10(X84,X85) | sP0(X85,X86,X87)) ) | $spl122),
  inference(resolution,[],[f20222,f331])).
fof(f26785,plain,(
  ( ! [X70,X68,X69] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0') | sP0(X68,X69,X70)) ) | ($spl1 | $spl122)),
  inference(subsumption_resolution,[],[f26760,f331])).
fof(f26760,plain,(
  ( ! [X70,X68,X69] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0') | sP0(X68,X69,X70) | p(s(bool,X68))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f24476,f14263])).
fof(f25295,plain,(
  ( ! [X30,X28,X29,X27] : (sP10(X27,X28) | ~p(s(bool,X27)) | sP0(X28,X29,X30)) ) | $spl260),
  inference(cnf_transformation,[],[f25295_D])).
fof(f25295_D,plain,(
  ( ! [X30,X28,X29,X27] : (sP10(X27,X28) | ~p(s(bool,X27)) | sP0(X28,X29,X30)) ) <=> ~$spl260),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl260])])).
fof(f26778,plain,(
  ( ! [X30,X33,X31,X29,X32] : (sP10(X29,X30) | sP0(X30,X31,X32) | ~p(s(bool,X29)) | sP10(X33,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(subsumption_resolution,[],[f26732,f20171])).
fof(f26732,plain,(
  ( ! [X30,X33,X31,X29,X32] : (sP10(X29,X30) | sP0(X30,X31,X32) | s(bool,'F_0') = s(bool,X33) | ~p(s(bool,X29)) | sP10(X33,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(resolution,[],[f24476,f16693])).
fof(f26707,plain,(
  ( ! [X26,X27,X25] : (sP10(X25,'T_0') | sP10(X25,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | p(s(bool,X26)) | p(s(bool,X27))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f24475,f14202])).
fof(f24475,plain,(
  ( ! [X80,X83,X81,X82] : (sP10(X80,'T_0') | sP10(X80,X81) | sP0(X82,X81,X83)) ) | $spl122),
  inference(resolution,[],[f20222,f330])).
fof(f26723,plain,(
  ( ! [X70,X68,X69] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0') | sP0(X68,X69,X70)) ) | ($spl1 | $spl122)),
  inference(subsumption_resolution,[],[f26698,f330])).
fof(f26698,plain,(
  ( ! [X70,X68,X69] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0') | sP0(X68,X69,X70) | p(s(bool,X69))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f24475,f14263])).
fof(f25235,plain,(
  ( ! [X30,X28,X29,X27] : (sP10(X27,X28) | ~p(s(bool,X27)) | sP0(X29,X28,X30)) ) | $spl258),
  inference(cnf_transformation,[],[f25235_D])).
fof(f25235_D,plain,(
  ( ! [X30,X28,X29,X27] : (sP10(X27,X28) | ~p(s(bool,X27)) | sP0(X29,X28,X30)) ) <=> ~$spl258),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl258])])).
fof(f26716,plain,(
  ( ! [X30,X33,X31,X29,X32] : (sP10(X29,X30) | sP0(X31,X30,X32) | ~p(s(bool,X29)) | sP10(X33,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(subsumption_resolution,[],[f26670,f20171])).
fof(f26670,plain,(
  ( ! [X30,X33,X31,X29,X32] : (sP10(X29,X30) | sP0(X31,X30,X32) | s(bool,'F_0') = s(bool,X33) | ~p(s(bool,X29)) | sP10(X33,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(resolution,[],[f24475,f16693])).
fof(f26642,plain,(
  ( ! [X17,X18,X16] : (sP10(X16,'T_0') | sP10(X16,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | p(s(bool,X17)) | p(s(bool,X18))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f24474,f14204])).
fof(f24474,plain,(
  ( ! [X78,X76,X79,X77] : (sP10(X76,'T_0') | sP10(X76,X77) | sP0(X78,X79,X77)) ) | $spl122),
  inference(resolution,[],[f20222,f329])).
fof(f26661,plain,(
  ( ! [X70,X68,X69] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0') | sP0(X68,X69,X70)) ) | ($spl1 | $spl122)),
  inference(subsumption_resolution,[],[f26636,f329])).
fof(f26636,plain,(
  ( ! [X70,X68,X69] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0') | sP0(X68,X69,X70) | p(s(bool,X70))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f24474,f14263])).
fof(f25143,plain,(
  ( ! [X30,X28,X29,X27] : (sP10(X27,X28) | ~p(s(bool,X27)) | sP0(X29,X30,X28)) ) | $spl254),
  inference(cnf_transformation,[],[f25143_D])).
fof(f25143_D,plain,(
  ( ! [X30,X28,X29,X27] : (sP10(X27,X28) | ~p(s(bool,X27)) | sP0(X29,X30,X28)) ) <=> ~$spl254),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl254])])).
fof(f26654,plain,(
  ( ! [X30,X33,X31,X29,X32] : (sP10(X29,X30) | sP0(X31,X32,X30) | ~p(s(bool,X29)) | sP10(X33,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(subsumption_resolution,[],[f26608,f20171])).
fof(f26608,plain,(
  ( ! [X30,X33,X31,X29,X32] : (sP10(X29,X30) | sP0(X31,X32,X30) | s(bool,'F_0') = s(bool,X33) | ~p(s(bool,X29)) | sP10(X33,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(resolution,[],[f24474,f16693])).
fof(f26586,plain,(
  ( ! [X26,X24,X25] : (sP10(X24,'T_0') | sP10(X24,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | ~p(s(bool,X25)) | ~p(s(bool,X26))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f24473,f14199])).
fof(f24473,plain,(
  ( ! [X74,X72,X75,X73] : (sP10(X72,'T_0') | sP10(X72,X73) | sP1(X74,X75,X73)) ) | $spl122),
  inference(resolution,[],[f20222,f325])).
fof(f26599,plain,(
  ( ! [X70,X68,X69] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0') | sP1(X68,X69,X70)) ) | ($spl1 | $spl122)),
  inference(subsumption_resolution,[],[f26578,f325])).
fof(f26578,plain,(
  ( ! [X70,X68,X69] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0') | sP1(X68,X69,X70) | p(s(bool,X70))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f24473,f14263])).
fof(f25084,plain,(
  ( ! [X30,X28,X29,X27] : (sP10(X27,X28) | ~p(s(bool,X27)) | sP1(X29,X30,X28)) ) | $spl250),
  inference(cnf_transformation,[],[f25084_D])).
fof(f25084_D,plain,(
  ( ! [X30,X28,X29,X27] : (sP10(X27,X28) | ~p(s(bool,X27)) | sP1(X29,X30,X28)) ) <=> ~$spl250),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl250])])).
fof(f26595,plain,(
  ( ! [X30,X33,X31,X29,X32] : (sP10(X29,X30) | sP1(X31,X32,X30) | ~p(s(bool,X29)) | sP10(X33,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(subsumption_resolution,[],[f26550,f20171])).
fof(f26550,plain,(
  ( ! [X30,X33,X31,X29,X32] : (sP10(X29,X30) | sP1(X31,X32,X30) | s(bool,'F_0') = s(bool,X33) | ~p(s(bool,X29)) | sP10(X33,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(resolution,[],[f24473,f16693])).
fof(f26528,plain,(
  ( ! [X37,X35,X36] : (sP10(X35,'T_0') | sP10(X35,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | ~p(s(bool,X36)) | ~p(s(bool,X37))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f24470,f14196])).
fof(f24470,plain,(
  ( ! [X61,X62,X60,X63] : (sP10(X60,'T_0') | sP10(X60,X61) | sP2(X62,X63,X61)) ) | $spl122),
  inference(resolution,[],[f20222,f321])).
fof(f26541,plain,(
  ( ! [X70,X68,X69] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0') | sP2(X68,X69,X70)) ) | ($spl1 | $spl122)),
  inference(subsumption_resolution,[],[f26516,f321])).
fof(f26516,plain,(
  ( ! [X70,X68,X69] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0') | sP2(X68,X69,X70) | p(s(bool,X70))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f24470,f14263])).
fof(f25024,plain,(
  ( ! [X30,X28,X29,X27] : (sP10(X27,X28) | ~p(s(bool,X27)) | sP2(X29,X30,X28)) ) | $spl246),
  inference(cnf_transformation,[],[f25024_D])).
fof(f25024_D,plain,(
  ( ! [X30,X28,X29,X27] : (sP10(X27,X28) | ~p(s(bool,X27)) | sP2(X29,X30,X28)) ) <=> ~$spl246),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl246])])).
fof(f26537,plain,(
  ( ! [X30,X33,X31,X29,X32] : (sP10(X29,X30) | sP2(X31,X32,X30) | ~p(s(bool,X29)) | sP10(X33,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(subsumption_resolution,[],[f26488,f20171])).
fof(f26488,plain,(
  ( ! [X30,X33,X31,X29,X32] : (sP10(X29,X30) | sP2(X31,X32,X30) | s(bool,'F_0') = s(bool,X33) | ~p(s(bool,X29)) | sP10(X33,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(resolution,[],[f24470,f16693])).
fof(f26432,plain,(
  ( ! [X2,X1] : (sP8(X1,X2,'T_0')) ) | $spl310),
  inference(cnf_transformation,[],[f26432_D])).
fof(f26432_D,plain,(
  ( ! [X2,X1] : (sP8(X1,X2,'T_0')) ) <=> ~$spl310),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl310])])).
fof(f23176,plain,(
  ( ! [X15] : (sP10('T_0',X15)) ) | $spl214),
  inference(cnf_transformation,[],[f23176_D])).
fof(f23176_D,plain,(
  ( ! [X15] : (sP10('T_0',X15)) ) <=> ~$spl214),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl214])])).
fof(f26419,plain,(
  ( ! [X2,X0,X1] : (sP10('T_0',X0) | sP8(X1,X2,'T_0')) ) | $spl122),
  inference(factoring,[],[f23282])).
fof(f26431,plain,(
  ( ! [X61,X62,X63] : (sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP8(X61,X62,X63)) ) | ($spl1 | $spl122)),
  inference(subsumption_resolution,[],[f26407,f409])).
fof(f26407,plain,(
  ( ! [X61,X62,X63] : (sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP8(X61,X62,X63) | p(s(bool,X63))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f23282,f14262])).
fof(f26424,plain,(
  ( ! [X30,X28,X31,X29,X27] : (sP10(X27,X28) | sP8(X29,X30,X27) | ~p(s(bool,X28)) | sP10(X31,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(subsumption_resolution,[],[f26382,f20171])).
fof(f26382,plain,(
  ( ! [X30,X28,X31,X29,X27] : (sP10(X27,X28) | sP8(X29,X30,X27) | s(bool,'F_0') = s(bool,X31) | ~p(s(bool,X28)) | sP10(X31,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(resolution,[],[f23282,f16692])).
fof(f26329,plain,(
  ( ! [X2,X1] : (sP6(X1,'T_0',X2)) ) | $spl306),
  inference(cnf_transformation,[],[f26329_D])).
fof(f26329_D,plain,(
  ( ! [X2,X1] : (sP6(X1,'T_0',X2)) ) <=> ~$spl306),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl306])])).
fof(f26315,plain,(
  ( ! [X2,X0,X1] : (sP10('T_0',X0) | sP6(X1,'T_0',X2)) ) | $spl122),
  inference(factoring,[],[f23266])).
fof(f26327,plain,(
  ( ! [X61,X62,X63] : (sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP6(X61,X62,X63)) ) | ($spl1 | $spl122)),
  inference(subsumption_resolution,[],[f26303,f393])).
fof(f26303,plain,(
  ( ! [X61,X62,X63] : (sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP6(X61,X62,X63) | p(s(bool,X62))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f23266,f14262])).
fof(f26320,plain,(
  ( ! [X30,X28,X31,X29,X27] : (sP10(X27,X28) | sP6(X29,X27,X30) | ~p(s(bool,X28)) | sP10(X31,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(subsumption_resolution,[],[f26278,f20171])).
fof(f26278,plain,(
  ( ! [X30,X28,X31,X29,X27] : (sP10(X27,X28) | sP6(X29,X27,X30) | s(bool,'F_0') = s(bool,X31) | ~p(s(bool,X28)) | sP10(X31,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(resolution,[],[f23266,f16692])).
fof(f26221,plain,(
  ( ! [X2,X1] : (sP6(X1,X2,'T_0')) ) | $spl302),
  inference(cnf_transformation,[],[f26221_D])).
fof(f26221_D,plain,(
  ( ! [X2,X1] : (sP6(X1,X2,'T_0')) ) <=> ~$spl302),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl302])])).
fof(f26206,plain,(
  ( ! [X2,X0,X1] : (sP10('T_0',X0) | sP6(X1,X2,'T_0')) ) | $spl122),
  inference(factoring,[],[f23265])).
fof(f26218,plain,(
  ( ! [X61,X62,X63] : (sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP6(X61,X62,X63)) ) | ($spl1 | $spl122)),
  inference(subsumption_resolution,[],[f26194,f392])).
fof(f26194,plain,(
  ( ! [X61,X62,X63] : (sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP6(X61,X62,X63) | p(s(bool,X63))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f23265,f14262])).
fof(f26211,plain,(
  ( ! [X30,X28,X31,X29,X27] : (sP10(X27,X28) | sP6(X29,X30,X27) | ~p(s(bool,X28)) | sP10(X31,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(subsumption_resolution,[],[f26169,f20171])).
fof(f26169,plain,(
  ( ! [X30,X28,X31,X29,X27] : (sP10(X27,X28) | sP6(X29,X30,X27) | s(bool,'F_0') = s(bool,X31) | ~p(s(bool,X28)) | sP10(X31,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(resolution,[],[f23265,f16692])).
fof(f26136,plain,(
  ( ! [X45,X43,X44] : (~p(s(bool,X43)) | ~p(s(bool,X44)) | sP10('T_0',X45) | sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X45)) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f14222,f23252])).
fof(f23252,plain,(
  ( ! [X152,X151,X153,X150] : (sP10('T_0',X150) | sP10(X151,X150) | sP4(X152,X151,X153)) ) | $spl122),
  inference(resolution,[],[f19619,f376])).
fof(f26135,plain,(
  ( ! [X41,X42,X40] : (~p(s(bool,X40)) | ~p(s(bool,X41)) | sP10('T_0',X42) | sP10(X42,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f14222,f22573])).
fof(f22573,plain,(
  ( ! [X152,X151,X153,X150] : (sP10('T_0',X150) | sP10(X150,X151) | sP4(X152,X151,X153)) ) | $spl122),
  inference(resolution,[],[f19618,f376])).
fof(f26132,plain,(
  ( ! [X35,X34] : (~p(s(bool,X34)) | ~p(s(bool,X35)) | sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0')) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f14222,f20182])).
fof(f20182,plain,(
  ( ! [X111,X112,X113] : (sP4(X112,X111,X113) | sP10(X111,'T_0')) ) | $spl122),
  inference(resolution,[],[f20069,f376])).
fof(f26129,plain,(
  ( ! [X26,X27] : (~p(s(bool,X26)) | ~p(s(bool,X27)) | sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f14222,f19578])).
fof(f19578,plain,(
  ( ! [X111,X112,X113] : (sP4(X112,X111,X113) | sP10('T_0',X111)) ) | $spl122),
  inference(resolution,[],[f19515,f376])).
fof(f26124,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,X13)) | ~p(s(bool,X14)) | sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X15) | s(bool,'F_0') = s(bool,X15)) ) | ($spl1 | $spl8 | $spl46)),
  inference(resolution,[],[f14222,f16879])).
fof(f16879,plain,(
  ( ! [X152,X151,X153,X150] : (sP10(X150,X151) | sP4(X152,X150,X153) | s(bool,'F_0') = s(bool,X151)) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16258,f376])).
fof(f26123,plain,(
  ( ! [X12,X10,X11] : (~p(s(bool,X10)) | ~p(s(bool,X11)) | sP10(X12,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | s(bool,'F_0') = s(bool,X12)) ) | ($spl1 | $spl8 | $spl46)),
  inference(resolution,[],[f14222,f16654])).
fof(f16654,plain,(
  ( ! [X152,X151,X153,X150] : (sP10(X150,X151) | sP4(X152,X151,X153) | s(bool,'F_0') = s(bool,X150)) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16255,f376])).
fof(f26108,plain,(
  ( ! [X2,X1] : (sP4(X1,'T_0',X2)) ) | $spl298),
  inference(cnf_transformation,[],[f26108_D])).
fof(f26108_D,plain,(
  ( ! [X2,X1] : (sP4(X1,'T_0',X2)) ) <=> ~$spl298),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl298])])).
fof(f26094,plain,(
  ( ! [X2,X0,X1] : (sP10('T_0',X0) | sP4(X1,'T_0',X2)) ) | $spl122),
  inference(factoring,[],[f23252])).
fof(f26106,plain,(
  ( ! [X61,X62,X63] : (sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP4(X61,X62,X63)) ) | ($spl1 | $spl122)),
  inference(subsumption_resolution,[],[f26081,f376])).
fof(f26081,plain,(
  ( ! [X61,X62,X63] : (sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP4(X61,X62,X63) | p(s(bool,X62))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f23252,f14262])).
fof(f26099,plain,(
  ( ! [X30,X28,X31,X29,X27] : (sP10(X27,X28) | sP4(X29,X27,X30) | ~p(s(bool,X28)) | sP10(X31,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(subsumption_resolution,[],[f26056,f20171])).
fof(f26056,plain,(
  ( ! [X30,X28,X31,X29,X27] : (sP10(X27,X28) | sP4(X29,X27,X30) | s(bool,'F_0') = s(bool,X31) | ~p(s(bool,X28)) | sP10(X31,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(resolution,[],[f23252,f16692])).
fof(f26019,plain,(
  ( ! [X2,X1] : (sP0('T_0',X1,X2)) ) | $spl294),
  inference(cnf_transformation,[],[f26019_D])).
fof(f26019_D,plain,(
  ( ! [X2,X1] : (sP0('T_0',X1,X2)) ) <=> ~$spl294),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl294])])).
fof(f26002,plain,(
  ( ! [X2,X0,X1] : (sP10('T_0',X0) | sP0('T_0',X1,X2)) ) | $spl122),
  inference(factoring,[],[f23232])).
fof(f23232,plain,(
  ( ! [X87,X85,X86,X84] : (sP10('T_0',X84) | sP10(X85,X84) | sP0(X85,X86,X87)) ) | $spl122),
  inference(resolution,[],[f19619,f331])).
fof(f26001,plain,(
  ( ! [X30,X31,X32] : (sP10('T_0',X30) | sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X30) | p(s(bool,X31)) | p(s(bool,X32))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f23232,f14203])).
fof(f26016,plain,(
  ( ! [X61,X62,X63] : (sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP0(X61,X62,X63)) ) | ($spl1 | $spl122)),
  inference(subsumption_resolution,[],[f25987,f331])).
fof(f25987,plain,(
  ( ! [X61,X62,X63] : (sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP0(X61,X62,X63) | p(s(bool,X61))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f23232,f14262])).
fof(f26007,plain,(
  ( ! [X30,X28,X31,X29,X27] : (sP10(X27,X28) | sP0(X27,X29,X30) | ~p(s(bool,X28)) | sP10(X31,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(subsumption_resolution,[],[f25962,f20171])).
fof(f25962,plain,(
  ( ! [X30,X28,X31,X29,X27] : (sP10(X27,X28) | sP0(X27,X29,X30) | s(bool,'F_0') = s(bool,X31) | ~p(s(bool,X28)) | sP10(X31,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(resolution,[],[f23232,f16692])).
fof(f25943,plain,(
  ( ! [X54,X55,X53] : (p(s(bool,X53)) | p(s(bool,X54)) | sP10('T_0',X55) | sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X55)) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f14204,f23230])).
fof(f23230,plain,(
  ( ! [X78,X76,X79,X77] : (sP10('T_0',X76) | sP10(X77,X76) | sP0(X78,X79,X77)) ) | $spl122),
  inference(resolution,[],[f19619,f329])).
fof(f25940,plain,(
  ( ! [X45,X46,X44] : (p(s(bool,X44)) | p(s(bool,X45)) | sP10('T_0',X46) | sP10(X46,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f14204,f22551])).
fof(f22551,plain,(
  ( ! [X78,X76,X79,X77] : (sP10('T_0',X76) | sP10(X76,X77) | sP0(X78,X79,X77)) ) | $spl122),
  inference(resolution,[],[f19618,f329])).
fof(f25937,plain,(
  ( ! [X39,X38] : (p(s(bool,X38)) | p(s(bool,X39)) | sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0')) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f14204,f20160])).
fof(f20160,plain,(
  ( ! [X61,X59,X60] : (sP0(X60,X61,X59) | sP10(X59,'T_0')) ) | $spl122),
  inference(resolution,[],[f20069,f329])).
fof(f25934,plain,(
  ( ! [X33,X32] : (p(s(bool,X32)) | p(s(bool,X33)) | sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f14204,f19556])).
fof(f19556,plain,(
  ( ! [X61,X59,X60] : (sP0(X60,X61,X59) | sP10('T_0',X59)) ) | $spl122),
  inference(resolution,[],[f19515,f329])).
fof(f25931,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | p(s(bool,X24)) | sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X25) | s(bool,'F_0') = s(bool,X25)) ) | ($spl1 | $spl8 | $spl46)),
  inference(resolution,[],[f14204,f16857])).
fof(f16857,plain,(
  ( ! [X78,X76,X79,X77] : (sP10(X76,X77) | sP0(X78,X79,X76) | s(bool,'F_0') = s(bool,X77)) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16258,f329])).
fof(f25928,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | p(s(bool,X15)) | sP10(X16,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | s(bool,'F_0') = s(bool,X16)) ) | ($spl1 | $spl8 | $spl46)),
  inference(resolution,[],[f14204,f16632])).
fof(f16632,plain,(
  ( ! [X78,X76,X79,X77] : (sP10(X76,X77) | sP0(X78,X79,X77) | s(bool,'F_0') = s(bool,X76)) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16255,f329])).
fof(f25913,plain,(
  ( ! [X2,X1] : (sP0(X1,'T_0',X2)) ) | $spl290),
  inference(cnf_transformation,[],[f25913_D])).
fof(f25913_D,plain,(
  ( ! [X2,X1] : (sP0(X1,'T_0',X2)) ) <=> ~$spl290),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl290])])).
fof(f25897,plain,(
  ( ! [X2,X0,X1] : (sP10('T_0',X0) | sP0(X1,'T_0',X2)) ) | $spl122),
  inference(factoring,[],[f23231])).
fof(f23231,plain,(
  ( ! [X80,X83,X81,X82] : (sP10('T_0',X80) | sP10(X81,X80) | sP0(X82,X81,X83)) ) | $spl122),
  inference(resolution,[],[f19619,f330])).
fof(f25894,plain,(
  ( ! [X24,X23,X22] : (sP10('T_0',X22) | sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X22) | p(s(bool,X23)) | p(s(bool,X24))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f23231,f14202])).
fof(f25911,plain,(
  ( ! [X61,X62,X63] : (sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP0(X61,X62,X63)) ) | ($spl1 | $spl122)),
  inference(subsumption_resolution,[],[f25883,f330])).
fof(f25883,plain,(
  ( ! [X61,X62,X63] : (sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP0(X61,X62,X63) | p(s(bool,X62))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f23231,f14262])).
fof(f25902,plain,(
  ( ! [X30,X28,X31,X29,X27] : (sP10(X27,X28) | sP0(X29,X27,X30) | ~p(s(bool,X28)) | sP10(X31,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(subsumption_resolution,[],[f25858,f20171])).
fof(f25858,plain,(
  ( ! [X30,X28,X31,X29,X27] : (sP10(X27,X28) | sP0(X29,X27,X30) | s(bool,'F_0') = s(bool,X31) | ~p(s(bool,X28)) | sP10(X31,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(resolution,[],[f23231,f16692])).
fof(f25811,plain,(
  ( ! [X2,X1] : (sP0(X1,X2,'T_0')) ) | $spl286),
  inference(cnf_transformation,[],[f25811_D])).
fof(f25811_D,plain,(
  ( ! [X2,X1] : (sP0(X1,X2,'T_0')) ) <=> ~$spl286),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl286])])).
fof(f25795,plain,(
  ( ! [X2,X0,X1] : (sP10('T_0',X0) | sP0(X1,X2,'T_0')) ) | $spl122),
  inference(factoring,[],[f23230])).
fof(f25809,plain,(
  ( ! [X61,X62,X63] : (sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP0(X61,X62,X63)) ) | ($spl1 | $spl122)),
  inference(subsumption_resolution,[],[f25781,f329])).
fof(f25781,plain,(
  ( ! [X61,X62,X63] : (sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP0(X61,X62,X63) | p(s(bool,X63))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f23230,f14262])).
fof(f25800,plain,(
  ( ! [X30,X28,X31,X29,X27] : (sP10(X27,X28) | sP0(X29,X30,X27) | ~p(s(bool,X28)) | sP10(X31,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(subsumption_resolution,[],[f25756,f20171])).
fof(f25756,plain,(
  ( ! [X30,X28,X31,X29,X27] : (sP10(X27,X28) | sP0(X29,X30,X27) | s(bool,'F_0') = s(bool,X31) | ~p(s(bool,X28)) | sP10(X31,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(resolution,[],[f23230,f16692])).
fof(f25739,plain,(
  ( ! [X52,X50,X51] : (p(s(bool,X50)) | p(s(bool,X51)) | sP10('T_0',X52) | sP10(X52,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f14203,f22553])).
fof(f22553,plain,(
  ( ! [X87,X85,X86,X84] : (sP10('T_0',X84) | sP10(X84,X85) | sP0(X85,X86,X87)) ) | $spl122),
  inference(resolution,[],[f19618,f331])).
fof(f25736,plain,(
  ( ! [X43,X42] : (p(s(bool,X42)) | p(s(bool,X43)) | sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0')) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f14203,f20162])).
fof(f20162,plain,(
  ( ! [X66,X67,X65] : (sP0(X65,X66,X67) | sP10(X65,'T_0')) ) | $spl122),
  inference(resolution,[],[f20069,f331])).
fof(f25733,plain,(
  ( ! [X37,X36] : (p(s(bool,X36)) | p(s(bool,X37)) | sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f14203,f19558])).
fof(f19558,plain,(
  ( ! [X66,X67,X65] : (sP0(X65,X66,X67) | sP10('T_0',X65)) ) | $spl122),
  inference(resolution,[],[f19515,f331])).
fof(f25730,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X29)) | p(s(bool,X30)) | sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X31) | s(bool,'F_0') = s(bool,X31)) ) | ($spl1 | $spl8 | $spl46)),
  inference(resolution,[],[f14203,f16859])).
fof(f16859,plain,(
  ( ! [X87,X85,X86,X84] : (sP10(X84,X85) | sP0(X84,X86,X87) | s(bool,'F_0') = s(bool,X85)) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16258,f331])).
fof(f25727,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | p(s(bool,X21)) | sP10(X22,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | s(bool,'F_0') = s(bool,X22)) ) | ($spl1 | $spl8 | $spl46)),
  inference(resolution,[],[f14203,f16634])).
fof(f16634,plain,(
  ( ! [X87,X85,X86,X84] : (sP10(X84,X85) | sP0(X85,X86,X87) | s(bool,'F_0') = s(bool,X84)) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16255,f331])).
fof(f25707,plain,(
  ( ! [X2,X1] : (sP1(X1,X2,'T_0')) ) | $spl282),
  inference(cnf_transformation,[],[f25707_D])).
fof(f25707_D,plain,(
  ( ! [X2,X1] : (sP1(X1,X2,'T_0')) ) <=> ~$spl282),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl282])])).
fof(f25693,plain,(
  ( ! [X2,X0,X1] : (sP10('T_0',X0) | sP1(X1,X2,'T_0')) ) | $spl122),
  inference(factoring,[],[f23229])).
fof(f23229,plain,(
  ( ! [X74,X72,X75,X73] : (sP10('T_0',X72) | sP10(X73,X72) | sP1(X74,X75,X73)) ) | $spl122),
  inference(resolution,[],[f19619,f325])).
fof(f25690,plain,(
  ( ! [X26,X24,X25] : (sP10('T_0',X24) | sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X24) | ~p(s(bool,X25)) | ~p(s(bool,X26))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f23229,f14199])).
fof(f25705,plain,(
  ( ! [X61,X62,X63] : (sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP1(X61,X62,X63)) ) | ($spl1 | $spl122)),
  inference(subsumption_resolution,[],[f25679,f325])).
fof(f25679,plain,(
  ( ! [X61,X62,X63] : (sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP1(X61,X62,X63) | p(s(bool,X63))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f23229,f14262])).
fof(f25698,plain,(
  ( ! [X30,X28,X31,X29,X27] : (sP10(X27,X28) | sP1(X29,X30,X27) | ~p(s(bool,X28)) | sP10(X31,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(subsumption_resolution,[],[f25654,f20171])).
fof(f25654,plain,(
  ( ! [X30,X28,X31,X29,X27] : (sP10(X27,X28) | sP1(X29,X30,X27) | s(bool,'F_0') = s(bool,X31) | ~p(s(bool,X28)) | sP10(X31,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(resolution,[],[f23229,f16692])).
fof(f25637,plain,(
  ( ! [X2,X1] : (sP2(X1,X2,'T_0')) ) | $spl278),
  inference(cnf_transformation,[],[f25637_D])).
fof(f25637_D,plain,(
  ( ! [X2,X1] : (sP2(X1,X2,'T_0')) ) <=> ~$spl278),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl278])])).
fof(f25623,plain,(
  ( ! [X2,X0,X1] : (sP10('T_0',X0) | sP2(X1,X2,'T_0')) ) | $spl122),
  inference(factoring,[],[f23226])).
fof(f23226,plain,(
  ( ! [X61,X62,X60,X63] : (sP10('T_0',X60) | sP10(X61,X60) | sP2(X62,X63,X61)) ) | $spl122),
  inference(resolution,[],[f19619,f321])).
fof(f25620,plain,(
  ( ! [X37,X35,X36] : (sP10('T_0',X35) | sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X35) | ~p(s(bool,X36)) | ~p(s(bool,X37))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f23226,f14196])).
fof(f25635,plain,(
  ( ! [X61,X62,X63] : (sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP2(X61,X62,X63)) ) | ($spl1 | $spl122)),
  inference(subsumption_resolution,[],[f25605,f321])).
fof(f25605,plain,(
  ( ! [X61,X62,X63] : (sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP2(X61,X62,X63) | p(s(bool,X63))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f23226,f14262])).
fof(f25628,plain,(
  ( ! [X30,X28,X31,X29,X27] : (sP10(X27,X28) | sP2(X29,X30,X27) | ~p(s(bool,X28)) | sP10(X31,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(subsumption_resolution,[],[f25580,f20171])).
fof(f25580,plain,(
  ( ! [X30,X28,X31,X29,X27] : (sP10(X27,X28) | sP2(X29,X30,X27) | s(bool,'F_0') = s(bool,X31) | ~p(s(bool,X28)) | sP10(X31,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(resolution,[],[f23226,f16692])).
fof(f25571,plain,(
  ( ! [X70,X68,X69] : (sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP8(X68,X69,X70)) ) | ($spl1 | $spl122)),
  inference(subsumption_resolution,[],[f25550,f409])).
fof(f25550,plain,(
  ( ! [X70,X68,X69] : (sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP8(X68,X69,X70) | p(s(bool,X70))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f22603,f14263])).
fof(f25564,plain,(
  ( ! [X30,X28,X31,X29,X27] : (sP10(X27,X28) | sP8(X29,X30,X28) | ~p(s(bool,X27)) | sP10(X31,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(subsumption_resolution,[],[f25522,f20171])).
fof(f25522,plain,(
  ( ! [X30,X28,X31,X29,X27] : (sP10(X27,X28) | sP8(X29,X30,X28) | s(bool,'F_0') = s(bool,X31) | ~p(s(bool,X27)) | sP10(X31,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(resolution,[],[f22603,f16692])).
fof(f25504,plain,(
  ( ! [X47,X48,X49] : (p(s(bool,X47)) | p(s(bool,X48)) | sP10('T_0',X49) | sP10(X49,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f14202,f22552])).
fof(f22552,plain,(
  ( ! [X80,X83,X81,X82] : (sP10('T_0',X80) | sP10(X80,X81) | sP0(X82,X81,X83)) ) | $spl122),
  inference(resolution,[],[f19618,f330])).
fof(f25501,plain,(
  ( ! [X41,X40] : (p(s(bool,X40)) | p(s(bool,X41)) | sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0')) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f14202,f20161])).
fof(f20161,plain,(
  ( ! [X64,X62,X63] : (sP0(X63,X62,X64) | sP10(X62,'T_0')) ) | $spl122),
  inference(resolution,[],[f20069,f330])).
fof(f25498,plain,(
  ( ! [X35,X34] : (p(s(bool,X34)) | p(s(bool,X35)) | sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f14202,f19557])).
fof(f19557,plain,(
  ( ! [X64,X62,X63] : (sP0(X63,X62,X64) | sP10('T_0',X62)) ) | $spl122),
  inference(resolution,[],[f19515,f330])).
fof(f25495,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | p(s(bool,X27)) | sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X28) | s(bool,'F_0') = s(bool,X28)) ) | ($spl1 | $spl8 | $spl46)),
  inference(resolution,[],[f14202,f16858])).
fof(f16858,plain,(
  ( ! [X80,X83,X81,X82] : (sP10(X80,X81) | sP0(X82,X80,X83) | s(bool,'F_0') = s(bool,X81)) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16258,f330])).
fof(f25492,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | p(s(bool,X18)) | sP10(X19,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | s(bool,'F_0') = s(bool,X19)) ) | ($spl1 | $spl8 | $spl46)),
  inference(resolution,[],[f14202,f16633])).
fof(f16633,plain,(
  ( ! [X80,X83,X81,X82] : (sP10(X80,X81) | sP0(X82,X81,X83) | s(bool,'F_0') = s(bool,X80)) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16255,f330])).
fof(f25482,plain,(
  ( ! [X70,X68,X69] : (sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP6(X68,X69,X70)) ) | ($spl1 | $spl122)),
  inference(subsumption_resolution,[],[f25461,f393])).
fof(f25461,plain,(
  ( ! [X70,X68,X69] : (sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP6(X68,X69,X70) | p(s(bool,X69))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f22587,f14263])).
fof(f25475,plain,(
  ( ! [X30,X28,X31,X29,X27] : (sP10(X27,X28) | sP6(X29,X28,X30) | ~p(s(bool,X27)) | sP10(X31,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(subsumption_resolution,[],[f25433,f20171])).
fof(f25433,plain,(
  ( ! [X30,X28,X31,X29,X27] : (sP10(X27,X28) | sP6(X29,X28,X30) | s(bool,'F_0') = s(bool,X31) | ~p(s(bool,X27)) | sP10(X31,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(resolution,[],[f22587,f16692])).
fof(f25422,plain,(
  ( ! [X70,X68,X69] : (sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP6(X68,X69,X70)) ) | ($spl1 | $spl122)),
  inference(subsumption_resolution,[],[f25401,f392])).
fof(f25401,plain,(
  ( ! [X70,X68,X69] : (sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP6(X68,X69,X70) | p(s(bool,X70))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f22586,f14263])).
fof(f25415,plain,(
  ( ! [X30,X28,X31,X29,X27] : (sP10(X27,X28) | sP6(X29,X30,X28) | ~p(s(bool,X27)) | sP10(X31,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(subsumption_resolution,[],[f25373,f20171])).
fof(f25373,plain,(
  ( ! [X30,X28,X31,X29,X27] : (sP10(X27,X28) | sP6(X29,X30,X28) | s(bool,'F_0') = s(bool,X31) | ~p(s(bool,X27)) | sP10(X31,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(resolution,[],[f22586,f16692])).
fof(f25363,plain,(
  ( ! [X70,X68,X69] : (sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP4(X68,X69,X70)) ) | ($spl1 | $spl122)),
  inference(subsumption_resolution,[],[f25341,f376])).
fof(f25341,plain,(
  ( ! [X70,X68,X69] : (sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP4(X68,X69,X70) | p(s(bool,X69))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f22573,f14263])).
fof(f25356,plain,(
  ( ! [X30,X28,X31,X29,X27] : (sP10(X27,X28) | sP4(X29,X28,X30) | ~p(s(bool,X27)) | sP10(X31,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(subsumption_resolution,[],[f25313,f20171])).
fof(f25313,plain,(
  ( ! [X30,X28,X31,X29,X27] : (sP10(X27,X28) | sP4(X29,X28,X30) | s(bool,'F_0') = s(bool,X31) | ~p(s(bool,X27)) | sP10(X31,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(resolution,[],[f22573,f16692])).
fof(f25303,plain,(
  ( ! [X70,X68,X69] : (sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP0(X68,X69,X70)) ) | ($spl1 | $spl122)),
  inference(subsumption_resolution,[],[f25280,f331])).
fof(f25280,plain,(
  ( ! [X70,X68,X69] : (sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP0(X68,X69,X70) | p(s(bool,X68))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f22553,f14263])).
fof(f25294,plain,(
  ( ! [X30,X28,X31,X29,X27] : (sP10(X27,X28) | sP0(X28,X29,X30) | ~p(s(bool,X27)) | sP10(X31,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(subsumption_resolution,[],[f25252,f20171])).
fof(f25252,plain,(
  ( ! [X30,X28,X31,X29,X27] : (sP10(X27,X28) | sP0(X28,X29,X30) | s(bool,'F_0') = s(bool,X31) | ~p(s(bool,X27)) | sP10(X31,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(resolution,[],[f22553,f16692])).
fof(f25243,plain,(
  ( ! [X70,X68,X69] : (sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP0(X68,X69,X70)) ) | ($spl1 | $spl122)),
  inference(subsumption_resolution,[],[f25220,f330])).
fof(f25220,plain,(
  ( ! [X70,X68,X69] : (sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP0(X68,X69,X70) | p(s(bool,X69))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f22552,f14263])).
fof(f25234,plain,(
  ( ! [X30,X28,X31,X29,X27] : (sP10(X27,X28) | sP0(X29,X28,X30) | ~p(s(bool,X27)) | sP10(X31,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(subsumption_resolution,[],[f25192,f20171])).
fof(f25192,plain,(
  ( ! [X30,X28,X31,X29,X27] : (sP10(X27,X28) | sP0(X29,X28,X30) | s(bool,'F_0') = s(bool,X31) | ~p(s(bool,X27)) | sP10(X31,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(resolution,[],[f22552,f16692])).
fof(f25171,plain,(
  ( ! [X41,X42,X40] : (~p(s(bool,X40)) | ~p(s(bool,X41)) | sP10('T_0',X42) | sP10(X42,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f14199,f22550])).
fof(f22550,plain,(
  ( ! [X74,X72,X75,X73] : (sP10('T_0',X72) | sP10(X72,X73) | sP1(X74,X75,X73)) ) | $spl122),
  inference(resolution,[],[f19618,f325])).
fof(f25168,plain,(
  ( ! [X35,X34] : (~p(s(bool,X34)) | ~p(s(bool,X35)) | sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0')) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f14199,f20159])).
fof(f20159,plain,(
  ( ! [X57,X58,X56] : (sP1(X57,X58,X56) | sP10(X56,'T_0')) ) | $spl122),
  inference(resolution,[],[f20069,f325])).
fof(f25165,plain,(
  ( ! [X26,X27] : (~p(s(bool,X26)) | ~p(s(bool,X27)) | sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f14199,f19555])).
fof(f19555,plain,(
  ( ! [X57,X58,X56] : (sP1(X57,X58,X56) | sP10('T_0',X56)) ) | $spl122),
  inference(resolution,[],[f19515,f325])).
fof(f25160,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,X13)) | ~p(s(bool,X14)) | sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X15) | s(bool,'F_0') = s(bool,X15)) ) | ($spl1 | $spl8 | $spl46)),
  inference(resolution,[],[f14199,f16856])).
fof(f16856,plain,(
  ( ! [X74,X72,X75,X73] : (sP10(X72,X73) | sP1(X74,X75,X72) | s(bool,'F_0') = s(bool,X73)) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16258,f325])).
fof(f25159,plain,(
  ( ! [X12,X10,X11] : (~p(s(bool,X10)) | ~p(s(bool,X11)) | sP10(X12,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | s(bool,'F_0') = s(bool,X12)) ) | ($spl1 | $spl8 | $spl46)),
  inference(resolution,[],[f14199,f16631])).
fof(f16631,plain,(
  ( ! [X74,X72,X75,X73] : (sP10(X72,X73) | sP1(X74,X75,X73) | s(bool,'F_0') = s(bool,X72)) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16255,f325])).
fof(f25151,plain,(
  ( ! [X70,X68,X69] : (sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP0(X68,X69,X70)) ) | ($spl1 | $spl122)),
  inference(subsumption_resolution,[],[f25128,f329])).
fof(f25128,plain,(
  ( ! [X70,X68,X69] : (sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP0(X68,X69,X70) | p(s(bool,X70))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f22551,f14263])).
fof(f25142,plain,(
  ( ! [X30,X28,X31,X29,X27] : (sP10(X27,X28) | sP0(X29,X30,X28) | ~p(s(bool,X27)) | sP10(X31,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(subsumption_resolution,[],[f25100,f20171])).
fof(f25100,plain,(
  ( ! [X30,X28,X31,X29,X27] : (sP10(X27,X28) | sP0(X29,X30,X28) | s(bool,'F_0') = s(bool,X31) | ~p(s(bool,X27)) | sP10(X31,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(resolution,[],[f22551,f16692])).
fof(f25090,plain,(
  ( ! [X70,X68,X69] : (sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP1(X68,X69,X70)) ) | ($spl1 | $spl122)),
  inference(subsumption_resolution,[],[f25068,f325])).
fof(f25068,plain,(
  ( ! [X70,X68,X69] : (sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP1(X68,X69,X70) | p(s(bool,X70))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f22550,f14263])).
fof(f25083,plain,(
  ( ! [X30,X28,X31,X29,X27] : (sP10(X27,X28) | sP1(X29,X30,X28) | ~p(s(bool,X27)) | sP10(X31,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(subsumption_resolution,[],[f25040,f20171])).
fof(f25040,plain,(
  ( ! [X30,X28,X31,X29,X27] : (sP10(X27,X28) | sP1(X29,X30,X28) | s(bool,'F_0') = s(bool,X31) | ~p(s(bool,X27)) | sP10(X31,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(resolution,[],[f22550,f16692])).
fof(f25015,plain,(
  ( ! [X37,X35,X36] : (sP10('T_0',X35) | sP10(X35,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | ~p(s(bool,X36)) | ~p(s(bool,X37))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f22547,f14196])).
fof(f22547,plain,(
  ( ! [X61,X62,X60,X63] : (sP10('T_0',X60) | sP10(X60,X61) | sP2(X62,X63,X61)) ) | $spl122),
  inference(resolution,[],[f19618,f321])).
fof(f25030,plain,(
  ( ! [X70,X68,X69] : (sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP2(X68,X69,X70)) ) | ($spl1 | $spl122)),
  inference(subsumption_resolution,[],[f25003,f321])).
fof(f25003,plain,(
  ( ! [X70,X68,X69] : (sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP2(X68,X69,X70) | p(s(bool,X70))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f22547,f14263])).
fof(f25023,plain,(
  ( ! [X30,X28,X31,X29,X27] : (sP10(X27,X28) | sP2(X29,X30,X28) | ~p(s(bool,X27)) | sP10(X31,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(subsumption_resolution,[],[f24975,f20171])).
fof(f24975,plain,(
  ( ! [X30,X28,X31,X29,X27] : (sP10(X27,X28) | sP2(X29,X30,X28) | s(bool,'F_0') = s(bool,X31) | ~p(s(bool,X27)) | sP10(X31,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(resolution,[],[f22547,f16692])).
fof(f24953,plain,(
  ( ! [X35,X34] : (~p(s(bool,X34)) | ~p(s(bool,X35)) | sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0')) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f14196,f20156])).
fof(f20156,plain,(
  ( ! [X47,X48,X49] : (sP2(X48,X49,X47) | sP10(X47,'T_0')) ) | $spl122),
  inference(resolution,[],[f20069,f321])).
fof(f24950,plain,(
  ( ! [X26,X27] : (~p(s(bool,X26)) | ~p(s(bool,X27)) | sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f14196,f19552])).
fof(f19552,plain,(
  ( ! [X47,X48,X49] : (sP2(X48,X49,X47) | sP10('T_0',X47)) ) | $spl122),
  inference(resolution,[],[f19515,f321])).
fof(f24945,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,X13)) | ~p(s(bool,X14)) | sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X15) | s(bool,'F_0') = s(bool,X15)) ) | ($spl1 | $spl8 | $spl46)),
  inference(resolution,[],[f14196,f16853])).
fof(f16853,plain,(
  ( ! [X61,X62,X60,X63] : (sP10(X60,X61) | sP2(X62,X63,X60) | s(bool,'F_0') = s(bool,X61)) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16258,f321])).
fof(f24944,plain,(
  ( ! [X12,X10,X11] : (~p(s(bool,X10)) | ~p(s(bool,X11)) | sP10(X12,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | s(bool,'F_0') = s(bool,X12)) ) | ($spl1 | $spl8 | $spl46)),
  inference(resolution,[],[f14196,f16628])).
fof(f16628,plain,(
  ( ! [X61,X62,X60,X63] : (sP10(X60,X61) | sP2(X62,X63,X61) | s(bool,'F_0') = s(bool,X60)) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16255,f321])).
fof(f24876,plain,(
  ( ! [X0,X3,X1] : (sP10(X0,'T_0') | s(bool,X1) != s(bool,X3) | p(s(bool,X3)) | sP10(X1,X0)) ) | $spl244),
  inference(cnf_transformation,[],[f24876_D])).
fof(f24876_D,plain,(
  ( ! [X0,X3,X1] : (sP10(X0,'T_0') | s(bool,X1) != s(bool,X3) | p(s(bool,X3)) | sP10(X1,X0)) ) <=> ~$spl244),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl244])])).
fof(f24790,plain,(
  ( ! [X356,X354,X357,X355] : (sP10(X354,'T_0') | sP10(X355,X354) | p(s(bool,X356)) | p(s(bool,X357)) | s(bool,X355) != s(bool,X357)) ) | $spl122),
  inference(resolution,[],[f20223,f536])).
fof(f24700,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP10(X35,'T_0') | sP10(X36,X35) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X37) | s(bool,X38) != s(bool,X39)) ) | $spl122),
  inference(resolution,[],[f20223,f312])).
fof(f24697,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,'T_0') | sP10(X21,X20) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X21) != s(bool,X24)) ) | $spl122),
  inference(resolution,[],[f20223,f311])).
fof(f24696,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP10(X15,'T_0') | sP10(X16,X15) | p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) ) | $spl122),
  inference(resolution,[],[f20223,f310])).
fof(f24695,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP10(X10,'T_0') | sP10(X11,X10) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X12) | s(bool,X13) != s(bool,X14)) ) | $spl122),
  inference(resolution,[],[f20223,f308])).
fof(f24693,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10(X0,'T_0') | sP10(X1,X0) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) ) | $spl122),
  inference(resolution,[],[f20223,f306])).
fof(f24861,plain,(
  ( ! [X31] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0') | p(s(bool,X31))) ) | ($spl1 | $spl122)),
  inference(duplicate_literal_removal,[],[f24688])).
fof(f24688,plain,(
  ( ! [X31] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0') | p(s(bool,X31)) | p(s(bool,X31))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f20223,f14262])).
fof(f24646,plain,(
  ( ! [X0,X3,X1] : (sP10(X0,'T_0') | s(bool,X1) != s(bool,X3) | p(s(bool,X3)) | sP10(X0,X1)) ) | $spl242),
  inference(cnf_transformation,[],[f24646_D])).
fof(f24646_D,plain,(
  ( ! [X0,X3,X1] : (sP10(X0,'T_0') | s(bool,X1) != s(bool,X3) | p(s(bool,X3)) | sP10(X0,X1)) ) <=> ~$spl242),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl242])])).
fof(f24554,plain,(
  ( ! [X356,X354,X357,X355] : (sP10(X354,'T_0') | sP10(X354,X355) | p(s(bool,X356)) | p(s(bool,X357)) | s(bool,X355) != s(bool,X357)) ) | $spl122),
  inference(resolution,[],[f20222,f536])).
fof(f24464,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP10(X35,'T_0') | sP10(X35,X36) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X37) | s(bool,X38) != s(bool,X39)) ) | $spl122),
  inference(resolution,[],[f20222,f312])).
fof(f24461,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,'T_0') | sP10(X20,X21) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X21) != s(bool,X24)) ) | $spl122),
  inference(resolution,[],[f20222,f311])).
fof(f24460,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP10(X15,'T_0') | sP10(X15,X16) | p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) ) | $spl122),
  inference(resolution,[],[f20222,f310])).
fof(f24459,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP10(X10,'T_0') | sP10(X10,X11) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X12) | s(bool,X13) != s(bool,X14)) ) | $spl122),
  inference(resolution,[],[f20222,f308])).
fof(f24457,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10(X0,'T_0') | sP10(X0,X1) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) ) | $spl122),
  inference(resolution,[],[f20222,f306])).
fof(f24632,plain,(
  ( ! [X32] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0') | p(s(bool,X32))) ) | ($spl1 | $spl122)),
  inference(duplicate_literal_removal,[],[f24455])).
fof(f24455,plain,(
  ( ! [X32] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0') | p(s(bool,X32)) | p(s(bool,X32))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f20222,f14263])).
fof(f24414,plain,(
  ( ! [X31] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0') | sP10('F_0',X31)) ) | ($spl1 | $spl122 | $spl133)),
  inference(subsumption_resolution,[],[f24378,f20649])).
fof(f24378,plain,(
  ( ! [X31] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0') | sP10('F_0',X31) | p(s(bool,X31))) ) | ($spl1 | $spl122 | $spl133)),
  inference(resolution,[],[f24016,f14262])).
fof(f24016,plain,(
  ( ! [X329,X328] : (sP10(X328,'T_0') | sP10(X329,X328) | sP10('F_0',X329)) ) | ($spl122 | $spl133)),
  inference(subsumption_resolution,[],[f23976,f537])).
fof(f23976,plain,(
  ( ! [X329,X328] : (~p(s(bool,'T_0')) | sP10(X328,'T_0') | sP10(X329,X328) | sP10('F_0',X329)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f20221,f23316])).
fof(f23316,plain,(
  ( ! [X374,X375] : (sP10('T_0',X374) | sP10(X375,X374) | sP10('F_0',X375)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f19619,f20649])).
fof(f20221,plain,(
  ( ! [X227,X228] : (~sP10(X228,X227) | ~p(s(bool,X228)) | sP10(X227,'T_0')) ) | $spl122),
  inference(resolution,[],[f20069,f439])).
fof(f24413,plain,(
  ( ! [X14,X12,X13] : (sP10(X12,'T_0') | sP10('F_0',X13) | s(bool,'F_0') = s(bool,X14) | sP10(X14,X13)) ) | ($spl8 | $spl46 | $spl122 | $spl133)),
  inference(subsumption_resolution,[],[f24369,f20069])).
fof(f24369,plain,(
  ( ! [X14,X12,X13] : (sP10(X12,'T_0') | sP10('F_0',X13) | s(bool,'F_0') = s(bool,X14) | ~p(s(bool,X12)) | sP10(X14,X13)) ) | ($spl8 | $spl46 | $spl122 | $spl133)),
  inference(resolution,[],[f24016,f16692])).
fof(f23676,plain,(
  ( ! [X14,X13] : (sP10(X13,X14) | ~p(s(bool,X14)) | sP10('F_0',X13)) ) | $spl230),
  inference(cnf_transformation,[],[f23676_D])).
fof(f23676_D,plain,(
  ( ! [X14,X13] : (sP10(X13,X14) | ~p(s(bool,X14)) | sP10('F_0',X13)) ) <=> ~$spl230),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl230])])).
fof(f24409,plain,(
  ( ! [X17,X15,X16] : (sP10(X15,X16) | sP10('F_0',X15) | ~p(s(bool,X16)) | sP10(X17,'T_0')) ) | ($spl8 | $spl46 | $spl122 | $spl133)),
  inference(subsumption_resolution,[],[f24353,f20171])).
fof(f24353,plain,(
  ( ! [X17,X15,X16] : (sP10(X15,X16) | sP10('F_0',X15) | s(bool,'F_0') = s(bool,X17) | ~p(s(bool,X16)) | sP10(X17,'T_0')) ) | ($spl8 | $spl46 | $spl122 | $spl133)),
  inference(resolution,[],[f24016,f16693])).
fof(f24341,plain,(
  ( ! [X32] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0') | sP10('F_0',X32)) ) | ($spl1 | $spl122 | $spl133)),
  inference(subsumption_resolution,[],[f24307,f20649])).
fof(f24307,plain,(
  ( ! [X32] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0') | sP10('F_0',X32) | p(s(bool,X32))) ) | ($spl1 | $spl122 | $spl133)),
  inference(resolution,[],[f24013,f14263])).
fof(f24013,plain,(
  ( ! [X321,X320] : (sP10(X320,'T_0') | sP10(X320,X321) | sP10('F_0',X321)) ) | ($spl122 | $spl133)),
  inference(subsumption_resolution,[],[f23973,f537])).
fof(f23973,plain,(
  ( ! [X321,X320] : (~p(s(bool,'T_0')) | sP10(X320,'T_0') | sP10(X320,X321) | sP10('F_0',X321)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f20221,f22637])).
fof(f22637,plain,(
  ( ! [X374,X375] : (sP10('T_0',X374) | sP10(X374,X375) | sP10('F_0',X375)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f19618,f20649])).
fof(f24340,plain,(
  ( ! [X17,X15,X16] : (sP10(X15,'T_0') | sP10('F_0',X16) | s(bool,'F_0') = s(bool,X17) | sP10(X17,X16)) ) | ($spl8 | $spl46 | $spl122 | $spl133)),
  inference(subsumption_resolution,[],[f24296,f20069])).
fof(f24296,plain,(
  ( ! [X17,X15,X16] : (sP10(X15,'T_0') | sP10('F_0',X16) | s(bool,'F_0') = s(bool,X17) | ~p(s(bool,X15)) | sP10(X17,X16)) ) | ($spl8 | $spl46 | $spl122 | $spl133)),
  inference(resolution,[],[f24013,f16693])).
fof(f23170,plain,(
  ( ! [X14,X13] : (sP10(X13,X14) | ~p(s(bool,X13)) | sP10('F_0',X14)) ) | $spl210),
  inference(cnf_transformation,[],[f23170_D])).
fof(f23170_D,plain,(
  ( ! [X14,X13] : (sP10(X13,X14) | ~p(s(bool,X13)) | sP10('F_0',X14)) ) <=> ~$spl210),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl210])])).
fof(f24336,plain,(
  ( ! [X17,X15,X16] : (sP10(X15,X16) | sP10('F_0',X16) | ~p(s(bool,X15)) | sP10(X17,'T_0')) ) | ($spl8 | $spl46 | $spl122 | $spl133)),
  inference(subsumption_resolution,[],[f24279,f20171])).
fof(f24279,plain,(
  ( ! [X17,X15,X16] : (sP10(X15,X16) | sP10('F_0',X16) | s(bool,'F_0') = s(bool,X17) | ~p(s(bool,X15)) | sP10(X17,'T_0')) ) | ($spl8 | $spl46 | $spl122 | $spl133)),
  inference(resolution,[],[f24013,f16693])).
fof(f24258,plain,(
  ( ! [X14,X15] : (sP10(X14,'T_0') | sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X14) | p(s(bool,X15))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f24015,f14313])).
fof(f24015,plain,(
  ( ! [X325,X327,X326] : (sP10(X325,'T_0') | sP10(X326,X325) | sP11(X326,X327)) ) | $spl122),
  inference(subsumption_resolution,[],[f23975,f537])).
fof(f23975,plain,(
  ( ! [X325,X327,X326] : (~p(s(bool,'T_0')) | sP10(X325,'T_0') | sP10(X326,X325) | sP11(X326,X327)) ) | $spl122),
  inference(resolution,[],[f20221,f23299])).
fof(f24212,plain,(
  ( ! [X0,X1] : (sP10(X0,'T_0') | p(s(bool,X1)) | sP10(X1,X0)) ) | $spl238),
  inference(cnf_transformation,[],[f24212_D])).
fof(f24212_D,plain,(
  ( ! [X0,X1] : (sP10(X0,'T_0') | p(s(bool,X1)) | sP10(X1,X0)) ) <=> ~$spl238),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl238])])).
fof(f24271,plain,(
  ( ! [X10,X11,X9] : (sP10(X9,'T_0') | sP10(X10,X9) | p(s(bool,X10)) | p(s(bool,X11))) ) | $spl122),
  inference(subsumption_resolution,[],[f24256,f492])).
fof(f24256,plain,(
  ( ! [X10,X11,X9] : (sP10(X9,'T_0') | sP10(X10,X9) | p(s(bool,'F_0')) | p(s(bool,X10)) | p(s(bool,X11))) ) | $spl122),
  inference(resolution,[],[f24015,f454])).
fof(f24270,plain,(
  ( ! [X6,X8,X7] : (sP10(X6,'T_0') | sP10(X7,X6) | p(s(bool,X7)) | p(s(bool,X8))) ) | $spl122),
  inference(subsumption_resolution,[],[f24255,f492])).
fof(f24255,plain,(
  ( ! [X6,X8,X7] : (sP10(X6,'T_0') | sP10(X7,X6) | p(s(bool,'F_0')) | p(s(bool,X7)) | p(s(bool,X8))) ) | $spl122),
  inference(resolution,[],[f24015,f458])).
fof(f24254,plain,(
  ( ! [X4,X5,X3] : (sP10(X3,'T_0') | sP10(X4,X3) | p(s(bool,X4)) | sP10(X5,'F_0')) ) | $spl122),
  inference(resolution,[],[f24015,f18976])).
fof(f16786,plain,(
  ( ! [X372,X371] : (sP11(X371,X372)) ) | $spl106),
  inference(cnf_transformation,[],[f16786_D])).
fof(f16786_D,plain,(
  ( ! [X372,X371] : (sP11(X371,X372)) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f24269,plain,(
  ( ! [X47,X46] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0') | sP11(X46,X47)) ) | ($spl1 | $spl122)),
  inference(subsumption_resolution,[],[f24248,f456])).
fof(f24248,plain,(
  ( ! [X47,X46] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0') | sP11(X46,X47) | p(s(bool,X46))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f24015,f14262])).
fof(f23563,plain,(
  ( ! [X21,X22,X20] : (sP10(X20,X21) | ~p(s(bool,X21)) | sP11(X20,X22)) ) | $spl226),
  inference(cnf_transformation,[],[f23563_D])).
fof(f23563_D,plain,(
  ( ! [X21,X22,X20] : (sP10(X20,X21) | ~p(s(bool,X21)) | sP11(X20,X22)) ) <=> ~$spl226),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl226])])).
fof(f24265,plain,(
  ( ! [X24,X23,X25,X22] : (sP10(X22,X23) | sP11(X22,X24) | ~p(s(bool,X23)) | sP10(X25,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(subsumption_resolution,[],[f24223,f20171])).
fof(f24223,plain,(
  ( ! [X24,X23,X25,X22] : (sP10(X22,X23) | sP11(X22,X24) | s(bool,'F_0') = s(bool,X25) | ~p(s(bool,X23)) | sP10(X25,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(resolution,[],[f24015,f16693])).
fof(f24199,plain,(
  ( ! [X12,X13] : (sP10(X12,'T_0') | sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X12) | p(s(bool,X13))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f24014,f14312])).
fof(f24014,plain,(
  ( ! [X323,X324,X322] : (sP10(X322,'T_0') | sP10(X323,X322) | sP11(X324,X323)) ) | $spl122),
  inference(subsumption_resolution,[],[f23974,f537])).
fof(f23974,plain,(
  ( ! [X323,X324,X322] : (~p(s(bool,'T_0')) | sP10(X322,'T_0') | sP10(X323,X322) | sP11(X324,X323)) ) | $spl122),
  inference(resolution,[],[f20221,f23298])).
fof(f24214,plain,(
  ( ! [X10,X11,X9] : (sP10(X9,'T_0') | sP10(X10,X9) | p(s(bool,X11)) | p(s(bool,X10))) ) | $spl122),
  inference(subsumption_resolution,[],[f24198,f492])).
fof(f24198,plain,(
  ( ! [X10,X11,X9] : (sP10(X9,'T_0') | sP10(X10,X9) | p(s(bool,'F_0')) | p(s(bool,X11)) | p(s(bool,X10))) ) | $spl122),
  inference(resolution,[],[f24014,f454])).
fof(f24213,plain,(
  ( ! [X6,X8,X7] : (sP10(X6,'T_0') | sP10(X7,X6) | p(s(bool,X8)) | p(s(bool,X7))) ) | $spl122),
  inference(subsumption_resolution,[],[f24197,f492])).
fof(f24197,plain,(
  ( ! [X6,X8,X7] : (sP10(X6,'T_0') | sP10(X7,X6) | p(s(bool,'F_0')) | p(s(bool,X8)) | p(s(bool,X7))) ) | $spl122),
  inference(resolution,[],[f24014,f458])).
fof(f24195,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,'T_0') | sP10(X1,X0) | p(s(bool,X1)) | sP10(X2,'F_0')) ) | $spl122),
  inference(resolution,[],[f24014,f18977])).
fof(f24211,plain,(
  ( ! [X47,X46] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0') | sP11(X46,X47)) ) | ($spl1 | $spl122)),
  inference(subsumption_resolution,[],[f24190,f455])).
fof(f24190,plain,(
  ( ! [X47,X46] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0') | sP11(X46,X47) | p(s(bool,X47))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f24014,f14262])).
fof(f23457,plain,(
  ( ! [X21,X22,X20] : (sP10(X20,X21) | ~p(s(bool,X21)) | sP11(X22,X20)) ) | $spl220),
  inference(cnf_transformation,[],[f23457_D])).
fof(f23457_D,plain,(
  ( ! [X21,X22,X20] : (sP10(X20,X21) | ~p(s(bool,X21)) | sP11(X22,X20)) ) <=> ~$spl220),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl220])])).
fof(f24207,plain,(
  ( ! [X24,X23,X25,X22] : (sP10(X22,X23) | sP11(X24,X22) | ~p(s(bool,X23)) | sP10(X25,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(subsumption_resolution,[],[f24165,f20171])).
fof(f24165,plain,(
  ( ! [X24,X23,X25,X22] : (sP10(X22,X23) | sP11(X24,X22) | s(bool,'F_0') = s(bool,X25) | ~p(s(bool,X23)) | sP10(X25,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(resolution,[],[f24014,f16693])).
fof(f24157,plain,(
  $false | ($spl8 | $spl234)),
  inference(subsumption_resolution,[],[f24150,f534])).
fof(f24150,plain,(
  ~p(s(bool,'T_0')) | ($spl8 | $spl234)),
  inference(resolution,[],[f24069,f6468])).
fof(f24148,plain,(
  $false | ($spl205 | $spl234)),
  inference(resolution,[],[f24069,f22771])).
fof(f24119,plain,(
  ( ! [X14,X15] : (sP10(X14,'T_0') | sP10(X14,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | p(s(bool,X15))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f24012,f14313])).
fof(f24012,plain,(
  ( ! [X318,X317,X319] : (sP10(X317,'T_0') | sP10(X317,X318) | sP11(X318,X319)) ) | $spl122),
  inference(subsumption_resolution,[],[f23972,f537])).
fof(f23972,plain,(
  ( ! [X318,X317,X319] : (~p(s(bool,'T_0')) | sP10(X317,'T_0') | sP10(X317,X318) | sP11(X318,X319)) ) | $spl122),
  inference(resolution,[],[f20221,f22620])).
fof(f24005,plain,(
  ( ! [X57,X56] : (sP10(X56,'T_0') | p(s(bool,X57)) | sP10(X56,X57)) ) | $spl232),
  inference(cnf_transformation,[],[f24005_D])).
fof(f24005_D,plain,(
  ( ! [X57,X56] : (sP10(X56,'T_0') | p(s(bool,X57)) | sP10(X56,X57)) ) <=> ~$spl232),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl232])])).
fof(f24132,plain,(
  ( ! [X10,X11,X9] : (sP10(X9,'T_0') | sP10(X9,X10) | p(s(bool,X10)) | p(s(bool,X11))) ) | $spl122),
  inference(subsumption_resolution,[],[f24117,f492])).
fof(f24117,plain,(
  ( ! [X10,X11,X9] : (sP10(X9,'T_0') | sP10(X9,X10) | p(s(bool,'F_0')) | p(s(bool,X10)) | p(s(bool,X11))) ) | $spl122),
  inference(resolution,[],[f24012,f454])).
fof(f24131,plain,(
  ( ! [X6,X8,X7] : (sP10(X6,'T_0') | sP10(X6,X7) | p(s(bool,X7)) | p(s(bool,X8))) ) | $spl122),
  inference(subsumption_resolution,[],[f24116,f492])).
fof(f24116,plain,(
  ( ! [X6,X8,X7] : (sP10(X6,'T_0') | sP10(X6,X7) | p(s(bool,'F_0')) | p(s(bool,X7)) | p(s(bool,X8))) ) | $spl122),
  inference(resolution,[],[f24012,f458])).
fof(f24115,plain,(
  ( ! [X4,X5,X3] : (sP10(X3,'T_0') | sP10(X3,X4) | p(s(bool,X4)) | sP10(X5,'F_0')) ) | $spl122),
  inference(resolution,[],[f24012,f18976])).
fof(f24130,plain,(
  ( ! [X50,X51] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0') | sP11(X50,X51)) ) | ($spl1 | $spl122)),
  inference(subsumption_resolution,[],[f24112,f456])).
fof(f24112,plain,(
  ( ! [X50,X51] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0') | sP11(X50,X51) | p(s(bool,X50))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f24012,f14263])).
fof(f22913,plain,(
  ( ! [X21,X22,X20] : (sP10(X20,X21) | ~p(s(bool,X20)) | sP11(X21,X22)) ) | $spl208),
  inference(cnf_transformation,[],[f22913_D])).
fof(f22913_D,plain,(
  ( ! [X21,X22,X20] : (sP10(X20,X21) | ~p(s(bool,X20)) | sP11(X21,X22)) ) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f24126,plain,(
  ( ! [X24,X23,X25,X22] : (sP10(X22,X23) | sP11(X23,X24) | ~p(s(bool,X22)) | sP10(X25,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(subsumption_resolution,[],[f24084,f20171])).
fof(f24084,plain,(
  ( ! [X24,X23,X25,X22] : (sP10(X22,X23) | sP11(X23,X24) | s(bool,'F_0') = s(bool,X25) | ~p(s(bool,X22)) | sP10(X25,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(resolution,[],[f24012,f16693])).
fof(f24060,plain,(
  ( ! [X12,X13] : (sP10(X12,'T_0') | sP10(X12,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | p(s(bool,X13))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f24011,f14312])).
fof(f24011,plain,(
  ( ! [X315,X316,X314] : (sP10(X314,'T_0') | sP10(X314,X315) | sP11(X316,X315)) ) | $spl122),
  inference(subsumption_resolution,[],[f23971,f537])).
fof(f23971,plain,(
  ( ! [X315,X316,X314] : (~p(s(bool,'T_0')) | sP10(X314,'T_0') | sP10(X314,X315) | sP11(X316,X315)) ) | $spl122),
  inference(resolution,[],[f20221,f22619])).
fof(f24075,plain,(
  ( ! [X10,X11,X9] : (sP10(X9,'T_0') | sP10(X9,X10) | p(s(bool,X11)) | p(s(bool,X10))) ) | $spl122),
  inference(subsumption_resolution,[],[f24059,f492])).
fof(f24059,plain,(
  ( ! [X10,X11,X9] : (sP10(X9,'T_0') | sP10(X9,X10) | p(s(bool,'F_0')) | p(s(bool,X11)) | p(s(bool,X10))) ) | $spl122),
  inference(resolution,[],[f24011,f454])).
fof(f24074,plain,(
  ( ! [X6,X8,X7] : (sP10(X6,'T_0') | sP10(X6,X7) | p(s(bool,X8)) | p(s(bool,X7))) ) | $spl122),
  inference(subsumption_resolution,[],[f24058,f492])).
fof(f24058,plain,(
  ( ! [X6,X8,X7] : (sP10(X6,'T_0') | sP10(X6,X7) | p(s(bool,'F_0')) | p(s(bool,X8)) | p(s(bool,X7))) ) | $spl122),
  inference(resolution,[],[f24011,f458])).
fof(f24056,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,'T_0') | sP10(X0,X1) | p(s(bool,X1)) | sP10(X2,'F_0')) ) | $spl122),
  inference(resolution,[],[f24011,f18977])).
fof(f24073,plain,(
  ( ! [X50,X51] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0') | sP11(X50,X51)) ) | ($spl1 | $spl122)),
  inference(subsumption_resolution,[],[f24054,f455])).
fof(f24054,plain,(
  ( ! [X50,X51] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0') | sP11(X50,X51) | p(s(bool,X51))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f24011,f14263])).
fof(f22769,plain,(
  ( ! [X19,X20,X18] : (sP10(X18,X19) | ~p(s(bool,X18)) | sP11(X20,X19)) ) | $spl202),
  inference(cnf_transformation,[],[f22769_D])).
fof(f22769_D,plain,(
  ( ! [X19,X20,X18] : (sP10(X18,X19) | ~p(s(bool,X18)) | sP11(X20,X19)) ) <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f24068,plain,(
  ( ! [X24,X23,X25,X22] : (sP10(X22,X23) | sP11(X24,X23) | ~p(s(bool,X22)) | sP10(X25,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(subsumption_resolution,[],[f24026,f20171])).
fof(f24026,plain,(
  ( ! [X24,X23,X25,X22] : (sP10(X22,X23) | sP11(X24,X23) | s(bool,'F_0') = s(bool,X25) | ~p(s(bool,X22)) | sP10(X25,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(resolution,[],[f24011,f16693])).
fof(f23988,plain,(
  ( ! [X301] : (~p(s(bool,X301)) | sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0')) ) | ($spl1 | $spl122)),
  inference(duplicate_literal_removal,[],[f23961])).
fof(f23961,plain,(
  ( ! [X301] : (~p(s(bool,X301)) | sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0') | ~p(s(bool,X301))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f20221,f14264])).
fof(f14264,plain,(
  ( ! [X151] : (sP10(X151,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | ~p(s(bool,X151))) ) | $spl1),
  inference(resolution,[],[f545,f440])).
fof(f24004,plain,(
  ( ! [X57,X56,X55] : (~p(s(bool,X55)) | sP10(X56,'T_0') | sP10(X56,X57) | p(s(bool,X57))) ) | ($spl8 | $spl46 | $spl122)),
  inference(subsumption_resolution,[],[f23884,f350])).
fof(f23884,plain,(
  ( ! [X57,X56,X55] : (~p(s(bool,X55)) | sP10(X56,'T_0') | sP10(X56,X57) | p(s(bool,X57)) | s(bool,'F_0') = s(bool,X55)) ) | ($spl8 | $spl46 | $spl122)),
  inference(resolution,[],[f20221,f16694])).
fof(f23841,plain,(
  ( ! [X337] : (~p(s(bool,X337)) | sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0')) ) | ($spl1 | $spl122)),
  inference(duplicate_literal_removal,[],[f23838])).
fof(f23838,plain,(
  ( ! [X337] : (~p(s(bool,X337)) | sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0') | ~p(s(bool,X337))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f20220,f14265])).
fof(f14265,plain,(
  ( ! [X152] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X152) | ~p(s(bool,X152))) ) | $spl1),
  inference(resolution,[],[f545,f443])).
fof(f20220,plain,(
  ( ! [X225,X226] : (~sP10(X225,X226) | ~p(s(bool,X226)) | sP10(X225,'T_0')) ) | $spl122),
  inference(resolution,[],[f20069,f439])).
fof(f23681,plain,(
  ( ! [X27] : (sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP10('F_0',X27)) ) | ($spl1 | $spl122 | $spl133)),
  inference(subsumption_resolution,[],[f23649,f20649])).
fof(f23649,plain,(
  ( ! [X27] : (sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP10('F_0',X27) | p(s(bool,X27))) ) | ($spl1 | $spl122 | $spl133)),
  inference(resolution,[],[f23316,f14262])).
fof(f23680,plain,(
  ( ! [X14,X12,X13] : (sP10('T_0',X12) | sP10('F_0',X13) | s(bool,'F_0') = s(bool,X14) | sP10(X14,X13)) ) | ($spl8 | $spl46 | $spl122 | $spl133)),
  inference(subsumption_resolution,[],[f23642,f19515])).
fof(f23642,plain,(
  ( ! [X14,X12,X13] : (sP10('T_0',X12) | sP10('F_0',X13) | s(bool,'F_0') = s(bool,X14) | ~p(s(bool,X12)) | sP10(X14,X13)) ) | ($spl8 | $spl46 | $spl122 | $spl133)),
  inference(resolution,[],[f23316,f16692])).
fof(f20653,plain,(
  ( ! [X9] : (s(bool,'F_0') = s(bool,X9) | sP10(X9,'T_0')) ) | $spl162),
  inference(cnf_transformation,[],[f20653_D])).
fof(f20653_D,plain,(
  ( ! [X9] : (s(bool,'F_0') = s(bool,X9) | sP10(X9,'T_0')) ) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f23628,plain,(
  ( ! [X14,X15,X13] : (sP10(X13,X14) | sP10('F_0',X13) | s(bool,'F_0') = s(bool,X15) | ~p(s(bool,X14)) | sP10(X15,'T_0')) ) | ($spl8 | $spl46 | $spl122 | $spl133)),
  inference(resolution,[],[f23316,f16692])).
fof(f23568,plain,(
  ( ! [X1] : (sP11('T_0',X1)) ) | $spl228),
  inference(cnf_transformation,[],[f23568_D])).
fof(f23568_D,plain,(
  ( ! [X1] : (sP11('T_0',X1)) ) <=> ~$spl228),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl228])])).
fof(f23558,plain,(
  ( ! [X0,X1] : (sP10('T_0',X0) | sP11('T_0',X1)) ) | $spl122),
  inference(factoring,[],[f23299])).
fof(f23557,plain,(
  ( ! [X14,X15] : (sP10('T_0',X14) | sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X14) | p(s(bool,X15))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f23299,f14313])).
fof(f23567,plain,(
  ( ! [X41,X40] : (sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP11(X40,X41)) ) | ($spl1 | $spl122)),
  inference(subsumption_resolution,[],[f23547,f456])).
fof(f23547,plain,(
  ( ! [X41,X40] : (sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP11(X40,X41) | p(s(bool,X40))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f23299,f14262])).
fof(f23526,plain,(
  ( ! [X23,X21,X22,X20] : (sP10(X20,X21) | sP11(X20,X22) | s(bool,'F_0') = s(bool,X23) | ~p(s(bool,X21)) | sP10(X23,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(resolution,[],[f23299,f16692])).
fof(f23463,plain,(
  ( ! [X1] : (sP11(X1,'T_0')) ) | $spl224),
  inference(cnf_transformation,[],[f23463_D])).
fof(f23463_D,plain,(
  ( ! [X1] : (sP11(X1,'T_0')) ) <=> ~$spl224),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl224])])).
fof(f23452,plain,(
  ( ! [X0,X1] : (sP10('T_0',X0) | sP11(X1,'T_0')) ) | $spl122),
  inference(factoring,[],[f23298])).
fof(f23450,plain,(
  ( ! [X12,X13] : (sP10('T_0',X12) | sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X12) | p(s(bool,X13))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f23298,f14312])).
fof(f23461,plain,(
  ( ! [X41,X40] : (sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP11(X40,X41)) ) | ($spl1 | $spl122)),
  inference(subsumption_resolution,[],[f23441,f455])).
fof(f23441,plain,(
  ( ! [X41,X40] : (sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP11(X40,X41) | p(s(bool,X41))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f23298,f14262])).
fof(f23420,plain,(
  ( ! [X23,X21,X22,X20] : (sP10(X20,X21) | sP11(X22,X20) | s(bool,'F_0') = s(bool,X23) | ~p(s(bool,X21)) | sP10(X23,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(resolution,[],[f23298,f16692])).
fof(f23402,plain,(
  ( ! [X0,X3,X1] : (sP10('T_0',X0) | s(bool,X1) != s(bool,X3) | p(s(bool,X3)) | sP10(X1,X0)) ) | $spl218),
  inference(cnf_transformation,[],[f23402_D])).
fof(f23402_D,plain,(
  ( ! [X0,X3,X1] : (sP10('T_0',X0) | s(bool,X1) != s(bool,X3) | p(s(bool,X3)) | sP10(X1,X0)) ) <=> ~$spl218),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl218])])).
fof(f23310,plain,(
  ( ! [X356,X354,X357,X355] : (sP10('T_0',X354) | sP10(X355,X354) | p(s(bool,X356)) | p(s(bool,X357)) | s(bool,X355) != s(bool,X357)) ) | $spl122),
  inference(resolution,[],[f19619,f536])).
fof(f23220,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP10('T_0',X35) | sP10(X36,X35) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X37) | s(bool,X38) != s(bool,X39)) ) | $spl122),
  inference(resolution,[],[f19619,f312])).
fof(f23217,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10('T_0',X20) | sP10(X21,X20) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X21) != s(bool,X24)) ) | $spl122),
  inference(resolution,[],[f19619,f311])).
fof(f23216,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP10('T_0',X15) | sP10(X16,X15) | p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) ) | $spl122),
  inference(resolution,[],[f19619,f310])).
fof(f23215,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP10('T_0',X10) | sP10(X11,X10) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X12) | s(bool,X13) != s(bool,X14)) ) | $spl122),
  inference(resolution,[],[f19619,f308])).
fof(f23213,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10('T_0',X0) | sP10(X1,X0) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) ) | $spl122),
  inference(resolution,[],[f19619,f306])).
fof(f23388,plain,(
  ( ! [X27] : (sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | p(s(bool,X27))) ) | ($spl1 | $spl122)),
  inference(duplicate_literal_removal,[],[f23208])).
fof(f23208,plain,(
  ( ! [X27] : (sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | p(s(bool,X27)) | p(s(bool,X27))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f19619,f14262])).
fof(f23177,plain,(
  ( ! [X28] : (sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP10('F_0',X28)) ) | ($spl1 | $spl122 | $spl133)),
  inference(subsumption_resolution,[],[f23146,f20649])).
fof(f23146,plain,(
  ( ! [X28] : (sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP10('F_0',X28) | p(s(bool,X28))) ) | ($spl1 | $spl122 | $spl133)),
  inference(resolution,[],[f22637,f14263])).
fof(f23174,plain,(
  ( ! [X17,X15,X16] : (sP10('T_0',X15) | sP10('F_0',X16) | s(bool,'F_0') = s(bool,X17) | sP10(X17,X16)) ) | ($spl8 | $spl46 | $spl122 | $spl133)),
  inference(subsumption_resolution,[],[f23137,f19515])).
fof(f23137,plain,(
  ( ! [X17,X15,X16] : (sP10('T_0',X15) | sP10('F_0',X16) | s(bool,'F_0') = s(bool,X17) | ~p(s(bool,X15)) | sP10(X17,X16)) ) | ($spl8 | $spl46 | $spl122 | $spl133)),
  inference(resolution,[],[f22637,f16693])).
fof(f23122,plain,(
  ( ! [X14,X15,X13] : (sP10(X13,X14) | sP10('F_0',X14) | s(bool,'F_0') = s(bool,X15) | ~p(s(bool,X13)) | sP10(X15,'T_0')) ) | ($spl8 | $spl46 | $spl122 | $spl133)),
  inference(resolution,[],[f22637,f16692])).
fof(f23114,plain,(
  $false | ($spl1 | $spl20)),
  inference(subsumption_resolution,[],[f23039,f22919])).
fof(f22919,plain,(
  ( ! [X1] : (p(s(bool,X1))) ) | ($spl1 | $spl20)),
  inference(subsumption_resolution,[],[f22918,f353])).
fof(f22918,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,X1))) ) | ($spl1 | $spl20)),
  inference(forward_demodulation,[],[f1007,f14213])).
fof(f23039,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl20)),
  inference(resolution,[],[f22919,f532])).
fof(f532,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f301])).
fof(f301,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f141])).
fof(f141,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f32])).
fof(f32,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.list.MEM_GENLIST','thm.bool.NOT_CLAUSES|split|2')).
fof(f23113,plain,(
  $false | ($spl1 | $spl20)),
  inference(subsumption_resolution,[],[f23038,f22919])).
fof(f23038,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl20)),
  inference(resolution,[],[f22919,f530])).
fof(f530,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f300])).
fof(f300,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f140])).
fof(f140,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f31])).
fof(f31,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.list.MEM_GENLIST','thm.bool.NOT_CLAUSES|split|1')).
fof(f23112,plain,(
  $false | ($spl1 | $spl20)),
  inference(subsumption_resolution,[],[f23035,f22919])).
fof(f23035,plain,(
  ( ! [X274] : (~p(s(bool,X274))) ) | ($spl1 | $spl20)),
  inference(resolution,[],[f22919,f524])).
fof(f524,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f170])).
fof(f170,plain,(
  ! [X0] : ((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f133])).
fof(f133,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) => ~p(s(bool,X0)))),
  inference(flattening,[],[f132])).
fof(f132,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) => ~p(s(bool,X0)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X11] : ((p(s(bool,X11)) => p(s(bool,'F_0'))) => ~p(s(bool,X11)))),
  file('thm.list.MEM_GENLIST','thm.bool.IMP_F')).
fof(f23111,plain,(
  $false | ($spl1 | $spl20)),
  inference(subsumption_resolution,[],[f23029,f22919])).
fof(f23029,plain,(
  ( ! [X268] : (~p(s(bool,X268))) ) | ($spl1 | $spl20)),
  inference(resolution,[],[f22919,f474])).
fof(f23037,plain,(
  $false | ($spl1 | $spl20)),
  inference(resolution,[],[f22919,f528])).
fof(f22907,plain,(
  ( ! [X14,X15] : (sP10('T_0',X14) | sP10(X14,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | p(s(bool,X15))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f22620,f14313])).
fof(f22917,plain,(
  ( ! [X45,X44] : (sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP11(X44,X45)) ) | ($spl1 | $spl122)),
  inference(subsumption_resolution,[],[f22900,f456])).
fof(f22900,plain,(
  ( ! [X45,X44] : (sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP11(X44,X45) | p(s(bool,X44))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f22620,f14263])).
fof(f22876,plain,(
  ( ! [X23,X21,X22,X20] : (sP10(X20,X21) | sP11(X21,X22) | s(bool,'F_0') = s(bool,X23) | ~p(s(bool,X20)) | sP10(X23,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(resolution,[],[f22620,f16692])).
fof(f22797,plain,(
  $false | ($spl8 | $spl46 | $spl204)),
  inference(subsumption_resolution,[],[f22796,f534])).
fof(f22796,plain,(
  ~p(s(bool,'T_0')) | ($spl8 | $spl46 | $spl204)),
  inference(subsumption_resolution,[],[f22781,f352])).
fof(f352,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f214])).
fof(f22781,plain,(
  ~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'F_0') | ($spl8 | $spl46 | $spl204)),
  inference(resolution,[],[f22772,f16252])).
fof(f16252,plain,(
  ( ! [X588,X589] : (~sP10(X589,X588) | ~p(s(bool,X589)) | s(bool,'F_0') = s(bool,X588)) ) | ($spl8 | $spl46)),
  inference(subsumption_resolution,[],[f16251,f537])).
fof(f16251,plain,(
  ( ! [X588,X589] : (~p(s(bool,'T_0')) | ~p(s(bool,X589)) | ~sP10(X589,X588) | s(bool,'F_0') = s(bool,X588)) ) | ($spl8 | $spl46)),
  inference(forward_demodulation,[],[f15608,f2535])).
fof(f15608,plain,(
  ( ! [X588,X589] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | ~p(s(bool,X589)) | ~sP10(X589,X588) | s(bool,'F_0') = s(bool,X588)) ) | ($spl8 | $spl46)),
  inference(superposition,[],[f439,f15318])).
fof(f22795,plain,(
  $false | ($spl8 | $spl46 | $spl204)),
  inference(subsumption_resolution,[],[f22794,f534])).
fof(f22794,plain,(
  ~p(s(bool,'T_0')) | ($spl8 | $spl46 | $spl204)),
  inference(subsumption_resolution,[],[f22780,f352])).
fof(f22780,plain,(
  ~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'F_0') | ($spl8 | $spl46 | $spl204)),
  inference(resolution,[],[f22772,f16250])).
fof(f16250,plain,(
  ( ! [X586,X587] : (~sP10(X586,X587) | ~p(s(bool,X587)) | s(bool,'F_0') = s(bool,X586)) ) | ($spl8 | $spl46)),
  inference(subsumption_resolution,[],[f16249,f537])).
fof(f16249,plain,(
  ( ! [X586,X587] : (~p(s(bool,'T_0')) | ~p(s(bool,X587)) | ~sP10(X586,X587) | s(bool,'F_0') = s(bool,X586)) ) | ($spl8 | $spl46)),
  inference(forward_demodulation,[],[f15607,f2535])).
fof(f15607,plain,(
  ( ! [X586,X587] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | ~p(s(bool,X587)) | ~sP10(X586,X587) | s(bool,'F_0') = s(bool,X586)) ) | ($spl8 | $spl46)),
  inference(superposition,[],[f439,f15318])).
fof(f22793,plain,(
  $false | $spl204),
  inference(subsumption_resolution,[],[f22789,f534])).
fof(f22789,plain,(
  ~p(s(bool,'T_0')) | $spl204),
  inference(duplicate_literal_removal,[],[f22778])).
fof(f22778,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'T_0')) | $spl204),
  inference(resolution,[],[f22772,f446])).
fof(f446,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | ~sP10(X1,X0)) )),
  inference(cnf_transformation,[],[f257])).
fof(f257,plain,(
  ! [X0,X1] : ((((p(s(bool,X1)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)))) | sP10(X1,X0)) & (~sP10(X1,X0) | ((~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X0)) | p(s(bool,X1))))))),
  inference(nnf_transformation,[],[f189])).
fof(f189,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> sP10(X1,X0))),
  inference(definition_folding,[],[f101,f188])).
fof(f101,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  inference(flattening,[],[f100])).
fof(f100,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  inference(rectify,[],[f60])).
fof(f60,axiom,(
  ! [X23,X24] : ((p(s(bool,X24)) <=> ~p(s(bool,X23))) <=> ((p(s(bool,X24)) | p(s(bool,X23))) & (~p(s(bool,X23)) | ~p(s(bool,X24)))))),
  file('thm.list.MEM_GENLIST','thm.sat.sat_29_prove_ltb')).
fof(f22792,plain,(
  $false | $spl204),
  inference(subsumption_resolution,[],[f22790,f534])).
fof(f22790,plain,(
  ~p(s(bool,'T_0')) | $spl204),
  inference(duplicate_literal_removal,[],[f22777])).
fof(f22777,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'T_0')) | $spl204),
  inference(resolution,[],[f22772,f439])).
fof(f22762,plain,(
  ( ! [X12,X13] : (sP10('T_0',X12) | sP10(X12,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | p(s(bool,X13))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f22619,f14312])).
fof(f22774,plain,(
  ( ! [X43,X44] : (sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP11(X43,X44)) ) | ($spl1 | $spl122)),
  inference(subsumption_resolution,[],[f22756,f455])).
fof(f22756,plain,(
  ( ! [X43,X44] : (sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP11(X43,X44) | p(s(bool,X44))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f22619,f14263])).
fof(f22772,plain,(
  sP10('T_0','T_0') | $spl204),
  inference(cnf_transformation,[],[f22772_D])).
fof(f22772_D,plain,(
  sP10('T_0','T_0') <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f22737,plain,(
  ( ! [X33,X34,X32] : (sP10(X32,X33) | sP11(X34,X33) | ~p(s(bool,X32)) | sP10('T_0','T_0')) ) | $spl122),
  inference(resolution,[],[f22619,f19616])).
fof(f19616,plain,(
  ( ! [X225,X226] : (~sP10(X225,X226) | ~p(s(bool,X226)) | sP10('T_0',X225)) ) | $spl122),
  inference(resolution,[],[f19515,f439])).
fof(f22733,plain,(
  ( ! [X21,X19,X20,X18] : (sP10(X18,X19) | sP11(X20,X19) | s(bool,'F_0') = s(bool,X21) | ~p(s(bool,X18)) | sP10(X21,'T_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(resolution,[],[f22619,f16692])).
fof(f22715,plain,(
  ( ! [X0,X3,X1] : (sP10('T_0',X0) | s(bool,X1) != s(bool,X3) | p(s(bool,X3)) | sP10(X0,X1)) ) | $spl200),
  inference(cnf_transformation,[],[f22715_D])).
fof(f22715_D,plain,(
  ( ! [X0,X3,X1] : (sP10('T_0',X0) | s(bool,X1) != s(bool,X3) | p(s(bool,X3)) | sP10(X0,X1)) ) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f22631,plain,(
  ( ! [X356,X354,X357,X355] : (sP10('T_0',X354) | sP10(X354,X355) | p(s(bool,X356)) | p(s(bool,X357)) | s(bool,X355) != s(bool,X357)) ) | $spl122),
  inference(resolution,[],[f19618,f536])).
fof(f22541,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP10('T_0',X35) | sP10(X35,X36) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X37) | s(bool,X38) != s(bool,X39)) ) | $spl122),
  inference(resolution,[],[f19618,f312])).
fof(f22538,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10('T_0',X20) | sP10(X20,X21) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X21) != s(bool,X24)) ) | $spl122),
  inference(resolution,[],[f19618,f311])).
fof(f22537,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP10('T_0',X15) | sP10(X15,X16) | p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) ) | $spl122),
  inference(resolution,[],[f19618,f310])).
fof(f22536,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP10('T_0',X10) | sP10(X10,X11) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X12) | s(bool,X13) != s(bool,X14)) ) | $spl122),
  inference(resolution,[],[f19618,f308])).
fof(f22534,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10('T_0',X0) | sP10(X0,X1) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) ) | $spl122),
  inference(resolution,[],[f19618,f306])).
fof(f22702,plain,(
  ( ! [X28] : (sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | p(s(bool,X28))) ) | ($spl1 | $spl122)),
  inference(duplicate_literal_removal,[],[f22532])).
fof(f22532,plain,(
  ( ! [X28] : (sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | p(s(bool,X28)) | p(s(bool,X28))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f19618,f14263])).
fof(f22480,plain,(
  ( ! [X279] : (~p(s(bool,X279)) | sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) ) | ($spl1 | $spl122)),
  inference(duplicate_literal_removal,[],[f22463])).
fof(f22463,plain,(
  ( ! [X279] : (~p(s(bool,X279)) | sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | ~p(s(bool,X279))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f19617,f14264])).
fof(f19617,plain,(
  ( ! [X227,X228] : (~sP10(X228,X227) | ~p(s(bool,X228)) | sP10('T_0',X227)) ) | $spl122),
  inference(resolution,[],[f19515,f439])).
fof(f22497,plain,(
  ( ! [X57,X56] : (sP10('T_0',X56) | p(s(bool,X57)) | sP10(X56,X57)) ) | $spl198),
  inference(cnf_transformation,[],[f22497_D])).
fof(f22497_D,plain,(
  ( ! [X57,X56] : (sP10('T_0',X56) | p(s(bool,X57)) | sP10(X56,X57)) ) <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f22496,plain,(
  ( ! [X57,X56,X55] : (~p(s(bool,X55)) | sP10('T_0',X56) | sP10(X56,X57) | p(s(bool,X57))) ) | ($spl8 | $spl46 | $spl122)),
  inference(subsumption_resolution,[],[f22396,f350])).
fof(f22396,plain,(
  ( ! [X57,X56,X55] : (~p(s(bool,X55)) | sP10('T_0',X56) | sP10(X56,X57) | p(s(bool,X57)) | s(bool,'F_0') = s(bool,X55)) ) | ($spl8 | $spl46 | $spl122)),
  inference(resolution,[],[f19617,f16694])).
fof(f22359,plain,(
  ( ! [X291] : (~p(s(bool,X291)) | sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) ) | ($spl1 | $spl122)),
  inference(duplicate_literal_removal,[],[f22356])).
fof(f22356,plain,(
  ( ! [X291] : (~p(s(bool,X291)) | sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | ~p(s(bool,X291))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f19616,f14265])).
fof(f22202,plain,(
  ( ! [X2,X1] : (sP8('F_0',X1,X2)) ) | $spl196),
  inference(cnf_transformation,[],[f22202_D])).
fof(f22202_D,plain,(
  ( ! [X2,X1] : (sP8('F_0',X1,X2)) ) <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f22195,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,'F_0') | sP8('F_0',X1,X2)) ) | $spl122),
  inference(factoring,[],[f19923])).
fof(f20054,plain,(
  ( ! [X76] : (sP10(X76,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | sP10(X76,'F_0')) ) | $spl156),
  inference(cnf_transformation,[],[f20054_D])).
fof(f20054_D,plain,(
  ( ! [X76] : (sP10(X76,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | sP10(X76,'F_0')) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f22194,plain,(
  ( ! [X24,X23,X25] : (sP10(X23,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | sP10(X23,'F_0') | p(s(bool,X24)) | ~p(s(bool,X25))) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f19923,f11976])).
fof(f22132,plain,(
  ( ! [X2,X1] : (sP8(X1,'F_0',X2)) ) | $spl194),
  inference(cnf_transformation,[],[f22132_D])).
fof(f22132_D,plain,(
  ( ! [X2,X1] : (sP8(X1,'F_0',X2)) ) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f22124,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,'F_0') | sP8(X1,'F_0',X2)) ) | $spl122),
  inference(factoring,[],[f19922])).
fof(f22122,plain,(
  ( ! [X21,X22,X20] : (sP10(X20,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | sP10(X20,'F_0') | p(s(bool,X21)) | ~p(s(bool,X22))) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f19922,f11977])).
fof(f22051,plain,(
  ( ! [X2,X1] : (sP6('F_0',X1,X2)) ) | $spl192),
  inference(cnf_transformation,[],[f22051_D])).
fof(f22051_D,plain,(
  ( ! [X2,X1] : (sP6('F_0',X1,X2)) ) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f22045,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,'F_0') | sP6('F_0',X1,X2)) ) | $spl122),
  inference(factoring,[],[f19909])).
fof(f22044,plain,(
  ( ! [X24,X23,X25] : (sP10(X23,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | sP10(X23,'F_0') | p(s(bool,X24)) | p(s(bool,X25))) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f19909,f11960])).
fof(f22043,plain,(
  ( ! [X21,X22,X20] : (sP10(X20,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | sP10(X20,'F_0') | sP7('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X21,X22)) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f19909,f12008])).
fof(f22001,plain,(
  ( ! [X2,X1] : (sP4('F_0',X1,X2)) ) | $spl190),
  inference(cnf_transformation,[],[f22001_D])).
fof(f22001_D,plain,(
  ( ! [X2,X1] : (sP4('F_0',X1,X2)) ) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f21995,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,'F_0') | sP4('F_0',X1,X2)) ) | $spl122),
  inference(factoring,[],[f19893])).
fof(f21994,plain,(
  ( ! [X28,X26,X27] : (sP10(X26,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | sP10(X26,'F_0') | ~p(s(bool,X27)) | p(s(bool,X28))) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f19893,f11946])).
fof(f21993,plain,(
  ( ! [X24,X23,X25] : (sP10(X23,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | sP10(X23,'F_0') | sP5('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X24,X25)) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f19893,f12007])).
fof(f21913,plain,(
  ( ! [X2,X1] : (sP4(X1,X2,'F_0')) ) | $spl188),
  inference(cnf_transformation,[],[f21913_D])).
fof(f21913_D,plain,(
  ( ! [X2,X1] : (sP4(X1,X2,'F_0')) ) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f21906,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,'F_0') | sP4(X1,X2,'F_0')) ) | $spl122),
  inference(factoring,[],[f19892])).
fof(f21903,plain,(
  ( ! [X21,X22,X20] : (sP10(X20,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | sP10(X20,'F_0') | p(s(bool,X21)) | ~p(s(bool,X22))) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f19892,f11947])).
fof(f21860,plain,(
  ( ! [X28,X29,X27] : (sP10(X27,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | sP10(X27,'F_0') | p(s(bool,X28)) | ~p(s(bool,X29))) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f19870,f11923])).
fof(f21868,plain,(
  ( ! [X19,X17,X18,X16] : (sP10(X16,'F_0') | sP3(X18,X17,X19) | p(s(bool,X17)) | ~sP2(X18,X19,X17) | ~sP0(X19,X17,X18)) ) | $spl122),
  inference(subsumption_resolution,[],[f21857,f18882])).
fof(f18882,plain,(
  ( ! [X227,X228] : (~sP10(X227,X228) | p(s(bool,X228)) | sP10(X227,'F_0')) ) | $spl122),
  inference(resolution,[],[f18767,f438])).
fof(f21857,plain,(
  ( ! [X19,X17,X18,X16] : (sP10(X16,X17) | sP10(X16,'F_0') | sP3(X18,X17,X19) | p(s(bool,X17)) | ~sP2(X18,X19,X17) | ~sP0(X19,X17,X18)) ) | $spl122),
  inference(resolution,[],[f19870,f318])).
fof(f21785,plain,(
  ( ! [X2,X1] : (sP1(X1,'F_0',X2)) ) | $spl186),
  inference(cnf_transformation,[],[f21785_D])).
fof(f21785_D,plain,(
  ( ! [X2,X1] : (sP1(X1,'F_0',X2)) ) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f21777,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,'F_0') | sP1(X1,'F_0',X2)) ) | $spl122),
  inference(factoring,[],[f19869])).
fof(f21775,plain,(
  ( ! [X26,X24,X25] : (sP10(X24,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | sP10(X24,'F_0') | p(s(bool,X25)) | ~p(s(bool,X26))) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f19869,f11924])).
fof(f21687,plain,(
  ( ! [X2,X1] : (sP2('F_0',X1,X2)) ) | $spl184),
  inference(cnf_transformation,[],[f21687_D])).
fof(f21687_D,plain,(
  ( ! [X2,X1] : (sP2('F_0',X1,X2)) ) <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f21678,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,'F_0') | sP2('F_0',X1,X2)) ) | $spl122),
  inference(factoring,[],[f19867])).
fof(f21677,plain,(
  ( ! [X39,X38,X40] : (sP10(X38,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | sP10(X38,'F_0') | p(s(bool,X39)) | ~p(s(bool,X40))) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f19867,f11920])).
fof(f21686,plain,(
  ( ! [X19,X17,X18,X16] : (sP10(X16,'F_0') | sP3(X17,X18,X19) | p(s(bool,X17)) | ~sP1(X18,X19,X17) | ~sP0(X19,X18,X17)) ) | $spl122),
  inference(subsumption_resolution,[],[f21670,f18882])).
fof(f21670,plain,(
  ( ! [X19,X17,X18,X16] : (sP10(X16,X17) | sP10(X16,'F_0') | sP3(X17,X18,X19) | p(s(bool,X17)) | ~sP1(X18,X19,X17) | ~sP0(X19,X18,X17)) ) | $spl122),
  inference(resolution,[],[f19867,f319])).
fof(f21535,plain,(
  ( ! [X2,X1] : (sP2(X1,'F_0',X2)) ) | $spl182),
  inference(cnf_transformation,[],[f21535_D])).
fof(f21535_D,plain,(
  ( ! [X2,X1] : (sP2(X1,'F_0',X2)) ) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f21527,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,'F_0') | sP2(X1,'F_0',X2)) ) | $spl122),
  inference(factoring,[],[f19866])).
fof(f21525,plain,(
  ( ! [X37,X35,X36] : (sP10(X35,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | sP10(X35,'F_0') | p(s(bool,X36)) | ~p(s(bool,X37))) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f19866,f11921])).
fof(f21534,plain,(
  ( ! [X24,X23] : (sP10(X23,X24) | s(bool,'T_0') != s(bool,X24) | sP10(X23,'F_0')) ) | $spl180),
  inference(cnf_transformation,[],[f21534_D])).
fof(f21534_D,plain,(
  ( ! [X24,X23] : (sP10(X23,X24) | s(bool,'T_0') != s(bool,X24) | sP10(X23,'F_0')) ) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f21521,plain,(
  ( ! [X24,X23,X25] : (sP10(X23,X24) | sP10(X23,'F_0') | s(bool,'T_0') != s(bool,X24) | ~p(s(bool,X25))) ) | ($spl6 | $spl122)),
  inference(resolution,[],[f19866,f14550])).
fof(f19487,plain,(
  ( ! [X76] : (sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X76) | sP10(X76,'F_0')) ) | $spl144),
  inference(cnf_transformation,[],[f19487_D])).
fof(f19487_D,plain,(
  ( ! [X76] : (sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X76) | sP10(X76,'F_0')) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f21461,plain,(
  ( ! [X24,X23,X25] : (sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X23) | sP10(X23,'F_0') | p(s(bool,X24)) | ~p(s(bool,X25))) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f19370,f11976])).
fof(f21417,plain,(
  ( ! [X21,X22,X20] : (sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X20) | sP10(X20,'F_0') | p(s(bool,X21)) | ~p(s(bool,X22))) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f19369,f11977])).
fof(f21377,plain,(
  ( ! [X24,X23,X25] : (sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X23) | sP10(X23,'F_0') | p(s(bool,X24)) | p(s(bool,X25))) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f19356,f11960])).
fof(f21376,plain,(
  ( ! [X21,X22,X20] : (sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X20) | sP10(X20,'F_0') | sP7('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X21,X22)) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f19356,f12008])).
fof(f21336,plain,(
  ( ! [X28,X26,X27] : (sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X26) | sP10(X26,'F_0') | ~p(s(bool,X27)) | p(s(bool,X28))) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f19340,f11946])).
fof(f21335,plain,(
  ( ! [X24,X23,X25] : (sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X23) | sP10(X23,'F_0') | sP5('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X24,X25)) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f19340,f12007])).
fof(f21291,plain,(
  ( ! [X21,X22,X20] : (sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X20) | sP10(X20,'F_0') | p(s(bool,X21)) | ~p(s(bool,X22))) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f19339,f11947])).
fof(f21229,plain,(
  ( ! [X28,X29,X27] : (sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X27) | sP10(X27,'F_0') | p(s(bool,X28)) | ~p(s(bool,X29))) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f19317,f11923])).
fof(f21237,plain,(
  ( ! [X19,X17,X18,X16] : (sP10(X17,'F_0') | sP3(X18,X16,X19) | p(s(bool,X16)) | ~sP2(X18,X19,X16) | ~sP0(X19,X16,X18)) ) | $spl122),
  inference(subsumption_resolution,[],[f21226,f18881])).
fof(f18881,plain,(
  ( ! [X225,X226] : (~sP10(X226,X225) | p(s(bool,X226)) | sP10(X225,'F_0')) ) | $spl122),
  inference(resolution,[],[f18767,f438])).
fof(f21226,plain,(
  ( ! [X19,X17,X18,X16] : (sP10(X16,X17) | sP10(X17,'F_0') | sP3(X18,X16,X19) | p(s(bool,X16)) | ~sP2(X18,X19,X16) | ~sP0(X19,X16,X18)) ) | $spl122),
  inference(resolution,[],[f19317,f318])).
fof(f21184,plain,(
  ( ! [X26,X24,X25] : (sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X24) | sP10(X24,'F_0') | p(s(bool,X25)) | ~p(s(bool,X26))) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f19316,f11924])).
fof(f21139,plain,(
  ( ! [X39,X38,X40] : (sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X38) | sP10(X38,'F_0') | p(s(bool,X39)) | ~p(s(bool,X40))) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f19314,f11920])).
fof(f21148,plain,(
  ( ! [X19,X17,X18,X16] : (sP10(X17,'F_0') | sP3(X16,X18,X19) | p(s(bool,X16)) | ~sP1(X18,X19,X16) | ~sP0(X19,X18,X16)) ) | $spl122),
  inference(subsumption_resolution,[],[f21132,f18881])).
fof(f21132,plain,(
  ( ! [X19,X17,X18,X16] : (sP10(X16,X17) | sP10(X17,'F_0') | sP3(X16,X18,X19) | p(s(bool,X16)) | ~sP1(X18,X19,X16) | ~sP0(X19,X18,X16)) ) | $spl122),
  inference(resolution,[],[f19314,f319])).
fof(f21090,plain,(
  ( ! [X37,X35,X36] : (sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X35) | sP10(X35,'F_0') | p(s(bool,X36)) | ~p(s(bool,X37))) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f19313,f11921])).
fof(f21099,plain,(
  ( ! [X24,X23] : (sP10(X23,X24) | s(bool,'T_0') != s(bool,X23) | sP10(X24,'F_0')) ) | $spl174),
  inference(cnf_transformation,[],[f21099_D])).
fof(f21099_D,plain,(
  ( ! [X24,X23] : (sP10(X23,X24) | s(bool,'T_0') != s(bool,X23) | sP10(X24,'F_0')) ) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f21086,plain,(
  ( ! [X24,X23,X25] : (sP10(X23,X24) | sP10(X24,'F_0') | s(bool,'T_0') != s(bool,X23) | ~p(s(bool,X25))) ) | ($spl6 | $spl122)),
  inference(resolution,[],[f19313,f14550])).
fof(f21029,plain,(
  ( ! [X24,X23] : (sP10(X23,'F_0') | s(bool,'F_0') = s(bool,X24) | sP10(X23,X24)) ) | $spl172),
  inference(cnf_transformation,[],[f21029_D])).
fof(f21029_D,plain,(
  ( ! [X24,X23] : (sP10(X23,'F_0') | s(bool,'F_0') = s(bool,X24) | sP10(X23,X24)) ) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f21042,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | sP10(X26,'F_0') | sP10(X26,X27) | s(bool,'F_0') = s(bool,X27)) ) | ($spl8 | $spl46 | $spl122)),
  inference(resolution,[],[f18977,f16926])).
fof(f21019,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | sP10(X23,'F_0') | sP10(X23,X24) | s(bool,'F_0') = s(bool,X24)) ) | ($spl8 | $spl46 | $spl122)),
  inference(resolution,[],[f18976,f16925])).
fof(f20841,plain,(
  sP10('F_0','const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | $spl168),
  inference(cnf_transformation,[],[f20841_D])).
fof(f20841_D,plain,(
  sP10('F_0','const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f20961,plain,(
  ( ! [X17,X18] : (sP10('F_0','const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X17)) | ~p(s(bool,X18))) ) | ($spl8 | $spl122 | $spl133)),
  inference(resolution,[],[f20726,f11976])).
fof(f20726,plain,(
  ( ! [X202,X203,X201] : (sP8(X201,X202,X203) | sP10('F_0',X201)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f20649,f411])).
fof(f20951,plain,(
  ( ! [X15,X16] : (sP10('F_0','const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X15)) | ~p(s(bool,X16))) ) | ($spl8 | $spl122 | $spl133)),
  inference(resolution,[],[f20725,f11977])).
fof(f20725,plain,(
  ( ! [X198,X200,X199] : (sP8(X199,X198,X200) | sP10('F_0',X198)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f20649,f410])).
fof(f20945,plain,(
  ( ! [X17,X18] : (sP10('F_0','const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X17)) | p(s(bool,X18))) ) | ($spl8 | $spl122 | $spl133)),
  inference(resolution,[],[f20712,f11960])).
fof(f20712,plain,(
  ( ! [X161,X159,X160] : (sP6(X159,X160,X161) | sP10('F_0',X159)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f20649,f394])).
fof(f20944,plain,(
  ( ! [X15,X16] : (sP10('F_0','const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | sP7('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X15,X16)) ) | ($spl8 | $spl122 | $spl133)),
  inference(resolution,[],[f20712,f12008])).
fof(f20938,plain,(
  ( ! [X19,X20] : (sP10('F_0','const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | ~p(s(bool,X19)) | p(s(bool,X20))) ) | ($spl8 | $spl122 | $spl133)),
  inference(resolution,[],[f20696,f11946])).
fof(f20696,plain,(
  ( ! [X111,X112,X113] : (sP4(X111,X112,X113) | sP10('F_0',X111)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f20649,f377])).
fof(f20937,plain,(
  ( ! [X17,X18] : (sP10('F_0','const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | sP5('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X17,X18)) ) | ($spl8 | $spl122 | $spl133)),
  inference(resolution,[],[f20696,f12007])).
fof(f20898,plain,(
  ( ! [X15,X16] : (sP10('F_0','const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X15)) | ~p(s(bool,X16))) ) | ($spl8 | $spl122 | $spl133)),
  inference(resolution,[],[f20695,f11947])).
fof(f20695,plain,(
  ( ! [X109,X110,X108] : (sP4(X109,X110,X108) | sP10('F_0',X108)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f20649,f375])).
fof(f20889,plain,(
  ( ! [X21,X20] : (sP10('F_0','const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X20)) | ~p(s(bool,X21))) ) | ($spl8 | $spl122 | $spl133)),
  inference(resolution,[],[f20673,f11923])).
fof(f20673,plain,(
  ( ! [X57,X58,X56] : (sP1(X56,X57,X58) | sP10('F_0',X56)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f20649,f327])).
fof(f20878,plain,(
  ( ! [X19,X18] : (sP10('F_0','const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X18)) | ~p(s(bool,X19))) ) | ($spl8 | $spl122 | $spl133)),
  inference(resolution,[],[f20672,f11924])).
fof(f20672,plain,(
  ( ! [X54,X55,X53] : (sP1(X54,X53,X55) | sP10('F_0',X53)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f20649,f326])).
fof(f20868,plain,(
  ( ! [X28,X27] : (sP10('F_0','const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X27)) | ~p(s(bool,X28))) ) | ($spl8 | $spl122 | $spl133)),
  inference(resolution,[],[f20670,f11920])).
fof(f20670,plain,(
  ( ! [X47,X48,X49] : (sP2(X47,X48,X49) | sP10('F_0',X47)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f20649,f323])).
fof(f20853,plain,(
  ( ! [X26,X25] : (sP10('F_0','const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X25)) | ~p(s(bool,X26))) ) | ($spl8 | $spl122 | $spl133)),
  inference(resolution,[],[f20669,f11921])).
fof(f20669,plain,(
  ( ! [X45,X46,X44] : (sP2(X45,X44,X46) | sP10('F_0',X44)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f20649,f322])).
fof(f20856,plain,(
  ( ! [X17] : (sP10('F_0',X17) | s(bool,'T_0') != s(bool,X17)) ) | $spl170),
  inference(cnf_transformation,[],[f20856_D])).
fof(f20856_D,plain,(
  ( ! [X17] : (sP10('F_0',X17) | s(bool,'T_0') != s(bool,X17)) ) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f20849,plain,(
  ( ! [X17,X18] : (sP10('F_0',X17) | s(bool,'T_0') != s(bool,X17) | ~p(s(bool,X18))) ) | ($spl6 | $spl122 | $spl133)),
  inference(resolution,[],[f20669,f14550])).
fof(f20839,plain,(
  ( ! [X45] : (~p(s(bool,X45)) | sP10('F_0','const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | ($spl8 | $spl46 | $spl122 | $spl133)),
  inference(subsumption_resolution,[],[f20818,f350])).
fof(f20818,plain,(
  ( ! [X45] : (~p(s(bool,X45)) | sP10('F_0','const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | s(bool,'F_0') = s(bool,X45)) ) | ($spl8 | $spl46 | $spl122 | $spl133)),
  inference(superposition,[],[f20649,f15318])).
fof(f20829,plain,(
  ( ! [X2,X0] : (sP10('F_0',X0) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) | $spl166),
  inference(cnf_transformation,[],[f20829_D])).
fof(f20829_D,plain,(
  ( ! [X2,X0] : (sP10('F_0',X0) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f20755,plain,(
  ( ! [X259,X260,X258] : (sP10('F_0',X258) | p(s(bool,X259)) | ~p(s(bool,X260)) | s(bool,X258) != s(bool,X260)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f20649,f536])).
fof(f20664,plain,(
  ( ! [X30,X28,X31,X29] : (sP10('F_0',X28) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f20649,f312])).
fof(f20662,plain,(
  ( ! [X23,X21,X22,X20] : (sP10('F_0',X20) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X20) != s(bool,X22)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f20649,f311])).
fof(f20660,plain,(
  ( ! [X14,X12,X15,X13] : (sP10('F_0',X12) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f20649,f310])).
fof(f20659,plain,(
  ( ! [X10,X8,X11,X9] : (sP10('F_0',X8) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f20649,f308])).
fof(f20657,plain,(
  ( ! [X2,X0,X3,X1] : (sP10('F_0',X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | ($spl122 | $spl133)),
  inference(resolution,[],[f20649,f306])).
fof(f20654,plain,(
  ( ! [X8] : (sP10('F_0',X8) | ~p(s(bool,X8))) ) | $spl164),
  inference(cnf_transformation,[],[f20654_D])).
fof(f20654_D,plain,(
  ( ! [X8] : (sP10('F_0',X8) | ~p(s(bool,X8))) ) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f20631,plain,(
  ( ! [X8,X9] : (sP10('F_0',X8) | s(bool,'F_0') = s(bool,X9) | ~p(s(bool,X8)) | sP10(X9,'T_0')) ) | ($spl8 | $spl46 | $spl122 | $spl133)),
  inference(resolution,[],[f20608,f16693])).
fof(f20619,plain,(
  ( ! [X23] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0') | sP10(X23,'F_0')) ) | ($spl1 | $spl122)),
  inference(subsumption_resolution,[],[f20594,f18767])).
fof(f20594,plain,(
  ( ! [X23] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0') | sP10(X23,'F_0') | p(s(bool,X23))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f20245,f14262])).
fof(f20566,plain,(
  ( ! [X24] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0') | sP10(X24,'F_0')) ) | ($spl1 | $spl122)),
  inference(subsumption_resolution,[],[f20544,f18767])).
fof(f20544,plain,(
  ( ! [X24] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0') | sP10(X24,'F_0') | p(s(bool,X24))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f20244,f14263])).
fof(f20244,plain,(
  ( ! [X267,X268] : (sP10(X267,'T_0') | sP10(X267,X268) | sP10(X268,'F_0')) ) | $spl122),
  inference(resolution,[],[f20069,f18883])).
fof(f20495,plain,(
  ( ! [X23] : (sP10(X23,'F_0') | sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) ) | ($spl1 | $spl122)),
  inference(subsumption_resolution,[],[f20461,f18767])).
fof(f20461,plain,(
  ( ! [X23] : (sP10(X23,'F_0') | sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | p(s(bool,X23))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f19958,f14262])).
fof(f19958,plain,(
  ( ! [X379,X378] : (sP10(X378,X379) | sP10(X378,'F_0') | sP10('T_0',X379)) ) | $spl122),
  inference(resolution,[],[f18884,f19515])).
fof(f20327,plain,(
  ( ! [X5] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0') | p(s(bool,X5))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f20229,f14313])).
fof(f20229,plain,(
  ( ! [X243,X244] : (sP11(X243,X244) | sP10(X243,'T_0')) ) | $spl122),
  inference(resolution,[],[f20069,f456])).
fof(f20320,plain,(
  ( ! [X4] : (sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),'T_0') | p(s(bool,X4))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f20228,f14312])).
fof(f20228,plain,(
  ( ! [X241,X242] : (sP11(X242,X241) | sP10(X241,'T_0')) ) | $spl122),
  inference(resolution,[],[f20069,f455])).
fof(f20309,plain,(
  ( ! [X2,X0] : (sP10(X0,'T_0') | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) | $spl158),
  inference(cnf_transformation,[],[f20309_D])).
fof(f20309_D,plain,(
  ( ! [X2,X0] : (sP10(X0,'T_0') | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f20240,plain,(
  ( ! [X257,X259,X258] : (sP10(X257,'T_0') | p(s(bool,X258)) | p(s(bool,X259)) | s(bool,X257) != s(bool,X259)) ) | $spl122),
  inference(resolution,[],[f20069,f536])).
fof(f20150,plain,(
  ( ! [X30,X28,X31,X29] : (sP10(X28,'T_0') | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) ) | $spl122),
  inference(resolution,[],[f20069,f312])).
fof(f20147,plain,(
  ( ! [X19,X17,X18,X16] : (sP10(X16,'T_0') | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) ) | $spl122),
  inference(resolution,[],[f20069,f311])).
fof(f20146,plain,(
  ( ! [X14,X12,X15,X13] : (sP10(X12,'T_0') | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl122),
  inference(resolution,[],[f20069,f310])).
fof(f20145,plain,(
  ( ! [X10,X8,X11,X9] : (sP10(X8,'T_0') | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) ) | $spl122),
  inference(resolution,[],[f20069,f308])).
fof(f20143,plain,(
  ( ! [X2,X0,X3,X1] : (sP10(X0,'T_0') | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl122),
  inference(resolution,[],[f20069,f306])).
fof(f20053,plain,(
  ( ! [X76,X75] : (~p(s(bool,X75)) | sP10(X76,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | sP10(X76,'F_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(subsumption_resolution,[],[f20014,f350])).
fof(f20014,plain,(
  ( ! [X76,X75] : (~p(s(bool,X75)) | sP10(X76,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | sP10(X76,'F_0') | s(bool,'F_0') = s(bool,X75)) ) | ($spl8 | $spl46 | $spl122)),
  inference(superposition,[],[f18884,f15318])).
fof(f20041,plain,(
  ( ! [X380] : (sP10(X380,'T_0') | sP10(X380,'F_0')) ) | $spl154),
  inference(cnf_transformation,[],[f20041_D])).
fof(f20041_D,plain,(
  ( ! [X380] : (sP10(X380,'T_0') | sP10(X380,'F_0')) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f20023,plain,(
  ( ! [X64,X65] : (~p(s(bool,X64)) | sP10(X65,'T_0') | sP10(X65,'F_0')) ) | $spl122),
  inference(duplicate_literal_removal,[],[f20008])).
fof(f20008,plain,(
  ( ! [X64,X65] : (~p(s(bool,X64)) | sP10(X65,'T_0') | sP10(X65,'F_0') | ~p(s(bool,X64))) ) | $spl122),
  inference(superposition,[],[f18884,f355])).
fof(f20024,plain,(
  ( ! [X62,X63] : (~p(s(bool,X62)) | sP10(X63,'T_0') | sP10(X63,'F_0')) ) | $spl122),
  inference(duplicate_literal_removal,[],[f20007])).
fof(f20007,plain,(
  ( ! [X62,X63] : (~p(s(bool,X62)) | sP10(X63,'T_0') | sP10(X63,'F_0') | ~p(s(bool,X62))) ) | $spl122),
  inference(superposition,[],[f18884,f357])).
fof(f20052,plain,(
  ( ! [X61,X60] : (~p(s(bool,X60)) | sP10(X61,'T_0') | sP10(X61,'F_0')) ) | $spl122),
  inference(subsumption_resolution,[],[f20006,f350])).
fof(f20006,plain,(
  ( ! [X61,X60] : (~p(s(bool,X60)) | sP10(X61,'T_0') | sP10(X61,'F_0') | s(bool,'F_0') = s(bool,X60)) ) | $spl122),
  inference(superposition,[],[f18884,f358])).
fof(f20051,plain,(
  ( ! [X59,X58] : (~p(s(bool,X58)) | sP10(X59,'T_0') | sP10(X59,'F_0')) ) | $spl122),
  inference(subsumption_resolution,[],[f20005,f350])).
fof(f20005,plain,(
  ( ! [X59,X58] : (~p(s(bool,X58)) | sP10(X59,'T_0') | sP10(X59,'F_0') | s(bool,'F_0') = s(bool,X58)) ) | $spl122),
  inference(superposition,[],[f18884,f359])).
fof(f20029,plain,(
  ( ! [X414,X413,X415] : (sP10(X413,'F_0') | p(s(bool,X414)) | p(s(bool,X415)) | ~sP11(X414,X415)) ) | $spl122),
  inference(duplicate_literal_removal,[],[f19979])).
fof(f19979,plain,(
  ( ! [X414,X413,X415] : (sP10(X413,'F_0') | sP10(X413,'F_0') | p(s(bool,X414)) | p(s(bool,X415)) | ~sP11(X414,X415)) ) | $spl122),
  inference(resolution,[],[f18884,f454])).
fof(f20031,plain,(
  ( ! [X408,X407,X409] : (sP10(X407,'F_0') | p(s(bool,X408)) | p(s(bool,X409)) | ~sP11(X408,X409)) ) | $spl122),
  inference(duplicate_literal_removal,[],[f19977])).
fof(f19977,plain,(
  ( ! [X408,X407,X409] : (sP10(X407,'F_0') | sP10(X407,'F_0') | p(s(bool,X408)) | p(s(bool,X409)) | ~sP11(X408,X409)) ) | $spl122),
  inference(resolution,[],[f18884,f458])).
fof(f19968,plain,(
  ( ! [X397,X398] : (sP10(X397,'T_0') | sP10(X397,'F_0') | ~p(s(bool,X398))) ) | $spl122),
  inference(resolution,[],[f18884,f507])).
fof(f507,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f293])).
fof(f293,plain,(
  ! [X0] : ((~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,'T_0')) & p(s(bool,X0)))))),
  inference(flattening,[],[f292])).
fof(f292,plain,(
  ! [X0] : (((~p(s(bool,'T_0')) | ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,'T_0')) & p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f124])).
fof(f124,plain,(
  ! [X0] : ((p(s(bool,'T_0')) & p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X11] : ((p(s(bool,'T_0')) & p(s(bool,X11))) <=> p(s(bool,X11)))),
  file('thm.list.MEM_GENLIST','thm.bool.AND_CLAUSES|split|0')).
fof(f19967,plain,(
  ( ! [X395,X396] : (sP10(X395,'T_0') | sP10(X395,'F_0') | ~p(s(bool,X396))) ) | $spl122),
  inference(resolution,[],[f18884,f505])).
fof(f505,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f291])).
fof(f291,plain,(
  ! [X0] : ((~p(s(bool,X0)) | ~p(s(bool,'T_0')) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,X0)) & p(s(bool,'T_0')))))),
  inference(flattening,[],[f290])).
fof(f290,plain,(
  ! [X0] : (((~p(s(bool,X0)) | ~p(s(bool,'T_0'))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,X0)) & p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0] : ((p(s(bool,X0)) & p(s(bool,'T_0'))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X11] : ((p(s(bool,X11)) & p(s(bool,'T_0'))) <=> p(s(bool,X11)))),
  file('thm.list.MEM_GENLIST','thm.bool.AND_CLAUSES|split|1')).
fof(f19966,plain,(
  ( ! [X393,X394] : (sP10(X393,'T_0') | sP10(X393,'F_0') | ~p(s(bool,X394))) ) | $spl122),
  inference(resolution,[],[f18884,f500])).
fof(f500,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f289])).
fof(f289,plain,(
  ! [X0] : (((~p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | p(s(bool,X0)) | p(s(bool,'T_0'))))),
  inference(flattening,[],[f288])).
fof(f288,plain,(
  ! [X0] : (((~p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (p(s(bool,X0)) | p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0] : ((p(s(bool,X0)) | p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X11] : ((p(s(bool,X11)) | p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  file('thm.list.MEM_GENLIST','thm.bool.OR_CLAUSES|split|1')).
fof(f19965,plain,(
  ( ! [X391,X392] : (sP10(X391,'T_0') | sP10(X391,'F_0') | ~p(s(bool,X392))) ) | $spl122),
  inference(resolution,[],[f18884,f498])).
fof(f498,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f287])).
fof(f287,plain,(
  ! [X0] : (((~p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | p(s(bool,'T_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f286])).
fof(f286,plain,(
  ! [X0] : (((~p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (p(s(bool,'T_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f121])).
fof(f121,plain,(
  ! [X0] : ((p(s(bool,'T_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X11] : ((p(s(bool,'T_0')) | p(s(bool,X11))) <=> p(s(bool,'T_0')))),
  file('thm.list.MEM_GENLIST','thm.bool.OR_CLAUSES|split|0')).
fof(f19964,plain,(
  ( ! [X389,X390] : (sP10(X389,'T_0') | sP10(X389,'F_0') | p(s(bool,X390))) ) | $spl122),
  inference(resolution,[],[f18884,f485])).
fof(f485,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f279])).
fof(f279,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,'T_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f278])).
fof(f278,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (~p(s(bool,'T_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f165])).
fof(f165,plain,(
  ! [X0] : ((~p(s(bool,'T_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(ennf_transformation,[],[f117])).
fof(f117,plain,(
  ! [X0] : ((p(s(bool,'T_0')) => p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X11] : ((p(s(bool,'T_0')) => p(s(bool,X11))) <=> p(s(bool,X11)))),
  file('thm.list.MEM_GENLIST','thm.bool.IMP_CLAUSES|split|0')).
fof(f19963,plain,(
  ( ! [X387,X388] : (sP10(X387,'T_0') | sP10(X387,'F_0') | ~p(s(bool,X388))) ) | $spl122),
  inference(resolution,[],[f18884,f483])).
fof(f483,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f277])).
fof(f277,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f276])).
fof(f276,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f164])).
fof(f164,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f116])).
fof(f116,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X11] : ((p(s(bool,'F_0')) => p(s(bool,X11))) <=> p(s(bool,'T_0')))),
  file('thm.list.MEM_GENLIST','thm.bool.IMP_CLAUSES|split|2')).
fof(f19961,plain,(
  ( ! [X385,X384] : (sP10(X384,'T_0') | sP10(X384,'F_0') | p(s(bool,X385))) ) | $spl122),
  inference(resolution,[],[f18884,f479])).
fof(f479,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f275])).
fof(f275,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,'T_0'))))),
  inference(flattening,[],[f274])).
fof(f274,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f163])).
fof(f163,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f115])).
fof(f115,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X11] : ((p(s(bool,X11)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  file('thm.list.MEM_GENLIST','thm.bool.IMP_CLAUSES|split|1')).
fof(f19960,plain,(
  ( ! [X383,X382] : (sP10(X382,'T_0') | sP10(X382,'F_0') | ~p(s(bool,X383))) ) | $spl122),
  inference(resolution,[],[f18884,f477])).
fof(f477,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f273])).
fof(f273,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))))),
  inference(flattening,[],[f272])).
fof(f272,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f162])).
fof(f162,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f114])).
fof(f114,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X11] : ((p(s(bool,X11)) => p(s(bool,X11))) <=> p(s(bool,'T_0')))),
  file('thm.list.MEM_GENLIST','thm.bool.IMP_CLAUSES|split|3')).
fof(f19959,plain,(
  ( ! [X380,X381] : (sP10(X380,'T_0') | sP10(X380,'F_0') | p(s(bool,X381))) ) | $spl122),
  inference(resolution,[],[f18884,f476])).
fof(f476,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f273])).
fof(f20035,plain,(
  ( ! [X0,X3,X1] : (sP10(X0,X1) | s(bool,X1) != s(bool,X3) | ~p(s(bool,X3)) | sP10(X0,'F_0')) ) | $spl152),
  inference(cnf_transformation,[],[f20035_D])).
fof(f20035_D,plain,(
  ( ! [X0,X3,X1] : (sP10(X0,X1) | s(bool,X1) != s(bool,X3) | ~p(s(bool,X3)) | sP10(X0,'F_0')) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f19952,plain,(
  ( ! [X356,X358,X357,X359] : (sP10(X356,X357) | sP10(X356,'F_0') | p(s(bool,X358)) | ~p(s(bool,X359)) | s(bool,X357) != s(bool,X359)) ) | $spl122),
  inference(resolution,[],[f18884,f536])).
fof(f19861,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP10(X35,X36) | sP10(X35,'F_0') | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X38) | s(bool,X37) != s(bool,X39)) ) | $spl122),
  inference(resolution,[],[f18884,f312])).
fof(f19859,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP10(X25,'F_0') | p(s(bool,X27)) | ~p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X26) != s(bool,X28)) ) | $spl122),
  inference(resolution,[],[f18884,f311])).
fof(f19857,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP10(X15,X16) | sP10(X15,'F_0') | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) ) | $spl122),
  inference(resolution,[],[f18884,f310])).
fof(f19856,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP10(X10,X11) | sP10(X10,'F_0') | ~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,X14)) ) | $spl122),
  inference(resolution,[],[f18884,f308])).
fof(f19854,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10(X0,X1) | sP10(X0,'F_0') | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) ) | $spl122),
  inference(resolution,[],[f18884,f306])).
fof(f19848,plain,(
  ( ! [X24] : (sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP10(X24,'F_0')) ) | ($spl1 | $spl122)),
  inference(subsumption_resolution,[],[f19825,f18767])).
fof(f19825,plain,(
  ( ! [X24] : (sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | sP10(X24,'F_0') | p(s(bool,X24))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f19640,f14263])).
fof(f19640,plain,(
  ( ! [X267,X268] : (sP10('T_0',X267) | sP10(X267,X268) | sP10(X268,'F_0')) ) | $spl122),
  inference(resolution,[],[f19515,f18883])).
fof(f19722,plain,(
  ( ! [X5] : (sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | p(s(bool,X5))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f19625,f14313])).
fof(f19625,plain,(
  ( ! [X243,X244] : (sP11(X243,X244) | sP10('T_0',X243)) ) | $spl122),
  inference(resolution,[],[f19515,f456])).
fof(f19715,plain,(
  ( ! [X4] : (sP10('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | p(s(bool,X4))) ) | ($spl1 | $spl122)),
  inference(resolution,[],[f19624,f14312])).
fof(f19624,plain,(
  ( ! [X241,X242] : (sP11(X242,X241) | sP10('T_0',X241)) ) | $spl122),
  inference(resolution,[],[f19515,f455])).
fof(f19704,plain,(
  ( ! [X2,X0] : (sP10('T_0',X0) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) | $spl146),
  inference(cnf_transformation,[],[f19704_D])).
fof(f19704_D,plain,(
  ( ! [X2,X0] : (sP10('T_0',X0) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f19636,plain,(
  ( ! [X257,X259,X258] : (sP10('T_0',X257) | p(s(bool,X258)) | p(s(bool,X259)) | s(bool,X257) != s(bool,X259)) ) | $spl122),
  inference(resolution,[],[f19515,f536])).
fof(f19546,plain,(
  ( ! [X30,X28,X31,X29] : (sP10('T_0',X28) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) ) | $spl122),
  inference(resolution,[],[f19515,f312])).
fof(f19543,plain,(
  ( ! [X19,X17,X18,X16] : (sP10('T_0',X16) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) ) | $spl122),
  inference(resolution,[],[f19515,f311])).
fof(f19542,plain,(
  ( ! [X14,X12,X15,X13] : (sP10('T_0',X12) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl122),
  inference(resolution,[],[f19515,f310])).
fof(f19541,plain,(
  ( ! [X10,X8,X11,X9] : (sP10('T_0',X8) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) ) | $spl122),
  inference(resolution,[],[f19515,f308])).
fof(f19539,plain,(
  ( ! [X2,X0,X3,X1] : (sP10('T_0',X0) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl122),
  inference(resolution,[],[f19515,f306])).
fof(f19530,plain,(
  ( ! [X14,X15] : (p(s(bool,X14)) | sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X15) | s(bool,'F_0') = s(bool,X15)) ) | ($spl1 | $spl8 | $spl46)),
  inference(resolution,[],[f14313,f16926])).
fof(f19528,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | sP10(X11,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | s(bool,'F_0') = s(bool,X11)) ) | ($spl1 | $spl8 | $spl46)),
  inference(resolution,[],[f14313,f16701])).
fof(f19486,plain,(
  ( ! [X76,X75] : (~p(s(bool,X75)) | sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X76) | sP10(X76,'F_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(subsumption_resolution,[],[f19460,f350])).
fof(f19460,plain,(
  ( ! [X76,X75] : (~p(s(bool,X75)) | sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X76) | sP10(X76,'F_0') | s(bool,'F_0') = s(bool,X75)) ) | ($spl8 | $spl46 | $spl122)),
  inference(superposition,[],[f18883,f15318])).
fof(f19477,plain,(
  ( ! [X378] : (sP10('T_0',X378) | sP10(X378,'F_0')) ) | $spl140),
  inference(cnf_transformation,[],[f19477_D])).
fof(f19477_D,plain,(
  ( ! [X378] : (sP10('T_0',X378) | sP10(X378,'F_0')) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f19465,plain,(
  ( ! [X64,X65] : (~p(s(bool,X64)) | sP10('T_0',X65) | sP10(X65,'F_0')) ) | $spl122),
  inference(duplicate_literal_removal,[],[f19454])).
fof(f19454,plain,(
  ( ! [X64,X65] : (~p(s(bool,X64)) | sP10('T_0',X65) | sP10(X65,'F_0') | ~p(s(bool,X64))) ) | $spl122),
  inference(superposition,[],[f18883,f355])).
fof(f19466,plain,(
  ( ! [X62,X63] : (~p(s(bool,X62)) | sP10('T_0',X63) | sP10(X63,'F_0')) ) | $spl122),
  inference(duplicate_literal_removal,[],[f19453])).
fof(f19453,plain,(
  ( ! [X62,X63] : (~p(s(bool,X62)) | sP10('T_0',X63) | sP10(X63,'F_0') | ~p(s(bool,X62))) ) | $spl122),
  inference(superposition,[],[f18883,f357])).
fof(f19485,plain,(
  ( ! [X61,X60] : (~p(s(bool,X60)) | sP10('T_0',X61) | sP10(X61,'F_0')) ) | $spl122),
  inference(subsumption_resolution,[],[f19452,f350])).
fof(f19452,plain,(
  ( ! [X61,X60] : (~p(s(bool,X60)) | sP10('T_0',X61) | sP10(X61,'F_0') | s(bool,'F_0') = s(bool,X60)) ) | $spl122),
  inference(superposition,[],[f18883,f358])).
fof(f19484,plain,(
  ( ! [X59,X58] : (~p(s(bool,X58)) | sP10('T_0',X59) | sP10(X59,'F_0')) ) | $spl122),
  inference(subsumption_resolution,[],[f19451,f350])).
fof(f19451,plain,(
  ( ! [X59,X58] : (~p(s(bool,X58)) | sP10('T_0',X59) | sP10(X59,'F_0') | s(bool,'F_0') = s(bool,X58)) ) | $spl122),
  inference(superposition,[],[f18883,f359])).
fof(f19425,plain,(
  ( ! [X412,X413,X411] : (sP10('F_0',X411) | sP10(X411,'F_0') | p(s(bool,X412)) | p(s(bool,X413)) | ~sP11(X412,X413)) ) | $spl122),
  inference(resolution,[],[f18883,f454])).
fof(f19423,plain,(
  ( ! [X405,X407,X406] : (sP10('F_0',X405) | sP10(X405,'F_0') | p(s(bool,X406)) | p(s(bool,X407)) | ~sP11(X406,X407)) ) | $spl122),
  inference(resolution,[],[f18883,f458])).
fof(f19414,plain,(
  ( ! [X395,X396] : (sP10('T_0',X395) | sP10(X395,'F_0') | ~p(s(bool,X396))) ) | $spl122),
  inference(resolution,[],[f18883,f507])).
fof(f19413,plain,(
  ( ! [X393,X394] : (sP10('T_0',X393) | sP10(X393,'F_0') | ~p(s(bool,X394))) ) | $spl122),
  inference(resolution,[],[f18883,f505])).
fof(f19412,plain,(
  ( ! [X391,X392] : (sP10('T_0',X391) | sP10(X391,'F_0') | ~p(s(bool,X392))) ) | $spl122),
  inference(resolution,[],[f18883,f500])).
fof(f19411,plain,(
  ( ! [X389,X390] : (sP10('T_0',X389) | sP10(X389,'F_0') | ~p(s(bool,X390))) ) | $spl122),
  inference(resolution,[],[f18883,f498])).
fof(f19410,plain,(
  ( ! [X387,X388] : (sP10('T_0',X387) | sP10(X387,'F_0') | p(s(bool,X388))) ) | $spl122),
  inference(resolution,[],[f18883,f485])).
fof(f19409,plain,(
  ( ! [X385,X386] : (sP10('T_0',X385) | sP10(X385,'F_0') | ~p(s(bool,X386))) ) | $spl122),
  inference(resolution,[],[f18883,f483])).
fof(f19407,plain,(
  ( ! [X383,X382] : (sP10('T_0',X382) | sP10(X382,'F_0') | p(s(bool,X383))) ) | $spl122),
  inference(resolution,[],[f18883,f479])).
fof(f19406,plain,(
  ( ! [X380,X381] : (sP10('T_0',X380) | sP10(X380,'F_0') | ~p(s(bool,X381))) ) | $spl122),
  inference(resolution,[],[f18883,f477])).
fof(f19405,plain,(
  ( ! [X379,X378] : (sP10('T_0',X378) | sP10(X378,'F_0') | p(s(bool,X379))) ) | $spl122),
  inference(resolution,[],[f18883,f476])).
fof(f19471,plain,(
  ( ! [X0,X3,X1] : (sP10(X0,X1) | s(bool,X0) != s(bool,X3) | ~p(s(bool,X3)) | sP10(X1,'F_0')) ) | $spl138),
  inference(cnf_transformation,[],[f19471_D])).
fof(f19471_D,plain,(
  ( ! [X0,X3,X1] : (sP10(X0,X1) | s(bool,X0) != s(bool,X3) | ~p(s(bool,X3)) | sP10(X1,'F_0')) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f19399,plain,(
  ( ! [X356,X358,X357,X359] : (sP10(X356,X357) | sP10(X357,'F_0') | p(s(bool,X358)) | ~p(s(bool,X359)) | s(bool,X356) != s(bool,X359)) ) | $spl122),
  inference(resolution,[],[f18883,f536])).
fof(f19308,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP10(X35,X36) | sP10(X36,'F_0') | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X35) != s(bool,X38) | s(bool,X37) != s(bool,X39)) ) | $spl122),
  inference(resolution,[],[f18883,f312])).
fof(f19306,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | sP10(X26,'F_0') | p(s(bool,X27)) | ~p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X25) != s(bool,X28)) ) | $spl122),
  inference(resolution,[],[f18883,f311])).
fof(f19304,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP10(X15,X16) | sP10(X16,'F_0') | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X15) != s(bool,X18)) ) | $spl122),
  inference(resolution,[],[f18883,f310])).
fof(f19303,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP10(X10,X11) | sP10(X11,'F_0') | ~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X10) != s(bool,X13) | s(bool,X12) != s(bool,X14)) ) | $spl122),
  inference(resolution,[],[f18883,f308])).
fof(f19301,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10(X0,X1) | sP10(X1,'F_0') | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X0) != s(bool,X3)) ) | $spl122),
  inference(resolution,[],[f18883,f306])).
fof(f19233,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X13) | s(bool,'F_0') = s(bool,X13)) ) | ($spl1 | $spl8 | $spl46)),
  inference(resolution,[],[f14312,f16925])).
fof(f19231,plain,(
  ( ! [X8,X9] : (p(s(bool,X8)) | sP10(X9,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | s(bool,'F_0') = s(bool,X9)) ) | ($spl1 | $spl8 | $spl46)),
  inference(resolution,[],[f14312,f16700])).
fof(f18773,plain,(
  sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),'F_0') | $spl126),
  inference(cnf_transformation,[],[f18773_D])).
fof(f18773_D,plain,(
  sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),'F_0') <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f19142,plain,(
  ( ! [X17,X18] : (sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),'F_0') | p(s(bool,X17)) | ~p(s(bool,X18))) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f18873,f11976])).
fof(f18873,plain,(
  ( ! [X202,X203,X201] : (sP8(X201,X202,X203) | sP10(X201,'F_0')) ) | $spl122),
  inference(resolution,[],[f18767,f411])).
fof(f19132,plain,(
  ( ! [X15,X16] : (sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),'F_0') | p(s(bool,X15)) | ~p(s(bool,X16))) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f18872,f11977])).
fof(f18872,plain,(
  ( ! [X198,X200,X199] : (sP8(X199,X198,X200) | sP10(X198,'F_0')) ) | $spl122),
  inference(resolution,[],[f18767,f410])).
fof(f19126,plain,(
  ( ! [X17,X18] : (sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),'F_0') | p(s(bool,X17)) | p(s(bool,X18))) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f18859,f11960])).
fof(f18859,plain,(
  ( ! [X161,X159,X160] : (sP6(X159,X160,X161) | sP10(X159,'F_0')) ) | $spl122),
  inference(resolution,[],[f18767,f394])).
fof(f19125,plain,(
  ( ! [X15,X16] : (sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),'F_0') | sP7('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X15,X16)) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f18859,f12008])).
fof(f19119,plain,(
  ( ! [X19,X20] : (sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),'F_0') | ~p(s(bool,X19)) | p(s(bool,X20))) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f18843,f11946])).
fof(f18843,plain,(
  ( ! [X111,X112,X113] : (sP4(X111,X112,X113) | sP10(X111,'F_0')) ) | $spl122),
  inference(resolution,[],[f18767,f377])).
fof(f19118,plain,(
  ( ! [X17,X18] : (sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),'F_0') | sP5('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X17,X18)) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f18843,f12007])).
fof(f19108,plain,(
  ( ! [X15,X16] : (sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),'F_0') | p(s(bool,X15)) | ~p(s(bool,X16))) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f18842,f11947])).
fof(f18842,plain,(
  ( ! [X109,X110,X108] : (sP4(X109,X110,X108) | sP10(X108,'F_0')) ) | $spl122),
  inference(resolution,[],[f18767,f375])).
fof(f19099,plain,(
  ( ! [X21,X20] : (sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),'F_0') | p(s(bool,X20)) | ~p(s(bool,X21))) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f18820,f11923])).
fof(f18820,plain,(
  ( ! [X57,X58,X56] : (sP1(X56,X57,X58) | sP10(X56,'F_0')) ) | $spl122),
  inference(resolution,[],[f18767,f327])).
fof(f19088,plain,(
  ( ! [X19,X18] : (sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),'F_0') | p(s(bool,X18)) | ~p(s(bool,X19))) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f18819,f11924])).
fof(f18819,plain,(
  ( ! [X54,X55,X53] : (sP1(X54,X53,X55) | sP10(X53,'F_0')) ) | $spl122),
  inference(resolution,[],[f18767,f326])).
fof(f19078,plain,(
  ( ! [X28,X27] : (sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),'F_0') | p(s(bool,X27)) | ~p(s(bool,X28))) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f18817,f11920])).
fof(f18817,plain,(
  ( ! [X47,X48,X49] : (sP2(X47,X48,X49) | sP10(X47,'F_0')) ) | $spl122),
  inference(resolution,[],[f18767,f323])).
fof(f19063,plain,(
  ( ! [X26,X25] : (sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),'F_0') | p(s(bool,X25)) | ~p(s(bool,X26))) ) | ($spl8 | $spl122)),
  inference(resolution,[],[f18816,f11921])).
fof(f18816,plain,(
  ( ! [X45,X46,X44] : (sP2(X45,X44,X46) | sP10(X44,'F_0')) ) | $spl122),
  inference(resolution,[],[f18767,f322])).
fof(f19066,plain,(
  ( ! [X17] : (sP10(X17,'F_0') | s(bool,'T_0') != s(bool,X17)) ) | $spl134),
  inference(cnf_transformation,[],[f19066_D])).
fof(f19066_D,plain,(
  ( ! [X17] : (sP10(X17,'F_0') | s(bool,'T_0') != s(bool,X17)) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f19059,plain,(
  ( ! [X17,X18] : (sP10(X17,'F_0') | s(bool,'T_0') != s(bool,X17) | ~p(s(bool,X18))) ) | ($spl6 | $spl122)),
  inference(resolution,[],[f18816,f14550])).
fof(f19016,plain,(
  $false | $spl132),
  inference(subsumption_resolution,[],[f19011,f528])).
fof(f19011,plain,(
  p(s(bool,'F_0')) | $spl132),
  inference(duplicate_literal_removal,[],[f19006])).
fof(f19006,plain,(
  p(s(bool,'F_0')) | p(s(bool,'F_0')) | $spl132),
  inference(resolution,[],[f18986,f447])).
fof(f447,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | ~sP10(X1,X0)) )),
  inference(cnf_transformation,[],[f257])).
fof(f19015,plain,(
  $false | $spl132),
  inference(subsumption_resolution,[],[f19014,f528])).
fof(f19014,plain,(
  p(s(bool,'F_0')) | $spl132),
  inference(duplicate_literal_removal,[],[f19003])).
fof(f19003,plain,(
  p(s(bool,'F_0')) | p(s(bool,'F_0')) | $spl132),
  inference(resolution,[],[f18986,f438])).
fof(f18991,plain,(
  ( ! [X45] : (~p(s(bool,X45)) | sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),'F_0')) ) | ($spl8 | $spl46 | $spl122)),
  inference(subsumption_resolution,[],[f18963,f350])).
fof(f18963,plain,(
  ( ! [X45] : (~p(s(bool,X45)) | sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),'F_0') | s(bool,'F_0') = s(bool,X45)) ) | ($spl8 | $spl46 | $spl122)),
  inference(superposition,[],[f18767,f15318])).
fof(f18981,plain,(
  sP10('T_0','F_0') | $spl130),
  inference(cnf_transformation,[],[f18981_D])).
fof(f18981_D,plain,(
  sP10('T_0','F_0') <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f18968,plain,(
  ( ! [X40] : (~p(s(bool,X40)) | sP10('T_0','F_0')) ) | $spl122),
  inference(duplicate_literal_removal,[],[f18957])).
fof(f18957,plain,(
  ( ! [X40] : (~p(s(bool,X40)) | sP10('T_0','F_0') | ~p(s(bool,X40))) ) | $spl122),
  inference(superposition,[],[f18767,f355])).
fof(f18969,plain,(
  ( ! [X39] : (~p(s(bool,X39)) | sP10('T_0','F_0')) ) | $spl122),
  inference(duplicate_literal_removal,[],[f18956])).
fof(f18956,plain,(
  ( ! [X39] : (~p(s(bool,X39)) | sP10('T_0','F_0') | ~p(s(bool,X39))) ) | $spl122),
  inference(superposition,[],[f18767,f357])).
fof(f18990,plain,(
  ( ! [X38] : (~p(s(bool,X38)) | sP10('T_0','F_0')) ) | $spl122),
  inference(subsumption_resolution,[],[f18955,f350])).
fof(f18955,plain,(
  ( ! [X38] : (~p(s(bool,X38)) | sP10('T_0','F_0') | s(bool,'F_0') = s(bool,X38)) ) | $spl122),
  inference(superposition,[],[f18767,f358])).
fof(f18989,plain,(
  ( ! [X37] : (~p(s(bool,X37)) | sP10('T_0','F_0')) ) | $spl122),
  inference(subsumption_resolution,[],[f18954,f350])).
fof(f18954,plain,(
  ( ! [X37] : (~p(s(bool,X37)) | sP10('T_0','F_0') | s(bool,'F_0') = s(bool,X37)) ) | $spl122),
  inference(superposition,[],[f18767,f359])).
fof(f18986,plain,(
  sP10('F_0','F_0') | $spl132),
  inference(cnf_transformation,[],[f18986_D])).
fof(f18986_D,plain,(
  sP10('F_0','F_0') <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f18928,plain,(
  ( ! [X288,X287] : (sP10('F_0','F_0') | p(s(bool,X287)) | p(s(bool,X288)) | ~sP11(X287,X288)) ) | $spl122),
  inference(resolution,[],[f18767,f454])).
fof(f18926,plain,(
  ( ! [X283,X284] : (sP10('F_0','F_0') | p(s(bool,X283)) | p(s(bool,X284)) | ~sP11(X283,X284)) ) | $spl122),
  inference(resolution,[],[f18767,f458])).
fof(f18917,plain,(
  ( ! [X282] : (sP10('T_0','F_0') | ~p(s(bool,X282))) ) | $spl122),
  inference(resolution,[],[f18767,f507])).
fof(f18916,plain,(
  ( ! [X281] : (sP10('T_0','F_0') | ~p(s(bool,X281))) ) | $spl122),
  inference(resolution,[],[f18767,f505])).
fof(f18915,plain,(
  ( ! [X280] : (sP10('T_0','F_0') | ~p(s(bool,X280))) ) | $spl122),
  inference(resolution,[],[f18767,f500])).
fof(f18914,plain,(
  ( ! [X279] : (sP10('T_0','F_0') | ~p(s(bool,X279))) ) | $spl122),
  inference(resolution,[],[f18767,f498])).
fof(f18913,plain,(
  ( ! [X278] : (sP10('T_0','F_0') | p(s(bool,X278))) ) | $spl122),
  inference(resolution,[],[f18767,f485])).
fof(f18912,plain,(
  ( ! [X277] : (sP10('T_0','F_0') | ~p(s(bool,X277))) ) | $spl122),
  inference(resolution,[],[f18767,f483])).
fof(f18910,plain,(
  ( ! [X276] : (sP10('T_0','F_0') | p(s(bool,X276))) ) | $spl122),
  inference(resolution,[],[f18767,f479])).
fof(f18909,plain,(
  ( ! [X275] : (sP10('T_0','F_0') | ~p(s(bool,X275))) ) | $spl122),
  inference(resolution,[],[f18767,f477])).
fof(f18908,plain,(
  ( ! [X274] : (sP10('T_0','F_0') | p(s(bool,X274))) ) | $spl122),
  inference(resolution,[],[f18767,f476])).
fof(f18974,plain,(
  ( ! [X2,X0] : (sP10(X0,'F_0') | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) | $spl128),
  inference(cnf_transformation,[],[f18974_D])).
fof(f18974_D,plain,(
  ( ! [X2,X0] : (sP10(X0,'F_0') | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f18902,plain,(
  ( ! [X259,X260,X258] : (sP10(X258,'F_0') | p(s(bool,X259)) | ~p(s(bool,X260)) | s(bool,X258) != s(bool,X260)) ) | $spl122),
  inference(resolution,[],[f18767,f536])).
fof(f18811,plain,(
  ( ! [X30,X28,X31,X29] : (sP10(X28,'F_0') | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) ) | $spl122),
  inference(resolution,[],[f18767,f312])).
fof(f18809,plain,(
  ( ! [X23,X21,X22,X20] : (sP10(X20,'F_0') | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X20) != s(bool,X22)) ) | $spl122),
  inference(resolution,[],[f18767,f311])).
fof(f18807,plain,(
  ( ! [X14,X12,X15,X13] : (sP10(X12,'F_0') | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl122),
  inference(resolution,[],[f18767,f310])).
fof(f18806,plain,(
  ( ! [X10,X8,X11,X9] : (sP10(X8,'F_0') | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) ) | $spl122),
  inference(resolution,[],[f18767,f308])).
fof(f18804,plain,(
  ( ! [X2,X0,X3,X1] : (sP10(X0,'F_0') | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl122),
  inference(resolution,[],[f18767,f306])).
fof(f18791,plain,(
  ( ! [X8,X7] : (~p(s(bool,X7)) | s(bool,'F_0') = s(bool,X8) | sP10(X8,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) ) | ($spl1 | $spl8 | $spl46)),
  inference(duplicate_literal_removal,[],[f18787])).
fof(f18787,plain,(
  ( ! [X8,X7] : (~p(s(bool,X7)) | s(bool,'F_0') = s(bool,X8) | ~p(s(bool,X7)) | sP10(X8,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) ) | ($spl1 | $spl8 | $spl46)),
  inference(resolution,[],[f14265,f16692])).
fof(f18771,plain,(
  ( ! [X18] : (sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),'F_0') | ~p(s(bool,X18))) ) | ($spl8 | $spl46)),
  inference(subsumption_resolution,[],[f18750,f350])).
fof(f18750,plain,(
  ( ! [X18] : (sP10('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),'F_0') | s(bool,'F_0') = s(bool,X18) | ~p(s(bool,X18))) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f18663,f6469])).
fof(f6469,plain,(
  ( ! [X151] : (~sP10(X151,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | ~p(s(bool,X151))) ) | $spl8),
  inference(resolution,[],[f2511,f439])).
fof(f18663,plain,(
  ( ! [X505,X506] : (sP10(X505,X506) | sP10(X506,'F_0') | s(bool,'F_0') = s(bool,X505)) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16694,f528])).
fof(f18770,plain,(
  ( ! [X14,X13] : (s(bool,'F_0') = s(bool,X14) | sP10(X14,X13)) ) | $spl124),
  inference(cnf_transformation,[],[f18770_D])).
fof(f18770_D,plain,(
  ( ! [X14,X13] : (s(bool,'F_0') = s(bool,X14) | sP10(X14,X13)) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f18767,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | sP10(X2,'F_0')) ) | $spl122),
  inference(cnf_transformation,[],[f18767_D])).
fof(f18767_D,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | sP10(X2,'F_0')) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f18769,plain,(
  ( ! [X14,X12,X13] : (sP10(X12,'F_0') | s(bool,'F_0') = s(bool,X14) | ~p(s(bool,X12)) | sP10(X14,X13)) ) | ($spl8 | $spl46)),
  inference(subsumption_resolution,[],[f18748,f16643])).
fof(f16643,plain,(
  ( ! [X114,X115] : (sP10(X114,X115) | s(bool,'F_0') = s(bool,X114) | s(bool,'F_0') != s(bool,X115)) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16255,f350])).
fof(f18748,plain,(
  ( ! [X14,X12,X13] : (sP10(X12,'F_0') | s(bool,'F_0') = s(bool,X13) | s(bool,'F_0') = s(bool,X14) | ~p(s(bool,X12)) | sP10(X14,X13)) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f18663,f16692])).
fof(f6288,plain,(
  ( ! [X37] : (s(bool,'F_0') = s(bool,X37)) ) | $spl62),
  inference(cnf_transformation,[],[f6288_D])).
fof(f6288_D,plain,(
  ( ! [X37] : (s(bool,'F_0') = s(bool,X37)) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f18765,plain,(
  ( ! [X8,X9] : (sP10(X8,'F_0') | s(bool,'F_0') = s(bool,X9) | ~p(s(bool,X8))) ) | ($spl8 | $spl46)),
  inference(duplicate_literal_removal,[],[f18746])).
fof(f18746,plain,(
  ( ! [X8,X9] : (sP10(X8,'F_0') | s(bool,'F_0') = s(bool,X9) | ~p(s(bool,X8)) | s(bool,'F_0') = s(bool,X9)) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f18663,f16250])).
fof(f18768,plain,(
  ( ! [X4,X5] : (sP10(X4,'F_0') | ~p(s(bool,X4)) | ~p(s(bool,X5))) ) | ($spl8 | $spl46)),
  inference(subsumption_resolution,[],[f18744,f350])).
fof(f18744,plain,(
  ( ! [X4,X5] : (sP10(X4,'F_0') | s(bool,'F_0') = s(bool,X5) | ~p(s(bool,X4)) | ~p(s(bool,X5))) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f18663,f446])).
fof(f18766,plain,(
  ( ! [X2,X3] : (sP10(X2,'F_0') | ~p(s(bool,X3)) | ~p(s(bool,X2))) ) | ($spl8 | $spl46)),
  inference(subsumption_resolution,[],[f18743,f350])).
fof(f18743,plain,(
  ( ! [X2,X3] : (sP10(X2,'F_0') | s(bool,'F_0') = s(bool,X3) | ~p(s(bool,X3)) | ~p(s(bool,X2))) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f18663,f439])).
fof(f18730,plain,(
  ( ! [X474,X473] : (sP10(X473,X474) | s(bool,'F_0') = s(bool,X473) | sP10(X474,'F_0')) ) | $spl120),
  inference(cnf_transformation,[],[f18730_D])).
fof(f18730_D,plain,(
  ( ! [X474,X473] : (sP10(X473,X474) | s(bool,'F_0') = s(bool,X473) | sP10(X474,'F_0')) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f18708,plain,(
  ( ! [X101,X102,X100] : (p(s(bool,X100)) | sP10(X101,X102) | sP10(X102,'F_0') | s(bool,'F_0') = s(bool,X101)) ) | ($spl8 | $spl46)),
  inference(duplicate_literal_removal,[],[f18701])).
fof(f18701,plain,(
  ( ! [X101,X102,X100] : (p(s(bool,X100)) | sP10(X101,X102) | sP10(X102,'F_0') | s(bool,'F_0') = s(bool,X101) | p(s(bool,X100))) ) | ($spl8 | $spl46)),
  inference(superposition,[],[f16694,f351])).
fof(f18709,plain,(
  ( ! [X99,X97,X98] : (p(s(bool,X97)) | sP10(X98,X99) | sP10(X99,'F_0') | s(bool,'F_0') = s(bool,X98)) ) | ($spl8 | $spl46)),
  inference(duplicate_literal_removal,[],[f18700])).
fof(f18700,plain,(
  ( ! [X99,X97,X98] : (p(s(bool,X97)) | sP10(X98,X99) | sP10(X99,'F_0') | s(bool,'F_0') = s(bool,X98) | p(s(bool,X97))) ) | ($spl8 | $spl46)),
  inference(superposition,[],[f16694,f353])).
fof(f18740,plain,(
  ( ! [X94,X95,X96] : (p(s(bool,X94)) | sP10(X95,X96) | sP10(X96,'F_0') | s(bool,'F_0') = s(bool,X95)) ) | ($spl8 | $spl46)),
  inference(subsumption_resolution,[],[f18699,f354])).
fof(f354,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f215])).
fof(f18699,plain,(
  ( ! [X94,X95,X96] : (p(s(bool,X94)) | sP10(X95,X96) | sP10(X96,'F_0') | s(bool,'F_0') = s(bool,X95) | s(bool,'T_0') = s(bool,X94)) ) | ($spl8 | $spl46)),
  inference(superposition,[],[f16694,f358])).
fof(f18739,plain,(
  ( ! [X92,X93,X91] : (p(s(bool,X91)) | sP10(X92,X93) | sP10(X93,'F_0') | s(bool,'F_0') = s(bool,X92)) ) | ($spl8 | $spl46)),
  inference(subsumption_resolution,[],[f18698,f354])).
fof(f18698,plain,(
  ( ! [X92,X93,X91] : (p(s(bool,X91)) | sP10(X92,X93) | sP10(X93,'F_0') | s(bool,'F_0') = s(bool,X92) | s(bool,'T_0') = s(bool,X91)) ) | ($spl8 | $spl46)),
  inference(superposition,[],[f16694,f359])).
fof(f18662,plain,(
  ( ! [X502,X504,X503] : (sP10(X502,X503) | sP10(X503,'F_0') | s(bool,'F_0') = s(bool,X502) | p(s(bool,X504))) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16694,f526])).
fof(f526,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f173])).
fof(f173,plain,(
  ! [X0] : (~p(s(bool,'F_0')) | p(s(bool,X0)))),
  inference(ennf_transformation,[],[f136])).
fof(f136,plain,(
  ! [X0] : (p(s(bool,'F_0')) => p(s(bool,X0)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X11] : (p(s(bool,'F_0')) => p(s(bool,X11)))),
  file('thm.list.MEM_GENLIST','thm.bool.FALSITY')).
fof(f18661,plain,(
  ( ! [X499,X500,X501] : (sP10(X499,X500) | sP10(X500,'F_0') | s(bool,'F_0') = s(bool,X499) | ~p(s(bool,X501))) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16694,f524])).
fof(f18660,plain,(
  ( ! [X496,X498,X497] : (sP10(X496,X497) | sP10(X497,'F_0') | s(bool,'F_0') = s(bool,X496) | p(s(bool,X498))) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16694,f516])).
fof(f516,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f299])).
fof(f299,plain,(
  ! [X0] : ((~p(s(bool,X0)) | ~p(s(bool,'F_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,X0)) & p(s(bool,'F_0')))))),
  inference(flattening,[],[f298])).
fof(f298,plain,(
  ! [X0] : (((~p(s(bool,X0)) | ~p(s(bool,'F_0'))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,X0)) & p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f127])).
fof(f127,plain,(
  ! [X0] : ((p(s(bool,X0)) & p(s(bool,'F_0'))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X11] : ((p(s(bool,X11)) & p(s(bool,'F_0'))) <=> p(s(bool,'F_0')))),
  file('thm.list.MEM_GENLIST','thm.bool.AND_CLAUSES|split|3')).
fof(f18659,plain,(
  ( ! [X494,X493,X495] : (sP10(X493,X494) | sP10(X494,'F_0') | s(bool,'F_0') = s(bool,X493) | p(s(bool,X495))) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16694,f514])).
fof(f514,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f297])).
fof(f297,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | ~p(s(bool,X0)) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,'F_0')) & p(s(bool,X0)))))),
  inference(flattening,[],[f296])).
fof(f296,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) | ~p(s(bool,X0))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,'F_0')) & p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0] : ((p(s(bool,'F_0')) & p(s(bool,X0))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X11] : ((p(s(bool,'F_0')) & p(s(bool,X11))) <=> p(s(bool,'F_0')))),
  file('thm.list.MEM_GENLIST','thm.bool.AND_CLAUSES|split|2')).
fof(f18658,plain,(
  ( ! [X492,X490,X491] : (sP10(X490,X491) | sP10(X491,'F_0') | s(bool,'F_0') = s(bool,X490) | p(s(bool,X492))) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16694,f494])).
fof(f494,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f285])).
fof(f285,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f284])).
fof(f284,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0] : ((p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X11] : ((p(s(bool,'F_0')) | p(s(bool,X11))) <=> p(s(bool,X11)))),
  file('thm.list.MEM_GENLIST','thm.bool.OR_CLAUSES|split|2')).
fof(f18657,plain,(
  ( ! [X488,X487,X489] : (sP10(X487,X488) | sP10(X488,'F_0') | s(bool,'F_0') = s(bool,X487) | p(s(bool,X489))) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16694,f492])).
fof(f18731,plain,(
  ( ! [X484,X486,X485] : (sP10(X484,X485) | sP10(X485,'F_0') | s(bool,'F_0') = s(bool,X484) | p(s(bool,X486))) ) | ($spl8 | $spl46)),
  inference(subsumption_resolution,[],[f18656,f537])).
fof(f18656,plain,(
  ( ! [X484,X486,X485] : (sP10(X484,X485) | sP10(X485,'F_0') | s(bool,'F_0') = s(bool,X484) | p(s(bool,X486)) | ~p(s(bool,'T_0'))) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16694,f484])).
fof(f484,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f277])).
fof(f18655,plain,(
  ( ! [X482,X481,X483] : (sP10(X481,X482) | sP10(X482,'F_0') | s(bool,'F_0') = s(bool,X481) | ~p(s(bool,X483))) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16694,f474])).
fof(f18654,plain,(
  ( ! [X478,X480,X477,X479] : (sP10(X477,X478) | sP10(X478,'F_0') | s(bool,'F_0') = s(bool,X477) | sP11(X479,X480)) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16694,f462])).
fof(f462,plain,(
  ( ! [X0,X1] : (sP11(X0,X1) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f263])).
fof(f18653,plain,(
  ( ! [X476,X474,X475,X473] : (sP10(X473,X474) | sP10(X474,'F_0') | s(bool,'F_0') = s(bool,X473) | sP11(X475,X476)) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16694,f457])).
fof(f457,plain,(
  ( ! [X0,X1] : (sP11(X0,X1) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f261])).
fof(f18727,plain,(
  ( ! [X4,X2,X0,X1] : (sP10(X0,X1) | s(bool,X2) != s(bool,X4) | p(s(bool,X4)) | s(bool,'F_0') = s(bool,X0) | sP10(X1,X2)) ) | $spl118),
  inference(cnf_transformation,[],[f18727_D])).
fof(f18727_D,plain,(
  ( ! [X4,X2,X0,X1] : (sP10(X0,X1) | s(bool,X2) != s(bool,X4) | p(s(bool,X4)) | s(bool,'F_0') = s(bool,X0) | sP10(X1,X2)) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f18647,plain,(
  ( ! [X453,X455,X451,X452,X454] : (sP10(X451,X452) | sP10(X452,X453) | s(bool,'F_0') = s(bool,X451) | p(s(bool,X454)) | p(s(bool,X455)) | s(bool,X453) != s(bool,X455)) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16694,f536])).
fof(f18557,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (sP10(X42,X43) | sP10(X43,X44) | s(bool,'F_0') = s(bool,X42) | p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X44) != s(bool,X45) | s(bool,X46) != s(bool,X47)) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16694,f312])).
fof(f18554,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (sP10(X24,X25) | sP10(X25,X26) | s(bool,'F_0') = s(bool,X24) | p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X26) != s(bool,X29)) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16694,f311])).
fof(f18553,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (sP10(X18,X19) | sP10(X19,X20) | s(bool,'F_0') = s(bool,X18) | p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X20) != s(bool,X22)) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16694,f310])).
fof(f18552,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (sP10(X12,X13) | sP10(X13,X14) | s(bool,'F_0') = s(bool,X12) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X14) != s(bool,X15) | s(bool,X16) != s(bool,X17)) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16694,f308])).
fof(f18550,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP10(X0,X1) | sP10(X1,X2) | s(bool,'F_0') = s(bool,X0) | p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X2) != s(bool,X4)) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16694,f306])).
fof(f18714,plain,(
  ( ! [X33,X32] : (sP10(X32,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | p(s(bool,X33)) | s(bool,'F_0') = s(bool,X32)) ) | ($spl1 | $spl8 | $spl46)),
  inference(duplicate_literal_removal,[],[f18549])).
fof(f18549,plain,(
  ( ! [X33,X32] : (sP10(X32,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | p(s(bool,X33)) | s(bool,'F_0') = s(bool,X32) | p(s(bool,X33))) ) | ($spl1 | $spl8 | $spl46)),
  inference(resolution,[],[f16694,f14263])).
fof(f18521,plain,(
  ( ! [X132,X133] : (s(bool,'F_0') = s(bool,X132) | ~p(s(bool,X133)) | sP10(X132,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) ) | ($spl1 | $spl8 | $spl46)),
  inference(duplicate_literal_removal,[],[f18517])).
fof(f18517,plain,(
  ( ! [X132,X133] : (s(bool,'F_0') = s(bool,X132) | ~p(s(bool,X133)) | sP10(X132,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | ~p(s(bool,X133))) ) | ($spl1 | $spl8 | $spl46)),
  inference(resolution,[],[f16693,f14264])).
fof(f17149,plain,(
  ( ! [X6,X7,X5] : (sP3(X6,X7,X5) | ~sP1(X7,X5,X6) | ~sP2(X6,X5,X7) | ~p(s(bool,X5))) ) | $spl114),
  inference(cnf_transformation,[],[f17149_D])).
fof(f17149_D,plain,(
  ( ! [X6,X7,X5] : (sP3(X6,X7,X5) | ~sP1(X7,X5,X6) | ~sP2(X6,X5,X7) | ~p(s(bool,X5))) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f17312,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'F_0') = s(bool,X5) | sP3(X6,X7,X4) | ~p(s(bool,X4)) | ~sP2(X6,X4,X7) | ~sP1(X7,X4,X6)) ) | ($spl8 | $spl46)),
  inference(subsumption_resolution,[],[f17304,f16252])).
fof(f17304,plain,(
  ( ! [X6,X4,X7,X5] : (sP10(X4,X5) | s(bool,'F_0') = s(bool,X5) | sP3(X6,X7,X4) | ~p(s(bool,X4)) | ~sP2(X6,X4,X7) | ~sP1(X7,X4,X6)) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16859,f317])).
fof(f317,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X2)) | ~sP2(X0,X2,X1) | ~sP1(X1,X2,X0) | ~sP0(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f199])).
fof(f17148,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'F_0') = s(bool,X4) | sP3(X6,X7,X5) | ~p(s(bool,X5)) | ~sP2(X6,X5,X7) | ~sP1(X7,X5,X6)) ) | ($spl8 | $spl46)),
  inference(subsumption_resolution,[],[f17140,f16250])).
fof(f17140,plain,(
  ( ! [X6,X4,X7,X5] : (sP10(X4,X5) | s(bool,'F_0') = s(bool,X4) | sP3(X6,X7,X5) | ~p(s(bool,X5)) | ~sP2(X6,X5,X7) | ~sP1(X7,X5,X6)) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16634,f317])).
fof(f17011,plain,(
  ( ! [X370] : (sP10('F_0',X370) | s(bool,'F_0') = s(bool,X370)) ) | $spl112),
  inference(cnf_transformation,[],[f17011_D])).
fof(f17011_D,plain,(
  ( ! [X370] : (sP10('F_0',X370) | s(bool,'F_0') = s(bool,X370)) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f16997,plain,(
  ( ! [X72,X73] : (p(s(bool,X72)) | sP10('F_0',X73) | s(bool,'F_0') = s(bool,X73)) ) | ($spl8 | $spl46)),
  inference(duplicate_literal_removal,[],[f16991])).
fof(f16991,plain,(
  ( ! [X72,X73] : (p(s(bool,X72)) | sP10('F_0',X73) | s(bool,'F_0') = s(bool,X73) | p(s(bool,X72))) ) | ($spl8 | $spl46)),
  inference(superposition,[],[f16258,f351])).
fof(f16998,plain,(
  ( ! [X70,X71] : (p(s(bool,X70)) | sP10('F_0',X71) | s(bool,'F_0') = s(bool,X71)) ) | ($spl8 | $spl46)),
  inference(duplicate_literal_removal,[],[f16990])).
fof(f16990,plain,(
  ( ! [X70,X71] : (p(s(bool,X70)) | sP10('F_0',X71) | s(bool,'F_0') = s(bool,X71) | p(s(bool,X70))) ) | ($spl8 | $spl46)),
  inference(superposition,[],[f16258,f353])).
fof(f17021,plain,(
  ( ! [X68,X69] : (p(s(bool,X68)) | sP10('F_0',X69) | s(bool,'F_0') = s(bool,X69)) ) | ($spl8 | $spl46)),
  inference(subsumption_resolution,[],[f16989,f354])).
fof(f16989,plain,(
  ( ! [X68,X69] : (p(s(bool,X68)) | sP10('F_0',X69) | s(bool,'F_0') = s(bool,X69) | s(bool,'T_0') = s(bool,X68)) ) | ($spl8 | $spl46)),
  inference(superposition,[],[f16258,f358])).
fof(f17020,plain,(
  ( ! [X66,X67] : (p(s(bool,X66)) | sP10('F_0',X67) | s(bool,'F_0') = s(bool,X67)) ) | ($spl8 | $spl46)),
  inference(subsumption_resolution,[],[f16988,f354])).
fof(f16988,plain,(
  ( ! [X66,X67] : (p(s(bool,X66)) | sP10('F_0',X67) | s(bool,'F_0') = s(bool,X67) | s(bool,'T_0') = s(bool,X66)) ) | ($spl8 | $spl46)),
  inference(superposition,[],[f16258,f359])).
fof(f16952,plain,(
  ( ! [X391,X390] : (sP10('F_0',X390) | s(bool,'F_0') = s(bool,X390) | p(s(bool,X391))) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16258,f526])).
fof(f16951,plain,(
  ( ! [X389,X388] : (sP10('F_0',X388) | s(bool,'F_0') = s(bool,X388) | ~p(s(bool,X389))) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16258,f524])).
fof(f16950,plain,(
  ( ! [X387,X386] : (sP10('F_0',X386) | s(bool,'F_0') = s(bool,X386) | p(s(bool,X387))) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16258,f516])).
fof(f16949,plain,(
  ( ! [X385,X384] : (sP10('F_0',X384) | s(bool,'F_0') = s(bool,X384) | p(s(bool,X385))) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16258,f514])).
fof(f16948,plain,(
  ( ! [X383,X382] : (sP10('F_0',X382) | s(bool,'F_0') = s(bool,X382) | p(s(bool,X383))) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16258,f494])).
fof(f16947,plain,(
  ( ! [X380,X381] : (sP10('F_0',X380) | s(bool,'F_0') = s(bool,X380) | p(s(bool,X381))) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16258,f492])).
fof(f17012,plain,(
  ( ! [X379,X378] : (sP10('F_0',X378) | s(bool,'F_0') = s(bool,X378) | p(s(bool,X379))) ) | ($spl8 | $spl46)),
  inference(subsumption_resolution,[],[f16946,f537])).
fof(f16946,plain,(
  ( ! [X379,X378] : (sP10('F_0',X378) | s(bool,'F_0') = s(bool,X378) | p(s(bool,X379)) | ~p(s(bool,'T_0'))) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16258,f484])).
fof(f16945,plain,(
  ( ! [X377,X376] : (sP10('F_0',X376) | s(bool,'F_0') = s(bool,X376) | ~p(s(bool,X377))) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16258,f474])).
fof(f16944,plain,(
  ( ! [X374,X373,X375] : (sP10('F_0',X373) | s(bool,'F_0') = s(bool,X373) | sP11(X374,X375)) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16258,f462])).
fof(f16943,plain,(
  ( ! [X372,X370,X371] : (sP10('F_0',X370) | s(bool,'F_0') = s(bool,X370) | sP11(X371,X372)) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16258,f457])).
fof(f17008,plain,(
  ( ! [X0,X3,X1] : (sP10(X0,X1) | s(bool,X0) != s(bool,X3) | p(s(bool,X3)) | s(bool,'F_0') = s(bool,X1)) ) | $spl110),
  inference(cnf_transformation,[],[f17008_D])).
fof(f17008_D,plain,(
  ( ! [X0,X3,X1] : (sP10(X0,X1) | s(bool,X0) != s(bool,X3) | p(s(bool,X3)) | s(bool,'F_0') = s(bool,X1)) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f16937,plain,(
  ( ! [X356,X354,X357,X355] : (sP10(X354,X355) | s(bool,'F_0') = s(bool,X355) | p(s(bool,X356)) | p(s(bool,X357)) | s(bool,X354) != s(bool,X357)) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16258,f536])).
fof(f16847,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP10(X35,X36) | s(bool,'F_0') = s(bool,X36) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X35) != s(bool,X37) | s(bool,X38) != s(bool,X39)) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16258,f312])).
fof(f16844,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | s(bool,'F_0') = s(bool,X21) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X20) != s(bool,X24)) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16258,f311])).
fof(f16843,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP10(X15,X16) | s(bool,'F_0') = s(bool,X16) | p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X15) != s(bool,X18)) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16258,f310])).
fof(f16842,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP10(X10,X11) | s(bool,'F_0') = s(bool,X11) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X10) != s(bool,X12) | s(bool,X13) != s(bool,X14)) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16258,f308])).
fof(f16840,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10(X0,X1) | s(bool,'F_0') = s(bool,X1) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X0) != s(bool,X3)) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16258,f306])).
fof(f16787,plain,(
  ( ! [X370] : (sP10(X370,'F_0') | s(bool,'F_0') = s(bool,X370)) ) | $spl108),
  inference(cnf_transformation,[],[f16787_D])).
fof(f16787_D,plain,(
  ( ! [X370] : (sP10(X370,'F_0') | s(bool,'F_0') = s(bool,X370)) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f16772,plain,(
  ( ! [X72,X73] : (p(s(bool,X72)) | sP10(X73,'F_0') | s(bool,'F_0') = s(bool,X73)) ) | ($spl8 | $spl46)),
  inference(duplicate_literal_removal,[],[f16766])).
fof(f16766,plain,(
  ( ! [X72,X73] : (p(s(bool,X72)) | sP10(X73,'F_0') | s(bool,'F_0') = s(bool,X73) | p(s(bool,X72))) ) | ($spl8 | $spl46)),
  inference(superposition,[],[f16255,f351])).
fof(f16773,plain,(
  ( ! [X70,X71] : (p(s(bool,X70)) | sP10(X71,'F_0') | s(bool,'F_0') = s(bool,X71)) ) | ($spl8 | $spl46)),
  inference(duplicate_literal_removal,[],[f16765])).
fof(f16765,plain,(
  ( ! [X70,X71] : (p(s(bool,X70)) | sP10(X71,'F_0') | s(bool,'F_0') = s(bool,X71) | p(s(bool,X70))) ) | ($spl8 | $spl46)),
  inference(superposition,[],[f16255,f353])).
fof(f16797,plain,(
  ( ! [X68,X69] : (p(s(bool,X68)) | sP10(X69,'F_0') | s(bool,'F_0') = s(bool,X69)) ) | ($spl8 | $spl46)),
  inference(subsumption_resolution,[],[f16764,f354])).
fof(f16764,plain,(
  ( ! [X68,X69] : (p(s(bool,X68)) | sP10(X69,'F_0') | s(bool,'F_0') = s(bool,X69) | s(bool,'T_0') = s(bool,X68)) ) | ($spl8 | $spl46)),
  inference(superposition,[],[f16255,f358])).
fof(f16796,plain,(
  ( ! [X66,X67] : (p(s(bool,X66)) | sP10(X67,'F_0') | s(bool,'F_0') = s(bool,X67)) ) | ($spl8 | $spl46)),
  inference(subsumption_resolution,[],[f16763,f354])).
fof(f16763,plain,(
  ( ! [X66,X67] : (p(s(bool,X66)) | sP10(X67,'F_0') | s(bool,'F_0') = s(bool,X67) | s(bool,'T_0') = s(bool,X66)) ) | ($spl8 | $spl46)),
  inference(superposition,[],[f16255,f359])).
fof(f16727,plain,(
  ( ! [X391,X390] : (sP10(X390,'F_0') | s(bool,'F_0') = s(bool,X390) | p(s(bool,X391))) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16255,f526])).
fof(f16726,plain,(
  ( ! [X389,X388] : (sP10(X388,'F_0') | s(bool,'F_0') = s(bool,X388) | ~p(s(bool,X389))) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16255,f524])).
fof(f16725,plain,(
  ( ! [X387,X386] : (sP10(X386,'F_0') | s(bool,'F_0') = s(bool,X386) | p(s(bool,X387))) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16255,f516])).
fof(f16724,plain,(
  ( ! [X385,X384] : (sP10(X384,'F_0') | s(bool,'F_0') = s(bool,X384) | p(s(bool,X385))) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16255,f514])).
fof(f16723,plain,(
  ( ! [X383,X382] : (sP10(X382,'F_0') | s(bool,'F_0') = s(bool,X382) | p(s(bool,X383))) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16255,f494])).
fof(f16722,plain,(
  ( ! [X380,X381] : (sP10(X380,'F_0') | s(bool,'F_0') = s(bool,X380) | p(s(bool,X381))) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16255,f492])).
fof(f16788,plain,(
  ( ! [X379,X378] : (sP10(X378,'F_0') | s(bool,'F_0') = s(bool,X378) | p(s(bool,X379))) ) | ($spl8 | $spl46)),
  inference(subsumption_resolution,[],[f16721,f537])).
fof(f16721,plain,(
  ( ! [X379,X378] : (sP10(X378,'F_0') | s(bool,'F_0') = s(bool,X378) | p(s(bool,X379)) | ~p(s(bool,'T_0'))) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16255,f484])).
fof(f16720,plain,(
  ( ! [X377,X376] : (sP10(X376,'F_0') | s(bool,'F_0') = s(bool,X376) | ~p(s(bool,X377))) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16255,f474])).
fof(f16719,plain,(
  ( ! [X374,X373,X375] : (sP10(X373,'F_0') | s(bool,'F_0') = s(bool,X373) | sP11(X374,X375)) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16255,f462])).
fof(f16718,plain,(
  ( ! [X372,X370,X371] : (sP10(X370,'F_0') | s(bool,'F_0') = s(bool,X370) | sP11(X371,X372)) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16255,f457])).
fof(f16783,plain,(
  ( ! [X0,X3,X1] : (sP10(X0,X1) | s(bool,X1) != s(bool,X3) | p(s(bool,X3)) | s(bool,'F_0') = s(bool,X0)) ) | $spl104),
  inference(cnf_transformation,[],[f16783_D])).
fof(f16783_D,plain,(
  ( ! [X0,X3,X1] : (sP10(X0,X1) | s(bool,X1) != s(bool,X3) | p(s(bool,X3)) | s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f16712,plain,(
  ( ! [X356,X354,X357,X355] : (sP10(X354,X355) | s(bool,'F_0') = s(bool,X354) | p(s(bool,X356)) | p(s(bool,X357)) | s(bool,X355) != s(bool,X357)) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16255,f536])).
fof(f16622,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP10(X35,X36) | s(bool,'F_0') = s(bool,X35) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X37) | s(bool,X38) != s(bool,X39)) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16255,f312])).
fof(f16619,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,X21) | s(bool,'F_0') = s(bool,X20) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X21) != s(bool,X24)) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16255,f311])).
fof(f16618,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP10(X15,X16) | s(bool,'F_0') = s(bool,X15) | p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16255,f310])).
fof(f16617,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP10(X10,X11) | s(bool,'F_0') = s(bool,X10) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X12) | s(bool,X13) != s(bool,X14)) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16255,f308])).
fof(f16615,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10(X0,X1) | s(bool,'F_0') = s(bool,X0) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) ) | ($spl8 | $spl46)),
  inference(resolution,[],[f16255,f306])).
fof(f16494,plain,(
  ( ! [X652,X653,X651] : (s(bool,X651) != s(bool,X652) | p(s(bool,X653)) | p(s(bool,X652)) | s(bool,'F_0') = s(bool,X651)) ) | ($spl8 | $spl46)),
  inference(subsumption_resolution,[],[f16493,f537])).
fof(f16493,plain,(
  ( ! [X652,X653,X651] : (~p(s(bool,'T_0')) | s(bool,X651) != s(bool,X652) | p(s(bool,X653)) | p(s(bool,X652)) | s(bool,'F_0') = s(bool,X651)) ) | ($spl8 | $spl46)),
  inference(forward_demodulation,[],[f15931,f2535])).
fof(f15931,plain,(
  ( ! [X652,X653,X651] : (s(bool,X651) != s(bool,X652) | p(s(bool,X653)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | p(s(bool,X652)) | s(bool,'F_0') = s(bool,X651)) ) | ($spl8 | $spl46)),
  inference(superposition,[],[f536,f15318])).
fof(f16460,plain,(
  ( ! [X512,X513] : (sP9(X512,X513,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | ~sP8(X513,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X512) | ~p(s(bool,X513))) ) | $spl102),
  inference(cnf_transformation,[],[f16460_D])).
fof(f16460_D,plain,(
  ( ! [X512,X513] : (sP9(X512,X513,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | ~sP8(X513,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X512) | ~p(s(bool,X513))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f16459,plain,(
  ( ! [X511,X512,X513] : (~p(s(bool,X511)) | sP9(X512,X513,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | ~p(s(bool,X513)) | ~sP8(X513,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X512)) ) | ($spl8 | $spl46)),
  inference(subsumption_resolution,[],[f15865,f350])).
fof(f15865,plain,(
  ( ! [X511,X512,X513] : (~p(s(bool,X511)) | sP9(X512,X513,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | ~p(s(bool,X513)) | ~sP8(X513,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X512) | s(bool,'F_0') = s(bool,X511)) ) | ($spl8 | $spl46)),
  inference(superposition,[],[f404,f15318])).
fof(f404,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP8(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f241])).
fof(f241,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | (sP8(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP8(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f240])).
fof(f240,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | (sP8(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP8(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP9(X2,X0,X1)))),
  inference(flattening,[],[f239])).
fof(f239,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | (sP8(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP8(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP9(X2,X0,X1)))),
  inference(nnf_transformation,[],[f186])).
fof(f186,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> (sP8(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f16458,plain,(
  ( ! [X509,X510] : (sP9(X509,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X510) | ~sP8('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X510,X509) | ~p(s(bool,X510))) ) | $spl100),
  inference(cnf_transformation,[],[f16458_D])).
fof(f16458_D,plain,(
  ( ! [X509,X510] : (sP9(X509,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X510) | ~sP8('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X510,X509) | ~p(s(bool,X510))) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f16457,plain,(
  ( ! [X509,X508,X510] : (~p(s(bool,X508)) | sP9(X509,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X510) | ~p(s(bool,X510)) | ~sP8('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X510,X509)) ) | ($spl8 | $spl46)),
  inference(subsumption_resolution,[],[f15864,f350])).
fof(f15864,plain,(
  ( ! [X509,X508,X510] : (~p(s(bool,X508)) | sP9(X509,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X510) | ~p(s(bool,X510)) | ~sP8('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X510,X509) | s(bool,'F_0') = s(bool,X508)) ) | ($spl8 | $spl46)),
  inference(superposition,[],[f404,f15318])).
fof(f16430,plain,(
  ( ! [X344,X345] : (sP5(X344,X345,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X345)) | ~sP4(X344,X345,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | $spl98),
  inference(cnf_transformation,[],[f16430_D])).
fof(f16430_D,plain,(
  ( ! [X344,X345] : (sP5(X344,X345,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X345)) | ~sP4(X344,X345,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f16429,plain,(
  ( ! [X344,X343,X345] : (~p(s(bool,X343)) | sP5(X344,X345,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | ~sP4(X344,X345,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X345))) ) | ($spl8 | $spl46)),
  inference(subsumption_resolution,[],[f15809,f350])).
fof(f15809,plain,(
  ( ! [X344,X343,X345] : (~p(s(bool,X343)) | sP5(X344,X345,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | ~sP4(X344,X345,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X345)) | s(bool,'F_0') = s(bool,X343)) ) | ($spl8 | $spl46)),
  inference(superposition,[],[f373,f15318])).
fof(f373,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | ~sP4(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f225])).
fof(f2413,plain,(
  ( ! [X2,X3] : (sP3(X2,X3,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | ~sP1(X3,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X2) | ~sP2(X2,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X3)) ) | $spl34),
  inference(cnf_transformation,[],[f2413_D])).
fof(f2413_D,plain,(
  ( ! [X2,X3] : (sP3(X2,X3,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | ~sP1(X3,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X2) | ~sP2(X2,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X3)) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f16375,plain,(
  ( ! [X151,X149,X150] : (~p(s(bool,X149)) | sP3(X150,X151,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | ~sP2(X150,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X151) | ~sP1(X151,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X150)) ) | ($spl6 | $spl8 | $spl46)),
  inference(subsumption_resolution,[],[f16374,f571])).
fof(f571,plain,(
  ( ! [X48,X49] : (sP0('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X48,X49)) ) | $spl6),
  inference(resolution,[],[f550,f331])).
fof(f16374,plain,(
  ( ! [X151,X149,X150] : (~p(s(bool,X149)) | sP3(X150,X151,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | ~sP2(X150,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X151) | ~sP1(X151,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X150) | ~sP0('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X151,X150)) ) | ($spl8 | $spl46)),
  inference(subsumption_resolution,[],[f15723,f350])).
fof(f15723,plain,(
  ( ! [X151,X149,X150] : (~p(s(bool,X149)) | sP3(X150,X151,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | ~sP2(X150,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X151) | ~sP1(X151,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X150) | ~sP0('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X151,X150) | s(bool,'F_0') = s(bool,X149)) ) | ($spl8 | $spl46)),
  inference(superposition,[],[f317,f15318])).
fof(f16368,plain,(
  ( ! [X127,X125,X126,X124] : (s(bool,X124) != s(bool,X125) | p(s(bool,X125)) | ~p(s(bool,X126)) | s(bool,X126) != s(bool,X127) | s(bool,'F_0') = s(bool,X124)) ) | ($spl8 | $spl46)),
  inference(subsumption_resolution,[],[f16367,f537])).
fof(f16367,plain,(
  ( ! [X127,X125,X126,X124] : (~p(s(bool,'T_0')) | s(bool,X124) != s(bool,X125) | p(s(bool,X125)) | ~p(s(bool,X126)) | s(bool,X126) != s(bool,X127) | s(bool,'F_0') = s(bool,X124)) ) | ($spl8 | $spl46)),
  inference(forward_demodulation,[],[f15716,f2535])).
fof(f15716,plain,(
  ( ! [X127,X125,X126,X124] : (s(bool,X124) != s(bool,X125) | p(s(bool,X125)) | ~p(s(bool,X126)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | s(bool,X126) != s(bool,X127) | s(bool,'F_0') = s(bool,X124)) ) | ($spl8 | $spl46)),
  inference(superposition,[],[f312,f15318])).
fof(f16356,plain,(
  ( ! [X90,X88,X91,X89] : (s(bool,X88) != s(bool,X89) | p(s(bool,X90)) | ~p(s(bool,X91)) | p(s(bool,X89)) | s(bool,'F_0') = s(bool,X88)) ) | ($spl8 | $spl46)),
  inference(subsumption_resolution,[],[f16355,f537])).
fof(f16355,plain,(
  ( ! [X90,X88,X91,X89] : (~p(s(bool,'T_0')) | s(bool,X88) != s(bool,X89) | p(s(bool,X90)) | ~p(s(bool,X91)) | p(s(bool,X89)) | s(bool,'F_0') = s(bool,X88)) ) | ($spl8 | $spl46)),
  inference(forward_demodulation,[],[f15707,f2535])).
fof(f15707,plain,(
  ( ! [X90,X88,X91,X89] : (s(bool,X88) != s(bool,X89) | p(s(bool,X90)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | ~p(s(bool,X91)) | p(s(bool,X89)) | s(bool,'F_0') = s(bool,X88)) ) | ($spl8 | $spl46)),
  inference(superposition,[],[f311,f15318])).
fof(f16350,plain,(
  ( ! [X61,X62,X60,X63] : (s(bool,X60) != s(bool,X61) | p(s(bool,X62)) | p(s(bool,X61)) | s(bool,X62) != s(bool,X63) | s(bool,'F_0') = s(bool,X60)) ) | ($spl8 | $spl46)),
  inference(subsumption_resolution,[],[f16349,f537])).
fof(f16349,plain,(
  ( ! [X61,X62,X60,X63] : (~p(s(bool,'T_0')) | s(bool,X60) != s(bool,X61) | p(s(bool,X62)) | p(s(bool,X61)) | s(bool,X62) != s(bool,X63) | s(bool,'F_0') = s(bool,X60)) ) | ($spl8 | $spl46)),
  inference(forward_demodulation,[],[f15700,f2535])).
fof(f15700,plain,(
  ( ! [X61,X62,X60,X63] : (s(bool,X60) != s(bool,X61) | p(s(bool,X62)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | p(s(bool,X61)) | s(bool,X62) != s(bool,X63) | s(bool,'F_0') = s(bool,X60)) ) | ($spl8 | $spl46)),
  inference(superposition,[],[f310,f15318])).
fof(f16342,plain,(
  ( ! [X43,X41,X42,X40] : (s(bool,X40) != s(bool,X41) | p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X42) != s(bool,X43) | s(bool,'F_0') = s(bool,X40)) ) | ($spl8 | $spl46)),
  inference(subsumption_resolution,[],[f16341,f537])).
fof(f16341,plain,(
  ( ! [X43,X41,X42,X40] : (~p(s(bool,'T_0')) | s(bool,X40) != s(bool,X41) | p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X42) != s(bool,X43) | s(bool,'F_0') = s(bool,X40)) ) | ($spl8 | $spl46)),
  inference(forward_demodulation,[],[f15695,f2535])).
fof(f15695,plain,(
  ( ! [X43,X41,X42,X40] : (s(bool,X40) != s(bool,X41) | p(s(bool,X41)) | ~p(s(bool,X42)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | s(bool,X42) != s(bool,X43) | s(bool,'F_0') = s(bool,X40)) ) | ($spl8 | $spl46)),
  inference(superposition,[],[f308,f15318])).
fof(f16336,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) != s(bool,X9) | p(s(bool,X10)) | p(s(bool,X9)) | s(bool,X10) != s(bool,X11) | s(bool,'F_0') = s(bool,X8)) ) | ($spl8 | $spl46)),
  inference(subsumption_resolution,[],[f16335,f537])).
fof(f16335,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,'T_0')) | s(bool,X8) != s(bool,X9) | p(s(bool,X10)) | p(s(bool,X9)) | s(bool,X10) != s(bool,X11) | s(bool,'F_0') = s(bool,X8)) ) | ($spl8 | $spl46)),
  inference(forward_demodulation,[],[f15687,f2535])).
fof(f15687,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) != s(bool,X9) | p(s(bool,X10)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | p(s(bool,X9)) | s(bool,X10) != s(bool,X11) | s(bool,'F_0') = s(bool,X8)) ) | ($spl8 | $spl46)),
  inference(superposition,[],[f306,f15318])).
fof(f16332,plain,(
  ( ! [X2,X3] : (s(X2,happ(s(fun('type.num.num',X2),X3),s('type.num.num',sK16))) = s(X2,'const.list.EL_2'(s('type.num.num',sK16),s('type.list.list'(X2),'const.list.GENLIST_2'(s(fun('type.num.num',X2),X3),s('type.num.num',sK14)))))) ) | $spl96),
  inference(cnf_transformation,[],[f16332_D])).
fof(f16332_D,plain,(
  ( ! [X2,X3] : (s(X2,happ(s(fun('type.num.num',X2),X3),s('type.num.num',sK16))) = s(X2,'const.list.EL_2'(s('type.num.num',sK16),s('type.list.list'(X2),'const.list.GENLIST_2'(s(fun('type.num.num',X2),X3),s('type.num.num',sK14)))))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f16331,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,X1)) | s(X2,happ(s(fun('type.num.num',X2),X3),s('type.num.num',sK16))) = s(X2,'const.list.EL_2'(s('type.num.num',sK16),s('type.list.list'(X2),'const.list.GENLIST_2'(s(fun('type.num.num',X2),X3),s('type.num.num',sK14)))))) ) | ($spl8 | $spl46)),
  inference(subsumption_resolution,[],[f15685,f350])).
fof(f15685,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,X1)) | s(X2,happ(s(fun('type.num.num',X2),X3),s('type.num.num',sK16))) = s(X2,'const.list.EL_2'(s('type.num.num',sK16),s('type.list.list'(X2),'const.list.GENLIST_2'(s(fun('type.num.num',X2),X3),s('type.num.num',sK14))))) | s(bool,'F_0') = s(bool,X1)) ) | ($spl8 | $spl46)),
  inference(superposition,[],[f360,f15318])).
fof(f16293,plain,(
  ( ! [X660,X661,X659] : (p(s(bool,X660)) | p(s(bool,X661)) | s(bool,X659) != s(bool,X661) | s(bool,'F_0') = s(bool,X659)) ) | ($spl8 | $spl46)),
  inference(subsumption_resolution,[],[f16292,f537])).
fof(f16292,plain,(
  ( ! [X660,X661,X659] : (~p(s(bool,'T_0')) | p(s(bool,X660)) | p(s(bool,X661)) | s(bool,X659) != s(bool,X661) | s(bool,'F_0') = s(bool,X659)) ) | ($spl8 | $spl46)),
  inference(forward_demodulation,[],[f15650,f2535])).
fof(f15650,plain,(
  ( ! [X660,X661,X659] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | p(s(bool,X660)) | p(s(bool,X661)) | s(bool,X659) != s(bool,X661) | s(bool,'F_0') = s(bool,X659)) ) | ($spl8 | $spl46)),
  inference(superposition,[],[f536,f15318])).
fof(f16028,plain,(
  ( ! [X134,X132,X135,X133] : (p(s(bool,X133)) | ~p(s(bool,X134)) | s(bool,X132) != s(bool,X133) | s(bool,X134) != s(bool,X135) | s(bool,'F_0') = s(bool,X132)) ) | ($spl8 | $spl46)),
  inference(subsumption_resolution,[],[f16027,f537])).
fof(f16027,plain,(
  ( ! [X134,X132,X135,X133] : (~p(s(bool,'T_0')) | p(s(bool,X133)) | ~p(s(bool,X134)) | s(bool,X132) != s(bool,X133) | s(bool,X134) != s(bool,X135) | s(bool,'F_0') = s(bool,X132)) ) | ($spl8 | $spl46)),
  inference(forward_demodulation,[],[f15434,f2535])).
fof(f15434,plain,(
  ( ! [X134,X132,X135,X133] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | p(s(bool,X133)) | ~p(s(bool,X134)) | s(bool,X132) != s(bool,X133) | s(bool,X134) != s(bool,X135) | s(bool,'F_0') = s(bool,X132)) ) | ($spl8 | $spl46)),
  inference(superposition,[],[f312,f15318])).
fof(f16021,plain,(
  ( ! [X103,X101,X102,X100] : (p(s(bool,X101)) | ~p(s(bool,X102)) | p(s(bool,X103)) | s(bool,X100) != s(bool,X103) | s(bool,'F_0') = s(bool,X100)) ) | ($spl8 | $spl46)),
  inference(subsumption_resolution,[],[f16020,f537])).
fof(f16020,plain,(
  ( ! [X103,X101,X102,X100] : (~p(s(bool,'T_0')) | p(s(bool,X101)) | ~p(s(bool,X102)) | p(s(bool,X103)) | s(bool,X100) != s(bool,X103) | s(bool,'F_0') = s(bool,X100)) ) | ($spl8 | $spl46)),
  inference(forward_demodulation,[],[f15426,f2535])).
fof(f15426,plain,(
  ( ! [X103,X101,X102,X100] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | p(s(bool,X101)) | ~p(s(bool,X102)) | p(s(bool,X103)) | s(bool,X100) != s(bool,X103) | s(bool,'F_0') = s(bool,X100)) ) | ($spl8 | $spl46)),
  inference(superposition,[],[f311,f15318])).
fof(f16014,plain,(
  ( ! [X80,X83,X81,X82] : (p(s(bool,X81)) | p(s(bool,X82)) | s(bool,X81) != s(bool,X83) | s(bool,X80) != s(bool,X82) | s(bool,'F_0') = s(bool,X80)) ) | ($spl8 | $spl46)),
  inference(subsumption_resolution,[],[f16013,f537])).
fof(f16013,plain,(
  ( ! [X80,X83,X81,X82] : (~p(s(bool,'T_0')) | p(s(bool,X81)) | p(s(bool,X82)) | s(bool,X81) != s(bool,X83) | s(bool,X80) != s(bool,X82) | s(bool,'F_0') = s(bool,X80)) ) | ($spl8 | $spl46)),
  inference(forward_demodulation,[],[f15421,f2535])).
fof(f15421,plain,(
  ( ! [X80,X83,X81,X82] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | p(s(bool,X81)) | p(s(bool,X82)) | s(bool,X81) != s(bool,X83) | s(bool,X80) != s(bool,X82) | s(bool,'F_0') = s(bool,X80)) ) | ($spl8 | $spl46)),
  inference(superposition,[],[f310,f15318])).
fof(f16008,plain,(
  ( ! [X54,X52,X55,X53] : (p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X52) != s(bool,X53) | s(bool,X54) != s(bool,X55) | s(bool,'F_0') = s(bool,X52)) ) | ($spl8 | $spl46)),
  inference(subsumption_resolution,[],[f16007,f537])).
fof(f16007,plain,(
  ( ! [X54,X52,X55,X53] : (~p(s(bool,'T_0')) | p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X52) != s(bool,X53) | s(bool,X54) != s(bool,X55) | s(bool,'F_0') = s(bool,X52)) ) | ($spl8 | $spl46)),
  inference(forward_demodulation,[],[f15414,f2535])).
fof(f15414,plain,(
  ( ! [X54,X52,X55,X53] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X52) != s(bool,X53) | s(bool,X54) != s(bool,X55) | s(bool,'F_0') = s(bool,X52)) ) | ($spl8 | $spl46)),
  inference(superposition,[],[f308,f15318])).
fof(f16000,plain,(
  ( ! [X26,X24,X27,X25] : (p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X24) != s(bool,X26) | s(bool,'F_0') = s(bool,X24)) ) | ($spl8 | $spl46)),
  inference(subsumption_resolution,[],[f15999,f537])).
fof(f15999,plain,(
  ( ! [X26,X24,X27,X25] : (~p(s(bool,'T_0')) | p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X24) != s(bool,X26) | s(bool,'F_0') = s(bool,X24)) ) | ($spl8 | $spl46)),
  inference(forward_demodulation,[],[f15407,f2535])).
fof(f15407,plain,(
  ( ! [X26,X24,X27,X25] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X24) != s(bool,X26) | s(bool,'F_0') = s(bool,X24)) ) | ($spl8 | $spl46)),
  inference(superposition,[],[f306,f15318])).
fof(f14590,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | $spl91),
  inference(cnf_transformation,[],[f14590_D])).
fof(f14590_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) <=> ~$spl91),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl91])])).
fof(f14810,plain,(
  ( ! [X8] : (sP3(X8,'F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) ) | $spl92),
  inference(cnf_transformation,[],[f14810_D])).
fof(f14810_D,plain,(
  ( ! [X8] : (sP3(X8,'F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f14807,plain,(
  ( ! [X8] : (sP3(X8,'F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) ) | ($spl1 | $spl6)),
  inference(resolution,[],[f14382,f723])).
fof(f723,plain,(
  ( ! [X33,X34] : (sP2(X33,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))),X34)) ) | $spl1),
  inference(resolution,[],[f545,f322])).
fof(f14382,plain,(
  ( ! [X4,X5] : (~sP2(X5,X4,'F_0') | sP3(X5,'F_0',X4) | s(bool,'T_0') != s(bool,X4)) ) | $spl6),
  inference(subsumption_resolution,[],[f14381,f528])).
fof(f14381,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | sP3(X5,'F_0',X4) | p(s(bool,'F_0')) | ~sP2(X5,X4,'F_0')) ) | $spl6),
  inference(subsumption_resolution,[],[f14378,f327])).
fof(f14378,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | sP3(X5,'F_0',X4) | p(s(bool,'F_0')) | ~sP2(X5,X4,'F_0') | ~sP1('F_0',X4,X5)) ) | $spl6),
  inference(resolution,[],[f14372,f318])).
fof(f14798,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f14703,f679])).
fof(f14703,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(resolution,[],[f679,f531])).
fof(f531,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f301])).
fof(f14797,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f14702,f679])).
fof(f14702,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(resolution,[],[f679,f529])).
fof(f529,plain,(
  p(s(bool,'F_0')) | p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f300])).
fof(f14796,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f14697,f679])).
fof(f14697,plain,(
  ( ! [X260] : (p(s(bool,X260))) ) | $spl14),
  inference(resolution,[],[f679,f485])).
fof(f14795,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f14695,f679])).
fof(f14695,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(resolution,[],[f679,f482])).
fof(f482,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f277])).
fof(f14794,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f14694,f679])).
fof(f14694,plain,(
  ( ! [X258] : (p(s(bool,X258))) ) | $spl14),
  inference(resolution,[],[f679,f479])).
fof(f14793,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f14692,f679])).
fof(f14692,plain,(
  ( ! [X256] : (p(s(bool,X256))) ) | $spl14),
  inference(resolution,[],[f679,f476])).
fof(f14704,plain,(
  $false | $spl14),
  inference(resolution,[],[f679,f533])).
fof(f14705,plain,(
  $false | $spl14),
  inference(resolution,[],[f679,f534])).
fof(f14706,plain,(
  $false | $spl14),
  inference(resolution,[],[f679,f537])).
fof(f14586,plain,(
  ( ! [X8] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | ~p(s(bool,X8))) ) | ($spl1 | $spl6)),
  inference(resolution,[],[f14550,f723])).
fof(f837,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) ) | $spl18),
  inference(cnf_transformation,[],[f837_D])).
fof(f837_D,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f14283,plain,(
  ( ! [X161,X162] : (p(s(bool,X161)) | ~p(s(bool,X162)) | s(bool,X162) != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) ) | $spl1),
  inference(resolution,[],[f545,f536])).
fof(f14192,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | s(bool,X22) != s(bool,X24)) ) | $spl1),
  inference(resolution,[],[f545,f312])).
fof(f14190,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) ) | $spl1),
  inference(resolution,[],[f545,f311])).
fof(f14188,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) ) | $spl1),
  inference(resolution,[],[f545,f310])).
fof(f14187,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | s(bool,X7) != s(bool,X9)) ) | $spl1),
  inference(resolution,[],[f545,f308])).
fof(f14185,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) ) | $spl1),
  inference(resolution,[],[f545,f306])).
fof(f14131,plain,(
  $false | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f14130,f1006])).
fof(f1006,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK19(sK13,'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)),sK12)),s('type.num.num',sK14)))) | $spl0),
  inference(forward_demodulation,[],[f880,f535])).
fof(f880,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK19(sK13,'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)),sK12)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) | $spl0),
  inference(resolution,[],[f544,f364])).
fof(f14130,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK19(sK13,'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)),sK12)),s('type.num.num',sK14)))) | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f14119,f546])).
fof(f14119,plain,(
  s(sK12,sK13) = s(sK12,happ(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK19(sK13,'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)),sK12)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK19(sK13,'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)),sK12)),s('type.num.num',sK14)))) | $spl0),
  inference(superposition,[],[f360,f2278])).
fof(f2278,plain,(
  s(sK12,sK13) = s(sK12,'const.list.EL_2'(s('type.num.num',sK19(sK13,'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)),sK12)),s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))) | $spl0),
  inference(resolution,[],[f544,f365])).
fof(f14129,plain,(
  $false | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f14128,f1006])).
fof(f14128,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK19(sK13,'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)),sK12)),s('type.num.num',sK14)))) | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f14118,f546])).
fof(f14118,plain,(
  s(sK12,sK13) = s(sK12,happ(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK19(sK13,'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)),sK12)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK19(sK13,'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)),sK12)),s('type.num.num',sK14)))) | $spl0),
  inference(superposition,[],[f2278,f360])).
fof(f13442,plain,(
  s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK19(sK13,'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)),sK12)),s('type.num.num',sK14))) | $spl89),
  inference(cnf_transformation,[],[f13442_D])).
fof(f13442_D,plain,(
  s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK19(sK13,'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)),sK12)),s('type.num.num',sK14))) <=> ~$spl89),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl89])])).
fof(f13434,plain,(
  ( ! [X15] : (p(s(bool,X15)) | s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK19(sK13,'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)),sK12)),s('type.num.num',sK14)))) ) | $spl0),
  inference(resolution,[],[f13259,f4175])).
fof(f4175,plain,(
  ( ! [X48,X49] : (sP0('const.prim_rec.<_2'(s('type.num.num',sK19(sK13,'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)),sK12)),s('type.num.num',sK14)),X48,X49)) ) | $spl0),
  inference(resolution,[],[f1006,f331])).
fof(f13259,plain,(
  ( ! [X2,X3] : (~sP0(X3,'F_0',X2) | p(s(bool,X2)) | s(bool,'F_0') != s(bool,X3)) ) | $spl0),
  inference(duplicate_literal_removal,[],[f13252])).
fof(f13252,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | ~sP0(X3,'F_0',X2) | p(s(bool,X2)) | s(bool,'F_0') != s(bool,X3)) ) | $spl0),
  inference(resolution,[],[f13245,f335])).
fof(f13245,plain,(
  ( ! [X6,X7] : (sP3(X6,'F_0',X7) | p(s(bool,X6)) | ~sP0(X7,'F_0',X6)) ) | $spl0),
  inference(subsumption_resolution,[],[f13242,f323])).
fof(f13242,plain,(
  ( ! [X6,X7] : (sP3(X6,'F_0',X7) | p(s(bool,X6)) | ~sP2(X6,X7,'F_0') | ~sP0(X7,'F_0',X6)) ) | $spl0),
  inference(resolution,[],[f13236,f319])).
fof(f13236,plain,(
  ( ! [X6,X7] : (sP1('F_0',X6,X7)) ) | $spl0),
  inference(subsumption_resolution,[],[f13235,f326])).
fof(f13235,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | sP1('F_0',X6,X7)) ) | $spl0),
  inference(subsumption_resolution,[],[f13229,f325])).
fof(f13229,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | p(s(bool,X7)) | sP1('F_0',X6,X7)) ) | $spl0),
  inference(resolution,[],[f13224,f314])).
fof(f314,plain,(
  ( ! [X2,X0,X1] : (sP1(X1,X2,X0) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f199])).
fof(f13224,plain,(
  ( ! [X4,X5] : (sP3(X5,'F_0',X4) | ~p(s(bool,X4)) | p(s(bool,X5))) ) | $spl0),
  inference(resolution,[],[f13199,f323])).
fof(f13199,plain,(
  ( ! [X2,X3] : (~sP2(X3,X2,'F_0') | ~p(s(bool,X2)) | sP3(X3,'F_0',X2)) ) | $spl0),
  inference(subsumption_resolution,[],[f13198,f355])).
fof(f13198,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | sP3(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP2(X3,X2,'F_0')) ) | $spl0),
  inference(subsumption_resolution,[],[f13194,f13182])).
fof(f13182,plain,(
  ( ! [X4,X5] : (sP0(X4,'F_0',X5) | s(bool,'T_0') != s(bool,X4)) ) | $spl0),
  inference(subsumption_resolution,[],[f13175,f329])).
fof(f13175,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | p(s(bool,X5)) | sP0(X4,'F_0',X5)) ) | $spl0),
  inference(resolution,[],[f11750,f313])).
fof(f11750,plain,(
  ( ! [X6,X7] : (sP3(X7,'F_0',X6) | s(bool,'T_0') != s(bool,X6) | p(s(bool,X7))) ) | $spl0),
  inference(superposition,[],[f8461,f332])).
fof(f8461,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl0),
  inference(backward_demodulation,[],[f2309,f2307])).
fof(f2307,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | $spl0),
  inference(resolution,[],[f544,f350])).
fof(f2309,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | $spl0),
  inference(resolution,[],[f544,f355])).
fof(f13194,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | sP3(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP2(X3,X2,'F_0') | ~sP0(X2,'F_0',X3)) ) | $spl0),
  inference(resolution,[],[f13183,f317])).
fof(f13183,plain,(
  ( ! [X6,X7] : (sP1('F_0',X6,X7) | s(bool,'T_0') != s(bool,X6)) ) | $spl0),
  inference(subsumption_resolution,[],[f13176,f325])).
fof(f13176,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X6) | p(s(bool,X7)) | sP1('F_0',X6,X7)) ) | $spl0),
  inference(resolution,[],[f11750,f314])).
fof(f12784,plain,(
  $false | ($spl0 | $spl18)),
  inference(subsumption_resolution,[],[f12687,f12574])).
fof(f12574,plain,(
  ( ! [X2] : (~p(s(bool,X2))) ) | ($spl0 | $spl18)),
  inference(subsumption_resolution,[],[f12573,f357])).
fof(f12573,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | ~p(s(bool,X2))) ) | ($spl0 | $spl18)),
  inference(forward_demodulation,[],[f837,f2309])).
fof(f12687,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl18)),
  inference(resolution,[],[f12574,f531])).
fof(f12783,plain,(
  $false | ($spl0 | $spl18)),
  inference(subsumption_resolution,[],[f12686,f12574])).
fof(f12686,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl18)),
  inference(resolution,[],[f12574,f529])).
fof(f12782,plain,(
  $false | ($spl0 | $spl18)),
  inference(subsumption_resolution,[],[f12681,f12574])).
fof(f12681,plain,(
  ( ! [X260] : (p(s(bool,X260))) ) | ($spl0 | $spl18)),
  inference(resolution,[],[f12574,f485])).
fof(f12781,plain,(
  $false | ($spl0 | $spl18)),
  inference(subsumption_resolution,[],[f12679,f12574])).
fof(f12679,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl18)),
  inference(resolution,[],[f12574,f482])).
fof(f12780,plain,(
  $false | ($spl0 | $spl18)),
  inference(subsumption_resolution,[],[f12678,f12574])).
fof(f12678,plain,(
  ( ! [X258] : (p(s(bool,X258))) ) | ($spl0 | $spl18)),
  inference(resolution,[],[f12574,f479])).
fof(f12779,plain,(
  $false | ($spl0 | $spl18)),
  inference(subsumption_resolution,[],[f12676,f12574])).
fof(f12676,plain,(
  ( ! [X256] : (p(s(bool,X256))) ) | ($spl0 | $spl18)),
  inference(resolution,[],[f12574,f476])).
fof(f12688,plain,(
  $false | ($spl0 | $spl18)),
  inference(resolution,[],[f12574,f533])).
fof(f12689,plain,(
  $false | ($spl0 | $spl18)),
  inference(resolution,[],[f12574,f534])).
fof(f12690,plain,(
  $false | ($spl0 | $spl18)),
  inference(resolution,[],[f12574,f537])).
fof(f12699,plain,(
  $false | ($spl0 | $spl18)),
  inference(resolution,[],[f12574,f1006])).
fof(f2535,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | $spl46),
  inference(cnf_transformation,[],[f2535_D])).
fof(f2535_D,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f11907,plain,(
  ( ! [X35] : (p(s(bool,X35)) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | $spl8),
  inference(subsumption_resolution,[],[f11884,f351])).
fof(f11884,plain,(
  ( ! [X35] : (s(bool,'F_0') != s(bool,X35) | p(s(bool,X35)) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | $spl8),
  inference(superposition,[],[f676,f359])).
fof(f11906,plain,(
  ( ! [X33] : (p(s(bool,X33)) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | $spl8),
  inference(subsumption_resolution,[],[f11882,f351])).
fof(f11882,plain,(
  ( ! [X33] : (s(bool,'F_0') != s(bool,X33) | p(s(bool,X33)) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | $spl8),
  inference(superposition,[],[f676,f358])).
fof(f550,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | $spl6),
  inference(cnf_transformation,[],[f550_D])).
fof(f550_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f11905,plain,(
  ( ! [X29] : (p(s(bool,X29)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))))) ) | $spl8),
  inference(subsumption_resolution,[],[f11878,f351])).
fof(f11878,plain,(
  ( ! [X29] : (s(bool,'F_0') != s(bool,X29) | p(s(bool,X29)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))))) ) | $spl8),
  inference(superposition,[],[f676,f353])).
fof(f11904,plain,(
  ( ! [X28] : (p(s(bool,X28)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))))) ) | $spl8),
  inference(subsumption_resolution,[],[f11877,f351])).
fof(f11877,plain,(
  ( ! [X28] : (s(bool,'F_0') != s(bool,X28) | p(s(bool,X28)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))))) ) | $spl8),
  inference(superposition,[],[f676,f351])).
fof(f11901,plain,(
  ( ! [X23,X22] : (p(s(bool,X23)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | p(s(bool,X22))) ) | $spl8),
  inference(subsumption_resolution,[],[f11874,f338])).
fof(f338,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f210])).
fof(f11874,plain,(
  ( ! [X23,X22] : (s(bool,X22) != s(bool,X23) | p(s(bool,X23)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | p(s(bool,X22))) ) | $spl8),
  inference(superposition,[],[f676,f342])).
fof(f342,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f149])).
fof(f149,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f148])).
fof(f148,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f68])).
fof(f68,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X12,X13] : ((p(s(bool,X12)) => p(s(bool,X13))) => ((p(s(bool,X13)) => p(s(bool,X12))) => s(bool,X12) = s(bool,X13)))),
  file('thm.list.MEM_GENLIST','thm.bool.IMP_ANTISYM_AX')).
fof(f11900,plain,(
  ( ! [X21,X20] : (p(s(bool,X21)) | p(s(bool,X20)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))))) ) | $spl8),
  inference(subsumption_resolution,[],[f11873,f338])).
fof(f11873,plain,(
  ( ! [X21,X20] : (s(bool,X20) != s(bool,X21) | p(s(bool,X21)) | p(s(bool,X20)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))))) ) | $spl8),
  inference(superposition,[],[f676,f342])).
fof(f11897,plain,(
  ( ! [X14,X15] : (p(s(bool,X15)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | p(s(bool,X14))) ) | $spl8),
  inference(subsumption_resolution,[],[f11870,f338])).
fof(f11870,plain,(
  ( ! [X14,X15] : (s(bool,X14) != s(bool,X15) | p(s(bool,X15)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | p(s(bool,X14))) ) | $spl8),
  inference(superposition,[],[f676,f338])).
fof(f11896,plain,(
  ( ! [X12,X13] : (p(s(bool,X13)) | p(s(bool,X12)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))))) ) | $spl8),
  inference(subsumption_resolution,[],[f11869,f338])).
fof(f11869,plain,(
  ( ! [X12,X13] : (s(bool,X12) != s(bool,X13) | p(s(bool,X13)) | p(s(bool,X12)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))))) ) | $spl8),
  inference(superposition,[],[f676,f338])).
fof(f2519,plain,(
  s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | $spl45),
  inference(cnf_transformation,[],[f2519_D])).
fof(f2519_D,plain,(
  s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) <=> ~$spl45),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl45])])).
fof(f11893,plain,(
  ( ! [X35] : (s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X35))) ) | $spl8),
  inference(subsumption_resolution,[],[f11860,f354])).
fof(f11860,plain,(
  ( ! [X35] : (s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X35)) | s(bool,'T_0') = s(bool,X35)) ) | $spl8),
  inference(superposition,[],[f676,f359])).
fof(f11892,plain,(
  ( ! [X33] : (s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X33))) ) | $spl8),
  inference(subsumption_resolution,[],[f11858,f354])).
fof(f11858,plain,(
  ( ! [X33] : (s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X33)) | s(bool,'T_0') = s(bool,X33)) ) | $spl8),
  inference(superposition,[],[f676,f358])).
fof(f11886,plain,(
  ( ! [X29] : (s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X29))) ) | $spl8),
  inference(duplicate_literal_removal,[],[f11854])).
fof(f11854,plain,(
  ( ! [X29] : (s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X29)) | p(s(bool,X29))) ) | $spl8),
  inference(superposition,[],[f676,f353])).
fof(f11887,plain,(
  ( ! [X28] : (s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X28))) ) | $spl8),
  inference(duplicate_literal_removal,[],[f11853])).
fof(f11853,plain,(
  ( ! [X28] : (s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X28)) | p(s(bool,X28))) ) | $spl8),
  inference(superposition,[],[f676,f351])).
fof(f549,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | $spl7),
  inference(cnf_transformation,[],[f549_D])).
fof(f549_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f11808,plain,(
  ( ! [X2,X3] : (sP3(X2,X3,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | ~sP2(X2,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X3) | ~sP1(X3,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X2)) ) | $spl6),
  inference(resolution,[],[f571,f317])).
fof(f2513,plain,(
  ( ! [X28] : (s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X28))) ) | $spl8),
  inference(duplicate_literal_removal,[],[f2480])).
fof(f2480,plain,(
  ( ! [X28] : (s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X28)) | p(s(bool,X28))) ) | $spl8),
  inference(superposition,[],[f676,f351])).
fof(f2536,plain,(
  ( ! [X35] : (p(s(bool,X35)) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | $spl8),
  inference(subsumption_resolution,[],[f2510,f351])).
fof(f2510,plain,(
  ( ! [X35] : (s(bool,'F_0') != s(bool,X35) | p(s(bool,X35)) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | $spl8),
  inference(superposition,[],[f676,f359])).
fof(f2531,plain,(
  ( ! [X28] : (p(s(bool,X28)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))))) ) | $spl8),
  inference(subsumption_resolution,[],[f2503,f351])).
fof(f2503,plain,(
  ( ! [X28] : (s(bool,'F_0') != s(bool,X28) | p(s(bool,X28)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))))) ) | $spl8),
  inference(superposition,[],[f676,f351])).
fof(f2532,plain,(
  ( ! [X29] : (p(s(bool,X29)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))))) ) | $spl8),
  inference(subsumption_resolution,[],[f2504,f351])).
fof(f2504,plain,(
  ( ! [X29] : (s(bool,'F_0') != s(bool,X29) | p(s(bool,X29)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))))) ) | $spl8),
  inference(superposition,[],[f676,f353])).
fof(f2533,plain,(
  ( ! [X33] : (p(s(bool,X33)) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | $spl8),
  inference(subsumption_resolution,[],[f2508,f351])).
fof(f2508,plain,(
  ( ! [X33] : (s(bool,'F_0') != s(bool,X33) | p(s(bool,X33)) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | $spl8),
  inference(superposition,[],[f676,f358])).
fof(f2520,plain,(
  ( ! [X33] : (s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X33))) ) | $spl8),
  inference(subsumption_resolution,[],[f2485,f354])).
fof(f2485,plain,(
  ( ! [X33] : (s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X33)) | s(bool,'T_0') = s(bool,X33)) ) | $spl8),
  inference(superposition,[],[f676,f358])).
fof(f2521,plain,(
  ( ! [X35] : (s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X35))) ) | $spl8),
  inference(subsumption_resolution,[],[f2487,f354])).
fof(f2487,plain,(
  ( ! [X35] : (s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X35)) | s(bool,'T_0') = s(bool,X35)) ) | $spl8),
  inference(superposition,[],[f676,f359])).
fof(f2512,plain,(
  ( ! [X29] : (s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X29))) ) | $spl8),
  inference(duplicate_literal_removal,[],[f2481])).
fof(f2481,plain,(
  ( ! [X29] : (s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X29)) | p(s(bool,X29))) ) | $spl8),
  inference(superposition,[],[f676,f353])).
fof(f6583,plain,(
  ( ! [X33] : (p(s(bool,X33)) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | $spl8),
  inference(subsumption_resolution,[],[f6560,f351])).
fof(f6560,plain,(
  ( ! [X33] : (s(bool,'F_0') != s(bool,X33) | p(s(bool,X33)) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | $spl8),
  inference(superposition,[],[f676,f358])).
fof(f6570,plain,(
  ( ! [X33] : (s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X33))) ) | $spl8),
  inference(subsumption_resolution,[],[f6537,f354])).
fof(f6537,plain,(
  ( ! [X33] : (s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X33)) | s(bool,'T_0') = s(bool,X33)) ) | $spl8),
  inference(superposition,[],[f676,f358])).
fof(f6571,plain,(
  ( ! [X35] : (s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X35))) ) | $spl8),
  inference(subsumption_resolution,[],[f6539,f354])).
fof(f6539,plain,(
  ( ! [X35] : (s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X35)) | s(bool,'T_0') = s(bool,X35)) ) | $spl8),
  inference(superposition,[],[f676,f359])).
fof(f6564,plain,(
  ( ! [X29] : (s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X29))) ) | $spl8),
  inference(duplicate_literal_removal,[],[f6533])).
fof(f6533,plain,(
  ( ! [X29] : (s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X29)) | p(s(bool,X29))) ) | $spl8),
  inference(superposition,[],[f676,f353])).
fof(f6565,plain,(
  ( ! [X28] : (s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X28))) ) | $spl8),
  inference(duplicate_literal_removal,[],[f6532])).
fof(f6532,plain,(
  ( ! [X28] : (s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X28)) | p(s(bool,X28))) ) | $spl8),
  inference(superposition,[],[f676,f351])).
fof(f6584,plain,(
  ( ! [X35] : (p(s(bool,X35)) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | $spl8),
  inference(subsumption_resolution,[],[f6562,f351])).
fof(f6562,plain,(
  ( ! [X35] : (s(bool,'F_0') != s(bool,X35) | p(s(bool,X35)) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | $spl8),
  inference(superposition,[],[f676,f359])).
fof(f11709,plain,(
  ( ! [X368] : (sP10('T_0',X368) | s(bool,'T_0') = s(bool,X368)) ) | $spl82),
  inference(cnf_transformation,[],[f11709_D])).
fof(f11709_D,plain,(
  ( ! [X368] : (sP10('T_0',X368) | s(bool,'T_0') = s(bool,X368)) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f11697,plain,(
  ( ! [X64,X65] : (~p(s(bool,X64)) | sP10('T_0',X65) | s(bool,'T_0') = s(bool,X65)) ) | ($spl24 | $spl44)),
  inference(duplicate_literal_removal,[],[f11686])).
fof(f11686,plain,(
  ( ! [X64,X65] : (~p(s(bool,X64)) | sP10('T_0',X65) | s(bool,'T_0') = s(bool,X65) | ~p(s(bool,X64))) ) | ($spl24 | $spl44)),
  inference(superposition,[],[f11186,f355])).
fof(f11186,plain,(
  ( ! [X590,X591] : (~p(s(bool,X591)) | sP10(X591,X590) | s(bool,'T_0') = s(bool,X590)) ) | ($spl24 | $spl44)),
  inference(subsumption_resolution,[],[f11185,f474])).
fof(f11185,plain,(
  ( ! [X590,X591] : (p(s(bool,'F_0')) | sP10(X591,X590) | ~p(s(bool,X591)) | s(bool,'T_0') = s(bool,X590)) ) | ($spl24 | $spl44)),
  inference(forward_demodulation,[],[f10515,f2518])).
fof(f10515,plain,(
  ( ! [X590,X591] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | sP10(X591,X590) | ~p(s(bool,X591)) | s(bool,'T_0') = s(bool,X590)) ) | ($spl24 | $spl44)),
  inference(superposition,[],[f440,f10208])).
fof(f10208,plain,(
  ( ! [X35] : (s(bool,X35) = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | s(bool,'T_0') = s(bool,X35)) ) | ($spl24 | $spl44)),
  inference(trivial_inequality_removal,[],[f10193])).
fof(f10193,plain,(
  ( ! [X35] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,X35) = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | s(bool,'T_0') = s(bool,X35)) ) | ($spl24 | $spl44)),
  inference(superposition,[],[f10149,f359])).
fof(f10149,plain,(
  ( ! [X23] : (s(bool,'F_0') != s(bool,X23) | s(bool,X23) = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | ($spl24 | $spl44)),
  inference(subsumption_resolution,[],[f10148,f528])).
fof(f10148,plain,(
  ( ! [X23] : (p(s(bool,'F_0')) | s(bool,'F_0') != s(bool,X23) | s(bool,X23) = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | ($spl24 | $spl44)),
  inference(forward_demodulation,[],[f10119,f2518])).
fof(f10119,plain,(
  ( ! [X23] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | s(bool,'F_0') != s(bool,X23) | s(bool,X23) = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | ($spl24 | $spl44)),
  inference(resolution,[],[f7706,f7461])).
fof(f7461,plain,(
  ( ! [X48,X49] : (sP3('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X48,X49) | s(bool,X48) = s(bool,X49)) ) | $spl24),
  inference(resolution,[],[f7423,f332])).
fof(f7423,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | $spl24),
  inference(equality_resolution,[],[f2011])).
fof(f7706,plain,(
  ( ! [X171,X170] : (~sP3(X171,X170,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X171)) | s(bool,'F_0') != s(bool,X170)) ) | $spl44),
  inference(superposition,[],[f335,f2518])).
fof(f11698,plain,(
  ( ! [X62,X63] : (~p(s(bool,X62)) | sP10('T_0',X63) | s(bool,'T_0') = s(bool,X63)) ) | ($spl24 | $spl44)),
  inference(duplicate_literal_removal,[],[f11685])).
fof(f11685,plain,(
  ( ! [X62,X63] : (~p(s(bool,X62)) | sP10('T_0',X63) | s(bool,'T_0') = s(bool,X63) | ~p(s(bool,X62))) ) | ($spl24 | $spl44)),
  inference(superposition,[],[f11186,f357])).
fof(f11721,plain,(
  ( ! [X61,X60] : (~p(s(bool,X60)) | sP10('T_0',X61) | s(bool,'T_0') = s(bool,X61)) ) | ($spl24 | $spl44)),
  inference(subsumption_resolution,[],[f11684,f350])).
fof(f11684,plain,(
  ( ! [X61,X60] : (~p(s(bool,X60)) | sP10('T_0',X61) | s(bool,'T_0') = s(bool,X61) | s(bool,'F_0') = s(bool,X60)) ) | ($spl24 | $spl44)),
  inference(superposition,[],[f11186,f358])).
fof(f11720,plain,(
  ( ! [X59,X58] : (~p(s(bool,X58)) | sP10('T_0',X59) | s(bool,'T_0') = s(bool,X59)) ) | ($spl24 | $spl44)),
  inference(subsumption_resolution,[],[f11683,f350])).
fof(f11683,plain,(
  ( ! [X59,X58] : (~p(s(bool,X58)) | sP10('T_0',X59) | s(bool,'T_0') = s(bool,X59) | s(bool,'F_0') = s(bool,X58)) ) | ($spl24 | $spl44)),
  inference(superposition,[],[f11186,f359])).
fof(f11714,plain,(
  ( ! [X395] : (sP10('F_0',X395) | s(bool,'T_0') = s(bool,X395)) ) | $spl86),
  inference(cnf_transformation,[],[f11714_D])).
fof(f11714_D,plain,(
  ( ! [X395] : (sP10('F_0',X395) | s(bool,'T_0') = s(bool,X395)) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f11656,plain,(
  ( ! [X401,X403,X402] : (sP10('F_0',X401) | s(bool,'T_0') = s(bool,X401) | p(s(bool,X402)) | p(s(bool,X403)) | ~sP11(X402,X403)) ) | ($spl24 | $spl44)),
  inference(resolution,[],[f11186,f454])).
fof(f11654,plain,(
  ( ! [X397,X395,X396] : (sP10('F_0',X395) | s(bool,'T_0') = s(bool,X395) | p(s(bool,X396)) | p(s(bool,X397)) | ~sP11(X396,X397)) ) | ($spl24 | $spl44)),
  inference(resolution,[],[f11186,f458])).
fof(f11645,plain,(
  ( ! [X385,X386] : (sP10('T_0',X385) | s(bool,'T_0') = s(bool,X385) | ~p(s(bool,X386))) ) | ($spl24 | $spl44)),
  inference(resolution,[],[f11186,f507])).
fof(f11644,plain,(
  ( ! [X383,X384] : (sP10('T_0',X383) | s(bool,'T_0') = s(bool,X383) | ~p(s(bool,X384))) ) | ($spl24 | $spl44)),
  inference(resolution,[],[f11186,f505])).
fof(f11643,plain,(
  ( ! [X382,X381] : (sP10('T_0',X381) | s(bool,'T_0') = s(bool,X381) | ~p(s(bool,X382))) ) | ($spl24 | $spl44)),
  inference(resolution,[],[f11186,f500])).
fof(f11642,plain,(
  ( ! [X379,X380] : (sP10('T_0',X379) | s(bool,'T_0') = s(bool,X379) | ~p(s(bool,X380))) ) | ($spl24 | $spl44)),
  inference(resolution,[],[f11186,f498])).
fof(f11641,plain,(
  ( ! [X377,X378] : (sP10('T_0',X377) | s(bool,'T_0') = s(bool,X377) | p(s(bool,X378))) ) | ($spl24 | $spl44)),
  inference(resolution,[],[f11186,f485])).
fof(f11640,plain,(
  ( ! [X376,X375] : (sP10('T_0',X375) | s(bool,'T_0') = s(bool,X375) | ~p(s(bool,X376))) ) | ($spl24 | $spl44)),
  inference(resolution,[],[f11186,f483])).
fof(f11638,plain,(
  ( ! [X372,X373] : (sP10('T_0',X372) | s(bool,'T_0') = s(bool,X372) | p(s(bool,X373))) ) | ($spl24 | $spl44)),
  inference(resolution,[],[f11186,f479])).
fof(f11637,plain,(
  ( ! [X370,X371] : (sP10('T_0',X370) | s(bool,'T_0') = s(bool,X370) | ~p(s(bool,X371))) ) | ($spl24 | $spl44)),
  inference(resolution,[],[f11186,f477])).
fof(f11636,plain,(
  ( ! [X368,X369] : (sP10('T_0',X368) | s(bool,'T_0') = s(bool,X368) | p(s(bool,X369))) ) | ($spl24 | $spl44)),
  inference(resolution,[],[f11186,f476])).
fof(f11703,plain,(
  ( ! [X0,X3,X1] : (sP10(X0,X1) | s(bool,X0) != s(bool,X3) | ~p(s(bool,X3)) | s(bool,'T_0') = s(bool,X1)) ) | $spl80),
  inference(cnf_transformation,[],[f11703_D])).
fof(f11703_D,plain,(
  ( ! [X0,X3,X1] : (sP10(X0,X1) | s(bool,X0) != s(bool,X3) | ~p(s(bool,X3)) | s(bool,'T_0') = s(bool,X1)) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f11633,plain,(
  ( ! [X356,X358,X357,X359] : (sP10(X356,X357) | s(bool,'T_0') = s(bool,X357) | p(s(bool,X358)) | ~p(s(bool,X359)) | s(bool,X356) != s(bool,X359)) ) | ($spl24 | $spl44)),
  inference(resolution,[],[f11186,f536])).
fof(f11542,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP10(X35,X36) | s(bool,'T_0') = s(bool,X36) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X35) != s(bool,X38) | s(bool,X37) != s(bool,X39)) ) | ($spl24 | $spl44)),
  inference(resolution,[],[f11186,f312])).
fof(f11540,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP10(X25,X26) | s(bool,'T_0') = s(bool,X26) | p(s(bool,X27)) | ~p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X25) != s(bool,X28)) ) | ($spl24 | $spl44)),
  inference(resolution,[],[f11186,f311])).
fof(f11538,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP10(X15,X16) | s(bool,'T_0') = s(bool,X16) | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X15) != s(bool,X18)) ) | ($spl24 | $spl44)),
  inference(resolution,[],[f11186,f310])).
fof(f11537,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP10(X10,X11) | s(bool,'T_0') = s(bool,X11) | ~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X10) != s(bool,X13) | s(bool,X12) != s(bool,X14)) ) | ($spl24 | $spl44)),
  inference(resolution,[],[f11186,f308])).
fof(f11535,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10(X0,X1) | s(bool,'T_0') = s(bool,X1) | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X0) != s(bool,X3)) ) | ($spl24 | $spl44)),
  inference(resolution,[],[f11186,f306])).
fof(f6224,plain,(
  ( ! [X438,X437] : (sP7(X437,X438,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X438)) | ~sP6(X437,X438,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | $spl58),
  inference(cnf_transformation,[],[f6224_D])).
fof(f6224_D,plain,(
  ( ! [X438,X437] : (sP7(X437,X438,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X438)) | ~sP6(X437,X438,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f11361,plain,(
  ( ! [X436,X438,X437] : (p(s(bool,X436)) | sP7(X437,X438,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | ~sP6(X437,X438,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X438))) ) | ($spl24 | $spl44)),
  inference(subsumption_resolution,[],[f10753,f354])).
fof(f10753,plain,(
  ( ! [X436,X438,X437] : (p(s(bool,X436)) | sP7(X437,X438,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | ~sP6(X437,X438,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X438)) | s(bool,'T_0') = s(bool,X436)) ) | ($spl24 | $spl44)),
  inference(superposition,[],[f390,f10208])).
fof(f390,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~sP6(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f233])).
fof(f6222,plain,(
  ( ! [X434,X435] : (sP7(X434,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X435) | p(s(bool,X435)) | ~sP6(X434,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X435)) ) | $spl56),
  inference(cnf_transformation,[],[f6222_D])).
fof(f6222_D,plain,(
  ( ! [X434,X435] : (sP7(X434,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X435) | p(s(bool,X435)) | ~sP6(X434,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X435)) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f11360,plain,(
  ( ! [X434,X433,X435] : (p(s(bool,X433)) | sP7(X434,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X435) | ~sP6(X434,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X435) | p(s(bool,X435))) ) | ($spl24 | $spl44)),
  inference(subsumption_resolution,[],[f10752,f354])).
fof(f10752,plain,(
  ( ! [X434,X433,X435] : (p(s(bool,X433)) | sP7(X434,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X435) | ~sP6(X434,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X435) | p(s(bool,X435)) | s(bool,'T_0') = s(bool,X433)) ) | ($spl24 | $spl44)),
  inference(superposition,[],[f390,f10208])).
fof(f6208,plain,(
  ( ! [X344,X345] : (sP5(X344,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X345) | ~p(s(bool,X345)) | ~sP4(X344,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X345)) ) | $spl54),
  inference(cnf_transformation,[],[f6208_D])).
fof(f6208_D,plain,(
  ( ! [X344,X345] : (sP5(X344,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X345) | ~p(s(bool,X345)) | ~sP4(X344,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X345)) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f11347,plain,(
  ( ! [X344,X343,X345] : (p(s(bool,X343)) | sP5(X344,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X345) | ~sP4(X344,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X345) | ~p(s(bool,X345))) ) | ($spl24 | $spl44)),
  inference(subsumption_resolution,[],[f10722,f354])).
fof(f10722,plain,(
  ( ! [X344,X343,X345] : (p(s(bool,X343)) | sP5(X344,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X345) | ~sP4(X344,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X345) | ~p(s(bool,X345)) | s(bool,'T_0') = s(bool,X343)) ) | ($spl24 | $spl44)),
  inference(superposition,[],[f373,f10208])).
fof(f6158,plain,(
  ( ! [X159,X160] : (sP3('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X159,X160) | ~sP0(X160,X159,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | ~sP1(X159,X160,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | $spl52),
  inference(cnf_transformation,[],[f6158_D])).
fof(f6158_D,plain,(
  ( ! [X159,X160] : (sP3('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X159,X160) | ~sP0(X160,X159,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | ~sP1(X159,X160,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f11298,plain,(
  ( ! [X158,X159,X160] : (p(s(bool,X158)) | sP3('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X159,X160) | ~sP1(X159,X160,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | ~sP0(X160,X159,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | ($spl7 | $spl24 | $spl44)),
  inference(subsumption_resolution,[],[f11297,f1899])).
fof(f1899,plain,(
  ( ! [X35,X34] : (sP2('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X34,X35)) ) | $spl7),
  inference(resolution,[],[f549,f323])).
fof(f11297,plain,(
  ( ! [X158,X159,X160] : (p(s(bool,X158)) | sP3('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X159,X160) | ~sP2('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X160,X159) | ~sP1(X159,X160,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | ~sP0(X160,X159,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | ($spl24 | $spl44)),
  inference(subsumption_resolution,[],[f10639,f354])).
fof(f10639,plain,(
  ( ! [X158,X159,X160] : (p(s(bool,X158)) | sP3('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X159,X160) | ~sP2('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X160,X159) | ~sP1(X159,X160,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | ~sP0(X160,X159,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | s(bool,'T_0') = s(bool,X158)) ) | ($spl24 | $spl44)),
  inference(superposition,[],[f319,f10208])).
fof(f6155,plain,(
  ( ! [X156,X157] : (sP3(X156,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X157) | ~sP0(X157,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X156) | ~sP2(X156,X157,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | $spl50),
  inference(cnf_transformation,[],[f6155_D])).
fof(f6155_D,plain,(
  ( ! [X156,X157] : (sP3(X156,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X157) | ~sP0(X157,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X156) | ~sP2(X156,X157,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f11296,plain,(
  ( ! [X156,X157,X155] : (p(s(bool,X155)) | sP3(X156,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X157) | ~sP2(X156,X157,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | ~sP0(X157,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X156)) ) | ($spl7 | $spl24 | $spl44)),
  inference(subsumption_resolution,[],[f11295,f1902])).
fof(f1902,plain,(
  ( ! [X41,X40] : (sP1('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X40,X41)) ) | $spl7),
  inference(resolution,[],[f549,f327])).
fof(f11295,plain,(
  ( ! [X156,X157,X155] : (p(s(bool,X155)) | sP3(X156,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X157) | ~sP2(X156,X157,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | ~sP1('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X157,X156) | ~sP0(X157,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X156)) ) | ($spl24 | $spl44)),
  inference(subsumption_resolution,[],[f10638,f354])).
fof(f10638,plain,(
  ( ! [X156,X157,X155] : (p(s(bool,X155)) | sP3(X156,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X157) | ~sP2(X156,X157,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | ~sP1('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X157,X156) | ~sP0(X157,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X156) | s(bool,'T_0') = s(bool,X155)) ) | ($spl24 | $spl44)),
  inference(superposition,[],[f318,f10208])).
fof(f2518,plain,(
  s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | $spl44),
  inference(cnf_transformation,[],[f2518_D])).
fof(f2518_D,plain,(
  s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f7440,plain,(
  ( ! [X34] : (~p(s(bool,X34)) | s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | $spl24),
  inference(subsumption_resolution,[],[f7421,f355])).
fof(f7421,plain,(
  ( ! [X34] : (s(bool,'T_0') != s(bool,X34) | ~p(s(bool,X34)) | s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | $spl24),
  inference(superposition,[],[f2011,f359])).
fof(f7439,plain,(
  ( ! [X32] : (~p(s(bool,X32)) | s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | $spl24),
  inference(subsumption_resolution,[],[f7419,f355])).
fof(f7419,plain,(
  ( ! [X32] : (s(bool,'T_0') != s(bool,X32) | ~p(s(bool,X32)) | s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | $spl24),
  inference(superposition,[],[f2011,f358])).
fof(f7438,plain,(
  ( ! [X31] : (~p(s(bool,X31)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))))) ) | $spl24),
  inference(subsumption_resolution,[],[f7418,f355])).
fof(f7418,plain,(
  ( ! [X31] : (s(bool,'T_0') != s(bool,X31) | ~p(s(bool,X31)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))))) ) | $spl24),
  inference(superposition,[],[f2011,f357])).
fof(f7437,plain,(
  ( ! [X30] : (~p(s(bool,X30)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))))) ) | $spl24),
  inference(subsumption_resolution,[],[f7417,f355])).
fof(f7417,plain,(
  ( ! [X30] : (s(bool,'T_0') != s(bool,X30) | ~p(s(bool,X30)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))))) ) | $spl24),
  inference(superposition,[],[f2011,f355])).
fof(f7436,plain,(
  ( ! [X26,X27] : (~p(s(bool,X27)) | ~p(s(bool,X26)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))))) ) | $spl24),
  inference(subsumption_resolution,[],[f7414,f341])).
fof(f341,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f210])).
fof(f7414,plain,(
  ( ! [X26,X27] : (s(bool,X26) != s(bool,X27) | ~p(s(bool,X27)) | ~p(s(bool,X26)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))))) ) | $spl24),
  inference(superposition,[],[f2011,f345])).
fof(f345,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f149])).
fof(f7435,plain,(
  ( ! [X24,X25] : (~p(s(bool,X25)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | ~p(s(bool,X24))) ) | $spl24),
  inference(subsumption_resolution,[],[f7413,f341])).
fof(f7413,plain,(
  ( ! [X24,X25] : (s(bool,X24) != s(bool,X25) | ~p(s(bool,X25)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | ~p(s(bool,X24))) ) | $spl24),
  inference(superposition,[],[f2011,f345])).
fof(f7434,plain,(
  ( ! [X19,X18] : (~p(s(bool,X19)) | ~p(s(bool,X18)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))))) ) | $spl24),
  inference(subsumption_resolution,[],[f7410,f341])).
fof(f7410,plain,(
  ( ! [X19,X18] : (s(bool,X18) != s(bool,X19) | ~p(s(bool,X19)) | ~p(s(bool,X18)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))))) ) | $spl24),
  inference(superposition,[],[f2011,f341])).
fof(f7433,plain,(
  ( ! [X17,X16] : (~p(s(bool,X17)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | ~p(s(bool,X16))) ) | $spl24),
  inference(subsumption_resolution,[],[f7409,f341])).
fof(f7409,plain,(
  ( ! [X17,X16] : (s(bool,X16) != s(bool,X17) | ~p(s(bool,X17)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | ~p(s(bool,X16))) ) | $spl24),
  inference(superposition,[],[f2011,f341])).
fof(f2534,plain,(
  s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | $spl47),
  inference(cnf_transformation,[],[f2534_D])).
fof(f2534_D,plain,(
  s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) <=> ~$spl47),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl47])])).
fof(f7431,plain,(
  ( ! [X34] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | ~p(s(bool,X34))) ) | $spl24),
  inference(subsumption_resolution,[],[f7398,f350])).
fof(f7398,plain,(
  ( ! [X34] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | ~p(s(bool,X34)) | s(bool,'F_0') = s(bool,X34)) ) | $spl24),
  inference(superposition,[],[f2011,f359])).
fof(f7430,plain,(
  ( ! [X32] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | ~p(s(bool,X32))) ) | $spl24),
  inference(subsumption_resolution,[],[f7396,f350])).
fof(f7396,plain,(
  ( ! [X32] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | ~p(s(bool,X32)) | s(bool,'F_0') = s(bool,X32)) ) | $spl24),
  inference(superposition,[],[f2011,f358])).
fof(f7424,plain,(
  ( ! [X31] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | ~p(s(bool,X31))) ) | $spl24),
  inference(duplicate_literal_removal,[],[f7395])).
fof(f7395,plain,(
  ( ! [X31] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | ~p(s(bool,X31)) | ~p(s(bool,X31))) ) | $spl24),
  inference(superposition,[],[f2011,f357])).
fof(f7425,plain,(
  ( ! [X30] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | ~p(s(bool,X30))) ) | $spl24),
  inference(duplicate_literal_removal,[],[f7394])).
fof(f7394,plain,(
  ( ! [X30] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | ~p(s(bool,X30)) | ~p(s(bool,X30))) ) | $spl24),
  inference(superposition,[],[f2011,f355])).
fof(f7317,plain,(
  ( ! [X4,X5] : (sP3(X4,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X5) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | ~sP2(X4,X5,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | ~sP0(X5,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X4)) ) | $spl7),
  inference(resolution,[],[f1902,f318])).
fof(f7303,plain,(
  ( ! [X6,X7] : (sP3('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X6,X7) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | ~sP1(X6,X7,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | ~sP0(X7,X6,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | $spl7),
  inference(resolution,[],[f1899,f319])).
fof(f6609,plain,(
  ( ! [X8] : (p(s(bool,X8)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))))) ) | $spl45),
  inference(subsumption_resolution,[],[f6589,f351])).
fof(f6589,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,X8) | p(s(bool,X8)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))))) ) | $spl45),
  inference(superposition,[],[f2519,f338])).
fof(f6610,plain,(
  ( ! [X9] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | p(s(bool,X9))) ) | $spl45),
  inference(subsumption_resolution,[],[f6590,f351])).
fof(f6590,plain,(
  ( ! [X9] : (s(bool,'F_0') != s(bool,X9) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | p(s(bool,X9))) ) | $spl45),
  inference(superposition,[],[f2519,f338])).
fof(f6611,plain,(
  ( ! [X12] : (p(s(bool,X12)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))))) ) | $spl45),
  inference(subsumption_resolution,[],[f6593,f351])).
fof(f6593,plain,(
  ( ! [X12] : (s(bool,'F_0') != s(bool,X12) | p(s(bool,X12)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))))) ) | $spl45),
  inference(superposition,[],[f2519,f342])).
fof(f6612,plain,(
  ( ! [X13] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | p(s(bool,X13))) ) | $spl45),
  inference(subsumption_resolution,[],[f6594,f351])).
fof(f6594,plain,(
  ( ! [X13] : (s(bool,'F_0') != s(bool,X13) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | p(s(bool,X13))) ) | $spl45),
  inference(superposition,[],[f2519,f342])).
fof(f6284,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl61),
  inference(cnf_transformation,[],[f6284_D])).
fof(f6284_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl61),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl61])])).
fof(f6599,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | $spl45),
  inference(superposition,[],[f2519,f355])).
fof(f6600,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | $spl45),
  inference(superposition,[],[f2519,f357])).
fof(f6601,plain,(
  s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | $spl45),
  inference(superposition,[],[f2519,f358])).
fof(f6603,plain,(
  s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | $spl45),
  inference(superposition,[],[f2519,f359])).
fof(f7253,plain,(
  $false | ($spl24 | $spl46)),
  inference(subsumption_resolution,[],[f7156,f7043])).
fof(f7043,plain,(
  ( ! [X1] : (~p(s(bool,X1))) ) | ($spl24 | $spl46)),
  inference(subsumption_resolution,[],[f7042,f357])).
fof(f7042,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,X1) | ~p(s(bool,X1))) ) | ($spl24 | $spl46)),
  inference(forward_demodulation,[],[f2011,f2535])).
fof(f7156,plain,(
  p(s(bool,'F_0')) | ($spl24 | $spl46)),
  inference(resolution,[],[f7043,f531])).
fof(f7252,plain,(
  $false | ($spl24 | $spl46)),
  inference(subsumption_resolution,[],[f7155,f7043])).
fof(f7155,plain,(
  p(s(bool,'F_0')) | ($spl24 | $spl46)),
  inference(resolution,[],[f7043,f529])).
fof(f7251,plain,(
  $false | ($spl24 | $spl46)),
  inference(subsumption_resolution,[],[f7150,f7043])).
fof(f7150,plain,(
  ( ! [X260] : (p(s(bool,X260))) ) | ($spl24 | $spl46)),
  inference(resolution,[],[f7043,f485])).
fof(f7250,plain,(
  $false | ($spl24 | $spl46)),
  inference(subsumption_resolution,[],[f7148,f7043])).
fof(f7148,plain,(
  p(s(bool,'F_0')) | ($spl24 | $spl46)),
  inference(resolution,[],[f7043,f482])).
fof(f7249,plain,(
  $false | ($spl24 | $spl46)),
  inference(subsumption_resolution,[],[f7147,f7043])).
fof(f7147,plain,(
  ( ! [X258] : (p(s(bool,X258))) ) | ($spl24 | $spl46)),
  inference(resolution,[],[f7043,f479])).
fof(f7248,plain,(
  $false | ($spl24 | $spl46)),
  inference(subsumption_resolution,[],[f7145,f7043])).
fof(f7145,plain,(
  ( ! [X256] : (p(s(bool,X256))) ) | ($spl24 | $spl46)),
  inference(resolution,[],[f7043,f476])).
fof(f7157,plain,(
  $false | ($spl24 | $spl46)),
  inference(resolution,[],[f7043,f533])).
fof(f7158,plain,(
  $false | ($spl24 | $spl46)),
  inference(resolution,[],[f7043,f534])).
fof(f7159,plain,(
  $false | ($spl24 | $spl46)),
  inference(resolution,[],[f7043,f537])).
fof(f7167,plain,(
  $false | ($spl0 | $spl24 | $spl46)),
  inference(resolution,[],[f7043,f544])).
fof(f7169,plain,(
  $false | ($spl0 | $spl24 | $spl46)),
  inference(resolution,[],[f7043,f1006])).
fof(f6357,plain,(
  ( ! [X2,X3] : (sP3(X2,X3,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | ~sP2(X2,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X3) | ~sP1(X3,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X2)) ) | $spl6),
  inference(resolution,[],[f571,f317])).
fof(f6293,plain,(
  ( ! [X28] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X28))) ) | $spl22),
  inference(subsumption_resolution,[],[f3551,f351])).
fof(f3551,plain,(
  ( ! [X28] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') != s(bool,X28) | p(s(bool,X28))) ) | $spl22),
  inference(superposition,[],[f1880,f351])).
fof(f6292,plain,(
  ( ! [X29] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X29))) ) | $spl22),
  inference(subsumption_resolution,[],[f3552,f351])).
fof(f3552,plain,(
  ( ! [X29] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') != s(bool,X29) | p(s(bool,X29))) ) | $spl22),
  inference(superposition,[],[f1880,f353])).
fof(f6290,plain,(
  ( ! [X35] : (s(bool,'F_0') != s(bool,X35)) ) | $spl64),
  inference(cnf_transformation,[],[f6290_D])).
fof(f6290_D,plain,(
  ( ! [X35] : (s(bool,'F_0') != s(bool,X35)) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f6291,plain,(
  ( ! [X33] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') != s(bool,X33)) ) | $spl22),
  inference(subsumption_resolution,[],[f3556,f1880])).
fof(f3556,plain,(
  ( ! [X33] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') != s(bool,X33) | s(bool,'T_0') = s(bool,X33)) ) | $spl22),
  inference(superposition,[],[f1880,f358])).
fof(f6289,plain,(
  ( ! [X35] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') != s(bool,X35)) ) | $spl22),
  inference(subsumption_resolution,[],[f3558,f1880])).
fof(f3558,plain,(
  ( ! [X35] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') != s(bool,X35) | s(bool,'T_0') = s(bool,X35)) ) | $spl22),
  inference(superposition,[],[f1880,f359])).
fof(f3577,plain,(
  ( ! [X36] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X36)) ) | $spl22),
  inference(trivial_inequality_removal,[],[f3559])).
fof(f3559,plain,(
  ( ! [X36] : (s(bool,X36) != s(bool,X36) | s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X36)) ) | $spl22),
  inference(superposition,[],[f1880,f359])).
fof(f3576,plain,(
  ( ! [X37] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X37)) ) | $spl22),
  inference(trivial_inequality_removal,[],[f3560])).
fof(f3560,plain,(
  ( ! [X37] : (s(bool,X37) != s(bool,X37) | s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X37)) ) | $spl22),
  inference(superposition,[],[f1880,f358])).
fof(f3575,plain,(
  ( ! [X38] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X38))) ) | $spl22),
  inference(trivial_inequality_removal,[],[f3561])).
fof(f3561,plain,(
  ( ! [X38] : (s(bool,X38) != s(bool,X38) | s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X38))) ) | $spl22),
  inference(superposition,[],[f1880,f357])).
fof(f3574,plain,(
  ( ! [X39] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X39))) ) | $spl22),
  inference(trivial_inequality_removal,[],[f3562])).
fof(f3562,plain,(
  ( ! [X39] : (s(bool,X39) != s(bool,X39) | s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X39))) ) | $spl22),
  inference(superposition,[],[f1880,f355])).
fof(f6287,plain,(
  ( ! [X45,X44] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X45)) | ~p(s(bool,X44))) ) | $spl22),
  inference(subsumption_resolution,[],[f4427,f341])).
fof(f4427,plain,(
  ( ! [X45,X44] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X44) != s(bool,X45) | ~p(s(bool,X45)) | ~p(s(bool,X44))) ) | $spl22),
  inference(superposition,[],[f3521,f355])).
fof(f6286,plain,(
  ( ! [X47,X46] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X47)) | ~p(s(bool,X46))) ) | $spl22),
  inference(subsumption_resolution,[],[f4428,f341])).
fof(f4428,plain,(
  ( ! [X47,X46] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X46) != s(bool,X47) | ~p(s(bool,X47)) | ~p(s(bool,X46))) ) | $spl22),
  inference(superposition,[],[f3521,f357])).
fof(f6285,plain,(
  ( ! [X48,X49] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X48) != s(bool,X49) | ~p(s(bool,X49))) ) | $spl22),
  inference(subsumption_resolution,[],[f4429,f3521])).
fof(f4429,plain,(
  ( ! [X48,X49] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X48) != s(bool,X49) | ~p(s(bool,X49)) | s(bool,'F_0') = s(bool,X48)) ) | $spl22),
  inference(superposition,[],[f3521,f358])).
fof(f6282,plain,(
  ( ! [X52,X53] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X52) != s(bool,X53) | ~p(s(bool,X53))) ) | $spl22),
  inference(subsumption_resolution,[],[f4431,f3521])).
fof(f4431,plain,(
  ( ! [X52,X53] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X52) != s(bool,X53) | ~p(s(bool,X53)) | s(bool,'F_0') = s(bool,X52)) ) | $spl22),
  inference(superposition,[],[f3521,f359])).
fof(f6223,plain,(
  ( ! [X436,X438,X437] : (p(s(bool,X436)) | sP7(X437,X438,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | ~sP6(X437,X438,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X438))) ) | $spl7),
  inference(subsumption_resolution,[],[f5610,f354])).
fof(f5610,plain,(
  ( ! [X436,X438,X437] : (p(s(bool,X436)) | sP7(X437,X438,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | ~sP6(X437,X438,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X438)) | s(bool,'T_0') = s(bool,X436)) ) | $spl7),
  inference(superposition,[],[f390,f5083])).
fof(f5083,plain,(
  ( ! [X35] : (s(bool,X35) = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | s(bool,'T_0') = s(bool,X35)) ) | $spl7),
  inference(trivial_inequality_removal,[],[f5069])).
fof(f5069,plain,(
  ( ! [X35] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,X35) = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | s(bool,'T_0') = s(bool,X35)) ) | $spl7),
  inference(superposition,[],[f5013,f359])).
fof(f5013,plain,(
  ( ! [X23] : (s(bool,'F_0') != s(bool,X23) | s(bool,X23) = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | $spl7),
  inference(subsumption_resolution,[],[f5012,f528])).
fof(f5012,plain,(
  ( ! [X23] : (p(s(bool,'F_0')) | s(bool,'F_0') != s(bool,X23) | s(bool,X23) = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | $spl7),
  inference(forward_demodulation,[],[f4986,f2565])).
fof(f2565,plain,(
  s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | $spl7),
  inference(resolution,[],[f549,f351])).
fof(f4986,plain,(
  ( ! [X23] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | s(bool,'F_0') != s(bool,X23) | s(bool,X23) = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | $spl7),
  inference(resolution,[],[f2812,f2557])).
fof(f2557,plain,(
  ( ! [X48,X49] : (sP3('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X48,X49) | s(bool,X48) = s(bool,X49)) ) | $spl7),
  inference(resolution,[],[f549,f332])).
fof(f2812,plain,(
  ( ! [X171,X170] : (~sP3(X171,X170,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X171)) | s(bool,'F_0') != s(bool,X170)) ) | $spl7),
  inference(superposition,[],[f335,f2565])).
fof(f6221,plain,(
  ( ! [X434,X433,X435] : (p(s(bool,X433)) | sP7(X434,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X435) | ~sP6(X434,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X435) | p(s(bool,X435))) ) | $spl7),
  inference(subsumption_resolution,[],[f5609,f354])).
fof(f5609,plain,(
  ( ! [X434,X433,X435] : (p(s(bool,X433)) | sP7(X434,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X435) | ~sP6(X434,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X435) | p(s(bool,X435)) | s(bool,'T_0') = s(bool,X433)) ) | $spl7),
  inference(superposition,[],[f390,f5083])).
fof(f6207,plain,(
  ( ! [X344,X343,X345] : (p(s(bool,X343)) | sP5(X344,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X345) | ~sP4(X344,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X345) | ~p(s(bool,X345))) ) | $spl7),
  inference(subsumption_resolution,[],[f5579,f354])).
fof(f5579,plain,(
  ( ! [X344,X343,X345] : (p(s(bool,X343)) | sP5(X344,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X345) | ~sP4(X344,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X345) | ~p(s(bool,X345)) | s(bool,'T_0') = s(bool,X343)) ) | $spl7),
  inference(superposition,[],[f373,f5083])).
fof(f6157,plain,(
  ( ! [X158,X159,X160] : (p(s(bool,X158)) | sP3('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X159,X160) | ~sP1(X159,X160,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | ~sP0(X160,X159,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | $spl7),
  inference(subsumption_resolution,[],[f6156,f1899])).
fof(f6156,plain,(
  ( ! [X158,X159,X160] : (p(s(bool,X158)) | sP3('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X159,X160) | ~sP2('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X160,X159) | ~sP1(X159,X160,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | ~sP0(X160,X159,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | $spl7),
  inference(subsumption_resolution,[],[f5496,f354])).
fof(f5496,plain,(
  ( ! [X158,X159,X160] : (p(s(bool,X158)) | sP3('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X159,X160) | ~sP2('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X160,X159) | ~sP1(X159,X160,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | ~sP0(X160,X159,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | s(bool,'T_0') = s(bool,X158)) ) | $spl7),
  inference(superposition,[],[f319,f5083])).
fof(f6154,plain,(
  ( ! [X156,X157,X155] : (p(s(bool,X155)) | sP3(X156,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X157) | ~sP2(X156,X157,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | ~sP0(X157,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X156)) ) | $spl7),
  inference(subsumption_resolution,[],[f6153,f1902])).
fof(f6153,plain,(
  ( ! [X156,X157,X155] : (p(s(bool,X155)) | sP3(X156,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X157) | ~sP2(X156,X157,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | ~sP1('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X157,X156) | ~sP0(X157,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X156)) ) | $spl7),
  inference(subsumption_resolution,[],[f5495,f354])).
fof(f5495,plain,(
  ( ! [X156,X157,X155] : (p(s(bool,X155)) | sP3(X156,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X157) | ~sP2(X156,X157,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | ~sP1('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X157,X156) | ~sP0(X157,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X156) | s(bool,'T_0') = s(bool,X155)) ) | $spl7),
  inference(superposition,[],[f318,f5083])).
fof(f4280,plain,(
  ( ! [X3] : (p(s(bool,X3)) | s(bool,X3) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK19(sK13,'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)),sK12)),s('type.num.num',sK14)))) ) | $spl48),
  inference(cnf_transformation,[],[f4280_D])).
fof(f4280_D,plain,(
  ( ! [X3] : (p(s(bool,X3)) | s(bool,X3) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK19(sK13,'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)),sK12)),s('type.num.num',sK14)))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f4253,plain,(
  ( ! [X163,X162] : (p(s(bool,X162)) | p(s(bool,X163)) | s(bool,X163) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK19(sK13,'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)),sK12)),s('type.num.num',sK14)))) ) | $spl0),
  inference(resolution,[],[f1006,f536])).
fof(f4163,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X23) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK19(sK13,'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)),sK12)),s('type.num.num',sK14))) | s(bool,X24) != s(bool,X25)) ) | $spl0),
  inference(resolution,[],[f1006,f312])).
fof(f4160,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X16) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK19(sK13,'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)),sK12)),s('type.num.num',sK14)))) ) | $spl0),
  inference(resolution,[],[f1006,f311])).
fof(f4159,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK19(sK13,'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)),sK12)),s('type.num.num',sK14)))) ) | $spl0),
  inference(resolution,[],[f1006,f310])).
fof(f4158,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X8) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK19(sK13,'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)),sK12)),s('type.num.num',sK14))) | s(bool,X9) != s(bool,X10)) ) | $spl0),
  inference(resolution,[],[f1006,f308])).
fof(f4156,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK19(sK13,'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)),sK12)),s('type.num.num',sK14)))) ) | $spl0),
  inference(resolution,[],[f1006,f306])).
fof(f2011,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | ~p(s(bool,X1))) ) | $spl24),
  inference(cnf_transformation,[],[f2011_D])).
fof(f2011_D,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | ~p(s(bool,X1))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f2635,plain,(
  ( ! [X161,X160] : (p(s(bool,X160)) | ~p(s(bool,X161)) | s(bool,X161) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | $spl7),
  inference(resolution,[],[f549,f536])).
fof(f2544,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | s(bool,X21) != s(bool,X23)) ) | $spl7),
  inference(resolution,[],[f549,f312])).
fof(f2542,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | $spl7),
  inference(resolution,[],[f549,f311])).
fof(f2540,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | $spl7),
  inference(resolution,[],[f549,f310])).
fof(f2539,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | s(bool,X6) != s(bool,X8)) ) | $spl7),
  inference(resolution,[],[f549,f308])).
fof(f2537,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | $spl7),
  inference(resolution,[],[f549,f306])).
fof(f2528,plain,(
  ( ! [X23,X22] : (p(s(bool,X23)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | p(s(bool,X22))) ) | $spl8),
  inference(subsumption_resolution,[],[f2500,f338])).
fof(f2500,plain,(
  ( ! [X23,X22] : (s(bool,X22) != s(bool,X23) | p(s(bool,X23)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | p(s(bool,X22))) ) | $spl8),
  inference(superposition,[],[f676,f342])).
fof(f2527,plain,(
  ( ! [X21,X20] : (p(s(bool,X21)) | p(s(bool,X20)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))))) ) | $spl8),
  inference(subsumption_resolution,[],[f2499,f338])).
fof(f2499,plain,(
  ( ! [X21,X20] : (s(bool,X20) != s(bool,X21) | p(s(bool,X21)) | p(s(bool,X20)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))))) ) | $spl8),
  inference(superposition,[],[f676,f342])).
fof(f2524,plain,(
  ( ! [X14,X15] : (p(s(bool,X15)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | p(s(bool,X14))) ) | $spl8),
  inference(subsumption_resolution,[],[f2496,f338])).
fof(f2496,plain,(
  ( ! [X14,X15] : (s(bool,X14) != s(bool,X15) | p(s(bool,X15)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | p(s(bool,X14))) ) | $spl8),
  inference(superposition,[],[f676,f338])).
fof(f2523,plain,(
  ( ! [X12,X13] : (p(s(bool,X13)) | p(s(bool,X12)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))))) ) | $spl8),
  inference(subsumption_resolution,[],[f2495,f338])).
fof(f2495,plain,(
  ( ! [X12,X13] : (s(bool,X12) != s(bool,X13) | p(s(bool,X13)) | p(s(bool,X12)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))))) ) | $spl8),
  inference(superposition,[],[f676,f338])).
fof(f2410,plain,(
  ( ! [X2,X3] : (sP3(X2,X3,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | ~sP2(X2,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X3) | ~sP1(X3,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X2)) ) | $spl6),
  inference(resolution,[],[f571,f317])).
fof(f1007,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) ) | $spl20),
  inference(cnf_transformation,[],[f1007_D])).
fof(f1007_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f2376,plain,(
  ( ! [X161,X160] : (p(s(bool,X160)) | p(s(bool,X161)) | s(bool,X161) != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) ) | $spl0),
  inference(resolution,[],[f544,f536])).
fof(f2286,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | s(bool,X22) != s(bool,X23)) ) | $spl0),
  inference(resolution,[],[f544,f312])).
fof(f2283,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) ) | $spl0),
  inference(resolution,[],[f544,f311])).
fof(f2282,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) ) | $spl0),
  inference(resolution,[],[f544,f310])).
fof(f2281,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | s(bool,X7) != s(bool,X8)) ) | $spl0),
  inference(resolution,[],[f544,f308])).
fof(f2279,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) ) | $spl0),
  inference(resolution,[],[f544,f306])).
fof(f2210,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2139,f680])).
fof(f2139,plain,(
  ~p(s(bool,'T_0')) | $spl16),
  inference(resolution,[],[f680,f532])).
fof(f2209,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2138,f680])).
fof(f2138,plain,(
  ~p(s(bool,'T_0')) | $spl16),
  inference(resolution,[],[f680,f530])).
fof(f2208,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2135,f680])).
fof(f2135,plain,(
  ( ! [X268] : (~p(s(bool,X268))) ) | $spl16),
  inference(resolution,[],[f680,f524])).
fof(f2207,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2129,f680])).
fof(f2129,plain,(
  ( ! [X262] : (~p(s(bool,X262))) ) | $spl16),
  inference(resolution,[],[f680,f474])).
fof(f2137,plain,(
  $false | $spl16),
  inference(resolution,[],[f680,f528])).
fof(f1984,plain,(
  ( ! [X161,X160] : (p(s(bool,X160)) | ~p(s(bool,X161)) | s(bool,X161) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | $spl7),
  inference(resolution,[],[f549,f536])).
fof(f1893,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | s(bool,X21) != s(bool,X23)) ) | $spl7),
  inference(resolution,[],[f549,f312])).
fof(f1891,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | $spl7),
  inference(resolution,[],[f549,f311])).
fof(f1889,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | $spl7),
  inference(resolution,[],[f549,f310])).
fof(f1888,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | s(bool,X6) != s(bool,X8)) ) | $spl7),
  inference(resolution,[],[f549,f308])).
fof(f1886,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | $spl7),
  inference(resolution,[],[f549,f306])).
fof(f1880,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | s(bool,'F_0') != s(bool,X2)) ) | $spl22),
  inference(cnf_transformation,[],[f1880_D])).
fof(f1880_D,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | s(bool,'F_0') != s(bool,X2)) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f1879,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) ) | $spl6),
  inference(duplicate_literal_removal,[],[f1873])).
fof(f1873,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) ) | $spl6),
  inference(resolution,[],[f1068,f335])).
fof(f1068,plain,(
  ( ! [X4,X5] : (sP3(X5,X4,'F_0') | s(bool,'T_0') != s(bool,X4) | p(s(bool,X5))) ) | $spl6),
  inference(superposition,[],[f1060,f332])).
fof(f1655,plain,(
  $false | ($spl2 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1654,f534])).
fof(f1654,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f1651,f582])).
fof(f582,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | $spl6),
  inference(resolution,[],[f550,f355])).
fof(f1651,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | ($spl2 | $spl4)),
  inference(trivial_inequality_removal,[],[f1650])).
fof(f1650,plain,(
  s(sK12,sK13) != s(sK12,sK13) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | ($spl2 | $spl4)),
  inference(superposition,[],[f546,f548])).
fof(f1189,plain,(
  ( ! [X161,X162] : (p(s(bool,X161)) | ~p(s(bool,X162)) | s(bool,X162) != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) ) | $spl1),
  inference(resolution,[],[f545,f536])).
fof(f1098,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | s(bool,X22) != s(bool,X24)) ) | $spl1),
  inference(resolution,[],[f545,f312])).
fof(f1096,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) ) | $spl1),
  inference(resolution,[],[f545,f311])).
fof(f1094,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) ) | $spl1),
  inference(resolution,[],[f545,f310])).
fof(f1093,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | s(bool,X7) != s(bool,X9)) ) | $spl1),
  inference(resolution,[],[f545,f308])).
fof(f1091,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) ) | $spl1),
  inference(resolution,[],[f545,f306])).
fof(f979,plain,(
  ( ! [X161,X160] : (p(s(bool,X160)) | p(s(bool,X161)) | s(bool,X161) != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) ) | $spl0),
  inference(resolution,[],[f544,f536])).
fof(f889,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | s(bool,X22) != s(bool,X23)) ) | $spl0),
  inference(resolution,[],[f544,f312])).
fof(f886,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) ) | $spl0),
  inference(resolution,[],[f544,f311])).
fof(f885,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) ) | $spl0),
  inference(resolution,[],[f544,f310])).
fof(f884,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | s(bool,X7) != s(bool,X8)) ) | $spl0),
  inference(resolution,[],[f544,f308])).
fof(f882,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) ) | $spl0),
  inference(resolution,[],[f544,f306])).
fof(f809,plain,(
  ( ! [X161,X162] : (p(s(bool,X161)) | ~p(s(bool,X162)) | s(bool,X162) != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) ) | $spl1),
  inference(resolution,[],[f545,f536])).
fof(f718,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | s(bool,X22) != s(bool,X24)) ) | $spl1),
  inference(resolution,[],[f545,f312])).
fof(f716,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) ) | $spl1),
  inference(resolution,[],[f545,f311])).
fof(f714,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) ) | $spl1),
  inference(resolution,[],[f545,f310])).
fof(f713,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))) | s(bool,X7) != s(bool,X9)) ) | $spl1),
  inference(resolution,[],[f545,f308])).
fof(f711,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) ) | $spl1),
  inference(resolution,[],[f545,f306])).
fof(f676,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X3))) ) | $spl8),
  inference(cnf_transformation,[],[f676_D])).
fof(f676_D,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X3))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f649,plain,(
  ( ! [X163,X162] : (p(s(bool,X162)) | p(s(bool,X163)) | s(bool,X163) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | $spl6),
  inference(resolution,[],[f550,f536])).
fof(f559,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X23) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | s(bool,X24) != s(bool,X25)) ) | $spl6),
  inference(resolution,[],[f550,f312])).
fof(f556,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X16) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | $spl6),
  inference(resolution,[],[f550,f311])).
fof(f555,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | $spl6),
  inference(resolution,[],[f550,f310])).
fof(f554,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X8) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | s(bool,X9) != s(bool,X10)) ) | $spl6),
  inference(resolution,[],[f550,f308])).
fof(f552,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | $spl6),
  inference(resolution,[],[f550,f306])).
fof(f544,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) | $spl0),
  inference(cnf_transformation,[],[f544_D])).
fof(f544_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f302,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))))),
  inference(cnf_transformation,[],[f194])).
fof(f194,plain,(
  (p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) | (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) & s(sK12,sK13) = s(sK12,happ(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK16))))) & (~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) | ! [X4] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num',sK14)))) | s(sK12,sK13) != s(sK12,happ(s(fun('type.num.num',sK12),sK15),s('type.num.num',X4)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13,sK14,sK15,sK16])],[f193])).
fof(f193,plain,(
  ? [X0,X1,X2,X3] : ((p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.GENLIST_2'(s(fun('type.num.num',X0),X3),s('type.num.num',X2)))))))) | ? [X5] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num',X2)))) & s(X0,X1) = s(X0,happ(s(fun('type.num.num',X0),X3),s('type.num.num',X5))))) & (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.GENLIST_2'(s(fun('type.num.num',X0),X3),s('type.num.num',X2)))))))) | ! [X4] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num',X2)))) | s(X0,X1) != s(X0,happ(s(fun('type.num.num',X0),X3),s('type.num.num',X4))))))),
  inference(rectify,[],[f192])).
fof(f192,plain,(
  ? [X0,X1,X2,X3] : ((p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.GENLIST_2'(s(fun('type.num.num',X0),X3),s('type.num.num',X2)))))))) | ? [X4] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num',X2)))) & s(X0,X1) = s(X0,happ(s(fun('type.num.num',X0),X3),s('type.num.num',X4))))) & (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.GENLIST_2'(s(fun('type.num.num',X0),X3),s('type.num.num',X2)))))))) | ! [X4] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num',X2)))) | s(X0,X1) != s(X0,happ(s(fun('type.num.num',X0),X3),s('type.num.num',X4))))))),
  inference(nnf_transformation,[],[f144])).
fof(f144,plain,(
  ? [X0,X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.GENLIST_2'(s(fun('type.num.num',X0),X3),s('type.num.num',X2)))))))) <~> ? [X4] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num',X2)))) & s(X0,X1) = s(X0,happ(s(fun('type.num.num',X0),X3),s('type.num.num',X4)))))),
  inference(ennf_transformation,[],[f63])).
fof(f63,plain,(
  ~! [X0,X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.GENLIST_2'(s(fun('type.num.num',X0),X3),s('type.num.num',X2)))))))) <=> ? [X4] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num',X2)))) & s(X0,X1) = s(X0,happ(s(fun('type.num.num',X0),X3),s('type.num.num',X4)))))),
  inference(flattening,[],[f62])).
fof(f62,plain,(
  ~! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.GENLIST_2'(s(fun('type.num.num',X0),X3),s('type.num.num',X2)))))))) <=> ? [X4] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num',X2)))) & s(X0,X1) = s(X0,happ(s(fun('type.num.num',X0),X3),s('type.num.num',X4)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : (p(s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X6),'const.list.GENLIST_2'(s(fun('type.num.num',X6),X9),s('type.num.num',X8)))))))) <=> ? [X10] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X10),s('type.num.num',X8)))) & s(X6,X7) = s(X6,happ(s(fun('type.num.num',X6),X9),s('type.num.num',X10)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : (p(s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X6),'const.list.GENLIST_2'(s(fun('type.num.num',X6),X9),s('type.num.num',X8)))))))) <=> ? [X10] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X10),s('type.num.num',X8)))) & s(X6,X7) = s(X6,happ(s(fun('type.num.num',X6),X9),s('type.num.num',X10)))))),
  file('thm.list.MEM_GENLIST','thm.list.MEM_GENLIST')).
fof(f548,plain,(
  s(sK12,sK13) = s(sK12,happ(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK16))) | $spl4),
  inference(cnf_transformation,[],[f548_D])).
fof(f548_D,plain,(
  s(sK12,sK13) = s(sK12,happ(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK16))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f303,plain,(
  s(sK12,sK13) = s(sK12,happ(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK16))) | p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))))),
  inference(cnf_transformation,[],[f194])).
fof(f545,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) | $spl1),
  inference(cnf_transformation,[],[f545_D])).
fof(f545_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14)))))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f546,plain,(
  ( ! [X4] : (s(sK12,sK13) != s(sK12,happ(s(fun('type.num.num',sK12),sK15),s('type.num.num',X4))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num',sK14))))) ) | $spl2),
  inference(cnf_transformation,[],[f546_D])).
fof(f546_D,plain,(
  ( ! [X4] : (s(sK12,sK13) != s(sK12,happ(s(fun('type.num.num',sK12),sK15),s('type.num.num',X4))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num',sK14))))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f304,plain,(
  ( ! [X4] : (s(sK12,sK13) != s(sK12,happ(s(fun('type.num.num',sK12),sK15),s('type.num.num',X4))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num',sK14)))) | ~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.GENLIST_2'(s(fun('type.num.num',sK12),sK15),s('type.num.num',sK14))))))))) )),
  inference(cnf_transformation,[],[f194])).
% SZS output end Proof for thm.list
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 1036
% Passive clauses: 6445
% Generated clauses: 33800
% Final active clauses: 648
% Final passive clauses: 3286
% Input formulas: 60
% Initial clauses: 234
% 
% Selected by SInE selection: 58
% SInE iterations: 1
% 
% Duplicate literals: 1657
% Trivial inequalities: 134
% Fw subsumption resolutions: 5070
% Fw demodulations: 2497
% Bw demodulations: 14
% 
% Simple tautologies: 2318
% Equational tautologies: 135
% Forward subsumptions: 15005
% 
% Binary resolution: 17717
% Factoring: 106
% Forward superposition: 2554
% Backward superposition: 3511
% Self superposition: 6
% Equality factoring: 6
% Equality resolution: 16
% 
% Split clauses: 1200
% Split components: 2648
% SAT solver clauses: 1342
% SAT solver unit clauses: 18
% SAT solver binary clauses: 38
% 
% Sat splits: 1200
% Sat splitting refutations: 54
% 
% Memory used [KB]: 9338
% Time elapsed: 0.599 s
% ------------------------------
% ------------------------------
% Success in time 12.347 s
