Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.real_sigma
% (95093)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 576
% Passive clauses: 8021
% Generated clauses: 16545
% Final active clauses: 572
% Final passive clauses: 7127
% Input formulas: 49
% Initial clauses: 192
% 
% Duplicate literals: 273
% Trivial inequalities: 14
% Fw subsumption resolutions: 1236
% Fw demodulations: 769
% Bw demodulations: 3
% 
% Simple tautologies: 206
% Equational tautologies: 56
% Forward subsumptions: 5705
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 7933
% Factoring: 1003
% Forward superposition: 2717
% Backward superposition: 2212
% Self superposition: 52
% Equality factoring: 26
% Equality resolution: 3
% 
% Split clauses: 241
% Split components: 488
% SAT solver clauses: 482
% SAT solver unit clauses: 17
% SAT solver binary clauses: 115
% 
% Sat splits: 241
% 
% TWLsolver clauses: 482
% TWLsolver calls for satisfiability: 310
% 
% Memory used [KB]: 10618
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.real_sigma
% (95114)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 338
% Passive clauses: 2626
% Generated clauses: 12106
% Final active clauses: 280
% Final passive clauses: 1178
% Input formulas: 49
% Initial clauses: 192
% 
% Duplicate literals: 93
% Trivial inequalities: 20
% Fw subsumption resolutions: 702
% Fw demodulations: 1618
% Bw demodulations: 1095
% Condensations: 52
% Global subsumptions: 14
% 
% Simple tautologies: 219
% Equational tautologies: 7
% Forward subsumptions: 4294
% Bw demodulations to eq. taut.: 68
% 
% Binary resolution: 2030
% Factoring: 72
% Forward superposition: 1167
% Backward superposition: 2045
% Self superposition: 36
% Equality factoring: 2
% Equality resolution: 4
% 
% Split clauses: 55
% Split components: 110
% SAT solver clauses: 2667
% SAT solver unit clauses: 216
% SAT solver binary clauses: 78
% SAT solver learnt clauses: 28
% SAT solver learnt literals: 12
% 
% Sat splits: 55
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 2611
% TWLsolver calls for satisfiability: 1924
% 
% Memory used [KB]: 14456
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.real_sigma
% (95126)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1508
% Passive clauses: 17208
% Generated clauses: 108961
% Final active clauses: 943
% Final passive clauses: 598
% Input formulas: 49
% Initial clauses: 192
% Discarded non-redundant clauses: 30026
% 
% Duplicate literals: 1017
% Trivial inequalities: 62
% Fw subsumption resolutions: 3611
% Fw demodulations: 23330
% Bw demodulations: 2656
% Condensations: 2358
% Global subsumptions: 239
% 
% Simple tautologies: 86
% Equational tautologies: 268
% Forward subsumptions: 57315
% Fw demodulations to eq. taut.: 78
% Bw demodulations to eq. taut.: 250
% 
% Unit resulting resolution: 9156
% Factoring: 4
% Forward superposition: 36559
% Backward superposition: 29031
% Self superposition: 190
% Equality factoring: 62
% Equality resolution: 2
% 
% Split clauses: 272
% Split components: 549
% SAT solver clauses: 18861
% SAT solver unit clauses: 1334
% SAT solver binary clauses: 3533
% 
% Sat splits: 272
% Sat splitting refutations: 867
% 
% Memory used [KB]: 120637
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.real_sigma
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.real_sigma
% SZS output start Proof for thm.real_sigma
fof(f2026,plain,(
  $false),
  inference(subsumption_resolution,[],[f1383,f698])).
fof(f698,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK17,sK16,sK15,sK14)))))),
  inference(subsumption_resolution,[],[f695,f239])).
fof(f239,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK14,bool),sK15))))),
  inference(cnf_transformation,[],[f156])).
fof(f156,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK14,bool),sK15)))) & ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(sK14,X4),s(fun(sK14,bool),sK15)))) | p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,X4))))) & s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK14,'type.realax.real'),sK17),s(fun(sK14,bool),sK15))) != s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK14,'type.realax.real'),'_dst_x0x1x1x0_4'(s(fun(bool,fun('type.realax.real',fun('type.realax.real','type.realax.real'))),'const.bool.COND_0'),s(fun(sK14,bool),sK16),s(fun(sK14,'type.realax.real'),sK17),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(fun(sK14,bool),sK15)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK14,sK15,sK16,sK17])],[f112])).
fof(f112,plain,(
  ? [X0,X1,X2,X3] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) & ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) & s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X3),s(fun(X0,bool),X1))) != s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),'_dst_x0x1x1x0_4'(s(fun(bool,fun('type.realax.real',fun('type.realax.real','type.realax.real'))),'const.bool.COND_0'),s(fun(X0,bool),X2),s(fun(X0,'type.realax.real'),X3),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(fun(X0,bool),X1))))),
  inference(flattening,[],[f111])).
fof(f111,plain,(
  ? [X0,X1,X2,X3] : ((p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) & ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X3),s(fun(X0,bool),X1))) != s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),'_dst_x0x1x1x0_4'(s(fun(bool,fun('type.realax.real',fun('type.realax.real','type.realax.real'))),'const.bool.COND_0'),s(fun(X0,bool),X2),s(fun(X0,'type.realax.real'),X3),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(fun(X0,bool),X1))))),
  inference(ennf_transformation,[],[f52])).
fof(f52,plain,(
  ~! [X0,X1,X2,X3] : ((p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) & ! [X4] : (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) => p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) => s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X3),s(fun(X0,bool),X1))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),'_dst_x0x1x1x0_4'(s(fun(bool,fun('type.realax.real',fun('type.realax.real','type.realax.real'))),'const.bool.COND_0'),s(fun(X0,bool),X2),s(fun(X0,'type.realax.real'),X3),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(fun(X0,bool),X1))))),
  inference(flattening,[],[f51])).
fof(f51,plain,(
  ~! [X0] : ! [X1,X2,X3] : ((p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) & ! [X4] : (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) => p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) => s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X3),s(fun(X0,bool),X1))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),'_dst_x0x1x1x0_4'(s(fun(bool,fun('type.realax.real',fun('type.realax.real','type.realax.real'))),'const.bool.COND_0'),s(fun(X0,bool),X2),s(fun(X0,'type.realax.real'),X3),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(fun(X0,bool),X1))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : ((p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X7)))) & ! [X10] : (p(s(bool,'const.bool.IN_2'(s(X6,X10),s(fun(X6,bool),X7)))) => p(s(bool,happ(s(fun(X6,bool),X8),s(X6,X10)))))) => s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X6,'type.realax.real'),X9),s(fun(X6,bool),X7))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X6,'type.realax.real'),'_dst_x0x1x1x0_4'(s(fun(bool,fun('type.realax.real',fun('type.realax.real','type.realax.real'))),'const.bool.COND_0'),s(fun(X6,bool),X8),s(fun(X6,'type.realax.real'),X9),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(fun(X6,bool),X7))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : ((p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X7)))) & ! [X10] : (p(s(bool,'const.bool.IN_2'(s(X6,X10),s(fun(X6,bool),X7)))) => p(s(bool,happ(s(fun(X6,bool),X8),s(X6,X10)))))) => s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X6,'type.realax.real'),X9),s(fun(X6,bool),X7))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X6,'type.realax.real'),'_dst_x0x1x1x0_4'(s(fun(bool,fun('type.realax.real',fun('type.realax.real','type.realax.real'))),'const.bool.COND_0'),s(fun(X6,bool),X8),s(fun(X6,'type.realax.real'),X9),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(fun(X6,bool),X7))))),
  file('thm.real_sigma.REAL_SUM_IMAGE_IF_ELIM','thm.real_sigma.REAL_SUM_IMAGE_IF_ELIM')).
fof(f695,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK17,sK16,sK15,sK14))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK14,bool),sK15))))),
  inference(trivial_inequality_removal,[],[f694])).
fof(f694,plain,(
  s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK14,'type.realax.real'),sK17),s(fun(sK14,bool),sK15))) != s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK14,'type.realax.real'),sK17),s(fun(sK14,bool),sK15))) | ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK17,sK16,sK15,sK14))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK14,bool),sK15))))),
  inference(superposition,[],[f241,f285])).
fof(f285,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X3),s(fun(X0,bool),X1))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),'_dst_x0x1x1x0_4'(s(fun(bool,fun('type.realax.real',fun('type.realax.real','type.realax.real'))),'const.bool.COND_0'),s(fun(X0,bool),X2),s(fun(X0,'type.realax.real'),X3),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(fun(X0,bool),X1))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK18(X3,X2,X1,X0))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f175])).
fof(f175,plain,(
  ! [X0,X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) | ! [X2,X3] : ((p(s(bool,'const.bool.IN_2'(s(X0,sK18(X3,X2,X1,X0)),s(fun(X0,bool),X1)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK18(X3,X2,X1,X0)))))) | s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X3),s(fun(X0,bool),X1))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),'_dst_x0x1x1x0_4'(s(fun(bool,fun('type.realax.real',fun('type.realax.real','type.realax.real'))),'const.bool.COND_0'),s(fun(X0,bool),X2),s(fun(X0,'type.realax.real'),X3),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(fun(X0,bool),X1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK18])],[f115])).
fof(f115,plain,(
  ! [X0,X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) | ! [X2,X3] : (? [X4] : (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X3),s(fun(X0,bool),X1))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),'_dst_x0x1x1x0_4'(s(fun(bool,fun('type.realax.real',fun('type.realax.real','type.realax.real'))),'const.bool.COND_0'),s(fun(X0,bool),X2),s(fun(X0,'type.realax.real'),X3),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(fun(X0,bool),X1)))))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0,X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) => ! [X2,X3] : (! [X4] : (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) => p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) => s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X3),s(fun(X0,bool),X1))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),'_dst_x0x1x1x0_4'(s(fun(bool,fun('type.realax.real',fun('type.realax.real','type.realax.real'))),'const.bool.COND_0'),s(fun(X0,bool),X2),s(fun(X0,'type.realax.real'),X3),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) => ! [X2,X3] : (! [X4] : (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) => p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) => s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X3),s(fun(X0,bool),X1))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),'_dst_x0x1x1x0_4'(s(fun(bool,fun('type.realax.real',fun('type.realax.real','type.realax.real'))),'const.bool.COND_0'),s(fun(X0,bool),X2),s(fun(X0,'type.realax.real'),X3),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X6] : ! [X7] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X7)))) => ! [X8,X9] : (! [X10] : (p(s(bool,'const.bool.IN_2'(s(X6,X10),s(fun(X6,bool),X7)))) => p(s(bool,happ(s(fun(X6,bool),X8),s(X6,X10))))) => s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X6,'type.realax.real'),X9),s(fun(X6,bool),X7))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X6,'type.realax.real'),'_dst_x0x1x1x0_4'(s(fun(bool,fun('type.realax.real',fun('type.realax.real','type.realax.real'))),'const.bool.COND_0'),s(fun(X6,bool),X8),s(fun(X6,'type.realax.real'),X9),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(fun(X6,bool),X7)))))),
  file('thm.real_sigma.REAL_SUM_IMAGE_IF_ELIM','thm.real_sigma.real_sigma_7_prove_ltb')).
fof(f241,plain,(
  s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK14,'type.realax.real'),sK17),s(fun(sK14,bool),sK15))) != s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK14,'type.realax.real'),'_dst_x0x1x1x0_4'(s(fun(bool,fun('type.realax.real',fun('type.realax.real','type.realax.real'))),'const.bool.COND_0'),s(fun(sK14,bool),sK16),s(fun(sK14,'type.realax.real'),sK17),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(fun(sK14,bool),sK15)))),
  inference(cnf_transformation,[],[f156])).
fof(f1383,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK17,sK16,sK15,sK14)))))),
  inference(resolution,[],[f697,f240])).
fof(f240,plain,(
  ( ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(sK14,X4),s(fun(sK14,bool),sK15)))) | p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,X4))))) )),
  inference(cnf_transformation,[],[f156])).
fof(f697,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK14,sK18(sK17,sK16,sK15,sK14)),s(fun(sK14,bool),sK15))))),
  inference(subsumption_resolution,[],[f696,f239])).
fof(f696,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK14,sK18(sK17,sK16,sK15,sK14)),s(fun(sK14,bool),sK15)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK14,bool),sK15))))),
  inference(trivial_inequality_removal,[],[f693])).
fof(f693,plain,(
  s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK14,'type.realax.real'),sK17),s(fun(sK14,bool),sK15))) != s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK14,'type.realax.real'),sK17),s(fun(sK14,bool),sK15))) | p(s(bool,'const.bool.IN_2'(s(sK14,sK18(sK17,sK16,sK15,sK14)),s(fun(sK14,bool),sK15)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK14,bool),sK15))))),
  inference(superposition,[],[f241,f284])).
fof(f284,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X3),s(fun(X0,bool),X1))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),'_dst_x0x1x1x0_4'(s(fun(bool,fun('type.realax.real',fun('type.realax.real','type.realax.real'))),'const.bool.COND_0'),s(fun(X0,bool),X2),s(fun(X0,'type.realax.real'),X3),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(fun(X0,bool),X1))) | p(s(bool,'const.bool.IN_2'(s(X0,sK18(X3,X2,X1,X0)),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f175])).
% SZS output end Proof for thm.real_sigma
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 216
% Passive clauses: 563
% Generated clauses: 1598
% Final active clauses: 214
% Final passive clauses: 414
% Input formulas: 49
% Initial clauses: 183
% 
% Selected by SInE selection: 40
% SInE iterations: 1
% 
% Duplicate literals: 73
% Trivial inequalities: 16
% Fw subsumption resolutions: 219
% Fw demodulations: 103
% Bw demodulations: 2
% 
% Simple tautologies: 174
% Equational tautologies: 2
% Forward subsumptions: 566
% 
% Binary resolution: 773
% Forward superposition: 183
% Backward superposition: 230
% Equality resolution: 2
% 
% Split clauses: 1
% Split components: 2
% SAT solver clauses: 2
% SAT solver binary clauses: 1
% 
% Sat splits: 1
% 
% Memory used [KB]: 1663
% Time elapsed: 0.028 s
% ------------------------------
% ------------------------------
% Success in time 11.735 s
