Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.util_prob
% (109747)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 201
% Passive clauses: 3218
% Generated clauses: 23844
% Final active clauses: 165
% Final passive clauses: 2787
% Input formulas: 16
% Initial clauses: 36
% 
% Duplicate literals: 728
% Trivial inequalities: 484
% Fw subsumption resolutions: 2941
% Fw demodulations: 3498
% Bw demodulations: 25
% 
% Simple tautologies: 384
% Equational tautologies: 26
% Forward subsumptions: 10889
% Bw demodulations to eq. taut.: 4
% 
% Binary resolution: 596
% Forward superposition: 8707
% Backward superposition: 6339
% Self superposition: 138
% Equality factoring: 49
% Equality resolution: 51
% 
% Split clauses: 816
% Split components: 1641
% SAT solver clauses: 1376
% SAT solver unit clauses: 13
% SAT solver binary clauses: 220
% 
% Sat splits: 816
% Sat splitting refutations: 6
% 
% TWLsolver clauses: 1376
% TWLsolver calls for satisfiability: 180
% 
% Memory used [KB]: 16247
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.util_prob
% (109756)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 155
% Passive clauses: 1051
% Generated clauses: 16349
% Final active clauses: 154
% Final passive clauses: 868
% Input formulas: 16
% Initial clauses: 36
% 
% Duplicate literals: 773
% Trivial inequalities: 398
% Fw subsumption resolutions: 1228
% Fw demodulations: 1766
% Bw demodulations: 28
% Condensations: 1348
% 
% Simple tautologies: 250
% Equational tautologies: 27
% Forward subsumptions: 6464
% 
% Binary resolution: 283
% Forward superposition: 5381
% Backward superposition: 3859
% Self superposition: 108
% Equality factoring: 47
% Equality resolution: 42
% 
% Split clauses: 62
% Split components: 132
% SAT solver clauses: 1606
% SAT solver unit clauses: 30
% SAT solver binary clauses: 15
% 
% Sat splits: 62
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 1606
% TWLsolver calls for satisfiability: 294
% 
% Memory used [KB]: 15735
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.util_prob
% (109772)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 780
% Passive clauses: 8826
% Generated clauses: 184223
% Final active clauses: 103
% Final passive clauses: 1022
% Input formulas: 16
% Initial clauses: 36
% 
% Duplicate literals: 5088
% Trivial inequalities: 908
% Fw subsumption resolutions: 15348
% Fw demodulations: 30743
% Bw demodulations: 6934
% Condensations: 20418
% Global subsumptions: 1034
% 
% Simple tautologies: 1337
% Equational tautologies: 635
% Forward subsumptions: 97934
% Fw demodulations to eq. taut.: 252
% Bw demodulations to eq. taut.: 265
% 
% Unit resulting resolution: 6560
% Forward superposition: 54237
% Backward superposition: 30831
% Self superposition: 760
% Equality factoring: 115
% Equality resolution: 187
% 
% Split clauses: 307
% Split components: 625
% SAT solver clauses: 11830
% SAT solver unit clauses: 173
% SAT solver binary clauses: 2129
% 
% Sat splits: 307
% Sat splitting refutations: 1354
% 
% Memory used [KB]: 99273
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.util_prob
% (110051)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1272
% Passive clauses: 25916
% Generated clauses: 75346
% Final active clauses: 1249
% Final passive clauses: 23779
% Input formulas: 16
% Initial clauses: 34
% 
% Selected by SInE selection: 14
% SInE iterations: 1
% 
% Duplicate literals: 6793
% Trivial inequalities: 5663
% Fw subsumption resolutions: 1385
% Fw demodulations: 8543
% Bw demodulations: 8
% 
% Simple tautologies: 895
% Equational tautologies: 1
% Forward subsumptions: 28375
% 
% Binary resolution: 15
% Forward superposition: 33036
% Backward superposition: 20551
% Self superposition: 96
% Equality factoring: 64
% Equality resolution: 609
% 
% Split clauses: 56
% Split components: 122
% SAT solver clauses: 63
% SAT solver unit clauses: 1
% SAT solver binary clauses: 14
% 
% Sat splits: 56
% Sat splitting refutations: 6
% 
% Memory used [KB]: 40041
% Time elapsed: 4.0000 s
% ------------------------------
% ------------------------------
% remaining time: 142 next slice time: 6
lrs+1002_2:3_br=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:sas=minisat:stl=30:sac=on:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none:sp=reverse_arity:urr=on_4 on thm.util_prob
% (110275)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 81
% Passive clauses: 1354
% Generated clauses: 11727
% Final active clauses: 70
% Final passive clauses: 1156
% Input formulas: 16
% Initial clauses: 36
% 
% Duplicate literals: 310
% Trivial inequalities: 297
% Fw subsumption resolutions: 1421
% Fw demodulations: 1078
% Bw demodulations: 79
% Global subsumptions: 3
% 
% Simple tautologies: 379
% Equational tautologies: 70
% Forward subsumptions: 6692
% Fw demodulations to eq. taut.: 25
% Bw demodulations to eq. taut.: 2
% 
% Unit resulting resolution: 426
% Forward superposition: 3592
% Backward superposition: 4178
% Self superposition: 113
% Equality factoring: 53
% Equality resolution: 20
% 
% Split clauses: 15
% Split components: 34
% SAT solver clauses: 1413
% SAT solver unit clauses: 35
% SAT solver binary clauses: 116
% 
% Sat splits: 15
% Sat splitting refutations: 39
% 
% Memory used [KB]: 10874
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 135 next slice time: 15
lrs+1011_12_bs=on:bsr=unit_only:cond=on:gs=on:gsaa=from_current:gsssp=full:nwc=1.1:sas=minisat:stl=60:sos=all:sac=on:sdd=off:sser=off:sfr=on:ssfp=100000:ssfq=1.2:ssnc=none:sp=reverse_arity:updr=off_14 on thm.util_prob
% (110293)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 321
% Passive clauses: 4067
% Generated clauses: 14713
% Final active clauses: 77
% Final passive clauses: 2022
% Input formulas: 16
% Initial clauses: 36
% 
% Duplicate literals: 267
% Trivial inequalities: 431
% Fw subsumption resolutions: 643
% Bw subsumption resolutions: 266
% Fw demodulations: 1966
% Bw demodulations: 692
% Condensations: 590
% Global subsumptions: 1
% 
% Simple tautologies: 281
% Equational tautologies: 34
% Forward subsumptions: 6184
% Backward subsumptions: 703
% 
% Binary resolution: 259
% Forward superposition: 4750
% Backward superposition: 2370
% Self superposition: 74
% Equality factoring: 32
% Equality resolution: 80
% 
% Split clauses: 7
% Split components: 14
% SAT solver clauses: 4092
% SAT solver unit clauses: 11
% SAT solver binary clauses: 144
% 
% Sat splits: 7
% Sat splitting refutations: 14
% 
% Memory used [KB]: 15351
% Time elapsed: 1.600 s
% ------------------------------
% ------------------------------
% remaining time: 119 next slice time: 9
lrs+10_8:1_bd=preordered:bs=on:ccuc=first:fsr=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:nicw=on:sas=minisat:stl=120:sos=on:sscc=on:sser=off:sfr=on:ssfp=4000:ssfq=1.0:ssnc=none:sp=reverse_arity:urr=on_7 on thm.util_prob
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.util_prob
% SZS output start Proof for thm.util_prob
fof(f9081,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f67,f106,f110,f112,f109,f114,f108,f116,f107,f118,f119,f123,f121,f124,f128,f126,f131,f138,f143,f145,f147,f148,f150,f151,f153,f287,f146,f366,f368,f370,f371,f373,f520,f369,f615,f617,f620,f622,f667,f721,f732,f733,f132,f735,f737,f781,f952,f954,f958,f139,f1123,f1606,f1605,f1604,f1135,f1607,f1608,f1609,f1610,f1611,f1612,f1613,f1614,f1615,f1616,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1634,f1635,f1636,f1637,f1638,f1639,f1640,f1641,f1642,f1643,f1644,f1645,f1646,f1647,f1648,f1649,f1650,f1651,f1652,f1653,f1654,f1655,f1656,f1657,f1658,f1659,f1660,f1661,f1662,f1663,f1664,f1665,f1666,f1667,f1668,f1669,f1670,f1671,f1672,f1673,f1674,f1675,f1676,f1677,f1678,f1679,f1680,f1681,f1682,f1683,f1684,f1685,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1698,f1699,f1700,f1701,f1702,f1703,f1704,f1705,f1706,f1707,f1708,f1709,f1710,f1711,f1712,f1713,f1714,f1715,f1716,f1717,f1718,f1719,f1720,f1721,f1722,f1723,f1724,f1725,f1726,f1727,f1728,f1729,f1730,f1731,f1732,f1733,f1734,f1735,f1736,f1737,f1738,f1739,f1740,f1741,f1742,f1743,f1744,f1745,f1746,f1747,f1748,f1749,f1750,f1751,f1752,f1753,f1754,f1755,f1756,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1766,f1767,f1768,f1769,f1770,f1771,f1772,f1773,f1774,f2039,f2150,f2153,f2154,f2152,f2440,f2450,f142,f2747,f2454,f2750,f2751,f2752,f2753,f2754,f2755,f2756,f2757,f2758,f2759,f2760,f2761,f2762,f2763,f2764,f2765,f2766,f2767,f2768,f2769,f2770,f2771,f2772,f2773,f2774,f2775,f2776,f2777,f2778,f2779,f2780,f2781,f2782,f2783,f2784,f2785,f2786,f2787,f2788,f2789,f2790,f2791,f2084,f836,f3107,f3100,f3099,f3098,f3097,f3096,f3095,f3094,f3083,f3069,f3061,f3060,f3059,f3058,f3053,f3052,f3051,f3050,f3043,f3042,f3041,f3040,f3039,f3038,f3037,f3026,f3018,f3017,f3016,f3015,f3112,f3109,f3168,f3363,f3362,f3361,f3355,f3354,f3353,f3347,f3346,f3345,f3344,f3343,f3342,f3341,f3340,f3339,f3338,f3337,f3336,f3335,f3334,f3333,f3332,f3331,f3330,f3329,f3328,f3327,f3326,f3325,f3324,f3323,f3322,f3321,f3301,f3300,f3299,f3295,f3294,f3293,f3289,f3288,f3287,f3277,f3276,f3275,f3274,f3273,f3272,f3271,f3270,f3269,f3259,f3257,f3255,f3249,f3248,f3247,f3246,f3245,f3244,f3243,f3242,f3241,f3240,f3239,f3238,f3232,f3231,f3230,f3229,f3228,f3227,f3226,f3225,f3224,f3223,f3222,f3221,f3220,f3219,f3218,f3217,f3216,f3215,f3214,f3213,f3212,f3211,f3210,f3209,f3208,f3207,f3206,f3196,f3194,f3192,f3186,f3185,f3184,f3183,f3182,f3181,f3180,f3179,f3178,f3177,f3176,f3175,f3167,f3936,f3935,f3927,f3926,f3925,f3924,f3923,f3922,f3921,f3920,f3919,f3918,f3917,f3916,f3915,f3914,f3913,f3912,f3911,f3910,f3909,f3908,f3895,f3891,f3869,f3868,f3853,f3852,f3851,f3850,f3849,f3848,f3847,f3846,f3840,f3839,f3838,f3837,f3836,f3835,f3834,f3833,f3832,f3831,f3823,f3822,f3821,f3820,f3819,f3818,f3817,f3816,f3815,f3814,f3813,f3812,f3811,f3810,f3809,f3808,f3807,f3806,f3805,f3804,f3791,f3787,f3778,f3777,f3776,f3775,f3774,f3773,f3772,f3771,f3770,f3943,f3769,f3945,f3939,f4021,f4024,f135,f4336,f4335,f4334,f4328,f4327,f4326,f4320,f4319,f4318,f4317,f4316,f4315,f4314,f4313,f4312,f4311,f4310,f4309,f4308,f4307,f4306,f4305,f4304,f4303,f4302,f4301,f4300,f4299,f4298,f4297,f4296,f4295,f4294,f4273,f4272,f4271,f4267,f4266,f4265,f4261,f4260,f4259,f4249,f4248,f4247,f4246,f4245,f4244,f4243,f4242,f4241,f4230,f4228,f4226,f4219,f4218,f4217,f4216,f4215,f4214,f4213,f4212,f4211,f4210,f4209,f4208,f4207,f4206,f4205,f4199,f4198,f4197,f4196,f4195,f4194,f4193,f4192,f4191,f4190,f4189,f4188,f4187,f4186,f4185,f4184,f4183,f4182,f4181,f4180,f4179,f4178,f4177,f4176,f4175,f4174,f4173,f4162,f4160,f4158,f4151,f4150,f4149,f4148,f4147,f4146,f4145,f4144,f4143,f4142,f4141,f4140,f4139,f4138,f4137,f4020,f4429,f3944,f4868,f2441,f4968,f4973,f4977,f4982,f5043,f5114,f5410,f5400,f5399,f5398,f5397,f5396,f5395,f5394,f5393,f5392,f5391,f5390,f5389,f5388,f5387,f5386,f5385,f5369,f5350,f5342,f5341,f5340,f5339,f5332,f5331,f5330,f5329,f5328,f5327,f5317,f5316,f5315,f5314,f5313,f5312,f5311,f5310,f5309,f5308,f5307,f5306,f5305,f5304,f5303,f5302,f5286,f5276,f5275,f5274,f5273,f5272,f5271,f5415,f5270,f5496,f5269,f5498,f5412,f5500,f4686,f1126,f5724,f5733,f5725,f5735,f5741,f5750,f5742,f1125,f5752,f5758,f5760,f5765,f5767,f5776,f5779,f5788,f5791,f5800,f5803,f5812,f5815,f6327,f6328,f6331,f6332,f6333,f6334,f6335,f6376,f6336,f6379,f6357,f6361,f6364,f6365,f6383,f6447,f6448,f6451,f6452,f6453,f6454,f6455,f6500,f6456,f6501,f6479,f6483,f6486,f6487,f6378,f6502,f6580,f6575,f103,f1124,f6626,f6642,f6635,f6648,f6661,f6680,f6671,f6685,f6816,f6842,f6857,f6864,f6867,f6869,f6871,f6873,f6875,f6877,f6879,f6881,f6883,f6885,f6892,f6975,f6977,f6979,f6981,f6983,f6985,f6987,f6989,f6991,f6993,f6995,f6996,f6997,f6998,f6999,f7000,f7001,f7002,f7003,f7004,f104,f102,f101,f6382,f7026,f98,f6377,f7027,f7025,f6381,f7046,f7119,f7118,f7073,f7130,f7129,f7074,f7131,f7094,f7147,f7095,f6380,f7148,f7206,f7957,f8008,f7958,f8009,f8069,f8084,f8070,f8085,f8213,f8264,f8214,f8265,f8917,f8932,f8918,f8933,f9036,f9038,f9041,f9044,f152,f723,f9046,f9050,f9052,f9053,f6858,f9054,f6613,f9055,f6260,f619,f9048,f9056,f133,f9059,f9058,f9061,f9065,f9067,f9068,f6859,f9069,f6612,f9070,f6258,f956,f9063,f9071,f140,f9074,f9073])).
fof(f9073,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))))))) | $spl172),
  inference(cnf_transformation,[],[f9073_D])).
fof(f9073_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))))))) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f9074,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))))))) | $spl86),
  inference(global_subsumption,[],[f67,f5752,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6816,f104,f102,f101,f100,f7026,f98,f96,f7027,f7591,f7593,f7594,f7596,f7598,f7599,f7600,f7667,f7670,f7642,f7682,f7687,f7688,f7689,f7691,f7698,f7699,f7702,f7695,f7703,f7704,f7705,f7706,f7710,f7711,f7712,f7766,f7769,f7718,f7770,f7779,f7780,f7781,f7785,f7786,f7787,f7837,f7838,f7841,f7858,f7859,f7860,f7863,f7864,f7918,f7923,f7924,f7925,f7926,f7928,f7935,f7936,f7939,f7932,f7940,f7941,f7942,f7943,f7947,f7948,f7949,f8003,f8006,f7955,f8007,f8020,f8021,f8022,f8024,f8025,f8026,f8078,f8067,f8083,f8195,f8196,f8198,f8203,f8204,f8205,f8255,f8259,f8277,f8278,f8279,f8325,f8331,f8332,f8333,f8336,f8338,f8397,f8401,f8402,f8403,f8450,f8451,f8453,f8462,f8465,f8457,f8466,f8467,f8470,f8471,f8472,f8474,f8485,f8488,f8478,f8489,f8490,f8493,f8494,f8495,f8496,f8497,f8499,f8500,f8501,f8502,f8504,f8505,f8506,f8507,f8508,f8555,f8558,f8565,f8566,f8567,f8568,f8570,f8571,f8572,f8573,f8575,f8576,f8577,f8578,f8579,f8624,f8637,f8638,f8639,f8640,f8642,f8643,f8644,f8645,f8646,f8647,f8648,f8688,f8689,f8690,f8691,f8693,f8694,f8695,f8696,f8698,f8699,f8700,f8701,f8702,f8703,f8705,f8706,f8707,f8708,f8710,f8711,f8712,f8713,f8714,f8715,f8716,f8717,f8718,f8788,f8760,f8800,f8805,f8807,f8808,f8809,f8810,f8858,f8859,f8861,f8862,f8863,f8864,f8869,f8870,f8871,f8872,f8874,f8931,f8966,f8969,f8936,f8937,f8963,f8989,f8990,f8991,f8992,f8993,f8995,f8996,f8997,f8998,f8999,f9001,f9002,f9003,f9004,f9005,f9022,f9025,f9008,f9009,f9018])).
fof(f9018,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))))))) | $spl86),
  inference(superposition,[],[f6755,f5752])).
fof(f9009,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X1),s(fun(fun(X0,bool),bool),'const.util_prob.DFUNSET_2'(s(fun(X0,bool),X3),s(fun(X0,fun(bool,bool)),'const.combin.K_1'(s(fun(bool,bool),X2)))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK7('const.combin.K_1'(s(fun(bool,bool),X2)),X3,X1,X0,bool))))) )),
  inference(superposition,[],[f6755,f71])).
fof(f71,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.util_prob.FUNSET_DFUNSET',aHL_BOOL_CASES)).
fof(f9008,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(bool,X0),X3),s(bool,'F_0'))),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(fun(bool,X0),X3),s(fun(fun(bool,X0),bool),'const.util_prob.DFUNSET_2'(s(fun(bool,bool),X2),s(fun(bool,fun(X0,bool)),'const.combin.K_1'(s(fun(X0,bool),X1)))))))) | s(bool,'T_0') = s(bool,sK7('const.combin.K_1'(s(fun(X0,bool),X1)),X2,X3,bool,X0))) )),
  inference(superposition,[],[f6755,f71])).
fof(f9025,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(fun(X7,X8),X10),s(fun(fun(X7,X8),bool),'const.util_prob.DFUNSET_2'(s(fun(X7,bool),X13),s(fun(X7,fun(X8,bool)),'const.combin.K_1'(s(fun(X8,bool),X12)))))))) | ~p(s(bool,'const.bool.IN_2'(s(X8,happ(s(fun(X7,X8),X10),s(X7,sK7('const.combin.K_1'(s(fun(X8,bool),X12)),X13,happ(s(fun(X9,fun(X7,X8)),'const.combin.K_1'(s(fun(X7,X8),X10))),s(X9,X11)),X7,X8)))),s(fun(X8,bool),X12))))) )),
  inference(forward_demodulation,[],[f9024,f80])).
fof(f80,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.combin.K_2'(s(X1,X2),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f29])).
fof(f29,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,'const.combin.K_2'(s(X1,X2),s(X0,X3)))),
  inference(flattening,[],[f28])).
fof(f28,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,'const.combin.K_2'(s(X1,X2),s(X0,X3)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X7,X6] : ! [X8,X9] : s(X6,X8) = s(X6,'const.combin.K_2'(s(X6,X8),s(X7,X9)))),
  file('thm.util_prob.FUNSET_DFUNSET','thm.combin.K_THM')).
fof(f9024,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(fun(X7,X8),'const.combin.K_2'(s(fun(X7,X8),X10),s(X9,X11))),s(fun(fun(X7,X8),bool),'const.util_prob.DFUNSET_2'(s(fun(X7,bool),X13),s(fun(X7,fun(X8,bool)),'const.combin.K_1'(s(fun(X8,bool),X12)))))))) | ~p(s(bool,'const.bool.IN_2'(s(X8,happ(s(fun(X7,X8),X10),s(X7,sK7('const.combin.K_1'(s(fun(X8,bool),X12)),X13,happ(s(fun(X9,fun(X7,X8)),'const.combin.K_1'(s(fun(X7,X8),X10))),s(X9,X11)),X7,X8)))),s(fun(X8,bool),X12))))) )),
  inference(forward_demodulation,[],[f9023,f81])).
fof(f81,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.combin.K_2'(s(X0,X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),'const.combin.K_1'(s(X0,X2))),s(X1,X3)))) )),
  inference(cnf_transformation,[],[f31])).
fof(f31,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.combin.K_2'(s(X0,X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),'const.combin.K_1'(s(X0,X2))),s(X1,X3)))),
  inference(flattening,[],[f30])).
fof(f30,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X0,'const.combin.K_2'(s(X0,X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),'const.combin.K_1'(s(X0,X2))),s(X1,X3)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X6,X7] : ! [X20,X21] : s(X6,'const.combin.K_2'(s(X6,X20),s(X7,X21))) = s(X6,happ(s(fun(X7,X6),'const.combin.K_1'(s(X6,X20))),s(X7,X21)))),
  file('thm.util_prob.FUNSET_DFUNSET','ARITY_1_2_const.combin.K')).
fof(f9023,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (~p(s(bool,'const.bool.IN_2'(s(X8,happ(s(fun(X7,X8),X10),s(X7,sK7('const.combin.K_1'(s(fun(X8,bool),X12)),X13,happ(s(fun(X9,fun(X7,X8)),'const.combin.K_1'(s(fun(X7,X8),X10))),s(X9,X11)),X7,X8)))),s(fun(X8,bool),X12)))) | p(s(bool,'const.bool.IN_2'(s(fun(X7,X8),happ(s(fun(X9,fun(X7,X8)),'const.combin.K_1'(s(fun(X7,X8),X10))),s(X9,X11))),s(fun(fun(X7,X8),bool),'const.util_prob.DFUNSET_2'(s(fun(X7,bool),X13),s(fun(X7,fun(X8,bool)),'const.combin.K_1'(s(fun(X8,bool),X12))))))))) )),
  inference(forward_demodulation,[],[f9007,f80])).
fof(f9007,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (~p(s(bool,'const.bool.IN_2'(s(X8,happ(s(fun(X7,X8),'const.combin.K_2'(s(fun(X7,X8),X10),s(X9,X11))),s(X7,sK7('const.combin.K_1'(s(fun(X8,bool),X12)),X13,happ(s(fun(X9,fun(X7,X8)),'const.combin.K_1'(s(fun(X7,X8),X10))),s(X9,X11)),X7,X8)))),s(fun(X8,bool),X12)))) | p(s(bool,'const.bool.IN_2'(s(fun(X7,X8),happ(s(fun(X9,fun(X7,X8)),'const.combin.K_1'(s(fun(X7,X8),X10))),s(X9,X11))),s(fun(fun(X7,X8),bool),'const.util_prob.DFUNSET_2'(s(fun(X7,bool),X13),s(fun(X7,fun(X8,bool)),'const.combin.K_1'(s(fun(X8,bool),X12))))))))) )),
  inference(superposition,[],[f6755,f81])).
fof(f9022,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,X1),X2),s(fun(fun(X0,X1),bool),'const.util_prob.DFUNSET_2'(s(fun(X0,bool),X6),s(fun(X0,fun(X1,bool)),'const.combin.K_1'(s(fun(X1,bool),X5)))))))) | ~p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,sK7('const.combin.K_1'(s(fun(X1,bool),X5)),X6,'const.combin.K_2'(s(fun(X0,X1),X2),s(X3,X4)),X0,X1)))),s(fun(X1,bool),X5))))) )),
  inference(forward_demodulation,[],[f9006,f80])).
fof(f9006,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,sK7('const.combin.K_1'(s(fun(X1,bool),X5)),X6,'const.combin.K_2'(s(fun(X0,X1),X2),s(X3,X4)),X0,X1)))),s(fun(X1,bool),X5)))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,X1),'const.combin.K_2'(s(fun(X0,X1),X2),s(X3,X4))),s(fun(fun(X0,X1),bool),'const.util_prob.DFUNSET_2'(s(fun(X0,bool),X6),s(fun(X0,fun(X1,bool)),'const.combin.K_1'(s(fun(X1,bool),X5))))))))) )),
  inference(superposition,[],[f6755,f80])).
fof(f9005,plain,(
  ( ! [X123,X121,X118,X122,X120,X119] : (p(s(bool,'const.bool.IN_2'(s(fun(X118,fun(bool,X119)),X120),s(fun(fun(X118,fun(bool,X119)),bool),'const.util_prob.DFUNSET_2'(s(fun(X118,bool),X121),s(fun(X118,fun(fun(bool,X119),bool)),'const.combin.K_1'(s(fun(fun(bool,X119),bool),'const.util_prob.DFUNSET_2'(s(fun(bool,bool),X122),s(fun(bool,fun(X119,bool)),X123)))))))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X122)))) | s(bool,'T_0') = s(bool,sK7(X123,X122,happ(s(fun(X118,fun(bool,X119)),X120),s(X118,sK7('const.combin.K_1'(s(fun(fun(bool,X119),bool),'const.util_prob.DFUNSET_2'(s(fun(bool,bool),X122),s(fun(bool,fun(X119,bool)),X123)))),X121,X120,X118,fun(bool,X119)))),bool,X119))) )),
  inference(resolution,[],[f6755,f6662])).
fof(f9004,plain,(
  ( ! [X116,X114,X112,X117,X115,X113] : (p(s(bool,'const.bool.IN_2'(s(fun(X112,fun(bool,X113)),X114),s(fun(fun(X112,fun(bool,X113)),bool),'const.util_prob.DFUNSET_2'(s(fun(X112,bool),X115),s(fun(X112,fun(fun(bool,X113),bool)),'const.combin.K_1'(s(fun(fun(bool,X113),bool),'const.util_prob.FUNSET_2'(s(fun(bool,bool),X116),s(fun(X113,bool),X117)))))))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X116)))) | s(bool,'T_0') = s(bool,sK6(X117,X116,happ(s(fun(X112,fun(bool,X113)),X114),s(X112,sK7('const.combin.K_1'(s(fun(fun(bool,X113),bool),'const.util_prob.FUNSET_2'(s(fun(bool,bool),X116),s(fun(X113,bool),X117)))),X115,X114,X112,fun(bool,X113)))),X113,bool))) )),
  inference(resolution,[],[f6755,f6627])).
fof(f9003,plain,(
  ( ! [X111,X109,X107,X110,X108,X106] : (p(s(bool,'const.bool.IN_2'(s(fun(X106,fun(bool,X107)),X108),s(fun(fun(X106,fun(bool,X107)),bool),'const.util_prob.DFUNSET_2'(s(fun(X106,bool),X109),s(fun(X106,fun(fun(bool,X107),bool)),'const.combin.K_1'(s(fun(fun(bool,X107),bool),'const.util_prob.FUNSET_2'(s(fun(bool,bool),X110),s(fun(X107,bool),X111)))))))))) | ~p(s(bool,'const.bool.IN_2'(s(X107,happ(s(fun(bool,X107),happ(s(fun(X106,fun(bool,X107)),X108),s(X106,sK7('const.combin.K_1'(s(fun(fun(bool,X107),bool),'const.util_prob.FUNSET_2'(s(fun(bool,bool),X110),s(fun(X107,bool),X111)))),X109,X108,X106,fun(bool,X107))))),s(bool,'F_0'))),s(fun(X107,bool),X111)))) | s(bool,'T_0') = s(bool,sK6(X111,X110,happ(s(fun(X106,fun(bool,X107)),X108),s(X106,sK7('const.combin.K_1'(s(fun(fun(bool,X107),bool),'const.util_prob.FUNSET_2'(s(fun(bool,bool),X110),s(fun(X107,bool),X111)))),X109,X108,X106,fun(bool,X107)))),X107,bool))) )),
  inference(resolution,[],[f6755,f6690])).
fof(f9002,plain,(
  ( ! [X103,X101,X105,X102,X100,X104] : (p(s(bool,'const.bool.IN_2'(s(fun(X100,fun(X101,bool)),X102),s(fun(fun(X100,fun(X101,bool)),bool),'const.util_prob.DFUNSET_2'(s(fun(X100,bool),X103),s(fun(X100,fun(fun(X101,bool),bool)),'const.combin.K_1'(s(fun(fun(X101,bool),bool),'const.util_prob.FUNSET_2'(s(fun(X101,bool),X104),s(fun(bool,bool),X105)))))))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X105)))) | s(bool,'T_0') = s(bool,happ(s(fun(X101,bool),happ(s(fun(X100,fun(X101,bool)),X102),s(X100,sK7('const.combin.K_1'(s(fun(fun(X101,bool),bool),'const.util_prob.FUNSET_2'(s(fun(X101,bool),X104),s(fun(bool,bool),X105)))),X103,X102,X100,fun(X101,bool))))),s(X101,sK6(X105,X104,happ(s(fun(X100,fun(X101,bool)),X102),s(X100,sK7('const.combin.K_1'(s(fun(fun(X101,bool),bool),'const.util_prob.FUNSET_2'(s(fun(X101,bool),X104),s(fun(bool,bool),X105)))),X103,X102,X100,fun(X101,bool)))),bool,X101))))) )),
  inference(resolution,[],[f6755,f6691])).
fof(f9001,plain,(
  ( ! [X94,X99,X97,X95,X93,X98,X96] : (p(s(bool,'const.bool.IN_2'(s(fun(X93,fun(X94,X95)),X96),s(fun(fun(X93,fun(X94,X95)),bool),'const.util_prob.DFUNSET_2'(s(fun(X93,bool),X97),s(fun(X93,fun(fun(X94,X95),bool)),'const.combin.K_1'(s(fun(fun(X94,X95),bool),'const.util_prob.DFUNSET_2'(s(fun(X94,bool),X98),s(fun(X94,fun(X95,bool)),X99)))))))))) | p(s(bool,'const.bool.IN_2'(s(X94,sK7(X99,X98,happ(s(fun(X93,fun(X94,X95)),X96),s(X93,sK7('const.combin.K_1'(s(fun(fun(X94,X95),bool),'const.util_prob.DFUNSET_2'(s(fun(X94,bool),X98),s(fun(X94,fun(X95,bool)),X99)))),X97,X96,X93,fun(X94,X95)))),X94,X95)),s(fun(X94,bool),X98))))) )),
  inference(resolution,[],[f6755,f87])).
fof(f8999,plain,(
  ( ! [X80,X78,X85,X83,X81,X79,X77,X84,X82] : (p(s(bool,'const.bool.IN_2'(s(fun(X77,fun(X78,X79)),X80),s(fun(fun(X77,fun(X78,X79)),bool),'const.util_prob.DFUNSET_2'(s(fun(X77,bool),X81),s(fun(X77,fun(fun(X78,X79),bool)),'const.combin.K_1'(s(fun(fun(X78,X79),bool),'const.util_prob.DFUNSET_2'(s(fun(X78,bool),X82),s(fun(X78,fun(X79,bool)),X83)))))))))) | p(s(bool,'const.bool.IN_2'(s(X78,sK7(X83,X82,'const.combin.K_2'(s(fun(X78,X79),happ(s(fun(X77,fun(X78,X79)),X80),s(X77,sK7('const.combin.K_1'(s(fun(fun(X78,X79),bool),'const.util_prob.DFUNSET_2'(s(fun(X78,bool),X82),s(fun(X78,fun(X79,bool)),X83)))),X81,X80,X77,fun(X78,X79))))),s(X84,X85)),X78,X79)),s(fun(X78,bool),X82))))) )),
  inference(resolution,[],[f6755,f6651])).
fof(f8998,plain,(
  ( ! [X70,X68,X76,X74,X72,X71,X69,X75,X73] : (p(s(bool,'const.bool.IN_2'(s(fun(X68,fun(X69,X70)),X71),s(fun(fun(X68,fun(X69,X70)),bool),'const.util_prob.DFUNSET_2'(s(fun(X68,bool),X72),s(fun(X68,fun(fun(X69,X70),bool)),'const.combin.K_1'(s(fun(fun(X69,X70),bool),'const.util_prob.DFUNSET_2'(s(fun(X69,bool),X73),s(fun(X69,fun(X70,bool)),X74)))))))))) | p(s(bool,'const.bool.IN_2'(s(X69,sK7(X74,'const.combin.K_2'(s(fun(X69,bool),X73),s(X75,X76)),happ(s(fun(X68,fun(X69,X70)),X71),s(X68,sK7('const.combin.K_1'(s(fun(fun(X69,X70),bool),'const.util_prob.DFUNSET_2'(s(fun(X69,bool),X73),s(fun(X69,fun(X70,bool)),X74)))),X72,X71,X68,fun(X69,X70)))),X69,X70)),s(fun(X69,bool),X73))))) )),
  inference(resolution,[],[f6755,f6673])).
fof(f8997,plain,(
  ( ! [X61,X59,X66,X64,X62,X60,X67,X65,X63] : (p(s(bool,'const.bool.IN_2'(s(fun(X59,fun(X60,X61)),X62),s(fun(fun(X59,fun(X60,X61)),bool),'const.util_prob.DFUNSET_2'(s(fun(X59,bool),X63),s(fun(X59,fun(fun(X60,X61),bool)),'const.combin.K_1'(s(fun(fun(X60,X61),bool),'const.util_prob.DFUNSET_2'(s(fun(X60,bool),X64),s(fun(X60,fun(X61,bool)),X65)))))))))) | p(s(bool,'const.bool.IN_2'(s(X60,sK7('const.combin.K_2'(s(fun(X60,fun(X61,bool)),X65),s(X66,X67)),X64,happ(s(fun(X59,fun(X60,X61)),X62),s(X59,sK7('const.combin.K_1'(s(fun(fun(X60,X61),bool),'const.util_prob.DFUNSET_2'(s(fun(X60,bool),X64),s(fun(X60,fun(X61,bool)),X65)))),X63,X62,X59,fun(X60,X61)))),X60,X61)),s(fun(X60,bool),X64))))) )),
  inference(resolution,[],[f6755,f6655])).
fof(f8996,plain,(
  ( ! [X57,X54,X52,X50,X58,X56,X55,X53,X51] : (p(s(bool,'const.bool.IN_2'(s(fun(X50,fun(X51,X52)),X53),s(fun(fun(X50,fun(X51,X52)),bool),'const.util_prob.DFUNSET_2'(s(fun(X50,bool),X54),s(fun(X50,fun(fun(X51,X52),bool)),'const.combin.K_1'(s(fun(fun(X51,X52),bool),'const.util_prob.DFUNSET_2'(s(fun(X51,bool),X55),s(fun(X51,fun(X52,bool)),X56)))))))))) | p(s(bool,'const.bool.IN_2'(s(X51,sK7(X56,happ(s(fun(X57,fun(X51,bool)),'const.combin.K_1'(s(fun(X51,bool),X55))),s(X57,X58)),happ(s(fun(X50,fun(X51,X52)),X53),s(X50,sK7('const.combin.K_1'(s(fun(fun(X51,X52),bool),'const.util_prob.DFUNSET_2'(s(fun(X51,bool),X55),s(fun(X51,fun(X52,bool)),X56)))),X54,X53,X50,fun(X51,X52)))),X51,X52)),s(fun(X51,bool),X55))))) )),
  inference(resolution,[],[f6755,f6676])).
fof(f8995,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X49] : (p(s(bool,'const.bool.IN_2'(s(fun(X43,fun(X44,X45)),X46),s(fun(fun(X43,fun(X44,X45)),bool),'const.util_prob.DFUNSET_2'(s(fun(X43,bool),X47),s(fun(X43,fun(fun(X44,X45),bool)),'const.combin.K_1'(s(fun(fun(X44,X45),bool),'const.util_prob.FUNSET_2'(s(fun(X44,bool),X48),s(fun(X45,bool),X49)))))))))) | p(s(bool,'const.bool.IN_2'(s(X44,sK6(X49,X48,happ(s(fun(X43,fun(X44,X45)),X46),s(X43,sK7('const.combin.K_1'(s(fun(fun(X44,X45),bool),'const.util_prob.FUNSET_2'(s(fun(X44,bool),X48),s(fun(X45,bool),X49)))),X47,X46,X43,fun(X44,X45)))),X45,X44)),s(fun(X44,bool),X48))))) )),
  inference(resolution,[],[f6755,f84])).
fof(f8993,plain,(
  ( ! [X30,X28,X35,X33,X31,X29,X27,X34,X32] : (p(s(bool,'const.bool.IN_2'(s(fun(X27,fun(X28,X29)),X30),s(fun(fun(X27,fun(X28,X29)),bool),'const.util_prob.DFUNSET_2'(s(fun(X27,bool),X31),s(fun(X27,fun(fun(X28,X29),bool)),'const.combin.K_1'(s(fun(fun(X28,X29),bool),'const.util_prob.FUNSET_2'(s(fun(X28,bool),X32),s(fun(X29,bool),X33)))))))))) | p(s(bool,'const.bool.IN_2'(s(X28,sK6(X33,X32,'const.combin.K_2'(s(fun(X28,X29),happ(s(fun(X27,fun(X28,X29)),X30),s(X27,sK7('const.combin.K_1'(s(fun(fun(X28,X29),bool),'const.util_prob.FUNSET_2'(s(fun(X28,bool),X32),s(fun(X29,bool),X33)))),X31,X30,X27,fun(X28,X29))))),s(X34,X35)),X29,X28)),s(fun(X28,bool),X32))))) )),
  inference(resolution,[],[f6755,f6616])).
fof(f8992,plain,(
  ( ! [X26,X24,X23,X21,X19,X25,X22,X20,X18] : (p(s(bool,'const.bool.IN_2'(s(fun(X18,fun(X19,X20)),X21),s(fun(fun(X18,fun(X19,X20)),bool),'const.util_prob.DFUNSET_2'(s(fun(X18,bool),X22),s(fun(X18,fun(fun(X19,X20),bool)),'const.combin.K_1'(s(fun(fun(X19,X20),bool),'const.util_prob.FUNSET_2'(s(fun(X19,bool),X23),s(fun(X20,bool),X24)))))))))) | p(s(bool,'const.bool.IN_2'(s(X19,sK6('const.combin.K_2'(s(fun(X20,bool),X24),s(X25,X26)),X23,happ(s(fun(X18,fun(X19,X20)),X21),s(X18,sK7('const.combin.K_1'(s(fun(fun(X19,X20),bool),'const.util_prob.FUNSET_2'(s(fun(X19,bool),X23),s(fun(X20,bool),X24)))),X22,X21,X18,fun(X19,X20)))),X20,X19)),s(fun(X19,bool),X23))))) )),
  inference(resolution,[],[f6755,f6620])).
fof(f8991,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X9,X16] : (p(s(bool,'const.bool.IN_2'(s(fun(X9,fun(X10,X11)),X12),s(fun(fun(X9,fun(X10,X11)),bool),'const.util_prob.DFUNSET_2'(s(fun(X9,bool),X13),s(fun(X9,fun(fun(X10,X11),bool)),'const.combin.K_1'(s(fun(fun(X10,X11),bool),'const.util_prob.FUNSET_2'(s(fun(X10,bool),X14),s(fun(X11,bool),X15)))))))))) | p(s(bool,'const.bool.IN_2'(s(X10,sK6(X15,'const.combin.K_2'(s(fun(X10,bool),X14),s(X16,X17)),happ(s(fun(X9,fun(X10,X11)),X12),s(X9,sK7('const.combin.K_1'(s(fun(fun(X10,X11),bool),'const.util_prob.FUNSET_2'(s(fun(X10,bool),X14),s(fun(X11,bool),X15)))),X13,X12,X9,fun(X10,X11)))),X11,X10)),s(fun(X10,bool),X14))))) )),
  inference(resolution,[],[f6755,f6637])).
fof(f8990,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,fun(X1,X2)),X3),s(fun(fun(X0,fun(X1,X2)),bool),'const.util_prob.DFUNSET_2'(s(fun(X0,bool),X4),s(fun(X0,fun(fun(X1,X2),bool)),'const.combin.K_1'(s(fun(fun(X1,X2),bool),'const.util_prob.FUNSET_2'(s(fun(X1,bool),X5),s(fun(X2,bool),X6)))))))))) | p(s(bool,'const.bool.IN_2'(s(X1,sK6(X6,happ(s(fun(X7,fun(X1,bool)),'const.combin.K_1'(s(fun(X1,bool),X5))),s(X7,X8)),happ(s(fun(X0,fun(X1,X2)),X3),s(X0,sK7('const.combin.K_1'(s(fun(fun(X1,X2),bool),'const.util_prob.FUNSET_2'(s(fun(X1,bool),X5),s(fun(X2,bool),X6)))),X4,X3,X0,fun(X1,X2)))),X2,X1)),s(fun(X1,bool),X5))))) )),
  inference(resolution,[],[f6755,f6640])).
fof(f8989,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),'const.combin.K_1'(s(X0,X4))),s(fun(fun(X1,X0),bool),'const.util_prob.DFUNSET_2'(s(fun(X1,bool),X3),s(fun(X1,fun(X0,bool)),'const.combin.K_1'(s(fun(X0,bool),X2))))))))) )),
  inference(forward_demodulation,[],[f8964,f80])).
fof(f8964,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),'const.combin.K_2'(s(fun(X0,bool),X2),s(X1,sK7('const.combin.K_1'(s(fun(X0,bool),X2)),X3,'const.combin.K_1'(s(X0,X4)),X1,X0))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),'const.combin.K_1'(s(X0,X4))),s(fun(fun(X1,X0),bool),'const.util_prob.DFUNSET_2'(s(fun(X1,bool),X3),s(fun(X1,fun(X0,bool)),'const.combin.K_1'(s(fun(X0,bool),X2))))))))) )),
  inference(superposition,[],[f6750,f81])).
fof(f8963,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X2,X3),s(fun(X2,bool),happ(s(fun(bool,fun(X2,bool)),X0),s(bool,'F_0')))))) | p(s(bool,'const.bool.IN_2'(s(fun(bool,X2),'const.combin.K_1'(s(X2,X3))),s(fun(fun(bool,X2),bool),'const.util_prob.DFUNSET_2'(s(fun(bool,bool),X1),s(fun(bool,fun(X2,bool)),X0)))))) | s(bool,'T_0') = s(bool,sK7(X0,X1,'const.combin.K_1'(s(X2,X3)),bool,X2))) )),
  inference(superposition,[],[f6750,f71])).
fof(f8937,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,'const.bool.IN_2'(s(bool,X4),s(fun(bool,bool),happ(s(fun(X5,fun(bool,bool)),X6),s(X5,sK7(X6,X7,'const.combin.K_1'(s(bool,X4)),X5,bool))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X5,bool),'const.combin.K_1'(s(bool,'F_0'))),s(fun(fun(X5,bool),bool),'const.util_prob.DFUNSET_2'(s(fun(X5,bool),X7),s(fun(X5,fun(bool,bool)),X6)))))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f6750,f71])).
fof(f8936,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),happ(s(fun(X1,fun(bool,bool)),X2),s(X1,sK7(X2,X3,'const.combin.K_1'(s(bool,'F_0')),X1,bool))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X1,bool),'const.combin.K_1'(s(bool,X0))),s(fun(fun(X1,bool),bool),'const.util_prob.DFUNSET_2'(s(fun(X1,bool),X3),s(fun(X1,fun(bool,bool)),X2)))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f6750,f71])).
fof(f8969,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(fun(X7,X8),'const.combin.K_1'(s(X8,X12))),s(fun(fun(X7,X8),bool),'const.util_prob.DFUNSET_2'(s(fun(X7,bool),X13),s(fun(X7,fun(X8,bool)),X10)))))) | ~p(s(bool,'const.bool.IN_2'(s(X8,X12),s(fun(X8,bool),happ(s(fun(X7,fun(X8,bool)),X10),s(X7,sK7(happ(s(fun(X9,fun(X7,fun(X8,bool))),'const.combin.K_1'(s(fun(X7,fun(X8,bool)),X10))),s(X9,X11)),X13,'const.combin.K_1'(s(X8,X12)),X7,X8)))))))) )),
  inference(forward_demodulation,[],[f8968,f80])).
fof(f8968,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(fun(X7,X8),'const.combin.K_1'(s(X8,X12))),s(fun(fun(X7,X8),bool),'const.util_prob.DFUNSET_2'(s(fun(X7,bool),X13),s(fun(X7,fun(X8,bool)),'const.combin.K_2'(s(fun(X7,fun(X8,bool)),X10),s(X9,X11)))))))) | ~p(s(bool,'const.bool.IN_2'(s(X8,X12),s(fun(X8,bool),happ(s(fun(X7,fun(X8,bool)),X10),s(X7,sK7(happ(s(fun(X9,fun(X7,fun(X8,bool))),'const.combin.K_1'(s(fun(X7,fun(X8,bool)),X10))),s(X9,X11)),X13,'const.combin.K_1'(s(X8,X12)),X7,X8)))))))) )),
  inference(forward_demodulation,[],[f8967,f81])).
fof(f8967,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (~p(s(bool,'const.bool.IN_2'(s(X8,X12),s(fun(X8,bool),happ(s(fun(X7,fun(X8,bool)),X10),s(X7,sK7(happ(s(fun(X9,fun(X7,fun(X8,bool))),'const.combin.K_1'(s(fun(X7,fun(X8,bool)),X10))),s(X9,X11)),X13,'const.combin.K_1'(s(X8,X12)),X7,X8))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X7,X8),'const.combin.K_1'(s(X8,X12))),s(fun(fun(X7,X8),bool),'const.util_prob.DFUNSET_2'(s(fun(X7,bool),X13),s(fun(X7,fun(X8,bool)),happ(s(fun(X9,fun(X7,fun(X8,bool))),'const.combin.K_1'(s(fun(X7,fun(X8,bool)),X10))),s(X9,X11))))))))) )),
  inference(forward_demodulation,[],[f8935,f80])).
fof(f8935,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (~p(s(bool,'const.bool.IN_2'(s(X8,X12),s(fun(X8,bool),happ(s(fun(X7,fun(X8,bool)),'const.combin.K_2'(s(fun(X7,fun(X8,bool)),X10),s(X9,X11))),s(X7,sK7(happ(s(fun(X9,fun(X7,fun(X8,bool))),'const.combin.K_1'(s(fun(X7,fun(X8,bool)),X10))),s(X9,X11)),X13,'const.combin.K_1'(s(X8,X12)),X7,X8))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X7,X8),'const.combin.K_1'(s(X8,X12))),s(fun(fun(X7,X8),bool),'const.util_prob.DFUNSET_2'(s(fun(X7,bool),X13),s(fun(X7,fun(X8,bool)),happ(s(fun(X9,fun(X7,fun(X8,bool))),'const.combin.K_1'(s(fun(X7,fun(X8,bool)),X10))),s(X9,X11))))))))) )),
  inference(superposition,[],[f6750,f81])).
fof(f8966,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,X1),'const.combin.K_1'(s(X1,X5))),s(fun(fun(X0,X1),bool),'const.util_prob.DFUNSET_2'(s(fun(X0,bool),X6),s(fun(X0,fun(X1,bool)),X2)))))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,sK7('const.combin.K_2'(s(fun(X0,fun(X1,bool)),X2),s(X3,X4)),X6,'const.combin.K_1'(s(X1,X5)),X0,X1)))))))) )),
  inference(forward_demodulation,[],[f8934,f80])).
fof(f8934,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,sK7('const.combin.K_2'(s(fun(X0,fun(X1,bool)),X2),s(X3,X4)),X6,'const.combin.K_1'(s(X1,X5)),X0,X1))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,X1),'const.combin.K_1'(s(X1,X5))),s(fun(fun(X0,X1),bool),'const.util_prob.DFUNSET_2'(s(fun(X0,bool),X6),s(fun(X0,fun(X1,bool)),'const.combin.K_2'(s(fun(X0,fun(X1,bool)),X2),s(X3,X4))))))))) )),
  inference(superposition,[],[f6750,f80])).
fof(f8931,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X9,X16] : (p(s(bool,'const.bool.IN_2'(s(fun(X9,X10),X14),s(fun(fun(X9,X10),bool),'const.util_prob.DFUNSET_2'(s(fun(X9,bool),X15),s(fun(X9,fun(X10,bool)),X12)))))) | p(s(bool,'const.bool.IN_2'(s(X9,sK7(happ(s(fun(X11,fun(X9,fun(X10,bool))),'const.combin.K_1'(s(fun(X9,fun(X10,bool)),X12))),s(X11,X13)),happ(s(fun(X16,fun(X9,bool)),'const.combin.K_1'(s(fun(X9,bool),X15))),s(X16,X17)),X14,X9,X10)),s(fun(X9,bool),X15))))) )),
  inference(forward_demodulation,[],[f8916,f80])).
fof(f8916,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X9,X16] : (p(s(bool,'const.bool.IN_2'(s(fun(X9,X10),X14),s(fun(fun(X9,X10),bool),'const.util_prob.DFUNSET_2'(s(fun(X9,bool),X15),s(fun(X9,fun(X10,bool)),'const.combin.K_2'(s(fun(X9,fun(X10,bool)),X12),s(X11,X13)))))))) | p(s(bool,'const.bool.IN_2'(s(X9,sK7(happ(s(fun(X11,fun(X9,fun(X10,bool))),'const.combin.K_1'(s(fun(X9,fun(X10,bool)),X12))),s(X11,X13)),happ(s(fun(X16,fun(X9,bool)),'const.combin.K_1'(s(fun(X9,bool),X15))),s(X16,X17)),X14,X9,X10)),s(fun(X9,bool),X15))))) )),
  inference(superposition,[],[f6676,f81])).
fof(f8874,plain,(
  ( ! [X61,X59,X66,X64,X62,X60,X58,X67,X65,X63] : (p(s(bool,'const.bool.IN_2'(s(X58,sK7(X59,happ(s(fun(X60,fun(X58,bool)),'const.combin.K_1'(s(fun(X58,bool),X61))),s(X60,X62)),X63,X58,X64)),s(fun(X58,bool),X61)))) | p(s(bool,'const.bool.IN_2'(s(X65,happ(s(fun(fun(X58,X64),X65),X66),s(fun(X58,X64),X63))),s(fun(X65,bool),X67)))) | p(s(bool,'const.bool.IN_2'(s(fun(X58,X64),sK6(X67,'const.util_prob.DFUNSET_2'(s(fun(X58,bool),X61),s(fun(X58,fun(X64,bool)),X59)),X66,X65,fun(X58,X64))),s(fun(fun(X58,X64),bool),'const.util_prob.DFUNSET_2'(s(fun(X58,bool),X61),s(fun(X58,fun(X64,bool)),X59))))))) )),
  inference(resolution,[],[f6676,f6711])).
fof(f8872,plain,(
  ( ! [X39,X47,X45,X43,X41,X48,X46,X44,X42,X40] : (p(s(bool,'const.bool.IN_2'(s(X39,sK7(X40,happ(s(fun(X41,fun(X39,bool)),'const.combin.K_1'(s(fun(X39,bool),X42))),s(X41,X43)),X44,X39,X45)),s(fun(X39,bool),X42)))) | p(s(bool,'const.bool.IN_2'(s(X46,happ(s(fun(fun(X39,X45),X46),X47),s(fun(X39,X45),X44))),s(fun(X46,bool),happ(s(fun(fun(X39,X45),fun(X46,bool)),X48),s(fun(X39,X45),X44)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X39,X45),sK7(X48,'const.util_prob.DFUNSET_2'(s(fun(X39,bool),X42),s(fun(X39,fun(X45,bool)),X40)),X47,fun(X39,X45),X46)),s(fun(fun(X39,X45),bool),'const.util_prob.DFUNSET_2'(s(fun(X39,bool),X42),s(fun(X39,fun(X45,bool)),X40))))))) )),
  inference(resolution,[],[f6676,f6756])).
fof(f8871,plain,(
  ( ! [X30,X28,X37,X35,X33,X31,X29,X38,X36,X34,X32] : (p(s(bool,'const.bool.IN_2'(s(X28,sK7(X29,happ(s(fun(X30,fun(X28,bool)),'const.combin.K_1'(s(fun(X28,bool),X31))),s(X30,X32)),happ(s(fun(X33,fun(X28,X34)),X35),s(X33,sK6('const.combin.K_2'(s(fun(fun(X28,X34),bool),'const.util_prob.DFUNSET_2'(s(fun(X28,bool),X31),s(fun(X28,fun(X34,bool)),X29))),s(X36,X37)),X38,X35,fun(X28,X34),X33))),X28,X34)),s(fun(X28,bool),X31)))) | p(s(bool,'const.bool.IN_2'(s(fun(X33,fun(X28,X34)),X35),s(fun(fun(X33,fun(X28,X34)),bool),'const.util_prob.FUNSET_2'(s(fun(X33,bool),X38),s(fun(fun(X28,X34),bool),'const.util_prob.DFUNSET_2'(s(fun(X28,bool),X31),s(fun(X28,fun(X34,bool)),X29))))))))) )),
  inference(resolution,[],[f6676,f6707])).
fof(f8870,plain,(
  ( ! [X26,X24,X23,X21,X19,X17,X27,X25,X22,X20,X18] : (p(s(bool,'const.bool.IN_2'(s(X17,sK7(X18,happ(s(fun(X19,fun(X17,bool)),'const.combin.K_1'(s(fun(X17,bool),X20))),s(X19,X21)),happ(s(fun(X22,fun(X17,X23)),X24),s(X22,sK6('const.util_prob.DFUNSET_2'(s(fun(X17,bool),X20),s(fun(X17,fun(X23,bool)),X18)),X25,'const.combin.K_2'(s(fun(X22,fun(X17,X23)),X24),s(X26,X27)),fun(X17,X23),X22))),X17,X23)),s(fun(X17,bool),X20)))) | p(s(bool,'const.bool.IN_2'(s(fun(X22,fun(X17,X23)),X24),s(fun(fun(X22,fun(X17,X23)),bool),'const.util_prob.FUNSET_2'(s(fun(X22,bool),X25),s(fun(fun(X17,X23),bool),'const.util_prob.DFUNSET_2'(s(fun(X17,bool),X20),s(fun(X17,fun(X23,bool)),X18))))))))) )),
  inference(resolution,[],[f6676,f6702])).
fof(f8869,plain,(
  ( ! [X14,X12,X10,X8,X15,X13,X11,X9,X16] : (p(s(bool,'const.bool.IN_2'(s(X8,sK7(X9,happ(s(fun(X10,fun(X8,bool)),'const.combin.K_1'(s(fun(X8,bool),X11))),s(X10,X12)),happ(s(fun(X13,fun(X8,X14)),X15),s(X13,sK6('const.util_prob.DFUNSET_2'(s(fun(X8,bool),X11),s(fun(X8,fun(X14,bool)),X9)),X16,X15,fun(X8,X14),X13))),X8,X14)),s(fun(X8,bool),X11)))) | p(s(bool,'const.bool.IN_2'(s(fun(X13,fun(X8,X14)),X15),s(fun(fun(X13,fun(X8,X14)),bool),'const.util_prob.FUNSET_2'(s(fun(X13,bool),X16),s(fun(fun(X8,X14),bool),'const.util_prob.DFUNSET_2'(s(fun(X8,bool),X11),s(fun(X8,fun(X14,bool)),X9))))))))) )),
  inference(resolution,[],[f6676,f85])).
fof(f8864,plain,(
  ( ! [X47,X54,X52,X50,X48,X56,X55,X53,X51,X49] : (p(s(bool,'const.bool.IN_2'(s(fun(fun(X47,X48),X49),X50),s(fun(fun(fun(X47,X48),X49),bool),'const.util_prob.DFUNSET_2'(s(fun(fun(X47,X48),bool),'const.util_prob.DFUNSET_2'(s(fun(X47,bool),X51),s(fun(X47,fun(X48,bool)),X52))),s(fun(fun(X47,X48),fun(X49,bool)),X53)))))) | ~p(s(bool,'const.bool.IN_2'(s(X47,X54),s(fun(X47,bool),X51)))) | p(s(bool,'const.bool.IN_2'(s(X48,happ(s(fun(X47,X48),sK7(X53,happ(s(fun(X55,fun(fun(X47,X48),bool)),'const.combin.K_1'(s(fun(fun(X47,X48),bool),'const.util_prob.DFUNSET_2'(s(fun(X47,bool),X51),s(fun(X47,fun(X48,bool)),X52))))),s(X55,X56)),X50,fun(X47,X48),X49)),s(X47,X54))),s(fun(X48,bool),happ(s(fun(X47,fun(X48,bool)),X52),s(X47,X54))))))) )),
  inference(resolution,[],[f6676,f86])).
fof(f8863,plain,(
  ( ! [X39,X37,X45,X43,X41,X38,X46,X44,X42,X40] : (p(s(bool,'const.bool.IN_2'(s(fun(fun(X37,X38),X39),X40),s(fun(fun(fun(X37,X38),X39),bool),'const.util_prob.DFUNSET_2'(s(fun(fun(X37,X38),bool),'const.util_prob.FUNSET_2'(s(fun(X37,bool),X41),s(fun(X38,bool),X42))),s(fun(fun(X37,X38),fun(X39,bool)),X43)))))) | ~p(s(bool,'const.bool.IN_2'(s(X37,X44),s(fun(X37,bool),X41)))) | p(s(bool,'const.bool.IN_2'(s(X38,happ(s(fun(X37,X38),sK7(X43,happ(s(fun(X45,fun(fun(X37,X38),bool)),'const.combin.K_1'(s(fun(fun(X37,X38),bool),'const.util_prob.FUNSET_2'(s(fun(X37,bool),X41),s(fun(X38,bool),X42))))),s(X45,X46)),X40,fun(X37,X38),X39)),s(X37,X44))),s(fun(X38,bool),X42))))) )),
  inference(resolution,[],[f6676,f83])).
fof(f8862,plain,(
  ( ! [X30,X35,X33,X31,X29,X36,X34,X32] : (p(s(bool,'const.bool.IN_2'(s(fun(bool,X29),X30),s(fun(fun(bool,X29),bool),'const.util_prob.DFUNSET_2'(s(fun(bool,bool),X31),s(fun(bool,fun(X29,bool)),X32)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X33,bool),'const.combin.K_1'(s(bool,'F_0'))),s(fun(fun(X33,bool),bool),'const.util_prob.FUNSET_2'(s(fun(X33,bool),X34),s(fun(bool,bool),X31)))))) | s(bool,'T_0') = s(bool,sK7(X32,happ(s(fun(X35,fun(bool,bool)),'const.combin.K_1'(s(fun(bool,bool),X31))),s(X35,X36)),X30,bool,X29))) )),
  inference(resolution,[],[f6676,f6783])).
fof(f8861,plain,(
  ( ! [X28,X26,X24,X23,X21,X19,X27,X25,X22,X20] : (p(s(bool,'const.bool.IN_2'(s(fun(X19,X20),X21),s(fun(fun(X19,X20),bool),'const.util_prob.DFUNSET_2'(s(fun(X19,bool),X22),s(fun(X19,fun(X20,bool)),X23)))))) | p(s(bool,'const.bool.IN_2'(s(X24,happ(s(fun(X19,X24),X25),s(X19,sK7(X23,happ(s(fun(X26,fun(X19,bool)),'const.combin.K_1'(s(fun(X19,bool),X22))),s(X26,X27)),X21,X19,X20)))),s(fun(X24,bool),X28)))) | p(s(bool,'const.bool.IN_2'(s(X19,sK6(X28,X22,X25,X24,X19)),s(fun(X19,bool),X22))))) )),
  inference(resolution,[],[f6676,f6711])).
fof(f8859,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,X1),X2),s(fun(fun(X0,X1),bool),'const.util_prob.DFUNSET_2'(s(fun(X0,bool),X3),s(fun(X0,fun(X1,bool)),X4)))))) | p(s(bool,'const.bool.IN_2'(s(X5,happ(s(fun(X0,X5),X6),s(X0,sK7(X4,happ(s(fun(X7,fun(X0,bool)),'const.combin.K_1'(s(fun(X0,bool),X3))),s(X7,X8)),X2,X0,X1)))),s(fun(X5,bool),happ(s(fun(X0,fun(X5,bool)),X9),s(X0,sK7(X4,happ(s(fun(X7,fun(X0,bool)),'const.combin.K_1'(s(fun(X0,bool),X3))),s(X7,X8)),X2,X0,X1))))))) | p(s(bool,'const.bool.IN_2'(s(X0,sK7(X9,X3,X6,X0,X5)),s(fun(X0,bool),X3))))) )),
  inference(resolution,[],[f6676,f6756])).
fof(f8858,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X9,X16] : (p(s(bool,'const.bool.IN_2'(s(fun(X9,X15),X16),s(fun(fun(X9,X15),bool),'const.util_prob.DFUNSET_2'(s(fun(X9,bool),X11),s(fun(X9,fun(X15,bool)),X13)))))) | p(s(bool,'const.bool.IN_2'(s(X9,sK7(X13,happ(s(fun(X10,fun(X9,bool)),'const.combin.K_1'(s(fun(X9,bool),X11))),s(X10,X12)),happ(s(fun(X14,fun(X9,X15)),'const.combin.K_1'(s(fun(X9,X15),X16))),s(X14,X17)),X9,X15)),s(fun(X9,bool),X11))))) )),
  inference(forward_demodulation,[],[f8857,f80])).
fof(f8857,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X9,X16] : (p(s(bool,'const.bool.IN_2'(s(fun(X9,X15),X16),s(fun(fun(X9,X15),bool),'const.util_prob.DFUNSET_2'(s(fun(X9,bool),'const.combin.K_2'(s(fun(X9,bool),X11),s(X10,X12))),s(fun(X9,fun(X15,bool)),X13)))))) | p(s(bool,'const.bool.IN_2'(s(X9,sK7(X13,happ(s(fun(X10,fun(X9,bool)),'const.combin.K_1'(s(fun(X9,bool),X11))),s(X10,X12)),happ(s(fun(X14,fun(X9,X15)),'const.combin.K_1'(s(fun(X9,X15),X16))),s(X14,X17)),X9,X15)),s(fun(X9,bool),X11))))) )),
  inference(forward_demodulation,[],[f8856,f81])).
fof(f8856,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X9,X16] : (p(s(bool,'const.bool.IN_2'(s(X9,sK7(X13,happ(s(fun(X10,fun(X9,bool)),'const.combin.K_1'(s(fun(X9,bool),X11))),s(X10,X12)),happ(s(fun(X14,fun(X9,X15)),'const.combin.K_1'(s(fun(X9,X15),X16))),s(X14,X17)),X9,X15)),s(fun(X9,bool),X11)))) | p(s(bool,'const.bool.IN_2'(s(fun(X9,X15),X16),s(fun(fun(X9,X15),bool),'const.util_prob.DFUNSET_2'(s(fun(X9,bool),happ(s(fun(X10,fun(X9,bool)),'const.combin.K_1'(s(fun(X9,bool),X11))),s(X10,X12))),s(fun(X9,fun(X15,bool)),X13))))))) )),
  inference(forward_demodulation,[],[f8845,f80])).
fof(f8845,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X9,X16] : (p(s(bool,'const.bool.IN_2'(s(X9,sK7(X13,happ(s(fun(X10,fun(X9,bool)),'const.combin.K_1'(s(fun(X9,bool),X11))),s(X10,X12)),happ(s(fun(X14,fun(X9,X15)),'const.combin.K_1'(s(fun(X9,X15),X16))),s(X14,X17)),X9,X15)),s(fun(X9,bool),'const.combin.K_2'(s(fun(X9,bool),X11),s(X10,X12)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X9,X15),X16),s(fun(fun(X9,X15),bool),'const.util_prob.DFUNSET_2'(s(fun(X9,bool),happ(s(fun(X10,fun(X9,bool)),'const.combin.K_1'(s(fun(X9,bool),X11))),s(X10,X12))),s(fun(X9,fun(X15,bool)),X13))))))) )),
  inference(superposition,[],[f6672,f81])).
fof(f8810,plain,(
  ( ! [X47,X54,X52,X50,X48,X56,X55,X53,X51,X49] : (p(s(bool,'const.bool.IN_2'(s(fun(fun(X47,X48),X49),X50),s(fun(fun(fun(X47,X48),X49),bool),'const.util_prob.DFUNSET_2'(s(fun(fun(X47,X48),bool),'const.util_prob.DFUNSET_2'(s(fun(X47,bool),X51),s(fun(X47,fun(X48,bool)),X52))),s(fun(fun(X47,X48),fun(X49,bool)),X53)))))) | ~p(s(bool,'const.bool.IN_2'(s(X47,X54),s(fun(X47,bool),X51)))) | p(s(bool,'const.bool.IN_2'(s(X48,happ(s(fun(X47,X48),sK7(X53,'const.util_prob.DFUNSET_2'(s(fun(X47,bool),X51),s(fun(X47,fun(X48,bool)),X52)),happ(s(fun(X55,fun(fun(X47,X48),X49)),'const.combin.K_1'(s(fun(fun(X47,X48),X49),X50))),s(X55,X56)),fun(X47,X48),X49)),s(X47,X54))),s(fun(X48,bool),happ(s(fun(X47,fun(X48,bool)),X52),s(X47,X54))))))) )),
  inference(resolution,[],[f6672,f86])).
fof(f8809,plain,(
  ( ! [X39,X37,X45,X43,X41,X38,X46,X44,X42,X40] : (p(s(bool,'const.bool.IN_2'(s(fun(fun(X37,X38),X39),X40),s(fun(fun(fun(X37,X38),X39),bool),'const.util_prob.DFUNSET_2'(s(fun(fun(X37,X38),bool),'const.util_prob.FUNSET_2'(s(fun(X37,bool),X41),s(fun(X38,bool),X42))),s(fun(fun(X37,X38),fun(X39,bool)),X43)))))) | ~p(s(bool,'const.bool.IN_2'(s(X37,X44),s(fun(X37,bool),X41)))) | p(s(bool,'const.bool.IN_2'(s(X38,happ(s(fun(X37,X38),sK7(X43,'const.util_prob.FUNSET_2'(s(fun(X37,bool),X41),s(fun(X38,bool),X42)),happ(s(fun(X45,fun(fun(X37,X38),X39)),'const.combin.K_1'(s(fun(fun(X37,X38),X39),X40))),s(X45,X46)),fun(X37,X38),X39)),s(X37,X44))),s(fun(X38,bool),X42))))) )),
  inference(resolution,[],[f6672,f83])).
fof(f8808,plain,(
  ( ! [X30,X35,X33,X31,X29,X36,X34,X32] : (p(s(bool,'const.bool.IN_2'(s(fun(bool,X29),X30),s(fun(fun(bool,X29),bool),'const.util_prob.DFUNSET_2'(s(fun(bool,bool),X31),s(fun(bool,fun(X29,bool)),X32)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X33,bool),'const.combin.K_1'(s(bool,'F_0'))),s(fun(fun(X33,bool),bool),'const.util_prob.FUNSET_2'(s(fun(X33,bool),X34),s(fun(bool,bool),X31)))))) | s(bool,'T_0') = s(bool,sK7(X32,X31,happ(s(fun(X35,fun(bool,X29)),'const.combin.K_1'(s(fun(bool,X29),X30))),s(X35,X36)),bool,X29))) )),
  inference(resolution,[],[f6672,f6783])).
fof(f8807,plain,(
  ( ! [X28,X26,X24,X23,X21,X19,X27,X25,X22,X20] : (p(s(bool,'const.bool.IN_2'(s(fun(X19,X20),X21),s(fun(fun(X19,X20),bool),'const.util_prob.DFUNSET_2'(s(fun(X19,bool),X22),s(fun(X19,fun(X20,bool)),X23)))))) | p(s(bool,'const.bool.IN_2'(s(X24,happ(s(fun(X19,X24),X25),s(X19,sK7(X23,X22,happ(s(fun(X26,fun(X19,X20)),'const.combin.K_1'(s(fun(X19,X20),X21))),s(X26,X27)),X19,X20)))),s(fun(X24,bool),X28)))) | p(s(bool,'const.bool.IN_2'(s(X19,sK6(X28,X22,X25,X24,X19)),s(fun(X19,bool),X22))))) )),
  inference(resolution,[],[f6672,f6711])).
fof(f8805,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,X1),X2),s(fun(fun(X0,X1),bool),'const.util_prob.DFUNSET_2'(s(fun(X0,bool),X3),s(fun(X0,fun(X1,bool)),X4)))))) | p(s(bool,'const.bool.IN_2'(s(X5,happ(s(fun(X0,X5),X6),s(X0,sK7(X4,X3,happ(s(fun(X7,fun(X0,X1)),'const.combin.K_1'(s(fun(X0,X1),X2))),s(X7,X8)),X0,X1)))),s(fun(X5,bool),happ(s(fun(X0,fun(X5,bool)),X9),s(X0,sK7(X4,X3,happ(s(fun(X7,fun(X0,X1)),'const.combin.K_1'(s(fun(X0,X1),X2))),s(X7,X8)),X0,X1))))))) | p(s(bool,'const.bool.IN_2'(s(X0,sK7(X9,X3,X6,X0,X5)),s(fun(X0,bool),X3))))) )),
  inference(resolution,[],[f6672,f6756])).
fof(f8800,plain,(
  ( ! [X54,X56,X55] : (p(s(bool,'const.bool.IN_2'(s(X54,happ(s(fun(sK1,X54),X55),s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)))))),s(fun(X54,bool),happ(s(fun(sK1,fun(X54,bool)),X56),s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1))))))))) | p(s(bool,'const.bool.IN_2'(s(sK1,sK7(X56,sK3,X55,sK1,X54)),s(fun(sK1,bool),sK3))))) ) | $spl86),
  inference(global_subsumption,[],[f67,f5752,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6816,f104,f102,f101,f100,f7026,f98,f96,f7027,f7591,f7593,f7594,f7596,f7598,f7599,f7600,f7667,f7670,f7642,f7682,f7687,f7688,f7689,f7691,f7698,f7699,f7702,f7695,f7703,f7704,f7705,f7706,f7710,f7711,f7712,f7766,f7769,f7718,f7770,f7779,f7780,f7781,f7785,f7786,f7787,f7837,f7838,f7841,f7858,f7859,f7860,f7863,f7864,f7918,f7923,f7924,f7925,f7926,f7928,f7935,f7936,f7939,f7932,f7940,f7941,f7942,f7943,f7947,f7948,f7949,f8003,f8006,f7955,f8007,f8020,f8021,f8022,f8024,f8025,f8026,f8078,f8067,f8083,f8195,f8196,f8198,f8203,f8204,f8205,f8255,f8259,f8277,f8278,f8279,f8325,f8331,f8332,f8333,f8336,f8338,f8397,f8401,f8402,f8403,f8450,f8451,f8453,f8462,f8465,f8457,f8466,f8467,f8470,f8471,f8472,f8474,f8485,f8488,f8478,f8489,f8490,f8493,f8494,f8495,f8496,f8497,f8499,f8500,f8501,f8502,f8504,f8505,f8506,f8507,f8508,f8555,f8558,f8565,f8566,f8567,f8568,f8570,f8571,f8572,f8573,f8575,f8576,f8577,f8578,f8579,f8624,f8637,f8638,f8639,f8640,f8642,f8643,f8644,f8645,f8646,f8647,f8648,f8688,f8689,f8690,f8691,f8693,f8694,f8695,f8696,f8698,f8699,f8700,f8701,f8702,f8703,f8705,f8706,f8707,f8708,f8710,f8711,f8712,f8713,f8714,f8715,f8716,f8717,f8718,f8788,f8760])).
fof(f8760,plain,(
  ( ! [X54,X56,X55] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(X54,happ(s(fun(sK1,X54),X55),s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)))))),s(fun(X54,bool),happ(s(fun(sK1,fun(X54,bool)),X56),s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1))))))))) | p(s(bool,'const.bool.IN_2'(s(sK1,sK7(X56,sK3,X55,sK1,X54)),s(fun(sK1,bool),sK3))))) ) | $spl86),
  inference(superposition,[],[f6756,f5752])).
fof(f8788,plain,(
  ( ! [X14,X12,X10,X8,X15,X13,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(X8,sK7(X15,happ(s(fun(X9,fun(X8,bool)),'const.combin.K_1'(s(fun(X8,bool),X10))),s(X9,X11)),X14,X8,X13)),s(fun(X8,bool),X10)))) | ~p(s(bool,'const.bool.IN_2'(s(X8,X12),s(fun(X8,bool),X10)))) | p(s(bool,'const.bool.IN_2'(s(X13,happ(s(fun(X8,X13),X14),s(X8,X12))),s(fun(X13,bool),happ(s(fun(X8,fun(X13,bool)),X15),s(X8,X12))))))) )),
  inference(forward_demodulation,[],[f8787,f80])).
fof(f8787,plain,(
  ( ! [X14,X12,X10,X8,X15,X13,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(X8,sK7(X15,happ(s(fun(X9,fun(X8,bool)),'const.combin.K_1'(s(fun(X8,bool),X10))),s(X9,X11)),X14,X8,X13)),s(fun(X8,bool),'const.combin.K_2'(s(fun(X8,bool),X10),s(X9,X11)))))) | ~p(s(bool,'const.bool.IN_2'(s(X8,X12),s(fun(X8,bool),X10)))) | p(s(bool,'const.bool.IN_2'(s(X13,happ(s(fun(X8,X13),X14),s(X8,X12))),s(fun(X13,bool),happ(s(fun(X8,fun(X13,bool)),X15),s(X8,X12))))))) )),
  inference(forward_demodulation,[],[f8786,f81])).
fof(f8786,plain,(
  ( ! [X14,X12,X10,X8,X15,X13,X11,X9] : (~p(s(bool,'const.bool.IN_2'(s(X8,X12),s(fun(X8,bool),X10)))) | p(s(bool,'const.bool.IN_2'(s(X13,happ(s(fun(X8,X13),X14),s(X8,X12))),s(fun(X13,bool),happ(s(fun(X8,fun(X13,bool)),X15),s(X8,X12)))))) | p(s(bool,'const.bool.IN_2'(s(X8,sK7(X15,happ(s(fun(X9,fun(X8,bool)),'const.combin.K_1'(s(fun(X8,bool),X10))),s(X9,X11)),X14,X8,X13)),s(fun(X8,bool),happ(s(fun(X9,fun(X8,bool)),'const.combin.K_1'(s(fun(X8,bool),X10))),s(X9,X11))))))) )),
  inference(forward_demodulation,[],[f8745,f80])).
fof(f8745,plain,(
  ( ! [X14,X12,X10,X8,X15,X13,X11,X9] : (~p(s(bool,'const.bool.IN_2'(s(X8,X12),s(fun(X8,bool),'const.combin.K_2'(s(fun(X8,bool),X10),s(X9,X11)))))) | p(s(bool,'const.bool.IN_2'(s(X13,happ(s(fun(X8,X13),X14),s(X8,X12))),s(fun(X13,bool),happ(s(fun(X8,fun(X13,bool)),X15),s(X8,X12)))))) | p(s(bool,'const.bool.IN_2'(s(X8,sK7(X15,happ(s(fun(X9,fun(X8,bool)),'const.combin.K_1'(s(fun(X8,bool),X10))),s(X9,X11)),X14,X8,X13)),s(fun(X8,bool),happ(s(fun(X9,fun(X8,bool)),'const.combin.K_1'(s(fun(X8,bool),X10))),s(X9,X11))))))) )),
  inference(superposition,[],[f6756,f81])).
fof(f8718,plain,(
  ( ! [X6,X8,X7,X5,X9] : (~p(s(bool,'const.bool.IN_2'(s(bool,X5),s(fun(bool,bool),X6)))) | p(s(bool,'const.bool.IN_2'(s(X7,happ(s(fun(bool,X7),X8),s(bool,'F_0'))),s(fun(X7,bool),happ(s(fun(bool,fun(X7,bool)),X9),s(bool,'F_0')))))) | p(s(bool,'const.bool.IN_2'(s(bool,sK7(X9,X6,X8,bool,X7)),s(fun(bool,bool),X6)))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f6756,f71])).
fof(f8717,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X2,happ(s(fun(bool,X2),X3),s(bool,X0))),s(fun(X2,bool),happ(s(fun(bool,fun(X2,bool)),X4),s(bool,X0)))))) | p(s(bool,'const.bool.IN_2'(s(bool,sK7(X4,X1,X3,bool,X2)),s(fun(bool,bool),X1)))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f6756,f71])).
fof(f8716,plain,(
  ( ! [X255,X253,X251,X257,X259,X254,X252,X250,X256,X258] : (p(s(bool,'const.bool.IN_2'(s(X250,happ(s(fun(X251,X250),X252),s(X251,sK6(X253,happ(s(fun(X254,fun(X251,bool)),'const.combin.K_1'(s(fun(X251,bool),X255))),s(X254,X256)),X257,X258,X251)))),s(fun(X250,bool),happ(s(fun(X251,fun(X250,bool)),X259),s(X251,sK6(X253,happ(s(fun(X254,fun(X251,bool)),'const.combin.K_1'(s(fun(X251,bool),X255))),s(X254,X256)),X257,X258,X251))))))) | p(s(bool,'const.bool.IN_2'(s(X251,sK7(X259,X255,X252,X251,X250)),s(fun(X251,bool),X255)))) | p(s(bool,'const.bool.IN_2'(s(fun(X251,X258),X257),s(fun(fun(X251,X258),bool),'const.util_prob.FUNSET_2'(s(fun(X251,bool),X255),s(fun(X258,bool),X253))))))) )),
  inference(resolution,[],[f6756,f6640])).
fof(f8715,plain,(
  ( ! [X243,X241,X249,X246,X244,X242,X240,X248,X247,X245] : (p(s(bool,'const.bool.IN_2'(s(X240,happ(s(fun(X241,X240),X242),s(X241,sK7('const.combin.K_2'(s(fun(X241,fun(X243,bool)),X244),s(X245,X246)),X247,X248,X241,X243)))),s(fun(X240,bool),happ(s(fun(X241,fun(X240,bool)),X249),s(X241,sK7('const.combin.K_2'(s(fun(X241,fun(X243,bool)),X244),s(X245,X246)),X247,X248,X241,X243))))))) | p(s(bool,'const.bool.IN_2'(s(X241,sK7(X249,X247,X242,X241,X240)),s(fun(X241,bool),X247)))) | p(s(bool,'const.bool.IN_2'(s(fun(X241,X243),X248),s(fun(fun(X241,X243),bool),'const.util_prob.DFUNSET_2'(s(fun(X241,bool),X247),s(fun(X241,fun(X243,bool)),X244))))))) )),
  inference(resolution,[],[f6756,f6655])).
fof(f8714,plain,(
  ( ! [X231,X239,X237,X235,X233,X230,X238,X236,X234,X232] : (p(s(bool,'const.bool.IN_2'(s(X230,happ(s(fun(X231,X230),X232),s(X231,sK6(happ(s(fun(X233,fun(X234,bool)),'const.combin.K_1'(s(fun(X234,bool),X235))),s(X233,X236)),X237,X238,X234,X231)))),s(fun(X230,bool),happ(s(fun(X231,fun(X230,bool)),X239),s(X231,sK6(happ(s(fun(X233,fun(X234,bool)),'const.combin.K_1'(s(fun(X234,bool),X235))),s(X233,X236)),X237,X238,X234,X231))))))) | p(s(bool,'const.bool.IN_2'(s(X231,sK7(X239,X237,X232,X231,X230)),s(fun(X231,bool),X237)))) | p(s(bool,'const.bool.IN_2'(s(fun(X231,X234),X238),s(fun(fun(X231,X234),bool),'const.util_prob.FUNSET_2'(s(fun(X231,bool),X237),s(fun(X234,bool),X235))))))) )),
  inference(resolution,[],[f6756,f6641])).
fof(f8713,plain,(
  ( ! [X222,X220,X229,X227,X225,X223,X221,X228,X226,X224] : (p(s(bool,'const.bool.IN_2'(s(X220,happ(s(fun(X221,X220),X222),s(X221,sK6('const.combin.K_2'(s(fun(X223,bool),X224),s(X225,X226)),X227,X228,X223,X221)))),s(fun(X220,bool),happ(s(fun(X221,fun(X220,bool)),X229),s(X221,sK6('const.combin.K_2'(s(fun(X223,bool),X224),s(X225,X226)),X227,X228,X223,X221))))))) | p(s(bool,'const.bool.IN_2'(s(X221,sK7(X229,X227,X222,X221,X220)),s(fun(X221,bool),X227)))) | p(s(bool,'const.bool.IN_2'(s(fun(X221,X223),X228),s(fun(fun(X221,X223),bool),'const.util_prob.FUNSET_2'(s(fun(X221,bool),X227),s(fun(X223,bool),X224))))))) )),
  inference(resolution,[],[f6756,f6620])).
fof(f8712,plain,(
  ( ! [X212,X210,X216,X218,X215,X213,X211,X217,X219,X214] : (p(s(bool,'const.bool.IN_2'(s(X210,happ(s(fun(X211,X210),X212),s(X211,sK7(X213,'const.combin.K_2'(s(fun(X211,bool),X214),s(X215,X216)),X217,X211,X218)))),s(fun(X210,bool),happ(s(fun(X211,fun(X210,bool)),X219),s(X211,sK7(X213,'const.combin.K_2'(s(fun(X211,bool),X214),s(X215,X216)),X217,X211,X218))))))) | p(s(bool,'const.bool.IN_2'(s(X211,sK7(X219,X214,X212,X211,X210)),s(fun(X211,bool),X214)))) | p(s(bool,'const.bool.IN_2'(s(fun(X211,X218),X217),s(fun(fun(X211,X218),bool),'const.util_prob.DFUNSET_2'(s(fun(X211,bool),X214),s(fun(X211,fun(X218,bool)),X213))))))) )),
  inference(resolution,[],[f6756,f6673])).
fof(f8711,plain,(
  ( ! [X208,X206,X204,X202,X200,X209,X207,X205,X203,X201] : (p(s(bool,'const.bool.IN_2'(s(X200,happ(s(fun(X201,X200),X202),s(X201,sK7(X203,X204,'const.combin.K_2'(s(fun(X201,X205),X206),s(X207,X208)),X201,X205)))),s(fun(X200,bool),happ(s(fun(X201,fun(X200,bool)),X209),s(X201,sK7(X203,X204,'const.combin.K_2'(s(fun(X201,X205),X206),s(X207,X208)),X201,X205))))))) | p(s(bool,'const.bool.IN_2'(s(X201,sK7(X209,X204,X202,X201,X200)),s(fun(X201,bool),X204)))) | p(s(bool,'const.bool.IN_2'(s(fun(X201,X205),X206),s(fun(fun(X201,X205),bool),'const.util_prob.DFUNSET_2'(s(fun(X201,bool),X204),s(fun(X201,fun(X205,bool)),X203))))))) )),
  inference(resolution,[],[f6756,f6651])).
fof(f8710,plain,(
  ( ! [X198,X196,X194,X192,X199,X197,X195,X193] : (p(s(bool,'const.bool.IN_2'(s(X192,happ(s(fun(X193,X192),X194),s(X193,sK7(X195,X196,X197,X193,X198)))),s(fun(X192,bool),happ(s(fun(X193,fun(X192,bool)),X199),s(X193,sK7(X195,X196,X197,X193,X198))))))) | p(s(bool,'const.bool.IN_2'(s(X193,sK7(X199,X196,X194,X193,X192)),s(fun(X193,bool),X196)))) | p(s(bool,'const.bool.IN_2'(s(fun(X193,X198),X197),s(fun(fun(X193,X198),bool),'const.util_prob.DFUNSET_2'(s(fun(X193,bool),X196),s(fun(X193,fun(X198,bool)),X195))))))) )),
  inference(resolution,[],[f6756,f87])).
fof(f8708,plain,(
  ( ! [X177,X175,X182,X180,X178,X176,X174,X183,X181,X179] : (p(s(bool,'const.bool.IN_2'(s(X174,happ(s(fun(X175,X174),X176),s(X175,sK6(X177,'const.combin.K_2'(s(fun(X175,bool),X178),s(X179,X180)),X181,X182,X175)))),s(fun(X174,bool),happ(s(fun(X175,fun(X174,bool)),X183),s(X175,sK6(X177,'const.combin.K_2'(s(fun(X175,bool),X178),s(X179,X180)),X181,X182,X175))))))) | p(s(bool,'const.bool.IN_2'(s(X175,sK7(X183,X178,X176,X175,X174)),s(fun(X175,bool),X178)))) | p(s(bool,'const.bool.IN_2'(s(fun(X175,X182),X181),s(fun(fun(X175,X182),bool),'const.util_prob.FUNSET_2'(s(fun(X175,bool),X178),s(fun(X182,bool),X177))))))) )),
  inference(resolution,[],[f6756,f6637])).
fof(f8707,plain,(
  ( ! [X167,X165,X173,X171,X169,X166,X164,X172,X170,X168] : (p(s(bool,'const.bool.IN_2'(s(X164,happ(s(fun(X165,X164),X166),s(X165,sK6(X167,X168,happ(s(fun(X169,fun(X165,X170)),'const.combin.K_1'(s(fun(X165,X170),X171))),s(X169,X172)),X170,X165)))),s(fun(X164,bool),happ(s(fun(X165,fun(X164,bool)),X173),s(X165,sK6(X167,X168,happ(s(fun(X169,fun(X165,X170)),'const.combin.K_1'(s(fun(X165,X170),X171))),s(X169,X172)),X170,X165))))))) | p(s(bool,'const.bool.IN_2'(s(X165,sK7(X173,X168,X166,X165,X164)),s(fun(X165,bool),X168)))) | p(s(bool,'const.bool.IN_2'(s(fun(X165,X170),X171),s(fun(fun(X165,X170),bool),'const.util_prob.FUNSET_2'(s(fun(X165,bool),X168),s(fun(X170,bool),X167))))))) )),
  inference(resolution,[],[f6756,f6636])).
fof(f8706,plain,(
  ( ! [X158,X156,X154,X163,X161,X159,X157,X155,X162,X160] : (p(s(bool,'const.bool.IN_2'(s(X154,happ(s(fun(X155,X154),X156),s(X155,sK6(X157,X158,'const.combin.K_2'(s(fun(X155,X159),X160),s(X161,X162)),X159,X155)))),s(fun(X154,bool),happ(s(fun(X155,fun(X154,bool)),X163),s(X155,sK6(X157,X158,'const.combin.K_2'(s(fun(X155,X159),X160),s(X161,X162)),X159,X155))))))) | p(s(bool,'const.bool.IN_2'(s(X155,sK7(X163,X158,X156,X155,X154)),s(fun(X155,bool),X158)))) | p(s(bool,'const.bool.IN_2'(s(fun(X155,X159),X160),s(fun(fun(X155,X159),bool),'const.util_prob.FUNSET_2'(s(fun(X155,bool),X158),s(fun(X159,bool),X157))))))) )),
  inference(resolution,[],[f6756,f6616])).
fof(f8705,plain,(
  ( ! [X146,X152,X151,X149,X147,X153,X150,X148] : (p(s(bool,'const.bool.IN_2'(s(X146,happ(s(fun(X147,X146),X148),s(X147,sK6(X149,X150,X151,X152,X147)))),s(fun(X146,bool),happ(s(fun(X147,fun(X146,bool)),X153),s(X147,sK6(X149,X150,X151,X152,X147))))))) | p(s(bool,'const.bool.IN_2'(s(X147,sK7(X153,X150,X148,X147,X146)),s(fun(X147,bool),X150)))) | p(s(bool,'const.bool.IN_2'(s(fun(X147,X152),X151),s(fun(fun(X147,X152),bool),'const.util_prob.FUNSET_2'(s(fun(X147,bool),X150),s(fun(X152,bool),X149))))))) )),
  inference(resolution,[],[f6756,f84])).
fof(f8703,plain,(
  ( ! [X134,X132,X136,X135,X133,X131,X137] : (p(s(bool,'const.bool.IN_2'(s(X131,happ(s(fun(fun(bool,X132),X131),X133),s(fun(bool,X132),X134))),s(fun(X131,bool),happ(s(fun(fun(bool,X132),fun(X131,bool)),X135),s(fun(bool,X132),X134)))))) | p(s(bool,'const.bool.IN_2'(s(fun(bool,X132),sK7(X135,'const.util_prob.DFUNSET_2'(s(fun(bool,bool),X136),s(fun(bool,fun(X132,bool)),X137)),X133,fun(bool,X132),X131)),s(fun(fun(bool,X132),bool),'const.util_prob.DFUNSET_2'(s(fun(bool,bool),X136),s(fun(bool,fun(X132,bool)),X137)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X136)))) | s(bool,'T_0') = s(bool,sK7(X137,X136,X134,bool,X132))) )),
  inference(resolution,[],[f6756,f6662])).
fof(f8702,plain,(
  ( ! [X127,X125,X130,X128,X126,X124,X129] : (p(s(bool,'const.bool.IN_2'(s(X124,happ(s(fun(fun(bool,X125),X124),X126),s(fun(bool,X125),X127))),s(fun(X124,bool),happ(s(fun(fun(bool,X125),fun(X124,bool)),X128),s(fun(bool,X125),X127)))))) | p(s(bool,'const.bool.IN_2'(s(fun(bool,X125),sK7(X128,'const.util_prob.FUNSET_2'(s(fun(bool,bool),X129),s(fun(X125,bool),X130)),X126,fun(bool,X125),X124)),s(fun(fun(bool,X125),bool),'const.util_prob.FUNSET_2'(s(fun(bool,bool),X129),s(fun(X125,bool),X130)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X129)))) | s(bool,'T_0') = s(bool,sK6(X130,X129,X127,X125,bool))) )),
  inference(resolution,[],[f6756,f6627])).
fof(f8701,plain,(
  ( ! [X123,X121,X118,X122,X120,X119,X117] : (p(s(bool,'const.bool.IN_2'(s(X117,happ(s(fun(fun(bool,X118),X117),X119),s(fun(bool,X118),X120))),s(fun(X117,bool),happ(s(fun(fun(bool,X118),fun(X117,bool)),X121),s(fun(bool,X118),X120)))))) | p(s(bool,'const.bool.IN_2'(s(fun(bool,X118),sK7(X121,'const.util_prob.FUNSET_2'(s(fun(bool,bool),X122),s(fun(X118,bool),X123)),X119,fun(bool,X118),X117)),s(fun(fun(bool,X118),bool),'const.util_prob.FUNSET_2'(s(fun(bool,bool),X122),s(fun(X118,bool),X123)))))) | ~p(s(bool,'const.bool.IN_2'(s(X118,happ(s(fun(bool,X118),X120),s(bool,'F_0'))),s(fun(X118,bool),X123)))) | s(bool,'T_0') = s(bool,sK6(X123,X122,X120,X118,bool))) )),
  inference(resolution,[],[f6756,f6690])).
fof(f8700,plain,(
  ( ! [X111,X116,X114,X112,X110,X115,X113] : (p(s(bool,'const.bool.IN_2'(s(X110,happ(s(fun(fun(X111,bool),X110),X112),s(fun(X111,bool),X113))),s(fun(X110,bool),happ(s(fun(fun(X111,bool),fun(X110,bool)),X114),s(fun(X111,bool),X113)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X111,bool),sK7(X114,'const.util_prob.FUNSET_2'(s(fun(X111,bool),X115),s(fun(bool,bool),X116)),X112,fun(X111,bool),X110)),s(fun(fun(X111,bool),bool),'const.util_prob.FUNSET_2'(s(fun(X111,bool),X115),s(fun(bool,bool),X116)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X116)))) | s(bool,'T_0') = s(bool,happ(s(fun(X111,bool),X113),s(X111,sK6(X116,X115,X113,bool,X111))))) )),
  inference(resolution,[],[f6756,f6691])).
fof(f8699,plain,(
  ( ! [X103,X109,X107,X105,X102,X108,X106,X104] : (p(s(bool,'const.bool.IN_2'(s(X102,happ(s(fun(fun(X103,X104),X102),X105),s(fun(X103,X104),'const.combin.K_1'(s(X104,X106))))),s(fun(X102,bool),happ(s(fun(fun(X103,X104),fun(X102,bool)),X107),s(fun(X103,X104),'const.combin.K_1'(s(X104,X106)))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X103,X104),sK7(X107,'const.util_prob.FUNSET_2'(s(fun(X103,bool),X108),s(fun(X104,bool),X109)),X105,fun(X103,X104),X102)),s(fun(fun(X103,X104),bool),'const.util_prob.FUNSET_2'(s(fun(X103,bool),X108),s(fun(X104,bool),X109)))))) | ~p(s(bool,'const.bool.IN_2'(s(X104,X106),s(fun(X104,bool),X109))))) )),
  inference(resolution,[],[f6756,f6706])).
fof(f8698,plain,(
  ( ! [X94,X101,X99,X97,X95,X100,X98,X96] : (p(s(bool,'const.bool.IN_2'(s(X94,happ(s(fun(fun(X95,X96),X94),X97),s(fun(X95,X96),X98))),s(fun(X94,bool),happ(s(fun(fun(X95,X96),fun(X94,bool)),X99),s(fun(X95,X96),X98)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X95,X96),sK7(X99,'const.util_prob.DFUNSET_2'(s(fun(X95,bool),X100),s(fun(X95,fun(X96,bool)),X101)),X97,fun(X95,X96),X94)),s(fun(fun(X95,X96),bool),'const.util_prob.DFUNSET_2'(s(fun(X95,bool),X100),s(fun(X95,fun(X96,bool)),X101)))))) | p(s(bool,'const.bool.IN_2'(s(X95,sK7(X101,X100,X98,X95,X96)),s(fun(X95,bool),X100))))) )),
  inference(resolution,[],[f6756,f87])).
fof(f8696,plain,(
  ( ! [X80,X78,X76,X85,X83,X81,X79,X77,X84,X82] : (p(s(bool,'const.bool.IN_2'(s(X76,happ(s(fun(fun(X77,X78),X76),X79),s(fun(X77,X78),X80))),s(fun(X76,bool),happ(s(fun(fun(X77,X78),fun(X76,bool)),X81),s(fun(X77,X78),X80)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X77,X78),sK7(X81,'const.util_prob.DFUNSET_2'(s(fun(X77,bool),X82),s(fun(X77,fun(X78,bool)),X83)),X79,fun(X77,X78),X76)),s(fun(fun(X77,X78),bool),'const.util_prob.DFUNSET_2'(s(fun(X77,bool),X82),s(fun(X77,fun(X78,bool)),X83)))))) | p(s(bool,'const.bool.IN_2'(s(X77,sK7(X83,X82,'const.combin.K_2'(s(fun(X77,X78),X80),s(X84,X85)),X77,X78)),s(fun(X77,bool),X82))))) )),
  inference(resolution,[],[f6756,f6651])).
fof(f8695,plain,(
  ( ! [X70,X68,X66,X74,X72,X71,X69,X67,X75,X73] : (p(s(bool,'const.bool.IN_2'(s(X66,happ(s(fun(fun(X67,X68),X66),X69),s(fun(X67,X68),X70))),s(fun(X66,bool),happ(s(fun(fun(X67,X68),fun(X66,bool)),X71),s(fun(X67,X68),X70)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X67,X68),sK7(X71,'const.util_prob.DFUNSET_2'(s(fun(X67,bool),X72),s(fun(X67,fun(X68,bool)),X73)),X69,fun(X67,X68),X66)),s(fun(fun(X67,X68),bool),'const.util_prob.DFUNSET_2'(s(fun(X67,bool),X72),s(fun(X67,fun(X68,bool)),X73)))))) | p(s(bool,'const.bool.IN_2'(s(X67,sK7(X73,'const.combin.K_2'(s(fun(X67,bool),X72),s(X74,X75)),X70,X67,X68)),s(fun(X67,bool),X72))))) )),
  inference(resolution,[],[f6756,f6673])).
fof(f8694,plain,(
  ( ! [X61,X59,X57,X64,X62,X60,X58,X56,X65,X63] : (p(s(bool,'const.bool.IN_2'(s(X56,happ(s(fun(fun(X57,X58),X56),X59),s(fun(X57,X58),X60))),s(fun(X56,bool),happ(s(fun(fun(X57,X58),fun(X56,bool)),X61),s(fun(X57,X58),X60)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X57,X58),sK7(X61,'const.util_prob.DFUNSET_2'(s(fun(X57,bool),X62),s(fun(X57,fun(X58,bool)),X63)),X59,fun(X57,X58),X56)),s(fun(fun(X57,X58),bool),'const.util_prob.DFUNSET_2'(s(fun(X57,bool),X62),s(fun(X57,fun(X58,bool)),X63)))))) | p(s(bool,'const.bool.IN_2'(s(X57,sK7('const.combin.K_2'(s(fun(X57,fun(X58,bool)),X63),s(X64,X65)),X62,X60,X57,X58)),s(fun(X57,bool),X62))))) )),
  inference(resolution,[],[f6756,f6655])).
fof(f8693,plain,(
  ( ! [X54,X52,X50,X48,X55,X53,X51,X49] : (p(s(bool,'const.bool.IN_2'(s(X48,happ(s(fun(fun(X49,X50),X48),X51),s(fun(X49,X50),X52))),s(fun(X48,bool),happ(s(fun(fun(X49,X50),fun(X48,bool)),X53),s(fun(X49,X50),X52)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X49,X50),sK7(X53,'const.util_prob.FUNSET_2'(s(fun(X49,bool),X54),s(fun(X50,bool),X55)),X51,fun(X49,X50),X48)),s(fun(fun(X49,X50),bool),'const.util_prob.FUNSET_2'(s(fun(X49,bool),X54),s(fun(X50,bool),X55)))))) | p(s(bool,'const.bool.IN_2'(s(X49,sK6(X55,X54,X52,X50,X49)),s(fun(X49,bool),X54))))) )),
  inference(resolution,[],[f6756,f84])).
fof(f8691,plain,(
  ( ! [X30,X39,X37,X35,X33,X31,X38,X36,X34,X32] : (p(s(bool,'const.bool.IN_2'(s(X30,happ(s(fun(fun(X31,X32),X30),X33),s(fun(X31,X32),X34))),s(fun(X30,bool),happ(s(fun(fun(X31,X32),fun(X30,bool)),X35),s(fun(X31,X32),X34)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X31,X32),sK7(X35,'const.util_prob.FUNSET_2'(s(fun(X31,bool),X36),s(fun(X32,bool),X37)),X33,fun(X31,X32),X30)),s(fun(fun(X31,X32),bool),'const.util_prob.FUNSET_2'(s(fun(X31,bool),X36),s(fun(X32,bool),X37)))))) | p(s(bool,'const.bool.IN_2'(s(X31,sK6(X37,X36,'const.combin.K_2'(s(fun(X31,X32),X34),s(X38,X39)),X32,X31)),s(fun(X31,bool),X36))))) )),
  inference(resolution,[],[f6756,f6616])).
fof(f8690,plain,(
  ( ! [X28,X26,X24,X23,X21,X29,X27,X25,X22,X20] : (p(s(bool,'const.bool.IN_2'(s(X20,happ(s(fun(fun(X21,X22),X20),X23),s(fun(X21,X22),X24))),s(fun(X20,bool),happ(s(fun(fun(X21,X22),fun(X20,bool)),X25),s(fun(X21,X22),X24)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X21,X22),sK7(X25,'const.util_prob.FUNSET_2'(s(fun(X21,bool),X26),s(fun(X22,bool),X27)),X23,fun(X21,X22),X20)),s(fun(fun(X21,X22),bool),'const.util_prob.FUNSET_2'(s(fun(X21,bool),X26),s(fun(X22,bool),X27)))))) | p(s(bool,'const.bool.IN_2'(s(X21,sK6('const.combin.K_2'(s(fun(X22,bool),X27),s(X28,X29)),X26,X24,X22,X21)),s(fun(X21,bool),X26))))) )),
  inference(resolution,[],[f6756,f6620])).
fof(f8689,plain,(
  ( ! [X14,X12,X10,X19,X17,X15,X13,X11,X18,X16] : (p(s(bool,'const.bool.IN_2'(s(X10,happ(s(fun(fun(X11,X12),X10),X13),s(fun(X11,X12),X14))),s(fun(X10,bool),happ(s(fun(fun(X11,X12),fun(X10,bool)),X15),s(fun(X11,X12),X14)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X11,X12),sK7(X15,'const.util_prob.FUNSET_2'(s(fun(X11,bool),X16),s(fun(X12,bool),X17)),X13,fun(X11,X12),X10)),s(fun(fun(X11,X12),bool),'const.util_prob.FUNSET_2'(s(fun(X11,bool),X16),s(fun(X12,bool),X17)))))) | p(s(bool,'const.bool.IN_2'(s(X11,sK6(X17,'const.combin.K_2'(s(fun(X11,bool),X16),s(X18,X19)),X14,X12,X11)),s(fun(X11,bool),X16))))) )),
  inference(resolution,[],[f6756,f6637])).
fof(f8688,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(fun(X1,X2),X0),X3),s(fun(X1,X2),X4))),s(fun(X0,bool),happ(s(fun(fun(X1,X2),fun(X0,bool)),X5),s(fun(X1,X2),X4)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X1,X2),sK7(X5,'const.util_prob.FUNSET_2'(s(fun(X1,bool),X6),s(fun(X2,bool),X7)),X3,fun(X1,X2),X0)),s(fun(fun(X1,X2),bool),'const.util_prob.FUNSET_2'(s(fun(X1,bool),X6),s(fun(X2,bool),X7)))))) | p(s(bool,'const.bool.IN_2'(s(X1,sK6(X7,happ(s(fun(X8,fun(X1,bool)),'const.combin.K_1'(s(fun(X1,bool),X6))),s(X8,X9)),X4,X2,X1)),s(fun(X1,bool),X6))))) )),
  inference(resolution,[],[f6756,f6640])).
fof(f8648,plain,(
  ( ! [X94,X92,X90,X88,X95,X93,X91,X89] : (p(s(bool,'const.bool.IN_2'(s(fun(X88,bool),'const.combin.K_1'(s(bool,'F_0'))),s(fun(fun(X88,bool),bool),'const.util_prob.FUNSET_2'(s(fun(X88,bool),X89),s(fun(bool,bool),X90)))))) | s(bool,'T_0') = s(bool,sK6(X91,happ(s(fun(X92,fun(bool,bool)),'const.combin.K_1'(s(fun(bool,bool),X90))),s(X92,X93)),X94,X95,bool)) | p(s(bool,'const.bool.IN_2'(s(fun(bool,X95),X94),s(fun(fun(bool,X95),bool),'const.util_prob.FUNSET_2'(s(fun(bool,bool),X90),s(fun(X95,bool),X91))))))) )),
  inference(resolution,[],[f6783,f6640])).
fof(f8647,plain,(
  ( ! [X80,X87,X85,X83,X81,X86,X84,X82] : (p(s(bool,'const.bool.IN_2'(s(fun(X80,bool),'const.combin.K_1'(s(bool,'F_0'))),s(fun(fun(X80,bool),bool),'const.util_prob.FUNSET_2'(s(fun(X80,bool),X81),s(fun(bool,bool),X82)))))) | s(bool,'T_0') = s(bool,sK7('const.combin.K_2'(s(fun(bool,fun(X83,bool)),X84),s(X85,X86)),X82,X87,bool,X83)) | p(s(bool,'const.bool.IN_2'(s(fun(bool,X83),X87),s(fun(fun(bool,X83),bool),'const.util_prob.DFUNSET_2'(s(fun(bool,bool),X82),s(fun(bool,fun(X83,bool)),X84))))))) )),
  inference(resolution,[],[f6783,f6655])).
fof(f8646,plain,(
  ( ! [X78,X76,X74,X72,X79,X77,X75,X73] : (p(s(bool,'const.bool.IN_2'(s(fun(X72,bool),'const.combin.K_1'(s(bool,'F_0'))),s(fun(fun(X72,bool),bool),'const.util_prob.FUNSET_2'(s(fun(X72,bool),X73),s(fun(bool,bool),X74)))))) | s(bool,'T_0') = s(bool,sK6(happ(s(fun(X75,fun(X76,bool)),'const.combin.K_1'(s(fun(X76,bool),X77))),s(X75,X78)),X74,X79,X76,bool)) | p(s(bool,'const.bool.IN_2'(s(fun(bool,X76),X79),s(fun(fun(bool,X76),bool),'const.util_prob.FUNSET_2'(s(fun(bool,bool),X74),s(fun(X76,bool),X77))))))) )),
  inference(resolution,[],[f6783,f6641])).
fof(f8645,plain,(
  ( ! [X70,X68,X66,X64,X71,X69,X67,X65] : (p(s(bool,'const.bool.IN_2'(s(fun(X64,bool),'const.combin.K_1'(s(bool,'F_0'))),s(fun(fun(X64,bool),bool),'const.util_prob.FUNSET_2'(s(fun(X64,bool),X65),s(fun(bool,bool),X66)))))) | s(bool,'T_0') = s(bool,sK6('const.combin.K_2'(s(fun(X67,bool),X68),s(X69,X70)),X66,X71,X67,bool)) | p(s(bool,'const.bool.IN_2'(s(fun(bool,X67),X71),s(fun(fun(bool,X67),bool),'const.util_prob.FUNSET_2'(s(fun(bool,bool),X66),s(fun(X67,bool),X68))))))) )),
  inference(resolution,[],[f6783,f6620])).
fof(f8644,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X56,X63] : (p(s(bool,'const.bool.IN_2'(s(fun(X56,bool),'const.combin.K_1'(s(bool,'F_0'))),s(fun(fun(X56,bool),bool),'const.util_prob.FUNSET_2'(s(fun(X56,bool),X57),s(fun(bool,bool),X58)))))) | s(bool,'T_0') = s(bool,sK7(X59,'const.combin.K_2'(s(fun(bool,bool),X58),s(X60,X61)),X62,bool,X63)) | p(s(bool,'const.bool.IN_2'(s(fun(bool,X63),X62),s(fun(fun(bool,X63),bool),'const.util_prob.DFUNSET_2'(s(fun(bool,bool),X58),s(fun(bool,fun(X63,bool)),X59))))))) )),
  inference(resolution,[],[f6783,f6673])).
fof(f8643,plain,(
  ( ! [X54,X52,X50,X48,X55,X53,X51,X49] : (p(s(bool,'const.bool.IN_2'(s(fun(X48,bool),'const.combin.K_1'(s(bool,'F_0'))),s(fun(fun(X48,bool),bool),'const.util_prob.FUNSET_2'(s(fun(X48,bool),X49),s(fun(bool,bool),X50)))))) | s(bool,'T_0') = s(bool,sK7(X51,X50,'const.combin.K_2'(s(fun(bool,X52),X53),s(X54,X55)),bool,X52)) | p(s(bool,'const.bool.IN_2'(s(fun(bool,X52),X53),s(fun(fun(bool,X52),bool),'const.util_prob.DFUNSET_2'(s(fun(bool,bool),X50),s(fun(bool,fun(X52,bool)),X51))))))) )),
  inference(resolution,[],[f6783,f6651])).
fof(f8642,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (p(s(bool,'const.bool.IN_2'(s(fun(X42,bool),'const.combin.K_1'(s(bool,'F_0'))),s(fun(fun(X42,bool),bool),'const.util_prob.FUNSET_2'(s(fun(X42,bool),X43),s(fun(bool,bool),X44)))))) | s(bool,'T_0') = s(bool,sK7(X45,X44,X46,bool,X47)) | p(s(bool,'const.bool.IN_2'(s(fun(bool,X47),X46),s(fun(fun(bool,X47),bool),'const.util_prob.DFUNSET_2'(s(fun(bool,bool),X44),s(fun(bool,fun(X47,bool)),X45))))))) )),
  inference(resolution,[],[f6783,f87])).
fof(f8640,plain,(
  ( ! [X30,X28,X35,X33,X31,X29,X34,X32] : (p(s(bool,'const.bool.IN_2'(s(fun(X28,bool),'const.combin.K_1'(s(bool,'F_0'))),s(fun(fun(X28,bool),bool),'const.util_prob.FUNSET_2'(s(fun(X28,bool),X29),s(fun(bool,bool),X30)))))) | s(bool,'T_0') = s(bool,sK6(X31,'const.combin.K_2'(s(fun(bool,bool),X30),s(X32,X33)),X34,X35,bool)) | p(s(bool,'const.bool.IN_2'(s(fun(bool,X35),X34),s(fun(fun(bool,X35),bool),'const.util_prob.FUNSET_2'(s(fun(bool,bool),X30),s(fun(X35,bool),X31))))))) )),
  inference(resolution,[],[f6783,f6637])).
fof(f8639,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22,X20] : (p(s(bool,'const.bool.IN_2'(s(fun(X20,bool),'const.combin.K_1'(s(bool,'F_0'))),s(fun(fun(X20,bool),bool),'const.util_prob.FUNSET_2'(s(fun(X20,bool),X21),s(fun(bool,bool),X22)))))) | s(bool,'T_0') = s(bool,sK6(X23,X22,happ(s(fun(X24,fun(bool,X25)),'const.combin.K_1'(s(fun(bool,X25),X26))),s(X24,X27)),X25,bool)) | p(s(bool,'const.bool.IN_2'(s(fun(bool,X25),X26),s(fun(fun(bool,X25),bool),'const.util_prob.FUNSET_2'(s(fun(bool,bool),X22),s(fun(X25,bool),X23))))))) )),
  inference(resolution,[],[f6783,f6636])).
fof(f8638,plain,(
  ( ! [X14,X12,X19,X17,X15,X13,X18,X16] : (p(s(bool,'const.bool.IN_2'(s(fun(X12,bool),'const.combin.K_1'(s(bool,'F_0'))),s(fun(fun(X12,bool),bool),'const.util_prob.FUNSET_2'(s(fun(X12,bool),X13),s(fun(bool,bool),X14)))))) | s(bool,'T_0') = s(bool,sK6(X15,X14,'const.combin.K_2'(s(fun(bool,X16),X17),s(X18,X19)),X16,bool)) | p(s(bool,'const.bool.IN_2'(s(fun(bool,X16),X17),s(fun(fun(bool,X16),bool),'const.util_prob.FUNSET_2'(s(fun(bool,bool),X14),s(fun(X16,bool),X15))))))) )),
  inference(resolution,[],[f6783,f6616])).
fof(f8637,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),'const.combin.K_1'(s(bool,'F_0'))),s(fun(fun(X6,bool),bool),'const.util_prob.FUNSET_2'(s(fun(X6,bool),X7),s(fun(bool,bool),X8)))))) | s(bool,'T_0') = s(bool,sK6(X9,X8,X10,X11,bool)) | p(s(bool,'const.bool.IN_2'(s(fun(bool,X11),X10),s(fun(fun(bool,X11),bool),'const.util_prob.FUNSET_2'(s(fun(bool,bool),X8),s(fun(X11,bool),X9))))))) )),
  inference(resolution,[],[f6783,f84])).
fof(f8624,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X9,X16] : (p(s(bool,'const.bool.IN_2'(s(fun(X9,X10),X12),s(fun(fun(X9,X10),bool),'const.util_prob.FUNSET_2'(s(fun(X9,bool),X17),s(fun(X10,bool),X14)))))) | ~p(s(bool,'const.bool.IN_2'(s(X10,happ(s(fun(X9,X10),X12),s(X9,sK6('const.combin.K_2'(s(fun(X10,bool),X14),s(X15,X16)),X17,happ(s(fun(X11,fun(X9,X10)),'const.combin.K_1'(s(fun(X9,X10),X12))),s(X11,X13)),X10,X9)))),s(fun(X10,bool),X14))))) )),
  inference(forward_demodulation,[],[f8623,f80])).
fof(f8623,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X9,X16] : (p(s(bool,'const.bool.IN_2'(s(fun(X9,X10),'const.combin.K_2'(s(fun(X9,X10),X12),s(X11,X13))),s(fun(fun(X9,X10),bool),'const.util_prob.FUNSET_2'(s(fun(X9,bool),X17),s(fun(X10,bool),X14)))))) | ~p(s(bool,'const.bool.IN_2'(s(X10,happ(s(fun(X9,X10),X12),s(X9,sK6('const.combin.K_2'(s(fun(X10,bool),X14),s(X15,X16)),X17,happ(s(fun(X11,fun(X9,X10)),'const.combin.K_1'(s(fun(X9,X10),X12))),s(X11,X13)),X10,X9)))),s(fun(X10,bool),X14))))) )),
  inference(forward_demodulation,[],[f8622,f81])).
fof(f8622,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X9,X16] : (~p(s(bool,'const.bool.IN_2'(s(X10,happ(s(fun(X9,X10),X12),s(X9,sK6('const.combin.K_2'(s(fun(X10,bool),X14),s(X15,X16)),X17,happ(s(fun(X11,fun(X9,X10)),'const.combin.K_1'(s(fun(X9,X10),X12))),s(X11,X13)),X10,X9)))),s(fun(X10,bool),X14)))) | p(s(bool,'const.bool.IN_2'(s(fun(X9,X10),happ(s(fun(X11,fun(X9,X10)),'const.combin.K_1'(s(fun(X9,X10),X12))),s(X11,X13))),s(fun(fun(X9,X10),bool),'const.util_prob.FUNSET_2'(s(fun(X9,bool),X17),s(fun(X10,bool),X14))))))) )),
  inference(forward_demodulation,[],[f8581,f80])).
fof(f8581,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X9,X16] : (~p(s(bool,'const.bool.IN_2'(s(X10,happ(s(fun(X9,X10),'const.combin.K_2'(s(fun(X9,X10),X12),s(X11,X13))),s(X9,sK6('const.combin.K_2'(s(fun(X10,bool),X14),s(X15,X16)),X17,happ(s(fun(X11,fun(X9,X10)),'const.combin.K_1'(s(fun(X9,X10),X12))),s(X11,X13)),X10,X9)))),s(fun(X10,bool),X14)))) | p(s(bool,'const.bool.IN_2'(s(fun(X9,X10),happ(s(fun(X11,fun(X9,X10)),'const.combin.K_1'(s(fun(X9,X10),X12))),s(X11,X13))),s(fun(fun(X9,X10),bool),'const.util_prob.FUNSET_2'(s(fun(X9,bool),X17),s(fun(X10,bool),X14))))))) )),
  inference(superposition,[],[f6707,f81])).
fof(f8579,plain,(
  ( ! [X144,X142,X140,X138,X143,X141,X139,X137] : (p(s(bool,'const.bool.IN_2'(s(fun(X137,fun(bool,X138)),X139),s(fun(fun(X137,fun(bool,X138)),bool),'const.util_prob.FUNSET_2'(s(fun(X137,bool),X140),s(fun(fun(bool,X138),bool),'const.util_prob.DFUNSET_2'(s(fun(bool,bool),X141),s(fun(bool,fun(X138,bool)),X142)))))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X141)))) | s(bool,'T_0') = s(bool,sK7(X142,X141,happ(s(fun(X137,fun(bool,X138)),X139),s(X137,sK6('const.combin.K_2'(s(fun(fun(bool,X138),bool),'const.util_prob.DFUNSET_2'(s(fun(bool,bool),X141),s(fun(bool,fun(X138,bool)),X142))),s(X143,X144)),X140,X139,fun(bool,X138),X137))),bool,X138))) )),
  inference(resolution,[],[f6707,f6662])).
fof(f8578,plain,(
  ( ! [X134,X132,X130,X136,X135,X133,X131,X129] : (p(s(bool,'const.bool.IN_2'(s(fun(X129,fun(bool,X130)),X131),s(fun(fun(X129,fun(bool,X130)),bool),'const.util_prob.FUNSET_2'(s(fun(X129,bool),X132),s(fun(fun(bool,X130),bool),'const.util_prob.FUNSET_2'(s(fun(bool,bool),X133),s(fun(X130,bool),X134)))))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X133)))) | s(bool,'T_0') = s(bool,sK6(X134,X133,happ(s(fun(X129,fun(bool,X130)),X131),s(X129,sK6('const.combin.K_2'(s(fun(fun(bool,X130),bool),'const.util_prob.FUNSET_2'(s(fun(bool,bool),X133),s(fun(X130,bool),X134))),s(X135,X136)),X132,X131,fun(bool,X130),X129))),X130,bool))) )),
  inference(resolution,[],[f6707,f6627])).
fof(f8577,plain,(
  ( ! [X127,X125,X123,X121,X128,X126,X124,X122] : (p(s(bool,'const.bool.IN_2'(s(fun(X121,fun(bool,X122)),X123),s(fun(fun(X121,fun(bool,X122)),bool),'const.util_prob.FUNSET_2'(s(fun(X121,bool),X124),s(fun(fun(bool,X122),bool),'const.util_prob.FUNSET_2'(s(fun(bool,bool),X125),s(fun(X122,bool),X126)))))))) | ~p(s(bool,'const.bool.IN_2'(s(X122,happ(s(fun(bool,X122),happ(s(fun(X121,fun(bool,X122)),X123),s(X121,sK6('const.combin.K_2'(s(fun(fun(bool,X122),bool),'const.util_prob.FUNSET_2'(s(fun(bool,bool),X125),s(fun(X122,bool),X126))),s(X127,X128)),X124,X123,fun(bool,X122),X121)))),s(bool,'F_0'))),s(fun(X122,bool),X126)))) | s(bool,'T_0') = s(bool,sK6(X126,X125,happ(s(fun(X121,fun(bool,X122)),X123),s(X121,sK6('const.combin.K_2'(s(fun(fun(bool,X122),bool),'const.util_prob.FUNSET_2'(s(fun(bool,bool),X125),s(fun(X122,bool),X126))),s(X127,X128)),X124,X123,fun(bool,X122),X121))),X122,bool))) )),
  inference(resolution,[],[f6707,f6690])).
fof(f8576,plain,(
  ( ! [X118,X116,X114,X120,X119,X117,X115,X113] : (p(s(bool,'const.bool.IN_2'(s(fun(X113,fun(X114,bool)),X115),s(fun(fun(X113,fun(X114,bool)),bool),'const.util_prob.FUNSET_2'(s(fun(X113,bool),X116),s(fun(fun(X114,bool),bool),'const.util_prob.FUNSET_2'(s(fun(X114,bool),X117),s(fun(bool,bool),X118)))))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X118)))) | s(bool,'T_0') = s(bool,happ(s(fun(X114,bool),happ(s(fun(X113,fun(X114,bool)),X115),s(X113,sK6('const.combin.K_2'(s(fun(fun(X114,bool),bool),'const.util_prob.FUNSET_2'(s(fun(X114,bool),X117),s(fun(bool,bool),X118))),s(X119,X120)),X116,X115,fun(X114,bool),X113)))),s(X114,sK6(X118,X117,happ(s(fun(X113,fun(X114,bool)),X115),s(X113,sK6('const.combin.K_2'(s(fun(fun(X114,bool),bool),'const.util_prob.FUNSET_2'(s(fun(X114,bool),X117),s(fun(bool,bool),X118))),s(X119,X120)),X116,X115,fun(X114,bool),X113))),bool,X114))))) )),
  inference(resolution,[],[f6707,f6691])).
fof(f8575,plain,(
  ( ! [X111,X109,X107,X105,X112,X110,X108,X106,X104] : (p(s(bool,'const.bool.IN_2'(s(fun(X104,fun(X105,X106)),X107),s(fun(fun(X104,fun(X105,X106)),bool),'const.util_prob.FUNSET_2'(s(fun(X104,bool),X108),s(fun(fun(X105,X106),bool),'const.util_prob.DFUNSET_2'(s(fun(X105,bool),X109),s(fun(X105,fun(X106,bool)),X110)))))))) | p(s(bool,'const.bool.IN_2'(s(X105,sK7(X110,X109,happ(s(fun(X104,fun(X105,X106)),X107),s(X104,sK6('const.combin.K_2'(s(fun(fun(X105,X106),bool),'const.util_prob.DFUNSET_2'(s(fun(X105,bool),X109),s(fun(X105,fun(X106,bool)),X110))),s(X111,X112)),X108,X107,fun(X105,X106),X104))),X105,X106)),s(fun(X105,bool),X109))))) )),
  inference(resolution,[],[f6707,f87])).
fof(f8573,plain,(
  ( ! [X94,X92,X90,X88,X87,X85,X93,X91,X89,X86,X84] : (p(s(bool,'const.bool.IN_2'(s(fun(X84,fun(X85,X86)),X87),s(fun(fun(X84,fun(X85,X86)),bool),'const.util_prob.FUNSET_2'(s(fun(X84,bool),X88),s(fun(fun(X85,X86),bool),'const.util_prob.DFUNSET_2'(s(fun(X85,bool),X89),s(fun(X85,fun(X86,bool)),X90)))))))) | p(s(bool,'const.bool.IN_2'(s(X85,sK7(X90,X89,'const.combin.K_2'(s(fun(X85,X86),happ(s(fun(X84,fun(X85,X86)),X87),s(X84,sK6('const.combin.K_2'(s(fun(fun(X85,X86),bool),'const.util_prob.DFUNSET_2'(s(fun(X85,bool),X89),s(fun(X85,fun(X86,bool)),X90))),s(X91,X92)),X88,X87,fun(X85,X86),X84)))),s(X93,X94)),X85,X86)),s(fun(X85,bool),X89))))) )),
  inference(resolution,[],[f6707,f6651])).
fof(f8572,plain,(
  ( ! [X80,X78,X76,X74,X83,X81,X79,X77,X75,X73,X82] : (p(s(bool,'const.bool.IN_2'(s(fun(X73,fun(X74,X75)),X76),s(fun(fun(X73,fun(X74,X75)),bool),'const.util_prob.FUNSET_2'(s(fun(X73,bool),X77),s(fun(fun(X74,X75),bool),'const.util_prob.DFUNSET_2'(s(fun(X74,bool),X78),s(fun(X74,fun(X75,bool)),X79)))))))) | p(s(bool,'const.bool.IN_2'(s(X74,sK7(X79,'const.combin.K_2'(s(fun(X74,bool),X78),s(X80,X81)),happ(s(fun(X73,fun(X74,X75)),X76),s(X73,sK6('const.combin.K_2'(s(fun(fun(X74,X75),bool),'const.util_prob.DFUNSET_2'(s(fun(X74,bool),X78),s(fun(X74,fun(X75,bool)),X79))),s(X82,X83)),X77,X76,fun(X74,X75),X73))),X74,X75)),s(fun(X74,bool),X78))))) )),
  inference(resolution,[],[f6707,f6673])).
fof(f8571,plain,(
  ( ! [X70,X68,X66,X64,X72,X62,X71,X69,X67,X65,X63] : (p(s(bool,'const.bool.IN_2'(s(fun(X62,fun(X63,X64)),X65),s(fun(fun(X62,fun(X63,X64)),bool),'const.util_prob.FUNSET_2'(s(fun(X62,bool),X66),s(fun(fun(X63,X64),bool),'const.util_prob.DFUNSET_2'(s(fun(X63,bool),X67),s(fun(X63,fun(X64,bool)),X68)))))))) | p(s(bool,'const.bool.IN_2'(s(X63,sK7('const.combin.K_2'(s(fun(X63,fun(X64,bool)),X68),s(X69,X70)),X67,happ(s(fun(X62,fun(X63,X64)),X65),s(X62,sK6('const.combin.K_2'(s(fun(fun(X63,X64),bool),'const.util_prob.DFUNSET_2'(s(fun(X63,bool),X67),s(fun(X63,fun(X64,bool)),X68))),s(X71,X72)),X66,X65,fun(X63,X64),X62))),X63,X64)),s(fun(X63,bool),X67))))) )),
  inference(resolution,[],[f6707,f6655])).
fof(f8570,plain,(
  ( ! [X61,X59,X57,X54,X60,X58,X56,X55,X53] : (p(s(bool,'const.bool.IN_2'(s(fun(X53,fun(X54,X55)),X56),s(fun(fun(X53,fun(X54,X55)),bool),'const.util_prob.FUNSET_2'(s(fun(X53,bool),X57),s(fun(fun(X54,X55),bool),'const.util_prob.FUNSET_2'(s(fun(X54,bool),X58),s(fun(X55,bool),X59)))))))) | p(s(bool,'const.bool.IN_2'(s(X54,sK6(X59,X58,happ(s(fun(X53,fun(X54,X55)),X56),s(X53,sK6('const.combin.K_2'(s(fun(fun(X54,X55),bool),'const.util_prob.FUNSET_2'(s(fun(X54,bool),X58),s(fun(X55,bool),X59))),s(X60,X61)),X57,X56,fun(X54,X55),X53))),X55,X54)),s(fun(X54,bool),X58))))) )),
  inference(resolution,[],[f6707,f84])).
fof(f8568,plain,(
  ( ! [X39,X37,X35,X33,X43,X41,X38,X36,X34,X42,X40] : (p(s(bool,'const.bool.IN_2'(s(fun(X33,fun(X34,X35)),X36),s(fun(fun(X33,fun(X34,X35)),bool),'const.util_prob.FUNSET_2'(s(fun(X33,bool),X37),s(fun(fun(X34,X35),bool),'const.util_prob.FUNSET_2'(s(fun(X34,bool),X38),s(fun(X35,bool),X39)))))))) | p(s(bool,'const.bool.IN_2'(s(X34,sK6(X39,X38,'const.combin.K_2'(s(fun(X34,X35),happ(s(fun(X33,fun(X34,X35)),X36),s(X33,sK6('const.combin.K_2'(s(fun(fun(X34,X35),bool),'const.util_prob.FUNSET_2'(s(fun(X34,bool),X38),s(fun(X35,bool),X39))),s(X40,X41)),X37,X36,fun(X34,X35),X33)))),s(X42,X43)),X35,X34)),s(fun(X34,bool),X38))))) )),
  inference(resolution,[],[f6707,f6616])).
fof(f8567,plain,(
  ( ! [X30,X28,X26,X24,X23,X31,X29,X27,X25,X32,X22] : (p(s(bool,'const.bool.IN_2'(s(fun(X22,fun(X23,X24)),X25),s(fun(fun(X22,fun(X23,X24)),bool),'const.util_prob.FUNSET_2'(s(fun(X22,bool),X26),s(fun(fun(X23,X24),bool),'const.util_prob.FUNSET_2'(s(fun(X23,bool),X27),s(fun(X24,bool),X28)))))))) | p(s(bool,'const.bool.IN_2'(s(X23,sK6('const.combin.K_2'(s(fun(X24,bool),X28),s(X29,X30)),X27,happ(s(fun(X22,fun(X23,X24)),X25),s(X22,sK6('const.combin.K_2'(s(fun(fun(X23,X24),bool),'const.util_prob.FUNSET_2'(s(fun(X23,bool),X27),s(fun(X24,bool),X28))),s(X31,X32)),X26,X25,fun(X23,X24),X22))),X24,X23)),s(fun(X23,bool),X27))))) )),
  inference(resolution,[],[f6707,f6620])).
fof(f8566,plain,(
  ( ! [X14,X12,X21,X19,X17,X15,X13,X11,X20,X18,X16] : (p(s(bool,'const.bool.IN_2'(s(fun(X11,fun(X12,X13)),X14),s(fun(fun(X11,fun(X12,X13)),bool),'const.util_prob.FUNSET_2'(s(fun(X11,bool),X15),s(fun(fun(X12,X13),bool),'const.util_prob.FUNSET_2'(s(fun(X12,bool),X16),s(fun(X13,bool),X17)))))))) | p(s(bool,'const.bool.IN_2'(s(X12,sK6(X17,'const.combin.K_2'(s(fun(X12,bool),X16),s(X18,X19)),happ(s(fun(X11,fun(X12,X13)),X14),s(X11,sK6('const.combin.K_2'(s(fun(fun(X12,X13),bool),'const.util_prob.FUNSET_2'(s(fun(X12,bool),X16),s(fun(X13,bool),X17))),s(X20,X21)),X15,X14,fun(X12,X13),X11))),X13,X12)),s(fun(X12,bool),X16))))) )),
  inference(resolution,[],[f6707,f6637])).
fof(f8565,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,fun(X1,X2)),X3),s(fun(fun(X0,fun(X1,X2)),bool),'const.util_prob.FUNSET_2'(s(fun(X0,bool),X4),s(fun(fun(X1,X2),bool),'const.util_prob.FUNSET_2'(s(fun(X1,bool),X5),s(fun(X2,bool),X6)))))))) | p(s(bool,'const.bool.IN_2'(s(X1,sK6(X6,happ(s(fun(X7,fun(X1,bool)),'const.combin.K_1'(s(fun(X1,bool),X5))),s(X7,X8)),happ(s(fun(X0,fun(X1,X2)),X3),s(X0,sK6('const.combin.K_2'(s(fun(fun(X1,X2),bool),'const.util_prob.FUNSET_2'(s(fun(X1,bool),X5),s(fun(X2,bool),X6))),s(X9,X10)),X4,X3,fun(X1,X2),X0))),X2,X1)),s(fun(X1,bool),X5))))) )),
  inference(resolution,[],[f6707,f6640])).
fof(f8558,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X9,X16] : (p(s(bool,'const.bool.IN_2'(s(fun(X13,X9),X14),s(fun(fun(X13,X9),bool),'const.util_prob.FUNSET_2'(s(fun(X13,bool),X15),s(fun(X9,bool),X11)))))) | ~p(s(bool,'const.bool.IN_2'(s(X9,happ(s(fun(X13,X9),X14),s(X13,sK6(happ(s(fun(X10,fun(X9,bool)),'const.combin.K_1'(s(fun(X9,bool),X11))),s(X10,X12)),X15,'const.combin.K_2'(s(fun(X13,X9),X14),s(X16,X17)),X9,X13)))),s(fun(X9,bool),X11))))) )),
  inference(forward_demodulation,[],[f8557,f80])).
fof(f8557,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X9,X16] : (p(s(bool,'const.bool.IN_2'(s(fun(X13,X9),X14),s(fun(fun(X13,X9),bool),'const.util_prob.FUNSET_2'(s(fun(X13,bool),X15),s(fun(X9,bool),'const.combin.K_2'(s(fun(X9,bool),X11),s(X10,X12)))))))) | ~p(s(bool,'const.bool.IN_2'(s(X9,happ(s(fun(X13,X9),X14),s(X13,sK6(happ(s(fun(X10,fun(X9,bool)),'const.combin.K_1'(s(fun(X9,bool),X11))),s(X10,X12)),X15,'const.combin.K_2'(s(fun(X13,X9),X14),s(X16,X17)),X9,X13)))),s(fun(X9,bool),X11))))) )),
  inference(forward_demodulation,[],[f8556,f81])).
fof(f8556,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X9,X16] : (~p(s(bool,'const.bool.IN_2'(s(X9,happ(s(fun(X13,X9),X14),s(X13,sK6(happ(s(fun(X10,fun(X9,bool)),'const.combin.K_1'(s(fun(X9,bool),X11))),s(X10,X12)),X15,'const.combin.K_2'(s(fun(X13,X9),X14),s(X16,X17)),X9,X13)))),s(fun(X9,bool),X11)))) | p(s(bool,'const.bool.IN_2'(s(fun(X13,X9),X14),s(fun(fun(X13,X9),bool),'const.util_prob.FUNSET_2'(s(fun(X13,bool),X15),s(fun(X9,bool),happ(s(fun(X10,fun(X9,bool)),'const.combin.K_1'(s(fun(X9,bool),X11))),s(X10,X12))))))))) )),
  inference(forward_demodulation,[],[f8542,f80])).
fof(f8542,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X9,X16] : (~p(s(bool,'const.bool.IN_2'(s(X9,happ(s(fun(X13,X9),X14),s(X13,sK6(happ(s(fun(X10,fun(X9,bool)),'const.combin.K_1'(s(fun(X9,bool),X11))),s(X10,X12)),X15,'const.combin.K_2'(s(fun(X13,X9),X14),s(X16,X17)),X9,X13)))),s(fun(X9,bool),'const.combin.K_2'(s(fun(X9,bool),X11),s(X10,X12)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X13,X9),X14),s(fun(fun(X13,X9),bool),'const.util_prob.FUNSET_2'(s(fun(X13,bool),X15),s(fun(X9,bool),happ(s(fun(X10,fun(X9,bool)),'const.combin.K_1'(s(fun(X9,bool),X11))),s(X10,X12))))))))) )),
  inference(superposition,[],[f6702,f81])).
fof(f8555,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(X4,X0),X5),s(fun(fun(X4,X0),bool),'const.util_prob.FUNSET_2'(s(fun(X4,bool),X6),s(fun(X0,bool),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X4,X0),X5),s(X4,sK6('const.combin.K_2'(s(fun(X0,bool),X1),s(X2,X3)),X6,'const.combin.K_2'(s(fun(X4,X0),X5),s(X7,X8)),X0,X4)))),s(fun(X0,bool),X1))))) )),
  inference(forward_demodulation,[],[f8541,f80])).
fof(f8541,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X4,X0),X5),s(X4,sK6('const.combin.K_2'(s(fun(X0,bool),X1),s(X2,X3)),X6,'const.combin.K_2'(s(fun(X4,X0),X5),s(X7,X8)),X0,X4)))),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(fun(X4,X0),X5),s(fun(fun(X4,X0),bool),'const.util_prob.FUNSET_2'(s(fun(X4,bool),X6),s(fun(X0,bool),'const.combin.K_2'(s(fun(X0,bool),X1),s(X2,X3))))))))) )),
  inference(superposition,[],[f6702,f80])).
fof(f8508,plain,(
  ( ! [X144,X142,X140,X138,X143,X141,X139,X137] : (p(s(bool,'const.bool.IN_2'(s(fun(X137,fun(bool,X138)),X139),s(fun(fun(X137,fun(bool,X138)),bool),'const.util_prob.FUNSET_2'(s(fun(X137,bool),X140),s(fun(fun(bool,X138),bool),'const.util_prob.DFUNSET_2'(s(fun(bool,bool),X141),s(fun(bool,fun(X138,bool)),X142)))))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X141)))) | s(bool,'T_0') = s(bool,sK7(X142,X141,happ(s(fun(X137,fun(bool,X138)),X139),s(X137,sK6('const.util_prob.DFUNSET_2'(s(fun(bool,bool),X141),s(fun(bool,fun(X138,bool)),X142)),X140,'const.combin.K_2'(s(fun(X137,fun(bool,X138)),X139),s(X143,X144)),fun(bool,X138),X137))),bool,X138))) )),
  inference(resolution,[],[f6702,f6662])).
fof(f8507,plain,(
  ( ! [X134,X132,X130,X136,X135,X133,X131,X129] : (p(s(bool,'const.bool.IN_2'(s(fun(X129,fun(bool,X130)),X131),s(fun(fun(X129,fun(bool,X130)),bool),'const.util_prob.FUNSET_2'(s(fun(X129,bool),X132),s(fun(fun(bool,X130),bool),'const.util_prob.FUNSET_2'(s(fun(bool,bool),X133),s(fun(X130,bool),X134)))))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X133)))) | s(bool,'T_0') = s(bool,sK6(X134,X133,happ(s(fun(X129,fun(bool,X130)),X131),s(X129,sK6('const.util_prob.FUNSET_2'(s(fun(bool,bool),X133),s(fun(X130,bool),X134)),X132,'const.combin.K_2'(s(fun(X129,fun(bool,X130)),X131),s(X135,X136)),fun(bool,X130),X129))),X130,bool))) )),
  inference(resolution,[],[f6702,f6627])).
fof(f8506,plain,(
  ( ! [X127,X125,X123,X121,X128,X126,X124,X122] : (p(s(bool,'const.bool.IN_2'(s(fun(X121,fun(bool,X122)),X123),s(fun(fun(X121,fun(bool,X122)),bool),'const.util_prob.FUNSET_2'(s(fun(X121,bool),X124),s(fun(fun(bool,X122),bool),'const.util_prob.FUNSET_2'(s(fun(bool,bool),X125),s(fun(X122,bool),X126)))))))) | ~p(s(bool,'const.bool.IN_2'(s(X122,happ(s(fun(bool,X122),happ(s(fun(X121,fun(bool,X122)),X123),s(X121,sK6('const.util_prob.FUNSET_2'(s(fun(bool,bool),X125),s(fun(X122,bool),X126)),X124,'const.combin.K_2'(s(fun(X121,fun(bool,X122)),X123),s(X127,X128)),fun(bool,X122),X121)))),s(bool,'F_0'))),s(fun(X122,bool),X126)))) | s(bool,'T_0') = s(bool,sK6(X126,X125,happ(s(fun(X121,fun(bool,X122)),X123),s(X121,sK6('const.util_prob.FUNSET_2'(s(fun(bool,bool),X125),s(fun(X122,bool),X126)),X124,'const.combin.K_2'(s(fun(X121,fun(bool,X122)),X123),s(X127,X128)),fun(bool,X122),X121))),X122,bool))) )),
  inference(resolution,[],[f6702,f6690])).
fof(f8505,plain,(
  ( ! [X118,X116,X114,X120,X119,X117,X115,X113] : (p(s(bool,'const.bool.IN_2'(s(fun(X113,fun(X114,bool)),X115),s(fun(fun(X113,fun(X114,bool)),bool),'const.util_prob.FUNSET_2'(s(fun(X113,bool),X116),s(fun(fun(X114,bool),bool),'const.util_prob.FUNSET_2'(s(fun(X114,bool),X117),s(fun(bool,bool),X118)))))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X118)))) | s(bool,'T_0') = s(bool,happ(s(fun(X114,bool),happ(s(fun(X113,fun(X114,bool)),X115),s(X113,sK6('const.util_prob.FUNSET_2'(s(fun(X114,bool),X117),s(fun(bool,bool),X118)),X116,'const.combin.K_2'(s(fun(X113,fun(X114,bool)),X115),s(X119,X120)),fun(X114,bool),X113)))),s(X114,sK6(X118,X117,happ(s(fun(X113,fun(X114,bool)),X115),s(X113,sK6('const.util_prob.FUNSET_2'(s(fun(X114,bool),X117),s(fun(bool,bool),X118)),X116,'const.combin.K_2'(s(fun(X113,fun(X114,bool)),X115),s(X119,X120)),fun(X114,bool),X113))),bool,X114))))) )),
  inference(resolution,[],[f6702,f6691])).
fof(f8504,plain,(
  ( ! [X111,X109,X107,X105,X112,X110,X108,X106,X104] : (p(s(bool,'const.bool.IN_2'(s(fun(X104,fun(X105,X106)),X107),s(fun(fun(X104,fun(X105,X106)),bool),'const.util_prob.FUNSET_2'(s(fun(X104,bool),X108),s(fun(fun(X105,X106),bool),'const.util_prob.DFUNSET_2'(s(fun(X105,bool),X109),s(fun(X105,fun(X106,bool)),X110)))))))) | p(s(bool,'const.bool.IN_2'(s(X105,sK7(X110,X109,happ(s(fun(X104,fun(X105,X106)),X107),s(X104,sK6('const.util_prob.DFUNSET_2'(s(fun(X105,bool),X109),s(fun(X105,fun(X106,bool)),X110)),X108,'const.combin.K_2'(s(fun(X104,fun(X105,X106)),X107),s(X111,X112)),fun(X105,X106),X104))),X105,X106)),s(fun(X105,bool),X109))))) )),
  inference(resolution,[],[f6702,f87])).
fof(f8502,plain,(
  ( ! [X94,X92,X90,X88,X87,X85,X93,X91,X89,X86,X84] : (p(s(bool,'const.bool.IN_2'(s(fun(X84,fun(X85,X86)),X87),s(fun(fun(X84,fun(X85,X86)),bool),'const.util_prob.FUNSET_2'(s(fun(X84,bool),X88),s(fun(fun(X85,X86),bool),'const.util_prob.DFUNSET_2'(s(fun(X85,bool),X89),s(fun(X85,fun(X86,bool)),X90)))))))) | p(s(bool,'const.bool.IN_2'(s(X85,sK7(X90,X89,'const.combin.K_2'(s(fun(X85,X86),happ(s(fun(X84,fun(X85,X86)),X87),s(X84,sK6('const.util_prob.DFUNSET_2'(s(fun(X85,bool),X89),s(fun(X85,fun(X86,bool)),X90)),X88,'const.combin.K_2'(s(fun(X84,fun(X85,X86)),X87),s(X91,X92)),fun(X85,X86),X84)))),s(X93,X94)),X85,X86)),s(fun(X85,bool),X89))))) )),
  inference(resolution,[],[f6702,f6651])).
fof(f8501,plain,(
  ( ! [X80,X78,X76,X74,X83,X81,X79,X77,X75,X73,X82] : (p(s(bool,'const.bool.IN_2'(s(fun(X73,fun(X74,X75)),X76),s(fun(fun(X73,fun(X74,X75)),bool),'const.util_prob.FUNSET_2'(s(fun(X73,bool),X77),s(fun(fun(X74,X75),bool),'const.util_prob.DFUNSET_2'(s(fun(X74,bool),X78),s(fun(X74,fun(X75,bool)),X79)))))))) | p(s(bool,'const.bool.IN_2'(s(X74,sK7(X79,'const.combin.K_2'(s(fun(X74,bool),X78),s(X80,X81)),happ(s(fun(X73,fun(X74,X75)),X76),s(X73,sK6('const.util_prob.DFUNSET_2'(s(fun(X74,bool),X78),s(fun(X74,fun(X75,bool)),X79)),X77,'const.combin.K_2'(s(fun(X73,fun(X74,X75)),X76),s(X82,X83)),fun(X74,X75),X73))),X74,X75)),s(fun(X74,bool),X78))))) )),
  inference(resolution,[],[f6702,f6673])).
fof(f8500,plain,(
  ( ! [X70,X68,X66,X64,X72,X62,X71,X69,X67,X65,X63] : (p(s(bool,'const.bool.IN_2'(s(fun(X62,fun(X63,X64)),X65),s(fun(fun(X62,fun(X63,X64)),bool),'const.util_prob.FUNSET_2'(s(fun(X62,bool),X66),s(fun(fun(X63,X64),bool),'const.util_prob.DFUNSET_2'(s(fun(X63,bool),X67),s(fun(X63,fun(X64,bool)),X68)))))))) | p(s(bool,'const.bool.IN_2'(s(X63,sK7('const.combin.K_2'(s(fun(X63,fun(X64,bool)),X68),s(X69,X70)),X67,happ(s(fun(X62,fun(X63,X64)),X65),s(X62,sK6('const.util_prob.DFUNSET_2'(s(fun(X63,bool),X67),s(fun(X63,fun(X64,bool)),X68)),X66,'const.combin.K_2'(s(fun(X62,fun(X63,X64)),X65),s(X71,X72)),fun(X63,X64),X62))),X63,X64)),s(fun(X63,bool),X67))))) )),
  inference(resolution,[],[f6702,f6655])).
fof(f8499,plain,(
  ( ! [X61,X59,X57,X54,X60,X58,X56,X55,X53] : (p(s(bool,'const.bool.IN_2'(s(fun(X53,fun(X54,X55)),X56),s(fun(fun(X53,fun(X54,X55)),bool),'const.util_prob.FUNSET_2'(s(fun(X53,bool),X57),s(fun(fun(X54,X55),bool),'const.util_prob.FUNSET_2'(s(fun(X54,bool),X58),s(fun(X55,bool),X59)))))))) | p(s(bool,'const.bool.IN_2'(s(X54,sK6(X59,X58,happ(s(fun(X53,fun(X54,X55)),X56),s(X53,sK6('const.util_prob.FUNSET_2'(s(fun(X54,bool),X58),s(fun(X55,bool),X59)),X57,'const.combin.K_2'(s(fun(X53,fun(X54,X55)),X56),s(X60,X61)),fun(X54,X55),X53))),X55,X54)),s(fun(X54,bool),X58))))) )),
  inference(resolution,[],[f6702,f84])).
fof(f8497,plain,(
  ( ! [X39,X37,X35,X33,X43,X41,X38,X36,X34,X42,X40] : (p(s(bool,'const.bool.IN_2'(s(fun(X33,fun(X34,X35)),X36),s(fun(fun(X33,fun(X34,X35)),bool),'const.util_prob.FUNSET_2'(s(fun(X33,bool),X37),s(fun(fun(X34,X35),bool),'const.util_prob.FUNSET_2'(s(fun(X34,bool),X38),s(fun(X35,bool),X39)))))))) | p(s(bool,'const.bool.IN_2'(s(X34,sK6(X39,X38,'const.combin.K_2'(s(fun(X34,X35),happ(s(fun(X33,fun(X34,X35)),X36),s(X33,sK6('const.util_prob.FUNSET_2'(s(fun(X34,bool),X38),s(fun(X35,bool),X39)),X37,'const.combin.K_2'(s(fun(X33,fun(X34,X35)),X36),s(X40,X41)),fun(X34,X35),X33)))),s(X42,X43)),X35,X34)),s(fun(X34,bool),X38))))) )),
  inference(resolution,[],[f6702,f6616])).
fof(f8496,plain,(
  ( ! [X30,X28,X26,X24,X23,X31,X29,X27,X25,X32,X22] : (p(s(bool,'const.bool.IN_2'(s(fun(X22,fun(X23,X24)),X25),s(fun(fun(X22,fun(X23,X24)),bool),'const.util_prob.FUNSET_2'(s(fun(X22,bool),X26),s(fun(fun(X23,X24),bool),'const.util_prob.FUNSET_2'(s(fun(X23,bool),X27),s(fun(X24,bool),X28)))))))) | p(s(bool,'const.bool.IN_2'(s(X23,sK6('const.combin.K_2'(s(fun(X24,bool),X28),s(X29,X30)),X27,happ(s(fun(X22,fun(X23,X24)),X25),s(X22,sK6('const.util_prob.FUNSET_2'(s(fun(X23,bool),X27),s(fun(X24,bool),X28)),X26,'const.combin.K_2'(s(fun(X22,fun(X23,X24)),X25),s(X31,X32)),fun(X23,X24),X22))),X24,X23)),s(fun(X23,bool),X27))))) )),
  inference(resolution,[],[f6702,f6620])).
fof(f8495,plain,(
  ( ! [X14,X12,X21,X19,X17,X15,X13,X11,X20,X18,X16] : (p(s(bool,'const.bool.IN_2'(s(fun(X11,fun(X12,X13)),X14),s(fun(fun(X11,fun(X12,X13)),bool),'const.util_prob.FUNSET_2'(s(fun(X11,bool),X15),s(fun(fun(X12,X13),bool),'const.util_prob.FUNSET_2'(s(fun(X12,bool),X16),s(fun(X13,bool),X17)))))))) | p(s(bool,'const.bool.IN_2'(s(X12,sK6(X17,'const.combin.K_2'(s(fun(X12,bool),X16),s(X18,X19)),happ(s(fun(X11,fun(X12,X13)),X14),s(X11,sK6('const.util_prob.FUNSET_2'(s(fun(X12,bool),X16),s(fun(X13,bool),X17)),X15,'const.combin.K_2'(s(fun(X11,fun(X12,X13)),X14),s(X20,X21)),fun(X12,X13),X11))),X13,X12)),s(fun(X12,bool),X16))))) )),
  inference(resolution,[],[f6702,f6637])).
fof(f8494,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,fun(X1,X2)),X3),s(fun(fun(X0,fun(X1,X2)),bool),'const.util_prob.FUNSET_2'(s(fun(X0,bool),X4),s(fun(fun(X1,X2),bool),'const.util_prob.FUNSET_2'(s(fun(X1,bool),X5),s(fun(X2,bool),X6)))))))) | p(s(bool,'const.bool.IN_2'(s(X1,sK6(X6,happ(s(fun(X7,fun(X1,bool)),'const.combin.K_1'(s(fun(X1,bool),X5))),s(X7,X8)),happ(s(fun(X0,fun(X1,X2)),X3),s(X0,sK6('const.util_prob.FUNSET_2'(s(fun(X1,bool),X5),s(fun(X2,bool),X6)),X4,'const.combin.K_2'(s(fun(X0,fun(X1,X2)),X3),s(X9,X10)),fun(X1,X2),X0))),X2,X1)),s(fun(X1,bool),X5))))) )),
  inference(resolution,[],[f6702,f6640])).
fof(f8493,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X7)))) | p(s(bool,'const.bool.IN_2'(s(fun(X9,bool),X10),s(fun(fun(X9,bool),bool),'const.util_prob.FUNSET_2'(s(fun(X9,bool),X11),s(fun(bool,bool),X7)))))) | s(bool,'T_0') = s(bool,happ(s(fun(X9,bool),X10),s(X9,sK6(happ(s(fun(X6,fun(bool,bool)),'const.combin.K_1'(s(fun(bool,bool),X7))),s(X6,X8)),X11,X10,bool,X9))))) )),
  inference(forward_demodulation,[],[f8492,f80])).
fof(f8492,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.combin.K_2'(s(fun(bool,bool),X7),s(X6,X8)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X9,bool),X10),s(fun(fun(X9,bool),bool),'const.util_prob.FUNSET_2'(s(fun(X9,bool),X11),s(fun(bool,bool),X7)))))) | s(bool,'T_0') = s(bool,happ(s(fun(X9,bool),X10),s(X9,sK6(happ(s(fun(X6,fun(bool,bool)),'const.combin.K_1'(s(fun(bool,bool),X7))),s(X6,X8)),X11,X10,bool,X9))))) )),
  inference(forward_demodulation,[],[f8491,f81])).
fof(f8491,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(fun(X9,bool),X10),s(fun(fun(X9,bool),bool),'const.util_prob.FUNSET_2'(s(fun(X9,bool),X11),s(fun(bool,bool),X7)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),happ(s(fun(X6,fun(bool,bool)),'const.combin.K_1'(s(fun(bool,bool),X7))),s(X6,X8)))))) | s(bool,'T_0') = s(bool,happ(s(fun(X9,bool),X10),s(X9,sK6(happ(s(fun(X6,fun(bool,bool)),'const.combin.K_1'(s(fun(bool,bool),X7))),s(X6,X8)),X11,X10,bool,X9))))) )),
  inference(forward_demodulation,[],[f8481,f80])).
fof(f8481,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(fun(X9,bool),X10),s(fun(fun(X9,bool),bool),'const.util_prob.FUNSET_2'(s(fun(X9,bool),X11),s(fun(bool,bool),'const.combin.K_2'(s(fun(bool,bool),X7),s(X6,X8)))))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),happ(s(fun(X6,fun(bool,bool)),'const.combin.K_1'(s(fun(bool,bool),X7))),s(X6,X8)))))) | s(bool,'T_0') = s(bool,happ(s(fun(X9,bool),X10),s(X9,sK6(happ(s(fun(X6,fun(bool,bool)),'const.combin.K_1'(s(fun(bool,bool),X7))),s(X6,X8)),X11,X10,bool,X9))))) )),
  inference(superposition,[],[f6691,f81])).
fof(f8490,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))) | p(s(bool,'const.bool.IN_2'(s(fun(X3,bool),X4),s(fun(fun(X3,bool),bool),'const.util_prob.FUNSET_2'(s(fun(X3,bool),X5),s(fun(bool,bool),X0)))))) | s(bool,'T_0') = s(bool,happ(s(fun(X3,bool),X4),s(X3,sK6('const.combin.K_2'(s(fun(bool,bool),X0),s(X1,X2)),X5,X4,bool,X3))))) )),
  inference(forward_demodulation,[],[f8480,f80])).
fof(f8480,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(X3,bool),X4),s(fun(fun(X3,bool),bool),'const.util_prob.FUNSET_2'(s(fun(X3,bool),X5),s(fun(bool,bool),X0)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.combin.K_2'(s(fun(bool,bool),X0),s(X1,X2)))))) | s(bool,'T_0') = s(bool,happ(s(fun(X3,bool),X4),s(X3,sK6('const.combin.K_2'(s(fun(bool,bool),X0),s(X1,X2)),X5,X4,bool,X3))))) )),
  inference(superposition,[],[f6691,f80])).
fof(f8489,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),X10),s(fun(fun(X6,bool),bool),'const.util_prob.FUNSET_2'(s(fun(X6,bool),X8),s(fun(bool,bool),X11)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X11)))) | s(bool,'T_0') = s(bool,happ(s(fun(X6,bool),X10),s(X6,sK6(X11,happ(s(fun(X7,fun(X6,bool)),'const.combin.K_1'(s(fun(X6,bool),X8))),s(X7,X9)),X10,bool,X6))))) )),
  inference(forward_demodulation,[],[f8479,f80])).
fof(f8479,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),X10),s(fun(fun(X6,bool),bool),'const.util_prob.FUNSET_2'(s(fun(X6,bool),'const.combin.K_2'(s(fun(X6,bool),X8),s(X7,X9))),s(fun(bool,bool),X11)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X11)))) | s(bool,'T_0') = s(bool,happ(s(fun(X6,bool),X10),s(X6,sK6(X11,happ(s(fun(X7,fun(X6,bool)),'const.combin.K_1'(s(fun(X6,bool),X8))),s(X7,X9)),X10,bool,X6))))) )),
  inference(superposition,[],[f6691,f81])).
fof(f8478,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X4),s(fun(fun(X0,bool),bool),'const.util_prob.FUNSET_2'(s(fun(X0,bool),X1),s(fun(bool,bool),X5)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X5)))) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X4),s(X0,sK6(X5,'const.combin.K_2'(s(fun(X0,bool),X1),s(X2,X3)),X4,bool,X0))))) )),
  inference(superposition,[],[f6691,f80])).
fof(f8488,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (s(bool,'T_0') = s(bool,happ(s(fun(X6,bool),X8),s(X6,sK6(X11,X10,happ(s(fun(X7,fun(X6,bool)),'const.combin.K_1'(s(fun(X6,bool),X8))),s(X7,X9)),bool,X6)))) | p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),X8),s(fun(fun(X6,bool),bool),'const.util_prob.FUNSET_2'(s(fun(X6,bool),X10),s(fun(bool,bool),X11)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X11))))) )),
  inference(forward_demodulation,[],[f8487,f80])).
fof(f8487,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (s(bool,'T_0') = s(bool,happ(s(fun(X6,bool),'const.combin.K_2'(s(fun(X6,bool),X8),s(X7,X9))),s(X6,sK6(X11,X10,happ(s(fun(X7,fun(X6,bool)),'const.combin.K_1'(s(fun(X6,bool),X8))),s(X7,X9)),bool,X6)))) | p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),X8),s(fun(fun(X6,bool),bool),'const.util_prob.FUNSET_2'(s(fun(X6,bool),X10),s(fun(bool,bool),X11)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X11))))) )),
  inference(forward_demodulation,[],[f8486,f81])).
fof(f8486,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),X8),s(fun(fun(X6,bool),bool),'const.util_prob.FUNSET_2'(s(fun(X6,bool),X10),s(fun(bool,bool),X11)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X11)))) | s(bool,'T_0') = s(bool,happ(s(fun(X6,bool),happ(s(fun(X7,fun(X6,bool)),'const.combin.K_1'(s(fun(X6,bool),X8))),s(X7,X9))),s(X6,sK6(X11,X10,happ(s(fun(X7,fun(X6,bool)),'const.combin.K_1'(s(fun(X6,bool),X8))),s(X7,X9)),bool,X6))))) )),
  inference(forward_demodulation,[],[f8477,f80])).
fof(f8477,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),'const.combin.K_2'(s(fun(X6,bool),X8),s(X7,X9))),s(fun(fun(X6,bool),bool),'const.util_prob.FUNSET_2'(s(fun(X6,bool),X10),s(fun(bool,bool),X11)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X11)))) | s(bool,'T_0') = s(bool,happ(s(fun(X6,bool),happ(s(fun(X7,fun(X6,bool)),'const.combin.K_1'(s(fun(X6,bool),X8))),s(X7,X9))),s(X6,sK6(X11,X10,happ(s(fun(X7,fun(X6,bool)),'const.combin.K_1'(s(fun(X6,bool),X8))),s(X7,X9)),bool,X6))))) )),
  inference(superposition,[],[f6691,f81])).
fof(f8485,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK6(X5,X4,'const.combin.K_2'(s(fun(X0,bool),X1),s(X2,X3)),bool,X0)))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X1),s(fun(fun(X0,bool),bool),'const.util_prob.FUNSET_2'(s(fun(X0,bool),X4),s(fun(bool,bool),X5)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X5))))) )),
  inference(forward_demodulation,[],[f8476,f80])).
fof(f8476,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X1),s(fun(fun(X0,bool),bool),'const.util_prob.FUNSET_2'(s(fun(X0,bool),X4),s(fun(bool,bool),X5)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X5)))) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),'const.combin.K_2'(s(fun(X0,bool),X1),s(X2,X3))),s(X0,sK6(X5,X4,'const.combin.K_2'(s(fun(X0,bool),X1),s(X2,X3)),bool,X0))))) )),
  inference(superposition,[],[f6691,f80])).
fof(f8474,plain,(
  ( ! [X23,X21,X19,X17,X22,X20,X18] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X17)))) | s(bool,'T_0') = s(bool,happ(s(fun(X18,bool),X19),s(X18,sK6(X17,X20,X19,bool,X18)))) | p(s(bool,'const.bool.IN_2'(s(X21,happ(s(fun(fun(X18,bool),X21),X22),s(fun(X18,bool),X19))),s(fun(X21,bool),X23)))) | p(s(bool,'const.bool.IN_2'(s(fun(X18,bool),sK6(X23,'const.util_prob.FUNSET_2'(s(fun(X18,bool),X20),s(fun(bool,bool),X17)),X22,X21,fun(X18,bool))),s(fun(fun(X18,bool),bool),'const.util_prob.FUNSET_2'(s(fun(X18,bool),X20),s(fun(bool,bool),X17))))))) )),
  inference(resolution,[],[f6691,f6711])).
fof(f8472,plain,(
  ( ! [X6,X10,X8,X7,X5,X9] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X5)))) | s(bool,'T_0') = s(bool,happ(s(fun(X6,bool),happ(s(fun(X7,fun(X6,bool)),X8),s(X7,sK6('const.util_prob.FUNSET_2'(s(fun(X6,bool),X9),s(fun(bool,bool),X5)),X10,X8,fun(X6,bool),X7)))),s(X6,sK6(X5,X9,happ(s(fun(X7,fun(X6,bool)),X8),s(X7,sK6('const.util_prob.FUNSET_2'(s(fun(X6,bool),X9),s(fun(bool,bool),X5)),X10,X8,fun(X6,bool),X7))),bool,X6)))) | p(s(bool,'const.bool.IN_2'(s(fun(X7,fun(X6,bool)),X8),s(fun(fun(X7,fun(X6,bool)),bool),'const.util_prob.FUNSET_2'(s(fun(X7,bool),X10),s(fun(fun(X6,bool),bool),'const.util_prob.FUNSET_2'(s(fun(X6,bool),X9),s(fun(bool,bool),X5))))))))) )),
  inference(resolution,[],[f6691,f85])).
fof(f8471,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))) | s(bool,'T_0') = s(bool,happ(s(fun(X1,bool),X2),s(X1,sK6(X0,X3,X2,bool,X1)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X4),s(fun(X1,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(bool,happ(s(fun(X1,bool),X2),s(X1,X4))),s(fun(bool,bool),X0))))) )),
  inference(resolution,[],[f6691,f83])).
fof(f8470,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (~p(s(bool,'const.bool.IN_2'(s(X6,happ(s(fun(bool,X6),X10),s(bool,'F_0'))),s(fun(X6,bool),X8)))) | p(s(bool,'const.bool.IN_2'(s(fun(bool,X6),X10),s(fun(fun(bool,X6),bool),'const.util_prob.FUNSET_2'(s(fun(bool,bool),X11),s(fun(X6,bool),X8)))))) | s(bool,'T_0') = s(bool,sK6(happ(s(fun(X7,fun(X6,bool)),'const.combin.K_1'(s(fun(X6,bool),X8))),s(X7,X9)),X11,X10,X6,bool))) )),
  inference(forward_demodulation,[],[f8469,f80])).
fof(f8469,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (~p(s(bool,'const.bool.IN_2'(s(X6,happ(s(fun(bool,X6),X10),s(bool,'F_0'))),s(fun(X6,bool),'const.combin.K_2'(s(fun(X6,bool),X8),s(X7,X9)))))) | p(s(bool,'const.bool.IN_2'(s(fun(bool,X6),X10),s(fun(fun(bool,X6),bool),'const.util_prob.FUNSET_2'(s(fun(bool,bool),X11),s(fun(X6,bool),X8)))))) | s(bool,'T_0') = s(bool,sK6(happ(s(fun(X7,fun(X6,bool)),'const.combin.K_1'(s(fun(X6,bool),X8))),s(X7,X9)),X11,X10,X6,bool))) )),
  inference(forward_demodulation,[],[f8468,f81])).
fof(f8468,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(fun(bool,X6),X10),s(fun(fun(bool,X6),bool),'const.util_prob.FUNSET_2'(s(fun(bool,bool),X11),s(fun(X6,bool),X8)))))) | ~p(s(bool,'const.bool.IN_2'(s(X6,happ(s(fun(bool,X6),X10),s(bool,'F_0'))),s(fun(X6,bool),happ(s(fun(X7,fun(X6,bool)),'const.combin.K_1'(s(fun(X6,bool),X8))),s(X7,X9)))))) | s(bool,'T_0') = s(bool,sK6(happ(s(fun(X7,fun(X6,bool)),'const.combin.K_1'(s(fun(X6,bool),X8))),s(X7,X9)),X11,X10,X6,bool))) )),
  inference(forward_demodulation,[],[f8460,f80])).
fof(f8460,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(fun(bool,X6),X10),s(fun(fun(bool,X6),bool),'const.util_prob.FUNSET_2'(s(fun(bool,bool),X11),s(fun(X6,bool),'const.combin.K_2'(s(fun(X6,bool),X8),s(X7,X9)))))))) | ~p(s(bool,'const.bool.IN_2'(s(X6,happ(s(fun(bool,X6),X10),s(bool,'F_0'))),s(fun(X6,bool),happ(s(fun(X7,fun(X6,bool)),'const.combin.K_1'(s(fun(X6,bool),X8))),s(X7,X9)))))) | s(bool,'T_0') = s(bool,sK6(happ(s(fun(X7,fun(X6,bool)),'const.combin.K_1'(s(fun(X6,bool),X8))),s(X7,X9)),X11,X10,X6,bool))) )),
  inference(superposition,[],[f6690,f81])).
fof(f8467,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(bool,X0),X4),s(bool,'F_0'))),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(fun(bool,X0),X4),s(fun(fun(bool,X0),bool),'const.util_prob.FUNSET_2'(s(fun(bool,bool),X5),s(fun(X0,bool),X1)))))) | s(bool,'T_0') = s(bool,sK6('const.combin.K_2'(s(fun(X0,bool),X1),s(X2,X3)),X5,X4,X0,bool))) )),
  inference(forward_demodulation,[],[f8459,f80])).
fof(f8459,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(bool,X0),X4),s(fun(fun(bool,X0),bool),'const.util_prob.FUNSET_2'(s(fun(bool,bool),X5),s(fun(X0,bool),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(bool,X0),X4),s(bool,'F_0'))),s(fun(X0,bool),'const.combin.K_2'(s(fun(X0,bool),X1),s(X2,X3)))))) | s(bool,'T_0') = s(bool,sK6('const.combin.K_2'(s(fun(X0,bool),X1),s(X2,X3)),X5,X4,X0,bool))) )),
  inference(superposition,[],[f6690,f80])).
fof(f8466,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(fun(bool,X9),X10),s(fun(fun(bool,X9),bool),'const.util_prob.FUNSET_2'(s(fun(bool,bool),X7),s(fun(X9,bool),X11)))))) | ~p(s(bool,'const.bool.IN_2'(s(X9,happ(s(fun(bool,X9),X10),s(bool,'F_0'))),s(fun(X9,bool),X11)))) | s(bool,'T_0') = s(bool,sK6(X11,happ(s(fun(X6,fun(bool,bool)),'const.combin.K_1'(s(fun(bool,bool),X7))),s(X6,X8)),X10,X9,bool))) )),
  inference(forward_demodulation,[],[f8458,f80])).
fof(f8458,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(fun(bool,X9),X10),s(fun(fun(bool,X9),bool),'const.util_prob.FUNSET_2'(s(fun(bool,bool),'const.combin.K_2'(s(fun(bool,bool),X7),s(X6,X8))),s(fun(X9,bool),X11)))))) | ~p(s(bool,'const.bool.IN_2'(s(X9,happ(s(fun(bool,X9),X10),s(bool,'F_0'))),s(fun(X9,bool),X11)))) | s(bool,'T_0') = s(bool,sK6(X11,happ(s(fun(X6,fun(bool,bool)),'const.combin.K_1'(s(fun(bool,bool),X7))),s(X6,X8)),X10,X9,bool))) )),
  inference(superposition,[],[f6690,f81])).
fof(f8457,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(bool,X3),X4),s(fun(fun(bool,X3),bool),'const.util_prob.FUNSET_2'(s(fun(bool,bool),X0),s(fun(X3,bool),X5)))))) | ~p(s(bool,'const.bool.IN_2'(s(X3,happ(s(fun(bool,X3),X4),s(bool,'F_0'))),s(fun(X3,bool),X5)))) | s(bool,'T_0') = s(bool,sK6(X5,'const.combin.K_2'(s(fun(bool,bool),X0),s(X1,X2)),X4,X3,bool))) )),
  inference(superposition,[],[f6690,f80])).
fof(f8465,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (~p(s(bool,'const.bool.IN_2'(s(X6,happ(s(fun(bool,X6),X8),s(bool,'F_0'))),s(fun(X6,bool),X11)))) | p(s(bool,'const.bool.IN_2'(s(fun(bool,X6),X8),s(fun(fun(bool,X6),bool),'const.util_prob.FUNSET_2'(s(fun(bool,bool),X10),s(fun(X6,bool),X11)))))) | s(bool,'T_0') = s(bool,sK6(X11,X10,happ(s(fun(X7,fun(bool,X6)),'const.combin.K_1'(s(fun(bool,X6),X8))),s(X7,X9)),X6,bool))) )),
  inference(forward_demodulation,[],[f8464,f80])).
fof(f8464,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (~p(s(bool,'const.bool.IN_2'(s(X6,happ(s(fun(bool,X6),'const.combin.K_2'(s(fun(bool,X6),X8),s(X7,X9))),s(bool,'F_0'))),s(fun(X6,bool),X11)))) | p(s(bool,'const.bool.IN_2'(s(fun(bool,X6),X8),s(fun(fun(bool,X6),bool),'const.util_prob.FUNSET_2'(s(fun(bool,bool),X10),s(fun(X6,bool),X11)))))) | s(bool,'T_0') = s(bool,sK6(X11,X10,happ(s(fun(X7,fun(bool,X6)),'const.combin.K_1'(s(fun(bool,X6),X8))),s(X7,X9)),X6,bool))) )),
  inference(forward_demodulation,[],[f8463,f81])).
fof(f8463,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(fun(bool,X6),X8),s(fun(fun(bool,X6),bool),'const.util_prob.FUNSET_2'(s(fun(bool,bool),X10),s(fun(X6,bool),X11)))))) | ~p(s(bool,'const.bool.IN_2'(s(X6,happ(s(fun(bool,X6),happ(s(fun(X7,fun(bool,X6)),'const.combin.K_1'(s(fun(bool,X6),X8))),s(X7,X9))),s(bool,'F_0'))),s(fun(X6,bool),X11)))) | s(bool,'T_0') = s(bool,sK6(X11,X10,happ(s(fun(X7,fun(bool,X6)),'const.combin.K_1'(s(fun(bool,X6),X8))),s(X7,X9)),X6,bool))) )),
  inference(forward_demodulation,[],[f8456,f80])).
fof(f8456,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(fun(bool,X6),'const.combin.K_2'(s(fun(bool,X6),X8),s(X7,X9))),s(fun(fun(bool,X6),bool),'const.util_prob.FUNSET_2'(s(fun(bool,bool),X10),s(fun(X6,bool),X11)))))) | ~p(s(bool,'const.bool.IN_2'(s(X6,happ(s(fun(bool,X6),happ(s(fun(X7,fun(bool,X6)),'const.combin.K_1'(s(fun(bool,X6),X8))),s(X7,X9))),s(bool,'F_0'))),s(fun(X6,bool),X11)))) | s(bool,'T_0') = s(bool,sK6(X11,X10,happ(s(fun(X7,fun(bool,X6)),'const.combin.K_1'(s(fun(bool,X6),X8))),s(X7,X9)),X6,bool))) )),
  inference(superposition,[],[f6690,f81])).
fof(f8462,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(bool,X0),X1),s(bool,'F_0'))),s(fun(X0,bool),X5)))) | p(s(bool,'const.bool.IN_2'(s(fun(bool,X0),X1),s(fun(fun(bool,X0),bool),'const.util_prob.FUNSET_2'(s(fun(bool,bool),X4),s(fun(X0,bool),X5)))))) | s(bool,'T_0') = s(bool,sK6(X5,X4,'const.combin.K_2'(s(fun(bool,X0),X1),s(X2,X3)),X0,bool))) )),
  inference(forward_demodulation,[],[f8455,f80])).
fof(f8455,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(bool,X0),X1),s(fun(fun(bool,X0),bool),'const.util_prob.FUNSET_2'(s(fun(bool,bool),X4),s(fun(X0,bool),X5)))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(bool,X0),'const.combin.K_2'(s(fun(bool,X0),X1),s(X2,X3))),s(bool,'F_0'))),s(fun(X0,bool),X5)))) | s(bool,'T_0') = s(bool,sK6(X5,X4,'const.combin.K_2'(s(fun(bool,X0),X1),s(X2,X3)),X0,bool))) )),
  inference(superposition,[],[f6690,f80])).
fof(f8453,plain,(
  ( ! [X23,X21,X19,X17,X22,X20,X18] : (~p(s(bool,'const.bool.IN_2'(s(X17,happ(s(fun(bool,X17),X18),s(bool,'F_0'))),s(fun(X17,bool),X19)))) | s(bool,'T_0') = s(bool,sK6(X19,X20,X18,X17,bool)) | p(s(bool,'const.bool.IN_2'(s(X21,happ(s(fun(fun(bool,X17),X21),X22),s(fun(bool,X17),X18))),s(fun(X21,bool),X23)))) | p(s(bool,'const.bool.IN_2'(s(fun(bool,X17),sK6(X23,'const.util_prob.FUNSET_2'(s(fun(bool,bool),X20),s(fun(X17,bool),X19)),X22,X21,fun(bool,X17))),s(fun(fun(bool,X17),bool),'const.util_prob.FUNSET_2'(s(fun(bool,bool),X20),s(fun(X17,bool),X19))))))) )),
  inference(resolution,[],[f6690,f6711])).
fof(f8451,plain,(
  ( ! [X6,X10,X8,X7,X5,X9] : (~p(s(bool,'const.bool.IN_2'(s(X5,happ(s(fun(bool,X5),happ(s(fun(X6,fun(bool,X5)),X7),s(X6,sK6('const.util_prob.FUNSET_2'(s(fun(bool,bool),X8),s(fun(X5,bool),X9)),X10,X7,fun(bool,X5),X6)))),s(bool,'F_0'))),s(fun(X5,bool),X9)))) | s(bool,'T_0') = s(bool,sK6(X9,X8,happ(s(fun(X6,fun(bool,X5)),X7),s(X6,sK6('const.util_prob.FUNSET_2'(s(fun(bool,bool),X8),s(fun(X5,bool),X9)),X10,X7,fun(bool,X5),X6))),X5,bool)) | p(s(bool,'const.bool.IN_2'(s(fun(X6,fun(bool,X5)),X7),s(fun(fun(X6,fun(bool,X5)),bool),'const.util_prob.FUNSET_2'(s(fun(X6,bool),X10),s(fun(fun(bool,X5),bool),'const.util_prob.FUNSET_2'(s(fun(bool,bool),X8),s(fun(X5,bool),X9))))))))) )),
  inference(resolution,[],[f6690,f85])).
fof(f8450,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(bool,X0),X1),s(bool,'F_0'))),s(fun(X0,bool),X2)))) | s(bool,'T_0') = s(bool,sK6(X2,X3,X1,X0,bool)) | ~p(s(bool,'const.bool.IN_2'(s(bool,X4),s(fun(bool,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(bool,X0),X1),s(bool,X4))),s(fun(X0,bool),X2))))) )),
  inference(resolution,[],[f6690,f83])).
fof(f8403,plain,(
  ( ! [X39,X47,X45,X43,X41,X48,X46,X44,X42,X40] : (p(s(bool,'const.bool.IN_2'(s(fun(fun(X39,X40),X41),X42),s(fun(fun(fun(X39,X40),X41),bool),'const.util_prob.FUNSET_2'(s(fun(fun(X39,X40),bool),'const.util_prob.DFUNSET_2'(s(fun(X39,bool),X43),s(fun(X39,fun(X40,bool)),X44))),s(fun(X41,bool),X45)))))) | ~p(s(bool,'const.bool.IN_2'(s(X39,X46),s(fun(X39,bool),X43)))) | p(s(bool,'const.bool.IN_2'(s(X40,happ(s(fun(X39,X40),sK6(happ(s(fun(X47,fun(X41,bool)),'const.combin.K_1'(s(fun(X41,bool),X45))),s(X47,X48)),'const.util_prob.DFUNSET_2'(s(fun(X39,bool),X43),s(fun(X39,fun(X40,bool)),X44)),X42,X41,fun(X39,X40))),s(X39,X46))),s(fun(X40,bool),happ(s(fun(X39,fun(X40,bool)),X44),s(X39,X46))))))) )),
  inference(resolution,[],[f6641,f86])).
fof(f8402,plain,(
  ( ! [X30,X37,X35,X33,X31,X29,X38,X36,X34,X32] : (p(s(bool,'const.bool.IN_2'(s(fun(fun(X29,X30),X31),X32),s(fun(fun(fun(X29,X30),X31),bool),'const.util_prob.FUNSET_2'(s(fun(fun(X29,X30),bool),'const.util_prob.FUNSET_2'(s(fun(X29,bool),X33),s(fun(X30,bool),X34))),s(fun(X31,bool),X35)))))) | ~p(s(bool,'const.bool.IN_2'(s(X29,X36),s(fun(X29,bool),X33)))) | p(s(bool,'const.bool.IN_2'(s(X30,happ(s(fun(X29,X30),sK6(happ(s(fun(X37,fun(X31,bool)),'const.combin.K_1'(s(fun(X31,bool),X35))),s(X37,X38)),'const.util_prob.FUNSET_2'(s(fun(X29,bool),X33),s(fun(X30,bool),X34)),X32,X31,fun(X29,X30))),s(X29,X36))),s(fun(X30,bool),X34))))) )),
  inference(resolution,[],[f6641,f83])).
fof(f8401,plain,(
  ( ! [X28,X26,X24,X23,X21,X19,X27,X25,X22,X20] : (p(s(bool,'const.bool.IN_2'(s(fun(X19,X20),X21),s(fun(fun(X19,X20),bool),'const.util_prob.FUNSET_2'(s(fun(X19,bool),X22),s(fun(X20,bool),X23)))))) | p(s(bool,'const.bool.IN_2'(s(X24,happ(s(fun(X19,X24),X25),s(X19,sK6(happ(s(fun(X26,fun(X20,bool)),'const.combin.K_1'(s(fun(X20,bool),X23))),s(X26,X27)),X22,X21,X20,X19)))),s(fun(X24,bool),X28)))) | p(s(bool,'const.bool.IN_2'(s(X19,sK6(X28,X22,X25,X24,X19)),s(fun(X19,bool),X22))))) )),
  inference(resolution,[],[f6641,f6711])).
fof(f8397,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X9,X16] : (p(s(bool,'const.bool.IN_2'(s(fun(X13,X9),X14),s(fun(fun(X13,X9),bool),'const.util_prob.FUNSET_2'(s(fun(X13,bool),X15),s(fun(X9,bool),X11)))))) | p(s(bool,'const.bool.IN_2'(s(X13,sK6(happ(s(fun(X10,fun(X9,bool)),'const.combin.K_1'(s(fun(X9,bool),X11))),s(X10,X12)),happ(s(fun(X16,fun(X13,bool)),'const.combin.K_1'(s(fun(X13,bool),X15))),s(X16,X17)),X14,X9,X13)),s(fun(X13,bool),X15))))) )),
  inference(forward_demodulation,[],[f8380,f80])).
fof(f8380,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X9,X16] : (p(s(bool,'const.bool.IN_2'(s(fun(X13,X9),X14),s(fun(fun(X13,X9),bool),'const.util_prob.FUNSET_2'(s(fun(X13,bool),X15),s(fun(X9,bool),'const.combin.K_2'(s(fun(X9,bool),X11),s(X10,X12)))))))) | p(s(bool,'const.bool.IN_2'(s(X13,sK6(happ(s(fun(X10,fun(X9,bool)),'const.combin.K_1'(s(fun(X9,bool),X11))),s(X10,X12)),happ(s(fun(X16,fun(X13,bool)),'const.combin.K_1'(s(fun(X13,bool),X15))),s(X16,X17)),X14,X9,X13)),s(fun(X13,bool),X15))))) )),
  inference(superposition,[],[f6640,f81])).
fof(f8338,plain,(
  ( ! [X30,X28,X26,X35,X33,X31,X29,X27,X34,X32] : (p(s(bool,'const.bool.IN_2'(s(X26,sK6(X27,happ(s(fun(X28,fun(X26,bool)),'const.combin.K_1'(s(fun(X26,bool),X29))),s(X28,X30)),X31,X32,X26)),s(fun(X26,bool),X29)))) | p(s(bool,'const.bool.IN_2'(s(X33,happ(s(fun(fun(X26,X32),X33),X34),s(fun(X26,X32),X31))),s(fun(X33,bool),X35)))) | p(s(bool,'const.bool.IN_2'(s(fun(X26,X32),sK6(X35,'const.util_prob.FUNSET_2'(s(fun(X26,bool),X29),s(fun(X32,bool),X27)),X34,X33,fun(X26,X32))),s(fun(fun(X26,X32),bool),'const.util_prob.FUNSET_2'(s(fun(X26,bool),X29),s(fun(X32,bool),X27))))))) )),
  inference(resolution,[],[f6640,f6711])).
fof(f8336,plain,(
  ( ! [X14,X12,X10,X8,X15,X13,X11,X9,X16] : (p(s(bool,'const.bool.IN_2'(s(X8,sK6(X9,happ(s(fun(X10,fun(X8,bool)),'const.combin.K_1'(s(fun(X8,bool),X11))),s(X10,X12)),happ(s(fun(X13,fun(X8,X14)),X15),s(X13,sK6('const.util_prob.FUNSET_2'(s(fun(X8,bool),X11),s(fun(X14,bool),X9)),X16,X15,fun(X8,X14),X13))),X14,X8)),s(fun(X8,bool),X11)))) | p(s(bool,'const.bool.IN_2'(s(fun(X13,fun(X8,X14)),X15),s(fun(fun(X13,fun(X8,X14)),bool),'const.util_prob.FUNSET_2'(s(fun(X13,bool),X16),s(fun(fun(X8,X14),bool),'const.util_prob.FUNSET_2'(s(fun(X8,bool),X11),s(fun(X14,bool),X9))))))))) )),
  inference(resolution,[],[f6640,f85])).
fof(f8333,plain,(
  ( ! [X47,X45,X43,X52,X50,X48,X46,X44,X51,X49] : (p(s(bool,'const.bool.IN_2'(s(fun(fun(X43,X44),X45),X46),s(fun(fun(fun(X43,X44),X45),bool),'const.util_prob.FUNSET_2'(s(fun(fun(X43,X44),bool),'const.util_prob.DFUNSET_2'(s(fun(X43,bool),X47),s(fun(X43,fun(X44,bool)),X48))),s(fun(X45,bool),X49)))))) | ~p(s(bool,'const.bool.IN_2'(s(X43,X50),s(fun(X43,bool),X47)))) | p(s(bool,'const.bool.IN_2'(s(X44,happ(s(fun(X43,X44),sK6(X49,happ(s(fun(X51,fun(fun(X43,X44),bool)),'const.combin.K_1'(s(fun(fun(X43,X44),bool),'const.util_prob.DFUNSET_2'(s(fun(X43,bool),X47),s(fun(X43,fun(X44,bool)),X48))))),s(X51,X52)),X46,X45,fun(X43,X44))),s(X43,X50))),s(fun(X44,bool),happ(s(fun(X43,fun(X44,bool)),X48),s(X43,X50))))))) )),
  inference(resolution,[],[f6640,f86])).
fof(f8332,plain,(
  ( ! [X39,X37,X35,X33,X41,X38,X36,X34,X42,X40] : (p(s(bool,'const.bool.IN_2'(s(fun(fun(X33,X34),X35),X36),s(fun(fun(fun(X33,X34),X35),bool),'const.util_prob.FUNSET_2'(s(fun(fun(X33,X34),bool),'const.util_prob.FUNSET_2'(s(fun(X33,bool),X37),s(fun(X34,bool),X38))),s(fun(X35,bool),X39)))))) | ~p(s(bool,'const.bool.IN_2'(s(X33,X40),s(fun(X33,bool),X37)))) | p(s(bool,'const.bool.IN_2'(s(X34,happ(s(fun(X33,X34),sK6(X39,happ(s(fun(X41,fun(fun(X33,X34),bool)),'const.combin.K_1'(s(fun(fun(X33,X34),bool),'const.util_prob.FUNSET_2'(s(fun(X33,bool),X37),s(fun(X34,bool),X38))))),s(X41,X42)),X36,X35,fun(X33,X34))),s(X33,X40))),s(fun(X34,bool),X38))))) )),
  inference(resolution,[],[f6640,f83])).
fof(f8331,plain,(
  ( ! [X30,X28,X26,X24,X23,X31,X29,X27,X25,X32] : (p(s(bool,'const.bool.IN_2'(s(fun(X23,X24),X25),s(fun(fun(X23,X24),bool),'const.util_prob.FUNSET_2'(s(fun(X23,bool),X26),s(fun(X24,bool),X27)))))) | p(s(bool,'const.bool.IN_2'(s(X28,happ(s(fun(X23,X28),X29),s(X23,sK6(X27,happ(s(fun(X30,fun(X23,bool)),'const.combin.K_1'(s(fun(X23,bool),X26))),s(X30,X31)),X25,X24,X23)))),s(fun(X28,bool),X32)))) | p(s(bool,'const.bool.IN_2'(s(X23,sK6(X32,X26,X29,X28,X23)),s(fun(X23,bool),X26))))) )),
  inference(resolution,[],[f6640,f6711])).
fof(f8325,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X9,X16] : (p(s(bool,'const.bool.IN_2'(s(fun(X9,X15),X16),s(fun(fun(X9,X15),bool),'const.util_prob.FUNSET_2'(s(fun(X9,bool),X11),s(fun(X15,bool),X13)))))) | p(s(bool,'const.bool.IN_2'(s(X9,sK6(X13,happ(s(fun(X10,fun(X9,bool)),'const.combin.K_1'(s(fun(X9,bool),X11))),s(X10,X12)),happ(s(fun(X14,fun(X9,X15)),'const.combin.K_1'(s(fun(X9,X15),X16))),s(X14,X17)),X15,X9)),s(fun(X9,bool),X11))))) )),
  inference(forward_demodulation,[],[f8324,f80])).
fof(f8324,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X9,X16] : (p(s(bool,'const.bool.IN_2'(s(fun(X9,X15),X16),s(fun(fun(X9,X15),bool),'const.util_prob.FUNSET_2'(s(fun(X9,bool),'const.combin.K_2'(s(fun(X9,bool),X11),s(X10,X12))),s(fun(X15,bool),X13)))))) | p(s(bool,'const.bool.IN_2'(s(X9,sK6(X13,happ(s(fun(X10,fun(X9,bool)),'const.combin.K_1'(s(fun(X9,bool),X11))),s(X10,X12)),happ(s(fun(X14,fun(X9,X15)),'const.combin.K_1'(s(fun(X9,X15),X16))),s(X14,X17)),X15,X9)),s(fun(X9,bool),X11))))) )),
  inference(forward_demodulation,[],[f8323,f81])).
fof(f8323,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X9,X16] : (p(s(bool,'const.bool.IN_2'(s(X9,sK6(X13,happ(s(fun(X10,fun(X9,bool)),'const.combin.K_1'(s(fun(X9,bool),X11))),s(X10,X12)),happ(s(fun(X14,fun(X9,X15)),'const.combin.K_1'(s(fun(X9,X15),X16))),s(X14,X17)),X15,X9)),s(fun(X9,bool),X11)))) | p(s(bool,'const.bool.IN_2'(s(fun(X9,X15),X16),s(fun(fun(X9,X15),bool),'const.util_prob.FUNSET_2'(s(fun(X9,bool),happ(s(fun(X10,fun(X9,bool)),'const.combin.K_1'(s(fun(X9,bool),X11))),s(X10,X12))),s(fun(X15,bool),X13))))))) )),
  inference(forward_demodulation,[],[f8312,f80])).
fof(f8312,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X9,X16] : (p(s(bool,'const.bool.IN_2'(s(X9,sK6(X13,happ(s(fun(X10,fun(X9,bool)),'const.combin.K_1'(s(fun(X9,bool),X11))),s(X10,X12)),happ(s(fun(X14,fun(X9,X15)),'const.combin.K_1'(s(fun(X9,X15),X16))),s(X14,X17)),X15,X9)),s(fun(X9,bool),'const.combin.K_2'(s(fun(X9,bool),X11),s(X10,X12)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X9,X15),X16),s(fun(fun(X9,X15),bool),'const.util_prob.FUNSET_2'(s(fun(X9,bool),happ(s(fun(X10,fun(X9,bool)),'const.combin.K_1'(s(fun(X9,bool),X11))),s(X10,X12))),s(fun(X15,bool),X13))))))) )),
  inference(superposition,[],[f6636,f81])).
fof(f8279,plain,(
  ( ! [X39,X47,X45,X43,X41,X48,X46,X44,X42,X40] : (p(s(bool,'const.bool.IN_2'(s(fun(fun(X39,X40),X41),X42),s(fun(fun(fun(X39,X40),X41),bool),'const.util_prob.FUNSET_2'(s(fun(fun(X39,X40),bool),'const.util_prob.DFUNSET_2'(s(fun(X39,bool),X43),s(fun(X39,fun(X40,bool)),X44))),s(fun(X41,bool),X45)))))) | ~p(s(bool,'const.bool.IN_2'(s(X39,X46),s(fun(X39,bool),X43)))) | p(s(bool,'const.bool.IN_2'(s(X40,happ(s(fun(X39,X40),sK6(X45,'const.util_prob.DFUNSET_2'(s(fun(X39,bool),X43),s(fun(X39,fun(X40,bool)),X44)),happ(s(fun(X47,fun(fun(X39,X40),X41)),'const.combin.K_1'(s(fun(fun(X39,X40),X41),X42))),s(X47,X48)),X41,fun(X39,X40))),s(X39,X46))),s(fun(X40,bool),happ(s(fun(X39,fun(X40,bool)),X44),s(X39,X46))))))) )),
  inference(resolution,[],[f6636,f86])).
fof(f8278,plain,(
  ( ! [X30,X37,X35,X33,X31,X29,X38,X36,X34,X32] : (p(s(bool,'const.bool.IN_2'(s(fun(fun(X29,X30),X31),X32),s(fun(fun(fun(X29,X30),X31),bool),'const.util_prob.FUNSET_2'(s(fun(fun(X29,X30),bool),'const.util_prob.FUNSET_2'(s(fun(X29,bool),X33),s(fun(X30,bool),X34))),s(fun(X31,bool),X35)))))) | ~p(s(bool,'const.bool.IN_2'(s(X29,X36),s(fun(X29,bool),X33)))) | p(s(bool,'const.bool.IN_2'(s(X30,happ(s(fun(X29,X30),sK6(X35,'const.util_prob.FUNSET_2'(s(fun(X29,bool),X33),s(fun(X30,bool),X34)),happ(s(fun(X37,fun(fun(X29,X30),X31)),'const.combin.K_1'(s(fun(fun(X29,X30),X31),X32))),s(X37,X38)),X31,fun(X29,X30))),s(X29,X36))),s(fun(X30,bool),X34))))) )),
  inference(resolution,[],[f6636,f83])).
fof(f8277,plain,(
  ( ! [X28,X26,X24,X23,X21,X19,X27,X25,X22,X20] : (p(s(bool,'const.bool.IN_2'(s(fun(X19,X20),X21),s(fun(fun(X19,X20),bool),'const.util_prob.FUNSET_2'(s(fun(X19,bool),X22),s(fun(X20,bool),X23)))))) | p(s(bool,'const.bool.IN_2'(s(X24,happ(s(fun(X19,X24),X25),s(X19,sK6(X23,X22,happ(s(fun(X26,fun(X19,X20)),'const.combin.K_1'(s(fun(X19,X20),X21))),s(X26,X27)),X20,X19)))),s(fun(X24,bool),X28)))) | p(s(bool,'const.bool.IN_2'(s(X19,sK6(X28,X22,X25,X24,X19)),s(fun(X19,bool),X22))))) )),
  inference(resolution,[],[f6636,f6711])).
fof(f8259,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X9,X16] : (p(s(bool,'const.bool.IN_2'(s(X9,sK7('const.combin.K_2'(s(fun(X9,fun(X13,bool)),X15),s(X16,X17)),happ(s(fun(X10,fun(X9,bool)),'const.combin.K_1'(s(fun(X9,bool),X11))),s(X10,X12)),X14,X9,X13)),s(fun(X9,bool),X11)))) | p(s(bool,'const.bool.IN_2'(s(fun(X9,X13),X14),s(fun(fun(X9,X13),bool),'const.util_prob.DFUNSET_2'(s(fun(X9,bool),X11),s(fun(X9,fun(X13,bool)),X15))))))) )),
  inference(forward_demodulation,[],[f8258,f80])).
fof(f8258,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X9,X16] : (p(s(bool,'const.bool.IN_2'(s(X9,sK7('const.combin.K_2'(s(fun(X9,fun(X13,bool)),X15),s(X16,X17)),happ(s(fun(X10,fun(X9,bool)),'const.combin.K_1'(s(fun(X9,bool),X11))),s(X10,X12)),X14,X9,X13)),s(fun(X9,bool),'const.combin.K_2'(s(fun(X9,bool),X11),s(X10,X12)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X9,X13),X14),s(fun(fun(X9,X13),bool),'const.util_prob.DFUNSET_2'(s(fun(X9,bool),X11),s(fun(X9,fun(X13,bool)),X15))))))) )),
  inference(forward_demodulation,[],[f8257,f81])).
fof(f8257,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X9,X16] : (p(s(bool,'const.bool.IN_2'(s(fun(X9,X13),X14),s(fun(fun(X9,X13),bool),'const.util_prob.DFUNSET_2'(s(fun(X9,bool),X11),s(fun(X9,fun(X13,bool)),X15)))))) | p(s(bool,'const.bool.IN_2'(s(X9,sK7('const.combin.K_2'(s(fun(X9,fun(X13,bool)),X15),s(X16,X17)),happ(s(fun(X10,fun(X9,bool)),'const.combin.K_1'(s(fun(X9,bool),X11))),s(X10,X12)),X14,X9,X13)),s(fun(X9,bool),happ(s(fun(X10,fun(X9,bool)),'const.combin.K_1'(s(fun(X9,bool),X11))),s(X10,X12))))))) )),
  inference(forward_demodulation,[],[f8210,f80])).
fof(f8210,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X9,X16] : (p(s(bool,'const.bool.IN_2'(s(fun(X9,X13),X14),s(fun(fun(X9,X13),bool),'const.util_prob.DFUNSET_2'(s(fun(X9,bool),'const.combin.K_2'(s(fun(X9,bool),X11),s(X10,X12))),s(fun(X9,fun(X13,bool)),X15)))))) | p(s(bool,'const.bool.IN_2'(s(X9,sK7('const.combin.K_2'(s(fun(X9,fun(X13,bool)),X15),s(X16,X17)),happ(s(fun(X10,fun(X9,bool)),'const.combin.K_1'(s(fun(X9,bool),X11))),s(X10,X12)),X14,X9,X13)),s(fun(X9,bool),happ(s(fun(X10,fun(X9,bool)),'const.combin.K_1'(s(fun(X9,bool),X11))),s(X10,X12))))))) )),
  inference(superposition,[],[f6655,f81])).
fof(f8255,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X9,X16] : (p(s(bool,'const.bool.IN_2'(s(fun(X9,X10),X12),s(fun(fun(X9,X10),bool),'const.util_prob.DFUNSET_2'(s(fun(X9,bool),X14),s(fun(X9,fun(X10,bool)),X15)))))) | p(s(bool,'const.bool.IN_2'(s(X9,sK7('const.combin.K_2'(s(fun(X9,fun(X10,bool)),X15),s(X16,X17)),X14,happ(s(fun(X11,fun(X9,X10)),'const.combin.K_1'(s(fun(X9,X10),X12))),s(X11,X13)),X9,X10)),s(fun(X9,bool),X14))))) )),
  inference(forward_demodulation,[],[f8208,f80])).
fof(f8208,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X9,X16] : (p(s(bool,'const.bool.IN_2'(s(fun(X9,X10),'const.combin.K_2'(s(fun(X9,X10),X12),s(X11,X13))),s(fun(fun(X9,X10),bool),'const.util_prob.DFUNSET_2'(s(fun(X9,bool),X14),s(fun(X9,fun(X10,bool)),X15)))))) | p(s(bool,'const.bool.IN_2'(s(X9,sK7('const.combin.K_2'(s(fun(X9,fun(X10,bool)),X15),s(X16,X17)),X14,happ(s(fun(X11,fun(X9,X10)),'const.combin.K_1'(s(fun(X9,X10),X12))),s(X11,X13)),X9,X10)),s(fun(X9,bool),X14))))) )),
  inference(superposition,[],[f6655,f81])).
fof(f8205,plain,(
  ( ! [X39,X47,X45,X43,X41,X48,X46,X44,X42,X40] : (p(s(bool,'const.bool.IN_2'(s(fun(fun(X39,X40),X41),X42),s(fun(fun(fun(X39,X40),X41),bool),'const.util_prob.DFUNSET_2'(s(fun(fun(X39,X40),bool),'const.util_prob.DFUNSET_2'(s(fun(X39,bool),X43),s(fun(X39,fun(X40,bool)),X44))),s(fun(fun(X39,X40),fun(X41,bool)),X45)))))) | ~p(s(bool,'const.bool.IN_2'(s(X39,X46),s(fun(X39,bool),X43)))) | p(s(bool,'const.bool.IN_2'(s(X40,happ(s(fun(X39,X40),sK7('const.combin.K_2'(s(fun(fun(X39,X40),fun(X41,bool)),X45),s(X47,X48)),'const.util_prob.DFUNSET_2'(s(fun(X39,bool),X43),s(fun(X39,fun(X40,bool)),X44)),X42,fun(X39,X40),X41)),s(X39,X46))),s(fun(X40,bool),happ(s(fun(X39,fun(X40,bool)),X44),s(X39,X46))))))) )),
  inference(resolution,[],[f6655,f86])).
fof(f8204,plain,(
  ( ! [X30,X37,X35,X33,X31,X29,X38,X36,X34,X32] : (p(s(bool,'const.bool.IN_2'(s(fun(fun(X29,X30),X31),X32),s(fun(fun(fun(X29,X30),X31),bool),'const.util_prob.DFUNSET_2'(s(fun(fun(X29,X30),bool),'const.util_prob.FUNSET_2'(s(fun(X29,bool),X33),s(fun(X30,bool),X34))),s(fun(fun(X29,X30),fun(X31,bool)),X35)))))) | ~p(s(bool,'const.bool.IN_2'(s(X29,X36),s(fun(X29,bool),X33)))) | p(s(bool,'const.bool.IN_2'(s(X30,happ(s(fun(X29,X30),sK7('const.combin.K_2'(s(fun(fun(X29,X30),fun(X31,bool)),X35),s(X37,X38)),'const.util_prob.FUNSET_2'(s(fun(X29,bool),X33),s(fun(X30,bool),X34)),X32,fun(X29,X30),X31)),s(X29,X36))),s(fun(X30,bool),X34))))) )),
  inference(resolution,[],[f6655,f83])).
fof(f8203,plain,(
  ( ! [X28,X26,X24,X23,X21,X19,X27,X25,X22,X20] : (p(s(bool,'const.bool.IN_2'(s(fun(X19,X20),X21),s(fun(fun(X19,X20),bool),'const.util_prob.DFUNSET_2'(s(fun(X19,bool),X22),s(fun(X19,fun(X20,bool)),X23)))))) | p(s(bool,'const.bool.IN_2'(s(X24,happ(s(fun(X19,X24),X25),s(X19,sK7('const.combin.K_2'(s(fun(X19,fun(X20,bool)),X23),s(X26,X27)),X22,X21,X19,X20)))),s(fun(X24,bool),X28)))) | p(s(bool,'const.bool.IN_2'(s(X19,sK6(X28,X22,X25,X24,X19)),s(fun(X19,bool),X22))))) )),
  inference(resolution,[],[f6655,f6711])).
fof(f8198,plain,(
  ( ! [X30,X28,X26,X35,X33,X31,X29,X27,X34,X32] : (p(s(bool,'const.bool.IN_2'(s(X26,sK7('const.combin.K_2'(s(fun(X26,fun(X27,bool)),X28),s(X29,X30)),X31,X32,X26,X27)),s(fun(X26,bool),X31)))) | p(s(bool,'const.bool.IN_2'(s(X33,happ(s(fun(fun(X26,X27),X33),X34),s(fun(X26,X27),X32))),s(fun(X33,bool),X35)))) | p(s(bool,'const.bool.IN_2'(s(fun(X26,X27),sK6(X35,'const.util_prob.DFUNSET_2'(s(fun(X26,bool),X31),s(fun(X26,fun(X27,bool)),X28)),X34,X33,fun(X26,X27))),s(fun(fun(X26,X27),bool),'const.util_prob.DFUNSET_2'(s(fun(X26,bool),X31),s(fun(X26,fun(X27,bool)),X28))))))) )),
  inference(resolution,[],[f6655,f6711])).
fof(f8196,plain,(
  ( ! [X14,X12,X10,X8,X15,X13,X11,X9,X16] : (p(s(bool,'const.bool.IN_2'(s(X8,sK7('const.combin.K_2'(s(fun(X8,fun(X9,bool)),X10),s(X11,X12)),X13,happ(s(fun(X14,fun(X8,X9)),X15),s(X14,sK6('const.util_prob.DFUNSET_2'(s(fun(X8,bool),X13),s(fun(X8,fun(X9,bool)),X10)),X16,X15,fun(X8,X9),X14))),X8,X9)),s(fun(X8,bool),X13)))) | p(s(bool,'const.bool.IN_2'(s(fun(X14,fun(X8,X9)),X15),s(fun(fun(X14,fun(X8,X9)),bool),'const.util_prob.FUNSET_2'(s(fun(X14,bool),X16),s(fun(fun(X8,X9),bool),'const.util_prob.DFUNSET_2'(s(fun(X8,bool),X13),s(fun(X8,fun(X9,bool)),X10))))))))) )),
  inference(resolution,[],[f6655,f85])).
fof(f8195,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK7('const.combin.K_2'(s(fun(X0,fun(X1,bool)),X2),s(X3,X4)),X5,X6,X0,X1)),s(fun(X0,bool),X5)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X7),s(fun(X0,bool),X5)))) | p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X6),s(X0,X7))),s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X7))))))) )),
  inference(resolution,[],[f6655,f86])).
fof(f8083,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X9,X16] : (p(s(bool,'const.bool.IN_2'(s(fun(X9,X10),X14),s(fun(fun(X9,X10),bool),'const.util_prob.DFUNSET_2'(s(fun(X9,bool),X15),s(fun(X9,fun(X10,bool)),X12)))))) | p(s(bool,'const.bool.IN_2'(s(X9,sK7(happ(s(fun(X11,fun(X9,fun(X10,bool))),'const.combin.K_1'(s(fun(X9,fun(X10,bool)),X12))),s(X11,X13)),'const.combin.K_2'(s(fun(X9,bool),X15),s(X16,X17)),X14,X9,X10)),s(fun(X9,bool),X15))))) )),
  inference(forward_demodulation,[],[f8068,f80])).
fof(f8068,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X9,X16] : (p(s(bool,'const.bool.IN_2'(s(fun(X9,X10),X14),s(fun(fun(X9,X10),bool),'const.util_prob.DFUNSET_2'(s(fun(X9,bool),X15),s(fun(X9,fun(X10,bool)),'const.combin.K_2'(s(fun(X9,fun(X10,bool)),X12),s(X11,X13)))))))) | p(s(bool,'const.bool.IN_2'(s(X9,sK7(happ(s(fun(X11,fun(X9,fun(X10,bool))),'const.combin.K_1'(s(fun(X9,fun(X10,bool)),X12))),s(X11,X13)),'const.combin.K_2'(s(fun(X9,bool),X15),s(X16,X17)),X14,X9,X10)),s(fun(X9,bool),X15))))) )),
  inference(superposition,[],[f6673,f81])).
fof(f8067,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,X1),X5),s(fun(fun(X0,X1),bool),'const.util_prob.DFUNSET_2'(s(fun(X0,bool),X6),s(fun(X0,fun(X1,bool)),X2)))))) | p(s(bool,'const.bool.IN_2'(s(X0,sK7('const.combin.K_2'(s(fun(X0,fun(X1,bool)),X2),s(X3,X4)),'const.combin.K_2'(s(fun(X0,bool),X6),s(X7,X8)),X5,X0,X1)),s(fun(X0,bool),X6))))) )),
  inference(superposition,[],[f6673,f80])).
fof(f8078,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X9,X16] : (p(s(bool,'const.bool.IN_2'(s(fun(X9,X10),X12),s(fun(fun(X9,X10),bool),'const.util_prob.DFUNSET_2'(s(fun(X9,bool),X14),s(fun(X9,fun(X10,bool)),X15)))))) | p(s(bool,'const.bool.IN_2'(s(X9,sK7(X15,'const.combin.K_2'(s(fun(X9,bool),X14),s(X16,X17)),happ(s(fun(X11,fun(X9,X10)),'const.combin.K_1'(s(fun(X9,X10),X12))),s(X11,X13)),X9,X10)),s(fun(X9,bool),X14))))) )),
  inference(forward_demodulation,[],[f8064,f80])).
fof(f8064,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X9,X16] : (p(s(bool,'const.bool.IN_2'(s(fun(X9,X10),'const.combin.K_2'(s(fun(X9,X10),X12),s(X11,X13))),s(fun(fun(X9,X10),bool),'const.util_prob.DFUNSET_2'(s(fun(X9,bool),X14),s(fun(X9,fun(X10,bool)),X15)))))) | p(s(bool,'const.bool.IN_2'(s(X9,sK7(X15,'const.combin.K_2'(s(fun(X9,bool),X14),s(X16,X17)),happ(s(fun(X11,fun(X9,X10)),'const.combin.K_1'(s(fun(X9,X10),X12))),s(X11,X13)),X9,X10)),s(fun(X9,bool),X14))))) )),
  inference(superposition,[],[f6673,f81])).
fof(f8026,plain,(
  ( ! [X26,X24,X23,X21,X19,X17,X25,X22,X20,X18] : (p(s(bool,'const.bool.IN_2'(s(X17,sK7(X18,'const.combin.K_2'(s(fun(X17,bool),X19),s(X20,X21)),X22,X17,X23)),s(fun(X17,bool),X19)))) | p(s(bool,'const.bool.IN_2'(s(X24,happ(s(fun(fun(X17,X23),X24),X25),s(fun(X17,X23),X22))),s(fun(X24,bool),X26)))) | p(s(bool,'const.bool.IN_2'(s(fun(X17,X23),sK6(X26,'const.util_prob.DFUNSET_2'(s(fun(X17,bool),X19),s(fun(X17,fun(X23,bool)),X18)),X25,X24,fun(X17,X23))),s(fun(fun(X17,X23),bool),'const.util_prob.DFUNSET_2'(s(fun(X17,bool),X19),s(fun(X17,fun(X23,bool)),X18))))))) )),
  inference(resolution,[],[f6673,f6711])).
fof(f8025,plain,(
  ( ! [X14,X12,X10,X8,X15,X13,X11,X9,X16] : (p(s(bool,'const.bool.IN_2'(s(X8,sK7(X9,'const.combin.K_2'(s(fun(X8,bool),X10),s(X11,X12)),happ(s(fun(X13,fun(X8,X14)),X15),s(X13,sK6('const.util_prob.DFUNSET_2'(s(fun(X8,bool),X10),s(fun(X8,fun(X14,bool)),X9)),X16,X15,fun(X8,X14),X13))),X8,X14)),s(fun(X8,bool),X10)))) | p(s(bool,'const.bool.IN_2'(s(fun(X13,fun(X8,X14)),X15),s(fun(fun(X13,fun(X8,X14)),bool),'const.util_prob.FUNSET_2'(s(fun(X13,bool),X16),s(fun(fun(X8,X14),bool),'const.util_prob.DFUNSET_2'(s(fun(X8,bool),X10),s(fun(X8,fun(X14,bool)),X9))))))))) )),
  inference(resolution,[],[f6673,f85])).
fof(f8024,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK7(X1,'const.combin.K_2'(s(fun(X0,bool),X2),s(X3,X4)),X5,X0,X6)),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X7),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X6,happ(s(fun(X0,X6),X5),s(X0,X7))),s(fun(X6,bool),happ(s(fun(X0,fun(X6,bool)),X1),s(X0,X7))))))) )),
  inference(resolution,[],[f6673,f86])).
fof(f8022,plain,(
  ( ! [X30,X39,X37,X35,X33,X31,X38,X36,X34,X32] : (p(s(bool,'const.bool.IN_2'(s(fun(fun(X30,X31),X32),X33),s(fun(fun(fun(X30,X31),X32),bool),'const.util_prob.DFUNSET_2'(s(fun(fun(X30,X31),bool),'const.util_prob.DFUNSET_2'(s(fun(X30,bool),X34),s(fun(X30,fun(X31,bool)),X35))),s(fun(fun(X30,X31),fun(X32,bool)),X36)))))) | ~p(s(bool,'const.bool.IN_2'(s(X30,X37),s(fun(X30,bool),X34)))) | p(s(bool,'const.bool.IN_2'(s(X31,happ(s(fun(X30,X31),sK7(X36,'const.combin.K_2'(s(fun(fun(X30,X31),bool),'const.util_prob.DFUNSET_2'(s(fun(X30,bool),X34),s(fun(X30,fun(X31,bool)),X35))),s(X38,X39)),X33,fun(X30,X31),X32)),s(X30,X37))),s(fun(X31,bool),happ(s(fun(X30,fun(X31,bool)),X35),s(X30,X37))))))) )),
  inference(resolution,[],[f6673,f86])).
fof(f8021,plain,(
  ( ! [X28,X26,X24,X23,X21,X29,X27,X25,X22,X20] : (p(s(bool,'const.bool.IN_2'(s(fun(fun(X20,X21),X22),X23),s(fun(fun(fun(X20,X21),X22),bool),'const.util_prob.DFUNSET_2'(s(fun(fun(X20,X21),bool),'const.util_prob.FUNSET_2'(s(fun(X20,bool),X24),s(fun(X21,bool),X25))),s(fun(fun(X20,X21),fun(X22,bool)),X26)))))) | ~p(s(bool,'const.bool.IN_2'(s(X20,X27),s(fun(X20,bool),X24)))) | p(s(bool,'const.bool.IN_2'(s(X21,happ(s(fun(X20,X21),sK7(X26,'const.combin.K_2'(s(fun(fun(X20,X21),bool),'const.util_prob.FUNSET_2'(s(fun(X20,bool),X24),s(fun(X21,bool),X25))),s(X28,X29)),X23,fun(X20,X21),X22)),s(X20,X27))),s(fun(X21,bool),X25))))) )),
  inference(resolution,[],[f6673,f83])).
fof(f8020,plain,(
  ( ! [X14,X12,X10,X19,X17,X15,X13,X11,X18,X16] : (p(s(bool,'const.bool.IN_2'(s(fun(X10,X11),X12),s(fun(fun(X10,X11),bool),'const.util_prob.DFUNSET_2'(s(fun(X10,bool),X13),s(fun(X10,fun(X11,bool)),X14)))))) | p(s(bool,'const.bool.IN_2'(s(X15,happ(s(fun(X10,X15),X16),s(X10,sK7(X14,'const.combin.K_2'(s(fun(X10,bool),X13),s(X17,X18)),X12,X10,X11)))),s(fun(X15,bool),X19)))) | p(s(bool,'const.bool.IN_2'(s(X10,sK6(X19,X13,X16,X15,X10)),s(fun(X10,bool),X13))))) )),
  inference(resolution,[],[f6673,f6711])).
fof(f8007,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X9,X16] : (p(s(bool,'const.bool.IN_2'(s(fun(X9,X10),X14),s(fun(fun(X9,X10),bool),'const.util_prob.DFUNSET_2'(s(fun(X9,bool),X15),s(fun(X9,fun(X10,bool)),X12)))))) | p(s(bool,'const.bool.IN_2'(s(X9,sK7(happ(s(fun(X11,fun(X9,fun(X10,bool))),'const.combin.K_1'(s(fun(X9,fun(X10,bool)),X12))),s(X11,X13)),X15,'const.combin.K_2'(s(fun(X9,X10),X14),s(X16,X17)),X9,X10)),s(fun(X9,bool),X15))))) )),
  inference(forward_demodulation,[],[f7956,f80])).
fof(f7956,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X9,X16] : (p(s(bool,'const.bool.IN_2'(s(fun(X9,X10),X14),s(fun(fun(X9,X10),bool),'const.util_prob.DFUNSET_2'(s(fun(X9,bool),X15),s(fun(X9,fun(X10,bool)),'const.combin.K_2'(s(fun(X9,fun(X10,bool)),X12),s(X11,X13)))))))) | p(s(bool,'const.bool.IN_2'(s(X9,sK7(happ(s(fun(X11,fun(X9,fun(X10,bool))),'const.combin.K_1'(s(fun(X9,fun(X10,bool)),X12))),s(X11,X13)),X15,'const.combin.K_2'(s(fun(X9,X10),X14),s(X16,X17)),X9,X10)),s(fun(X9,bool),X15))))) )),
  inference(superposition,[],[f6651,f81])).
fof(f7955,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,X1),X5),s(fun(fun(X0,X1),bool),'const.util_prob.DFUNSET_2'(s(fun(X0,bool),X6),s(fun(X0,fun(X1,bool)),X2)))))) | p(s(bool,'const.bool.IN_2'(s(X0,sK7('const.combin.K_2'(s(fun(X0,fun(X1,bool)),X2),s(X3,X4)),X6,'const.combin.K_2'(s(fun(X0,X1),X5),s(X7,X8)),X0,X1)),s(fun(X0,bool),X6))))) )),
  inference(superposition,[],[f6651,f80])).
fof(f8006,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X9,X16] : (p(s(bool,'const.bool.IN_2'(s(X9,sK7(X15,happ(s(fun(X10,fun(X9,bool)),'const.combin.K_1'(s(fun(X9,bool),X11))),s(X10,X12)),'const.combin.K_2'(s(fun(X9,X13),X14),s(X16,X17)),X9,X13)),s(fun(X9,bool),X11)))) | p(s(bool,'const.bool.IN_2'(s(fun(X9,X13),X14),s(fun(fun(X9,X13),bool),'const.util_prob.DFUNSET_2'(s(fun(X9,bool),X11),s(fun(X9,fun(X13,bool)),X15))))))) )),
  inference(forward_demodulation,[],[f8005,f80])).
fof(f8005,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X9,X16] : (p(s(bool,'const.bool.IN_2'(s(X9,sK7(X15,happ(s(fun(X10,fun(X9,bool)),'const.combin.K_1'(s(fun(X9,bool),X11))),s(X10,X12)),'const.combin.K_2'(s(fun(X9,X13),X14),s(X16,X17)),X9,X13)),s(fun(X9,bool),'const.combin.K_2'(s(fun(X9,bool),X11),s(X10,X12)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X9,X13),X14),s(fun(fun(X9,X13),bool),'const.util_prob.DFUNSET_2'(s(fun(X9,bool),X11),s(fun(X9,fun(X13,bool)),X15))))))) )),
  inference(forward_demodulation,[],[f8004,f81])).
fof(f8004,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X9,X16] : (p(s(bool,'const.bool.IN_2'(s(fun(X9,X13),X14),s(fun(fun(X9,X13),bool),'const.util_prob.DFUNSET_2'(s(fun(X9,bool),X11),s(fun(X9,fun(X13,bool)),X15)))))) | p(s(bool,'const.bool.IN_2'(s(X9,sK7(X15,happ(s(fun(X10,fun(X9,bool)),'const.combin.K_1'(s(fun(X9,bool),X11))),s(X10,X12)),'const.combin.K_2'(s(fun(X9,X13),X14),s(X16,X17)),X9,X13)),s(fun(X9,bool),happ(s(fun(X10,fun(X9,bool)),'const.combin.K_1'(s(fun(X9,bool),X11))),s(X10,X12))))))) )),
  inference(forward_demodulation,[],[f7954,f80])).
fof(f7954,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X9,X16] : (p(s(bool,'const.bool.IN_2'(s(fun(X9,X13),X14),s(fun(fun(X9,X13),bool),'const.util_prob.DFUNSET_2'(s(fun(X9,bool),'const.combin.K_2'(s(fun(X9,bool),X11),s(X10,X12))),s(fun(X9,fun(X13,bool)),X15)))))) | p(s(bool,'const.bool.IN_2'(s(X9,sK7(X15,happ(s(fun(X10,fun(X9,bool)),'const.combin.K_1'(s(fun(X9,bool),X11))),s(X10,X12)),'const.combin.K_2'(s(fun(X9,X13),X14),s(X16,X17)),X9,X13)),s(fun(X9,bool),happ(s(fun(X10,fun(X9,bool)),'const.combin.K_1'(s(fun(X9,bool),X11))),s(X10,X12))))))) )),
  inference(superposition,[],[f6651,f81])).
fof(f8003,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK7(X6,'const.combin.K_2'(s(fun(X0,bool),X1),s(X2,X3)),'const.combin.K_2'(s(fun(X0,X4),X5),s(X7,X8)),X0,X4)),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,X4),X5),s(fun(fun(X0,X4),bool),'const.util_prob.DFUNSET_2'(s(fun(X0,bool),X1),s(fun(X0,fun(X4,bool)),X6))))))) )),
  inference(forward_demodulation,[],[f7953,f80])).
fof(f7953,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,X4),X5),s(fun(fun(X0,X4),bool),'const.util_prob.DFUNSET_2'(s(fun(X0,bool),X1),s(fun(X0,fun(X4,bool)),X6)))))) | p(s(bool,'const.bool.IN_2'(s(X0,sK7(X6,'const.combin.K_2'(s(fun(X0,bool),X1),s(X2,X3)),'const.combin.K_2'(s(fun(X0,X4),X5),s(X7,X8)),X0,X4)),s(fun(X0,bool),'const.combin.K_2'(s(fun(X0,bool),X1),s(X2,X3))))))) )),
  inference(superposition,[],[f6651,f80])).
fof(f7949,plain,(
  ( ! [X30,X39,X37,X35,X33,X31,X38,X36,X34,X32] : (p(s(bool,'const.bool.IN_2'(s(fun(fun(X30,X31),X32),X33),s(fun(fun(fun(X30,X31),X32),bool),'const.util_prob.DFUNSET_2'(s(fun(fun(X30,X31),bool),'const.util_prob.DFUNSET_2'(s(fun(X30,bool),X34),s(fun(X30,fun(X31,bool)),X35))),s(fun(fun(X30,X31),fun(X32,bool)),X36)))))) | ~p(s(bool,'const.bool.IN_2'(s(X30,X37),s(fun(X30,bool),X34)))) | p(s(bool,'const.bool.IN_2'(s(X31,happ(s(fun(X30,X31),sK7(X36,'const.util_prob.DFUNSET_2'(s(fun(X30,bool),X34),s(fun(X30,fun(X31,bool)),X35)),'const.combin.K_2'(s(fun(fun(X30,X31),X32),X33),s(X38,X39)),fun(X30,X31),X32)),s(X30,X37))),s(fun(X31,bool),happ(s(fun(X30,fun(X31,bool)),X35),s(X30,X37))))))) )),
  inference(resolution,[],[f6651,f86])).
fof(f7948,plain,(
  ( ! [X28,X26,X24,X23,X21,X29,X27,X25,X22,X20] : (p(s(bool,'const.bool.IN_2'(s(fun(fun(X20,X21),X22),X23),s(fun(fun(fun(X20,X21),X22),bool),'const.util_prob.DFUNSET_2'(s(fun(fun(X20,X21),bool),'const.util_prob.FUNSET_2'(s(fun(X20,bool),X24),s(fun(X21,bool),X25))),s(fun(fun(X20,X21),fun(X22,bool)),X26)))))) | ~p(s(bool,'const.bool.IN_2'(s(X20,X27),s(fun(X20,bool),X24)))) | p(s(bool,'const.bool.IN_2'(s(X21,happ(s(fun(X20,X21),sK7(X26,'const.util_prob.FUNSET_2'(s(fun(X20,bool),X24),s(fun(X21,bool),X25)),'const.combin.K_2'(s(fun(fun(X20,X21),X22),X23),s(X28,X29)),fun(X20,X21),X22)),s(X20,X27))),s(fun(X21,bool),X25))))) )),
  inference(resolution,[],[f6651,f83])).
fof(f7947,plain,(
  ( ! [X14,X12,X10,X19,X17,X15,X13,X11,X18,X16] : (p(s(bool,'const.bool.IN_2'(s(fun(X10,X11),X12),s(fun(fun(X10,X11),bool),'const.util_prob.DFUNSET_2'(s(fun(X10,bool),X13),s(fun(X10,fun(X11,bool)),X14)))))) | p(s(bool,'const.bool.IN_2'(s(X15,happ(s(fun(X10,X15),X16),s(X10,sK7(X14,X13,'const.combin.K_2'(s(fun(X10,X11),X12),s(X17,X18)),X10,X11)))),s(fun(X15,bool),X19)))) | p(s(bool,'const.bool.IN_2'(s(X10,sK6(X19,X13,X16,X15,X10)),s(fun(X10,bool),X13))))) )),
  inference(resolution,[],[f6651,f6711])).
fof(f7943,plain,(
  ( ! [X26,X24,X23,X21,X19,X17,X25,X22,X20,X18] : (p(s(bool,'const.bool.IN_2'(s(X17,sK7(X18,X19,'const.combin.K_2'(s(fun(X17,X20),X21),s(X22,X23)),X17,X20)),s(fun(X17,bool),X19)))) | p(s(bool,'const.bool.IN_2'(s(X24,happ(s(fun(fun(X17,X20),X24),X25),s(fun(X17,X20),X21))),s(fun(X24,bool),X26)))) | p(s(bool,'const.bool.IN_2'(s(fun(X17,X20),sK6(X26,'const.util_prob.DFUNSET_2'(s(fun(X17,bool),X19),s(fun(X17,fun(X20,bool)),X18)),X25,X24,fun(X17,X20))),s(fun(fun(X17,X20),bool),'const.util_prob.DFUNSET_2'(s(fun(X17,bool),X19),s(fun(X17,fun(X20,bool)),X18))))))) )),
  inference(resolution,[],[f6651,f6711])).
fof(f7942,plain,(
  ( ! [X14,X12,X10,X8,X15,X13,X11,X9,X16] : (p(s(bool,'const.bool.IN_2'(s(X8,sK7(X9,X10,'const.combin.K_2'(s(fun(X8,X11),happ(s(fun(X12,fun(X8,X11)),X13),s(X12,sK6('const.util_prob.DFUNSET_2'(s(fun(X8,bool),X10),s(fun(X8,fun(X11,bool)),X9)),X14,X13,fun(X8,X11),X12)))),s(X15,X16)),X8,X11)),s(fun(X8,bool),X10)))) | p(s(bool,'const.bool.IN_2'(s(fun(X12,fun(X8,X11)),X13),s(fun(fun(X12,fun(X8,X11)),bool),'const.util_prob.FUNSET_2'(s(fun(X12,bool),X14),s(fun(fun(X8,X11),bool),'const.util_prob.DFUNSET_2'(s(fun(X8,bool),X10),s(fun(X8,fun(X11,bool)),X9))))))))) )),
  inference(resolution,[],[f6651,f85])).
fof(f7941,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK7(X1,X2,'const.combin.K_2'(s(fun(X0,X3),X4),s(X5,X6)),X0,X3)),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X7),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X3,happ(s(fun(X0,X3),X4),s(X0,X7))),s(fun(X3,bool),happ(s(fun(X0,fun(X3,bool)),X1),s(X0,X7))))))) )),
  inference(resolution,[],[f6651,f86])).
fof(f7940,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(fun(bool,X6),X10),s(fun(fun(bool,X6),bool),'const.util_prob.DFUNSET_2'(s(fun(bool,bool),X11),s(fun(bool,fun(X6,bool)),X8)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X11)))) | s(bool,'T_0') = s(bool,sK7(happ(s(fun(X7,fun(bool,fun(X6,bool))),'const.combin.K_1'(s(fun(bool,fun(X6,bool)),X8))),s(X7,X9)),X11,X10,bool,X6))) )),
  inference(forward_demodulation,[],[f7933,f80])).
fof(f7933,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(fun(bool,X6),X10),s(fun(fun(bool,X6),bool),'const.util_prob.DFUNSET_2'(s(fun(bool,bool),X11),s(fun(bool,fun(X6,bool)),'const.combin.K_2'(s(fun(bool,fun(X6,bool)),X8),s(X7,X9)))))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X11)))) | s(bool,'T_0') = s(bool,sK7(happ(s(fun(X7,fun(bool,fun(X6,bool))),'const.combin.K_1'(s(fun(bool,fun(X6,bool)),X8))),s(X7,X9)),X11,X10,bool,X6))) )),
  inference(superposition,[],[f6662,f81])).
fof(f7932,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(bool,X0),X4),s(fun(fun(bool,X0),bool),'const.util_prob.DFUNSET_2'(s(fun(bool,bool),X5),s(fun(bool,fun(X0,bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X5)))) | s(bool,'T_0') = s(bool,sK7('const.combin.K_2'(s(fun(bool,fun(X0,bool)),X1),s(X2,X3)),X5,X4,bool,X0))) )),
  inference(superposition,[],[f6662,f80])).
fof(f7939,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X7)))) | p(s(bool,'const.bool.IN_2'(s(fun(bool,X9),X10),s(fun(fun(bool,X9),bool),'const.util_prob.DFUNSET_2'(s(fun(bool,bool),X7),s(fun(bool,fun(X9,bool)),X11)))))) | s(bool,'T_0') = s(bool,sK7(X11,happ(s(fun(X6,fun(bool,bool)),'const.combin.K_1'(s(fun(bool,bool),X7))),s(X6,X8)),X10,bool,X9))) )),
  inference(forward_demodulation,[],[f7938,f80])).
fof(f7938,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.combin.K_2'(s(fun(bool,bool),X7),s(X6,X8)))))) | p(s(bool,'const.bool.IN_2'(s(fun(bool,X9),X10),s(fun(fun(bool,X9),bool),'const.util_prob.DFUNSET_2'(s(fun(bool,bool),X7),s(fun(bool,fun(X9,bool)),X11)))))) | s(bool,'T_0') = s(bool,sK7(X11,happ(s(fun(X6,fun(bool,bool)),'const.combin.K_1'(s(fun(bool,bool),X7))),s(X6,X8)),X10,bool,X9))) )),
  inference(forward_demodulation,[],[f7937,f81])).
fof(f7937,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(fun(bool,X9),X10),s(fun(fun(bool,X9),bool),'const.util_prob.DFUNSET_2'(s(fun(bool,bool),X7),s(fun(bool,fun(X9,bool)),X11)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),happ(s(fun(X6,fun(bool,bool)),'const.combin.K_1'(s(fun(bool,bool),X7))),s(X6,X8)))))) | s(bool,'T_0') = s(bool,sK7(X11,happ(s(fun(X6,fun(bool,bool)),'const.combin.K_1'(s(fun(bool,bool),X7))),s(X6,X8)),X10,bool,X9))) )),
  inference(forward_demodulation,[],[f7931,f80])).
fof(f7931,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(fun(bool,X9),X10),s(fun(fun(bool,X9),bool),'const.util_prob.DFUNSET_2'(s(fun(bool,bool),'const.combin.K_2'(s(fun(bool,bool),X7),s(X6,X8))),s(fun(bool,fun(X9,bool)),X11)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),happ(s(fun(X6,fun(bool,bool)),'const.combin.K_1'(s(fun(bool,bool),X7))),s(X6,X8)))))) | s(bool,'T_0') = s(bool,sK7(X11,happ(s(fun(X6,fun(bool,bool)),'const.combin.K_1'(s(fun(bool,bool),X7))),s(X6,X8)),X10,bool,X9))) )),
  inference(superposition,[],[f6662,f81])).
fof(f7936,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))) | p(s(bool,'const.bool.IN_2'(s(fun(bool,X3),X4),s(fun(fun(bool,X3),bool),'const.util_prob.DFUNSET_2'(s(fun(bool,bool),X0),s(fun(bool,fun(X3,bool)),X5)))))) | s(bool,'T_0') = s(bool,sK7(X5,'const.combin.K_2'(s(fun(bool,bool),X0),s(X1,X2)),X4,bool,X3))) )),
  inference(forward_demodulation,[],[f7930,f80])).
fof(f7930,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(bool,X3),X4),s(fun(fun(bool,X3),bool),'const.util_prob.DFUNSET_2'(s(fun(bool,bool),X0),s(fun(bool,fun(X3,bool)),X5)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.combin.K_2'(s(fun(bool,bool),X0),s(X1,X2)))))) | s(bool,'T_0') = s(bool,sK7(X5,'const.combin.K_2'(s(fun(bool,bool),X0),s(X1,X2)),X4,bool,X3))) )),
  inference(superposition,[],[f6662,f80])).
fof(f7935,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(fun(bool,X6),X8),s(fun(fun(bool,X6),bool),'const.util_prob.DFUNSET_2'(s(fun(bool,bool),X10),s(fun(bool,fun(X6,bool)),X11)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X10)))) | s(bool,'T_0') = s(bool,sK7(X11,X10,happ(s(fun(X7,fun(bool,X6)),'const.combin.K_1'(s(fun(bool,X6),X8))),s(X7,X9)),bool,X6))) )),
  inference(forward_demodulation,[],[f7929,f80])).
fof(f7929,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(fun(bool,X6),'const.combin.K_2'(s(fun(bool,X6),X8),s(X7,X9))),s(fun(fun(bool,X6),bool),'const.util_prob.DFUNSET_2'(s(fun(bool,bool),X10),s(fun(bool,fun(X6,bool)),X11)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X10)))) | s(bool,'T_0') = s(bool,sK7(X11,X10,happ(s(fun(X7,fun(bool,X6)),'const.combin.K_1'(s(fun(bool,X6),X8))),s(X7,X9)),bool,X6))) )),
  inference(superposition,[],[f6662,f81])).
fof(f7928,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(bool,X0),X1),s(fun(fun(bool,X0),bool),'const.util_prob.DFUNSET_2'(s(fun(bool,bool),X4),s(fun(bool,fun(X0,bool)),X5)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X4)))) | s(bool,'T_0') = s(bool,sK7(X5,X4,'const.combin.K_2'(s(fun(bool,X0),X1),s(X2,X3)),bool,X0))) )),
  inference(superposition,[],[f6662,f80])).
fof(f7926,plain,(
  ( ! [X14,X12,X17,X15,X13,X11,X16] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X11)))) | s(bool,'T_0') = s(bool,sK7(X12,X11,X13,bool,X14)) | p(s(bool,'const.bool.IN_2'(s(X15,happ(s(fun(fun(bool,X14),X15),X16),s(fun(bool,X14),X13))),s(fun(X15,bool),X17)))) | p(s(bool,'const.bool.IN_2'(s(fun(bool,X14),sK6(X17,'const.util_prob.DFUNSET_2'(s(fun(bool,bool),X11),s(fun(bool,fun(X14,bool)),X12)),X16,X15,fun(bool,X14))),s(fun(fun(bool,X14),bool),'const.util_prob.DFUNSET_2'(s(fun(bool,bool),X11),s(fun(bool,fun(X14,bool)),X12))))))) )),
  inference(resolution,[],[f6662,f6711])).
fof(f7925,plain,(
  ( ! [X6,X10,X8,X7,X5,X9] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X5)))) | s(bool,'T_0') = s(bool,sK7(X6,X5,happ(s(fun(X7,fun(bool,X8)),X9),s(X7,sK6('const.util_prob.DFUNSET_2'(s(fun(bool,bool),X5),s(fun(bool,fun(X8,bool)),X6)),X10,X9,fun(bool,X8),X7))),bool,X8)) | p(s(bool,'const.bool.IN_2'(s(fun(X7,fun(bool,X8)),X9),s(fun(fun(X7,fun(bool,X8)),bool),'const.util_prob.FUNSET_2'(s(fun(X7,bool),X10),s(fun(fun(bool,X8),bool),'const.util_prob.DFUNSET_2'(s(fun(bool,bool),X5),s(fun(bool,fun(X8,bool)),X6))))))))) )),
  inference(resolution,[],[f6662,f85])).
fof(f7924,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))) | s(bool,'T_0') = s(bool,sK7(X1,X0,X2,bool,X3)) | ~p(s(bool,'const.bool.IN_2'(s(bool,X4),s(fun(bool,bool),X0)))) | p(s(bool,'const.bool.IN_2'(s(X3,happ(s(fun(bool,X3),X2),s(bool,X4))),s(fun(X3,bool),happ(s(fun(bool,fun(X3,bool)),X1),s(bool,X4))))))) )),
  inference(resolution,[],[f6662,f86])).
fof(f7923,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X9,X16] : (p(s(bool,'const.bool.IN_2'(s(fun(X13,X9),X14),s(fun(fun(X13,X9),bool),'const.util_prob.FUNSET_2'(s(fun(X13,bool),X15),s(fun(X9,bool),X11)))))) | p(s(bool,'const.bool.IN_2'(s(X13,sK6(happ(s(fun(X10,fun(X9,bool)),'const.combin.K_1'(s(fun(X9,bool),X11))),s(X10,X12)),'const.combin.K_2'(s(fun(X13,bool),X15),s(X16,X17)),X14,X9,X13)),s(fun(X13,bool),X15))))) )),
  inference(forward_demodulation,[],[f7906,f80])).
fof(f7906,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X9,X16] : (p(s(bool,'const.bool.IN_2'(s(fun(X13,X9),X14),s(fun(fun(X13,X9),bool),'const.util_prob.FUNSET_2'(s(fun(X13,bool),X15),s(fun(X9,bool),'const.combin.K_2'(s(fun(X9,bool),X11),s(X10,X12)))))))) | p(s(bool,'const.bool.IN_2'(s(X13,sK6(happ(s(fun(X10,fun(X9,bool)),'const.combin.K_1'(s(fun(X9,bool),X11))),s(X10,X12)),'const.combin.K_2'(s(fun(X13,bool),X15),s(X16,X17)),X14,X9,X13)),s(fun(X13,bool),X15))))) )),
  inference(superposition,[],[f6637,f81])).
fof(f7918,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X9,X16] : (p(s(bool,'const.bool.IN_2'(s(fun(X9,X10),X12),s(fun(fun(X9,X10),bool),'const.util_prob.FUNSET_2'(s(fun(X9,bool),X14),s(fun(X10,bool),X15)))))) | p(s(bool,'const.bool.IN_2'(s(X9,sK6(X15,'const.combin.K_2'(s(fun(X9,bool),X14),s(X16,X17)),happ(s(fun(X11,fun(X9,X10)),'const.combin.K_1'(s(fun(X9,X10),X12))),s(X11,X13)),X10,X9)),s(fun(X9,bool),X14))))) )),
  inference(forward_demodulation,[],[f7902,f80])).
fof(f7902,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X9,X16] : (p(s(bool,'const.bool.IN_2'(s(fun(X9,X10),'const.combin.K_2'(s(fun(X9,X10),X12),s(X11,X13))),s(fun(fun(X9,X10),bool),'const.util_prob.FUNSET_2'(s(fun(X9,bool),X14),s(fun(X10,bool),X15)))))) | p(s(bool,'const.bool.IN_2'(s(X9,sK6(X15,'const.combin.K_2'(s(fun(X9,bool),X14),s(X16,X17)),happ(s(fun(X11,fun(X9,X10)),'const.combin.K_1'(s(fun(X9,X10),X12))),s(X11,X13)),X10,X9)),s(fun(X9,bool),X14))))) )),
  inference(superposition,[],[f6637,f81])).
fof(f7864,plain,(
  ( ! [X26,X24,X23,X21,X19,X17,X25,X22,X20,X18] : (p(s(bool,'const.bool.IN_2'(s(X17,sK6(X18,'const.combin.K_2'(s(fun(X17,bool),X19),s(X20,X21)),X22,X23,X17)),s(fun(X17,bool),X19)))) | p(s(bool,'const.bool.IN_2'(s(X24,happ(s(fun(fun(X17,X23),X24),X25),s(fun(X17,X23),X22))),s(fun(X24,bool),X26)))) | p(s(bool,'const.bool.IN_2'(s(fun(X17,X23),sK6(X26,'const.util_prob.FUNSET_2'(s(fun(X17,bool),X19),s(fun(X23,bool),X18)),X25,X24,fun(X17,X23))),s(fun(fun(X17,X23),bool),'const.util_prob.FUNSET_2'(s(fun(X17,bool),X19),s(fun(X23,bool),X18))))))) )),
  inference(resolution,[],[f6637,f6711])).
fof(f7863,plain,(
  ( ! [X14,X12,X10,X8,X15,X13,X11,X9,X16] : (p(s(bool,'const.bool.IN_2'(s(X8,sK6(X9,'const.combin.K_2'(s(fun(X8,bool),X10),s(X11,X12)),happ(s(fun(X13,fun(X8,X14)),X15),s(X13,sK6('const.util_prob.FUNSET_2'(s(fun(X8,bool),X10),s(fun(X14,bool),X9)),X16,X15,fun(X8,X14),X13))),X14,X8)),s(fun(X8,bool),X10)))) | p(s(bool,'const.bool.IN_2'(s(fun(X13,fun(X8,X14)),X15),s(fun(fun(X13,fun(X8,X14)),bool),'const.util_prob.FUNSET_2'(s(fun(X13,bool),X16),s(fun(fun(X8,X14),bool),'const.util_prob.FUNSET_2'(s(fun(X8,bool),X10),s(fun(X14,bool),X9))))))))) )),
  inference(resolution,[],[f6637,f85])).
fof(f7860,plain,(
  ( ! [X39,X37,X35,X43,X41,X38,X36,X34,X42,X40] : (p(s(bool,'const.bool.IN_2'(s(fun(fun(X34,X35),X36),X37),s(fun(fun(fun(X34,X35),X36),bool),'const.util_prob.FUNSET_2'(s(fun(fun(X34,X35),bool),'const.util_prob.DFUNSET_2'(s(fun(X34,bool),X38),s(fun(X34,fun(X35,bool)),X39))),s(fun(X36,bool),X40)))))) | ~p(s(bool,'const.bool.IN_2'(s(X34,X41),s(fun(X34,bool),X38)))) | p(s(bool,'const.bool.IN_2'(s(X35,happ(s(fun(X34,X35),sK6(X40,'const.combin.K_2'(s(fun(fun(X34,X35),bool),'const.util_prob.DFUNSET_2'(s(fun(X34,bool),X38),s(fun(X34,fun(X35,bool)),X39))),s(X42,X43)),X37,X36,fun(X34,X35))),s(X34,X41))),s(fun(X35,bool),happ(s(fun(X34,fun(X35,bool)),X39),s(X34,X41))))))) )),
  inference(resolution,[],[f6637,f86])).
fof(f7859,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (p(s(bool,'const.bool.IN_2'(s(fun(fun(X24,X25),X26),X27),s(fun(fun(fun(X24,X25),X26),bool),'const.util_prob.FUNSET_2'(s(fun(fun(X24,X25),bool),'const.util_prob.FUNSET_2'(s(fun(X24,bool),X28),s(fun(X25,bool),X29))),s(fun(X26,bool),X30)))))) | ~p(s(bool,'const.bool.IN_2'(s(X24,X31),s(fun(X24,bool),X28)))) | p(s(bool,'const.bool.IN_2'(s(X25,happ(s(fun(X24,X25),sK6(X30,'const.combin.K_2'(s(fun(fun(X24,X25),bool),'const.util_prob.FUNSET_2'(s(fun(X24,bool),X28),s(fun(X25,bool),X29))),s(X32,X33)),X27,X26,fun(X24,X25))),s(X24,X31))),s(fun(X25,bool),X29))))) )),
  inference(resolution,[],[f6637,f83])).
fof(f7858,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (p(s(bool,'const.bool.IN_2'(s(fun(X14,X15),X16),s(fun(fun(X14,X15),bool),'const.util_prob.FUNSET_2'(s(fun(X14,bool),X17),s(fun(X15,bool),X18)))))) | p(s(bool,'const.bool.IN_2'(s(X19,happ(s(fun(X14,X19),X20),s(X14,sK6(X18,'const.combin.K_2'(s(fun(X14,bool),X17),s(X21,X22)),X16,X15,X14)))),s(fun(X19,bool),X23)))) | p(s(bool,'const.bool.IN_2'(s(X14,sK6(X23,X17,X20,X19,X14)),s(fun(X14,bool),X17))))) )),
  inference(resolution,[],[f6637,f6711])).
fof(f7841,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X9,X16] : (p(s(bool,'const.bool.IN_2'(s(X9,sK6('const.combin.K_2'(s(fun(X13,bool),X15),s(X16,X17)),happ(s(fun(X10,fun(X9,bool)),'const.combin.K_1'(s(fun(X9,bool),X11))),s(X10,X12)),X14,X13,X9)),s(fun(X9,bool),X11)))) | p(s(bool,'const.bool.IN_2'(s(fun(X9,X13),X14),s(fun(fun(X9,X13),bool),'const.util_prob.FUNSET_2'(s(fun(X9,bool),X11),s(fun(X13,bool),X15))))))) )),
  inference(forward_demodulation,[],[f7840,f80])).
fof(f7840,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X9,X16] : (p(s(bool,'const.bool.IN_2'(s(X9,sK6('const.combin.K_2'(s(fun(X13,bool),X15),s(X16,X17)),happ(s(fun(X10,fun(X9,bool)),'const.combin.K_1'(s(fun(X9,bool),X11))),s(X10,X12)),X14,X13,X9)),s(fun(X9,bool),'const.combin.K_2'(s(fun(X9,bool),X11),s(X10,X12)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X9,X13),X14),s(fun(fun(X9,X13),bool),'const.util_prob.FUNSET_2'(s(fun(X9,bool),X11),s(fun(X13,bool),X15))))))) )),
  inference(forward_demodulation,[],[f7839,f81])).
fof(f7839,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X9,X16] : (p(s(bool,'const.bool.IN_2'(s(fun(X9,X13),X14),s(fun(fun(X9,X13),bool),'const.util_prob.FUNSET_2'(s(fun(X9,bool),X11),s(fun(X13,bool),X15)))))) | p(s(bool,'const.bool.IN_2'(s(X9,sK6('const.combin.K_2'(s(fun(X13,bool),X15),s(X16,X17)),happ(s(fun(X10,fun(X9,bool)),'const.combin.K_1'(s(fun(X9,bool),X11))),s(X10,X12)),X14,X13,X9)),s(fun(X9,bool),happ(s(fun(X10,fun(X9,bool)),'const.combin.K_1'(s(fun(X9,bool),X11))),s(X10,X12))))))) )),
  inference(forward_demodulation,[],[f7792,f80])).
fof(f7792,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X9,X16] : (p(s(bool,'const.bool.IN_2'(s(fun(X9,X13),X14),s(fun(fun(X9,X13),bool),'const.util_prob.FUNSET_2'(s(fun(X9,bool),'const.combin.K_2'(s(fun(X9,bool),X11),s(X10,X12))),s(fun(X13,bool),X15)))))) | p(s(bool,'const.bool.IN_2'(s(X9,sK6('const.combin.K_2'(s(fun(X13,bool),X15),s(X16,X17)),happ(s(fun(X10,fun(X9,bool)),'const.combin.K_1'(s(fun(X9,bool),X11))),s(X10,X12)),X14,X13,X9)),s(fun(X9,bool),happ(s(fun(X10,fun(X9,bool)),'const.combin.K_1'(s(fun(X9,bool),X11))),s(X10,X12))))))) )),
  inference(superposition,[],[f6620,f81])).
fof(f7838,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK6('const.combin.K_2'(s(fun(X4,bool),X6),s(X7,X8)),'const.combin.K_2'(s(fun(X0,bool),X1),s(X2,X3)),X5,X4,X0)),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,X4),X5),s(fun(fun(X0,X4),bool),'const.util_prob.FUNSET_2'(s(fun(X0,bool),X1),s(fun(X4,bool),X6))))))) )),
  inference(forward_demodulation,[],[f7791,f80])).
fof(f7791,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,X4),X5),s(fun(fun(X0,X4),bool),'const.util_prob.FUNSET_2'(s(fun(X0,bool),X1),s(fun(X4,bool),X6)))))) | p(s(bool,'const.bool.IN_2'(s(X0,sK6('const.combin.K_2'(s(fun(X4,bool),X6),s(X7,X8)),'const.combin.K_2'(s(fun(X0,bool),X1),s(X2,X3)),X5,X4,X0)),s(fun(X0,bool),'const.combin.K_2'(s(fun(X0,bool),X1),s(X2,X3))))))) )),
  inference(superposition,[],[f6620,f80])).
fof(f7837,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X9,X16] : (p(s(bool,'const.bool.IN_2'(s(fun(X9,X10),X12),s(fun(fun(X9,X10),bool),'const.util_prob.FUNSET_2'(s(fun(X9,bool),X14),s(fun(X10,bool),X15)))))) | p(s(bool,'const.bool.IN_2'(s(X9,sK6('const.combin.K_2'(s(fun(X10,bool),X15),s(X16,X17)),X14,happ(s(fun(X11,fun(X9,X10)),'const.combin.K_1'(s(fun(X9,X10),X12))),s(X11,X13)),X10,X9)),s(fun(X9,bool),X14))))) )),
  inference(forward_demodulation,[],[f7790,f80])).
fof(f7790,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X9,X16] : (p(s(bool,'const.bool.IN_2'(s(fun(X9,X10),'const.combin.K_2'(s(fun(X9,X10),X12),s(X11,X13))),s(fun(fun(X9,X10),bool),'const.util_prob.FUNSET_2'(s(fun(X9,bool),X14),s(fun(X10,bool),X15)))))) | p(s(bool,'const.bool.IN_2'(s(X9,sK6('const.combin.K_2'(s(fun(X10,bool),X15),s(X16,X17)),X14,happ(s(fun(X11,fun(X9,X10)),'const.combin.K_1'(s(fun(X9,X10),X12))),s(X11,X13)),X10,X9)),s(fun(X9,bool),X14))))) )),
  inference(superposition,[],[f6620,f81])).
fof(f7787,plain,(
  ( ! [X30,X39,X37,X35,X33,X31,X38,X36,X34,X32] : (p(s(bool,'const.bool.IN_2'(s(fun(fun(X30,X31),X32),X33),s(fun(fun(fun(X30,X31),X32),bool),'const.util_prob.FUNSET_2'(s(fun(fun(X30,X31),bool),'const.util_prob.DFUNSET_2'(s(fun(X30,bool),X34),s(fun(X30,fun(X31,bool)),X35))),s(fun(X32,bool),X36)))))) | ~p(s(bool,'const.bool.IN_2'(s(X30,X37),s(fun(X30,bool),X34)))) | p(s(bool,'const.bool.IN_2'(s(X31,happ(s(fun(X30,X31),sK6('const.combin.K_2'(s(fun(X32,bool),X36),s(X38,X39)),'const.util_prob.DFUNSET_2'(s(fun(X30,bool),X34),s(fun(X30,fun(X31,bool)),X35)),X33,X32,fun(X30,X31))),s(X30,X37))),s(fun(X31,bool),happ(s(fun(X30,fun(X31,bool)),X35),s(X30,X37))))))) )),
  inference(resolution,[],[f6620,f86])).
fof(f7786,plain,(
  ( ! [X28,X26,X24,X23,X21,X29,X27,X25,X22,X20] : (p(s(bool,'const.bool.IN_2'(s(fun(fun(X20,X21),X22),X23),s(fun(fun(fun(X20,X21),X22),bool),'const.util_prob.FUNSET_2'(s(fun(fun(X20,X21),bool),'const.util_prob.FUNSET_2'(s(fun(X20,bool),X24),s(fun(X21,bool),X25))),s(fun(X22,bool),X26)))))) | ~p(s(bool,'const.bool.IN_2'(s(X20,X27),s(fun(X20,bool),X24)))) | p(s(bool,'const.bool.IN_2'(s(X21,happ(s(fun(X20,X21),sK6('const.combin.K_2'(s(fun(X22,bool),X26),s(X28,X29)),'const.util_prob.FUNSET_2'(s(fun(X20,bool),X24),s(fun(X21,bool),X25)),X23,X22,fun(X20,X21))),s(X20,X27))),s(fun(X21,bool),X25))))) )),
  inference(resolution,[],[f6620,f83])).
fof(f7785,plain,(
  ( ! [X14,X12,X10,X19,X17,X15,X13,X11,X18,X16] : (p(s(bool,'const.bool.IN_2'(s(fun(X10,X11),X12),s(fun(fun(X10,X11),bool),'const.util_prob.FUNSET_2'(s(fun(X10,bool),X13),s(fun(X11,bool),X14)))))) | p(s(bool,'const.bool.IN_2'(s(X15,happ(s(fun(X10,X15),X16),s(X10,sK6('const.combin.K_2'(s(fun(X11,bool),X14),s(X17,X18)),X13,X12,X11,X10)))),s(fun(X15,bool),X19)))) | p(s(bool,'const.bool.IN_2'(s(X10,sK6(X19,X13,X16,X15,X10)),s(fun(X10,bool),X13))))) )),
  inference(resolution,[],[f6620,f6711])).
fof(f7781,plain,(
  ( ! [X26,X24,X23,X21,X19,X17,X25,X22,X20,X18] : (p(s(bool,'const.bool.IN_2'(s(X17,sK6('const.combin.K_2'(s(fun(X18,bool),X19),s(X20,X21)),X22,X23,X18,X17)),s(fun(X17,bool),X22)))) | p(s(bool,'const.bool.IN_2'(s(X24,happ(s(fun(fun(X17,X18),X24),X25),s(fun(X17,X18),X23))),s(fun(X24,bool),X26)))) | p(s(bool,'const.bool.IN_2'(s(fun(X17,X18),sK6(X26,'const.util_prob.FUNSET_2'(s(fun(X17,bool),X22),s(fun(X18,bool),X19)),X25,X24,fun(X17,X18))),s(fun(fun(X17,X18),bool),'const.util_prob.FUNSET_2'(s(fun(X17,bool),X22),s(fun(X18,bool),X19))))))) )),
  inference(resolution,[],[f6620,f6711])).
fof(f7780,plain,(
  ( ! [X14,X12,X10,X8,X15,X13,X11,X9,X16] : (p(s(bool,'const.bool.IN_2'(s(X8,sK6('const.combin.K_2'(s(fun(X9,bool),X10),s(X11,X12)),X13,happ(s(fun(X14,fun(X8,X9)),X15),s(X14,sK6('const.util_prob.FUNSET_2'(s(fun(X8,bool),X13),s(fun(X9,bool),X10)),X16,X15,fun(X8,X9),X14))),X9,X8)),s(fun(X8,bool),X13)))) | p(s(bool,'const.bool.IN_2'(s(fun(X14,fun(X8,X9)),X15),s(fun(fun(X14,fun(X8,X9)),bool),'const.util_prob.FUNSET_2'(s(fun(X14,bool),X16),s(fun(fun(X8,X9),bool),'const.util_prob.FUNSET_2'(s(fun(X8,bool),X13),s(fun(X9,bool),X10))))))))) )),
  inference(resolution,[],[f6620,f85])).
fof(f7779,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK6('const.combin.K_2'(s(fun(X1,bool),X2),s(X3,X4)),X5,X6,X1,X0)),s(fun(X0,bool),X5)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X7),s(fun(X0,bool),X5)))) | p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X6),s(X0,X7))),s(fun(X1,bool),X2))))) )),
  inference(resolution,[],[f6620,f83])).
fof(f7770,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X9,X16] : (p(s(bool,'const.bool.IN_2'(s(fun(X13,X9),X14),s(fun(fun(X13,X9),bool),'const.util_prob.FUNSET_2'(s(fun(X13,bool),X15),s(fun(X9,bool),X11)))))) | p(s(bool,'const.bool.IN_2'(s(X13,sK6(happ(s(fun(X10,fun(X9,bool)),'const.combin.K_1'(s(fun(X9,bool),X11))),s(X10,X12)),X15,'const.combin.K_2'(s(fun(X13,X9),X14),s(X16,X17)),X9,X13)),s(fun(X13,bool),X15))))) )),
  inference(forward_demodulation,[],[f7719,f80])).
fof(f7719,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X9,X16] : (p(s(bool,'const.bool.IN_2'(s(fun(X13,X9),X14),s(fun(fun(X13,X9),bool),'const.util_prob.FUNSET_2'(s(fun(X13,bool),X15),s(fun(X9,bool),'const.combin.K_2'(s(fun(X9,bool),X11),s(X10,X12)))))))) | p(s(bool,'const.bool.IN_2'(s(X13,sK6(happ(s(fun(X10,fun(X9,bool)),'const.combin.K_1'(s(fun(X9,bool),X11))),s(X10,X12)),X15,'const.combin.K_2'(s(fun(X13,X9),X14),s(X16,X17)),X9,X13)),s(fun(X13,bool),X15))))) )),
  inference(superposition,[],[f6616,f81])).
fof(f7718,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(X4,X0),X5),s(fun(fun(X4,X0),bool),'const.util_prob.FUNSET_2'(s(fun(X4,bool),X6),s(fun(X0,bool),X1)))))) | p(s(bool,'const.bool.IN_2'(s(X4,sK6('const.combin.K_2'(s(fun(X0,bool),X1),s(X2,X3)),X6,'const.combin.K_2'(s(fun(X4,X0),X5),s(X7,X8)),X0,X4)),s(fun(X4,bool),X6))))) )),
  inference(superposition,[],[f6616,f80])).
fof(f7769,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X9,X16] : (p(s(bool,'const.bool.IN_2'(s(X9,sK6(X15,happ(s(fun(X10,fun(X9,bool)),'const.combin.K_1'(s(fun(X9,bool),X11))),s(X10,X12)),'const.combin.K_2'(s(fun(X9,X13),X14),s(X16,X17)),X13,X9)),s(fun(X9,bool),X11)))) | p(s(bool,'const.bool.IN_2'(s(fun(X9,X13),X14),s(fun(fun(X9,X13),bool),'const.util_prob.FUNSET_2'(s(fun(X9,bool),X11),s(fun(X13,bool),X15))))))) )),
  inference(forward_demodulation,[],[f7768,f80])).
fof(f7768,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X9,X16] : (p(s(bool,'const.bool.IN_2'(s(X9,sK6(X15,happ(s(fun(X10,fun(X9,bool)),'const.combin.K_1'(s(fun(X9,bool),X11))),s(X10,X12)),'const.combin.K_2'(s(fun(X9,X13),X14),s(X16,X17)),X13,X9)),s(fun(X9,bool),'const.combin.K_2'(s(fun(X9,bool),X11),s(X10,X12)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X9,X13),X14),s(fun(fun(X9,X13),bool),'const.util_prob.FUNSET_2'(s(fun(X9,bool),X11),s(fun(X13,bool),X15))))))) )),
  inference(forward_demodulation,[],[f7767,f81])).
fof(f7767,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X9,X16] : (p(s(bool,'const.bool.IN_2'(s(fun(X9,X13),X14),s(fun(fun(X9,X13),bool),'const.util_prob.FUNSET_2'(s(fun(X9,bool),X11),s(fun(X13,bool),X15)))))) | p(s(bool,'const.bool.IN_2'(s(X9,sK6(X15,happ(s(fun(X10,fun(X9,bool)),'const.combin.K_1'(s(fun(X9,bool),X11))),s(X10,X12)),'const.combin.K_2'(s(fun(X9,X13),X14),s(X16,X17)),X13,X9)),s(fun(X9,bool),happ(s(fun(X10,fun(X9,bool)),'const.combin.K_1'(s(fun(X9,bool),X11))),s(X10,X12))))))) )),
  inference(forward_demodulation,[],[f7717,f80])).
fof(f7717,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X9,X16] : (p(s(bool,'const.bool.IN_2'(s(fun(X9,X13),X14),s(fun(fun(X9,X13),bool),'const.util_prob.FUNSET_2'(s(fun(X9,bool),'const.combin.K_2'(s(fun(X9,bool),X11),s(X10,X12))),s(fun(X13,bool),X15)))))) | p(s(bool,'const.bool.IN_2'(s(X9,sK6(X15,happ(s(fun(X10,fun(X9,bool)),'const.combin.K_1'(s(fun(X9,bool),X11))),s(X10,X12)),'const.combin.K_2'(s(fun(X9,X13),X14),s(X16,X17)),X13,X9)),s(fun(X9,bool),happ(s(fun(X10,fun(X9,bool)),'const.combin.K_1'(s(fun(X9,bool),X11))),s(X10,X12))))))) )),
  inference(superposition,[],[f6616,f81])).
fof(f7766,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK6(X6,'const.combin.K_2'(s(fun(X0,bool),X1),s(X2,X3)),'const.combin.K_2'(s(fun(X0,X4),X5),s(X7,X8)),X4,X0)),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,X4),X5),s(fun(fun(X0,X4),bool),'const.util_prob.FUNSET_2'(s(fun(X0,bool),X1),s(fun(X4,bool),X6))))))) )),
  inference(forward_demodulation,[],[f7716,f80])).
fof(f7716,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,X4),X5),s(fun(fun(X0,X4),bool),'const.util_prob.FUNSET_2'(s(fun(X0,bool),X1),s(fun(X4,bool),X6)))))) | p(s(bool,'const.bool.IN_2'(s(X0,sK6(X6,'const.combin.K_2'(s(fun(X0,bool),X1),s(X2,X3)),'const.combin.K_2'(s(fun(X0,X4),X5),s(X7,X8)),X4,X0)),s(fun(X0,bool),'const.combin.K_2'(s(fun(X0,bool),X1),s(X2,X3))))))) )),
  inference(superposition,[],[f6616,f80])).
fof(f7712,plain,(
  ( ! [X30,X39,X37,X35,X33,X31,X38,X36,X34,X32] : (p(s(bool,'const.bool.IN_2'(s(fun(fun(X30,X31),X32),X33),s(fun(fun(fun(X30,X31),X32),bool),'const.util_prob.FUNSET_2'(s(fun(fun(X30,X31),bool),'const.util_prob.DFUNSET_2'(s(fun(X30,bool),X34),s(fun(X30,fun(X31,bool)),X35))),s(fun(X32,bool),X36)))))) | ~p(s(bool,'const.bool.IN_2'(s(X30,X37),s(fun(X30,bool),X34)))) | p(s(bool,'const.bool.IN_2'(s(X31,happ(s(fun(X30,X31),sK6(X36,'const.util_prob.DFUNSET_2'(s(fun(X30,bool),X34),s(fun(X30,fun(X31,bool)),X35)),'const.combin.K_2'(s(fun(fun(X30,X31),X32),X33),s(X38,X39)),X32,fun(X30,X31))),s(X30,X37))),s(fun(X31,bool),happ(s(fun(X30,fun(X31,bool)),X35),s(X30,X37))))))) )),
  inference(resolution,[],[f6616,f86])).
fof(f7711,plain,(
  ( ! [X28,X26,X24,X23,X21,X29,X27,X25,X22,X20] : (p(s(bool,'const.bool.IN_2'(s(fun(fun(X20,X21),X22),X23),s(fun(fun(fun(X20,X21),X22),bool),'const.util_prob.FUNSET_2'(s(fun(fun(X20,X21),bool),'const.util_prob.FUNSET_2'(s(fun(X20,bool),X24),s(fun(X21,bool),X25))),s(fun(X22,bool),X26)))))) | ~p(s(bool,'const.bool.IN_2'(s(X20,X27),s(fun(X20,bool),X24)))) | p(s(bool,'const.bool.IN_2'(s(X21,happ(s(fun(X20,X21),sK6(X26,'const.util_prob.FUNSET_2'(s(fun(X20,bool),X24),s(fun(X21,bool),X25)),'const.combin.K_2'(s(fun(fun(X20,X21),X22),X23),s(X28,X29)),X22,fun(X20,X21))),s(X20,X27))),s(fun(X21,bool),X25))))) )),
  inference(resolution,[],[f6616,f83])).
fof(f7710,plain,(
  ( ! [X14,X12,X10,X19,X17,X15,X13,X11,X18,X16] : (p(s(bool,'const.bool.IN_2'(s(fun(X10,X11),X12),s(fun(fun(X10,X11),bool),'const.util_prob.FUNSET_2'(s(fun(X10,bool),X13),s(fun(X11,bool),X14)))))) | p(s(bool,'const.bool.IN_2'(s(X15,happ(s(fun(X10,X15),X16),s(X10,sK6(X14,X13,'const.combin.K_2'(s(fun(X10,X11),X12),s(X17,X18)),X11,X10)))),s(fun(X15,bool),X19)))) | p(s(bool,'const.bool.IN_2'(s(X10,sK6(X19,X13,X16,X15,X10)),s(fun(X10,bool),X13))))) )),
  inference(resolution,[],[f6616,f6711])).
fof(f7706,plain,(
  ( ! [X26,X24,X23,X21,X19,X17,X25,X22,X20,X18] : (p(s(bool,'const.bool.IN_2'(s(X17,sK6(X18,X19,'const.combin.K_2'(s(fun(X17,X20),X21),s(X22,X23)),X20,X17)),s(fun(X17,bool),X19)))) | p(s(bool,'const.bool.IN_2'(s(X24,happ(s(fun(fun(X17,X20),X24),X25),s(fun(X17,X20),X21))),s(fun(X24,bool),X26)))) | p(s(bool,'const.bool.IN_2'(s(fun(X17,X20),sK6(X26,'const.util_prob.FUNSET_2'(s(fun(X17,bool),X19),s(fun(X20,bool),X18)),X25,X24,fun(X17,X20))),s(fun(fun(X17,X20),bool),'const.util_prob.FUNSET_2'(s(fun(X17,bool),X19),s(fun(X20,bool),X18))))))) )),
  inference(resolution,[],[f6616,f6711])).
fof(f7705,plain,(
  ( ! [X14,X12,X10,X8,X15,X13,X11,X9,X16] : (p(s(bool,'const.bool.IN_2'(s(X8,sK6(X9,X10,'const.combin.K_2'(s(fun(X8,X11),happ(s(fun(X12,fun(X8,X11)),X13),s(X12,sK6('const.util_prob.FUNSET_2'(s(fun(X8,bool),X10),s(fun(X11,bool),X9)),X14,X13,fun(X8,X11),X12)))),s(X15,X16)),X11,X8)),s(fun(X8,bool),X10)))) | p(s(bool,'const.bool.IN_2'(s(fun(X12,fun(X8,X11)),X13),s(fun(fun(X12,fun(X8,X11)),bool),'const.util_prob.FUNSET_2'(s(fun(X12,bool),X14),s(fun(fun(X8,X11),bool),'const.util_prob.FUNSET_2'(s(fun(X8,bool),X10),s(fun(X11,bool),X9))))))))) )),
  inference(resolution,[],[f6616,f85])).
fof(f7704,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK6(X1,X2,'const.combin.K_2'(s(fun(X0,X3),X4),s(X5,X6)),X3,X0)),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X7),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X3,happ(s(fun(X0,X3),X4),s(X0,X7))),s(fun(X3,bool),X1))))) )),
  inference(resolution,[],[f6616,f83])).
fof(f7703,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(fun(bool,X6),X10),s(fun(fun(bool,X6),bool),'const.util_prob.FUNSET_2'(s(fun(bool,bool),X11),s(fun(X6,bool),X8)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X11)))) | s(bool,'T_0') = s(bool,sK6(happ(s(fun(X7,fun(X6,bool)),'const.combin.K_1'(s(fun(X6,bool),X8))),s(X7,X9)),X11,X10,X6,bool))) )),
  inference(forward_demodulation,[],[f7696,f80])).
fof(f7696,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(fun(bool,X6),X10),s(fun(fun(bool,X6),bool),'const.util_prob.FUNSET_2'(s(fun(bool,bool),X11),s(fun(X6,bool),'const.combin.K_2'(s(fun(X6,bool),X8),s(X7,X9)))))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X11)))) | s(bool,'T_0') = s(bool,sK6(happ(s(fun(X7,fun(X6,bool)),'const.combin.K_1'(s(fun(X6,bool),X8))),s(X7,X9)),X11,X10,X6,bool))) )),
  inference(superposition,[],[f6627,f81])).
fof(f7695,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(bool,X0),X4),s(fun(fun(bool,X0),bool),'const.util_prob.FUNSET_2'(s(fun(bool,bool),X5),s(fun(X0,bool),X1)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X5)))) | s(bool,'T_0') = s(bool,sK6('const.combin.K_2'(s(fun(X0,bool),X1),s(X2,X3)),X5,X4,X0,bool))) )),
  inference(superposition,[],[f6627,f80])).
fof(f7702,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X7)))) | p(s(bool,'const.bool.IN_2'(s(fun(bool,X9),X10),s(fun(fun(bool,X9),bool),'const.util_prob.FUNSET_2'(s(fun(bool,bool),X7),s(fun(X9,bool),X11)))))) | s(bool,'T_0') = s(bool,sK6(X11,happ(s(fun(X6,fun(bool,bool)),'const.combin.K_1'(s(fun(bool,bool),X7))),s(X6,X8)),X10,X9,bool))) )),
  inference(forward_demodulation,[],[f7701,f80])).
fof(f7701,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.combin.K_2'(s(fun(bool,bool),X7),s(X6,X8)))))) | p(s(bool,'const.bool.IN_2'(s(fun(bool,X9),X10),s(fun(fun(bool,X9),bool),'const.util_prob.FUNSET_2'(s(fun(bool,bool),X7),s(fun(X9,bool),X11)))))) | s(bool,'T_0') = s(bool,sK6(X11,happ(s(fun(X6,fun(bool,bool)),'const.combin.K_1'(s(fun(bool,bool),X7))),s(X6,X8)),X10,X9,bool))) )),
  inference(forward_demodulation,[],[f7700,f81])).
fof(f7700,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(fun(bool,X9),X10),s(fun(fun(bool,X9),bool),'const.util_prob.FUNSET_2'(s(fun(bool,bool),X7),s(fun(X9,bool),X11)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),happ(s(fun(X6,fun(bool,bool)),'const.combin.K_1'(s(fun(bool,bool),X7))),s(X6,X8)))))) | s(bool,'T_0') = s(bool,sK6(X11,happ(s(fun(X6,fun(bool,bool)),'const.combin.K_1'(s(fun(bool,bool),X7))),s(X6,X8)),X10,X9,bool))) )),
  inference(forward_demodulation,[],[f7694,f80])).
fof(f7694,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(fun(bool,X9),X10),s(fun(fun(bool,X9),bool),'const.util_prob.FUNSET_2'(s(fun(bool,bool),'const.combin.K_2'(s(fun(bool,bool),X7),s(X6,X8))),s(fun(X9,bool),X11)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),happ(s(fun(X6,fun(bool,bool)),'const.combin.K_1'(s(fun(bool,bool),X7))),s(X6,X8)))))) | s(bool,'T_0') = s(bool,sK6(X11,happ(s(fun(X6,fun(bool,bool)),'const.combin.K_1'(s(fun(bool,bool),X7))),s(X6,X8)),X10,X9,bool))) )),
  inference(superposition,[],[f6627,f81])).
fof(f7699,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))) | p(s(bool,'const.bool.IN_2'(s(fun(bool,X3),X4),s(fun(fun(bool,X3),bool),'const.util_prob.FUNSET_2'(s(fun(bool,bool),X0),s(fun(X3,bool),X5)))))) | s(bool,'T_0') = s(bool,sK6(X5,'const.combin.K_2'(s(fun(bool,bool),X0),s(X1,X2)),X4,X3,bool))) )),
  inference(forward_demodulation,[],[f7693,f80])).
fof(f7693,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(bool,X3),X4),s(fun(fun(bool,X3),bool),'const.util_prob.FUNSET_2'(s(fun(bool,bool),X0),s(fun(X3,bool),X5)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.combin.K_2'(s(fun(bool,bool),X0),s(X1,X2)))))) | s(bool,'T_0') = s(bool,sK6(X5,'const.combin.K_2'(s(fun(bool,bool),X0),s(X1,X2)),X4,X3,bool))) )),
  inference(superposition,[],[f6627,f80])).
fof(f7698,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(fun(bool,X6),X8),s(fun(fun(bool,X6),bool),'const.util_prob.FUNSET_2'(s(fun(bool,bool),X10),s(fun(X6,bool),X11)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X10)))) | s(bool,'T_0') = s(bool,sK6(X11,X10,happ(s(fun(X7,fun(bool,X6)),'const.combin.K_1'(s(fun(bool,X6),X8))),s(X7,X9)),X6,bool))) )),
  inference(forward_demodulation,[],[f7692,f80])).
fof(f7692,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(fun(bool,X6),'const.combin.K_2'(s(fun(bool,X6),X8),s(X7,X9))),s(fun(fun(bool,X6),bool),'const.util_prob.FUNSET_2'(s(fun(bool,bool),X10),s(fun(X6,bool),X11)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X10)))) | s(bool,'T_0') = s(bool,sK6(X11,X10,happ(s(fun(X7,fun(bool,X6)),'const.combin.K_1'(s(fun(bool,X6),X8))),s(X7,X9)),X6,bool))) )),
  inference(superposition,[],[f6627,f81])).
fof(f7691,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(bool,X0),X1),s(fun(fun(bool,X0),bool),'const.util_prob.FUNSET_2'(s(fun(bool,bool),X4),s(fun(X0,bool),X5)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X4)))) | s(bool,'T_0') = s(bool,sK6(X5,X4,'const.combin.K_2'(s(fun(bool,X0),X1),s(X2,X3)),X0,bool))) )),
  inference(superposition,[],[f6627,f80])).
fof(f7689,plain,(
  ( ! [X14,X12,X17,X15,X13,X11,X16] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X11)))) | s(bool,'T_0') = s(bool,sK6(X12,X11,X13,X14,bool)) | p(s(bool,'const.bool.IN_2'(s(X15,happ(s(fun(fun(bool,X14),X15),X16),s(fun(bool,X14),X13))),s(fun(X15,bool),X17)))) | p(s(bool,'const.bool.IN_2'(s(fun(bool,X14),sK6(X17,'const.util_prob.FUNSET_2'(s(fun(bool,bool),X11),s(fun(X14,bool),X12)),X16,X15,fun(bool,X14))),s(fun(fun(bool,X14),bool),'const.util_prob.FUNSET_2'(s(fun(bool,bool),X11),s(fun(X14,bool),X12))))))) )),
  inference(resolution,[],[f6627,f6711])).
fof(f7688,plain,(
  ( ! [X6,X10,X8,X7,X5,X9] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X5)))) | s(bool,'T_0') = s(bool,sK6(X6,X5,happ(s(fun(X7,fun(bool,X8)),X9),s(X7,sK6('const.util_prob.FUNSET_2'(s(fun(bool,bool),X5),s(fun(X8,bool),X6)),X10,X9,fun(bool,X8),X7))),X8,bool)) | p(s(bool,'const.bool.IN_2'(s(fun(X7,fun(bool,X8)),X9),s(fun(fun(X7,fun(bool,X8)),bool),'const.util_prob.FUNSET_2'(s(fun(X7,bool),X10),s(fun(fun(bool,X8),bool),'const.util_prob.FUNSET_2'(s(fun(bool,bool),X5),s(fun(X8,bool),X6))))))))) )),
  inference(resolution,[],[f6627,f85])).
fof(f7687,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))) | s(bool,'T_0') = s(bool,sK6(X1,X0,X2,X3,bool)) | ~p(s(bool,'const.bool.IN_2'(s(bool,X4),s(fun(bool,bool),X0)))) | p(s(bool,'const.bool.IN_2'(s(X3,happ(s(fun(bool,X3),X2),s(bool,X4))),s(fun(X3,bool),X1))))) )),
  inference(resolution,[],[f6627,f83])).
fof(f7682,plain,(
  ( ! [X54,X56,X55] : (p(s(bool,'const.bool.IN_2'(s(X54,happ(s(fun(sK1,X54),X55),s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)))))),s(fun(X54,bool),X56)))) | p(s(bool,'const.bool.IN_2'(s(sK1,sK6(X56,sK3,X55,X54,sK1)),s(fun(sK1,bool),sK3))))) ) | $spl86),
  inference(global_subsumption,[],[f67,f5752,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6816,f104,f102,f101,f100,f7026,f98,f96,f7027,f7591,f7593,f7594,f7596,f7598,f7599,f7600,f7667,f7670,f7642])).
fof(f7642,plain,(
  ( ! [X54,X56,X55] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(X54,happ(s(fun(sK1,X54),X55),s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)))))),s(fun(X54,bool),X56)))) | p(s(bool,'const.bool.IN_2'(s(sK1,sK6(X56,sK3,X55,X54,sK1)),s(fun(sK1,bool),sK3))))) ) | $spl86),
  inference(superposition,[],[f6711,f5752])).
fof(f7670,plain,(
  ( ! [X14,X12,X10,X8,X15,X13,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(X8,sK6(X15,happ(s(fun(X9,fun(X8,bool)),'const.combin.K_1'(s(fun(X8,bool),X10))),s(X9,X11)),X14,X13,X8)),s(fun(X8,bool),X10)))) | ~p(s(bool,'const.bool.IN_2'(s(X8,X12),s(fun(X8,bool),X10)))) | p(s(bool,'const.bool.IN_2'(s(X13,happ(s(fun(X8,X13),X14),s(X8,X12))),s(fun(X13,bool),X15))))) )),
  inference(forward_demodulation,[],[f7669,f80])).
fof(f7669,plain,(
  ( ! [X14,X12,X10,X8,X15,X13,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(X8,sK6(X15,happ(s(fun(X9,fun(X8,bool)),'const.combin.K_1'(s(fun(X8,bool),X10))),s(X9,X11)),X14,X13,X8)),s(fun(X8,bool),'const.combin.K_2'(s(fun(X8,bool),X10),s(X9,X11)))))) | ~p(s(bool,'const.bool.IN_2'(s(X8,X12),s(fun(X8,bool),X10)))) | p(s(bool,'const.bool.IN_2'(s(X13,happ(s(fun(X8,X13),X14),s(X8,X12))),s(fun(X13,bool),X15))))) )),
  inference(forward_demodulation,[],[f7668,f81])).
fof(f7668,plain,(
  ( ! [X14,X12,X10,X8,X15,X13,X11,X9] : (~p(s(bool,'const.bool.IN_2'(s(X8,X12),s(fun(X8,bool),X10)))) | p(s(bool,'const.bool.IN_2'(s(X13,happ(s(fun(X8,X13),X14),s(X8,X12))),s(fun(X13,bool),X15)))) | p(s(bool,'const.bool.IN_2'(s(X8,sK6(X15,happ(s(fun(X9,fun(X8,bool)),'const.combin.K_1'(s(fun(X8,bool),X10))),s(X9,X11)),X14,X13,X8)),s(fun(X8,bool),happ(s(fun(X9,fun(X8,bool)),'const.combin.K_1'(s(fun(X8,bool),X10))),s(X9,X11))))))) )),
  inference(forward_demodulation,[],[f7627,f80])).
fof(f7627,plain,(
  ( ! [X14,X12,X10,X8,X15,X13,X11,X9] : (~p(s(bool,'const.bool.IN_2'(s(X8,X12),s(fun(X8,bool),'const.combin.K_2'(s(fun(X8,bool),X10),s(X9,X11)))))) | p(s(bool,'const.bool.IN_2'(s(X13,happ(s(fun(X8,X13),X14),s(X8,X12))),s(fun(X13,bool),X15)))) | p(s(bool,'const.bool.IN_2'(s(X8,sK6(X15,happ(s(fun(X9,fun(X8,bool)),'const.combin.K_1'(s(fun(X8,bool),X10))),s(X9,X11)),X14,X13,X8)),s(fun(X8,bool),happ(s(fun(X9,fun(X8,bool)),'const.combin.K_1'(s(fun(X8,bool),X10))),s(X9,X11))))))) )),
  inference(superposition,[],[f6711,f81])).
fof(f7667,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK6(X7,'const.combin.K_2'(s(fun(X0,bool),X1),s(X2,X3)),X6,X5,X0)),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X5,happ(s(fun(X0,X5),X6),s(X0,X4))),s(fun(X5,bool),X7))))) )),
  inference(forward_demodulation,[],[f7626,f80])).
fof(f7626,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X5,happ(s(fun(X0,X5),X6),s(X0,X4))),s(fun(X5,bool),X7)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK6(X7,'const.combin.K_2'(s(fun(X0,bool),X1),s(X2,X3)),X6,X5,X0)),s(fun(X0,bool),'const.combin.K_2'(s(fun(X0,bool),X1),s(X2,X3))))))) )),
  inference(superposition,[],[f6711,f80])).
fof(f7600,plain,(
  ( ! [X6,X8,X7,X5,X9] : (~p(s(bool,'const.bool.IN_2'(s(bool,X5),s(fun(bool,bool),X6)))) | p(s(bool,'const.bool.IN_2'(s(X7,happ(s(fun(bool,X7),X8),s(bool,'F_0'))),s(fun(X7,bool),X9)))) | p(s(bool,'const.bool.IN_2'(s(bool,sK6(X9,X6,X8,X7,bool)),s(fun(bool,bool),X6)))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f6711,f71])).
fof(f7599,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X2,happ(s(fun(bool,X2),X3),s(bool,X0))),s(fun(X2,bool),X4)))) | p(s(bool,'const.bool.IN_2'(s(bool,sK6(X4,X1,X3,X2,bool)),s(fun(bool,bool),X1)))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f6711,f71])).
fof(f7598,plain,(
  ( ! [X70,X68,X66,X64,X71,X69,X67,X65] : (p(s(bool,'const.bool.IN_2'(s(X64,happ(s(fun(X65,X64),X66),s(X65,sK7(X67,X68,X69,X65,X70)))),s(fun(X64,bool),X71)))) | p(s(bool,'const.bool.IN_2'(s(X65,sK6(X71,X68,X66,X64,X65)),s(fun(X65,bool),X68)))) | p(s(bool,'const.bool.IN_2'(s(fun(X65,X70),X69),s(fun(fun(X65,X70),bool),'const.util_prob.DFUNSET_2'(s(fun(X65,bool),X68),s(fun(X65,fun(X70,bool)),X67))))))) )),
  inference(resolution,[],[f6711,f87])).
fof(f7596,plain,(
  ( ! [X54,X52,X50,X48,X55,X53,X51,X49] : (p(s(bool,'const.bool.IN_2'(s(X48,happ(s(fun(X49,X48),X50),s(X49,sK6(X51,X52,X53,X54,X49)))),s(fun(X48,bool),X55)))) | p(s(bool,'const.bool.IN_2'(s(X49,sK6(X55,X52,X50,X48,X49)),s(fun(X49,bool),X52)))) | p(s(bool,'const.bool.IN_2'(s(fun(X49,X54),X53),s(fun(fun(X49,X54),bool),'const.util_prob.FUNSET_2'(s(fun(X49,bool),X52),s(fun(X54,bool),X51))))))) )),
  inference(resolution,[],[f6711,f84])).
fof(f7594,plain,(
  ( ! [X39,X37,X35,X33,X38,X36,X34,X32] : (p(s(bool,'const.bool.IN_2'(s(X32,happ(s(fun(fun(X33,X34),X32),X35),s(fun(X33,X34),'const.combin.K_1'(s(X34,X36))))),s(fun(X32,bool),X37)))) | p(s(bool,'const.bool.IN_2'(s(fun(X33,X34),sK6(X37,'const.util_prob.FUNSET_2'(s(fun(X33,bool),X38),s(fun(X34,bool),X39)),X35,X32,fun(X33,X34))),s(fun(fun(X33,X34),bool),'const.util_prob.FUNSET_2'(s(fun(X33,bool),X38),s(fun(X34,bool),X39)))))) | ~p(s(bool,'const.bool.IN_2'(s(X34,X36),s(fun(X34,bool),X39))))) )),
  inference(resolution,[],[f6711,f6706])).
fof(f7593,plain,(
  ( ! [X30,X28,X26,X24,X31,X29,X27,X25] : (p(s(bool,'const.bool.IN_2'(s(X24,happ(s(fun(fun(X25,X26),X24),X27),s(fun(X25,X26),X28))),s(fun(X24,bool),X29)))) | p(s(bool,'const.bool.IN_2'(s(fun(X25,X26),sK6(X29,'const.util_prob.DFUNSET_2'(s(fun(X25,bool),X30),s(fun(X25,fun(X26,bool)),X31)),X27,X24,fun(X25,X26))),s(fun(fun(X25,X26),bool),'const.util_prob.DFUNSET_2'(s(fun(X25,bool),X30),s(fun(X25,fun(X26,bool)),X31)))))) | p(s(bool,'const.bool.IN_2'(s(X25,sK7(X31,X30,X28,X25,X26)),s(fun(X25,bool),X30))))) )),
  inference(resolution,[],[f6711,f87])).
fof(f7591,plain,(
  ( ! [X14,X12,X10,X8,X15,X13,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(X8,happ(s(fun(fun(X9,X10),X8),X11),s(fun(X9,X10),X12))),s(fun(X8,bool),X13)))) | p(s(bool,'const.bool.IN_2'(s(fun(X9,X10),sK6(X13,'const.util_prob.FUNSET_2'(s(fun(X9,bool),X14),s(fun(X10,bool),X15)),X11,X8,fun(X9,X10))),s(fun(fun(X9,X10),bool),'const.util_prob.FUNSET_2'(s(fun(X9,bool),X14),s(fun(X10,bool),X15)))))) | p(s(bool,'const.bool.IN_2'(s(X9,sK6(X15,X14,X12,X10,X9)),s(fun(X9,bool),X14))))) )),
  inference(resolution,[],[f6711,f84])).
fof(f96,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X2) != s(bool,X3) | ~p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,X3)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f66])).
fof(f66,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f65])).
fof(f65,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f46])).
fof(f46,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f39])).
fof(f39,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X8,X14,X9,X15] : ((s(bool,X8) = s(bool,X14) & (p(s(bool,X14)) => s(bool,X9) = s(bool,X15))) => ((p(s(bool,X8)) => p(s(bool,X9))) <=> (p(s(bool,X14)) => p(s(bool,X15)))))),
  file('thm.util_prob.FUNSET_DFUNSET','thm.bool.IMP_CONG')).
fof(f100,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X2) != s(bool,X3) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f66])).
fof(f6784,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,'const.bool.IN_2'(s(fun(X5,bool),'const.combin.K_1'(s(bool,X4))),s(fun(fun(X5,bool),bool),'const.util_prob.FUNSET_2'(s(fun(X5,bool),X6),s(fun(bool,bool),X7)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X7)))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f6706,f71])).
fof(f6783,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(fun(X1,bool),'const.combin.K_1'(s(bool,'F_0'))),s(fun(fun(X1,bool),bool),'const.util_prob.FUNSET_2'(s(fun(X1,bool),X2),s(fun(bool,bool),X3)))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f6706,f71])).
fof(f6758,plain,(
  ( ! [X14,X21,X19,X17,X15,X20,X18,X16] : (~p(s(bool,'const.bool.IN_2'(s(X14,X15),s(fun(X14,bool),X16)))) | p(s(bool,'const.bool.IN_2'(s(X17,happ(s(fun(X14,X17),sK7(X18,'const.util_prob.DFUNSET_2'(s(fun(X14,bool),X16),s(fun(X14,fun(X17,bool)),X19)),X20,fun(X14,X17),X21)),s(X14,X15))),s(fun(X17,bool),happ(s(fun(X14,fun(X17,bool)),X19),s(X14,X15)))))) | p(s(bool,'const.bool.IN_2'(s(fun(fun(X14,X17),X21),X20),s(fun(fun(fun(X14,X17),X21),bool),'const.util_prob.DFUNSET_2'(s(fun(fun(X14,X17),bool),'const.util_prob.DFUNSET_2'(s(fun(X14,bool),X16),s(fun(X14,fun(X17,bool)),X19))),s(fun(fun(X14,X17),fun(X21,bool)),X18))))))) )),
  inference(resolution,[],[f86,f87])).
fof(f6757,plain,(
  ( ! [X6,X12,X10,X8,X7,X13,X11,X9] : (~p(s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),X8)))) | p(s(bool,'const.bool.IN_2'(s(X9,happ(s(fun(X6,X9),sK6(X10,'const.util_prob.DFUNSET_2'(s(fun(X6,bool),X8),s(fun(X6,fun(X9,bool)),X11)),X12,X13,fun(X6,X9))),s(X6,X7))),s(fun(X9,bool),happ(s(fun(X6,fun(X9,bool)),X11),s(X6,X7)))))) | p(s(bool,'const.bool.IN_2'(s(fun(fun(X6,X9),X13),X12),s(fun(fun(fun(X6,X9),X13),bool),'const.util_prob.FUNSET_2'(s(fun(fun(X6,X9),bool),'const.util_prob.DFUNSET_2'(s(fun(X6,bool),X8),s(fun(X6,fun(X9,bool)),X11))),s(fun(X13,bool),X10))))))) )),
  inference(resolution,[],[f86,f84])).
fof(f6756,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X3,happ(s(fun(X0,X3),X4),s(X0,X1))),s(fun(X3,bool),happ(s(fun(X0,fun(X3,bool)),X5),s(X0,X1)))))) | p(s(bool,'const.bool.IN_2'(s(X0,sK7(X5,X2,X4,X0,X3)),s(fun(X0,bool),X2))))) )),
  inference(resolution,[],[f86,f87])).
fof(f6755,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X4),s(X1,sK7('const.combin.K_1'(s(fun(X0,bool),X2)),X3,X4,X1,X0)))),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X4),s(fun(fun(X1,X0),bool),'const.util_prob.DFUNSET_2'(s(fun(X1,bool),X3),s(fun(X1,fun(X0,bool)),'const.combin.K_1'(s(fun(X0,bool),X2))))))))) )),
  inference(forward_demodulation,[],[f6744,f80])).
fof(f6744,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X4),s(X1,sK7('const.combin.K_1'(s(fun(X0,bool),X2)),X3,X4,X1,X0)))),s(fun(X0,bool),'const.combin.K_2'(s(fun(X0,bool),X2),s(X1,sK7('const.combin.K_1'(s(fun(X0,bool),X2)),X3,X4,X1,X0))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X4),s(fun(fun(X1,X0),bool),'const.util_prob.DFUNSET_2'(s(fun(X1,bool),X3),s(fun(X1,fun(X0,bool)),'const.combin.K_1'(s(fun(X0,bool),X2))))))))) )),
  inference(superposition,[],[f88,f81])).
fof(f6743,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X3,happ(s(fun(bool,X3),X2),s(bool,'F_0'))),s(fun(X3,bool),happ(s(fun(bool,fun(X3,bool)),X0),s(bool,'F_0')))))) | p(s(bool,'const.bool.IN_2'(s(fun(bool,X3),X2),s(fun(fun(bool,X3),bool),'const.util_prob.DFUNSET_2'(s(fun(bool,bool),X1),s(fun(bool,fun(X3,bool)),X0)))))) | s(bool,'T_0') = s(bool,sK7(X0,X1,X2,bool,X3))) )),
  inference(superposition,[],[f88,f71])).
fof(f6754,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(fun(X7,X8),X12),s(fun(fun(X7,X8),bool),'const.util_prob.DFUNSET_2'(s(fun(X7,bool),X13),s(fun(X7,fun(X8,bool)),X10)))))) | ~p(s(bool,'const.bool.IN_2'(s(X8,happ(s(fun(X7,X8),X12),s(X7,sK7(happ(s(fun(X9,fun(X7,fun(X8,bool))),'const.combin.K_1'(s(fun(X7,fun(X8,bool)),X10))),s(X9,X11)),X13,X12,X7,X8)))),s(fun(X8,bool),happ(s(fun(X7,fun(X8,bool)),X10),s(X7,sK7(happ(s(fun(X9,fun(X7,fun(X8,bool))),'const.combin.K_1'(s(fun(X7,fun(X8,bool)),X10))),s(X9,X11)),X13,X12,X7,X8)))))))) )),
  inference(forward_demodulation,[],[f6753,f80])).
fof(f6753,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(fun(X7,X8),X12),s(fun(fun(X7,X8),bool),'const.util_prob.DFUNSET_2'(s(fun(X7,bool),X13),s(fun(X7,fun(X8,bool)),'const.combin.K_2'(s(fun(X7,fun(X8,bool)),X10),s(X9,X11)))))))) | ~p(s(bool,'const.bool.IN_2'(s(X8,happ(s(fun(X7,X8),X12),s(X7,sK7(happ(s(fun(X9,fun(X7,fun(X8,bool))),'const.combin.K_1'(s(fun(X7,fun(X8,bool)),X10))),s(X9,X11)),X13,X12,X7,X8)))),s(fun(X8,bool),happ(s(fun(X7,fun(X8,bool)),X10),s(X7,sK7(happ(s(fun(X9,fun(X7,fun(X8,bool))),'const.combin.K_1'(s(fun(X7,fun(X8,bool)),X10))),s(X9,X11)),X13,X12,X7,X8)))))))) )),
  inference(forward_demodulation,[],[f6752,f81])).
fof(f6752,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (~p(s(bool,'const.bool.IN_2'(s(X8,happ(s(fun(X7,X8),X12),s(X7,sK7(happ(s(fun(X9,fun(X7,fun(X8,bool))),'const.combin.K_1'(s(fun(X7,fun(X8,bool)),X10))),s(X9,X11)),X13,X12,X7,X8)))),s(fun(X8,bool),happ(s(fun(X7,fun(X8,bool)),X10),s(X7,sK7(happ(s(fun(X9,fun(X7,fun(X8,bool))),'const.combin.K_1'(s(fun(X7,fun(X8,bool)),X10))),s(X9,X11)),X13,X12,X7,X8))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X7,X8),X12),s(fun(fun(X7,X8),bool),'const.util_prob.DFUNSET_2'(s(fun(X7,bool),X13),s(fun(X7,fun(X8,bool)),happ(s(fun(X9,fun(X7,fun(X8,bool))),'const.combin.K_1'(s(fun(X7,fun(X8,bool)),X10))),s(X9,X11))))))))) )),
  inference(forward_demodulation,[],[f6742,f80])).
fof(f6742,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (~p(s(bool,'const.bool.IN_2'(s(X8,happ(s(fun(X7,X8),X12),s(X7,sK7(happ(s(fun(X9,fun(X7,fun(X8,bool))),'const.combin.K_1'(s(fun(X7,fun(X8,bool)),X10))),s(X9,X11)),X13,X12,X7,X8)))),s(fun(X8,bool),happ(s(fun(X7,fun(X8,bool)),'const.combin.K_2'(s(fun(X7,fun(X8,bool)),X10),s(X9,X11))),s(X7,sK7(happ(s(fun(X9,fun(X7,fun(X8,bool))),'const.combin.K_1'(s(fun(X7,fun(X8,bool)),X10))),s(X9,X11)),X13,X12,X7,X8))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X7,X8),X12),s(fun(fun(X7,X8),bool),'const.util_prob.DFUNSET_2'(s(fun(X7,bool),X13),s(fun(X7,fun(X8,bool)),happ(s(fun(X9,fun(X7,fun(X8,bool))),'const.combin.K_1'(s(fun(X7,fun(X8,bool)),X10))),s(X9,X11))))))))) )),
  inference(superposition,[],[f88,f81])).
fof(f6751,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,X1),X5),s(fun(fun(X0,X1),bool),'const.util_prob.DFUNSET_2'(s(fun(X0,bool),X6),s(fun(X0,fun(X1,bool)),X2)))))) | ~p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X5),s(X0,sK7('const.combin.K_2'(s(fun(X0,fun(X1,bool)),X2),s(X3,X4)),X6,X5,X0,X1)))),s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,sK7('const.combin.K_2'(s(fun(X0,fun(X1,bool)),X2),s(X3,X4)),X6,X5,X0,X1)))))))) )),
  inference(forward_demodulation,[],[f6741,f80])).
fof(f6741,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X5),s(X0,sK7('const.combin.K_2'(s(fun(X0,fun(X1,bool)),X2),s(X3,X4)),X6,X5,X0,X1)))),s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,sK7('const.combin.K_2'(s(fun(X0,fun(X1,bool)),X2),s(X3,X4)),X6,X5,X0,X1))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,X1),X5),s(fun(fun(X0,X1),bool),'const.util_prob.DFUNSET_2'(s(fun(X0,bool),X6),s(fun(X0,fun(X1,bool)),'const.combin.K_2'(s(fun(X0,fun(X1,bool)),X2),s(X3,X4))))))))) )),
  inference(superposition,[],[f88,f80])).
fof(f6750,plain,(
  ( ! [X6,X4,X8,X7,X5] : (~p(s(bool,'const.bool.IN_2'(s(X4,X6),s(fun(X4,bool),happ(s(fun(X5,fun(X4,bool)),X7),s(X5,sK7(X7,X8,'const.combin.K_1'(s(X4,X6)),X5,X4))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X5,X4),'const.combin.K_1'(s(X4,X6))),s(fun(fun(X5,X4),bool),'const.util_prob.DFUNSET_2'(s(fun(X5,bool),X8),s(fun(X5,fun(X4,bool)),X7))))))) )),
  inference(forward_demodulation,[],[f6740,f80])).
fof(f6740,plain,(
  ( ! [X6,X4,X8,X7,X5] : (~p(s(bool,'const.bool.IN_2'(s(X4,'const.combin.K_2'(s(X4,X6),s(X5,sK7(X7,X8,'const.combin.K_1'(s(X4,X6)),X5,X4)))),s(fun(X4,bool),happ(s(fun(X5,fun(X4,bool)),X7),s(X5,sK7(X7,X8,'const.combin.K_1'(s(X4,X6)),X5,X4))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X5,X4),'const.combin.K_1'(s(X4,X6))),s(fun(fun(X5,X4),bool),'const.util_prob.DFUNSET_2'(s(fun(X5,bool),X8),s(fun(X5,fun(X4,bool)),X7))))))) )),
  inference(superposition,[],[f88,f81])).
fof(f6739,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),happ(s(fun(X0,fun(bool,bool)),X2),s(X0,sK7(X2,X3,X1,X0,bool))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X1),s(fun(fun(X0,bool),bool),'const.util_prob.DFUNSET_2'(s(fun(X0,bool),X3),s(fun(X0,fun(bool,bool)),X2)))))) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK7(X2,X3,X1,X0,bool))))) )),
  inference(superposition,[],[f88,f71])).
fof(f6749,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(fun(X7,X8),X10),s(fun(fun(X7,X8),bool),'const.util_prob.DFUNSET_2'(s(fun(X7,bool),X13),s(fun(X7,fun(X8,bool)),X12)))))) | ~p(s(bool,'const.bool.IN_2'(s(X8,happ(s(fun(X7,X8),X10),s(X7,sK7(X12,X13,happ(s(fun(X9,fun(X7,X8)),'const.combin.K_1'(s(fun(X7,X8),X10))),s(X9,X11)),X7,X8)))),s(fun(X8,bool),happ(s(fun(X7,fun(X8,bool)),X12),s(X7,sK7(X12,X13,happ(s(fun(X9,fun(X7,X8)),'const.combin.K_1'(s(fun(X7,X8),X10))),s(X9,X11)),X7,X8)))))))) )),
  inference(forward_demodulation,[],[f6748,f80])).
fof(f6748,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(fun(X7,X8),'const.combin.K_2'(s(fun(X7,X8),X10),s(X9,X11))),s(fun(fun(X7,X8),bool),'const.util_prob.DFUNSET_2'(s(fun(X7,bool),X13),s(fun(X7,fun(X8,bool)),X12)))))) | ~p(s(bool,'const.bool.IN_2'(s(X8,happ(s(fun(X7,X8),X10),s(X7,sK7(X12,X13,happ(s(fun(X9,fun(X7,X8)),'const.combin.K_1'(s(fun(X7,X8),X10))),s(X9,X11)),X7,X8)))),s(fun(X8,bool),happ(s(fun(X7,fun(X8,bool)),X12),s(X7,sK7(X12,X13,happ(s(fun(X9,fun(X7,X8)),'const.combin.K_1'(s(fun(X7,X8),X10))),s(X9,X11)),X7,X8)))))))) )),
  inference(forward_demodulation,[],[f6747,f81])).
fof(f6747,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (~p(s(bool,'const.bool.IN_2'(s(X8,happ(s(fun(X7,X8),X10),s(X7,sK7(X12,X13,happ(s(fun(X9,fun(X7,X8)),'const.combin.K_1'(s(fun(X7,X8),X10))),s(X9,X11)),X7,X8)))),s(fun(X8,bool),happ(s(fun(X7,fun(X8,bool)),X12),s(X7,sK7(X12,X13,happ(s(fun(X9,fun(X7,X8)),'const.combin.K_1'(s(fun(X7,X8),X10))),s(X9,X11)),X7,X8))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X7,X8),happ(s(fun(X9,fun(X7,X8)),'const.combin.K_1'(s(fun(X7,X8),X10))),s(X9,X11))),s(fun(fun(X7,X8),bool),'const.util_prob.DFUNSET_2'(s(fun(X7,bool),X13),s(fun(X7,fun(X8,bool)),X12))))))) )),
  inference(forward_demodulation,[],[f6738,f80])).
fof(f6738,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (~p(s(bool,'const.bool.IN_2'(s(X8,happ(s(fun(X7,X8),'const.combin.K_2'(s(fun(X7,X8),X10),s(X9,X11))),s(X7,sK7(X12,X13,happ(s(fun(X9,fun(X7,X8)),'const.combin.K_1'(s(fun(X7,X8),X10))),s(X9,X11)),X7,X8)))),s(fun(X8,bool),happ(s(fun(X7,fun(X8,bool)),X12),s(X7,sK7(X12,X13,happ(s(fun(X9,fun(X7,X8)),'const.combin.K_1'(s(fun(X7,X8),X10))),s(X9,X11)),X7,X8))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X7,X8),happ(s(fun(X9,fun(X7,X8)),'const.combin.K_1'(s(fun(X7,X8),X10))),s(X9,X11))),s(fun(fun(X7,X8),bool),'const.util_prob.DFUNSET_2'(s(fun(X7,bool),X13),s(fun(X7,fun(X8,bool)),X12))))))) )),
  inference(superposition,[],[f88,f81])).
fof(f6746,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,X1),X2),s(fun(fun(X0,X1),bool),'const.util_prob.DFUNSET_2'(s(fun(X0,bool),X6),s(fun(X0,fun(X1,bool)),X5)))))) | ~p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,sK7(X5,X6,'const.combin.K_2'(s(fun(X0,X1),X2),s(X3,X4)),X0,X1)))),s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X5),s(X0,sK7(X5,X6,'const.combin.K_2'(s(fun(X0,X1),X2),s(X3,X4)),X0,X1)))))))) )),
  inference(forward_demodulation,[],[f6737,f80])).
fof(f6737,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,sK7(X5,X6,'const.combin.K_2'(s(fun(X0,X1),X2),s(X3,X4)),X0,X1)))),s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X5),s(X0,sK7(X5,X6,'const.combin.K_2'(s(fun(X0,X1),X2),s(X3,X4)),X0,X1))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,X1),'const.combin.K_2'(s(fun(X0,X1),X2),s(X3,X4))),s(fun(fun(X0,X1),bool),'const.util_prob.DFUNSET_2'(s(fun(X0,bool),X6),s(fun(X0,fun(X1,bool)),X5))))))) )),
  inference(superposition,[],[f88,f80])).
fof(f6713,plain,(
  ( ! [X14,X21,X19,X17,X15,X20,X18,X16] : (~p(s(bool,'const.bool.IN_2'(s(X14,X15),s(fun(X14,bool),X16)))) | p(s(bool,'const.bool.IN_2'(s(X17,happ(s(fun(X14,X17),sK7(X18,'const.util_prob.FUNSET_2'(s(fun(X14,bool),X16),s(fun(X17,bool),X19)),X20,fun(X14,X17),X21)),s(X14,X15))),s(fun(X17,bool),X19)))) | p(s(bool,'const.bool.IN_2'(s(fun(fun(X14,X17),X21),X20),s(fun(fun(fun(X14,X17),X21),bool),'const.util_prob.DFUNSET_2'(s(fun(fun(X14,X17),bool),'const.util_prob.FUNSET_2'(s(fun(X14,bool),X16),s(fun(X17,bool),X19))),s(fun(fun(X14,X17),fun(X21,bool)),X18))))))) )),
  inference(resolution,[],[f83,f87])).
fof(f6712,plain,(
  ( ! [X6,X12,X10,X8,X7,X13,X11,X9] : (~p(s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),X8)))) | p(s(bool,'const.bool.IN_2'(s(X9,happ(s(fun(X6,X9),sK6(X10,'const.util_prob.FUNSET_2'(s(fun(X6,bool),X8),s(fun(X9,bool),X11)),X12,X13,fun(X6,X9))),s(X6,X7))),s(fun(X9,bool),X11)))) | p(s(bool,'const.bool.IN_2'(s(fun(fun(X6,X9),X13),X12),s(fun(fun(fun(X6,X9),X13),bool),'const.util_prob.FUNSET_2'(s(fun(fun(X6,X9),bool),'const.util_prob.FUNSET_2'(s(fun(X6,bool),X8),s(fun(X9,bool),X11))),s(fun(X13,bool),X10))))))) )),
  inference(resolution,[],[f83,f84])).
fof(f6711,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X3,happ(s(fun(X0,X3),X4),s(X0,X1))),s(fun(X3,bool),X5)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK6(X5,X2,X4,X3,X0)),s(fun(X0,bool),X2))))) )),
  inference(resolution,[],[f83,f84])).
fof(f6710,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(fun(X11,X7),X12),s(fun(fun(X11,X7),bool),'const.util_prob.FUNSET_2'(s(fun(X11,bool),X13),s(fun(X7,bool),X9)))))) | ~p(s(bool,'const.bool.IN_2'(s(X7,happ(s(fun(X11,X7),X12),s(X11,sK6(happ(s(fun(X8,fun(X7,bool)),'const.combin.K_1'(s(fun(X7,bool),X9))),s(X8,X10)),X13,X12,X7,X11)))),s(fun(X7,bool),X9))))) )),
  inference(forward_demodulation,[],[f6709,f80])).
fof(f6709,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(fun(X11,X7),X12),s(fun(fun(X11,X7),bool),'const.util_prob.FUNSET_2'(s(fun(X11,bool),X13),s(fun(X7,bool),'const.combin.K_2'(s(fun(X7,bool),X9),s(X8,X10)))))))) | ~p(s(bool,'const.bool.IN_2'(s(X7,happ(s(fun(X11,X7),X12),s(X11,sK6(happ(s(fun(X8,fun(X7,bool)),'const.combin.K_1'(s(fun(X7,bool),X9))),s(X8,X10)),X13,X12,X7,X11)))),s(fun(X7,bool),X9))))) )),
  inference(forward_demodulation,[],[f6708,f81])).
fof(f6708,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (~p(s(bool,'const.bool.IN_2'(s(X7,happ(s(fun(X11,X7),X12),s(X11,sK6(happ(s(fun(X8,fun(X7,bool)),'const.combin.K_1'(s(fun(X7,bool),X9))),s(X8,X10)),X13,X12,X7,X11)))),s(fun(X7,bool),X9)))) | p(s(bool,'const.bool.IN_2'(s(fun(X11,X7),X12),s(fun(fun(X11,X7),bool),'const.util_prob.FUNSET_2'(s(fun(X11,bool),X13),s(fun(X7,bool),happ(s(fun(X8,fun(X7,bool)),'const.combin.K_1'(s(fun(X7,bool),X9))),s(X8,X10))))))))) )),
  inference(forward_demodulation,[],[f6694,f80])).
fof(f6694,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (~p(s(bool,'const.bool.IN_2'(s(X7,happ(s(fun(X11,X7),X12),s(X11,sK6(happ(s(fun(X8,fun(X7,bool)),'const.combin.K_1'(s(fun(X7,bool),X9))),s(X8,X10)),X13,X12,X7,X11)))),s(fun(X7,bool),'const.combin.K_2'(s(fun(X7,bool),X9),s(X8,X10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X11,X7),X12),s(fun(fun(X11,X7),bool),'const.util_prob.FUNSET_2'(s(fun(X11,bool),X13),s(fun(X7,bool),happ(s(fun(X8,fun(X7,bool)),'const.combin.K_1'(s(fun(X7,bool),X9))),s(X8,X10))))))))) )),
  inference(superposition,[],[f85,f81])).
fof(f6707,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X4,X0),X5),s(X4,sK6('const.combin.K_2'(s(fun(X0,bool),X1),s(X2,X3)),X6,X5,X0,X4)))),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(fun(X4,X0),X5),s(fun(fun(X4,X0),bool),'const.util_prob.FUNSET_2'(s(fun(X4,bool),X6),s(fun(X0,bool),X1))))))) )),
  inference(forward_demodulation,[],[f6693,f80])).
fof(f6693,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X4,X0),X5),s(X4,sK6('const.combin.K_2'(s(fun(X0,bool),X1),s(X2,X3)),X6,X5,X0,X4)))),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(fun(X4,X0),X5),s(fun(fun(X4,X0),bool),'const.util_prob.FUNSET_2'(s(fun(X4,bool),X6),s(fun(X0,bool),'const.combin.K_2'(s(fun(X0,bool),X1),s(X2,X3))))))))) )),
  inference(superposition,[],[f85,f80])).
fof(f6706,plain,(
  ( ! [X6,X4,X8,X7,X5] : (p(s(bool,'const.bool.IN_2'(s(fun(X5,X4),'const.combin.K_1'(s(X4,X6))),s(fun(fun(X5,X4),bool),'const.util_prob.FUNSET_2'(s(fun(X5,bool),X8),s(fun(X4,bool),X7)))))) | ~p(s(bool,'const.bool.IN_2'(s(X4,X6),s(fun(X4,bool),X7))))) )),
  inference(forward_demodulation,[],[f6692,f80])).
fof(f6692,plain,(
  ( ! [X6,X4,X8,X7,X5] : (~p(s(bool,'const.bool.IN_2'(s(X4,'const.combin.K_2'(s(X4,X6),s(X5,sK6(X7,X8,'const.combin.K_1'(s(X4,X6)),X4,X5)))),s(fun(X4,bool),X7)))) | p(s(bool,'const.bool.IN_2'(s(fun(X5,X4),'const.combin.K_1'(s(X4,X6))),s(fun(fun(X5,X4),bool),'const.util_prob.FUNSET_2'(s(fun(X5,bool),X8),s(fun(X4,bool),X7))))))) )),
  inference(superposition,[],[f85,f81])).
fof(f6691,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X1),s(fun(fun(X0,bool),bool),'const.util_prob.FUNSET_2'(s(fun(X0,bool),X3),s(fun(bool,bool),X2)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X2)))) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK6(X2,X3,X1,bool,X0))))) )),
  inference(superposition,[],[f85,f71])).
fof(f6690,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(bool,X3),X2),s(fun(fun(bool,X3),bool),'const.util_prob.FUNSET_2'(s(fun(bool,bool),X1),s(fun(X3,bool),X0)))))) | ~p(s(bool,'const.bool.IN_2'(s(X3,happ(s(fun(bool,X3),X2),s(bool,'F_0'))),s(fun(X3,bool),X0)))) | s(bool,'T_0') = s(bool,sK6(X0,X1,X2,X3,bool))) )),
  inference(superposition,[],[f85,f71])).
fof(f6705,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(fun(X7,X8),X10),s(fun(fun(X7,X8),bool),'const.util_prob.FUNSET_2'(s(fun(X7,bool),X13),s(fun(X8,bool),X12)))))) | ~p(s(bool,'const.bool.IN_2'(s(X8,happ(s(fun(X7,X8),X10),s(X7,sK6(X12,X13,happ(s(fun(X9,fun(X7,X8)),'const.combin.K_1'(s(fun(X7,X8),X10))),s(X9,X11)),X8,X7)))),s(fun(X8,bool),X12))))) )),
  inference(forward_demodulation,[],[f6704,f80])).
fof(f6704,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(fun(X7,X8),'const.combin.K_2'(s(fun(X7,X8),X10),s(X9,X11))),s(fun(fun(X7,X8),bool),'const.util_prob.FUNSET_2'(s(fun(X7,bool),X13),s(fun(X8,bool),X12)))))) | ~p(s(bool,'const.bool.IN_2'(s(X8,happ(s(fun(X7,X8),X10),s(X7,sK6(X12,X13,happ(s(fun(X9,fun(X7,X8)),'const.combin.K_1'(s(fun(X7,X8),X10))),s(X9,X11)),X8,X7)))),s(fun(X8,bool),X12))))) )),
  inference(forward_demodulation,[],[f6703,f81])).
fof(f6703,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (~p(s(bool,'const.bool.IN_2'(s(X8,happ(s(fun(X7,X8),X10),s(X7,sK6(X12,X13,happ(s(fun(X9,fun(X7,X8)),'const.combin.K_1'(s(fun(X7,X8),X10))),s(X9,X11)),X8,X7)))),s(fun(X8,bool),X12)))) | p(s(bool,'const.bool.IN_2'(s(fun(X7,X8),happ(s(fun(X9,fun(X7,X8)),'const.combin.K_1'(s(fun(X7,X8),X10))),s(X9,X11))),s(fun(fun(X7,X8),bool),'const.util_prob.FUNSET_2'(s(fun(X7,bool),X13),s(fun(X8,bool),X12))))))) )),
  inference(forward_demodulation,[],[f6689,f80])).
fof(f6689,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (~p(s(bool,'const.bool.IN_2'(s(X8,happ(s(fun(X7,X8),'const.combin.K_2'(s(fun(X7,X8),X10),s(X9,X11))),s(X7,sK6(X12,X13,happ(s(fun(X9,fun(X7,X8)),'const.combin.K_1'(s(fun(X7,X8),X10))),s(X9,X11)),X8,X7)))),s(fun(X8,bool),X12)))) | p(s(bool,'const.bool.IN_2'(s(fun(X7,X8),happ(s(fun(X9,fun(X7,X8)),'const.combin.K_1'(s(fun(X7,X8),X10))),s(X9,X11))),s(fun(fun(X7,X8),bool),'const.util_prob.FUNSET_2'(s(fun(X7,bool),X13),s(fun(X8,bool),X12))))))) )),
  inference(superposition,[],[f85,f81])).
fof(f6702,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,sK6(X5,X6,'const.combin.K_2'(s(fun(X0,X1),X2),s(X3,X4)),X1,X0)))),s(fun(X1,bool),X5)))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,X1),X2),s(fun(fun(X0,X1),bool),'const.util_prob.FUNSET_2'(s(fun(X0,bool),X6),s(fun(X1,bool),X5))))))) )),
  inference(forward_demodulation,[],[f6688,f80])).
fof(f6688,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,sK6(X5,X6,'const.combin.K_2'(s(fun(X0,X1),X2),s(X3,X4)),X1,X0)))),s(fun(X1,bool),X5)))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,X1),'const.combin.K_2'(s(fun(X0,X1),X2),s(X3,X4))),s(fun(fun(X0,X1),bool),'const.util_prob.FUNSET_2'(s(fun(X0,bool),X6),s(fun(X1,bool),X5))))))) )),
  inference(superposition,[],[f85,f80])).
fof(f6687,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(fun(X7,fun(X8,X9)),X10),s(fun(fun(X7,fun(X8,X9)),bool),'const.util_prob.FUNSET_2'(s(fun(X7,bool),X11),s(fun(fun(X8,X9),bool),'const.util_prob.DFUNSET_2'(s(fun(X8,bool),X12),s(fun(X8,fun(X9,bool)),X13)))))))) | p(s(bool,'const.bool.IN_2'(s(X8,sK7(X13,X12,happ(s(fun(X7,fun(X8,X9)),X10),s(X7,sK6('const.util_prob.DFUNSET_2'(s(fun(X8,bool),X12),s(fun(X8,fun(X9,bool)),X13)),X11,X10,fun(X8,X9),X7))),X8,X9)),s(fun(X8,bool),X12))))) )),
  inference(resolution,[],[f85,f87])).
fof(f6686,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,fun(X1,X2)),X3),s(fun(fun(X0,fun(X1,X2)),bool),'const.util_prob.FUNSET_2'(s(fun(X0,bool),X4),s(fun(fun(X1,X2),bool),'const.util_prob.FUNSET_2'(s(fun(X1,bool),X5),s(fun(X2,bool),X6)))))))) | p(s(bool,'const.bool.IN_2'(s(X1,sK6(X6,X5,happ(s(fun(X0,fun(X1,X2)),X3),s(X0,sK6('const.util_prob.FUNSET_2'(s(fun(X1,bool),X5),s(fun(X2,bool),X6)),X4,X3,fun(X1,X2),X0))),X2,X1)),s(fun(X1,bool),X5))))) )),
  inference(resolution,[],[f85,f84])).
fof(f6662,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(bool,X3),X2),s(fun(fun(bool,X3),bool),'const.util_prob.DFUNSET_2'(s(fun(bool,bool),X1),s(fun(bool,fun(X3,bool)),X0)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))) | s(bool,'T_0') = s(bool,sK7(X0,X1,X2,bool,X3))) )),
  inference(superposition,[],[f87,f71])).
fof(f6677,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(fun(X7,X8),X12),s(fun(fun(X7,X8),bool),'const.util_prob.DFUNSET_2'(s(fun(X7,bool),X13),s(fun(X7,fun(X8,bool)),X10)))))) | p(s(bool,'const.bool.IN_2'(s(X7,sK7(happ(s(fun(X9,fun(X7,fun(X8,bool))),'const.combin.K_1'(s(fun(X7,fun(X8,bool)),X10))),s(X9,X11)),X13,X12,X7,X8)),s(fun(X7,bool),X13))))) )),
  inference(forward_demodulation,[],[f6656,f80])).
fof(f6656,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(fun(X7,X8),X12),s(fun(fun(X7,X8),bool),'const.util_prob.DFUNSET_2'(s(fun(X7,bool),X13),s(fun(X7,fun(X8,bool)),'const.combin.K_2'(s(fun(X7,fun(X8,bool)),X10),s(X9,X11)))))))) | p(s(bool,'const.bool.IN_2'(s(X7,sK7(happ(s(fun(X9,fun(X7,fun(X8,bool))),'const.combin.K_1'(s(fun(X7,fun(X8,bool)),X10))),s(X9,X11)),X13,X12,X7,X8)),s(fun(X7,bool),X13))))) )),
  inference(superposition,[],[f87,f81])).
fof(f6655,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,X1),X5),s(fun(fun(X0,X1),bool),'const.util_prob.DFUNSET_2'(s(fun(X0,bool),X6),s(fun(X0,fun(X1,bool)),X2)))))) | p(s(bool,'const.bool.IN_2'(s(X0,sK7('const.combin.K_2'(s(fun(X0,fun(X1,bool)),X2),s(X3,X4)),X6,X5,X0,X1)),s(fun(X0,bool),X6))))) )),
  inference(superposition,[],[f87,f80])).
fof(f6676,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(X7,sK7(X13,happ(s(fun(X8,fun(X7,bool)),'const.combin.K_1'(s(fun(X7,bool),X9))),s(X8,X10)),X12,X7,X11)),s(fun(X7,bool),X9)))) | p(s(bool,'const.bool.IN_2'(s(fun(X7,X11),X12),s(fun(fun(X7,X11),bool),'const.util_prob.DFUNSET_2'(s(fun(X7,bool),X9),s(fun(X7,fun(X11,bool)),X13))))))) )),
  inference(forward_demodulation,[],[f6675,f80])).
fof(f6675,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(X7,sK7(X13,happ(s(fun(X8,fun(X7,bool)),'const.combin.K_1'(s(fun(X7,bool),X9))),s(X8,X10)),X12,X7,X11)),s(fun(X7,bool),'const.combin.K_2'(s(fun(X7,bool),X9),s(X8,X10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X7,X11),X12),s(fun(fun(X7,X11),bool),'const.util_prob.DFUNSET_2'(s(fun(X7,bool),X9),s(fun(X7,fun(X11,bool)),X13))))))) )),
  inference(forward_demodulation,[],[f6674,f81])).
fof(f6674,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(fun(X7,X11),X12),s(fun(fun(X7,X11),bool),'const.util_prob.DFUNSET_2'(s(fun(X7,bool),X9),s(fun(X7,fun(X11,bool)),X13)))))) | p(s(bool,'const.bool.IN_2'(s(X7,sK7(X13,happ(s(fun(X8,fun(X7,bool)),'const.combin.K_1'(s(fun(X7,bool),X9))),s(X8,X10)),X12,X7,X11)),s(fun(X7,bool),happ(s(fun(X8,fun(X7,bool)),'const.combin.K_1'(s(fun(X7,bool),X9))),s(X8,X10))))))) )),
  inference(forward_demodulation,[],[f6654,f80])).
fof(f6654,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(fun(X7,X11),X12),s(fun(fun(X7,X11),bool),'const.util_prob.DFUNSET_2'(s(fun(X7,bool),'const.combin.K_2'(s(fun(X7,bool),X9),s(X8,X10))),s(fun(X7,fun(X11,bool)),X13)))))) | p(s(bool,'const.bool.IN_2'(s(X7,sK7(X13,happ(s(fun(X8,fun(X7,bool)),'const.combin.K_1'(s(fun(X7,bool),X9))),s(X8,X10)),X12,X7,X11)),s(fun(X7,bool),happ(s(fun(X8,fun(X7,bool)),'const.combin.K_1'(s(fun(X7,bool),X9))),s(X8,X10))))))) )),
  inference(superposition,[],[f87,f81])).
fof(f6673,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK7(X6,'const.combin.K_2'(s(fun(X0,bool),X1),s(X2,X3)),X5,X0,X4)),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,X4),X5),s(fun(fun(X0,X4),bool),'const.util_prob.DFUNSET_2'(s(fun(X0,bool),X1),s(fun(X0,fun(X4,bool)),X6))))))) )),
  inference(forward_demodulation,[],[f6653,f80])).
fof(f6653,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,X4),X5),s(fun(fun(X0,X4),bool),'const.util_prob.DFUNSET_2'(s(fun(X0,bool),X1),s(fun(X0,fun(X4,bool)),X6)))))) | p(s(bool,'const.bool.IN_2'(s(X0,sK7(X6,'const.combin.K_2'(s(fun(X0,bool),X1),s(X2,X3)),X5,X0,X4)),s(fun(X0,bool),'const.combin.K_2'(s(fun(X0,bool),X1),s(X2,X3))))))) )),
  inference(superposition,[],[f87,f80])).
fof(f6672,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(X7,sK7(X13,X12,happ(s(fun(X9,fun(X7,X8)),'const.combin.K_1'(s(fun(X7,X8),X10))),s(X9,X11)),X7,X8)),s(fun(X7,bool),X12)))) | p(s(bool,'const.bool.IN_2'(s(fun(X7,X8),X10),s(fun(fun(X7,X8),bool),'const.util_prob.DFUNSET_2'(s(fun(X7,bool),X12),s(fun(X7,fun(X8,bool)),X13))))))) )),
  inference(forward_demodulation,[],[f6652,f80])).
fof(f6652,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(fun(X7,X8),'const.combin.K_2'(s(fun(X7,X8),X10),s(X9,X11))),s(fun(fun(X7,X8),bool),'const.util_prob.DFUNSET_2'(s(fun(X7,bool),X12),s(fun(X7,fun(X8,bool)),X13)))))) | p(s(bool,'const.bool.IN_2'(s(X7,sK7(X13,X12,happ(s(fun(X9,fun(X7,X8)),'const.combin.K_1'(s(fun(X7,X8),X10))),s(X9,X11)),X7,X8)),s(fun(X7,bool),X12))))) )),
  inference(superposition,[],[f87,f81])).
fof(f6651,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,X1),X2),s(fun(fun(X0,X1),bool),'const.util_prob.DFUNSET_2'(s(fun(X0,bool),X5),s(fun(X0,fun(X1,bool)),X6)))))) | p(s(bool,'const.bool.IN_2'(s(X0,sK7(X6,X5,'const.combin.K_2'(s(fun(X0,X1),X2),s(X3,X4)),X0,X1)),s(fun(X0,bool),X5))))) )),
  inference(superposition,[],[f87,f80])).
fof(f6627,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(bool,X3),X2),s(fun(fun(bool,X3),bool),'const.util_prob.FUNSET_2'(s(fun(bool,bool),X1),s(fun(X3,bool),X0)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))) | s(bool,'T_0') = s(bool,sK6(X0,X1,X2,X3,bool))) )),
  inference(superposition,[],[f84,f71])).
fof(f6641,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(X11,sK6(happ(s(fun(X8,fun(X7,bool)),'const.combin.K_1'(s(fun(X7,bool),X9))),s(X8,X10)),X13,X12,X7,X11)),s(fun(X11,bool),X13)))) | p(s(bool,'const.bool.IN_2'(s(fun(X11,X7),X12),s(fun(fun(X11,X7),bool),'const.util_prob.FUNSET_2'(s(fun(X11,bool),X13),s(fun(X7,bool),X9))))))) )),
  inference(forward_demodulation,[],[f6621,f80])).
fof(f6621,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(fun(X11,X7),X12),s(fun(fun(X11,X7),bool),'const.util_prob.FUNSET_2'(s(fun(X11,bool),X13),s(fun(X7,bool),'const.combin.K_2'(s(fun(X7,bool),X9),s(X8,X10)))))))) | p(s(bool,'const.bool.IN_2'(s(X11,sK6(happ(s(fun(X8,fun(X7,bool)),'const.combin.K_1'(s(fun(X7,bool),X9))),s(X8,X10)),X13,X12,X7,X11)),s(fun(X11,bool),X13))))) )),
  inference(superposition,[],[f84,f81])).
fof(f6620,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(X4,X0),X5),s(fun(fun(X4,X0),bool),'const.util_prob.FUNSET_2'(s(fun(X4,bool),X6),s(fun(X0,bool),X1)))))) | p(s(bool,'const.bool.IN_2'(s(X4,sK6('const.combin.K_2'(s(fun(X0,bool),X1),s(X2,X3)),X6,X5,X0,X4)),s(fun(X4,bool),X6))))) )),
  inference(superposition,[],[f84,f80])).
fof(f6640,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(X7,sK6(X13,happ(s(fun(X8,fun(X7,bool)),'const.combin.K_1'(s(fun(X7,bool),X9))),s(X8,X10)),X12,X11,X7)),s(fun(X7,bool),X9)))) | p(s(bool,'const.bool.IN_2'(s(fun(X7,X11),X12),s(fun(fun(X7,X11),bool),'const.util_prob.FUNSET_2'(s(fun(X7,bool),X9),s(fun(X11,bool),X13))))))) )),
  inference(forward_demodulation,[],[f6639,f80])).
fof(f6639,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(X7,sK6(X13,happ(s(fun(X8,fun(X7,bool)),'const.combin.K_1'(s(fun(X7,bool),X9))),s(X8,X10)),X12,X11,X7)),s(fun(X7,bool),'const.combin.K_2'(s(fun(X7,bool),X9),s(X8,X10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X7,X11),X12),s(fun(fun(X7,X11),bool),'const.util_prob.FUNSET_2'(s(fun(X7,bool),X9),s(fun(X11,bool),X13))))))) )),
  inference(forward_demodulation,[],[f6638,f81])).
fof(f6638,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(fun(X7,X11),X12),s(fun(fun(X7,X11),bool),'const.util_prob.FUNSET_2'(s(fun(X7,bool),X9),s(fun(X11,bool),X13)))))) | p(s(bool,'const.bool.IN_2'(s(X7,sK6(X13,happ(s(fun(X8,fun(X7,bool)),'const.combin.K_1'(s(fun(X7,bool),X9))),s(X8,X10)),X12,X11,X7)),s(fun(X7,bool),happ(s(fun(X8,fun(X7,bool)),'const.combin.K_1'(s(fun(X7,bool),X9))),s(X8,X10))))))) )),
  inference(forward_demodulation,[],[f6619,f80])).
fof(f6619,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(fun(X7,X11),X12),s(fun(fun(X7,X11),bool),'const.util_prob.FUNSET_2'(s(fun(X7,bool),'const.combin.K_2'(s(fun(X7,bool),X9),s(X8,X10))),s(fun(X11,bool),X13)))))) | p(s(bool,'const.bool.IN_2'(s(X7,sK6(X13,happ(s(fun(X8,fun(X7,bool)),'const.combin.K_1'(s(fun(X7,bool),X9))),s(X8,X10)),X12,X11,X7)),s(fun(X7,bool),happ(s(fun(X8,fun(X7,bool)),'const.combin.K_1'(s(fun(X7,bool),X9))),s(X8,X10))))))) )),
  inference(superposition,[],[f84,f81])).
fof(f6637,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK6(X6,'const.combin.K_2'(s(fun(X0,bool),X1),s(X2,X3)),X5,X4,X0)),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,X4),X5),s(fun(fun(X0,X4),bool),'const.util_prob.FUNSET_2'(s(fun(X0,bool),X1),s(fun(X4,bool),X6))))))) )),
  inference(forward_demodulation,[],[f6618,f80])).
fof(f6618,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,X4),X5),s(fun(fun(X0,X4),bool),'const.util_prob.FUNSET_2'(s(fun(X0,bool),X1),s(fun(X4,bool),X6)))))) | p(s(bool,'const.bool.IN_2'(s(X0,sK6(X6,'const.combin.K_2'(s(fun(X0,bool),X1),s(X2,X3)),X5,X4,X0)),s(fun(X0,bool),'const.combin.K_2'(s(fun(X0,bool),X1),s(X2,X3))))))) )),
  inference(superposition,[],[f84,f80])).
fof(f6636,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(X7,sK6(X13,X12,happ(s(fun(X9,fun(X7,X8)),'const.combin.K_1'(s(fun(X7,X8),X10))),s(X9,X11)),X8,X7)),s(fun(X7,bool),X12)))) | p(s(bool,'const.bool.IN_2'(s(fun(X7,X8),X10),s(fun(fun(X7,X8),bool),'const.util_prob.FUNSET_2'(s(fun(X7,bool),X12),s(fun(X8,bool),X13))))))) )),
  inference(forward_demodulation,[],[f6617,f80])).
fof(f6617,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(fun(X7,X8),'const.combin.K_2'(s(fun(X7,X8),X10),s(X9,X11))),s(fun(fun(X7,X8),bool),'const.util_prob.FUNSET_2'(s(fun(X7,bool),X12),s(fun(X8,bool),X13)))))) | p(s(bool,'const.bool.IN_2'(s(X7,sK6(X13,X12,happ(s(fun(X9,fun(X7,X8)),'const.combin.K_1'(s(fun(X7,X8),X10))),s(X9,X11)),X8,X7)),s(fun(X7,bool),X12))))) )),
  inference(superposition,[],[f84,f81])).
fof(f6616,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,X1),X2),s(fun(fun(X0,X1),bool),'const.util_prob.FUNSET_2'(s(fun(X0,bool),X5),s(fun(X1,bool),X6)))))) | p(s(bool,'const.bool.IN_2'(s(X0,sK6(X6,X5,'const.combin.K_2'(s(fun(X0,X1),X2),s(X3,X4)),X1,X0)),s(fun(X0,bool),X5))))) )),
  inference(superposition,[],[f84,f80])).
fof(f88,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,sK7(X4,X3,X2,X1,X0)))),s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X4),s(X1,sK7(X4,X3,X2,X1,X0))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X2),s(fun(fun(X1,X0),bool),'const.util_prob.DFUNSET_2'(s(fun(X1,bool),X3),s(fun(X1,fun(X0,bool)),X4))))))) )),
  inference(cnf_transformation,[],[f62])).
fof(f62,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X2),s(fun(fun(X1,X0),bool),'const.util_prob.DFUNSET_2'(s(fun(X1,bool),X3),s(fun(X1,fun(X0,bool)),X4)))))) | ! [X6] : (~p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))),s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X4),s(X1,X6)))))))) & ((p(s(bool,'const.bool.IN_2'(s(X1,sK7(X4,X3,X2,X1,X0)),s(fun(X1,bool),X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,sK7(X4,X3,X2,X1,X0)))),s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X4),s(X1,sK7(X4,X3,X2,X1,X0)))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X2),s(fun(fun(X1,X0),bool),'const.util_prob.DFUNSET_2'(s(fun(X1,bool),X3),s(fun(X1,fun(X0,bool)),X4))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK7])],[f61])).
fof(f61,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X2),s(fun(fun(X1,X0),bool),'const.util_prob.DFUNSET_2'(s(fun(X1,bool),X3),s(fun(X1,fun(X0,bool)),X4)))))) | ! [X6] : (~p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))),s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X4),s(X1,X6)))))))) & (? [X5] : (p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))),s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X4),s(X1,X5))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X2),s(fun(fun(X1,X0),bool),'const.util_prob.DFUNSET_2'(s(fun(X1,bool),X3),s(fun(X1,fun(X0,bool)),X4))))))))),
  inference(rectify,[],[f60])).
fof(f60,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X2),s(fun(fun(X1,X0),bool),'const.util_prob.DFUNSET_2'(s(fun(X1,bool),X3),s(fun(X1,fun(X0,bool)),X4)))))) | ! [X5] : (~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))),s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X4),s(X1,X5)))))))) & (? [X5] : (p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))),s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X4),s(X1,X5))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X2),s(fun(fun(X1,X0),bool),'const.util_prob.DFUNSET_2'(s(fun(X1,bool),X3),s(fun(X1,fun(X0,bool)),X4))))))))),
  inference(nnf_transformation,[],[f43])).
fof(f43,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X2),s(fun(fun(X1,X0),bool),'const.util_prob.DFUNSET_2'(s(fun(X1,bool),X3),s(fun(X1,fun(X0,bool)),X4)))))) <=> ! [X5] : (~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))),s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X4),s(X1,X5))))))))),
  inference(ennf_transformation,[],[f37])).
fof(f37,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X2),s(fun(fun(X1,X0),bool),'const.util_prob.DFUNSET_2'(s(fun(X1,bool),X3),s(fun(X1,fun(X0,bool)),X4)))))) <=> ! [X5] : (p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) => p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))),s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X4),s(X1,X5))))))))),
  inference(flattening,[],[f36])).
fof(f36,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X2),s(fun(fun(X1,X0),bool),'const.util_prob.DFUNSET_2'(s(fun(X1,bool),X3),s(fun(X1,fun(X0,bool)),X4)))))) <=> ! [X5] : (p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) => p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))),s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X4),s(X1,X5))))))))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X7,X6] : ! [X17,X18,X19] : (p(s(bool,'const.bool.IN_2'(s(fun(X6,X7),X17),s(fun(fun(X6,X7),bool),'const.util_prob.DFUNSET_2'(s(fun(X6,bool),X18),s(fun(X6,fun(X7,bool)),X19)))))) <=> ! [X8] : (p(s(bool,'const.bool.IN_2'(s(X6,X8),s(fun(X6,bool),X18)))) => p(s(bool,'const.bool.IN_2'(s(X7,happ(s(fun(X6,X7),X17),s(X6,X8))),s(fun(X7,bool),happ(s(fun(X6,fun(X7,bool)),X19),s(X6,X8))))))))),
  file('thm.util_prob.FUNSET_DFUNSET','thm.util_prob.IN_DFUNSET')).
fof(f85,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,sK6(X4,X3,X2,X1,X0)))),s(fun(X1,bool),X4)))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,X1),X2),s(fun(fun(X0,X1),bool),'const.util_prob.FUNSET_2'(s(fun(X0,bool),X3),s(fun(X1,bool),X4))))))) )),
  inference(cnf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.bool.IN_2'(s(fun(X0,X1),X2),s(fun(fun(X0,X1),bool),'const.util_prob.FUNSET_2'(s(fun(X0,bool),X3),s(fun(X1,bool),X4)))))) | ! [X6] : (~p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X6))),s(fun(X1,bool),X4)))))) & ((p(s(bool,'const.bool.IN_2'(s(X0,sK6(X4,X3,X2,X1,X0)),s(fun(X0,bool),X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,sK6(X4,X3,X2,X1,X0)))),s(fun(X1,bool),X4))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,X1),X2),s(fun(fun(X0,X1),bool),'const.util_prob.FUNSET_2'(s(fun(X0,bool),X3),s(fun(X1,bool),X4))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK6])],[f58])).
fof(f58,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.bool.IN_2'(s(fun(X0,X1),X2),s(fun(fun(X0,X1),bool),'const.util_prob.FUNSET_2'(s(fun(X0,bool),X3),s(fun(X1,bool),X4)))))) | ! [X6] : (~p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X6))),s(fun(X1,bool),X4)))))) & (? [X5] : (p(s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X5))),s(fun(X1,bool),X4))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,X1),X2),s(fun(fun(X0,X1),bool),'const.util_prob.FUNSET_2'(s(fun(X0,bool),X3),s(fun(X1,bool),X4))))))))),
  inference(rectify,[],[f57])).
fof(f57,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.bool.IN_2'(s(fun(X0,X1),X2),s(fun(fun(X0,X1),bool),'const.util_prob.FUNSET_2'(s(fun(X0,bool),X3),s(fun(X1,bool),X4)))))) | ! [X5] : (~p(s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X5))),s(fun(X1,bool),X4)))))) & (? [X5] : (p(s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X5))),s(fun(X1,bool),X4))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,X1),X2),s(fun(fun(X0,X1),bool),'const.util_prob.FUNSET_2'(s(fun(X0,bool),X3),s(fun(X1,bool),X4))))))))),
  inference(nnf_transformation,[],[f42])).
fof(f42,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,X1),X2),s(fun(fun(X0,X1),bool),'const.util_prob.FUNSET_2'(s(fun(X0,bool),X3),s(fun(X1,bool),X4)))))) <=> ! [X5] : (~p(s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X5))),s(fun(X1,bool),X4))))))),
  inference(ennf_transformation,[],[f35])).
fof(f35,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,X1),X2),s(fun(fun(X0,X1),bool),'const.util_prob.FUNSET_2'(s(fun(X0,bool),X3),s(fun(X1,bool),X4)))))) <=> ! [X5] : (p(s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),X3)))) => p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X5))),s(fun(X1,bool),X4))))))),
  inference(flattening,[],[f34])).
fof(f34,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,X1),X2),s(fun(fun(X0,X1),bool),'const.util_prob.FUNSET_2'(s(fun(X0,bool),X3),s(fun(X1,bool),X4)))))) <=> ! [X5] : (p(s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),X3)))) => p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X5))),s(fun(X1,bool),X4))))))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X6,X7] : ! [X17,X18,X19] : (p(s(bool,'const.bool.IN_2'(s(fun(X6,X7),X17),s(fun(fun(X6,X7),bool),'const.util_prob.FUNSET_2'(s(fun(X6,bool),X18),s(fun(X7,bool),X19)))))) <=> ! [X8] : (p(s(bool,'const.bool.IN_2'(s(X6,X8),s(fun(X6,bool),X18)))) => p(s(bool,'const.bool.IN_2'(s(X7,happ(s(fun(X6,X7),X17),s(X6,X8))),s(fun(X7,bool),X19))))))),
  file('thm.util_prob.FUNSET_DFUNSET','thm.util_prob.IN_FUNSET')).
fof(f87,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X2),s(fun(fun(X1,X0),bool),'const.util_prob.DFUNSET_2'(s(fun(X1,bool),X3),s(fun(X1,fun(X0,bool)),X4)))))) | p(s(bool,'const.bool.IN_2'(s(X1,sK7(X4,X3,X2,X1,X0)),s(fun(X1,bool),X3))))) )),
  inference(cnf_transformation,[],[f62])).
fof(f86,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X2),s(fun(fun(X1,X0),bool),'const.util_prob.DFUNSET_2'(s(fun(X1,bool),X3),s(fun(X1,fun(X0,bool)),X4)))))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))),s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X4),s(X1,X6))))))) )),
  inference(cnf_transformation,[],[f62])).
fof(f83,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,X1),X2),s(fun(fun(X0,X1),bool),'const.util_prob.FUNSET_2'(s(fun(X0,bool),X3),s(fun(X1,bool),X4)))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X6))),s(fun(X1,bool),X4))))) )),
  inference(cnf_transformation,[],[f59])).
fof(f84,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,X1),X2),s(fun(fun(X0,X1),bool),'const.util_prob.FUNSET_2'(s(fun(X0,bool),X3),s(fun(X1,bool),X4)))))) | p(s(bool,'const.bool.IN_2'(s(X0,sK6(X4,X3,X2,X1,X0)),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f59])).
fof(f140,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3))))))) | $spl22),
  inference(cnf_transformation,[],[f140_D])).
fof(f140_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3))))))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f9071,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3))))))) | $spl98),
  inference(global_subsumption,[],[f67,f5815,f6536,f6561,f6538,f6562,f6552,f6553,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6795,f6805,f6816,f104,f102,f101,f100,f7026,f98,f96,f7027,f7591,f7593,f7594,f7596,f7598,f7599,f7600,f7611,f7621,f7667,f7670,f7634,f7675,f7644,f7684,f7687,f7688,f7689,f7691,f7698,f7699,f7702,f7695,f7703,f7704,f7705,f7706,f7710,f7711,f7712,f7766,f7769,f7718,f7770,f7779,f7780,f7781,f7785,f7786,f7787,f7837,f7838,f7841,f7858,f7859,f7860,f7863,f7864,f7918,f7923,f7924,f7925,f7926,f7928,f7935,f7936,f7939,f7932,f7940,f7941,f7942,f7943,f7947,f7948,f7949,f8003,f8006,f7955,f8007,f8020,f8021,f8022,f8024,f8025,f8026,f8078,f8067,f8083,f8195,f8196,f8198,f8203,f8204,f8205,f8255,f8259,f8277,f8278,f8279,f8325,f8331,f8332,f8333,f8336,f8338,f8397,f8401,f8402,f8403,f8450,f8451,f8453,f8462,f8465,f8457,f8466,f8467,f8470,f8471,f8472,f8474,f8485,f8488,f8478,f8489,f8490,f8493,f8494,f8495,f8496,f8497,f8499,f8500,f8501,f8502,f8504,f8505,f8506,f8507,f8508,f8555,f8558,f8545,f8561,f8565,f8566,f8567,f8568,f8570,f8571,f8572,f8573,f8575,f8576,f8577,f8578,f8579,f8624,f8616,f8632,f8637,f8638,f8639,f8640,f8642,f8643,f8644,f8645,f8646,f8647,f8648,f8688,f8689,f8690,f8691,f8693,f8694,f8695,f8696,f8698,f8699,f8700,f8701,f8702,f8703,f8705,f8706,f8707,f8708,f8710,f8711,f8712,f8713,f8714,f8715,f8716,f8717,f8718,f8729,f8739,f8788,f8752,f8793,f8762,f8802,f8805,f8807,f8808,f8809,f8810,f8858,f8859,f8861,f8862,f8863,f8864,f8869,f8870,f8871,f8872,f8874,f8931,f8966,f8969,f8936,f8937,f8948,f8958,f8963,f8989,f8990,f8991,f8992,f8993,f8995,f8996,f8997,f8998,f8999,f9001,f9002,f9003,f9004,f9005,f9022,f9025,f9008,f9009,f9015,f9033,f9020])).
fof(f9020,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3))))))) | $spl98),
  inference(superposition,[],[f6755,f5815])).
fof(f9033,plain,(
  ( ! [X14,X12,X10,X13,X11] : (p(s(bool,'const.bool.IN_2'(s(fun(X11,fun(sK0,X10)),X12),s(fun(fun(X11,fun(sK0,X10)),bool),'const.util_prob.DFUNSET_2'(s(fun(X11,bool),X14),s(fun(X11,fun(fun(sK0,X10),bool)),'const.combin.K_1'(s(fun(fun(sK0,X10),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X10,bool)),X13)))))))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK7(X13,sK2,happ(s(fun(X11,fun(sK0,X10)),X12),s(X11,sK7('const.combin.K_1'(s(fun(fun(sK0,X10),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X10,bool)),X13)))),X14,X12,X11,fun(sK0,X10)))),sK0,X10)))),s(fun(sK1,bool),sK3)))) ) | $spl98),
  inference(global_subsumption,[],[f67,f5815,f6536,f6561,f6538,f6562,f6552,f6553,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6795,f6805,f6816,f104,f102,f101,f100,f7026,f98,f96,f7027,f7591,f7593,f7594,f7596,f7598,f7599,f7600,f7611,f7621,f7667,f7670,f7634,f7675,f7644,f7684,f7687,f7688,f7689,f7691,f7698,f7699,f7702,f7695,f7703,f7704,f7705,f7706,f7710,f7711,f7712,f7766,f7769,f7718,f7770,f7779,f7780,f7781,f7785,f7786,f7787,f7837,f7838,f7841,f7858,f7859,f7860,f7863,f7864,f7918,f7923,f7924,f7925,f7926,f7928,f7935,f7936,f7939,f7932,f7940,f7941,f7942,f7943,f7947,f7948,f7949,f8003,f8006,f7955,f8007,f8020,f8021,f8022,f8024,f8025,f8026,f8078,f8067,f8083,f8195,f8196,f8198,f8203,f8204,f8205,f8255,f8259,f8277,f8278,f8279,f8325,f8331,f8332,f8333,f8336,f8338,f8397,f8401,f8402,f8403,f8450,f8451,f8453,f8462,f8465,f8457,f8466,f8467,f8470,f8471,f8472,f8474,f8485,f8488,f8478,f8489,f8490,f8493,f8494,f8495,f8496,f8497,f8499,f8500,f8501,f8502,f8504,f8505,f8506,f8507,f8508,f8555,f8558,f8545,f8561,f8565,f8566,f8567,f8568,f8570,f8571,f8572,f8573,f8575,f8576,f8577,f8578,f8579,f8624,f8616,f8632,f8637,f8638,f8639,f8640,f8642,f8643,f8644,f8645,f8646,f8647,f8648,f8688,f8689,f8690,f8691,f8693,f8694,f8695,f8696,f8698,f8699,f8700,f8701,f8702,f8703,f8705,f8706,f8707,f8708,f8710,f8711,f8712,f8713,f8714,f8715,f8716,f8717,f8718,f8729,f8739,f8788,f8752,f8793,f8762,f8802,f8805,f8807,f8808,f8809,f8810,f8858,f8859,f8861,f8862,f8863,f8864,f8869,f8870,f8871,f8872,f8874,f8931,f8966,f8969,f8936,f8937,f8948,f8958,f8963,f8989,f8990,f8991,f8992,f8993,f8995,f8996,f8997,f8998,f8999,f9001,f9002,f9003,f9004,f9005,f9022,f9025,f9008,f9009,f9015])).
fof(f9015,plain,(
  ( ! [X14,X12,X10,X13,X11] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(X11,fun(sK0,X10)),X12),s(fun(fun(X11,fun(sK0,X10)),bool),'const.util_prob.DFUNSET_2'(s(fun(X11,bool),X14),s(fun(X11,fun(fun(sK0,X10),bool)),'const.combin.K_1'(s(fun(fun(sK0,X10),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X10,bool)),X13)))))))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK7(X13,sK2,happ(s(fun(X11,fun(sK0,X10)),X12),s(X11,sK7('const.combin.K_1'(s(fun(fun(sK0,X10),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X10,bool)),X13)))),X14,X12,X11,fun(sK0,X10)))),sK0,X10)))),s(fun(sK1,bool),sK3)))) ) | $spl98),
  inference(superposition,[],[f6755,f5815])).
fof(f8958,plain,(
  ( ! [X88,X87,X85,X83,X86,X84] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),happ(s(fun(X86,fun(bool,bool)),X87),s(X86,sK7(X87,X88,'const.combin.K_1'(s(bool,'T_0')),X86,bool))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X86,bool),'const.combin.K_1'(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK7(X83,sK2,X84,sK0,X85)))),s(fun(sK1,bool),sK3))))),s(fun(fun(X86,bool),bool),'const.util_prob.DFUNSET_2'(s(fun(X86,bool),X88),s(fun(X86,fun(bool,bool)),X87)))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X85),X84),s(fun(fun(sK0,X85),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X85,bool)),X83)))))) ) | $spl98),
  inference(superposition,[],[f6750,f5815])).
fof(f8948,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),happ(s(fun(X47,fun(bool,bool)),X48),s(X47,sK7(X48,X49,'const.combin.K_1'(s(bool,'T_0')),X47,bool))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X47,bool),'const.combin.K_1'(s(bool,'const.bool.IN_2'(s(fun(sK0,X44),X45),s(fun(fun(sK0,X44),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X44,bool)),X46))))))),s(fun(fun(X47,bool),bool),'const.util_prob.DFUNSET_2'(s(fun(X47,bool),X49),s(fun(X47,fun(bool,bool)),X48)))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK7(X46,sK2,X45,sK0,X44)))),s(fun(sK1,bool),sK3)))) ) | $spl98),
  inference(superposition,[],[f6750,f5815])).
fof(f8802,plain,(
  ( ! [X68,X66,X64,X67,X65,X63] : (p(s(bool,'const.bool.IN_2'(s(X66,happ(s(fun(sK1,X66),X67),s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK7(X63,sK2,X64,sK0,X65)))))),s(fun(X66,bool),happ(s(fun(sK1,fun(X66,bool)),X68),s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK7(X63,sK2,X64,sK0,X65))))))))) | p(s(bool,'const.bool.IN_2'(s(sK1,sK7(X68,sK3,X67,sK1,X66)),s(fun(sK1,bool),sK3)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X65),X64),s(fun(fun(sK0,X65),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X65,bool)),X63)))))) ) | $spl98),
  inference(global_subsumption,[],[f67,f5815,f6536,f6561,f6538,f6562,f6552,f6553,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6795,f6805,f6816,f104,f102,f101,f100,f7026,f98,f96,f7027,f7591,f7593,f7594,f7596,f7598,f7599,f7600,f7611,f7621,f7667,f7670,f7634,f7675,f7644,f7684,f7687,f7688,f7689,f7691,f7698,f7699,f7702,f7695,f7703,f7704,f7705,f7706,f7710,f7711,f7712,f7766,f7769,f7718,f7770,f7779,f7780,f7781,f7785,f7786,f7787,f7837,f7838,f7841,f7858,f7859,f7860,f7863,f7864,f7918,f7923,f7924,f7925,f7926,f7928,f7935,f7936,f7939,f7932,f7940,f7941,f7942,f7943,f7947,f7948,f7949,f8003,f8006,f7955,f8007,f8020,f8021,f8022,f8024,f8025,f8026,f8078,f8067,f8083,f8195,f8196,f8198,f8203,f8204,f8205,f8255,f8259,f8277,f8278,f8279,f8325,f8331,f8332,f8333,f8336,f8338,f8397,f8401,f8402,f8403,f8450,f8451,f8453,f8462,f8465,f8457,f8466,f8467,f8470,f8471,f8472,f8474,f8485,f8488,f8478,f8489,f8490,f8493,f8494,f8495,f8496,f8497,f8499,f8500,f8501,f8502,f8504,f8505,f8506,f8507,f8508,f8555,f8558,f8545,f8561,f8565,f8566,f8567,f8568,f8570,f8571,f8572,f8573,f8575,f8576,f8577,f8578,f8579,f8624,f8616,f8632,f8637,f8638,f8639,f8640,f8642,f8643,f8644,f8645,f8646,f8647,f8648,f8688,f8689,f8690,f8691,f8693,f8694,f8695,f8696,f8698,f8699,f8700,f8701,f8702,f8703,f8705,f8706,f8707,f8708,f8710,f8711,f8712,f8713,f8714,f8715,f8716,f8717,f8718,f8729,f8739,f8788,f8752,f8793,f8762])).
fof(f8762,plain,(
  ( ! [X68,X66,X64,X67,X65,X63] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(X66,happ(s(fun(sK1,X66),X67),s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK7(X63,sK2,X64,sK0,X65)))))),s(fun(X66,bool),happ(s(fun(sK1,fun(X66,bool)),X68),s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK7(X63,sK2,X64,sK0,X65))))))))) | p(s(bool,'const.bool.IN_2'(s(sK1,sK7(X68,sK3,X67,sK1,X66)),s(fun(sK1,bool),sK3)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X65),X64),s(fun(fun(sK0,X65),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X65,bool)),X63)))))) ) | $spl98),
  inference(superposition,[],[f6756,f5815])).
fof(f8793,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (p(s(bool,'const.bool.IN_2'(s(X27,happ(s(fun(fun(sK0,X24),X27),X28),s(fun(sK0,X24),X25))),s(fun(X27,bool),happ(s(fun(fun(sK0,X24),fun(X27,bool)),X29),s(fun(sK0,X24),X25)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,X24),sK7(X29,'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X24,bool)),X26)),X28,fun(sK0,X24),X27)),s(fun(fun(sK0,X24),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X24,bool)),X26)))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK7(X26,sK2,X25,sK0,X24)))),s(fun(sK1,bool),sK3)))) ) | $spl98),
  inference(global_subsumption,[],[f67,f5815,f6536,f6561,f6538,f6562,f6552,f6553,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6795,f6805,f6816,f104,f102,f101,f100,f7026,f98,f96,f7027,f7591,f7593,f7594,f7596,f7598,f7599,f7600,f7611,f7621,f7667,f7670,f7634,f7675,f7644,f7684,f7687,f7688,f7689,f7691,f7698,f7699,f7702,f7695,f7703,f7704,f7705,f7706,f7710,f7711,f7712,f7766,f7769,f7718,f7770,f7779,f7780,f7781,f7785,f7786,f7787,f7837,f7838,f7841,f7858,f7859,f7860,f7863,f7864,f7918,f7923,f7924,f7925,f7926,f7928,f7935,f7936,f7939,f7932,f7940,f7941,f7942,f7943,f7947,f7948,f7949,f8003,f8006,f7955,f8007,f8020,f8021,f8022,f8024,f8025,f8026,f8078,f8067,f8083,f8195,f8196,f8198,f8203,f8204,f8205,f8255,f8259,f8277,f8278,f8279,f8325,f8331,f8332,f8333,f8336,f8338,f8397,f8401,f8402,f8403,f8450,f8451,f8453,f8462,f8465,f8457,f8466,f8467,f8470,f8471,f8472,f8474,f8485,f8488,f8478,f8489,f8490,f8493,f8494,f8495,f8496,f8497,f8499,f8500,f8501,f8502,f8504,f8505,f8506,f8507,f8508,f8555,f8558,f8545,f8561,f8565,f8566,f8567,f8568,f8570,f8571,f8572,f8573,f8575,f8576,f8577,f8578,f8579,f8624,f8616,f8632,f8637,f8638,f8639,f8640,f8642,f8643,f8644,f8645,f8646,f8647,f8648,f8688,f8689,f8690,f8691,f8693,f8694,f8695,f8696,f8698,f8699,f8700,f8701,f8702,f8703,f8705,f8706,f8707,f8708,f8710,f8711,f8712,f8713,f8714,f8715,f8716,f8717,f8718,f8729,f8739,f8788,f8752])).
fof(f8752,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(X27,happ(s(fun(fun(sK0,X24),X27),X28),s(fun(sK0,X24),X25))),s(fun(X27,bool),happ(s(fun(fun(sK0,X24),fun(X27,bool)),X29),s(fun(sK0,X24),X25)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,X24),sK7(X29,'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X24,bool)),X26)),X28,fun(sK0,X24),X27)),s(fun(fun(sK0,X24),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X24,bool)),X26)))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK7(X26,sK2,X25,sK0,X24)))),s(fun(sK1,bool),sK3)))) ) | $spl98),
  inference(superposition,[],[f6756,f5815])).
fof(f8739,plain,(
  ( ! [X111,X109,X107,X105,X110,X108,X106] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X108)))) | p(s(bool,'const.bool.IN_2'(s(X109,happ(s(fun(bool,X109),X110),s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK7(X105,sK2,X106,sK0,X107)))),s(fun(sK1,bool),sK3))))),s(fun(X109,bool),happ(s(fun(bool,fun(X109,bool)),X111),s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK7(X105,sK2,X106,sK0,X107)))),s(fun(sK1,bool),sK3)))))))) | p(s(bool,'const.bool.IN_2'(s(bool,sK7(X111,X108,X110,bool,X109)),s(fun(bool,bool),X108)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X107),X106),s(fun(fun(sK0,X107),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X107,bool)),X105)))))) ) | $spl98),
  inference(superposition,[],[f6756,f5815])).
fof(f8729,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X56] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X59)))) | p(s(bool,'const.bool.IN_2'(s(X60,happ(s(fun(bool,X60),X61),s(bool,'const.bool.IN_2'(s(fun(sK0,X56),X57),s(fun(fun(sK0,X56),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X56,bool)),X58))))))),s(fun(X60,bool),happ(s(fun(bool,fun(X60,bool)),X62),s(bool,'const.bool.IN_2'(s(fun(sK0,X56),X57),s(fun(fun(sK0,X56),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X56,bool)),X58)))))))))) | p(s(bool,'const.bool.IN_2'(s(bool,sK7(X62,X59,X61,bool,X60)),s(fun(bool,bool),X59)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK7(X58,sK2,X57,sK0,X56)))),s(fun(sK1,bool),sK3)))) ) | $spl98),
  inference(superposition,[],[f6756,f5815])).
fof(f8632,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (p(s(bool,'const.bool.IN_2'(s(fun(X15,fun(sK0,X14)),X16),s(fun(fun(X15,fun(sK0,X14)),bool),'const.util_prob.FUNSET_2'(s(fun(X15,bool),X20),s(fun(fun(sK0,X14),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X14,bool)),X17)))))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK7(X17,sK2,happ(s(fun(X15,fun(sK0,X14)),X16),s(X15,sK6('const.combin.K_2'(s(fun(fun(sK0,X14),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X14,bool)),X17))),s(X18,X19)),X20,X16,fun(sK0,X14),X15))),sK0,X14)))),s(fun(sK1,bool),sK3)))) ) | $spl98),
  inference(global_subsumption,[],[f67,f5815,f6536,f6561,f6538,f6562,f6552,f6553,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6795,f6805,f6816,f104,f102,f101,f100,f7026,f98,f96,f7027,f7591,f7593,f7594,f7596,f7598,f7599,f7600,f7611,f7621,f7667,f7670,f7634,f7675,f7644,f7684,f7687,f7688,f7689,f7691,f7698,f7699,f7702,f7695,f7703,f7704,f7705,f7706,f7710,f7711,f7712,f7766,f7769,f7718,f7770,f7779,f7780,f7781,f7785,f7786,f7787,f7837,f7838,f7841,f7858,f7859,f7860,f7863,f7864,f7918,f7923,f7924,f7925,f7926,f7928,f7935,f7936,f7939,f7932,f7940,f7941,f7942,f7943,f7947,f7948,f7949,f8003,f8006,f7955,f8007,f8020,f8021,f8022,f8024,f8025,f8026,f8078,f8067,f8083,f8195,f8196,f8198,f8203,f8204,f8205,f8255,f8259,f8277,f8278,f8279,f8325,f8331,f8332,f8333,f8336,f8338,f8397,f8401,f8402,f8403,f8450,f8451,f8453,f8462,f8465,f8457,f8466,f8467,f8470,f8471,f8472,f8474,f8485,f8488,f8478,f8489,f8490,f8493,f8494,f8495,f8496,f8497,f8499,f8500,f8501,f8502,f8504,f8505,f8506,f8507,f8508,f8555,f8558,f8545,f8561,f8565,f8566,f8567,f8568,f8570,f8571,f8572,f8573,f8575,f8576,f8577,f8578,f8579,f8624,f8616])).
fof(f8616,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(X15,fun(sK0,X14)),X16),s(fun(fun(X15,fun(sK0,X14)),bool),'const.util_prob.FUNSET_2'(s(fun(X15,bool),X20),s(fun(fun(sK0,X14),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X14,bool)),X17)))))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK7(X17,sK2,happ(s(fun(X15,fun(sK0,X14)),X16),s(X15,sK6('const.combin.K_2'(s(fun(fun(sK0,X14),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X14,bool)),X17))),s(X18,X19)),X20,X16,fun(sK0,X14),X15))),sK0,X14)))),s(fun(sK1,bool),sK3)))) ) | $spl98),
  inference(superposition,[],[f6707,f5815])).
fof(f8561,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (p(s(bool,'const.bool.IN_2'(s(fun(X15,fun(sK0,X14)),X16),s(fun(fun(X15,fun(sK0,X14)),bool),'const.util_prob.FUNSET_2'(s(fun(X15,bool),X18),s(fun(fun(sK0,X14),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X14,bool)),X17)))))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK7(X17,sK2,happ(s(fun(X15,fun(sK0,X14)),X16),s(X15,sK6('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X14,bool)),X17)),X18,'const.combin.K_2'(s(fun(X15,fun(sK0,X14)),X16),s(X19,X20)),fun(sK0,X14),X15))),sK0,X14)))),s(fun(sK1,bool),sK3)))) ) | $spl98),
  inference(global_subsumption,[],[f67,f5815,f6536,f6561,f6538,f6562,f6552,f6553,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6795,f6805,f6816,f104,f102,f101,f100,f7026,f98,f96,f7027,f7591,f7593,f7594,f7596,f7598,f7599,f7600,f7611,f7621,f7667,f7670,f7634,f7675,f7644,f7684,f7687,f7688,f7689,f7691,f7698,f7699,f7702,f7695,f7703,f7704,f7705,f7706,f7710,f7711,f7712,f7766,f7769,f7718,f7770,f7779,f7780,f7781,f7785,f7786,f7787,f7837,f7838,f7841,f7858,f7859,f7860,f7863,f7864,f7918,f7923,f7924,f7925,f7926,f7928,f7935,f7936,f7939,f7932,f7940,f7941,f7942,f7943,f7947,f7948,f7949,f8003,f8006,f7955,f8007,f8020,f8021,f8022,f8024,f8025,f8026,f8078,f8067,f8083,f8195,f8196,f8198,f8203,f8204,f8205,f8255,f8259,f8277,f8278,f8279,f8325,f8331,f8332,f8333,f8336,f8338,f8397,f8401,f8402,f8403,f8450,f8451,f8453,f8462,f8465,f8457,f8466,f8467,f8470,f8471,f8472,f8474,f8485,f8488,f8478,f8489,f8490,f8493,f8494,f8495,f8496,f8497,f8499,f8500,f8501,f8502,f8504,f8505,f8506,f8507,f8508,f8555,f8558,f8545])).
fof(f8545,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(X15,fun(sK0,X14)),X16),s(fun(fun(X15,fun(sK0,X14)),bool),'const.util_prob.FUNSET_2'(s(fun(X15,bool),X18),s(fun(fun(sK0,X14),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X14,bool)),X17)))))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK7(X17,sK2,happ(s(fun(X15,fun(sK0,X14)),X16),s(X15,sK6('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X14,bool)),X17)),X18,'const.combin.K_2'(s(fun(X15,fun(sK0,X14)),X16),s(X19,X20)),fun(sK0,X14),X15))),sK0,X14)))),s(fun(sK1,bool),sK3)))) ) | $spl98),
  inference(superposition,[],[f6702,f5815])).
fof(f7684,plain,(
  ( ! [X68,X66,X64,X67,X65,X63] : (p(s(bool,'const.bool.IN_2'(s(X66,happ(s(fun(sK1,X66),X67),s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK7(X63,sK2,X64,sK0,X65)))))),s(fun(X66,bool),X68)))) | p(s(bool,'const.bool.IN_2'(s(sK1,sK6(X68,sK3,X67,X66,sK1)),s(fun(sK1,bool),sK3)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X65),X64),s(fun(fun(sK0,X65),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X65,bool)),X63)))))) ) | $spl98),
  inference(global_subsumption,[],[f67,f5815,f6536,f6561,f6538,f6562,f6552,f6553,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6795,f6805,f6816,f104,f102,f101,f100,f7026,f98,f96,f7027,f7591,f7593,f7594,f7596,f7598,f7599,f7600,f7611,f7621,f7667,f7670,f7634,f7675,f7644])).
fof(f7644,plain,(
  ( ! [X68,X66,X64,X67,X65,X63] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(X66,happ(s(fun(sK1,X66),X67),s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK7(X63,sK2,X64,sK0,X65)))))),s(fun(X66,bool),X68)))) | p(s(bool,'const.bool.IN_2'(s(sK1,sK6(X68,sK3,X67,X66,sK1)),s(fun(sK1,bool),sK3)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X65),X64),s(fun(fun(sK0,X65),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X65,bool)),X63)))))) ) | $spl98),
  inference(superposition,[],[f6711,f5815])).
fof(f7675,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (p(s(bool,'const.bool.IN_2'(s(X27,happ(s(fun(fun(sK0,X24),X27),X28),s(fun(sK0,X24),X25))),s(fun(X27,bool),X29)))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,X24),sK6(X29,'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X24,bool)),X26)),X28,X27,fun(sK0,X24))),s(fun(fun(sK0,X24),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X24,bool)),X26)))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK7(X26,sK2,X25,sK0,X24)))),s(fun(sK1,bool),sK3)))) ) | $spl98),
  inference(global_subsumption,[],[f67,f5815,f6536,f6561,f6538,f6562,f6552,f6553,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6795,f6805,f6816,f104,f102,f101,f100,f7026,f98,f96,f7027,f7591,f7593,f7594,f7596,f7598,f7599,f7600,f7611,f7621,f7667,f7670,f7634])).
fof(f7634,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(X27,happ(s(fun(fun(sK0,X24),X27),X28),s(fun(sK0,X24),X25))),s(fun(X27,bool),X29)))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,X24),sK6(X29,'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X24,bool)),X26)),X28,X27,fun(sK0,X24))),s(fun(fun(sK0,X24),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X24,bool)),X26)))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK7(X26,sK2,X25,sK0,X24)))),s(fun(sK1,bool),sK3)))) ) | $spl98),
  inference(superposition,[],[f6711,f5815])).
fof(f7621,plain,(
  ( ! [X111,X109,X107,X105,X110,X108,X106] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X108)))) | p(s(bool,'const.bool.IN_2'(s(X109,happ(s(fun(bool,X109),X110),s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK7(X105,sK2,X106,sK0,X107)))),s(fun(sK1,bool),sK3))))),s(fun(X109,bool),X111)))) | p(s(bool,'const.bool.IN_2'(s(bool,sK6(X111,X108,X110,X109,bool)),s(fun(bool,bool),X108)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X107),X106),s(fun(fun(sK0,X107),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X107,bool)),X105)))))) ) | $spl98),
  inference(superposition,[],[f6711,f5815])).
fof(f7611,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X56] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X59)))) | p(s(bool,'const.bool.IN_2'(s(X60,happ(s(fun(bool,X60),X61),s(bool,'const.bool.IN_2'(s(fun(sK0,X56),X57),s(fun(fun(sK0,X56),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X56,bool)),X58))))))),s(fun(X60,bool),X62)))) | p(s(bool,'const.bool.IN_2'(s(bool,sK6(X62,X59,X61,X60,bool)),s(fun(bool,bool),X59)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK7(X58,sK2,X57,sK0,X56)))),s(fun(sK1,bool),sK3)))) ) | $spl98),
  inference(superposition,[],[f6711,f5815])).
fof(f6805,plain,(
  ( ! [X88,X87,X85,X83,X86,X84] : (p(s(bool,'const.bool.IN_2'(s(fun(X86,bool),'const.combin.K_1'(s(bool,'T_0'))),s(fun(fun(X86,bool),bool),'const.util_prob.FUNSET_2'(s(fun(X86,bool),X87),s(fun(bool,bool),X88)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK7(X83,sK2,X84,sK0,X85)))),s(fun(sK1,bool),sK3))),s(fun(bool,bool),X88)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X85),X84),s(fun(fun(sK0,X85),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X85,bool)),X83)))))) ) | $spl98),
  inference(superposition,[],[f6706,f5815])).
fof(f6795,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (p(s(bool,'const.bool.IN_2'(s(fun(X47,bool),'const.combin.K_1'(s(bool,'T_0'))),s(fun(fun(X47,bool),bool),'const.util_prob.FUNSET_2'(s(fun(X47,bool),X48),s(fun(bool,bool),X49)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(fun(sK0,X44),X45),s(fun(fun(sK0,X44),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X44,bool)),X46))))),s(fun(bool,bool),X49)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK7(X46,sK2,X45,sK0,X44)))),s(fun(sK1,bool),sK3)))) ) | $spl98),
  inference(superposition,[],[f6706,f5815])).
fof(f6553,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,X6),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X6,bool)),X7)),X8,fun(sK0,X6))),s(fun(fun(sK0,X6),bool),X8))) | s(fun(fun(sK0,X6),bool),X8) = s(fun(fun(sK0,X6),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X6,bool)),X7))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK7(X7,sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X6,bool)),X7)),X8,fun(sK0,X6)),sK0,X6)))),s(fun(sK1,bool),sK3)))) ) | $spl98),
  inference(superposition,[],[f79,f5815])).
fof(f79,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) )),
  inference(cnf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f54])).
fof(f54,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(rectify,[],[f53])).
fof(f53,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(nnf_transformation,[],[f27])).
fof(f27,plain,(
  ! [X0,X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(flattening,[],[f26])).
fof(f26,plain,(
  ! [X0] : ! [X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X6] : ! [X16,X10] : (s(fun(X6,bool),X10) = s(fun(X6,bool),X16) <=> ! [X8] : s(bool,'const.bool.IN_2'(s(X6,X8),s(fun(X6,bool),X10))) = s(bool,'const.bool.IN_2'(s(X6,X8),s(fun(X6,bool),X16))))),
  file('thm.util_prob.FUNSET_DFUNSET','thm.pred_set.EXTENSION')).
fof(f6552,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,X3),sK4(X4,'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X3,bool)),X5)),fun(sK0,X3))),s(fun(fun(sK0,X3),bool),X4))) | s(fun(fun(sK0,X3),bool),X4) = s(fun(fun(sK0,X3),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X3,bool)),X5))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK7(X5,sK2,sK4(X4,'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X3,bool)),X5)),fun(sK0,X3)),sK0,X3)))),s(fun(sK1,bool),sK3)))) ) | $spl98),
  inference(superposition,[],[f79,f5815])).
fof(f6562,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X5),X9),s(fun(fun(sK0,X5),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X5,bool)),X7))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK7(happ(s(fun(X6,fun(sK0,fun(X5,bool))),'const.combin.K_1'(s(fun(sK0,fun(X5,bool)),X7))),s(X6,X8)),sK2,X9,sK0,X5)))),s(fun(sK1,bool),sK3)))) ) | $spl98),
  inference(forward_demodulation,[],[f6539,f80])).
fof(f6539,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X5),X9),s(fun(fun(sK0,X5),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X5,bool)),'const.combin.K_2'(s(fun(sK0,fun(X5,bool)),X7),s(X6,X8))))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK7(happ(s(fun(X6,fun(sK0,fun(X5,bool))),'const.combin.K_1'(s(fun(sK0,fun(X5,bool)),X7))),s(X6,X8)),sK2,X9,sK0,X5)))),s(fun(sK1,bool),sK3)))) ) | $spl98),
  inference(superposition,[],[f5815,f81])).
fof(f6538,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X0),X4),s(fun(fun(sK0,X0),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X0,bool)),X1))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK7('const.combin.K_2'(s(fun(sK0,fun(X0,bool)),X1),s(X2,X3)),sK2,X4,sK0,X0)))),s(fun(sK1,bool),sK3)))) ) | $spl98),
  inference(superposition,[],[f5815,f80])).
fof(f6561,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X5),X7),s(fun(fun(sK0,X5),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X5,bool)),X9))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK7(X9,sK2,happ(s(fun(X6,fun(sK0,X5)),'const.combin.K_1'(s(fun(sK0,X5),X7))),s(X6,X8)),sK0,X5)))),s(fun(sK1,bool),sK3)))) ) | $spl98),
  inference(forward_demodulation,[],[f6537,f80])).
fof(f6537,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X5),'const.combin.K_2'(s(fun(sK0,X5),X7),s(X6,X8))),s(fun(fun(sK0,X5),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X5,bool)),X9))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK7(X9,sK2,happ(s(fun(X6,fun(sK0,X5)),'const.combin.K_1'(s(fun(sK0,X5),X7))),s(X6,X8)),sK0,X5)))),s(fun(sK1,bool),sK3)))) ) | $spl98),
  inference(superposition,[],[f5815,f81])).
fof(f6536,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X0),X1),s(fun(fun(sK0,X0),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X0,bool)),X4))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK7(X4,sK2,'const.combin.K_2'(s(fun(sK0,X0),X1),s(X2,X3)),sK0,X0)))),s(fun(sK1,bool),sK3)))) ) | $spl98),
  inference(superposition,[],[f5815,f80])).
fof(f9063,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)))))) | $spl170),
  inference(cnf_transformation,[],[f9063_D])).
fof(f9063_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)))))) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f956,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))))) | $spl46),
  inference(cnf_transformation,[],[f956_D])).
fof(f956_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))))) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f6258,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)))))) | $spl88),
  inference(resolution,[],[f5760,f84])).
fof(f9070,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))))) | $spl94),
  inference(global_subsumption,[],[f67,f5791,f6437,f6498,f6439,f6499,f6447,f6448,f6451,f6452,f6453,f6454,f6455,f6456,f6472,f6473,f6479,f6483,f6486,f6487,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6793,f6807,f6816,f104,f102,f101,f100,f7026,f98,f96,f7027,f7591,f7593,f7594,f7596,f7598,f7599,f7600,f7609,f7623,f7667,f7670,f7632,f7673,f7646,f7686,f7687,f7688,f7689,f7691,f7698,f7699,f7702,f7695,f7703,f7704,f7705,f7706,f7710,f7711,f7712,f7766,f7769,f7718,f7770,f7779,f7780,f7781,f7785,f7786,f7787,f7837,f7838,f7841,f7858,f7859,f7860,f7863,f7864,f7918,f7923,f7924,f7925,f7926,f7928,f7935,f7936,f7939,f7932,f7940,f7941,f7942,f7943,f7947,f7948,f7949,f8003,f8006,f7955,f8007,f8020,f8021,f8022,f8024,f8025,f8026,f8078,f8067,f8083,f8195,f8196,f8198,f8203,f8204,f8205,f8255,f8259,f8277,f8278,f8279,f8325,f8331,f8332,f8333,f8336,f8338,f8397,f8401,f8402,f8403,f8450,f8451,f8453,f8462,f8465,f8457,f8466,f8467,f8470,f8471,f8472,f8474,f8485,f8488,f8478,f8489,f8490,f8493,f8494,f8495,f8496,f8497,f8499,f8500,f8501,f8502,f8504,f8505,f8506,f8507,f8508,f8555,f8558,f8543,f8559,f8565,f8566,f8567,f8568,f8570,f8571,f8572,f8573,f8575,f8576,f8577,f8578,f8579,f8624,f8614,f8630,f8637,f8638,f8639,f8640,f8642,f8643,f8644,f8645,f8646,f8647,f8648,f8688,f8689,f8690,f8691,f8693,f8694,f8695,f8696,f8698,f8699,f8700,f8701,f8702,f8703,f8705,f8706,f8707,f8708,f8710,f8711,f8712,f8713,f8714,f8715,f8716,f8717,f8718,f8727,f8741,f8788,f8750,f8791,f8764,f8804,f8805,f8807,f8808,f8809,f8810,f8858,f8859,f8861,f8862,f8863,f8864,f8869,f8870,f8871,f8872,f8874,f8931,f8966,f8969,f8936,f8937,f8946,f8960,f8963,f8989,f8990,f8991,f8992,f8993,f8995,f8996,f8997,f8998,f8999,f9001,f9002,f9003,f9004,f9005,f9022,f9025,f9008,f9009,f9013,f9031,f9060,f9061,f9064,f9065,f6700,f9069,f6443])).
fof(f6443,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))))) | $spl94),
  inference(superposition,[],[f85,f5791])).
fof(f6700,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))))) | $spl94),
  inference(superposition,[],[f85,f5791])).
fof(f9064,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))))) | ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)))))) | $spl94),
  inference(forward_demodulation,[],[f8548,f80])).
fof(f8548,plain,(
  ( ! [X30,X31] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),'const.combin.K_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(X30,X31))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)))))) ) | $spl94),
  inference(superposition,[],[f6702,f5791])).
fof(f9060,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))))) | ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)))))) | $spl94),
  inference(forward_demodulation,[],[f8619,f80])).
fof(f8619,plain,(
  ( ! [X30,X31] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),'const.combin.K_2'(s(fun(sK1,bool),sK3),s(X30,X31)))))))) ) | $spl94),
  inference(superposition,[],[f6707,f5791])).
fof(f9031,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(X1,fun(sK0,X0)),X2),s(fun(fun(X1,fun(sK0,X0)),bool),'const.util_prob.DFUNSET_2'(s(fun(X1,bool),X4),s(fun(X1,fun(fun(sK0,X0),bool)),'const.combin.K_1'(s(fun(fun(sK0,X0),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X0,bool),X3)))))))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK6(X3,sK2,happ(s(fun(X1,fun(sK0,X0)),X2),s(X1,sK7('const.combin.K_1'(s(fun(fun(sK0,X0),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X0,bool),X3)))),X4,X2,X1,fun(sK0,X0)))),X0,sK0)))),s(fun(sK1,bool),sK3)))) ) | $spl94),
  inference(global_subsumption,[],[f67,f5791,f6437,f6498,f6439,f6499,f6447,f6448,f6451,f6452,f6453,f6454,f6455,f6456,f6472,f6473,f6479,f6483,f6486,f6487,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6793,f6807,f6816,f104,f102,f101,f100,f7026,f98,f96,f7027,f7591,f7593,f7594,f7596,f7598,f7599,f7600,f7609,f7623,f7667,f7670,f7632,f7673,f7646,f7686,f7687,f7688,f7689,f7691,f7698,f7699,f7702,f7695,f7703,f7704,f7705,f7706,f7710,f7711,f7712,f7766,f7769,f7718,f7770,f7779,f7780,f7781,f7785,f7786,f7787,f7837,f7838,f7841,f7858,f7859,f7860,f7863,f7864,f7918,f7923,f7924,f7925,f7926,f7928,f7935,f7936,f7939,f7932,f7940,f7941,f7942,f7943,f7947,f7948,f7949,f8003,f8006,f7955,f8007,f8020,f8021,f8022,f8024,f8025,f8026,f8078,f8067,f8083,f8195,f8196,f8198,f8203,f8204,f8205,f8255,f8259,f8277,f8278,f8279,f8325,f8331,f8332,f8333,f8336,f8338,f8397,f8401,f8402,f8403,f8450,f8451,f8453,f8462,f8465,f8457,f8466,f8467,f8470,f8471,f8472,f8474,f8485,f8488,f8478,f8489,f8490,f8493,f8494,f8495,f8496,f8497,f8499,f8500,f8501,f8502,f8504,f8505,f8506,f8507,f8508,f8555,f8558,f8543,f8559,f8565,f8566,f8567,f8568,f8570,f8571,f8572,f8573,f8575,f8576,f8577,f8578,f8579,f8624,f8614,f8630,f8637,f8638,f8639,f8640,f8642,f8643,f8644,f8645,f8646,f8647,f8648,f8688,f8689,f8690,f8691,f8693,f8694,f8695,f8696,f8698,f8699,f8700,f8701,f8702,f8703,f8705,f8706,f8707,f8708,f8710,f8711,f8712,f8713,f8714,f8715,f8716,f8717,f8718,f8727,f8741,f8788,f8750,f8791,f8764,f8804,f8805,f8807,f8808,f8809,f8810,f8858,f8859,f8861,f8862,f8863,f8864,f8869,f8870,f8871,f8872,f8874,f8931,f8966,f8969,f8936,f8937,f8946,f8960,f8963,f8989,f8990,f8991,f8992,f8993,f8995,f8996,f8997,f8998,f8999,f9001,f9002,f9003,f9004,f9005,f9022,f9025,f9008,f9009,f9013])).
fof(f9013,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(X1,fun(sK0,X0)),X2),s(fun(fun(X1,fun(sK0,X0)),bool),'const.util_prob.DFUNSET_2'(s(fun(X1,bool),X4),s(fun(X1,fun(fun(sK0,X0),bool)),'const.combin.K_1'(s(fun(fun(sK0,X0),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X0,bool),X3)))))))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK6(X3,sK2,happ(s(fun(X1,fun(sK0,X0)),X2),s(X1,sK7('const.combin.K_1'(s(fun(fun(sK0,X0),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X0,bool),X3)))),X4,X2,X1,fun(sK0,X0)))),X0,sK0)))),s(fun(sK1,bool),sK3)))) ) | $spl94),
  inference(superposition,[],[f6755,f5791])).
fof(f8960,plain,(
  ( ! [X99,X97,X95,X100,X98,X96] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),happ(s(fun(X98,fun(bool,bool)),X99),s(X98,sK7(X99,X100,'const.combin.K_1'(s(bool,'T_0')),X98,bool))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X98,bool),'const.combin.K_1'(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK6(X95,sK2,X96,X97,sK0)))),s(fun(sK1,bool),sK3))))),s(fun(fun(X98,bool),bool),'const.util_prob.DFUNSET_2'(s(fun(X98,bool),X100),s(fun(X98,fun(bool,bool)),X99)))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X97),X96),s(fun(fun(sK0,X97),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X97,bool),X95)))))) ) | $spl94),
  inference(superposition,[],[f6750,f5791])).
fof(f8946,plain,(
  ( ! [X37,X35,X33,X36,X34,X32] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),happ(s(fun(X35,fun(bool,bool)),X36),s(X35,sK7(X36,X37,'const.combin.K_1'(s(bool,'T_0')),X35,bool))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X35,bool),'const.combin.K_1'(s(bool,'const.bool.IN_2'(s(fun(sK0,X32),X33),s(fun(fun(sK0,X32),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X32,bool),X34))))))),s(fun(fun(X35,bool),bool),'const.util_prob.DFUNSET_2'(s(fun(X35,bool),X37),s(fun(X35,fun(bool,bool)),X36)))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK6(X34,sK2,X33,X32,sK0)))),s(fun(sK1,bool),sK3)))) ) | $spl94),
  inference(superposition,[],[f6750,f5791])).
fof(f8804,plain,(
  ( ! [X80,X78,X76,X79,X77,X75] : (p(s(bool,'const.bool.IN_2'(s(X78,happ(s(fun(sK1,X78),X79),s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK6(X75,sK2,X76,X77,sK0)))))),s(fun(X78,bool),happ(s(fun(sK1,fun(X78,bool)),X80),s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK6(X75,sK2,X76,X77,sK0))))))))) | p(s(bool,'const.bool.IN_2'(s(sK1,sK7(X80,sK3,X79,sK1,X78)),s(fun(sK1,bool),sK3)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X77),X76),s(fun(fun(sK0,X77),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X77,bool),X75)))))) ) | $spl94),
  inference(global_subsumption,[],[f67,f5791,f6437,f6498,f6439,f6499,f6447,f6448,f6451,f6452,f6453,f6454,f6455,f6456,f6472,f6473,f6479,f6483,f6486,f6487,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6793,f6807,f6816,f104,f102,f101,f100,f7026,f98,f96,f7027,f7591,f7593,f7594,f7596,f7598,f7599,f7600,f7609,f7623,f7667,f7670,f7632,f7673,f7646,f7686,f7687,f7688,f7689,f7691,f7698,f7699,f7702,f7695,f7703,f7704,f7705,f7706,f7710,f7711,f7712,f7766,f7769,f7718,f7770,f7779,f7780,f7781,f7785,f7786,f7787,f7837,f7838,f7841,f7858,f7859,f7860,f7863,f7864,f7918,f7923,f7924,f7925,f7926,f7928,f7935,f7936,f7939,f7932,f7940,f7941,f7942,f7943,f7947,f7948,f7949,f8003,f8006,f7955,f8007,f8020,f8021,f8022,f8024,f8025,f8026,f8078,f8067,f8083,f8195,f8196,f8198,f8203,f8204,f8205,f8255,f8259,f8277,f8278,f8279,f8325,f8331,f8332,f8333,f8336,f8338,f8397,f8401,f8402,f8403,f8450,f8451,f8453,f8462,f8465,f8457,f8466,f8467,f8470,f8471,f8472,f8474,f8485,f8488,f8478,f8489,f8490,f8493,f8494,f8495,f8496,f8497,f8499,f8500,f8501,f8502,f8504,f8505,f8506,f8507,f8508,f8555,f8558,f8543,f8559,f8565,f8566,f8567,f8568,f8570,f8571,f8572,f8573,f8575,f8576,f8577,f8578,f8579,f8624,f8614,f8630,f8637,f8638,f8639,f8640,f8642,f8643,f8644,f8645,f8646,f8647,f8648,f8688,f8689,f8690,f8691,f8693,f8694,f8695,f8696,f8698,f8699,f8700,f8701,f8702,f8703,f8705,f8706,f8707,f8708,f8710,f8711,f8712,f8713,f8714,f8715,f8716,f8717,f8718,f8727,f8741,f8788,f8750,f8791,f8764])).
fof(f8764,plain,(
  ( ! [X80,X78,X76,X79,X77,X75] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(X78,happ(s(fun(sK1,X78),X79),s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK6(X75,sK2,X76,X77,sK0)))))),s(fun(X78,bool),happ(s(fun(sK1,fun(X78,bool)),X80),s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK6(X75,sK2,X76,X77,sK0))))))))) | p(s(bool,'const.bool.IN_2'(s(sK1,sK7(X80,sK3,X79,sK1,X78)),s(fun(sK1,bool),sK3)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X77),X76),s(fun(fun(sK0,X77),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X77,bool),X75)))))) ) | $spl94),
  inference(superposition,[],[f6756,f5791])).
fof(f8791,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (p(s(bool,'const.bool.IN_2'(s(X15,happ(s(fun(fun(sK0,X12),X15),X16),s(fun(sK0,X12),X13))),s(fun(X15,bool),happ(s(fun(fun(sK0,X12),fun(X15,bool)),X17),s(fun(sK0,X12),X13)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,X12),sK7(X17,'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X12,bool),X14)),X16,fun(sK0,X12),X15)),s(fun(fun(sK0,X12),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X12,bool),X14)))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK6(X14,sK2,X13,X12,sK0)))),s(fun(sK1,bool),sK3)))) ) | $spl94),
  inference(global_subsumption,[],[f67,f5791,f6437,f6498,f6439,f6499,f6447,f6448,f6451,f6452,f6453,f6454,f6455,f6456,f6472,f6473,f6479,f6483,f6486,f6487,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6793,f6807,f6816,f104,f102,f101,f100,f7026,f98,f96,f7027,f7591,f7593,f7594,f7596,f7598,f7599,f7600,f7609,f7623,f7667,f7670,f7632,f7673,f7646,f7686,f7687,f7688,f7689,f7691,f7698,f7699,f7702,f7695,f7703,f7704,f7705,f7706,f7710,f7711,f7712,f7766,f7769,f7718,f7770,f7779,f7780,f7781,f7785,f7786,f7787,f7837,f7838,f7841,f7858,f7859,f7860,f7863,f7864,f7918,f7923,f7924,f7925,f7926,f7928,f7935,f7936,f7939,f7932,f7940,f7941,f7942,f7943,f7947,f7948,f7949,f8003,f8006,f7955,f8007,f8020,f8021,f8022,f8024,f8025,f8026,f8078,f8067,f8083,f8195,f8196,f8198,f8203,f8204,f8205,f8255,f8259,f8277,f8278,f8279,f8325,f8331,f8332,f8333,f8336,f8338,f8397,f8401,f8402,f8403,f8450,f8451,f8453,f8462,f8465,f8457,f8466,f8467,f8470,f8471,f8472,f8474,f8485,f8488,f8478,f8489,f8490,f8493,f8494,f8495,f8496,f8497,f8499,f8500,f8501,f8502,f8504,f8505,f8506,f8507,f8508,f8555,f8558,f8543,f8559,f8565,f8566,f8567,f8568,f8570,f8571,f8572,f8573,f8575,f8576,f8577,f8578,f8579,f8624,f8614,f8630,f8637,f8638,f8639,f8640,f8642,f8643,f8644,f8645,f8646,f8647,f8648,f8688,f8689,f8690,f8691,f8693,f8694,f8695,f8696,f8698,f8699,f8700,f8701,f8702,f8703,f8705,f8706,f8707,f8708,f8710,f8711,f8712,f8713,f8714,f8715,f8716,f8717,f8718,f8727,f8741,f8788,f8750])).
fof(f8750,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(X15,happ(s(fun(fun(sK0,X12),X15),X16),s(fun(sK0,X12),X13))),s(fun(X15,bool),happ(s(fun(fun(sK0,X12),fun(X15,bool)),X17),s(fun(sK0,X12),X13)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,X12),sK7(X17,'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X12,bool),X14)),X16,fun(sK0,X12),X15)),s(fun(fun(sK0,X12),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X12,bool),X14)))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK6(X14,sK2,X13,X12,sK0)))),s(fun(sK1,bool),sK3)))) ) | $spl94),
  inference(superposition,[],[f6756,f5791])).
fof(f8741,plain,(
  ( ! [X125,X123,X121,X124,X122,X120,X119] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X122)))) | p(s(bool,'const.bool.IN_2'(s(X123,happ(s(fun(bool,X123),X124),s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK6(X119,sK2,X120,X121,sK0)))),s(fun(sK1,bool),sK3))))),s(fun(X123,bool),happ(s(fun(bool,fun(X123,bool)),X125),s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK6(X119,sK2,X120,X121,sK0)))),s(fun(sK1,bool),sK3)))))))) | p(s(bool,'const.bool.IN_2'(s(bool,sK7(X125,X122,X124,bool,X123)),s(fun(bool,bool),X122)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X121),X120),s(fun(fun(sK0,X121),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X121,bool),X119)))))) ) | $spl94),
  inference(superposition,[],[f6756,f5791])).
fof(f8727,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X45)))) | p(s(bool,'const.bool.IN_2'(s(X46,happ(s(fun(bool,X46),X47),s(bool,'const.bool.IN_2'(s(fun(sK0,X42),X43),s(fun(fun(sK0,X42),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X42,bool),X44))))))),s(fun(X46,bool),happ(s(fun(bool,fun(X46,bool)),X48),s(bool,'const.bool.IN_2'(s(fun(sK0,X42),X43),s(fun(fun(sK0,X42),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X42,bool),X44)))))))))) | p(s(bool,'const.bool.IN_2'(s(bool,sK7(X48,X45,X47,bool,X46)),s(fun(bool,bool),X45)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK6(X44,sK2,X43,X42,sK0)))),s(fun(sK1,bool),sK3)))) ) | $spl94),
  inference(superposition,[],[f6756,f5791])).
fof(f8630,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(X1,fun(sK0,X0)),X2),s(fun(fun(X1,fun(sK0,X0)),bool),'const.util_prob.FUNSET_2'(s(fun(X1,bool),X6),s(fun(fun(sK0,X0),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X0,bool),X3)))))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK6(X3,sK2,happ(s(fun(X1,fun(sK0,X0)),X2),s(X1,sK6('const.combin.K_2'(s(fun(fun(sK0,X0),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X0,bool),X3))),s(X4,X5)),X6,X2,fun(sK0,X0),X1))),X0,sK0)))),s(fun(sK1,bool),sK3)))) ) | $spl94),
  inference(global_subsumption,[],[f67,f5791,f6437,f6498,f6439,f6499,f6447,f6448,f6451,f6452,f6453,f6454,f6455,f6456,f6472,f6473,f6479,f6483,f6486,f6487,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6793,f6807,f6816,f104,f102,f101,f100,f7026,f98,f96,f7027,f7591,f7593,f7594,f7596,f7598,f7599,f7600,f7609,f7623,f7667,f7670,f7632,f7673,f7646,f7686,f7687,f7688,f7689,f7691,f7698,f7699,f7702,f7695,f7703,f7704,f7705,f7706,f7710,f7711,f7712,f7766,f7769,f7718,f7770,f7779,f7780,f7781,f7785,f7786,f7787,f7837,f7838,f7841,f7858,f7859,f7860,f7863,f7864,f7918,f7923,f7924,f7925,f7926,f7928,f7935,f7936,f7939,f7932,f7940,f7941,f7942,f7943,f7947,f7948,f7949,f8003,f8006,f7955,f8007,f8020,f8021,f8022,f8024,f8025,f8026,f8078,f8067,f8083,f8195,f8196,f8198,f8203,f8204,f8205,f8255,f8259,f8277,f8278,f8279,f8325,f8331,f8332,f8333,f8336,f8338,f8397,f8401,f8402,f8403,f8450,f8451,f8453,f8462,f8465,f8457,f8466,f8467,f8470,f8471,f8472,f8474,f8485,f8488,f8478,f8489,f8490,f8493,f8494,f8495,f8496,f8497,f8499,f8500,f8501,f8502,f8504,f8505,f8506,f8507,f8508,f8555,f8558,f8543,f8559,f8565,f8566,f8567,f8568,f8570,f8571,f8572,f8573,f8575,f8576,f8577,f8578,f8579,f8624,f8614])).
fof(f8614,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(X1,fun(sK0,X0)),X2),s(fun(fun(X1,fun(sK0,X0)),bool),'const.util_prob.FUNSET_2'(s(fun(X1,bool),X6),s(fun(fun(sK0,X0),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X0,bool),X3)))))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK6(X3,sK2,happ(s(fun(X1,fun(sK0,X0)),X2),s(X1,sK6('const.combin.K_2'(s(fun(fun(sK0,X0),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X0,bool),X3))),s(X4,X5)),X6,X2,fun(sK0,X0),X1))),X0,sK0)))),s(fun(sK1,bool),sK3)))) ) | $spl94),
  inference(superposition,[],[f6707,f5791])).
fof(f8559,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(X1,fun(sK0,X0)),X2),s(fun(fun(X1,fun(sK0,X0)),bool),'const.util_prob.FUNSET_2'(s(fun(X1,bool),X4),s(fun(fun(sK0,X0),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X0,bool),X3)))))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK6(X3,sK2,happ(s(fun(X1,fun(sK0,X0)),X2),s(X1,sK6('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X0,bool),X3)),X4,'const.combin.K_2'(s(fun(X1,fun(sK0,X0)),X2),s(X5,X6)),fun(sK0,X0),X1))),X0,sK0)))),s(fun(sK1,bool),sK3)))) ) | $spl94),
  inference(global_subsumption,[],[f67,f5791,f6437,f6498,f6439,f6499,f6447,f6448,f6451,f6452,f6453,f6454,f6455,f6456,f6472,f6473,f6479,f6483,f6486,f6487,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6793,f6807,f6816,f104,f102,f101,f100,f7026,f98,f96,f7027,f7591,f7593,f7594,f7596,f7598,f7599,f7600,f7609,f7623,f7667,f7670,f7632,f7673,f7646,f7686,f7687,f7688,f7689,f7691,f7698,f7699,f7702,f7695,f7703,f7704,f7705,f7706,f7710,f7711,f7712,f7766,f7769,f7718,f7770,f7779,f7780,f7781,f7785,f7786,f7787,f7837,f7838,f7841,f7858,f7859,f7860,f7863,f7864,f7918,f7923,f7924,f7925,f7926,f7928,f7935,f7936,f7939,f7932,f7940,f7941,f7942,f7943,f7947,f7948,f7949,f8003,f8006,f7955,f8007,f8020,f8021,f8022,f8024,f8025,f8026,f8078,f8067,f8083,f8195,f8196,f8198,f8203,f8204,f8205,f8255,f8259,f8277,f8278,f8279,f8325,f8331,f8332,f8333,f8336,f8338,f8397,f8401,f8402,f8403,f8450,f8451,f8453,f8462,f8465,f8457,f8466,f8467,f8470,f8471,f8472,f8474,f8485,f8488,f8478,f8489,f8490,f8493,f8494,f8495,f8496,f8497,f8499,f8500,f8501,f8502,f8504,f8505,f8506,f8507,f8508,f8555,f8558,f8543])).
fof(f8543,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(X1,fun(sK0,X0)),X2),s(fun(fun(X1,fun(sK0,X0)),bool),'const.util_prob.FUNSET_2'(s(fun(X1,bool),X4),s(fun(fun(sK0,X0),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X0,bool),X3)))))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK6(X3,sK2,happ(s(fun(X1,fun(sK0,X0)),X2),s(X1,sK6('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X0,bool),X3)),X4,'const.combin.K_2'(s(fun(X1,fun(sK0,X0)),X2),s(X5,X6)),fun(sK0,X0),X1))),X0,sK0)))),s(fun(sK1,bool),sK3)))) ) | $spl94),
  inference(superposition,[],[f6702,f5791])).
fof(f7686,plain,(
  ( ! [X80,X78,X76,X79,X77,X75] : (p(s(bool,'const.bool.IN_2'(s(X78,happ(s(fun(sK1,X78),X79),s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK6(X75,sK2,X76,X77,sK0)))))),s(fun(X78,bool),X80)))) | p(s(bool,'const.bool.IN_2'(s(sK1,sK6(X80,sK3,X79,X78,sK1)),s(fun(sK1,bool),sK3)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X77),X76),s(fun(fun(sK0,X77),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X77,bool),X75)))))) ) | $spl94),
  inference(global_subsumption,[],[f67,f5791,f6437,f6498,f6439,f6499,f6447,f6448,f6451,f6452,f6453,f6454,f6455,f6456,f6472,f6473,f6479,f6483,f6486,f6487,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6793,f6807,f6816,f104,f102,f101,f100,f7026,f98,f96,f7027,f7591,f7593,f7594,f7596,f7598,f7599,f7600,f7609,f7623,f7667,f7670,f7632,f7673,f7646])).
fof(f7646,plain,(
  ( ! [X80,X78,X76,X79,X77,X75] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(X78,happ(s(fun(sK1,X78),X79),s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK6(X75,sK2,X76,X77,sK0)))))),s(fun(X78,bool),X80)))) | p(s(bool,'const.bool.IN_2'(s(sK1,sK6(X80,sK3,X79,X78,sK1)),s(fun(sK1,bool),sK3)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X77),X76),s(fun(fun(sK0,X77),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X77,bool),X75)))))) ) | $spl94),
  inference(superposition,[],[f6711,f5791])).
fof(f7673,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (p(s(bool,'const.bool.IN_2'(s(X15,happ(s(fun(fun(sK0,X12),X15),X16),s(fun(sK0,X12),X13))),s(fun(X15,bool),X17)))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,X12),sK6(X17,'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X12,bool),X14)),X16,X15,fun(sK0,X12))),s(fun(fun(sK0,X12),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X12,bool),X14)))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK6(X14,sK2,X13,X12,sK0)))),s(fun(sK1,bool),sK3)))) ) | $spl94),
  inference(global_subsumption,[],[f67,f5791,f6437,f6498,f6439,f6499,f6447,f6448,f6451,f6452,f6453,f6454,f6455,f6456,f6472,f6473,f6479,f6483,f6486,f6487,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6793,f6807,f6816,f104,f102,f101,f100,f7026,f98,f96,f7027,f7591,f7593,f7594,f7596,f7598,f7599,f7600,f7609,f7623,f7667,f7670,f7632])).
fof(f7632,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(X15,happ(s(fun(fun(sK0,X12),X15),X16),s(fun(sK0,X12),X13))),s(fun(X15,bool),X17)))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,X12),sK6(X17,'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X12,bool),X14)),X16,X15,fun(sK0,X12))),s(fun(fun(sK0,X12),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X12,bool),X14)))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK6(X14,sK2,X13,X12,sK0)))),s(fun(sK1,bool),sK3)))) ) | $spl94),
  inference(superposition,[],[f6711,f5791])).
fof(f7623,plain,(
  ( ! [X125,X123,X121,X124,X122,X120,X119] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X122)))) | p(s(bool,'const.bool.IN_2'(s(X123,happ(s(fun(bool,X123),X124),s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK6(X119,sK2,X120,X121,sK0)))),s(fun(sK1,bool),sK3))))),s(fun(X123,bool),X125)))) | p(s(bool,'const.bool.IN_2'(s(bool,sK6(X125,X122,X124,X123,bool)),s(fun(bool,bool),X122)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X121),X120),s(fun(fun(sK0,X121),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X121,bool),X119)))))) ) | $spl94),
  inference(superposition,[],[f6711,f5791])).
fof(f7609,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X45)))) | p(s(bool,'const.bool.IN_2'(s(X46,happ(s(fun(bool,X46),X47),s(bool,'const.bool.IN_2'(s(fun(sK0,X42),X43),s(fun(fun(sK0,X42),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X42,bool),X44))))))),s(fun(X46,bool),X48)))) | p(s(bool,'const.bool.IN_2'(s(bool,sK6(X48,X45,X47,X46,bool)),s(fun(bool,bool),X45)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK6(X44,sK2,X43,X42,sK0)))),s(fun(sK1,bool),sK3)))) ) | $spl94),
  inference(superposition,[],[f6711,f5791])).
fof(f6807,plain,(
  ( ! [X99,X97,X95,X100,X98,X96] : (p(s(bool,'const.bool.IN_2'(s(fun(X98,bool),'const.combin.K_1'(s(bool,'T_0'))),s(fun(fun(X98,bool),bool),'const.util_prob.FUNSET_2'(s(fun(X98,bool),X99),s(fun(bool,bool),X100)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK6(X95,sK2,X96,X97,sK0)))),s(fun(sK1,bool),sK3))),s(fun(bool,bool),X100)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X97),X96),s(fun(fun(sK0,X97),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X97,bool),X95)))))) ) | $spl94),
  inference(superposition,[],[f6706,f5791])).
fof(f6793,plain,(
  ( ! [X37,X35,X33,X36,X34,X32] : (p(s(bool,'const.bool.IN_2'(s(fun(X35,bool),'const.combin.K_1'(s(bool,'T_0'))),s(fun(fun(X35,bool),bool),'const.util_prob.FUNSET_2'(s(fun(X35,bool),X36),s(fun(bool,bool),X37)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(fun(sK0,X32),X33),s(fun(fun(sK0,X32),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X32,bool),X34))))),s(fun(bool,bool),X37)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK6(X34,sK2,X33,X32,sK0)))),s(fun(sK1,bool),sK3)))) ) | $spl94),
  inference(superposition,[],[f6706,f5791])).
fof(f6473,plain,(
  ( ! [X26,X24,X25] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,X24),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X24,bool),X25)),X26,fun(sK0,X24))),s(fun(fun(sK0,X24),bool),X26))) | s(fun(fun(sK0,X24),bool),X26) = s(fun(fun(sK0,X24),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X24,bool),X25))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK6(X25,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X24,bool),X25)),X26,fun(sK0,X24)),X24,sK0)))),s(fun(sK1,bool),sK3)))) ) | $spl94),
  inference(superposition,[],[f79,f5791])).
fof(f6472,plain,(
  ( ! [X23,X21,X22] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,X21),sK4(X22,'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X21,bool),X23)),fun(sK0,X21))),s(fun(fun(sK0,X21),bool),X22))) | s(fun(fun(sK0,X21),bool),X22) = s(fun(fun(sK0,X21),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X21,bool),X23))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK6(X23,sK2,sK4(X22,'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X21,bool),X23)),fun(sK0,X21)),X21,sK0)))),s(fun(sK1,bool),sK3)))) ) | $spl94),
  inference(superposition,[],[f79,f5791])).
fof(f6499,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X5),X9),s(fun(fun(sK0,X5),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X5,bool),X7))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK6(happ(s(fun(X6,fun(X5,bool)),'const.combin.K_1'(s(fun(X5,bool),X7))),s(X6,X8)),sK2,X9,X5,sK0)))),s(fun(sK1,bool),sK3)))) ) | $spl94),
  inference(forward_demodulation,[],[f6440,f80])).
fof(f6440,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X5),X9),s(fun(fun(sK0,X5),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X5,bool),'const.combin.K_2'(s(fun(X5,bool),X7),s(X6,X8))))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK6(happ(s(fun(X6,fun(X5,bool)),'const.combin.K_1'(s(fun(X5,bool),X7))),s(X6,X8)),sK2,X9,X5,sK0)))),s(fun(sK1,bool),sK3)))) ) | $spl94),
  inference(superposition,[],[f5791,f81])).
fof(f6439,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X0),X4),s(fun(fun(sK0,X0),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X0,bool),X1))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK6('const.combin.K_2'(s(fun(X0,bool),X1),s(X2,X3)),sK2,X4,X0,sK0)))),s(fun(sK1,bool),sK3)))) ) | $spl94),
  inference(superposition,[],[f5791,f80])).
fof(f6498,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X5),X7),s(fun(fun(sK0,X5),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X5,bool),X9))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK6(X9,sK2,happ(s(fun(X6,fun(sK0,X5)),'const.combin.K_1'(s(fun(sK0,X5),X7))),s(X6,X8)),X5,sK0)))),s(fun(sK1,bool),sK3)))) ) | $spl94),
  inference(forward_demodulation,[],[f6438,f80])).
fof(f6438,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X5),'const.combin.K_2'(s(fun(sK0,X5),X7),s(X6,X8))),s(fun(fun(sK0,X5),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X5,bool),X9))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK6(X9,sK2,happ(s(fun(X6,fun(sK0,X5)),'const.combin.K_1'(s(fun(sK0,X5),X7))),s(X6,X8)),X5,sK0)))),s(fun(sK1,bool),sK3)))) ) | $spl94),
  inference(superposition,[],[f5791,f81])).
fof(f6437,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X0),X1),s(fun(fun(sK0,X0),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X0,bool),X4))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK6(X4,sK2,'const.combin.K_2'(s(fun(sK0,X0),X1),s(X2,X3)),X0,sK0)))),s(fun(sK1,bool),sK3)))) ) | $spl94),
  inference(superposition,[],[f5791,f80])).
fof(f6612,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))))) | $spl88),
  inference(resolution,[],[f84,f5760])).
fof(f9069,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))))) | $spl94),
  inference(global_subsumption,[],[f67,f5791,f6437,f6498,f6439,f6499,f6447,f6448,f6451,f6452,f6453,f6454,f6455,f6456,f6472,f6473,f6479,f6483,f6486,f6487,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6793,f6807,f6816,f104,f102,f101,f100,f7026,f98,f96,f7027,f7591,f7593,f7594,f7596,f7598,f7599,f7600,f7609,f7623,f7667,f7670,f7632,f7673,f7646,f7686,f7687,f7688,f7689,f7691,f7698,f7699,f7702,f7695,f7703,f7704,f7705,f7706,f7710,f7711,f7712,f7766,f7769,f7718,f7770,f7779,f7780,f7781,f7785,f7786,f7787,f7837,f7838,f7841,f7858,f7859,f7860,f7863,f7864,f7918,f7923,f7924,f7925,f7926,f7928,f7935,f7936,f7939,f7932,f7940,f7941,f7942,f7943,f7947,f7948,f7949,f8003,f8006,f7955,f8007,f8020,f8021,f8022,f8024,f8025,f8026,f8078,f8067,f8083,f8195,f8196,f8198,f8203,f8204,f8205,f8255,f8259,f8277,f8278,f8279,f8325,f8331,f8332,f8333,f8336,f8338,f8397,f8401,f8402,f8403,f8450,f8451,f8453,f8462,f8465,f8457,f8466,f8467,f8470,f8471,f8472,f8474,f8485,f8488,f8478,f8489,f8490,f8493,f8494,f8495,f8496,f8497,f8499,f8500,f8501,f8502,f8504,f8505,f8506,f8507,f8508,f8555,f8558,f8543,f8559,f8565,f8566,f8567,f8568,f8570,f8571,f8572,f8573,f8575,f8576,f8577,f8578,f8579,f8624,f8614,f8630,f8637,f8638,f8639,f8640,f8642,f8643,f8644,f8645,f8646,f8647,f8648,f8688,f8689,f8690,f8691,f8693,f8694,f8695,f8696,f8698,f8699,f8700,f8701,f8702,f8703,f8705,f8706,f8707,f8708,f8710,f8711,f8712,f8713,f8714,f8715,f8716,f8717,f8718,f8727,f8741,f8788,f8750,f8791,f8764,f8804,f8805,f8807,f8808,f8809,f8810,f8858,f8859,f8861,f8862,f8863,f8864,f8869,f8870,f8871,f8872,f8874,f8931,f8966,f8969,f8936,f8937,f8946,f8960,f8963,f8989,f8990,f8991,f8992,f8993,f8995,f8996,f8997,f8998,f8999,f9001,f9002,f9003,f9004,f9005,f9022,f9025,f9008,f9009,f9013,f9031,f9060,f9061,f9064,f9065,f6700])).
fof(f6859,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))))) | ($spl88 | $spl122)),
  inference(duplicate_literal_removal,[],[f6843])).
fof(f6843,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))))) | ($spl88 | $spl122)),
  inference(resolution,[],[f6642,f5760])).
fof(f9068,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)))))) | $spl88),
  inference(forward_demodulation,[],[f7854,f80])).
fof(f7854,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),'const.combin.K_2'(s(fun(sK0,bool),sK2),s(X0,X1))),s(fun(sK1,bool),sK3)))))) ) | $spl88),
  inference(resolution,[],[f6637,f5760])).
fof(f9067,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)))))) | $spl88),
  inference(forward_demodulation,[],[f9066,f80])).
fof(f9066,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),'const.combin.K_2'(s(fun(sK0,bool),sK2),s(X0,X1))),s(fun(sK1,bool),sK3))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))))))) ) | $spl88),
  inference(forward_demodulation,[],[f8326,f81])).
fof(f8326,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),happ(s(fun(X0,fun(sK0,bool)),'const.combin.K_1'(s(fun(sK0,bool),sK2))),s(X0,X1))),s(fun(sK1,bool),sK3)))))) ) | $spl88),
  inference(resolution,[],[f6640,f5760])).
fof(f9065,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)))))) | $spl94),
  inference(global_subsumption,[],[f67,f5791,f6437,f6498,f6439,f6499,f6447,f6448,f6451,f6452,f6453,f6454,f6455,f6456,f6472,f6473,f6479,f6483,f6486,f6487,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6793,f6807,f6816,f104,f102,f101,f100,f7026,f98,f96,f7027,f7591,f7593,f7594,f7596,f7598,f7599,f7600,f7609,f7623,f7667,f7670,f7632,f7673,f7646,f7686,f7687,f7688,f7689,f7691,f7698,f7699,f7702,f7695,f7703,f7704,f7705,f7706,f7710,f7711,f7712,f7766,f7769,f7718,f7770,f7779,f7780,f7781,f7785,f7786,f7787,f7837,f7838,f7841,f7858,f7859,f7860,f7863,f7864,f7918,f7923,f7924,f7925,f7926,f7928,f7935,f7936,f7939,f7932,f7940,f7941,f7942,f7943,f7947,f7948,f7949,f8003,f8006,f7955,f8007,f8020,f8021,f8022,f8024,f8025,f8026,f8078,f8067,f8083,f8195,f8196,f8198,f8203,f8204,f8205,f8255,f8259,f8277,f8278,f8279,f8325,f8331,f8332,f8333,f8336,f8338,f8397,f8401,f8402,f8403,f8450,f8451,f8453,f8462,f8465,f8457,f8466,f8467,f8470,f8471,f8472,f8474,f8485,f8488,f8478,f8489,f8490,f8493,f8494,f8495,f8496,f8497,f8499,f8500,f8501,f8502,f8504,f8505,f8506,f8507,f8508,f8555,f8558,f8543,f8559,f8565,f8566,f8567,f8568,f8570,f8571,f8572,f8573,f8575,f8576,f8577,f8578,f8579,f8624,f8614,f8630,f8637,f8638,f8639,f8640,f8642,f8643,f8644,f8645,f8646,f8647,f8648,f8688,f8689,f8690,f8691,f8693,f8694,f8695,f8696,f8698,f8699,f8700,f8701,f8702,f8703,f8705,f8706,f8707,f8708,f8710,f8711,f8712,f8713,f8714,f8715,f8716,f8717,f8718,f8727,f8741,f8788,f8750,f8791,f8764,f8804,f8805,f8807,f8808,f8809,f8810,f8858,f8859,f8861,f8862,f8863,f8864,f8869,f8870,f8871,f8872,f8874,f8931,f8966,f8969,f8936,f8937,f8946,f8960,f8963,f8989,f8990,f8991,f8992,f8993,f8995,f8996,f8997,f8998,f8999,f9001,f9002,f9003,f9004,f9005,f9022,f9025,f9008,f9009,f9013,f9031,f9060,f9061,f9064])).
fof(f9061,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)))))) | $spl94),
  inference(global_subsumption,[],[f67,f5791,f6437,f6498,f6439,f6499,f6447,f6448,f6451,f6452,f6453,f6454,f6455,f6456,f6472,f6473,f6479,f6483,f6486,f6487,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6793,f6807,f6816,f104,f102,f101,f100,f7026,f98,f96,f7027,f7591,f7593,f7594,f7596,f7598,f7599,f7600,f7609,f7623,f7667,f7670,f7632,f7673,f7646,f7686,f7687,f7688,f7689,f7691,f7698,f7699,f7702,f7695,f7703,f7704,f7705,f7706,f7710,f7711,f7712,f7766,f7769,f7718,f7770,f7779,f7780,f7781,f7785,f7786,f7787,f7837,f7838,f7841,f7858,f7859,f7860,f7863,f7864,f7918,f7923,f7924,f7925,f7926,f7928,f7935,f7936,f7939,f7932,f7940,f7941,f7942,f7943,f7947,f7948,f7949,f8003,f8006,f7955,f8007,f8020,f8021,f8022,f8024,f8025,f8026,f8078,f8067,f8083,f8195,f8196,f8198,f8203,f8204,f8205,f8255,f8259,f8277,f8278,f8279,f8325,f8331,f8332,f8333,f8336,f8338,f8397,f8401,f8402,f8403,f8450,f8451,f8453,f8462,f8465,f8457,f8466,f8467,f8470,f8471,f8472,f8474,f8485,f8488,f8478,f8489,f8490,f8493,f8494,f8495,f8496,f8497,f8499,f8500,f8501,f8502,f8504,f8505,f8506,f8507,f8508,f8555,f8558,f8543,f8559,f8565,f8566,f8567,f8568,f8570,f8571,f8572,f8573,f8575,f8576,f8577,f8578,f8579,f8624,f8614,f8630,f8637,f8638,f8639,f8640,f8642,f8643,f8644,f8645,f8646,f8647,f8648,f8688,f8689,f8690,f8691,f8693,f8694,f8695,f8696,f8698,f8699,f8700,f8701,f8702,f8703,f8705,f8706,f8707,f8708,f8710,f8711,f8712,f8713,f8714,f8715,f8716,f8717,f8718,f8727,f8741,f8788,f8750,f8791,f8764,f8804,f8805,f8807,f8808,f8809,f8810,f8858,f8859,f8861,f8862,f8863,f8864,f8869,f8870,f8871,f8872,f8874,f8931,f8966,f8969,f8936,f8937,f8946,f8960,f8963,f8989,f8990,f8991,f8992,f8993,f8995,f8996,f8997,f8998,f8999,f9001,f9002,f9003,f9004,f9005,f9022,f9025,f9008,f9009,f9013,f9031,f9060])).
fof(f9058,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))))))) | $spl168),
  inference(cnf_transformation,[],[f9058_D])).
fof(f9058_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))))))) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f9059,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))))))) | $spl80),
  inference(global_subsumption,[],[f67,f5733,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6816,f104,f102,f101,f100,f7026,f98,f96,f7027,f7591,f7593,f7594,f7596,f7598,f7599,f7600,f7667,f7670,f7639,f7679,f7687,f7688,f7689,f7691,f7698,f7699,f7702,f7695,f7703,f7704,f7705,f7706,f7710,f7711,f7712,f7766,f7769,f7718,f7770,f7779,f7780,f7781,f7785,f7786,f7787,f7837,f7838,f7841,f7858,f7859,f7860,f7863,f7864,f7918,f7923,f7924,f7925,f7926,f7928,f7935,f7936,f7939,f7932,f7940,f7941,f7942,f7943,f7947,f7948,f7949,f8003,f8006,f7955,f8007,f8020,f8021,f8022,f8024,f8025,f8026,f8078,f8067,f8083,f8195,f8196,f8198,f8203,f8204,f8205,f8255,f8259,f8277,f8278,f8279,f8325,f8331,f8332,f8333,f8336,f8338,f8397,f8401,f8402,f8403,f8450,f8451,f8453,f8462,f8465,f8457,f8466,f8467,f8470,f8471,f8472,f8474,f8485,f8488,f8478,f8489,f8490,f8493,f8494,f8495,f8496,f8497,f8499,f8500,f8501,f8502,f8504,f8505,f8506,f8507,f8508,f8555,f8558,f8565,f8566,f8567,f8568,f8570,f8571,f8572,f8573,f8575,f8576,f8577,f8578,f8579,f8624,f8637,f8638,f8639,f8640,f8642,f8643,f8644,f8645,f8646,f8647,f8648,f8688,f8689,f8690,f8691,f8693,f8694,f8695,f8696,f8698,f8699,f8700,f8701,f8702,f8703,f8705,f8706,f8707,f8708,f8710,f8711,f8712,f8713,f8714,f8715,f8716,f8717,f8718,f8788,f8757,f8797,f8805,f8807,f8808,f8809,f8810,f8858,f8859,f8861,f8862,f8863,f8864,f8869,f8870,f8871,f8872,f8874,f8931,f8966,f8969,f8936,f8937,f8963,f8989,f8990,f8991,f8992,f8993,f8995,f8996,f8997,f8998,f8999,f9001,f9002,f9003,f9004,f9005,f9022,f9025,f9008,f9009,f9017])).
fof(f9017,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))))))) | $spl80),
  inference(superposition,[],[f6755,f5733])).
fof(f8797,plain,(
  ( ! [X47,X45,X46] : (p(s(bool,'const.bool.IN_2'(s(X45,happ(s(fun(sK1,X45),X46),s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)))))),s(fun(X45,bool),happ(s(fun(sK1,fun(X45,bool)),X47),s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1))))))))) | p(s(bool,'const.bool.IN_2'(s(sK1,sK7(X47,sK3,X46,sK1,X45)),s(fun(sK1,bool),sK3))))) ) | $spl80),
  inference(global_subsumption,[],[f67,f5733,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6816,f104,f102,f101,f100,f7026,f98,f96,f7027,f7591,f7593,f7594,f7596,f7598,f7599,f7600,f7667,f7670,f7639,f7679,f7687,f7688,f7689,f7691,f7698,f7699,f7702,f7695,f7703,f7704,f7705,f7706,f7710,f7711,f7712,f7766,f7769,f7718,f7770,f7779,f7780,f7781,f7785,f7786,f7787,f7837,f7838,f7841,f7858,f7859,f7860,f7863,f7864,f7918,f7923,f7924,f7925,f7926,f7928,f7935,f7936,f7939,f7932,f7940,f7941,f7942,f7943,f7947,f7948,f7949,f8003,f8006,f7955,f8007,f8020,f8021,f8022,f8024,f8025,f8026,f8078,f8067,f8083,f8195,f8196,f8198,f8203,f8204,f8205,f8255,f8259,f8277,f8278,f8279,f8325,f8331,f8332,f8333,f8336,f8338,f8397,f8401,f8402,f8403,f8450,f8451,f8453,f8462,f8465,f8457,f8466,f8467,f8470,f8471,f8472,f8474,f8485,f8488,f8478,f8489,f8490,f8493,f8494,f8495,f8496,f8497,f8499,f8500,f8501,f8502,f8504,f8505,f8506,f8507,f8508,f8555,f8558,f8565,f8566,f8567,f8568,f8570,f8571,f8572,f8573,f8575,f8576,f8577,f8578,f8579,f8624,f8637,f8638,f8639,f8640,f8642,f8643,f8644,f8645,f8646,f8647,f8648,f8688,f8689,f8690,f8691,f8693,f8694,f8695,f8696,f8698,f8699,f8700,f8701,f8702,f8703,f8705,f8706,f8707,f8708,f8710,f8711,f8712,f8713,f8714,f8715,f8716,f8717,f8718,f8788,f8757])).
fof(f8757,plain,(
  ( ! [X47,X45,X46] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(X45,happ(s(fun(sK1,X45),X46),s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)))))),s(fun(X45,bool),happ(s(fun(sK1,fun(X45,bool)),X47),s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1))))))))) | p(s(bool,'const.bool.IN_2'(s(sK1,sK7(X47,sK3,X46,sK1,X45)),s(fun(sK1,bool),sK3))))) ) | $spl80),
  inference(superposition,[],[f6756,f5733])).
fof(f7679,plain,(
  ( ! [X47,X45,X46] : (p(s(bool,'const.bool.IN_2'(s(X45,happ(s(fun(sK1,X45),X46),s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)))))),s(fun(X45,bool),X47)))) | p(s(bool,'const.bool.IN_2'(s(sK1,sK6(X47,sK3,X46,X45,sK1)),s(fun(sK1,bool),sK3))))) ) | $spl80),
  inference(global_subsumption,[],[f67,f5733,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6816,f104,f102,f101,f100,f7026,f98,f96,f7027,f7591,f7593,f7594,f7596,f7598,f7599,f7600,f7667,f7670,f7639])).
fof(f7639,plain,(
  ( ! [X47,X45,X46] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(X45,happ(s(fun(sK1,X45),X46),s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)))))),s(fun(X45,bool),X47)))) | p(s(bool,'const.bool.IN_2'(s(sK1,sK6(X47,sK3,X46,X45,sK1)),s(fun(sK1,bool),sK3))))) ) | $spl80),
  inference(superposition,[],[f6711,f5733])).
fof(f133,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3))))))) | $spl18),
  inference(cnf_transformation,[],[f133_D])).
fof(f133_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3))))))) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f9056,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3))))))) | $spl96),
  inference(global_subsumption,[],[f67,f5803,f6506,f6531,f6508,f6532,f6522,f6523,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6796,f6804,f6816,f104,f102,f101,f100,f7026,f98,f96,f7027,f7591,f7593,f7594,f7596,f7598,f7599,f7600,f7612,f7620,f7667,f7670,f7635,f7676,f7643,f7683,f7687,f7688,f7689,f7691,f7698,f7699,f7702,f7695,f7703,f7704,f7705,f7706,f7710,f7711,f7712,f7766,f7769,f7718,f7770,f7779,f7780,f7781,f7785,f7786,f7787,f7837,f7838,f7841,f7858,f7859,f7860,f7863,f7864,f7918,f7923,f7924,f7925,f7926,f7928,f7935,f7936,f7939,f7932,f7940,f7941,f7942,f7943,f7947,f7948,f7949,f8003,f8006,f7955,f8007,f8020,f8021,f8022,f8024,f8025,f8026,f8078,f8067,f8083,f8195,f8196,f8198,f8203,f8204,f8205,f8255,f8259,f8277,f8278,f8279,f8325,f8331,f8332,f8333,f8336,f8338,f8397,f8401,f8402,f8403,f8450,f8451,f8453,f8462,f8465,f8457,f8466,f8467,f8470,f8471,f8472,f8474,f8485,f8488,f8478,f8489,f8490,f8493,f8494,f8495,f8496,f8497,f8499,f8500,f8501,f8502,f8504,f8505,f8506,f8507,f8508,f8555,f8558,f8546,f8562,f8565,f8566,f8567,f8568,f8570,f8571,f8572,f8573,f8575,f8576,f8577,f8578,f8579,f8624,f8617,f8633,f8637,f8638,f8639,f8640,f8642,f8643,f8644,f8645,f8646,f8647,f8648,f8688,f8689,f8690,f8691,f8693,f8694,f8695,f8696,f8698,f8699,f8700,f8701,f8702,f8703,f8705,f8706,f8707,f8708,f8710,f8711,f8712,f8713,f8714,f8715,f8716,f8717,f8718,f8730,f8738,f8788,f8753,f8794,f8761,f8801,f8805,f8807,f8808,f8809,f8810,f8858,f8859,f8861,f8862,f8863,f8864,f8869,f8870,f8871,f8872,f8874,f8931,f8966,f8969,f8936,f8937,f8949,f8957,f8963,f8989,f8990,f8991,f8992,f8993,f8995,f8996,f8997,f8998,f8999,f9001,f9002,f9003,f9004,f9005,f9022,f9025,f9008,f9009,f9016,f9034,f9019])).
fof(f9019,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3))))))) | $spl96),
  inference(superposition,[],[f6755,f5803])).
fof(f9034,plain,(
  ( ! [X19,X17,X15,X18,X16] : (p(s(bool,'const.bool.IN_2'(s(fun(X16,fun(sK0,X15)),X17),s(fun(fun(X16,fun(sK0,X15)),bool),'const.util_prob.DFUNSET_2'(s(fun(X16,bool),X19),s(fun(X16,fun(fun(sK0,X15),bool)),'const.combin.K_1'(s(fun(fun(sK0,X15),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X15,bool)),X18)))))))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK7(X18,sK2,happ(s(fun(X16,fun(sK0,X15)),X17),s(X16,sK7('const.combin.K_1'(s(fun(fun(sK0,X15),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X15,bool)),X18)))),X19,X17,X16,fun(sK0,X15)))),sK0,X15)))),s(fun(sK1,bool),sK3)))) ) | $spl96),
  inference(global_subsumption,[],[f67,f5803,f6506,f6531,f6508,f6532,f6522,f6523,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6796,f6804,f6816,f104,f102,f101,f100,f7026,f98,f96,f7027,f7591,f7593,f7594,f7596,f7598,f7599,f7600,f7612,f7620,f7667,f7670,f7635,f7676,f7643,f7683,f7687,f7688,f7689,f7691,f7698,f7699,f7702,f7695,f7703,f7704,f7705,f7706,f7710,f7711,f7712,f7766,f7769,f7718,f7770,f7779,f7780,f7781,f7785,f7786,f7787,f7837,f7838,f7841,f7858,f7859,f7860,f7863,f7864,f7918,f7923,f7924,f7925,f7926,f7928,f7935,f7936,f7939,f7932,f7940,f7941,f7942,f7943,f7947,f7948,f7949,f8003,f8006,f7955,f8007,f8020,f8021,f8022,f8024,f8025,f8026,f8078,f8067,f8083,f8195,f8196,f8198,f8203,f8204,f8205,f8255,f8259,f8277,f8278,f8279,f8325,f8331,f8332,f8333,f8336,f8338,f8397,f8401,f8402,f8403,f8450,f8451,f8453,f8462,f8465,f8457,f8466,f8467,f8470,f8471,f8472,f8474,f8485,f8488,f8478,f8489,f8490,f8493,f8494,f8495,f8496,f8497,f8499,f8500,f8501,f8502,f8504,f8505,f8506,f8507,f8508,f8555,f8558,f8546,f8562,f8565,f8566,f8567,f8568,f8570,f8571,f8572,f8573,f8575,f8576,f8577,f8578,f8579,f8624,f8617,f8633,f8637,f8638,f8639,f8640,f8642,f8643,f8644,f8645,f8646,f8647,f8648,f8688,f8689,f8690,f8691,f8693,f8694,f8695,f8696,f8698,f8699,f8700,f8701,f8702,f8703,f8705,f8706,f8707,f8708,f8710,f8711,f8712,f8713,f8714,f8715,f8716,f8717,f8718,f8730,f8738,f8788,f8753,f8794,f8761,f8801,f8805,f8807,f8808,f8809,f8810,f8858,f8859,f8861,f8862,f8863,f8864,f8869,f8870,f8871,f8872,f8874,f8931,f8966,f8969,f8936,f8937,f8949,f8957,f8963,f8989,f8990,f8991,f8992,f8993,f8995,f8996,f8997,f8998,f8999,f9001,f9002,f9003,f9004,f9005,f9022,f9025,f9008,f9009,f9016])).
fof(f9016,plain,(
  ( ! [X19,X17,X15,X18,X16] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(X16,fun(sK0,X15)),X17),s(fun(fun(X16,fun(sK0,X15)),bool),'const.util_prob.DFUNSET_2'(s(fun(X16,bool),X19),s(fun(X16,fun(fun(sK0,X15),bool)),'const.combin.K_1'(s(fun(fun(sK0,X15),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X15,bool)),X18)))))))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK7(X18,sK2,happ(s(fun(X16,fun(sK0,X15)),X17),s(X16,sK7('const.combin.K_1'(s(fun(fun(sK0,X15),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X15,bool)),X18)))),X19,X17,X16,fun(sK0,X15)))),sK0,X15)))),s(fun(sK1,bool),sK3)))) ) | $spl96),
  inference(superposition,[],[f6755,f5803])).
fof(f8957,plain,(
  ( ! [X80,X78,X81,X79,X77,X82] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),happ(s(fun(X80,fun(bool,bool)),X81),s(X80,sK7(X81,X82,'const.combin.K_1'(s(bool,'T_0')),X80,bool))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X80,bool),'const.combin.K_1'(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK7(X77,sK2,X78,sK0,X79)))),s(fun(sK1,bool),sK3))))),s(fun(fun(X80,bool),bool),'const.util_prob.DFUNSET_2'(s(fun(X80,bool),X82),s(fun(X80,fun(bool,bool)),X81)))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X79),X78),s(fun(fun(sK0,X79),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X79,bool)),X77)))))) ) | $spl96),
  inference(superposition,[],[f6750,f5803])).
fof(f8949,plain,(
  ( ! [X54,X52,X50,X55,X53,X51] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),happ(s(fun(X53,fun(bool,bool)),X54),s(X53,sK7(X54,X55,'const.combin.K_1'(s(bool,'T_0')),X53,bool))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X53,bool),'const.combin.K_1'(s(bool,'const.bool.IN_2'(s(fun(sK0,X50),X51),s(fun(fun(sK0,X50),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X50,bool)),X52))))))),s(fun(fun(X53,bool),bool),'const.util_prob.DFUNSET_2'(s(fun(X53,bool),X55),s(fun(X53,fun(bool,bool)),X54)))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK7(X52,sK2,X51,sK0,X50)))),s(fun(sK1,bool),sK3)))) ) | $spl96),
  inference(superposition,[],[f6750,f5803])).
fof(f8801,plain,(
  ( ! [X61,X59,X57,X62,X60,X58] : (p(s(bool,'const.bool.IN_2'(s(X60,happ(s(fun(sK1,X60),X61),s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK7(X57,sK2,X58,sK0,X59)))))),s(fun(X60,bool),happ(s(fun(sK1,fun(X60,bool)),X62),s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK7(X57,sK2,X58,sK0,X59))))))))) | p(s(bool,'const.bool.IN_2'(s(sK1,sK7(X62,sK3,X61,sK1,X60)),s(fun(sK1,bool),sK3)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X59),X58),s(fun(fun(sK0,X59),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X59,bool)),X57)))))) ) | $spl96),
  inference(global_subsumption,[],[f67,f5803,f6506,f6531,f6508,f6532,f6522,f6523,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6796,f6804,f6816,f104,f102,f101,f100,f7026,f98,f96,f7027,f7591,f7593,f7594,f7596,f7598,f7599,f7600,f7612,f7620,f7667,f7670,f7635,f7676,f7643,f7683,f7687,f7688,f7689,f7691,f7698,f7699,f7702,f7695,f7703,f7704,f7705,f7706,f7710,f7711,f7712,f7766,f7769,f7718,f7770,f7779,f7780,f7781,f7785,f7786,f7787,f7837,f7838,f7841,f7858,f7859,f7860,f7863,f7864,f7918,f7923,f7924,f7925,f7926,f7928,f7935,f7936,f7939,f7932,f7940,f7941,f7942,f7943,f7947,f7948,f7949,f8003,f8006,f7955,f8007,f8020,f8021,f8022,f8024,f8025,f8026,f8078,f8067,f8083,f8195,f8196,f8198,f8203,f8204,f8205,f8255,f8259,f8277,f8278,f8279,f8325,f8331,f8332,f8333,f8336,f8338,f8397,f8401,f8402,f8403,f8450,f8451,f8453,f8462,f8465,f8457,f8466,f8467,f8470,f8471,f8472,f8474,f8485,f8488,f8478,f8489,f8490,f8493,f8494,f8495,f8496,f8497,f8499,f8500,f8501,f8502,f8504,f8505,f8506,f8507,f8508,f8555,f8558,f8546,f8562,f8565,f8566,f8567,f8568,f8570,f8571,f8572,f8573,f8575,f8576,f8577,f8578,f8579,f8624,f8617,f8633,f8637,f8638,f8639,f8640,f8642,f8643,f8644,f8645,f8646,f8647,f8648,f8688,f8689,f8690,f8691,f8693,f8694,f8695,f8696,f8698,f8699,f8700,f8701,f8702,f8703,f8705,f8706,f8707,f8708,f8710,f8711,f8712,f8713,f8714,f8715,f8716,f8717,f8718,f8730,f8738,f8788,f8753,f8794,f8761])).
fof(f8761,plain,(
  ( ! [X61,X59,X57,X62,X60,X58] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(X60,happ(s(fun(sK1,X60),X61),s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK7(X57,sK2,X58,sK0,X59)))))),s(fun(X60,bool),happ(s(fun(sK1,fun(X60,bool)),X62),s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK7(X57,sK2,X58,sK0,X59))))))))) | p(s(bool,'const.bool.IN_2'(s(sK1,sK7(X62,sK3,X61,sK1,X60)),s(fun(sK1,bool),sK3)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X59),X58),s(fun(fun(sK0,X59),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X59,bool)),X57)))))) ) | $spl96),
  inference(superposition,[],[f6756,f5803])).
fof(f8794,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (p(s(bool,'const.bool.IN_2'(s(X33,happ(s(fun(fun(sK0,X30),X33),X34),s(fun(sK0,X30),X31))),s(fun(X33,bool),happ(s(fun(fun(sK0,X30),fun(X33,bool)),X35),s(fun(sK0,X30),X31)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,X30),sK7(X35,'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X30,bool)),X32)),X34,fun(sK0,X30),X33)),s(fun(fun(sK0,X30),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X30,bool)),X32)))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK7(X32,sK2,X31,sK0,X30)))),s(fun(sK1,bool),sK3)))) ) | $spl96),
  inference(global_subsumption,[],[f67,f5803,f6506,f6531,f6508,f6532,f6522,f6523,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6796,f6804,f6816,f104,f102,f101,f100,f7026,f98,f96,f7027,f7591,f7593,f7594,f7596,f7598,f7599,f7600,f7612,f7620,f7667,f7670,f7635,f7676,f7643,f7683,f7687,f7688,f7689,f7691,f7698,f7699,f7702,f7695,f7703,f7704,f7705,f7706,f7710,f7711,f7712,f7766,f7769,f7718,f7770,f7779,f7780,f7781,f7785,f7786,f7787,f7837,f7838,f7841,f7858,f7859,f7860,f7863,f7864,f7918,f7923,f7924,f7925,f7926,f7928,f7935,f7936,f7939,f7932,f7940,f7941,f7942,f7943,f7947,f7948,f7949,f8003,f8006,f7955,f8007,f8020,f8021,f8022,f8024,f8025,f8026,f8078,f8067,f8083,f8195,f8196,f8198,f8203,f8204,f8205,f8255,f8259,f8277,f8278,f8279,f8325,f8331,f8332,f8333,f8336,f8338,f8397,f8401,f8402,f8403,f8450,f8451,f8453,f8462,f8465,f8457,f8466,f8467,f8470,f8471,f8472,f8474,f8485,f8488,f8478,f8489,f8490,f8493,f8494,f8495,f8496,f8497,f8499,f8500,f8501,f8502,f8504,f8505,f8506,f8507,f8508,f8555,f8558,f8546,f8562,f8565,f8566,f8567,f8568,f8570,f8571,f8572,f8573,f8575,f8576,f8577,f8578,f8579,f8624,f8617,f8633,f8637,f8638,f8639,f8640,f8642,f8643,f8644,f8645,f8646,f8647,f8648,f8688,f8689,f8690,f8691,f8693,f8694,f8695,f8696,f8698,f8699,f8700,f8701,f8702,f8703,f8705,f8706,f8707,f8708,f8710,f8711,f8712,f8713,f8714,f8715,f8716,f8717,f8718,f8730,f8738,f8788,f8753])).
fof(f8753,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(X33,happ(s(fun(fun(sK0,X30),X33),X34),s(fun(sK0,X30),X31))),s(fun(X33,bool),happ(s(fun(fun(sK0,X30),fun(X33,bool)),X35),s(fun(sK0,X30),X31)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,X30),sK7(X35,'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X30,bool)),X32)),X34,fun(sK0,X30),X33)),s(fun(fun(sK0,X30),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X30,bool)),X32)))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK7(X32,sK2,X31,sK0,X30)))),s(fun(sK1,bool),sK3)))) ) | $spl96),
  inference(superposition,[],[f6756,f5803])).
fof(f8738,plain,(
  ( ! [X103,X101,X99,X102,X100,X98,X104] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X101)))) | p(s(bool,'const.bool.IN_2'(s(X102,happ(s(fun(bool,X102),X103),s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK7(X98,sK2,X99,sK0,X100)))),s(fun(sK1,bool),sK3))))),s(fun(X102,bool),happ(s(fun(bool,fun(X102,bool)),X104),s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK7(X98,sK2,X99,sK0,X100)))),s(fun(sK1,bool),sK3)))))))) | p(s(bool,'const.bool.IN_2'(s(bool,sK7(X104,X101,X103,bool,X102)),s(fun(bool,bool),X101)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X100),X99),s(fun(fun(sK0,X100),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X100,bool)),X98)))))) ) | $spl96),
  inference(superposition,[],[f6756,f5803])).
fof(f8730,plain,(
  ( ! [X68,X66,X64,X69,X67,X65,X63] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X66)))) | p(s(bool,'const.bool.IN_2'(s(X67,happ(s(fun(bool,X67),X68),s(bool,'const.bool.IN_2'(s(fun(sK0,X63),X64),s(fun(fun(sK0,X63),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X63,bool)),X65))))))),s(fun(X67,bool),happ(s(fun(bool,fun(X67,bool)),X69),s(bool,'const.bool.IN_2'(s(fun(sK0,X63),X64),s(fun(fun(sK0,X63),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X63,bool)),X65)))))))))) | p(s(bool,'const.bool.IN_2'(s(bool,sK7(X69,X66,X68,bool,X67)),s(fun(bool,bool),X66)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK7(X65,sK2,X64,sK0,X63)))),s(fun(sK1,bool),sK3)))) ) | $spl96),
  inference(superposition,[],[f6756,f5803])).
fof(f8633,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (p(s(bool,'const.bool.IN_2'(s(fun(X22,fun(sK0,X21)),X23),s(fun(fun(X22,fun(sK0,X21)),bool),'const.util_prob.FUNSET_2'(s(fun(X22,bool),X27),s(fun(fun(sK0,X21),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X21,bool)),X24)))))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK7(X24,sK2,happ(s(fun(X22,fun(sK0,X21)),X23),s(X22,sK6('const.combin.K_2'(s(fun(fun(sK0,X21),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X21,bool)),X24))),s(X25,X26)),X27,X23,fun(sK0,X21),X22))),sK0,X21)))),s(fun(sK1,bool),sK3)))) ) | $spl96),
  inference(global_subsumption,[],[f67,f5803,f6506,f6531,f6508,f6532,f6522,f6523,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6796,f6804,f6816,f104,f102,f101,f100,f7026,f98,f96,f7027,f7591,f7593,f7594,f7596,f7598,f7599,f7600,f7612,f7620,f7667,f7670,f7635,f7676,f7643,f7683,f7687,f7688,f7689,f7691,f7698,f7699,f7702,f7695,f7703,f7704,f7705,f7706,f7710,f7711,f7712,f7766,f7769,f7718,f7770,f7779,f7780,f7781,f7785,f7786,f7787,f7837,f7838,f7841,f7858,f7859,f7860,f7863,f7864,f7918,f7923,f7924,f7925,f7926,f7928,f7935,f7936,f7939,f7932,f7940,f7941,f7942,f7943,f7947,f7948,f7949,f8003,f8006,f7955,f8007,f8020,f8021,f8022,f8024,f8025,f8026,f8078,f8067,f8083,f8195,f8196,f8198,f8203,f8204,f8205,f8255,f8259,f8277,f8278,f8279,f8325,f8331,f8332,f8333,f8336,f8338,f8397,f8401,f8402,f8403,f8450,f8451,f8453,f8462,f8465,f8457,f8466,f8467,f8470,f8471,f8472,f8474,f8485,f8488,f8478,f8489,f8490,f8493,f8494,f8495,f8496,f8497,f8499,f8500,f8501,f8502,f8504,f8505,f8506,f8507,f8508,f8555,f8558,f8546,f8562,f8565,f8566,f8567,f8568,f8570,f8571,f8572,f8573,f8575,f8576,f8577,f8578,f8579,f8624,f8617])).
fof(f8617,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(X22,fun(sK0,X21)),X23),s(fun(fun(X22,fun(sK0,X21)),bool),'const.util_prob.FUNSET_2'(s(fun(X22,bool),X27),s(fun(fun(sK0,X21),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X21,bool)),X24)))))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK7(X24,sK2,happ(s(fun(X22,fun(sK0,X21)),X23),s(X22,sK6('const.combin.K_2'(s(fun(fun(sK0,X21),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X21,bool)),X24))),s(X25,X26)),X27,X23,fun(sK0,X21),X22))),sK0,X21)))),s(fun(sK1,bool),sK3)))) ) | $spl96),
  inference(superposition,[],[f6707,f5803])).
fof(f8562,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (p(s(bool,'const.bool.IN_2'(s(fun(X22,fun(sK0,X21)),X23),s(fun(fun(X22,fun(sK0,X21)),bool),'const.util_prob.FUNSET_2'(s(fun(X22,bool),X25),s(fun(fun(sK0,X21),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X21,bool)),X24)))))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK7(X24,sK2,happ(s(fun(X22,fun(sK0,X21)),X23),s(X22,sK6('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X21,bool)),X24)),X25,'const.combin.K_2'(s(fun(X22,fun(sK0,X21)),X23),s(X26,X27)),fun(sK0,X21),X22))),sK0,X21)))),s(fun(sK1,bool),sK3)))) ) | $spl96),
  inference(global_subsumption,[],[f67,f5803,f6506,f6531,f6508,f6532,f6522,f6523,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6796,f6804,f6816,f104,f102,f101,f100,f7026,f98,f96,f7027,f7591,f7593,f7594,f7596,f7598,f7599,f7600,f7612,f7620,f7667,f7670,f7635,f7676,f7643,f7683,f7687,f7688,f7689,f7691,f7698,f7699,f7702,f7695,f7703,f7704,f7705,f7706,f7710,f7711,f7712,f7766,f7769,f7718,f7770,f7779,f7780,f7781,f7785,f7786,f7787,f7837,f7838,f7841,f7858,f7859,f7860,f7863,f7864,f7918,f7923,f7924,f7925,f7926,f7928,f7935,f7936,f7939,f7932,f7940,f7941,f7942,f7943,f7947,f7948,f7949,f8003,f8006,f7955,f8007,f8020,f8021,f8022,f8024,f8025,f8026,f8078,f8067,f8083,f8195,f8196,f8198,f8203,f8204,f8205,f8255,f8259,f8277,f8278,f8279,f8325,f8331,f8332,f8333,f8336,f8338,f8397,f8401,f8402,f8403,f8450,f8451,f8453,f8462,f8465,f8457,f8466,f8467,f8470,f8471,f8472,f8474,f8485,f8488,f8478,f8489,f8490,f8493,f8494,f8495,f8496,f8497,f8499,f8500,f8501,f8502,f8504,f8505,f8506,f8507,f8508,f8555,f8558,f8546])).
fof(f8546,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(X22,fun(sK0,X21)),X23),s(fun(fun(X22,fun(sK0,X21)),bool),'const.util_prob.FUNSET_2'(s(fun(X22,bool),X25),s(fun(fun(sK0,X21),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X21,bool)),X24)))))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK7(X24,sK2,happ(s(fun(X22,fun(sK0,X21)),X23),s(X22,sK6('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X21,bool)),X24)),X25,'const.combin.K_2'(s(fun(X22,fun(sK0,X21)),X23),s(X26,X27)),fun(sK0,X21),X22))),sK0,X21)))),s(fun(sK1,bool),sK3)))) ) | $spl96),
  inference(superposition,[],[f6702,f5803])).
fof(f7683,plain,(
  ( ! [X61,X59,X57,X62,X60,X58] : (p(s(bool,'const.bool.IN_2'(s(X60,happ(s(fun(sK1,X60),X61),s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK7(X57,sK2,X58,sK0,X59)))))),s(fun(X60,bool),X62)))) | p(s(bool,'const.bool.IN_2'(s(sK1,sK6(X62,sK3,X61,X60,sK1)),s(fun(sK1,bool),sK3)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X59),X58),s(fun(fun(sK0,X59),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X59,bool)),X57)))))) ) | $spl96),
  inference(global_subsumption,[],[f67,f5803,f6506,f6531,f6508,f6532,f6522,f6523,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6796,f6804,f6816,f104,f102,f101,f100,f7026,f98,f96,f7027,f7591,f7593,f7594,f7596,f7598,f7599,f7600,f7612,f7620,f7667,f7670,f7635,f7676,f7643])).
fof(f7643,plain,(
  ( ! [X61,X59,X57,X62,X60,X58] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(X60,happ(s(fun(sK1,X60),X61),s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK7(X57,sK2,X58,sK0,X59)))))),s(fun(X60,bool),X62)))) | p(s(bool,'const.bool.IN_2'(s(sK1,sK6(X62,sK3,X61,X60,sK1)),s(fun(sK1,bool),sK3)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X59),X58),s(fun(fun(sK0,X59),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X59,bool)),X57)))))) ) | $spl96),
  inference(superposition,[],[f6711,f5803])).
fof(f7676,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (p(s(bool,'const.bool.IN_2'(s(X33,happ(s(fun(fun(sK0,X30),X33),X34),s(fun(sK0,X30),X31))),s(fun(X33,bool),X35)))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,X30),sK6(X35,'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X30,bool)),X32)),X34,X33,fun(sK0,X30))),s(fun(fun(sK0,X30),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X30,bool)),X32)))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK7(X32,sK2,X31,sK0,X30)))),s(fun(sK1,bool),sK3)))) ) | $spl96),
  inference(global_subsumption,[],[f67,f5803,f6506,f6531,f6508,f6532,f6522,f6523,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6796,f6804,f6816,f104,f102,f101,f100,f7026,f98,f96,f7027,f7591,f7593,f7594,f7596,f7598,f7599,f7600,f7612,f7620,f7667,f7670,f7635])).
fof(f7635,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(X33,happ(s(fun(fun(sK0,X30),X33),X34),s(fun(sK0,X30),X31))),s(fun(X33,bool),X35)))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,X30),sK6(X35,'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X30,bool)),X32)),X34,X33,fun(sK0,X30))),s(fun(fun(sK0,X30),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X30,bool)),X32)))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK7(X32,sK2,X31,sK0,X30)))),s(fun(sK1,bool),sK3)))) ) | $spl96),
  inference(superposition,[],[f6711,f5803])).
fof(f7620,plain,(
  ( ! [X103,X101,X99,X102,X100,X98,X104] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X101)))) | p(s(bool,'const.bool.IN_2'(s(X102,happ(s(fun(bool,X102),X103),s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK7(X98,sK2,X99,sK0,X100)))),s(fun(sK1,bool),sK3))))),s(fun(X102,bool),X104)))) | p(s(bool,'const.bool.IN_2'(s(bool,sK6(X104,X101,X103,X102,bool)),s(fun(bool,bool),X101)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X100),X99),s(fun(fun(sK0,X100),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X100,bool)),X98)))))) ) | $spl96),
  inference(superposition,[],[f6711,f5803])).
fof(f7612,plain,(
  ( ! [X68,X66,X64,X69,X67,X65,X63] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X66)))) | p(s(bool,'const.bool.IN_2'(s(X67,happ(s(fun(bool,X67),X68),s(bool,'const.bool.IN_2'(s(fun(sK0,X63),X64),s(fun(fun(sK0,X63),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X63,bool)),X65))))))),s(fun(X67,bool),X69)))) | p(s(bool,'const.bool.IN_2'(s(bool,sK6(X69,X66,X68,X67,bool)),s(fun(bool,bool),X66)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK7(X65,sK2,X64,sK0,X63)))),s(fun(sK1,bool),sK3)))) ) | $spl96),
  inference(superposition,[],[f6711,f5803])).
fof(f6804,plain,(
  ( ! [X80,X78,X81,X79,X77,X82] : (p(s(bool,'const.bool.IN_2'(s(fun(X80,bool),'const.combin.K_1'(s(bool,'T_0'))),s(fun(fun(X80,bool),bool),'const.util_prob.FUNSET_2'(s(fun(X80,bool),X81),s(fun(bool,bool),X82)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK7(X77,sK2,X78,sK0,X79)))),s(fun(sK1,bool),sK3))),s(fun(bool,bool),X82)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X79),X78),s(fun(fun(sK0,X79),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X79,bool)),X77)))))) ) | $spl96),
  inference(superposition,[],[f6706,f5803])).
fof(f6796,plain,(
  ( ! [X54,X52,X50,X55,X53,X51] : (p(s(bool,'const.bool.IN_2'(s(fun(X53,bool),'const.combin.K_1'(s(bool,'T_0'))),s(fun(fun(X53,bool),bool),'const.util_prob.FUNSET_2'(s(fun(X53,bool),X54),s(fun(bool,bool),X55)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(fun(sK0,X50),X51),s(fun(fun(sK0,X50),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X50,bool)),X52))))),s(fun(bool,bool),X55)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK7(X52,sK2,X51,sK0,X50)))),s(fun(sK1,bool),sK3)))) ) | $spl96),
  inference(superposition,[],[f6706,f5803])).
fof(f6523,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,X3),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X3,bool)),X4)),X5,fun(sK0,X3))),s(fun(fun(sK0,X3),bool),X5))) | s(fun(fun(sK0,X3),bool),X5) = s(fun(fun(sK0,X3),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X3,bool)),X4))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK7(X4,sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X3,bool)),X4)),X5,fun(sK0,X3)),sK0,X3)))),s(fun(sK1,bool),sK3)))) ) | $spl96),
  inference(superposition,[],[f79,f5803])).
fof(f6522,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,X0),sK4(X1,'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X0,bool)),X2)),fun(sK0,X0))),s(fun(fun(sK0,X0),bool),X1))) | s(fun(fun(sK0,X0),bool),X1) = s(fun(fun(sK0,X0),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X0,bool)),X2))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK7(X2,sK2,sK4(X1,'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X0,bool)),X2)),fun(sK0,X0)),sK0,X0)))),s(fun(sK1,bool),sK3)))) ) | $spl96),
  inference(superposition,[],[f79,f5803])).
fof(f6532,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X5),X9),s(fun(fun(sK0,X5),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X5,bool)),X7))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK7(happ(s(fun(X6,fun(sK0,fun(X5,bool))),'const.combin.K_1'(s(fun(sK0,fun(X5,bool)),X7))),s(X6,X8)),sK2,X9,sK0,X5)))),s(fun(sK1,bool),sK3)))) ) | $spl96),
  inference(forward_demodulation,[],[f6509,f80])).
fof(f6509,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X5),X9),s(fun(fun(sK0,X5),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X5,bool)),'const.combin.K_2'(s(fun(sK0,fun(X5,bool)),X7),s(X6,X8))))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK7(happ(s(fun(X6,fun(sK0,fun(X5,bool))),'const.combin.K_1'(s(fun(sK0,fun(X5,bool)),X7))),s(X6,X8)),sK2,X9,sK0,X5)))),s(fun(sK1,bool),sK3)))) ) | $spl96),
  inference(superposition,[],[f5803,f81])).
fof(f6508,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X0),X4),s(fun(fun(sK0,X0),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X0,bool)),X1))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK7('const.combin.K_2'(s(fun(sK0,fun(X0,bool)),X1),s(X2,X3)),sK2,X4,sK0,X0)))),s(fun(sK1,bool),sK3)))) ) | $spl96),
  inference(superposition,[],[f5803,f80])).
fof(f6531,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X5),X7),s(fun(fun(sK0,X5),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X5,bool)),X9))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK7(X9,sK2,happ(s(fun(X6,fun(sK0,X5)),'const.combin.K_1'(s(fun(sK0,X5),X7))),s(X6,X8)),sK0,X5)))),s(fun(sK1,bool),sK3)))) ) | $spl96),
  inference(forward_demodulation,[],[f6507,f80])).
fof(f6507,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X5),'const.combin.K_2'(s(fun(sK0,X5),X7),s(X6,X8))),s(fun(fun(sK0,X5),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X5,bool)),X9))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK7(X9,sK2,happ(s(fun(X6,fun(sK0,X5)),'const.combin.K_1'(s(fun(sK0,X5),X7))),s(X6,X8)),sK0,X5)))),s(fun(sK1,bool),sK3)))) ) | $spl96),
  inference(superposition,[],[f5803,f81])).
fof(f6506,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X0),X1),s(fun(fun(sK0,X0),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X0,bool)),X4))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK7(X4,sK2,'const.combin.K_2'(s(fun(sK0,X0),X1),s(X2,X3)),sK0,X0)))),s(fun(sK1,bool),sK3)))) ) | $spl96),
  inference(superposition,[],[f5803,f80])).
fof(f9048,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)))))) | $spl166),
  inference(cnf_transformation,[],[f9048_D])).
fof(f9048_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)))))) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f619,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))))) | $spl36),
  inference(cnf_transformation,[],[f619_D])).
fof(f619_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))))) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f6260,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)))))) | $spl90),
  inference(resolution,[],[f5767,f84])).
fof(f9055,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))))) | $spl92),
  inference(global_subsumption,[],[f67,f5779,f6317,f6374,f6319,f6375,f6327,f6328,f6331,f6332,f6333,f6334,f6335,f6336,f6350,f6351,f6357,f6361,f6364,f6365,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6794,f6806,f6816,f104,f102,f101,f100,f7026,f98,f96,f7027,f7591,f7593,f7594,f7596,f7598,f7599,f7600,f7610,f7622,f7667,f7670,f7633,f7674,f7645,f7685,f7687,f7688,f7689,f7691,f7698,f7699,f7702,f7695,f7703,f7704,f7705,f7706,f7710,f7711,f7712,f7766,f7769,f7718,f7770,f7779,f7780,f7781,f7785,f7786,f7787,f7837,f7838,f7841,f7858,f7859,f7860,f7863,f7864,f7918,f7923,f7924,f7925,f7926,f7928,f7935,f7936,f7939,f7932,f7940,f7941,f7942,f7943,f7947,f7948,f7949,f8003,f8006,f7955,f8007,f8020,f8021,f8022,f8024,f8025,f8026,f8078,f8067,f8083,f8195,f8196,f8198,f8203,f8204,f8205,f8255,f8259,f8277,f8278,f8279,f8325,f8331,f8332,f8333,f8336,f8338,f8397,f8401,f8402,f8403,f8450,f8451,f8453,f8462,f8465,f8457,f8466,f8467,f8470,f8471,f8472,f8474,f8485,f8488,f8478,f8489,f8490,f8493,f8494,f8495,f8496,f8497,f8499,f8500,f8501,f8502,f8504,f8505,f8506,f8507,f8508,f8555,f8558,f8544,f8560,f8565,f8566,f8567,f8568,f8570,f8571,f8572,f8573,f8575,f8576,f8577,f8578,f8579,f8624,f8615,f8631,f8637,f8638,f8639,f8640,f8642,f8643,f8644,f8645,f8646,f8647,f8648,f8688,f8689,f8690,f8691,f8693,f8694,f8695,f8696,f8698,f8699,f8700,f8701,f8702,f8703,f8705,f8706,f8707,f8708,f8710,f8711,f8712,f8713,f8714,f8715,f8716,f8717,f8718,f8728,f8740,f8788,f8751,f8792,f8763,f8803,f8805,f8807,f8808,f8809,f8810,f8858,f8859,f8861,f8862,f8863,f8864,f8869,f8870,f8871,f8872,f8874,f8931,f8966,f8969,f8936,f8937,f8947,f8959,f8963,f8989,f8990,f8991,f8992,f8993,f8995,f8996,f8997,f8998,f8999,f9001,f9002,f9003,f9004,f9005,f9022,f9025,f9008,f9009,f9014,f9032,f9045,f9046,f9049,f9050,f6699,f9054,f6323])).
fof(f6323,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))))) | $spl92),
  inference(superposition,[],[f85,f5779])).
fof(f6699,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))))) | $spl92),
  inference(superposition,[],[f85,f5779])).
fof(f9049,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))))) | ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)))))) | $spl92),
  inference(forward_demodulation,[],[f8547,f80])).
fof(f8547,plain,(
  ( ! [X28,X29] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),'const.combin.K_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(X28,X29))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)))))) ) | $spl92),
  inference(superposition,[],[f6702,f5779])).
fof(f9045,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))))) | ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)))))) | $spl92),
  inference(forward_demodulation,[],[f8618,f80])).
fof(f8618,plain,(
  ( ! [X28,X29] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),'const.combin.K_2'(s(fun(sK1,bool),sK3),s(X28,X29)))))))) ) | $spl92),
  inference(superposition,[],[f6707,f5779])).
fof(f9032,plain,(
  ( ! [X6,X8,X7,X5,X9] : (p(s(bool,'const.bool.IN_2'(s(fun(X6,fun(sK0,X5)),X7),s(fun(fun(X6,fun(sK0,X5)),bool),'const.util_prob.DFUNSET_2'(s(fun(X6,bool),X9),s(fun(X6,fun(fun(sK0,X5),bool)),'const.combin.K_1'(s(fun(fun(sK0,X5),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X5,bool),X8)))))))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK6(X8,sK2,happ(s(fun(X6,fun(sK0,X5)),X7),s(X6,sK7('const.combin.K_1'(s(fun(fun(sK0,X5),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X5,bool),X8)))),X9,X7,X6,fun(sK0,X5)))),X5,sK0)))),s(fun(sK1,bool),sK3)))) ) | $spl92),
  inference(global_subsumption,[],[f67,f5779,f6317,f6374,f6319,f6375,f6327,f6328,f6331,f6332,f6333,f6334,f6335,f6336,f6350,f6351,f6357,f6361,f6364,f6365,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6794,f6806,f6816,f104,f102,f101,f100,f7026,f98,f96,f7027,f7591,f7593,f7594,f7596,f7598,f7599,f7600,f7610,f7622,f7667,f7670,f7633,f7674,f7645,f7685,f7687,f7688,f7689,f7691,f7698,f7699,f7702,f7695,f7703,f7704,f7705,f7706,f7710,f7711,f7712,f7766,f7769,f7718,f7770,f7779,f7780,f7781,f7785,f7786,f7787,f7837,f7838,f7841,f7858,f7859,f7860,f7863,f7864,f7918,f7923,f7924,f7925,f7926,f7928,f7935,f7936,f7939,f7932,f7940,f7941,f7942,f7943,f7947,f7948,f7949,f8003,f8006,f7955,f8007,f8020,f8021,f8022,f8024,f8025,f8026,f8078,f8067,f8083,f8195,f8196,f8198,f8203,f8204,f8205,f8255,f8259,f8277,f8278,f8279,f8325,f8331,f8332,f8333,f8336,f8338,f8397,f8401,f8402,f8403,f8450,f8451,f8453,f8462,f8465,f8457,f8466,f8467,f8470,f8471,f8472,f8474,f8485,f8488,f8478,f8489,f8490,f8493,f8494,f8495,f8496,f8497,f8499,f8500,f8501,f8502,f8504,f8505,f8506,f8507,f8508,f8555,f8558,f8544,f8560,f8565,f8566,f8567,f8568,f8570,f8571,f8572,f8573,f8575,f8576,f8577,f8578,f8579,f8624,f8615,f8631,f8637,f8638,f8639,f8640,f8642,f8643,f8644,f8645,f8646,f8647,f8648,f8688,f8689,f8690,f8691,f8693,f8694,f8695,f8696,f8698,f8699,f8700,f8701,f8702,f8703,f8705,f8706,f8707,f8708,f8710,f8711,f8712,f8713,f8714,f8715,f8716,f8717,f8718,f8728,f8740,f8788,f8751,f8792,f8763,f8803,f8805,f8807,f8808,f8809,f8810,f8858,f8859,f8861,f8862,f8863,f8864,f8869,f8870,f8871,f8872,f8874,f8931,f8966,f8969,f8936,f8937,f8947,f8959,f8963,f8989,f8990,f8991,f8992,f8993,f8995,f8996,f8997,f8998,f8999,f9001,f9002,f9003,f9004,f9005,f9022,f9025,f9008,f9009,f9014])).
fof(f9014,plain,(
  ( ! [X6,X8,X7,X5,X9] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(X6,fun(sK0,X5)),X7),s(fun(fun(X6,fun(sK0,X5)),bool),'const.util_prob.DFUNSET_2'(s(fun(X6,bool),X9),s(fun(X6,fun(fun(sK0,X5),bool)),'const.combin.K_1'(s(fun(fun(sK0,X5),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X5,bool),X8)))))))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK6(X8,sK2,happ(s(fun(X6,fun(sK0,X5)),X7),s(X6,sK7('const.combin.K_1'(s(fun(fun(sK0,X5),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X5,bool),X8)))),X9,X7,X6,fun(sK0,X5)))),X5,sK0)))),s(fun(sK1,bool),sK3)))) ) | $spl92),
  inference(superposition,[],[f6755,f5779])).
fof(f8959,plain,(
  ( ! [X94,X92,X90,X93,X91,X89] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),happ(s(fun(X92,fun(bool,bool)),X93),s(X92,sK7(X93,X94,'const.combin.K_1'(s(bool,'T_0')),X92,bool))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X92,bool),'const.combin.K_1'(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK6(X89,sK2,X90,X91,sK0)))),s(fun(sK1,bool),sK3))))),s(fun(fun(X92,bool),bool),'const.util_prob.DFUNSET_2'(s(fun(X92,bool),X94),s(fun(X92,fun(bool,bool)),X93)))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X91),X90),s(fun(fun(sK0,X91),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X91,bool),X89)))))) ) | $spl92),
  inference(superposition,[],[f6750,f5779])).
fof(f8947,plain,(
  ( ! [X39,X43,X41,X38,X42,X40] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),happ(s(fun(X41,fun(bool,bool)),X42),s(X41,sK7(X42,X43,'const.combin.K_1'(s(bool,'T_0')),X41,bool))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X41,bool),'const.combin.K_1'(s(bool,'const.bool.IN_2'(s(fun(sK0,X38),X39),s(fun(fun(sK0,X38),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X38,bool),X40))))))),s(fun(fun(X41,bool),bool),'const.util_prob.DFUNSET_2'(s(fun(X41,bool),X43),s(fun(X41,fun(bool,bool)),X42)))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK6(X40,sK2,X39,X38,sK0)))),s(fun(sK1,bool),sK3)))) ) | $spl92),
  inference(superposition,[],[f6750,f5779])).
fof(f8803,plain,(
  ( ! [X70,X74,X72,X71,X69,X73] : (p(s(bool,'const.bool.IN_2'(s(X72,happ(s(fun(sK1,X72),X73),s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK6(X69,sK2,X70,X71,sK0)))))),s(fun(X72,bool),happ(s(fun(sK1,fun(X72,bool)),X74),s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK6(X69,sK2,X70,X71,sK0))))))))) | p(s(bool,'const.bool.IN_2'(s(sK1,sK7(X74,sK3,X73,sK1,X72)),s(fun(sK1,bool),sK3)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X71),X70),s(fun(fun(sK0,X71),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X71,bool),X69)))))) ) | $spl92),
  inference(global_subsumption,[],[f67,f5779,f6317,f6374,f6319,f6375,f6327,f6328,f6331,f6332,f6333,f6334,f6335,f6336,f6350,f6351,f6357,f6361,f6364,f6365,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6794,f6806,f6816,f104,f102,f101,f100,f7026,f98,f96,f7027,f7591,f7593,f7594,f7596,f7598,f7599,f7600,f7610,f7622,f7667,f7670,f7633,f7674,f7645,f7685,f7687,f7688,f7689,f7691,f7698,f7699,f7702,f7695,f7703,f7704,f7705,f7706,f7710,f7711,f7712,f7766,f7769,f7718,f7770,f7779,f7780,f7781,f7785,f7786,f7787,f7837,f7838,f7841,f7858,f7859,f7860,f7863,f7864,f7918,f7923,f7924,f7925,f7926,f7928,f7935,f7936,f7939,f7932,f7940,f7941,f7942,f7943,f7947,f7948,f7949,f8003,f8006,f7955,f8007,f8020,f8021,f8022,f8024,f8025,f8026,f8078,f8067,f8083,f8195,f8196,f8198,f8203,f8204,f8205,f8255,f8259,f8277,f8278,f8279,f8325,f8331,f8332,f8333,f8336,f8338,f8397,f8401,f8402,f8403,f8450,f8451,f8453,f8462,f8465,f8457,f8466,f8467,f8470,f8471,f8472,f8474,f8485,f8488,f8478,f8489,f8490,f8493,f8494,f8495,f8496,f8497,f8499,f8500,f8501,f8502,f8504,f8505,f8506,f8507,f8508,f8555,f8558,f8544,f8560,f8565,f8566,f8567,f8568,f8570,f8571,f8572,f8573,f8575,f8576,f8577,f8578,f8579,f8624,f8615,f8631,f8637,f8638,f8639,f8640,f8642,f8643,f8644,f8645,f8646,f8647,f8648,f8688,f8689,f8690,f8691,f8693,f8694,f8695,f8696,f8698,f8699,f8700,f8701,f8702,f8703,f8705,f8706,f8707,f8708,f8710,f8711,f8712,f8713,f8714,f8715,f8716,f8717,f8718,f8728,f8740,f8788,f8751,f8792,f8763])).
fof(f8763,plain,(
  ( ! [X70,X74,X72,X71,X69,X73] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(X72,happ(s(fun(sK1,X72),X73),s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK6(X69,sK2,X70,X71,sK0)))))),s(fun(X72,bool),happ(s(fun(sK1,fun(X72,bool)),X74),s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK6(X69,sK2,X70,X71,sK0))))))))) | p(s(bool,'const.bool.IN_2'(s(sK1,sK7(X74,sK3,X73,sK1,X72)),s(fun(sK1,bool),sK3)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X71),X70),s(fun(fun(sK0,X71),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X71,bool),X69)))))) ) | $spl92),
  inference(superposition,[],[f6756,f5779])).
fof(f8792,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (p(s(bool,'const.bool.IN_2'(s(X21,happ(s(fun(fun(sK0,X18),X21),X22),s(fun(sK0,X18),X19))),s(fun(X21,bool),happ(s(fun(fun(sK0,X18),fun(X21,bool)),X23),s(fun(sK0,X18),X19)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,X18),sK7(X23,'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X18,bool),X20)),X22,fun(sK0,X18),X21)),s(fun(fun(sK0,X18),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X18,bool),X20)))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK6(X20,sK2,X19,X18,sK0)))),s(fun(sK1,bool),sK3)))) ) | $spl92),
  inference(global_subsumption,[],[f67,f5779,f6317,f6374,f6319,f6375,f6327,f6328,f6331,f6332,f6333,f6334,f6335,f6336,f6350,f6351,f6357,f6361,f6364,f6365,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6794,f6806,f6816,f104,f102,f101,f100,f7026,f98,f96,f7027,f7591,f7593,f7594,f7596,f7598,f7599,f7600,f7610,f7622,f7667,f7670,f7633,f7674,f7645,f7685,f7687,f7688,f7689,f7691,f7698,f7699,f7702,f7695,f7703,f7704,f7705,f7706,f7710,f7711,f7712,f7766,f7769,f7718,f7770,f7779,f7780,f7781,f7785,f7786,f7787,f7837,f7838,f7841,f7858,f7859,f7860,f7863,f7864,f7918,f7923,f7924,f7925,f7926,f7928,f7935,f7936,f7939,f7932,f7940,f7941,f7942,f7943,f7947,f7948,f7949,f8003,f8006,f7955,f8007,f8020,f8021,f8022,f8024,f8025,f8026,f8078,f8067,f8083,f8195,f8196,f8198,f8203,f8204,f8205,f8255,f8259,f8277,f8278,f8279,f8325,f8331,f8332,f8333,f8336,f8338,f8397,f8401,f8402,f8403,f8450,f8451,f8453,f8462,f8465,f8457,f8466,f8467,f8470,f8471,f8472,f8474,f8485,f8488,f8478,f8489,f8490,f8493,f8494,f8495,f8496,f8497,f8499,f8500,f8501,f8502,f8504,f8505,f8506,f8507,f8508,f8555,f8558,f8544,f8560,f8565,f8566,f8567,f8568,f8570,f8571,f8572,f8573,f8575,f8576,f8577,f8578,f8579,f8624,f8615,f8631,f8637,f8638,f8639,f8640,f8642,f8643,f8644,f8645,f8646,f8647,f8648,f8688,f8689,f8690,f8691,f8693,f8694,f8695,f8696,f8698,f8699,f8700,f8701,f8702,f8703,f8705,f8706,f8707,f8708,f8710,f8711,f8712,f8713,f8714,f8715,f8716,f8717,f8718,f8728,f8740,f8788,f8751])).
fof(f8751,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(X21,happ(s(fun(fun(sK0,X18),X21),X22),s(fun(sK0,X18),X19))),s(fun(X21,bool),happ(s(fun(fun(sK0,X18),fun(X21,bool)),X23),s(fun(sK0,X18),X19)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,X18),sK7(X23,'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X18,bool),X20)),X22,fun(sK0,X18),X21)),s(fun(fun(sK0,X18),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X18,bool),X20)))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK6(X20,sK2,X19,X18,sK0)))),s(fun(sK1,bool),sK3)))) ) | $spl92),
  inference(superposition,[],[f6756,f5779])).
fof(f8740,plain,(
  ( ! [X118,X116,X114,X112,X117,X115,X113] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X115)))) | p(s(bool,'const.bool.IN_2'(s(X116,happ(s(fun(bool,X116),X117),s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK6(X112,sK2,X113,X114,sK0)))),s(fun(sK1,bool),sK3))))),s(fun(X116,bool),happ(s(fun(bool,fun(X116,bool)),X118),s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK6(X112,sK2,X113,X114,sK0)))),s(fun(sK1,bool),sK3)))))))) | p(s(bool,'const.bool.IN_2'(s(bool,sK7(X118,X115,X117,bool,X116)),s(fun(bool,bool),X115)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X114),X113),s(fun(fun(sK0,X114),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X114,bool),X112)))))) ) | $spl92),
  inference(superposition,[],[f6756,f5779])).
fof(f8728,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X52)))) | p(s(bool,'const.bool.IN_2'(s(X53,happ(s(fun(bool,X53),X54),s(bool,'const.bool.IN_2'(s(fun(sK0,X49),X50),s(fun(fun(sK0,X49),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X49,bool),X51))))))),s(fun(X53,bool),happ(s(fun(bool,fun(X53,bool)),X55),s(bool,'const.bool.IN_2'(s(fun(sK0,X49),X50),s(fun(fun(sK0,X49),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X49,bool),X51)))))))))) | p(s(bool,'const.bool.IN_2'(s(bool,sK7(X55,X52,X54,bool,X53)),s(fun(bool,bool),X52)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK6(X51,sK2,X50,X49,sK0)))),s(fun(sK1,bool),sK3)))) ) | $spl92),
  inference(superposition,[],[f6756,f5779])).
fof(f8631,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(fun(X8,fun(sK0,X7)),X9),s(fun(fun(X8,fun(sK0,X7)),bool),'const.util_prob.FUNSET_2'(s(fun(X8,bool),X13),s(fun(fun(sK0,X7),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X7,bool),X10)))))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK6(X10,sK2,happ(s(fun(X8,fun(sK0,X7)),X9),s(X8,sK6('const.combin.K_2'(s(fun(fun(sK0,X7),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X7,bool),X10))),s(X11,X12)),X13,X9,fun(sK0,X7),X8))),X7,sK0)))),s(fun(sK1,bool),sK3)))) ) | $spl92),
  inference(global_subsumption,[],[f67,f5779,f6317,f6374,f6319,f6375,f6327,f6328,f6331,f6332,f6333,f6334,f6335,f6336,f6350,f6351,f6357,f6361,f6364,f6365,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6794,f6806,f6816,f104,f102,f101,f100,f7026,f98,f96,f7027,f7591,f7593,f7594,f7596,f7598,f7599,f7600,f7610,f7622,f7667,f7670,f7633,f7674,f7645,f7685,f7687,f7688,f7689,f7691,f7698,f7699,f7702,f7695,f7703,f7704,f7705,f7706,f7710,f7711,f7712,f7766,f7769,f7718,f7770,f7779,f7780,f7781,f7785,f7786,f7787,f7837,f7838,f7841,f7858,f7859,f7860,f7863,f7864,f7918,f7923,f7924,f7925,f7926,f7928,f7935,f7936,f7939,f7932,f7940,f7941,f7942,f7943,f7947,f7948,f7949,f8003,f8006,f7955,f8007,f8020,f8021,f8022,f8024,f8025,f8026,f8078,f8067,f8083,f8195,f8196,f8198,f8203,f8204,f8205,f8255,f8259,f8277,f8278,f8279,f8325,f8331,f8332,f8333,f8336,f8338,f8397,f8401,f8402,f8403,f8450,f8451,f8453,f8462,f8465,f8457,f8466,f8467,f8470,f8471,f8472,f8474,f8485,f8488,f8478,f8489,f8490,f8493,f8494,f8495,f8496,f8497,f8499,f8500,f8501,f8502,f8504,f8505,f8506,f8507,f8508,f8555,f8558,f8544,f8560,f8565,f8566,f8567,f8568,f8570,f8571,f8572,f8573,f8575,f8576,f8577,f8578,f8579,f8624,f8615])).
fof(f8615,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(X8,fun(sK0,X7)),X9),s(fun(fun(X8,fun(sK0,X7)),bool),'const.util_prob.FUNSET_2'(s(fun(X8,bool),X13),s(fun(fun(sK0,X7),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X7,bool),X10)))))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK6(X10,sK2,happ(s(fun(X8,fun(sK0,X7)),X9),s(X8,sK6('const.combin.K_2'(s(fun(fun(sK0,X7),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X7,bool),X10))),s(X11,X12)),X13,X9,fun(sK0,X7),X8))),X7,sK0)))),s(fun(sK1,bool),sK3)))) ) | $spl92),
  inference(superposition,[],[f6707,f5779])).
fof(f8560,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(fun(X8,fun(sK0,X7)),X9),s(fun(fun(X8,fun(sK0,X7)),bool),'const.util_prob.FUNSET_2'(s(fun(X8,bool),X11),s(fun(fun(sK0,X7),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X7,bool),X10)))))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK6(X10,sK2,happ(s(fun(X8,fun(sK0,X7)),X9),s(X8,sK6('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X7,bool),X10)),X11,'const.combin.K_2'(s(fun(X8,fun(sK0,X7)),X9),s(X12,X13)),fun(sK0,X7),X8))),X7,sK0)))),s(fun(sK1,bool),sK3)))) ) | $spl92),
  inference(global_subsumption,[],[f67,f5779,f6317,f6374,f6319,f6375,f6327,f6328,f6331,f6332,f6333,f6334,f6335,f6336,f6350,f6351,f6357,f6361,f6364,f6365,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6794,f6806,f6816,f104,f102,f101,f100,f7026,f98,f96,f7027,f7591,f7593,f7594,f7596,f7598,f7599,f7600,f7610,f7622,f7667,f7670,f7633,f7674,f7645,f7685,f7687,f7688,f7689,f7691,f7698,f7699,f7702,f7695,f7703,f7704,f7705,f7706,f7710,f7711,f7712,f7766,f7769,f7718,f7770,f7779,f7780,f7781,f7785,f7786,f7787,f7837,f7838,f7841,f7858,f7859,f7860,f7863,f7864,f7918,f7923,f7924,f7925,f7926,f7928,f7935,f7936,f7939,f7932,f7940,f7941,f7942,f7943,f7947,f7948,f7949,f8003,f8006,f7955,f8007,f8020,f8021,f8022,f8024,f8025,f8026,f8078,f8067,f8083,f8195,f8196,f8198,f8203,f8204,f8205,f8255,f8259,f8277,f8278,f8279,f8325,f8331,f8332,f8333,f8336,f8338,f8397,f8401,f8402,f8403,f8450,f8451,f8453,f8462,f8465,f8457,f8466,f8467,f8470,f8471,f8472,f8474,f8485,f8488,f8478,f8489,f8490,f8493,f8494,f8495,f8496,f8497,f8499,f8500,f8501,f8502,f8504,f8505,f8506,f8507,f8508,f8555,f8558,f8544])).
fof(f8544,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(X8,fun(sK0,X7)),X9),s(fun(fun(X8,fun(sK0,X7)),bool),'const.util_prob.FUNSET_2'(s(fun(X8,bool),X11),s(fun(fun(sK0,X7),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X7,bool),X10)))))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK6(X10,sK2,happ(s(fun(X8,fun(sK0,X7)),X9),s(X8,sK6('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X7,bool),X10)),X11,'const.combin.K_2'(s(fun(X8,fun(sK0,X7)),X9),s(X12,X13)),fun(sK0,X7),X8))),X7,sK0)))),s(fun(sK1,bool),sK3)))) ) | $spl92),
  inference(superposition,[],[f6702,f5779])).
fof(f7685,plain,(
  ( ! [X70,X74,X72,X71,X69,X73] : (p(s(bool,'const.bool.IN_2'(s(X72,happ(s(fun(sK1,X72),X73),s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK6(X69,sK2,X70,X71,sK0)))))),s(fun(X72,bool),X74)))) | p(s(bool,'const.bool.IN_2'(s(sK1,sK6(X74,sK3,X73,X72,sK1)),s(fun(sK1,bool),sK3)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X71),X70),s(fun(fun(sK0,X71),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X71,bool),X69)))))) ) | $spl92),
  inference(global_subsumption,[],[f67,f5779,f6317,f6374,f6319,f6375,f6327,f6328,f6331,f6332,f6333,f6334,f6335,f6336,f6350,f6351,f6357,f6361,f6364,f6365,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6794,f6806,f6816,f104,f102,f101,f100,f7026,f98,f96,f7027,f7591,f7593,f7594,f7596,f7598,f7599,f7600,f7610,f7622,f7667,f7670,f7633,f7674,f7645])).
fof(f7645,plain,(
  ( ! [X70,X74,X72,X71,X69,X73] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(X72,happ(s(fun(sK1,X72),X73),s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK6(X69,sK2,X70,X71,sK0)))))),s(fun(X72,bool),X74)))) | p(s(bool,'const.bool.IN_2'(s(sK1,sK6(X74,sK3,X73,X72,sK1)),s(fun(sK1,bool),sK3)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X71),X70),s(fun(fun(sK0,X71),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X71,bool),X69)))))) ) | $spl92),
  inference(superposition,[],[f6711,f5779])).
fof(f7674,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (p(s(bool,'const.bool.IN_2'(s(X21,happ(s(fun(fun(sK0,X18),X21),X22),s(fun(sK0,X18),X19))),s(fun(X21,bool),X23)))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,X18),sK6(X23,'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X18,bool),X20)),X22,X21,fun(sK0,X18))),s(fun(fun(sK0,X18),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X18,bool),X20)))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK6(X20,sK2,X19,X18,sK0)))),s(fun(sK1,bool),sK3)))) ) | $spl92),
  inference(global_subsumption,[],[f67,f5779,f6317,f6374,f6319,f6375,f6327,f6328,f6331,f6332,f6333,f6334,f6335,f6336,f6350,f6351,f6357,f6361,f6364,f6365,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6794,f6806,f6816,f104,f102,f101,f100,f7026,f98,f96,f7027,f7591,f7593,f7594,f7596,f7598,f7599,f7600,f7610,f7622,f7667,f7670,f7633])).
fof(f7633,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(X21,happ(s(fun(fun(sK0,X18),X21),X22),s(fun(sK0,X18),X19))),s(fun(X21,bool),X23)))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,X18),sK6(X23,'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X18,bool),X20)),X22,X21,fun(sK0,X18))),s(fun(fun(sK0,X18),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X18,bool),X20)))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK6(X20,sK2,X19,X18,sK0)))),s(fun(sK1,bool),sK3)))) ) | $spl92),
  inference(superposition,[],[f6711,f5779])).
fof(f7622,plain,(
  ( ! [X118,X116,X114,X112,X117,X115,X113] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X115)))) | p(s(bool,'const.bool.IN_2'(s(X116,happ(s(fun(bool,X116),X117),s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK6(X112,sK2,X113,X114,sK0)))),s(fun(sK1,bool),sK3))))),s(fun(X116,bool),X118)))) | p(s(bool,'const.bool.IN_2'(s(bool,sK6(X118,X115,X117,X116,bool)),s(fun(bool,bool),X115)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X114),X113),s(fun(fun(sK0,X114),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X114,bool),X112)))))) ) | $spl92),
  inference(superposition,[],[f6711,f5779])).
fof(f7610,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X52)))) | p(s(bool,'const.bool.IN_2'(s(X53,happ(s(fun(bool,X53),X54),s(bool,'const.bool.IN_2'(s(fun(sK0,X49),X50),s(fun(fun(sK0,X49),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X49,bool),X51))))))),s(fun(X53,bool),X55)))) | p(s(bool,'const.bool.IN_2'(s(bool,sK6(X55,X52,X54,X53,bool)),s(fun(bool,bool),X52)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK6(X51,sK2,X50,X49,sK0)))),s(fun(sK1,bool),sK3)))) ) | $spl92),
  inference(superposition,[],[f6711,f5779])).
fof(f6806,plain,(
  ( ! [X94,X92,X90,X93,X91,X89] : (p(s(bool,'const.bool.IN_2'(s(fun(X92,bool),'const.combin.K_1'(s(bool,'T_0'))),s(fun(fun(X92,bool),bool),'const.util_prob.FUNSET_2'(s(fun(X92,bool),X93),s(fun(bool,bool),X94)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK6(X89,sK2,X90,X91,sK0)))),s(fun(sK1,bool),sK3))),s(fun(bool,bool),X94)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X91),X90),s(fun(fun(sK0,X91),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X91,bool),X89)))))) ) | $spl92),
  inference(superposition,[],[f6706,f5779])).
fof(f6794,plain,(
  ( ! [X39,X43,X41,X38,X42,X40] : (p(s(bool,'const.bool.IN_2'(s(fun(X41,bool),'const.combin.K_1'(s(bool,'T_0'))),s(fun(fun(X41,bool),bool),'const.util_prob.FUNSET_2'(s(fun(X41,bool),X42),s(fun(bool,bool),X43)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(fun(sK0,X38),X39),s(fun(fun(sK0,X38),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X38,bool),X40))))),s(fun(bool,bool),X43)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK6(X40,sK2,X39,X38,sK0)))),s(fun(sK1,bool),sK3)))) ) | $spl92),
  inference(superposition,[],[f6706,f5779])).
fof(f6351,plain,(
  ( ! [X23,X21,X22] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,X21),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X21,bool),X22)),X23,fun(sK0,X21))),s(fun(fun(sK0,X21),bool),X23))) | s(fun(fun(sK0,X21),bool),X23) = s(fun(fun(sK0,X21),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X21,bool),X22))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK6(X22,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X21,bool),X22)),X23,fun(sK0,X21)),X21,sK0)))),s(fun(sK1,bool),sK3)))) ) | $spl92),
  inference(superposition,[],[f79,f5779])).
fof(f6350,plain,(
  ( ! [X19,X20,X18] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,X18),sK4(X19,'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X18,bool),X20)),fun(sK0,X18))),s(fun(fun(sK0,X18),bool),X19))) | s(fun(fun(sK0,X18),bool),X19) = s(fun(fun(sK0,X18),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X18,bool),X20))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK6(X20,sK2,sK4(X19,'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X18,bool),X20)),fun(sK0,X18)),X18,sK0)))),s(fun(sK1,bool),sK3)))) ) | $spl92),
  inference(superposition,[],[f79,f5779])).
fof(f6375,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X5),X9),s(fun(fun(sK0,X5),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X5,bool),X7))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK6(happ(s(fun(X6,fun(X5,bool)),'const.combin.K_1'(s(fun(X5,bool),X7))),s(X6,X8)),sK2,X9,X5,sK0)))),s(fun(sK1,bool),sK3)))) ) | $spl92),
  inference(forward_demodulation,[],[f6320,f80])).
fof(f6320,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X5),X9),s(fun(fun(sK0,X5),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X5,bool),'const.combin.K_2'(s(fun(X5,bool),X7),s(X6,X8))))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK6(happ(s(fun(X6,fun(X5,bool)),'const.combin.K_1'(s(fun(X5,bool),X7))),s(X6,X8)),sK2,X9,X5,sK0)))),s(fun(sK1,bool),sK3)))) ) | $spl92),
  inference(superposition,[],[f5779,f81])).
fof(f6319,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X0),X4),s(fun(fun(sK0,X0),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X0,bool),X1))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK6('const.combin.K_2'(s(fun(X0,bool),X1),s(X2,X3)),sK2,X4,X0,sK0)))),s(fun(sK1,bool),sK3)))) ) | $spl92),
  inference(superposition,[],[f5779,f80])).
fof(f6374,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X5),X7),s(fun(fun(sK0,X5),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X5,bool),X9))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK6(X9,sK2,happ(s(fun(X6,fun(sK0,X5)),'const.combin.K_1'(s(fun(sK0,X5),X7))),s(X6,X8)),X5,sK0)))),s(fun(sK1,bool),sK3)))) ) | $spl92),
  inference(forward_demodulation,[],[f6318,f80])).
fof(f6318,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X5),'const.combin.K_2'(s(fun(sK0,X5),X7),s(X6,X8))),s(fun(fun(sK0,X5),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X5,bool),X9))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK6(X9,sK2,happ(s(fun(X6,fun(sK0,X5)),'const.combin.K_1'(s(fun(sK0,X5),X7))),s(X6,X8)),X5,sK0)))),s(fun(sK1,bool),sK3)))) ) | $spl92),
  inference(superposition,[],[f5779,f81])).
fof(f6317,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X0),X1),s(fun(fun(sK0,X0),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X0,bool),X4))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK6(X4,sK2,'const.combin.K_2'(s(fun(sK0,X0),X1),s(X2,X3)),X0,sK0)))),s(fun(sK1,bool),sK3)))) ) | $spl92),
  inference(superposition,[],[f5779,f80])).
fof(f6613,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))))) | $spl90),
  inference(resolution,[],[f84,f5767])).
fof(f9054,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))))) | $spl92),
  inference(global_subsumption,[],[f67,f5779,f6317,f6374,f6319,f6375,f6327,f6328,f6331,f6332,f6333,f6334,f6335,f6336,f6350,f6351,f6357,f6361,f6364,f6365,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6794,f6806,f6816,f104,f102,f101,f100,f7026,f98,f96,f7027,f7591,f7593,f7594,f7596,f7598,f7599,f7600,f7610,f7622,f7667,f7670,f7633,f7674,f7645,f7685,f7687,f7688,f7689,f7691,f7698,f7699,f7702,f7695,f7703,f7704,f7705,f7706,f7710,f7711,f7712,f7766,f7769,f7718,f7770,f7779,f7780,f7781,f7785,f7786,f7787,f7837,f7838,f7841,f7858,f7859,f7860,f7863,f7864,f7918,f7923,f7924,f7925,f7926,f7928,f7935,f7936,f7939,f7932,f7940,f7941,f7942,f7943,f7947,f7948,f7949,f8003,f8006,f7955,f8007,f8020,f8021,f8022,f8024,f8025,f8026,f8078,f8067,f8083,f8195,f8196,f8198,f8203,f8204,f8205,f8255,f8259,f8277,f8278,f8279,f8325,f8331,f8332,f8333,f8336,f8338,f8397,f8401,f8402,f8403,f8450,f8451,f8453,f8462,f8465,f8457,f8466,f8467,f8470,f8471,f8472,f8474,f8485,f8488,f8478,f8489,f8490,f8493,f8494,f8495,f8496,f8497,f8499,f8500,f8501,f8502,f8504,f8505,f8506,f8507,f8508,f8555,f8558,f8544,f8560,f8565,f8566,f8567,f8568,f8570,f8571,f8572,f8573,f8575,f8576,f8577,f8578,f8579,f8624,f8615,f8631,f8637,f8638,f8639,f8640,f8642,f8643,f8644,f8645,f8646,f8647,f8648,f8688,f8689,f8690,f8691,f8693,f8694,f8695,f8696,f8698,f8699,f8700,f8701,f8702,f8703,f8705,f8706,f8707,f8708,f8710,f8711,f8712,f8713,f8714,f8715,f8716,f8717,f8718,f8728,f8740,f8788,f8751,f8792,f8763,f8803,f8805,f8807,f8808,f8809,f8810,f8858,f8859,f8861,f8862,f8863,f8864,f8869,f8870,f8871,f8872,f8874,f8931,f8966,f8969,f8936,f8937,f8947,f8959,f8963,f8989,f8990,f8991,f8992,f8993,f8995,f8996,f8997,f8998,f8999,f9001,f9002,f9003,f9004,f9005,f9022,f9025,f9008,f9009,f9014,f9032,f9045,f9046,f9049,f9050,f6699])).
fof(f6858,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))))) | ($spl90 | $spl122)),
  inference(duplicate_literal_removal,[],[f6844])).
fof(f6844,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))))) | ($spl90 | $spl122)),
  inference(resolution,[],[f6642,f5767])).
fof(f9053,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)))))) | $spl90),
  inference(forward_demodulation,[],[f7855,f80])).
fof(f7855,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),'const.combin.K_2'(s(fun(sK0,bool),sK2),s(X2,X3))),s(fun(sK1,bool),sK3)))))) ) | $spl90),
  inference(resolution,[],[f6637,f5767])).
fof(f9052,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)))))) | $spl90),
  inference(forward_demodulation,[],[f9051,f80])).
fof(f9051,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),'const.combin.K_2'(s(fun(sK0,bool),sK2),s(X2,X3))),s(fun(sK1,bool),sK3))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))))))) ) | $spl90),
  inference(forward_demodulation,[],[f8327,f81])).
fof(f8327,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),happ(s(fun(X2,fun(sK0,bool)),'const.combin.K_1'(s(fun(sK0,bool),sK2))),s(X2,X3))),s(fun(sK1,bool),sK3)))))) ) | $spl90),
  inference(resolution,[],[f6640,f5767])).
fof(f9050,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)))))) | $spl92),
  inference(global_subsumption,[],[f67,f5779,f6317,f6374,f6319,f6375,f6327,f6328,f6331,f6332,f6333,f6334,f6335,f6336,f6350,f6351,f6357,f6361,f6364,f6365,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6794,f6806,f6816,f104,f102,f101,f100,f7026,f98,f96,f7027,f7591,f7593,f7594,f7596,f7598,f7599,f7600,f7610,f7622,f7667,f7670,f7633,f7674,f7645,f7685,f7687,f7688,f7689,f7691,f7698,f7699,f7702,f7695,f7703,f7704,f7705,f7706,f7710,f7711,f7712,f7766,f7769,f7718,f7770,f7779,f7780,f7781,f7785,f7786,f7787,f7837,f7838,f7841,f7858,f7859,f7860,f7863,f7864,f7918,f7923,f7924,f7925,f7926,f7928,f7935,f7936,f7939,f7932,f7940,f7941,f7942,f7943,f7947,f7948,f7949,f8003,f8006,f7955,f8007,f8020,f8021,f8022,f8024,f8025,f8026,f8078,f8067,f8083,f8195,f8196,f8198,f8203,f8204,f8205,f8255,f8259,f8277,f8278,f8279,f8325,f8331,f8332,f8333,f8336,f8338,f8397,f8401,f8402,f8403,f8450,f8451,f8453,f8462,f8465,f8457,f8466,f8467,f8470,f8471,f8472,f8474,f8485,f8488,f8478,f8489,f8490,f8493,f8494,f8495,f8496,f8497,f8499,f8500,f8501,f8502,f8504,f8505,f8506,f8507,f8508,f8555,f8558,f8544,f8560,f8565,f8566,f8567,f8568,f8570,f8571,f8572,f8573,f8575,f8576,f8577,f8578,f8579,f8624,f8615,f8631,f8637,f8638,f8639,f8640,f8642,f8643,f8644,f8645,f8646,f8647,f8648,f8688,f8689,f8690,f8691,f8693,f8694,f8695,f8696,f8698,f8699,f8700,f8701,f8702,f8703,f8705,f8706,f8707,f8708,f8710,f8711,f8712,f8713,f8714,f8715,f8716,f8717,f8718,f8728,f8740,f8788,f8751,f8792,f8763,f8803,f8805,f8807,f8808,f8809,f8810,f8858,f8859,f8861,f8862,f8863,f8864,f8869,f8870,f8871,f8872,f8874,f8931,f8966,f8969,f8936,f8937,f8947,f8959,f8963,f8989,f8990,f8991,f8992,f8993,f8995,f8996,f8997,f8998,f8999,f9001,f9002,f9003,f9004,f9005,f9022,f9025,f9008,f9009,f9014,f9032,f9045,f9046,f9049])).
fof(f9046,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)))))) | $spl92),
  inference(global_subsumption,[],[f67,f5779,f6317,f6374,f6319,f6375,f6327,f6328,f6331,f6332,f6333,f6334,f6335,f6336,f6350,f6351,f6357,f6361,f6364,f6365,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6794,f6806,f6816,f104,f102,f101,f100,f7026,f98,f96,f7027,f7591,f7593,f7594,f7596,f7598,f7599,f7600,f7610,f7622,f7667,f7670,f7633,f7674,f7645,f7685,f7687,f7688,f7689,f7691,f7698,f7699,f7702,f7695,f7703,f7704,f7705,f7706,f7710,f7711,f7712,f7766,f7769,f7718,f7770,f7779,f7780,f7781,f7785,f7786,f7787,f7837,f7838,f7841,f7858,f7859,f7860,f7863,f7864,f7918,f7923,f7924,f7925,f7926,f7928,f7935,f7936,f7939,f7932,f7940,f7941,f7942,f7943,f7947,f7948,f7949,f8003,f8006,f7955,f8007,f8020,f8021,f8022,f8024,f8025,f8026,f8078,f8067,f8083,f8195,f8196,f8198,f8203,f8204,f8205,f8255,f8259,f8277,f8278,f8279,f8325,f8331,f8332,f8333,f8336,f8338,f8397,f8401,f8402,f8403,f8450,f8451,f8453,f8462,f8465,f8457,f8466,f8467,f8470,f8471,f8472,f8474,f8485,f8488,f8478,f8489,f8490,f8493,f8494,f8495,f8496,f8497,f8499,f8500,f8501,f8502,f8504,f8505,f8506,f8507,f8508,f8555,f8558,f8544,f8560,f8565,f8566,f8567,f8568,f8570,f8571,f8572,f8573,f8575,f8576,f8577,f8578,f8579,f8624,f8615,f8631,f8637,f8638,f8639,f8640,f8642,f8643,f8644,f8645,f8646,f8647,f8648,f8688,f8689,f8690,f8691,f8693,f8694,f8695,f8696,f8698,f8699,f8700,f8701,f8702,f8703,f8705,f8706,f8707,f8708,f8710,f8711,f8712,f8713,f8714,f8715,f8716,f8717,f8718,f8728,f8740,f8788,f8751,f8792,f8763,f8803,f8805,f8807,f8808,f8809,f8810,f8858,f8859,f8861,f8862,f8863,f8864,f8869,f8870,f8871,f8872,f8874,f8931,f8966,f8969,f8936,f8937,f8947,f8959,f8963,f8989,f8990,f8991,f8992,f8993,f8995,f8996,f8997,f8998,f8999,f9001,f9002,f9003,f9004,f9005,f9022,f9025,f9008,f9009,f9014,f9032,f9045])).
fof(f723,plain,(
  p(s(bool,'F_0')) | $spl42),
  inference(cnf_transformation,[],[f723_D])).
fof(f723_D,plain,(
  p(s(bool,'F_0')) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f152,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl30),
  inference(cnf_transformation,[],[f152_D])).
fof(f152_D,plain,(
  s(bool,'T_0') = s(bool,'F_0') <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f9044,plain,(
  s(bool,'T_0') = s(bool,'F_0') | p(s(bool,'F_0')) | ($spl52 | $spl98)),
  inference(global_subsumption,[],[f67,f2153,f5043,f5815,f6536,f6561,f6538,f6562,f6552,f6553,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6795,f6805,f6816,f104,f102,f101,f100,f7026,f98,f96,f7027,f7591,f7593,f7594,f7596,f7598,f7599,f7600,f7611,f7621,f7667,f7670,f7634,f7675,f7644,f7684,f7687,f7688,f7689,f7691,f7698,f7699,f7702,f7695,f7703,f7704,f7705,f7706,f7710,f7711,f7712,f7766,f7769,f7718,f7770,f7779,f7780,f7781,f7785,f7786,f7787,f7837,f7838,f7841,f7858,f7859,f7860,f7863,f7864,f7918,f7923,f7924,f7925,f7926,f7928,f7935,f7936,f7939,f7932,f7940,f7941,f7942,f7943,f7947,f7948,f7949,f8003,f8006,f7955,f8007,f7958,f8020,f8021,f8022,f8024,f8025,f8026,f8078,f8067,f8083,f8070,f8195,f8196,f8198,f8203,f8204,f8205,f8255,f8259,f8214,f8277,f8278,f8279,f8325,f8331,f8332,f8333,f8336,f8338,f8397,f8401,f8402,f8403,f8450,f8451,f8453,f8462,f8465,f8457,f8466,f8467,f8470,f8471,f8472,f8474,f8485,f8488,f8478,f8489,f8490,f8493,f8494,f8495,f8496,f8497,f8499,f8500,f8501,f8502,f8504,f8505,f8506,f8507,f8508,f8555,f8558,f8545,f8561,f8565,f8566,f8567,f8568,f8570,f8571,f8572,f8573,f8575,f8576,f8577,f8578,f8579,f8624,f8616,f8632,f8637,f8638,f8639,f8640,f8642,f8643,f8644,f8645,f8646,f8647,f8648,f8688,f8689,f8690,f8691,f8693,f8694,f8695,f8696,f8698,f8699,f8700,f8701,f8702,f8703,f8705,f8706,f8707,f8708,f8710,f8711,f8712,f8713,f8714,f8715,f8716,f8717,f8718,f8729,f8739,f8788,f8752,f8793,f8762,f8802,f8805,f8807,f8808,f8809,f8810,f8858,f8859,f8861,f8862,f8863,f8864,f8869,f8870,f8871,f8872,f8874,f8931,f8918,f8966,f8969,f8936,f8937,f8948,f8958,f8963,f8989,f8990,f8991,f8992,f8993,f8995,f8996,f8997,f8998,f8999,f9001,f9002,f9003,f9004,f9005,f9022,f9025,f9008,f9009,f9015,f9033,f9043])).
fof(f9043,plain,(
  s(bool,'T_0') = s(bool,'F_0') | p(s(bool,'F_0')) | ~p(s(bool,'T_0')) | ($spl52 | $spl98)),
  inference(forward_demodulation,[],[f9042,f2153])).
fof(f9042,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3))))))) | ($spl52 | $spl98)),
  inference(forward_demodulation,[],[f9020,f2153])).
fof(f9041,plain,(
  s(bool,'T_0') = s(bool,'F_0') | p(s(bool,'F_0')) | ($spl38 | $spl96)),
  inference(global_subsumption,[],[f67,f781,f2084,f621,f5803,f6506,f6531,f6508,f6532,f6522,f6523,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6796,f6804,f6816,f104,f102,f101,f100,f7026,f98,f96,f7027,f7591,f7593,f7594,f7596,f7598,f7599,f7600,f7612,f7620,f7667,f7670,f7635,f7676,f7643,f7683,f7687,f7688,f7689,f7691,f7698,f7699,f7702,f7695,f7703,f7704,f7705,f7706,f7710,f7711,f7712,f7766,f7769,f7718,f7770,f7779,f7780,f7781,f7785,f7786,f7787,f7837,f7838,f7841,f7858,f7859,f7860,f7863,f7864,f7918,f7923,f7924,f7925,f7926,f7928,f7935,f7936,f7939,f7932,f7940,f7941,f7942,f7943,f7947,f7948,f7949,f8003,f8006,f7955,f8007,f7957,f8020,f8021,f8022,f8024,f8025,f8026,f8078,f8067,f8083,f8069,f8195,f8196,f8198,f8203,f8204,f8205,f8255,f8259,f8213,f8277,f8278,f8279,f8325,f8331,f8332,f8333,f8336,f8338,f8397,f8401,f8402,f8403,f8450,f8451,f8453,f8462,f8465,f8457,f8466,f8467,f8470,f8471,f8472,f8474,f8485,f8488,f8478,f8489,f8490,f8493,f8494,f8495,f8496,f8497,f8499,f8500,f8501,f8502,f8504,f8505,f8506,f8507,f8508,f8555,f8558,f8546,f8562,f8565,f8566,f8567,f8568,f8570,f8571,f8572,f8573,f8575,f8576,f8577,f8578,f8579,f8624,f8617,f8633,f8637,f8638,f8639,f8640,f8642,f8643,f8644,f8645,f8646,f8647,f8648,f8688,f8689,f8690,f8691,f8693,f8694,f8695,f8696,f8698,f8699,f8700,f8701,f8702,f8703,f8705,f8706,f8707,f8708,f8710,f8711,f8712,f8713,f8714,f8715,f8716,f8717,f8718,f8730,f8738,f8788,f8753,f8794,f8761,f8801,f8805,f8807,f8808,f8809,f8810,f8858,f8859,f8861,f8862,f8863,f8864,f8869,f8870,f8871,f8872,f8874,f8931,f8917,f8966,f8969,f8936,f8937,f8949,f8957,f8963,f8989,f8990,f8991,f8992,f8993,f8995,f8996,f8997,f8998,f8999,f9001,f9002,f9003,f9004,f9005,f9022,f9025,f9008,f9009,f9016,f9034,f9040])).
fof(f9040,plain,(
  s(bool,'T_0') = s(bool,'F_0') | p(s(bool,'F_0')) | ~p(s(bool,'T_0')) | ($spl38 | $spl96)),
  inference(forward_demodulation,[],[f9039,f621])).
fof(f9039,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3))))))) | ($spl38 | $spl96)),
  inference(forward_demodulation,[],[f9019,f621])).
fof(f621,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3))))))) | $spl38),
  inference(cnf_transformation,[],[f621_D])).
fof(f621_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3))))))) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f9038,plain,(
  p(s(bool,'F_0')) | ($spl52 | $spl86)),
  inference(global_subsumption,[],[f67,f2153,f5043,f5752,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6816,f104,f102,f101,f100,f7026,f98,f96,f7027,f7591,f7593,f7594,f7596,f7598,f7599,f7600,f7667,f7670,f7642,f7682,f7687,f7688,f7689,f7691,f7698,f7699,f7702,f7695,f7703,f7704,f7705,f7706,f7710,f7711,f7712,f7766,f7769,f7718,f7770,f7779,f7780,f7781,f7785,f7786,f7787,f7837,f7838,f7841,f7858,f7859,f7860,f7863,f7864,f7918,f7923,f7924,f7925,f7926,f7928,f7935,f7936,f7939,f7932,f7940,f7941,f7942,f7943,f7947,f7948,f7949,f8003,f8006,f7955,f8007,f7958,f8020,f8021,f8022,f8024,f8025,f8026,f8078,f8067,f8083,f8070,f8195,f8196,f8198,f8203,f8204,f8205,f8255,f8259,f8214,f8277,f8278,f8279,f8325,f8331,f8332,f8333,f8336,f8338,f8397,f8401,f8402,f8403,f8450,f8451,f8453,f8462,f8465,f8457,f8466,f8467,f8470,f8471,f8472,f8474,f8485,f8488,f8478,f8489,f8490,f8493,f8494,f8495,f8496,f8497,f8499,f8500,f8501,f8502,f8504,f8505,f8506,f8507,f8508,f8555,f8558,f8565,f8566,f8567,f8568,f8570,f8571,f8572,f8573,f8575,f8576,f8577,f8578,f8579,f8624,f8637,f8638,f8639,f8640,f8642,f8643,f8644,f8645,f8646,f8647,f8648,f8688,f8689,f8690,f8691,f8693,f8694,f8695,f8696,f8698,f8699,f8700,f8701,f8702,f8703,f8705,f8706,f8707,f8708,f8710,f8711,f8712,f8713,f8714,f8715,f8716,f8717,f8718,f8788,f8760,f8800,f8805,f8807,f8808,f8809,f8810,f8858,f8859,f8861,f8862,f8863,f8864,f8869,f8870,f8871,f8872,f8874,f8931,f8918,f8966,f8969,f8936,f8937,f8963,f8989,f8990,f8991,f8992,f8993,f8995,f8996,f8997,f8998,f8999,f9001,f9002,f9003,f9004,f9005,f9022,f9025,f9008,f9009,f9037])).
fof(f9037,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'T_0')) | ($spl52 | $spl86)),
  inference(forward_demodulation,[],[f9018,f2153])).
fof(f9036,plain,(
  p(s(bool,'F_0')) | ($spl38 | $spl80)),
  inference(global_subsumption,[],[f67,f781,f2084,f621,f5733,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6816,f104,f102,f101,f100,f7026,f98,f96,f7027,f7591,f7593,f7594,f7596,f7598,f7599,f7600,f7667,f7670,f7639,f7679,f7687,f7688,f7689,f7691,f7698,f7699,f7702,f7695,f7703,f7704,f7705,f7706,f7710,f7711,f7712,f7766,f7769,f7718,f7770,f7779,f7780,f7781,f7785,f7786,f7787,f7837,f7838,f7841,f7858,f7859,f7860,f7863,f7864,f7918,f7923,f7924,f7925,f7926,f7928,f7935,f7936,f7939,f7932,f7940,f7941,f7942,f7943,f7947,f7948,f7949,f8003,f8006,f7955,f8007,f7957,f8020,f8021,f8022,f8024,f8025,f8026,f8078,f8067,f8083,f8069,f8195,f8196,f8198,f8203,f8204,f8205,f8255,f8259,f8213,f8277,f8278,f8279,f8325,f8331,f8332,f8333,f8336,f8338,f8397,f8401,f8402,f8403,f8450,f8451,f8453,f8462,f8465,f8457,f8466,f8467,f8470,f8471,f8472,f8474,f8485,f8488,f8478,f8489,f8490,f8493,f8494,f8495,f8496,f8497,f8499,f8500,f8501,f8502,f8504,f8505,f8506,f8507,f8508,f8555,f8558,f8565,f8566,f8567,f8568,f8570,f8571,f8572,f8573,f8575,f8576,f8577,f8578,f8579,f8624,f8637,f8638,f8639,f8640,f8642,f8643,f8644,f8645,f8646,f8647,f8648,f8688,f8689,f8690,f8691,f8693,f8694,f8695,f8696,f8698,f8699,f8700,f8701,f8702,f8703,f8705,f8706,f8707,f8708,f8710,f8711,f8712,f8713,f8714,f8715,f8716,f8717,f8718,f8788,f8757,f8797,f8805,f8807,f8808,f8809,f8810,f8858,f8859,f8861,f8862,f8863,f8864,f8869,f8870,f8871,f8872,f8874,f8931,f8917,f8966,f8969,f8936,f8937,f8963,f8989,f8990,f8991,f8992,f8993,f8995,f8996,f8997,f8998,f8999,f9001,f9002,f9003,f9004,f9005,f9022,f9025,f9008,f9009,f9035])).
fof(f9035,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'T_0')) | ($spl38 | $spl80)),
  inference(forward_demodulation,[],[f9017,f621])).
fof(f8933,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),happ(s(fun(X2,fun(sK0,bool)),'const.combin.K_1'(s(fun(sK0,bool),sK2))),s(X2,X3)),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))))) ) | $spl164),
  inference(cnf_transformation,[],[f8933_D])).
fof(f8933_D,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),happ(s(fun(X2,fun(sK0,bool)),'const.combin.K_1'(s(fun(sK0,bool),sK2))),s(X2,X3)),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))))) ) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f8918,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),happ(s(fun(X2,fun(sK0,bool)),'const.combin.K_1'(s(fun(sK0,bool),sK2))),s(X2,X3)),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))))) ) | $spl52),
  inference(superposition,[],[f6676,f2153])).
fof(f8932,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),happ(s(fun(X0,fun(sK0,bool)),'const.combin.K_1'(s(fun(sK0,bool),sK2))),s(X0,X1)),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))))) ) | $spl162),
  inference(cnf_transformation,[],[f8932_D])).
fof(f8932_D,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),happ(s(fun(X0,fun(sK0,bool)),'const.combin.K_1'(s(fun(sK0,bool),sK2))),s(X0,X1)),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))))) ) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f8917,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),happ(s(fun(X0,fun(sK0,bool)),'const.combin.K_1'(s(fun(sK0,bool),sK2))),s(X0,X1)),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))))) ) | $spl38),
  inference(superposition,[],[f6676,f621])).
fof(f8265,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_2'(s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3))),s(X2,X3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))))) ) | $spl160),
  inference(cnf_transformation,[],[f8265_D])).
fof(f8265_D,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_2'(s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3))),s(X2,X3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))))) ) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f8214,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_2'(s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3))),s(X2,X3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))))) ) | $spl52),
  inference(superposition,[],[f6655,f2153])).
fof(f8264,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_2'(s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3))),s(X0,X1)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))))) ) | $spl158),
  inference(cnf_transformation,[],[f8264_D])).
fof(f8264_D,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_2'(s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3))),s(X0,X1)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))))) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f8213,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_2'(s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3))),s(X0,X1)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))))) ) | $spl38),
  inference(superposition,[],[f6655,f621])).
fof(f8085,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),'const.combin.K_2'(s(fun(sK0,bool),sK2),s(X2,X3)),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))))) ) | $spl156),
  inference(cnf_transformation,[],[f8085_D])).
fof(f8085_D,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),'const.combin.K_2'(s(fun(sK0,bool),sK2),s(X2,X3)),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))))) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f8070,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),'const.combin.K_2'(s(fun(sK0,bool),sK2),s(X2,X3)),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))))) ) | $spl52),
  inference(superposition,[],[f6673,f2153])).
fof(f8084,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),'const.combin.K_2'(s(fun(sK0,bool),sK2),s(X0,X1)),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))))) ) | $spl154),
  inference(cnf_transformation,[],[f8084_D])).
fof(f8084_D,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),'const.combin.K_2'(s(fun(sK0,bool),sK2),s(X0,X1)),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))))) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f8069,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),'const.combin.K_2'(s(fun(sK0,bool),sK2),s(X0,X1)),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))))) ) | $spl38),
  inference(superposition,[],[f6673,f621])).
fof(f8009,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,'const.combin.K_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(X2,X3)),sK0,sK1)),s(fun(sK0,bool),sK2))))) ) | $spl152),
  inference(cnf_transformation,[],[f8009_D])).
fof(f8009_D,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,'const.combin.K_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(X2,X3)),sK0,sK1)),s(fun(sK0,bool),sK2))))) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f7958,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,'const.combin.K_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(X2,X3)),sK0,sK1)),s(fun(sK0,bool),sK2))))) ) | $spl52),
  inference(superposition,[],[f6651,f2153])).
fof(f8008,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,'const.combin.K_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(X0,X1)),sK0,sK1)),s(fun(sK0,bool),sK2))))) ) | $spl150),
  inference(cnf_transformation,[],[f8008_D])).
fof(f8008_D,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,'const.combin.K_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(X0,X1)),sK0,sK1)),s(fun(sK0,bool),sK2))))) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f7957,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,'const.combin.K_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(X0,X1)),sK0,sK1)),s(fun(sK0,bool),sK2))))) ) | $spl38),
  inference(superposition,[],[f6651,f621])).
fof(f7206,plain,(
  $false | ($spl31 | $spl140)),
  inference(unit_resulting_resolution,[],[f70,f153,f7129])).
fof(f70,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f20])).
fof(f20,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.util_prob.FUNSET_DFUNSET',aHL_FALSITY)).
fof(f7148,plain,(
  ( ! [X39,X41,X40] : (p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK7(X39,sK2,X40,sK0,X41)))),s(fun(sK1,bool),sK3)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X41),X40),s(fun(fun(sK0,X41),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X41,bool)),X39)))))) ) | $spl148),
  inference(cnf_transformation,[],[f7148_D])).
fof(f7148_D,plain,(
  ( ! [X39,X41,X40] : (p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK7(X39,sK2,X40,sK0,X41)))),s(fun(sK1,bool),sK3)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X41),X40),s(fun(fun(sK0,X41),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X41,bool)),X39)))))) ) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f6380,plain,(
  ( ! [X27] : (s(bool,'T_0') != s(bool,X27) | ~p(s(bool,X27))) ) | $spl108),
  inference(cnf_transformation,[],[f6380_D])).
fof(f6380_D,plain,(
  ( ! [X27] : (s(bool,'T_0') != s(bool,X27) | ~p(s(bool,X27))) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f7095,plain,(
  ( ! [X39,X41,X42,X40] : (s(bool,'T_0') != s(bool,X42) | p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK7(X39,sK2,X40,sK0,X41)))),s(fun(sK1,bool),sK3)))) | ~p(s(bool,X42)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X41),X40),s(fun(fun(sK0,X41),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X41,bool)),X39)))))) ) | ($spl98 | $spl134)),
  inference(superposition,[],[f7025,f5815])).
fof(f7147,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK7(X35,sK2,X36,sK0,X37)))),s(fun(sK1,bool),sK3)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X37),X36),s(fun(fun(sK0,X37),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X37,bool)),X35)))))) ) | $spl146),
  inference(cnf_transformation,[],[f7147_D])).
fof(f7147_D,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK7(X35,sK2,X36,sK0,X37)))),s(fun(sK1,bool),sK3)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X37),X36),s(fun(fun(sK0,X37),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X37,bool)),X35)))))) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f7094,plain,(
  ( ! [X37,X35,X38,X36] : (s(bool,'T_0') != s(bool,X38) | p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK7(X35,sK2,X36,sK0,X37)))),s(fun(sK1,bool),sK3)))) | ~p(s(bool,X38)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X37),X36),s(fun(fun(sK0,X37),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X37,bool)),X35)))))) ) | ($spl96 | $spl134)),
  inference(superposition,[],[f7025,f5803])).
fof(f7131,plain,(
  ( ! [X2,X3] : (s(bool,X2) != s(bool,X3) | s(bool,'T_0') = s(bool,X2) | ~p(s(bool,X3))) ) | $spl144),
  inference(cnf_transformation,[],[f7131_D])).
fof(f7131_D,plain,(
  ( ! [X2,X3] : (s(bool,X2) != s(bool,X3) | s(bool,'T_0') = s(bool,X2) | ~p(s(bool,X3))) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f7074,plain,(
  ( ! [X2,X3] : (s(bool,X2) != s(bool,X3) | p(s(bool,'F_0')) | ~p(s(bool,X3)) | s(bool,'T_0') = s(bool,X2)) ) | $spl134),
  inference(superposition,[],[f7025,f71])).
fof(f7129,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl140),
  inference(cnf_transformation,[],[f7129_D])).
fof(f7129_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f7130,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | ~p(s(bool,X1))) ) | $spl142),
  inference(cnf_transformation,[],[f7130_D])).
fof(f7130_D,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | ~p(s(bool,X1))) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f7073,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,'T_0') = s(bool,X0)) ) | $spl134),
  inference(superposition,[],[f7025,f71])).
fof(f7118,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl136),
  inference(cnf_transformation,[],[f7118_D])).
fof(f7118_D,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f7119,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,X1))) ) | $spl138),
  inference(cnf_transformation,[],[f7119_D])).
fof(f7119_D,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,X1))) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f7046,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl134),
  inference(superposition,[],[f7025,f71])).
fof(f6381,plain,(
  ( ! [X47] : (p(s(bool,X47))) ) | $spl110),
  inference(cnf_transformation,[],[f6381_D])).
fof(f6381_D,plain,(
  ( ! [X47] : (p(s(bool,X47))) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f7025,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | p(s(bool,X1)) | ~p(s(bool,X0))) ) | $spl134),
  inference(cnf_transformation,[],[f7025_D])).
fof(f7025_D,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | p(s(bool,X1)) | ~p(s(bool,X0))) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f7027,plain,(
  ( ! [X0,X3,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,X3)) | s(bool,X0) != s(bool,X1)) )),
  inference(global_subsumption,[],[f67,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6816,f104,f102,f101,f100,f7026,f98,f96])).
fof(f6377,plain,(
  ( ! [X23,X22] : (s(bool,X22) != s(bool,X23) | ~p(s(bool,X22))) ) | $spl102),
  inference(cnf_transformation,[],[f6377_D])).
fof(f6377_D,plain,(
  ( ! [X23,X22] : (s(bool,X22) != s(bool,X23) | ~p(s(bool,X22))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f98,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X2) != s(bool,X3) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f66])).
fof(f7026,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X0) != s(bool,X1)) )),
  inference(global_subsumption,[],[f67,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6816,f104,f102,f101,f100])).
fof(f6382,plain,(
  ( ! [X46] : (~p(s(bool,X46))) ) | $spl112),
  inference(cnf_transformation,[],[f6382_D])).
fof(f6382_D,plain,(
  ( ! [X46] : (~p(s(bool,X46))) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f101,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f66])).
fof(f102,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X2) != s(bool,X3) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X2)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f66])).
fof(f104,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(duplicate_literal_removal,[],[f99])).
fof(f99,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f66])).
fof(f7004,plain,(
  $false | ($spl10 | $spl102)),
  inference(global_subsumption,[],[f67,f121,f232,f287,f6377,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6816,f6967,f6968,f6992,f6994,f6891,f6866])).
fof(f6866,plain,(
  ~p(s(bool,'T_0')) | ($spl10 | $spl102)),
  inference(unit_resulting_resolution,[],[f121,f6377])).
fof(f6891,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl102),
  inference(unit_resulting_resolution,[],[f80,f6377])).
fof(f6994,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl102),
  inference(forward_demodulation,[],[f6893,f80])).
fof(f6893,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.combin.K_2'(s(bool,X0),s(X1,X2))))) ) | $spl102),
  inference(unit_resulting_resolution,[],[f81,f6377])).
fof(f6992,plain,(
  ~p(s(bool,'T_0')) | ($spl10 | $spl102)),
  inference(forward_demodulation,[],[f6894,f121])).
fof(f6894,plain,(
  ~p(s(bool,g(s(fun(sK0,sK1),sK5('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),bool,fun(sK0,sK1)))))) | ($spl10 | $spl102)),
  inference(unit_resulting_resolution,[],[f121,f6377])).
fof(f6968,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl102),
  inference(forward_demodulation,[],[f6909,f80])).
fof(f6909,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.combin.K_2'(s(bool,X0),s(X1,X2))))) ) | $spl102),
  inference(unit_resulting_resolution,[],[f80,f6377])).
fof(f6967,plain,(
  ( ! [X1] : (~p(s(bool,X1))) ) | $spl102),
  inference(forward_demodulation,[],[f6966,f80])).
fof(f6966,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.combin.K_2'(s(bool,X1),s(X0,X2))))) ) | $spl102),
  inference(forward_demodulation,[],[f6910,f81])).
fof(f6910,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),'const.combin.K_1'(s(bool,X1))),s(X0,X2))))) ) | $spl102),
  inference(unit_resulting_resolution,[],[f81,f6377])).
fof(f232,plain,(
  s(bool,'T_0') != s(bool,f(s(fun(sK0,sK1),sK5('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),bool,fun(sK0,sK1))))) | s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))) = s(fun(fun(sK0,sK1),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3))))) | $spl10),
  inference(superposition,[],[f82,f121])).
fof(f82,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK5(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK5(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f56])).
fof(f56,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK5(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK5(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5])],[f41])).
fof(f41,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f33])).
fof(f33,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f32])).
fof(f32,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.util_prob.FUNSET_DFUNSET',aHL_EXT)).
fof(f7003,plain,(
  $false | ($spl14 | $spl102)),
  inference(global_subsumption,[],[f67,f126,f457,f520,f6377,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6816,f6967,f6968,f6990,f6994,f6891,f6868])).
fof(f6868,plain,(
  ~p(s(bool,'T_0')) | ($spl14 | $spl102)),
  inference(unit_resulting_resolution,[],[f126,f6377])).
fof(f6990,plain,(
  ~p(s(bool,'T_0')) | ($spl14 | $spl102)),
  inference(forward_demodulation,[],[f6895,f126])).
fof(f6895,plain,(
  ~p(s(bool,g(s(fun(sK0,sK1),sK5('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),bool,fun(sK0,sK1)))))) | ($spl14 | $spl102)),
  inference(unit_resulting_resolution,[],[f126,f6377])).
fof(f457,plain,(
  s(bool,'T_0') != s(bool,f(s(fun(sK0,sK1),sK5('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),bool,fun(sK0,sK1))))) | s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))) = s(fun(fun(sK0,sK1),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3))))) | $spl14),
  inference(superposition,[],[f82,f126])).
fof(f7002,plain,(
  $false | ($spl58 | $spl102)),
  inference(global_subsumption,[],[f67,f3168,f6377,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6816,f6967,f6968,f6988,f6994,f6891,f6870])).
fof(f6870,plain,(
  ~p(s(bool,'T_0')) | ($spl58 | $spl102)),
  inference(unit_resulting_resolution,[],[f3168,f6377])).
fof(f6988,plain,(
  ~p(s(bool,'T_0')) | ($spl58 | $spl102)),
  inference(forward_demodulation,[],[f6896,f3168])).
fof(f6896,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2)))) | ($spl58 | $spl102)),
  inference(unit_resulting_resolution,[],[f3168,f6377])).
fof(f7001,plain,(
  $false | ($spl78 | $spl102)),
  inference(global_subsumption,[],[f67,f5500,f6377,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6816,f6967,f6968,f6986,f6994,f6891,f6872])).
fof(f6872,plain,(
  ~p(s(bool,'T_0')) | ($spl78 | $spl102)),
  inference(unit_resulting_resolution,[],[f5500,f6377])).
fof(f6986,plain,(
  ~p(s(bool,'T_0')) | ($spl78 | $spl102)),
  inference(forward_demodulation,[],[f6897,f5500])).
fof(f6897,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2)))) | ($spl78 | $spl102)),
  inference(unit_resulting_resolution,[],[f5500,f6377])).
fof(f7000,plain,(
  $false | ($spl36 | $spl102)),
  inference(global_subsumption,[],[f67,f619,f4686,f6377,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6816,f6967,f6968,f6984,f6994,f6891,f6874])).
fof(f6874,plain,(
  ~p(s(bool,'T_0')) | ($spl36 | $spl102)),
  inference(unit_resulting_resolution,[],[f619,f6377])).
fof(f6984,plain,(
  ~p(s(bool,'T_0')) | ($spl36 | $spl102)),
  inference(forward_demodulation,[],[f6898,f619])).
fof(f6898,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)))))) | ($spl36 | $spl102)),
  inference(unit_resulting_resolution,[],[f619,f6377])).
fof(f6999,plain,(
  $false | ($spl46 | $spl102)),
  inference(global_subsumption,[],[f67,f956,f4868,f6377,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6816,f6967,f6968,f6982,f6994,f6891,f6876])).
fof(f6876,plain,(
  ~p(s(bool,'T_0')) | ($spl46 | $spl102)),
  inference(unit_resulting_resolution,[],[f956,f6377])).
fof(f6982,plain,(
  ~p(s(bool,'T_0')) | ($spl46 | $spl102)),
  inference(forward_demodulation,[],[f6899,f956])).
fof(f6899,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)))))) | ($spl46 | $spl102)),
  inference(unit_resulting_resolution,[],[f956,f6377])).
fof(f6998,plain,(
  $false | ($spl80 | $spl102)),
  inference(global_subsumption,[],[f67,f5733,f6377,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6816,f6967,f6968,f6980,f6994,f6891,f6878])).
fof(f6878,plain,(
  ~p(s(bool,'T_0')) | ($spl80 | $spl102)),
  inference(unit_resulting_resolution,[],[f5733,f6377])).
fof(f6980,plain,(
  ~p(s(bool,'T_0')) | ($spl80 | $spl102)),
  inference(forward_demodulation,[],[f6900,f5733])).
fof(f6900,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)))),s(fun(sK1,bool),sK3)))) | ($spl80 | $spl102)),
  inference(unit_resulting_resolution,[],[f5733,f6377])).
fof(f6997,plain,(
  $false | ($spl84 | $spl102)),
  inference(global_subsumption,[],[f67,f5750,f6377,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6816,f6967,f6968,f6978,f6994,f6891,f6880])).
fof(f6880,plain,(
  ~p(s(bool,'T_0')) | ($spl84 | $spl102)),
  inference(unit_resulting_resolution,[],[f5750,f6377])).
fof(f6978,plain,(
  ~p(s(bool,'T_0')) | ($spl84 | $spl102)),
  inference(forward_demodulation,[],[f6901,f5750])).
fof(f6901,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)))),s(fun(sK1,bool),sK3)))) | ($spl84 | $spl102)),
  inference(unit_resulting_resolution,[],[f5750,f6377])).
fof(f6996,plain,(
  $false | ($spl82 | $spl102)),
  inference(global_subsumption,[],[f67,f5735,f6377,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6816,f6967,f6968,f6976,f6994,f6891,f6882])).
fof(f6882,plain,(
  ~p(s(bool,'T_0')) | ($spl82 | $spl102)),
  inference(unit_resulting_resolution,[],[f5735,f6377])).
fof(f6976,plain,(
  ~p(s(bool,'T_0')) | ($spl82 | $spl102)),
  inference(forward_demodulation,[],[f6902,f5735])).
fof(f6902,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)))),s(fun(sK1,bool),sK3)))) | ($spl82 | $spl102)),
  inference(unit_resulting_resolution,[],[f5735,f6377])).
fof(f6995,plain,(
  $false | ($spl86 | $spl102)),
  inference(global_subsumption,[],[f67,f5752,f6377,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6816,f6967,f6968,f6974,f6994,f6891,f6884])).
fof(f6884,plain,(
  ~p(s(bool,'T_0')) | ($spl86 | $spl102)),
  inference(unit_resulting_resolution,[],[f5752,f6377])).
fof(f6974,plain,(
  ~p(s(bool,'T_0')) | ($spl86 | $spl102)),
  inference(forward_demodulation,[],[f6903,f5752])).
fof(f6903,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)))),s(fun(sK1,bool),sK3)))) | ($spl86 | $spl102)),
  inference(unit_resulting_resolution,[],[f5752,f6377])).
fof(f6993,plain,(
  $false | ($spl10 | $spl102)),
  inference(global_subsumption,[],[f67,f121,f232,f287,f6377,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6816,f6967,f6968,f6992])).
fof(f6991,plain,(
  $false | ($spl14 | $spl102)),
  inference(global_subsumption,[],[f67,f126,f457,f520,f6377,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6816,f6967,f6968,f6990])).
fof(f6989,plain,(
  $false | ($spl58 | $spl102)),
  inference(global_subsumption,[],[f67,f3168,f6377,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6816,f6967,f6968,f6988])).
fof(f6987,plain,(
  $false | ($spl78 | $spl102)),
  inference(global_subsumption,[],[f67,f5500,f6377,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6816,f6967,f6968,f6986])).
fof(f6985,plain,(
  $false | ($spl36 | $spl102)),
  inference(global_subsumption,[],[f67,f619,f4686,f6377,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6816,f6967,f6968,f6984])).
fof(f6983,plain,(
  $false | ($spl46 | $spl102)),
  inference(global_subsumption,[],[f67,f956,f4868,f6377,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6816,f6967,f6968,f6982])).
fof(f6981,plain,(
  $false | ($spl80 | $spl102)),
  inference(global_subsumption,[],[f67,f5733,f6377,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6816,f6967,f6968,f6980])).
fof(f6979,plain,(
  $false | ($spl84 | $spl102)),
  inference(global_subsumption,[],[f67,f5750,f6377,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6816,f6967,f6968,f6978])).
fof(f6977,plain,(
  $false | ($spl82 | $spl102)),
  inference(global_subsumption,[],[f67,f5735,f6377,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6816,f6967,f6968,f6976])).
fof(f6975,plain,(
  $false | ($spl86 | $spl102)),
  inference(global_subsumption,[],[f67,f5752,f6377,f103,f84,f83,f86,f87,f85,f88,f6616,f6636,f6637,f6640,f6620,f6641,f6626,f6627,f6635,f6651,f6672,f6673,f6676,f6655,f6677,f6661,f6662,f6671,f6686,f6687,f6702,f6705,f6690,f6691,f6706,f6707,f6710,f6711,f6712,f6713,f6746,f6749,f6739,f6750,f6751,f6754,f6743,f6755,f6756,f6757,f6758,f6783,f6784,f6816,f6967,f6968,f6974])).
fof(f6892,plain,(
  $false | ($spl48 | $spl102)),
  inference(unit_resulting_resolution,[],[f1124,f80,f6377])).
fof(f6885,plain,(
  $false | ($spl48 | $spl86 | $spl102)),
  inference(unit_resulting_resolution,[],[f1124,f5752,f6377])).
fof(f6883,plain,(
  $false | ($spl48 | $spl82 | $spl102)),
  inference(unit_resulting_resolution,[],[f1124,f5735,f6377])).
fof(f6881,plain,(
  $false | ($spl48 | $spl84 | $spl102)),
  inference(unit_resulting_resolution,[],[f1124,f5750,f6377])).
fof(f6879,plain,(
  $false | ($spl48 | $spl80 | $spl102)),
  inference(unit_resulting_resolution,[],[f1124,f5733,f6377])).
fof(f6877,plain,(
  $false | ($spl46 | $spl48 | $spl102)),
  inference(unit_resulting_resolution,[],[f1124,f956,f6377])).
fof(f6875,plain,(
  $false | ($spl36 | $spl48 | $spl102)),
  inference(unit_resulting_resolution,[],[f1124,f619,f6377])).
fof(f6873,plain,(
  $false | ($spl48 | $spl78 | $spl102)),
  inference(unit_resulting_resolution,[],[f1124,f5500,f6377])).
fof(f6871,plain,(
  $false | ($spl48 | $spl58 | $spl102)),
  inference(unit_resulting_resolution,[],[f1124,f3168,f6377])).
fof(f6869,plain,(
  $false | ($spl14 | $spl48 | $spl102)),
  inference(unit_resulting_resolution,[],[f1124,f126,f6377])).
fof(f6867,plain,(
  $false | ($spl10 | $spl48 | $spl102)),
  inference(unit_resulting_resolution,[],[f1124,f121,f6377])).
fof(f6864,plain,(
  ( ! [X24,X23,X21,X22,X20] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(X20,X24),X23),s(fun(fun(X20,X24),bool),'const.util_prob.FUNSET_2'(s(fun(X20,bool),X22),s(fun(X24,bool),X21))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X20,sK6(X21,X22,X23,X24,X20)),s(fun(X20,bool),X22)))) ) | $spl132),
  inference(cnf_transformation,[],[f6864_D])).
fof(f6864_D,plain,(
  ( ! [X24,X23,X21,X22,X20] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(X20,X24),X23),s(fun(fun(X20,X24),bool),'const.util_prob.FUNSET_2'(s(fun(X20,bool),X22),s(fun(X24,bool),X21))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X20,sK6(X21,X22,X23,X24,X20)),s(fun(X20,bool),X22)))) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f6857,plain,(
  ( ! [X24,X23,X21,X22,X20] : (p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(X20,X24),X23),s(fun(fun(X20,X24),bool),'const.util_prob.FUNSET_2'(s(fun(X20,bool),X22),s(fun(X24,bool),X21))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X20,sK6(X21,X22,X23,X24,X20)),s(fun(X20,bool),X22)))) ) | $spl122),
  inference(superposition,[],[f6642,f71])).
fof(f6842,plain,(
  ( ! [X6,X10,X8,X7,X9] : (~p(s(bool,'const.bool.IN_2'(s(X7,X8),s(fun(X7,bool),X10)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(X6,X7),'const.combin.K_1'(s(X7,X8))),s(fun(fun(X6,X7),bool),'const.util_prob.FUNSET_2'(s(fun(X6,bool),X9),s(fun(X7,bool),X10)))))) ) | $spl130),
  inference(cnf_transformation,[],[f6842_D])).
fof(f6842_D,plain,(
  ( ! [X6,X10,X8,X7,X9] : (~p(s(bool,'const.bool.IN_2'(s(X7,X8),s(fun(X7,bool),X10)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(X6,X7),'const.combin.K_1'(s(X7,X8))),s(fun(fun(X6,X7),bool),'const.util_prob.FUNSET_2'(s(fun(X6,bool),X9),s(fun(X7,bool),X10)))))) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f6816,plain,(
  ( ! [X6,X10,X8,X7,X9] : (p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(X7,X8),s(fun(X7,bool),X10)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(X6,X7),'const.combin.K_1'(s(X7,X8))),s(fun(fun(X6,X7),bool),'const.util_prob.FUNSET_2'(s(fun(X6,bool),X9),s(fun(X7,bool),X10)))))) )),
  inference(superposition,[],[f6706,f71])).
fof(f6685,plain,(
  ( ! [X24,X23,X21,X22,X20] : (p(s(bool,'const.bool.IN_2'(s(fun(X20,X24),X23),s(fun(fun(X20,X24),bool),'const.util_prob.DFUNSET_2'(s(fun(X20,bool),X22),s(fun(X20,fun(X24,bool)),X21)))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X20,sK7(X21,X22,X23,X20,X24)),s(fun(X20,bool),X22)))) ) | $spl128),
  inference(cnf_transformation,[],[f6685_D])).
fof(f6685_D,plain,(
  ( ! [X24,X23,X21,X22,X20] : (p(s(bool,'const.bool.IN_2'(s(fun(X20,X24),X23),s(fun(fun(X20,X24),bool),'const.util_prob.DFUNSET_2'(s(fun(X20,bool),X22),s(fun(X20,fun(X24,bool)),X21)))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X20,sK7(X21,X22,X23,X20,X24)),s(fun(X20,bool),X22)))) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f6671,plain,(
  ( ! [X24,X23,X21,X22,X20] : (p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(fun(X20,X24),X23),s(fun(fun(X20,X24),bool),'const.util_prob.DFUNSET_2'(s(fun(X20,bool),X22),s(fun(X20,fun(X24,bool)),X21)))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X20,sK7(X21,X22,X23,X20,X24)),s(fun(X20,bool),X22)))) )),
  inference(superposition,[],[f87,f71])).
fof(f6680,plain,(
  ( ! [X6,X10,X8,X7,X9] : (p(s(bool,'const.bool.IN_2'(s(X6,sK7(X10,X9,X8,X6,X7)),s(fun(X6,bool),X9)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(X6,X7),X8),s(fun(fun(X6,X7),bool),'const.util_prob.DFUNSET_2'(s(fun(X6,bool),X9),s(fun(X6,fun(X7,bool)),X10)))))) ) | $spl126),
  inference(cnf_transformation,[],[f6680_D])).
fof(f6680_D,plain,(
  ( ! [X6,X10,X8,X7,X9] : (p(s(bool,'const.bool.IN_2'(s(X6,sK7(X10,X9,X8,X6,X7)),s(fun(X6,bool),X9)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(X6,X7),X8),s(fun(fun(X6,X7),bool),'const.util_prob.DFUNSET_2'(s(fun(X6,bool),X9),s(fun(X6,fun(X7,bool)),X10)))))) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f6661,plain,(
  ( ! [X6,X10,X8,X7,X9] : (p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(X6,sK7(X10,X9,X8,X6,X7)),s(fun(X6,bool),X9)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(X6,X7),X8),s(fun(fun(X6,X7),bool),'const.util_prob.DFUNSET_2'(s(fun(X6,bool),X9),s(fun(X6,fun(X7,bool)),X10)))))) )),
  inference(superposition,[],[f87,f71])).
fof(f6648,plain,(
  ( ! [X24,X23,X21,X22,X20] : (p(s(bool,'const.bool.IN_2'(s(fun(X20,X24),X23),s(fun(fun(X20,X24),bool),'const.util_prob.FUNSET_2'(s(fun(X20,bool),X22),s(fun(X24,bool),X21)))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X20,sK6(X21,X22,X23,X24,X20)),s(fun(X20,bool),X22)))) ) | $spl124),
  inference(cnf_transformation,[],[f6648_D])).
fof(f6648_D,plain,(
  ( ! [X24,X23,X21,X22,X20] : (p(s(bool,'const.bool.IN_2'(s(fun(X20,X24),X23),s(fun(fun(X20,X24),bool),'const.util_prob.FUNSET_2'(s(fun(X20,bool),X22),s(fun(X24,bool),X21)))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X20,sK6(X21,X22,X23,X24,X20)),s(fun(X20,bool),X22)))) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f6635,plain,(
  ( ! [X24,X23,X21,X22,X20] : (p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(fun(X20,X24),X23),s(fun(fun(X20,X24),bool),'const.util_prob.FUNSET_2'(s(fun(X20,bool),X22),s(fun(X24,bool),X21)))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X20,sK6(X21,X22,X23,X24,X20)),s(fun(X20,bool),X22)))) )),
  inference(superposition,[],[f84,f71])).
fof(f6642,plain,(
  ( ! [X6,X10,X8,X7,X9] : (p(s(bool,'const.bool.IN_2'(s(X6,sK6(X10,X9,X8,X7,X6)),s(fun(X6,bool),X9)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(X6,X7),X8),s(fun(fun(X6,X7),bool),'const.util_prob.FUNSET_2'(s(fun(X6,bool),X9),s(fun(X7,bool),X10)))))) ) | $spl122),
  inference(cnf_transformation,[],[f6642_D])).
fof(f6642_D,plain,(
  ( ! [X6,X10,X8,X7,X9] : (p(s(bool,'const.bool.IN_2'(s(X6,sK6(X10,X9,X8,X7,X6)),s(fun(X6,bool),X9)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(X6,X7),X8),s(fun(fun(X6,X7),bool),'const.util_prob.FUNSET_2'(s(fun(X6,bool),X9),s(fun(X7,bool),X10)))))) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f6626,plain,(
  ( ! [X6,X10,X8,X7,X9] : (p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(X6,sK6(X10,X9,X8,X7,X6)),s(fun(X6,bool),X9)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(X6,X7),X8),s(fun(fun(X6,X7),bool),'const.util_prob.FUNSET_2'(s(fun(X6,bool),X9),s(fun(X7,bool),X10)))))) )),
  inference(superposition,[],[f84,f71])).
fof(f1124,plain,(
  p(s(bool,'T_0')) | $spl48),
  inference(cnf_transformation,[],[f1124_D])).
fof(f1124_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f103,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f72])).
fof(f72,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f22])).
fof(f22,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f21])).
fof(f21,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6] : ! [X8] : (s(X6,X8) = s(X6,X8) <=> p(s(bool,'T_0')))),
  file('thm.util_prob.FUNSET_DFUNSET','thm.bool.REFL_CLAUSE')).
fof(f6575,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f70,f6381])).
fof(f6580,plain,(
  $false | $spl110),
  inference(resolution,[],[f6381,f70])).
fof(f6502,plain,(
  ( ! [X54,X52,X53] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK0,X52),X53),s(fun(fun(sK0,X52),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X52,bool),X54)))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK6(X54,sK2,X53,X52,sK0)))),s(fun(sK1,bool),sK3)))) ) | $spl120),
  inference(cnf_transformation,[],[f6502_D])).
fof(f6502_D,plain,(
  ( ! [X54,X52,X53] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK0,X52),X53),s(fun(fun(sK0,X52),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X52,bool),X54)))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK6(X54,sK2,X53,X52,sK0)))),s(fun(sK1,bool),sK3)))) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f6378,plain,(
  ( ! [X21] : (s(bool,'T_0') != s(bool,X21) | p(s(bool,X21))) ) | $spl104),
  inference(cnf_transformation,[],[f6378_D])).
fof(f6378_D,plain,(
  ( ! [X21] : (s(bool,'T_0') != s(bool,X21) | p(s(bool,X21))) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f6487,plain,(
  ( ! [X103,X101,X102,X100,X104] : (s(bool,'T_0') != s(bool,X103) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,X100),X101),s(fun(fun(sK0,X100),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X100,bool),X102)))))) | p(s(bool,X103)) | p(s(bool,X104)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK6(X102,sK2,X101,X100,sK0)))),s(fun(sK1,bool),sK3)))) ) | $spl94),
  inference(superposition,[],[f104,f5791])).
fof(f6486,plain,(
  ( ! [X94,X99,X97,X95,X98,X96] : (s(bool,'T_0') != s(bool,X97) | ~p(s(bool,X98)) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,X94),X95),s(fun(fun(sK0,X94),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X94,bool),X96)))))) | p(s(bool,X97)) | s(bool,X98) != s(bool,X99) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK6(X96,sK2,X95,X94,sK0)))),s(fun(sK1,bool),sK3)))) ) | $spl94),
  inference(superposition,[],[f102,f5791])).
fof(f6483,plain,(
  ( ! [X80,X78,X76,X81,X79,X77] : (s(bool,'T_0') != s(bool,X79) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,X76),X77),s(fun(fun(sK0,X76),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X76,bool),X78)))))) | ~p(s(bool,X80)) | p(s(bool,X79)) | p(s(bool,X81)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK6(X78,sK2,X77,X76,sK0)))),s(fun(sK1,bool),sK3)))) ) | $spl94),
  inference(superposition,[],[f101,f5791])).
fof(f6479,plain,(
  ( ! [X57,X54,X52,X56,X55,X53] : (s(bool,'T_0') != s(bool,X55) | ~p(s(bool,X56)) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,X52),X53),s(fun(fun(sK0,X52),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X52,bool),X54)))))) | p(s(bool,X55)) | s(bool,X56) != s(bool,X57) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK6(X54,sK2,X53,X52,sK0)))),s(fun(sK1,bool),sK3)))) ) | $spl94),
  inference(superposition,[],[f98,f5791])).
fof(f6501,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK6(X24,sK2,X25,X26,sK0)))),s(fun(sK1,bool),sK3)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X26),X25),s(fun(fun(sK0,X26),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X26,bool),X24)))))) ) | $spl118),
  inference(cnf_transformation,[],[f6501_D])).
fof(f6501_D,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK6(X24,sK2,X25,X26,sK0)))),s(fun(sK1,bool),sK3)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X26),X25),s(fun(fun(sK0,X26),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X26,bool),X24)))))) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f6456,plain,(
  ( ! [X74,X72,X71,X75,X73] : (s(bool,'T_0') != s(bool,X74) | ~p(s(bool,X74)) | p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK6(X71,sK2,X72,X73,sK0)))),s(fun(sK1,bool),sK3)))) | p(s(bool,X75)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X73),X72),s(fun(fun(sK0,X73),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X73,bool),X71)))))) ) | $spl94),
  inference(superposition,[],[f104,f5791])).
fof(f6500,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK6(X18,sK2,X19,X20,sK0)))),s(fun(sK1,bool),sK3)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X20),X19),s(fun(fun(sK0,X20),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X20,bool),X18)))))) ) | $spl116),
  inference(cnf_transformation,[],[f6500_D])).
fof(f6500_D,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK6(X18,sK2,X19,X20,sK0)))),s(fun(sK1,bool),sK3)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X20),X19),s(fun(fun(sK0,X20),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X20,bool),X18)))))) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f6455,plain,(
  ( ! [X70,X68,X66,X69,X67] : (s(bool,'T_0') != s(bool,X69) | ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK6(X66,sK2,X67,X68,sK0)))),s(fun(sK1,bool),sK3)))) | p(s(bool,X69)) | p(s(bool,X70)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X68),X67),s(fun(fun(sK0,X68),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X68,bool),X66)))))) ) | $spl94),
  inference(superposition,[],[f104,f5791])).
fof(f6454,plain,(
  ( ! [X61,X64,X62,X60,X65,X63] : (s(bool,'T_0') != s(bool,X63) | ~p(s(bool,X64)) | ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK6(X60,sK2,X61,X62,sK0)))),s(fun(sK1,bool),sK3)))) | p(s(bool,X63)) | s(bool,X64) != s(bool,X65) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X62),X61),s(fun(fun(sK0,X62),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X62,bool),X60)))))) ) | $spl94),
  inference(superposition,[],[f102,f5791])).
fof(f6453,plain,(
  ( ! [X59,X57,X54,X58,X56,X55] : (s(bool,'T_0') != s(bool,X57) | ~p(s(bool,X58)) | ~p(s(bool,X57)) | p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK6(X54,sK2,X55,X56,sK0)))),s(fun(sK1,bool),sK3)))) | s(bool,X58) != s(bool,X59) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X56),X55),s(fun(fun(sK0,X56),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X56,bool),X54)))))) ) | $spl94),
  inference(superposition,[],[f102,f5791])).
fof(f6452,plain,(
  ( ! [X52,X50,X48,X53,X51,X49] : (s(bool,'T_0') != s(bool,X51) | ~p(s(bool,X51)) | ~p(s(bool,X52)) | p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK6(X48,sK2,X49,X50,sK0)))),s(fun(sK1,bool),sK3)))) | p(s(bool,X53)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X50),X49),s(fun(fun(sK0,X50),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X50,bool),X48)))))) ) | $spl94),
  inference(superposition,[],[f101,f5791])).
fof(f6451,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(bool,'T_0') != s(bool,X45) | ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK6(X42,sK2,X43,X44,sK0)))),s(fun(sK1,bool),sK3)))) | ~p(s(bool,X46)) | p(s(bool,X45)) | p(s(bool,X47)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X44),X43),s(fun(fun(sK0,X44),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X44,bool),X42)))))) ) | $spl94),
  inference(superposition,[],[f101,f5791])).
fof(f6448,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (s(bool,'T_0') != s(bool,X27) | ~p(s(bool,X28)) | ~p(s(bool,X27)) | p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK6(X24,sK2,X25,X26,sK0)))),s(fun(sK1,bool),sK3)))) | s(bool,X28) != s(bool,X29) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X26),X25),s(fun(fun(sK0,X26),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X26,bool),X24)))))) ) | $spl94),
  inference(superposition,[],[f98,f5791])).
fof(f6447,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,'T_0') != s(bool,X21) | ~p(s(bool,X22)) | ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK6(X18,sK2,X19,X20,sK0)))),s(fun(sK1,bool),sK3)))) | p(s(bool,X21)) | s(bool,X22) != s(bool,X23) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X20),X19),s(fun(fun(sK0,X20),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X20,bool),X18)))))) ) | $spl94),
  inference(superposition,[],[f98,f5791])).
fof(f6383,plain,(
  ( ! [X50,X51,X49] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK0,X49),X50),s(fun(fun(sK0,X49),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X49,bool),X51)))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK6(X51,sK2,X50,X49,sK0)))),s(fun(sK1,bool),sK3)))) ) | $spl114),
  inference(cnf_transformation,[],[f6383_D])).
fof(f6383_D,plain,(
  ( ! [X50,X51,X49] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK0,X49),X50),s(fun(fun(sK0,X49),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X49,bool),X51)))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK6(X51,sK2,X50,X49,sK0)))),s(fun(sK1,bool),sK3)))) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f6365,plain,(
  ( ! [X101,X99,X97,X100,X98] : (s(bool,'T_0') != s(bool,X100) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,X97),X98),s(fun(fun(sK0,X97),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X97,bool),X99)))))) | p(s(bool,X100)) | p(s(bool,X101)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK6(X99,sK2,X98,X97,sK0)))),s(fun(sK1,bool),sK3)))) ) | $spl92),
  inference(superposition,[],[f104,f5779])).
fof(f6364,plain,(
  ( ! [X94,X92,X95,X93,X91,X96] : (s(bool,'T_0') != s(bool,X94) | ~p(s(bool,X95)) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,X91),X92),s(fun(fun(sK0,X91),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X91,bool),X93)))))) | p(s(bool,X94)) | s(bool,X95) != s(bool,X96) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK6(X93,sK2,X92,X91,sK0)))),s(fun(sK1,bool),sK3)))) ) | $spl92),
  inference(superposition,[],[f102,f5779])).
fof(f6361,plain,(
  ( ! [X78,X76,X74,X77,X75,X73] : (s(bool,'T_0') != s(bool,X76) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,X73),X74),s(fun(fun(sK0,X73),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X73,bool),X75)))))) | ~p(s(bool,X77)) | p(s(bool,X76)) | p(s(bool,X78)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK6(X75,sK2,X74,X73,sK0)))),s(fun(sK1,bool),sK3)))) ) | $spl92),
  inference(superposition,[],[f101,f5779])).
fof(f6357,plain,(
  ( ! [X54,X52,X50,X53,X51,X49] : (s(bool,'T_0') != s(bool,X52) | ~p(s(bool,X53)) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,X49),X50),s(fun(fun(sK0,X49),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X49,bool),X51)))))) | p(s(bool,X52)) | s(bool,X53) != s(bool,X54) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK6(X51,sK2,X50,X49,sK0)))),s(fun(sK1,bool),sK3)))) ) | $spl92),
  inference(superposition,[],[f98,f5779])).
fof(f6379,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK6(X24,sK2,X25,X26,sK0)))),s(fun(sK1,bool),sK3)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X26),X25),s(fun(fun(sK0,X26),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X26,bool),X24)))))) ) | $spl106),
  inference(cnf_transformation,[],[f6379_D])).
fof(f6379_D,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK6(X24,sK2,X25,X26,sK0)))),s(fun(sK1,bool),sK3)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X26),X25),s(fun(fun(sK0,X26),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X26,bool),X24)))))) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f6336,plain,(
  ( ! [X74,X72,X71,X75,X73] : (s(bool,'T_0') != s(bool,X74) | ~p(s(bool,X74)) | p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK6(X71,sK2,X72,X73,sK0)))),s(fun(sK1,bool),sK3)))) | p(s(bool,X75)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X73),X72),s(fun(fun(sK0,X73),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X73,bool),X71)))))) ) | $spl92),
  inference(superposition,[],[f104,f5779])).
fof(f6376,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK6(X18,sK2,X19,X20,sK0)))),s(fun(sK1,bool),sK3)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X20),X19),s(fun(fun(sK0,X20),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X20,bool),X18)))))) ) | $spl100),
  inference(cnf_transformation,[],[f6376_D])).
fof(f6376_D,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK6(X18,sK2,X19,X20,sK0)))),s(fun(sK1,bool),sK3)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X20),X19),s(fun(fun(sK0,X20),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X20,bool),X18)))))) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f6335,plain,(
  ( ! [X70,X68,X66,X69,X67] : (s(bool,'T_0') != s(bool,X69) | ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK6(X66,sK2,X67,X68,sK0)))),s(fun(sK1,bool),sK3)))) | p(s(bool,X69)) | p(s(bool,X70)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X68),X67),s(fun(fun(sK0,X68),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X68,bool),X66)))))) ) | $spl92),
  inference(superposition,[],[f104,f5779])).
fof(f6334,plain,(
  ( ! [X61,X64,X62,X60,X65,X63] : (s(bool,'T_0') != s(bool,X63) | ~p(s(bool,X64)) | ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK6(X60,sK2,X61,X62,sK0)))),s(fun(sK1,bool),sK3)))) | p(s(bool,X63)) | s(bool,X64) != s(bool,X65) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X62),X61),s(fun(fun(sK0,X62),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X62,bool),X60)))))) ) | $spl92),
  inference(superposition,[],[f102,f5779])).
fof(f6333,plain,(
  ( ! [X59,X57,X54,X58,X56,X55] : (s(bool,'T_0') != s(bool,X57) | ~p(s(bool,X58)) | ~p(s(bool,X57)) | p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK6(X54,sK2,X55,X56,sK0)))),s(fun(sK1,bool),sK3)))) | s(bool,X58) != s(bool,X59) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X56),X55),s(fun(fun(sK0,X56),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X56,bool),X54)))))) ) | $spl92),
  inference(superposition,[],[f102,f5779])).
fof(f6332,plain,(
  ( ! [X52,X50,X48,X53,X51,X49] : (s(bool,'T_0') != s(bool,X51) | ~p(s(bool,X51)) | ~p(s(bool,X52)) | p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK6(X48,sK2,X49,X50,sK0)))),s(fun(sK1,bool),sK3)))) | p(s(bool,X53)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X50),X49),s(fun(fun(sK0,X50),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X50,bool),X48)))))) ) | $spl92),
  inference(superposition,[],[f101,f5779])).
fof(f6331,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(bool,'T_0') != s(bool,X45) | ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK6(X42,sK2,X43,X44,sK0)))),s(fun(sK1,bool),sK3)))) | ~p(s(bool,X46)) | p(s(bool,X45)) | p(s(bool,X47)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X44),X43),s(fun(fun(sK0,X44),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X44,bool),X42)))))) ) | $spl92),
  inference(superposition,[],[f101,f5779])).
fof(f6328,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (s(bool,'T_0') != s(bool,X27) | ~p(s(bool,X28)) | ~p(s(bool,X27)) | p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK6(X24,sK2,X25,X26,sK0)))),s(fun(sK1,bool),sK3)))) | s(bool,X28) != s(bool,X29) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X26),X25),s(fun(fun(sK0,X26),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X26,bool),X24)))))) ) | $spl92),
  inference(superposition,[],[f98,f5779])).
fof(f6327,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,'T_0') != s(bool,X21) | ~p(s(bool,X22)) | ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK6(X18,sK2,X19,X20,sK0)))),s(fun(sK1,bool),sK3)))) | p(s(bool,X21)) | s(bool,X22) != s(bool,X23) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X20),X19),s(fun(fun(sK0,X20),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X20,bool),X18)))))) ) | $spl92),
  inference(superposition,[],[f98,f5779])).
fof(f5815,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK7(X2,sK2,X1,sK0,X0)))),s(fun(sK1,bool),sK3))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X0),X1),s(fun(fun(sK0,X0),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X0,bool)),X2)))))) ) | $spl98),
  inference(cnf_transformation,[],[f5815_D])).
fof(f5815_D,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK7(X2,sK2,X1,sK0,X0)))),s(fun(sK1,bool),sK3))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X0),X1),s(fun(fun(sK0,X0),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X0,bool)),X2)))))) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f5812,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK7(X2,sK2,X1,sK0,X0)))),s(fun(sK1,bool),sK3))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X0),X1),s(fun(fun(sK0,X0),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X0,bool)),X2)))))) ) | $spl68),
  inference(superposition,[],[f5738,f71])).
fof(f5738,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,X5),X4),s(fun(fun(sK0,X5),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X5,bool)),X3)))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK7(X3,sK2,X4,sK0,X5)))),s(fun(sK1,bool),sK3)))) ) | $spl68),
  inference(resolution,[],[f4982,f87])).
fof(f5803,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK7(X2,sK2,X1,sK0,X0)))),s(fun(sK1,bool),sK3))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X0),X1),s(fun(fun(sK0,X0),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X0,bool)),X2)))))) ) | $spl96),
  inference(cnf_transformation,[],[f5803_D])).
fof(f5803_D,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK7(X2,sK2,X1,sK0,X0)))),s(fun(sK1,bool),sK3))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X0),X1),s(fun(fun(sK0,X0),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X0,bool)),X2)))))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f5800,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK7(X2,sK2,X1,sK0,X0)))),s(fun(sK1,bool),sK3))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X0),X1),s(fun(fun(sK0,X0),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X0,bool)),X2)))))) ) | $spl66),
  inference(superposition,[],[f5721,f71])).
fof(f5721,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,X5),X4),s(fun(fun(sK0,X5),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(X5,bool)),X3)))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK7(X3,sK2,X4,sK0,X5)))),s(fun(sK1,bool),sK3)))) ) | $spl66),
  inference(resolution,[],[f4973,f87])).
fof(f5791,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK6(X2,sK2,X1,X0,sK0)))),s(fun(sK1,bool),sK3))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X0),X1),s(fun(fun(sK0,X0),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X0,bool),X2)))))) ) | $spl94),
  inference(cnf_transformation,[],[f5791_D])).
fof(f5791_D,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK6(X2,sK2,X1,X0,sK0)))),s(fun(sK1,bool),sK3))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X0),X1),s(fun(fun(sK0,X0),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X0,bool),X2)))))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f5788,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK6(X2,sK2,X1,X0,sK0)))),s(fun(sK1,bool),sK3))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X0),X1),s(fun(fun(sK0,X0),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X0,bool),X2)))))) ) | $spl68),
  inference(superposition,[],[f5737,f71])).
fof(f5737,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,X2),X1),s(fun(fun(sK0,X2),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X2,bool),X0)))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK6(X0,sK2,X1,X2,sK0)))),s(fun(sK1,bool),sK3)))) ) | $spl68),
  inference(resolution,[],[f4982,f84])).
fof(f5779,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK6(X2,sK2,X1,X0,sK0)))),s(fun(sK1,bool),sK3))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X0),X1),s(fun(fun(sK0,X0),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X0,bool),X2)))))) ) | $spl92),
  inference(cnf_transformation,[],[f5779_D])).
fof(f5779_D,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK6(X2,sK2,X1,X0,sK0)))),s(fun(sK1,bool),sK3))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X0),X1),s(fun(fun(sK0,X0),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X0,bool),X2)))))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f5776,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK6(X2,sK2,X1,X0,sK0)))),s(fun(sK1,bool),sK3))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,X0),X1),s(fun(fun(sK0,X0),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X0,bool),X2)))))) ) | $spl66),
  inference(superposition,[],[f5720,f71])).
fof(f5720,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,X2),X1),s(fun(fun(sK0,X2),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X2,bool),X0)))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK6(X0,sK2,X1,X2,sK0)))),s(fun(sK1,bool),sK3)))) ) | $spl66),
  inference(resolution,[],[f4973,f84])).
fof(f5767,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,X0,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),X0),s(fun(sK1,bool),sK3)))))) ) | $spl90),
  inference(cnf_transformation,[],[f5767_D])).
fof(f5767_D,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,X0,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),X0),s(fun(sK1,bool),sK3)))))) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f5765,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,X0,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),X0),s(fun(sK1,bool),sK3)))))) ) | $spl50),
  inference(superposition,[],[f5712,f71])).
fof(f5712,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),X0),s(fun(sK1,bool),sK3)))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,X0,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))))) ) | $spl50),
  inference(resolution,[],[f1126,f85])).
fof(f5760,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,X0,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),X0),s(fun(sK1,bool),sK3)))))) ) | $spl88),
  inference(cnf_transformation,[],[f5760_D])).
fof(f5760_D,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,X0,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),X0),s(fun(sK1,bool),sK3)))))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f5758,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,X0,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),X0),s(fun(sK1,bool),sK3)))))) ) | $spl54),
  inference(superposition,[],[f4974,f71])).
fof(f4974,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),X0),s(fun(sK1,bool),sK3)))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,X0,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))))) ) | $spl54),
  inference(resolution,[],[f2441,f85])).
fof(f5752,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)))),s(fun(sK1,bool),sK3))) | $spl86),
  inference(cnf_transformation,[],[f5752_D])).
fof(f5752_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)))),s(fun(sK1,bool),sK3))) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f1125,plain,(
  ~p(s(bool,'T_0')) | $spl49),
  inference(cnf_transformation,[],[f1125_D])).
fof(f1125_D,plain,(
  ~p(s(bool,'T_0')) <=> ~$spl49),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl49])])).
fof(f5742,plain,(
  ~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)))),s(fun(sK1,bool),sK3))) | ($spl68 | $spl78)),
  inference(superposition,[],[f4982,f5500])).
fof(f5750,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)))),s(fun(sK1,bool),sK3))) | $spl84),
  inference(cnf_transformation,[],[f5750_D])).
fof(f5750_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)))),s(fun(sK1,bool),sK3))) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f5741,plain,(
  ~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)))),s(fun(sK1,bool),sK3))) | ($spl58 | $spl68)),
  inference(superposition,[],[f4982,f3168])).
fof(f5735,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)))),s(fun(sK1,bool),sK3))) | $spl82),
  inference(cnf_transformation,[],[f5735_D])).
fof(f5735_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)))),s(fun(sK1,bool),sK3))) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f5725,plain,(
  ~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)))),s(fun(sK1,bool),sK3))) | ($spl66 | $spl78)),
  inference(superposition,[],[f4973,f5500])).
fof(f5733,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)))),s(fun(sK1,bool),sK3))) | $spl80),
  inference(cnf_transformation,[],[f5733_D])).
fof(f5733_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)))),s(fun(sK1,bool),sK3))) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f5724,plain,(
  ~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)))),s(fun(sK1,bool),sK3))) | ($spl58 | $spl66)),
  inference(superposition,[],[f4973,f3168])).
fof(f1126,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,X0))),s(fun(sK1,bool),sK3)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK2))))) ) | $spl50),
  inference(cnf_transformation,[],[f1126_D])).
fof(f1126_D,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,X0))),s(fun(sK1,bool),sK3)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK2))))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f4686,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK2)))) | p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,X0))),s(fun(sK1,bool),sK3))))) ) | $spl36),
  inference(superposition,[],[f83,f619])).
fof(f5500,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl78),
  inference(cnf_transformation,[],[f5500_D])).
fof(f5500_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f5412,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl70),
  inference(superposition,[],[f5114,f71])).
fof(f5498,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)))),s(fun(sK1,bool),sK3)))) | $spl76),
  inference(cnf_transformation,[],[f5498_D])).
fof(f5498_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)))),s(fun(sK1,bool),sK3)))) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f5269,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)))),s(fun(sK1,bool),sK3)))) | ($spl54 | $spl70)),
  inference(unit_resulting_resolution,[],[f5114,f2441])).
fof(f5496,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)))),s(fun(sK1,bool),sK3)))) | $spl74),
  inference(cnf_transformation,[],[f5496_D])).
fof(f5496_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)))),s(fun(sK1,bool),sK3)))) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f5270,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)))),s(fun(sK1,bool),sK3)))) | ($spl50 | $spl70)),
  inference(unit_resulting_resolution,[],[f5114,f1126])).
fof(f5415,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl73),
  inference(cnf_transformation,[],[f5415_D])).
fof(f5415_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) <=> ~$spl73),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl73])])).
fof(f5271,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl12 | $spl70)),
  inference(unit_resulting_resolution,[],[f70,f70,f122,f5114,f96])).
fof(f122,plain,(
  s(bool,'F_0') = s(bool,f(s(fun(sK0,sK1),sK5('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),bool,fun(sK0,sK1))))) | $spl12),
  inference(cnf_transformation,[],[f122_D])).
fof(f122_D,plain,(
  s(bool,'F_0') = s(bool,f(s(fun(sK0,sK1),sK5('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),bool,fun(sK0,sK1))))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f5272,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl16 | $spl70)),
  inference(unit_resulting_resolution,[],[f70,f70,f127,f5114,f96])).
fof(f127,plain,(
  s(bool,'F_0') = s(bool,f(s(fun(sK0,sK1),sK5('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),bool,fun(sK0,sK1))))) | $spl16),
  inference(cnf_transformation,[],[f127_D])).
fof(f127_D,plain,(
  s(bool,'F_0') = s(bool,f(s(fun(sK0,sK1),sK5('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),bool,fun(sK0,sK1))))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f5273,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl38 | $spl70)),
  inference(unit_resulting_resolution,[],[f70,f70,f621,f5114,f96])).
fof(f5274,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl52 | $spl70)),
  inference(unit_resulting_resolution,[],[f70,f70,f2153,f5114,f96])).
fof(f5275,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl62 | $spl70)),
  inference(unit_resulting_resolution,[],[f70,f70,f3944,f5114,f96])).
fof(f5276,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl70),
  inference(unit_resulting_resolution,[],[f70,f70,f80,f5114,f96])).
fof(f5286,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl70),
  inference(unit_resulting_resolution,[],[f5114,f70,f80,f5114,f98])).
fof(f5302,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl10 | $spl70)),
  inference(unit_resulting_resolution,[],[f70,f103,f121,f5114,f98])).
fof(f5303,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl10 | $spl70)),
  inference(unit_resulting_resolution,[],[f70,f69,f121,f5114,f98])).
fof(f69,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  p(s(bool,'T_0'))),
  file('thm.util_prob.FUNSET_DFUNSET','thm.bool.TRUTH')).
fof(f5304,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl10 | $spl70)),
  inference(unit_resulting_resolution,[],[f70,f68,f121,f5114,f98])).
fof(f68,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.util_prob.FUNSET_DFUNSET',aHL_TRUTH)).
fof(f5305,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl14 | $spl70)),
  inference(unit_resulting_resolution,[],[f70,f103,f126,f5114,f98])).
fof(f5306,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl14 | $spl70)),
  inference(unit_resulting_resolution,[],[f70,f69,f126,f5114,f98])).
fof(f5307,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl14 | $spl70)),
  inference(unit_resulting_resolution,[],[f70,f68,f126,f5114,f98])).
fof(f5308,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl58 | $spl70)),
  inference(unit_resulting_resolution,[],[f70,f103,f3168,f5114,f98])).
fof(f5309,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl58 | $spl70)),
  inference(unit_resulting_resolution,[],[f70,f69,f3168,f5114,f98])).
fof(f5310,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl58 | $spl70)),
  inference(unit_resulting_resolution,[],[f70,f68,f3168,f5114,f98])).
fof(f5311,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl36 | $spl70)),
  inference(unit_resulting_resolution,[],[f70,f103,f619,f5114,f98])).
fof(f5312,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl36 | $spl70)),
  inference(unit_resulting_resolution,[],[f70,f69,f619,f5114,f98])).
fof(f5313,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl36 | $spl70)),
  inference(unit_resulting_resolution,[],[f70,f68,f619,f5114,f98])).
fof(f5314,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl46 | $spl70)),
  inference(unit_resulting_resolution,[],[f70,f103,f956,f5114,f98])).
fof(f5315,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl46 | $spl70)),
  inference(unit_resulting_resolution,[],[f70,f69,f956,f5114,f98])).
fof(f5316,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl46 | $spl70)),
  inference(unit_resulting_resolution,[],[f70,f68,f956,f5114,f98])).
fof(f5317,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl70),
  inference(unit_resulting_resolution,[],[f103,f70,f80,f5114,f98])).
fof(f5327,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl12 | $spl70)),
  inference(unit_resulting_resolution,[],[f70,f70,f122,f5114,f100])).
fof(f5328,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl16 | $spl70)),
  inference(unit_resulting_resolution,[],[f70,f70,f127,f5114,f100])).
fof(f5329,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl38 | $spl70)),
  inference(unit_resulting_resolution,[],[f70,f70,f621,f5114,f100])).
fof(f5330,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl52 | $spl70)),
  inference(unit_resulting_resolution,[],[f70,f70,f2153,f5114,f100])).
fof(f5331,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl62 | $spl70)),
  inference(unit_resulting_resolution,[],[f70,f70,f3944,f5114,f100])).
fof(f5332,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl70),
  inference(unit_resulting_resolution,[],[f70,f70,f80,f5114,f100])).
fof(f5339,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl70),
  inference(unit_resulting_resolution,[],[f70,f70,f103,f5114,f101])).
fof(f5340,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl70),
  inference(unit_resulting_resolution,[],[f70,f70,f69,f5114,f101])).
fof(f5341,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl70),
  inference(unit_resulting_resolution,[],[f70,f70,f68,f5114,f101])).
fof(f5342,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl70),
  inference(unit_resulting_resolution,[],[f70,f70,f5114,f5114,f101])).
fof(f5350,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl70),
  inference(unit_resulting_resolution,[],[f70,f70,f5114,f5114,f101])).
fof(f5369,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl70),
  inference(unit_resulting_resolution,[],[f5114,f70,f80,f5114,f102])).
fof(f5385,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl10 | $spl70)),
  inference(unit_resulting_resolution,[],[f70,f103,f121,f5114,f102])).
fof(f5386,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl10 | $spl70)),
  inference(unit_resulting_resolution,[],[f70,f69,f121,f5114,f102])).
fof(f5387,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl10 | $spl70)),
  inference(unit_resulting_resolution,[],[f70,f68,f121,f5114,f102])).
fof(f5388,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl14 | $spl70)),
  inference(unit_resulting_resolution,[],[f70,f103,f126,f5114,f102])).
fof(f5389,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl14 | $spl70)),
  inference(unit_resulting_resolution,[],[f70,f69,f126,f5114,f102])).
fof(f5390,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl14 | $spl70)),
  inference(unit_resulting_resolution,[],[f70,f68,f126,f5114,f102])).
fof(f5391,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl58 | $spl70)),
  inference(unit_resulting_resolution,[],[f70,f103,f3168,f5114,f102])).
fof(f5392,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl58 | $spl70)),
  inference(unit_resulting_resolution,[],[f70,f69,f3168,f5114,f102])).
fof(f5393,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl58 | $spl70)),
  inference(unit_resulting_resolution,[],[f70,f68,f3168,f5114,f102])).
fof(f5394,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl36 | $spl70)),
  inference(unit_resulting_resolution,[],[f70,f103,f619,f5114,f102])).
fof(f5395,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl36 | $spl70)),
  inference(unit_resulting_resolution,[],[f70,f69,f619,f5114,f102])).
fof(f5396,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl36 | $spl70)),
  inference(unit_resulting_resolution,[],[f70,f68,f619,f5114,f102])).
fof(f5397,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl46 | $spl70)),
  inference(unit_resulting_resolution,[],[f70,f103,f956,f5114,f102])).
fof(f5398,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl46 | $spl70)),
  inference(unit_resulting_resolution,[],[f70,f69,f956,f5114,f102])).
fof(f5399,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl46 | $spl70)),
  inference(unit_resulting_resolution,[],[f70,f68,f956,f5114,f102])).
fof(f5400,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl70),
  inference(unit_resulting_resolution,[],[f103,f70,f80,f5114,f102])).
fof(f5410,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl70),
  inference(unit_resulting_resolution,[],[f70,f70,f5114,f104])).
fof(f5114,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2)))) | $spl70),
  inference(cnf_transformation,[],[f5114_D])).
fof(f5114_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2)))) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f5043,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2)))) | $spl52),
  inference(superposition,[],[f87,f2153])).
fof(f4982,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK2)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,X0))),s(fun(sK1,bool),sK3)))) ) | $spl68),
  inference(cnf_transformation,[],[f4982_D])).
fof(f4982_D,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK2)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,X0))),s(fun(sK1,bool),sK3)))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f4977,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK2)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,X0))),s(fun(sK1,bool),sK3)))) ) | $spl54),
  inference(superposition,[],[f2441,f71])).
fof(f4973,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK2)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,X0))),s(fun(sK1,bool),sK3)))) ) | $spl66),
  inference(cnf_transformation,[],[f4973_D])).
fof(f4973_D,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK2)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,X0))),s(fun(sK1,bool),sK3)))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f4968,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK2)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,X0))),s(fun(sK1,bool),sK3)))) ) | $spl50),
  inference(superposition,[],[f1126,f71])).
fof(f2441,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,X0))),s(fun(sK1,bool),sK3)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK2))))) ) | $spl54),
  inference(cnf_transformation,[],[f2441_D])).
fof(f2441_D,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,X0))),s(fun(sK1,bool),sK3)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK2))))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f4868,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK2)))) | p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,X0))),s(fun(sK1,bool),sK3))))) ) | $spl46),
  inference(superposition,[],[f83,f956])).
fof(f3944,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl62),
  inference(cnf_transformation,[],[f3944_D])).
fof(f3944_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f4429,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl65),
  inference(unit_resulting_resolution,[],[f4020,f71])).
fof(f4020,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl65),
  inference(cnf_transformation,[],[f4020_D])).
fof(f4020_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) <=> ~$spl65),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl65])])).
fof(f4137,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl12 | $spl41)),
  inference(unit_resulting_resolution,[],[f103,f70,f122,f720,f96])).
fof(f720,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2)))) | $spl41),
  inference(cnf_transformation,[],[f720_D])).
fof(f720_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2)))) <=> ~$spl41),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl41])])).
fof(f4138,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl12 | $spl41)),
  inference(unit_resulting_resolution,[],[f69,f70,f122,f720,f96])).
fof(f4139,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl12 | $spl41)),
  inference(unit_resulting_resolution,[],[f68,f70,f122,f720,f96])).
fof(f4140,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl16 | $spl41)),
  inference(unit_resulting_resolution,[],[f103,f70,f127,f720,f96])).
fof(f4141,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl16 | $spl41)),
  inference(unit_resulting_resolution,[],[f69,f70,f127,f720,f96])).
fof(f4142,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl16 | $spl41)),
  inference(unit_resulting_resolution,[],[f68,f70,f127,f720,f96])).
fof(f4143,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl38 | $spl41)),
  inference(unit_resulting_resolution,[],[f103,f70,f621,f720,f96])).
fof(f4144,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl38 | $spl41)),
  inference(unit_resulting_resolution,[],[f69,f70,f621,f720,f96])).
fof(f4145,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl38 | $spl41)),
  inference(unit_resulting_resolution,[],[f68,f70,f621,f720,f96])).
fof(f4146,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl41),
  inference(unit_resulting_resolution,[],[f103,f70,f80,f720,f96])).
fof(f4147,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl41),
  inference(unit_resulting_resolution,[],[f69,f70,f80,f720,f96])).
fof(f4148,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl41),
  inference(unit_resulting_resolution,[],[f68,f70,f80,f720,f96])).
fof(f4149,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl41),
  inference(unit_resulting_resolution,[],[f103,f720,f80,f720,f96])).
fof(f4150,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl41),
  inference(unit_resulting_resolution,[],[f69,f720,f80,f720,f96])).
fof(f4151,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl41),
  inference(unit_resulting_resolution,[],[f68,f720,f80,f720,f96])).
fof(f4158,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl41),
  inference(unit_resulting_resolution,[],[f720,f103,f80,f720,f96])).
fof(f4160,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl41),
  inference(unit_resulting_resolution,[],[f720,f69,f80,f720,f96])).
fof(f4162,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl41),
  inference(unit_resulting_resolution,[],[f720,f68,f80,f720,f96])).
fof(f4173,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl10 | $spl41)),
  inference(unit_resulting_resolution,[],[f103,f103,f121,f720,f98])).
fof(f4174,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl10 | $spl41)),
  inference(unit_resulting_resolution,[],[f69,f103,f121,f720,f98])).
fof(f4175,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl10 | $spl41)),
  inference(unit_resulting_resolution,[],[f68,f103,f121,f720,f98])).
fof(f4176,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl10 | $spl41)),
  inference(unit_resulting_resolution,[],[f103,f69,f121,f720,f98])).
fof(f4177,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl10 | $spl41)),
  inference(unit_resulting_resolution,[],[f69,f69,f121,f720,f98])).
fof(f4178,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl10 | $spl41)),
  inference(unit_resulting_resolution,[],[f68,f69,f121,f720,f98])).
fof(f4179,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl10 | $spl41)),
  inference(unit_resulting_resolution,[],[f103,f68,f121,f720,f98])).
fof(f4180,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl10 | $spl41)),
  inference(unit_resulting_resolution,[],[f69,f68,f121,f720,f98])).
fof(f4181,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl10 | $spl41)),
  inference(unit_resulting_resolution,[],[f68,f68,f121,f720,f98])).
fof(f4182,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl14 | $spl41)),
  inference(unit_resulting_resolution,[],[f103,f103,f126,f720,f98])).
fof(f4183,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl14 | $spl41)),
  inference(unit_resulting_resolution,[],[f69,f103,f126,f720,f98])).
fof(f4184,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl14 | $spl41)),
  inference(unit_resulting_resolution,[],[f68,f103,f126,f720,f98])).
fof(f4185,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl14 | $spl41)),
  inference(unit_resulting_resolution,[],[f103,f69,f126,f720,f98])).
fof(f4186,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl14 | $spl41)),
  inference(unit_resulting_resolution,[],[f69,f69,f126,f720,f98])).
fof(f4187,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl14 | $spl41)),
  inference(unit_resulting_resolution,[],[f68,f69,f126,f720,f98])).
fof(f4188,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl14 | $spl41)),
  inference(unit_resulting_resolution,[],[f103,f68,f126,f720,f98])).
fof(f4189,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl14 | $spl41)),
  inference(unit_resulting_resolution,[],[f69,f68,f126,f720,f98])).
fof(f4190,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl14 | $spl41)),
  inference(unit_resulting_resolution,[],[f68,f68,f126,f720,f98])).
fof(f4191,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl41),
  inference(unit_resulting_resolution,[],[f103,f103,f80,f720,f98])).
fof(f4192,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl41),
  inference(unit_resulting_resolution,[],[f69,f103,f80,f720,f98])).
fof(f4193,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl41),
  inference(unit_resulting_resolution,[],[f68,f103,f80,f720,f98])).
fof(f4194,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl41),
  inference(unit_resulting_resolution,[],[f103,f69,f80,f720,f98])).
fof(f4195,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl41),
  inference(unit_resulting_resolution,[],[f69,f69,f80,f720,f98])).
fof(f4196,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl41),
  inference(unit_resulting_resolution,[],[f68,f69,f80,f720,f98])).
fof(f4197,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl41),
  inference(unit_resulting_resolution,[],[f103,f68,f80,f720,f98])).
fof(f4198,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl41),
  inference(unit_resulting_resolution,[],[f69,f68,f80,f720,f98])).
fof(f4199,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl41),
  inference(unit_resulting_resolution,[],[f68,f68,f80,f720,f98])).
fof(f4205,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl12 | $spl41)),
  inference(unit_resulting_resolution,[],[f103,f70,f122,f720,f100])).
fof(f4206,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl12 | $spl41)),
  inference(unit_resulting_resolution,[],[f69,f70,f122,f720,f100])).
fof(f4207,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl12 | $spl41)),
  inference(unit_resulting_resolution,[],[f68,f70,f122,f720,f100])).
fof(f4208,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl16 | $spl41)),
  inference(unit_resulting_resolution,[],[f103,f70,f127,f720,f100])).
fof(f4209,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl16 | $spl41)),
  inference(unit_resulting_resolution,[],[f69,f70,f127,f720,f100])).
fof(f4210,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl16 | $spl41)),
  inference(unit_resulting_resolution,[],[f68,f70,f127,f720,f100])).
fof(f4211,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl38 | $spl41)),
  inference(unit_resulting_resolution,[],[f103,f70,f621,f720,f100])).
fof(f4212,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl38 | $spl41)),
  inference(unit_resulting_resolution,[],[f69,f70,f621,f720,f100])).
fof(f4213,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl38 | $spl41)),
  inference(unit_resulting_resolution,[],[f68,f70,f621,f720,f100])).
fof(f4214,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl41),
  inference(unit_resulting_resolution,[],[f103,f70,f80,f720,f100])).
fof(f4215,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl41),
  inference(unit_resulting_resolution,[],[f69,f70,f80,f720,f100])).
fof(f4216,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl41),
  inference(unit_resulting_resolution,[],[f68,f70,f80,f720,f100])).
fof(f4217,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl41),
  inference(unit_resulting_resolution,[],[f103,f720,f80,f720,f100])).
fof(f4218,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl41),
  inference(unit_resulting_resolution,[],[f69,f720,f80,f720,f100])).
fof(f4219,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl41),
  inference(unit_resulting_resolution,[],[f68,f720,f80,f720,f100])).
fof(f4226,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl41),
  inference(unit_resulting_resolution,[],[f720,f103,f80,f720,f100])).
fof(f4228,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl41),
  inference(unit_resulting_resolution,[],[f720,f69,f80,f720,f100])).
fof(f4230,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl41),
  inference(unit_resulting_resolution,[],[f720,f68,f80,f720,f100])).
fof(f4241,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl41),
  inference(unit_resulting_resolution,[],[f70,f103,f103,f720,f101])).
fof(f4242,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl41),
  inference(unit_resulting_resolution,[],[f70,f69,f103,f720,f101])).
fof(f4243,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl41),
  inference(unit_resulting_resolution,[],[f70,f68,f103,f720,f101])).
fof(f4244,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl41),
  inference(unit_resulting_resolution,[],[f70,f103,f69,f720,f101])).
fof(f4245,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl41),
  inference(unit_resulting_resolution,[],[f70,f69,f69,f720,f101])).
fof(f4246,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl41),
  inference(unit_resulting_resolution,[],[f70,f68,f69,f720,f101])).
fof(f4247,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl41),
  inference(unit_resulting_resolution,[],[f70,f103,f68,f720,f101])).
fof(f4248,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl41),
  inference(unit_resulting_resolution,[],[f70,f69,f68,f720,f101])).
fof(f4249,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl41),
  inference(unit_resulting_resolution,[],[f70,f68,f68,f720,f101])).
fof(f4259,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl41),
  inference(unit_resulting_resolution,[],[f103,f720,f103,f720,f101])).
fof(f4260,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl41),
  inference(unit_resulting_resolution,[],[f69,f720,f103,f720,f101])).
fof(f4261,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl41),
  inference(unit_resulting_resolution,[],[f68,f720,f103,f720,f101])).
fof(f4265,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl41),
  inference(unit_resulting_resolution,[],[f103,f720,f69,f720,f101])).
fof(f4266,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl41),
  inference(unit_resulting_resolution,[],[f69,f720,f69,f720,f101])).
fof(f4267,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl41),
  inference(unit_resulting_resolution,[],[f68,f720,f69,f720,f101])).
fof(f4271,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl41),
  inference(unit_resulting_resolution,[],[f103,f720,f68,f720,f101])).
fof(f4272,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl41),
  inference(unit_resulting_resolution,[],[f69,f720,f68,f720,f101])).
fof(f4273,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl41),
  inference(unit_resulting_resolution,[],[f68,f720,f68,f720,f101])).
fof(f4294,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl10 | $spl41)),
  inference(unit_resulting_resolution,[],[f103,f103,f121,f720,f102])).
fof(f4295,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl10 | $spl41)),
  inference(unit_resulting_resolution,[],[f69,f103,f121,f720,f102])).
fof(f4296,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl10 | $spl41)),
  inference(unit_resulting_resolution,[],[f68,f103,f121,f720,f102])).
fof(f4297,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl10 | $spl41)),
  inference(unit_resulting_resolution,[],[f103,f69,f121,f720,f102])).
fof(f4298,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl10 | $spl41)),
  inference(unit_resulting_resolution,[],[f69,f69,f121,f720,f102])).
fof(f4299,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl10 | $spl41)),
  inference(unit_resulting_resolution,[],[f68,f69,f121,f720,f102])).
fof(f4300,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl10 | $spl41)),
  inference(unit_resulting_resolution,[],[f103,f68,f121,f720,f102])).
fof(f4301,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl10 | $spl41)),
  inference(unit_resulting_resolution,[],[f69,f68,f121,f720,f102])).
fof(f4302,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl10 | $spl41)),
  inference(unit_resulting_resolution,[],[f68,f68,f121,f720,f102])).
fof(f4303,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl14 | $spl41)),
  inference(unit_resulting_resolution,[],[f103,f103,f126,f720,f102])).
fof(f4304,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl14 | $spl41)),
  inference(unit_resulting_resolution,[],[f69,f103,f126,f720,f102])).
fof(f4305,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl14 | $spl41)),
  inference(unit_resulting_resolution,[],[f68,f103,f126,f720,f102])).
fof(f4306,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl14 | $spl41)),
  inference(unit_resulting_resolution,[],[f103,f69,f126,f720,f102])).
fof(f4307,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl14 | $spl41)),
  inference(unit_resulting_resolution,[],[f69,f69,f126,f720,f102])).
fof(f4308,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl14 | $spl41)),
  inference(unit_resulting_resolution,[],[f68,f69,f126,f720,f102])).
fof(f4309,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl14 | $spl41)),
  inference(unit_resulting_resolution,[],[f103,f68,f126,f720,f102])).
fof(f4310,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl14 | $spl41)),
  inference(unit_resulting_resolution,[],[f69,f68,f126,f720,f102])).
fof(f4311,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl14 | $spl41)),
  inference(unit_resulting_resolution,[],[f68,f68,f126,f720,f102])).
fof(f4312,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl41),
  inference(unit_resulting_resolution,[],[f103,f103,f80,f720,f102])).
fof(f4313,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl41),
  inference(unit_resulting_resolution,[],[f69,f103,f80,f720,f102])).
fof(f4314,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl41),
  inference(unit_resulting_resolution,[],[f68,f103,f80,f720,f102])).
fof(f4315,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl41),
  inference(unit_resulting_resolution,[],[f103,f69,f80,f720,f102])).
fof(f4316,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl41),
  inference(unit_resulting_resolution,[],[f69,f69,f80,f720,f102])).
fof(f4317,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl41),
  inference(unit_resulting_resolution,[],[f68,f69,f80,f720,f102])).
fof(f4318,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl41),
  inference(unit_resulting_resolution,[],[f103,f68,f80,f720,f102])).
fof(f4319,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl41),
  inference(unit_resulting_resolution,[],[f69,f68,f80,f720,f102])).
fof(f4320,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl41),
  inference(unit_resulting_resolution,[],[f68,f68,f80,f720,f102])).
fof(f4326,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl41),
  inference(unit_resulting_resolution,[],[f70,f103,f720,f104])).
fof(f4327,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl41),
  inference(unit_resulting_resolution,[],[f70,f69,f720,f104])).
fof(f4328,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl41),
  inference(unit_resulting_resolution,[],[f70,f68,f720,f104])).
fof(f4334,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl41),
  inference(unit_resulting_resolution,[],[f103,f720,f720,f104])).
fof(f4335,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl41),
  inference(unit_resulting_resolution,[],[f69,f720,f720,f104])).
fof(f4336,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl41),
  inference(unit_resulting_resolution,[],[f68,f720,f720,f104])).
fof(f135,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))))) | $spl21),
  inference(cnf_transformation,[],[f135_D])).
fof(f135_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))))) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f4024,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))))) | ($spl5 | $spl38)),
  inference(backward_demodulation,[],[f621,f114])).
fof(f4021,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl64),
  inference(cnf_transformation,[],[f4021_D])).
fof(f4021_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f3939,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl40),
  inference(superposition,[],[f721,f71])).
fof(f3945,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl63),
  inference(cnf_transformation,[],[f3945_D])).
fof(f3945_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) <=> ~$spl63),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl63])])).
fof(f3769,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl12 | $spl40)),
  inference(unit_resulting_resolution,[],[f70,f70,f122,f721,f96])).
fof(f3943,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl61),
  inference(cnf_transformation,[],[f3943_D])).
fof(f3943_D,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) <=> ~$spl61),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl61])])).
fof(f3770,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl12 | $spl40 | $spl45)),
  inference(unit_resulting_resolution,[],[f835,f70,f122,f721,f96])).
fof(f835,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2)))) | $spl45),
  inference(cnf_transformation,[],[f835_D])).
fof(f835_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2)))) <=> ~$spl45),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl45])])).
fof(f3771,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl16 | $spl40)),
  inference(unit_resulting_resolution,[],[f70,f70,f127,f721,f96])).
fof(f3772,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl16 | $spl40 | $spl45)),
  inference(unit_resulting_resolution,[],[f835,f70,f127,f721,f96])).
fof(f3773,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl20 | $spl40)),
  inference(unit_resulting_resolution,[],[f70,f70,f134,f721,f96])).
fof(f134,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))))) | $spl20),
  inference(cnf_transformation,[],[f134_D])).
fof(f134_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))))) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f3774,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl20 | $spl40 | $spl45)),
  inference(unit_resulting_resolution,[],[f835,f70,f134,f721,f96])).
fof(f3775,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl40),
  inference(unit_resulting_resolution,[],[f70,f70,f80,f721,f96])).
fof(f3776,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl40 | $spl45)),
  inference(unit_resulting_resolution,[],[f835,f70,f80,f721,f96])).
fof(f3777,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl40 | $spl45)),
  inference(unit_resulting_resolution,[],[f70,f835,f80,f721,f96])).
fof(f3778,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl40 | $spl45)),
  inference(unit_resulting_resolution,[],[f835,f835,f80,f721,f96])).
fof(f3787,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl40),
  inference(unit_resulting_resolution,[],[f721,f70,f80,f721,f98])).
fof(f3791,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl40 | $spl45)),
  inference(unit_resulting_resolution,[],[f721,f835,f80,f721,f98])).
fof(f3804,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl10 | $spl40)),
  inference(unit_resulting_resolution,[],[f70,f103,f121,f721,f98])).
fof(f3805,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl10 | $spl40 | $spl45)),
  inference(unit_resulting_resolution,[],[f835,f103,f121,f721,f98])).
fof(f3806,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl10 | $spl40)),
  inference(unit_resulting_resolution,[],[f70,f69,f121,f721,f98])).
fof(f3807,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl10 | $spl40 | $spl45)),
  inference(unit_resulting_resolution,[],[f835,f69,f121,f721,f98])).
fof(f3808,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl10 | $spl40)),
  inference(unit_resulting_resolution,[],[f70,f68,f121,f721,f98])).
fof(f3809,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl10 | $spl40 | $spl45)),
  inference(unit_resulting_resolution,[],[f835,f68,f121,f721,f98])).
fof(f3810,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl14 | $spl40)),
  inference(unit_resulting_resolution,[],[f70,f103,f126,f721,f98])).
fof(f3811,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl14 | $spl40 | $spl45)),
  inference(unit_resulting_resolution,[],[f835,f103,f126,f721,f98])).
fof(f3812,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl14 | $spl40)),
  inference(unit_resulting_resolution,[],[f70,f69,f126,f721,f98])).
fof(f3813,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl14 | $spl40 | $spl45)),
  inference(unit_resulting_resolution,[],[f835,f69,f126,f721,f98])).
fof(f3814,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl14 | $spl40)),
  inference(unit_resulting_resolution,[],[f70,f68,f126,f721,f98])).
fof(f3815,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl14 | $spl40 | $spl45)),
  inference(unit_resulting_resolution,[],[f835,f68,f126,f721,f98])).
fof(f3816,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl18 | $spl40)),
  inference(unit_resulting_resolution,[],[f70,f103,f133,f721,f98])).
fof(f3817,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl18 | $spl40 | $spl45)),
  inference(unit_resulting_resolution,[],[f835,f103,f133,f721,f98])).
fof(f3818,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl18 | $spl40)),
  inference(unit_resulting_resolution,[],[f70,f69,f133,f721,f98])).
fof(f3819,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl18 | $spl40 | $spl45)),
  inference(unit_resulting_resolution,[],[f835,f69,f133,f721,f98])).
fof(f3820,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl18 | $spl40)),
  inference(unit_resulting_resolution,[],[f70,f68,f133,f721,f98])).
fof(f3821,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl18 | $spl40 | $spl45)),
  inference(unit_resulting_resolution,[],[f835,f68,f133,f721,f98])).
fof(f3822,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl40),
  inference(unit_resulting_resolution,[],[f103,f70,f80,f721,f98])).
fof(f3823,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl40 | $spl45)),
  inference(unit_resulting_resolution,[],[f103,f835,f80,f721,f98])).
fof(f3831,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl12 | $spl40)),
  inference(unit_resulting_resolution,[],[f70,f70,f122,f721,f100])).
fof(f3832,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl12 | $spl40 | $spl45)),
  inference(unit_resulting_resolution,[],[f835,f70,f122,f721,f100])).
fof(f3833,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl16 | $spl40)),
  inference(unit_resulting_resolution,[],[f70,f70,f127,f721,f100])).
fof(f3834,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl16 | $spl40 | $spl45)),
  inference(unit_resulting_resolution,[],[f835,f70,f127,f721,f100])).
fof(f3835,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl20 | $spl40)),
  inference(unit_resulting_resolution,[],[f70,f70,f134,f721,f100])).
fof(f3836,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl20 | $spl40 | $spl45)),
  inference(unit_resulting_resolution,[],[f835,f70,f134,f721,f100])).
fof(f3837,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl40),
  inference(unit_resulting_resolution,[],[f70,f70,f80,f721,f100])).
fof(f3838,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl40 | $spl45)),
  inference(unit_resulting_resolution,[],[f835,f70,f80,f721,f100])).
fof(f3839,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl40 | $spl45)),
  inference(unit_resulting_resolution,[],[f70,f835,f80,f721,f100])).
fof(f3840,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl40 | $spl45)),
  inference(unit_resulting_resolution,[],[f835,f835,f80,f721,f100])).
fof(f3846,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl40),
  inference(unit_resulting_resolution,[],[f70,f70,f103,f721,f101])).
fof(f3847,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl40 | $spl45)),
  inference(unit_resulting_resolution,[],[f70,f835,f103,f721,f101])).
fof(f3848,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl40),
  inference(unit_resulting_resolution,[],[f70,f70,f69,f721,f101])).
fof(f3849,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl40 | $spl45)),
  inference(unit_resulting_resolution,[],[f70,f835,f69,f721,f101])).
fof(f3850,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl40),
  inference(unit_resulting_resolution,[],[f70,f70,f68,f721,f101])).
fof(f3851,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl40 | $spl45)),
  inference(unit_resulting_resolution,[],[f70,f835,f68,f721,f101])).
fof(f3852,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl40),
  inference(unit_resulting_resolution,[],[f70,f70,f721,f721,f101])).
fof(f3853,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl40 | $spl45)),
  inference(unit_resulting_resolution,[],[f70,f835,f721,f721,f101])).
fof(f3868,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl40),
  inference(unit_resulting_resolution,[],[f70,f70,f721,f721,f101])).
fof(f3869,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl40 | $spl45)),
  inference(unit_resulting_resolution,[],[f70,f835,f721,f721,f101])).
fof(f3891,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl40),
  inference(unit_resulting_resolution,[],[f721,f70,f80,f721,f102])).
fof(f3895,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl40 | $spl45)),
  inference(unit_resulting_resolution,[],[f721,f835,f80,f721,f102])).
fof(f3908,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl10 | $spl40)),
  inference(unit_resulting_resolution,[],[f70,f103,f121,f721,f102])).
fof(f3909,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl10 | $spl40 | $spl45)),
  inference(unit_resulting_resolution,[],[f835,f103,f121,f721,f102])).
fof(f3910,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl10 | $spl40)),
  inference(unit_resulting_resolution,[],[f70,f69,f121,f721,f102])).
fof(f3911,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl10 | $spl40 | $spl45)),
  inference(unit_resulting_resolution,[],[f835,f69,f121,f721,f102])).
fof(f3912,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl10 | $spl40)),
  inference(unit_resulting_resolution,[],[f70,f68,f121,f721,f102])).
fof(f3913,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl10 | $spl40 | $spl45)),
  inference(unit_resulting_resolution,[],[f835,f68,f121,f721,f102])).
fof(f3914,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl14 | $spl40)),
  inference(unit_resulting_resolution,[],[f70,f103,f126,f721,f102])).
fof(f3915,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl14 | $spl40 | $spl45)),
  inference(unit_resulting_resolution,[],[f835,f103,f126,f721,f102])).
fof(f3916,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl14 | $spl40)),
  inference(unit_resulting_resolution,[],[f70,f69,f126,f721,f102])).
fof(f3917,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl14 | $spl40 | $spl45)),
  inference(unit_resulting_resolution,[],[f835,f69,f126,f721,f102])).
fof(f3918,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl14 | $spl40)),
  inference(unit_resulting_resolution,[],[f70,f68,f126,f721,f102])).
fof(f3919,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl14 | $spl40 | $spl45)),
  inference(unit_resulting_resolution,[],[f835,f68,f126,f721,f102])).
fof(f3920,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl18 | $spl40)),
  inference(unit_resulting_resolution,[],[f70,f103,f133,f721,f102])).
fof(f3921,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl18 | $spl40 | $spl45)),
  inference(unit_resulting_resolution,[],[f835,f103,f133,f721,f102])).
fof(f3922,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl18 | $spl40)),
  inference(unit_resulting_resolution,[],[f70,f69,f133,f721,f102])).
fof(f3923,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl18 | $spl40 | $spl45)),
  inference(unit_resulting_resolution,[],[f835,f69,f133,f721,f102])).
fof(f3924,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl18 | $spl40)),
  inference(unit_resulting_resolution,[],[f70,f68,f133,f721,f102])).
fof(f3925,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl18 | $spl40 | $spl45)),
  inference(unit_resulting_resolution,[],[f835,f68,f133,f721,f102])).
fof(f3926,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl40),
  inference(unit_resulting_resolution,[],[f103,f70,f80,f721,f102])).
fof(f3927,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl40 | $spl45)),
  inference(unit_resulting_resolution,[],[f103,f835,f80,f721,f102])).
fof(f3935,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl40),
  inference(unit_resulting_resolution,[],[f70,f70,f721,f104])).
fof(f3936,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl40 | $spl45)),
  inference(unit_resulting_resolution,[],[f70,f835,f721,f104])).
fof(f3167,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl59),
  inference(cnf_transformation,[],[f3167_D])).
fof(f3167_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) <=> ~$spl59),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl59])])).
fof(f3175,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl12 | $spl45)),
  inference(unit_resulting_resolution,[],[f103,f70,f122,f835,f96])).
fof(f3176,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl12 | $spl45)),
  inference(unit_resulting_resolution,[],[f69,f70,f122,f835,f96])).
fof(f3177,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl12 | $spl45)),
  inference(unit_resulting_resolution,[],[f68,f70,f122,f835,f96])).
fof(f3178,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl16 | $spl45)),
  inference(unit_resulting_resolution,[],[f103,f70,f127,f835,f96])).
fof(f3179,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl16 | $spl45)),
  inference(unit_resulting_resolution,[],[f69,f70,f127,f835,f96])).
fof(f3180,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl16 | $spl45)),
  inference(unit_resulting_resolution,[],[f68,f70,f127,f835,f96])).
fof(f3181,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl45),
  inference(unit_resulting_resolution,[],[f103,f70,f80,f835,f96])).
fof(f3182,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl45),
  inference(unit_resulting_resolution,[],[f69,f70,f80,f835,f96])).
fof(f3183,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl45),
  inference(unit_resulting_resolution,[],[f68,f70,f80,f835,f96])).
fof(f3184,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl45),
  inference(unit_resulting_resolution,[],[f103,f835,f80,f835,f96])).
fof(f3185,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl45),
  inference(unit_resulting_resolution,[],[f69,f835,f80,f835,f96])).
fof(f3186,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl45),
  inference(unit_resulting_resolution,[],[f68,f835,f80,f835,f96])).
fof(f3192,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl45),
  inference(unit_resulting_resolution,[],[f835,f103,f80,f835,f96])).
fof(f3194,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl45),
  inference(unit_resulting_resolution,[],[f835,f69,f80,f835,f96])).
fof(f3196,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl45),
  inference(unit_resulting_resolution,[],[f835,f68,f80,f835,f96])).
fof(f3206,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl10 | $spl45)),
  inference(unit_resulting_resolution,[],[f103,f103,f121,f835,f98])).
fof(f3207,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl10 | $spl45)),
  inference(unit_resulting_resolution,[],[f69,f103,f121,f835,f98])).
fof(f3208,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl10 | $spl45)),
  inference(unit_resulting_resolution,[],[f68,f103,f121,f835,f98])).
fof(f3209,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl10 | $spl45)),
  inference(unit_resulting_resolution,[],[f103,f69,f121,f835,f98])).
fof(f3210,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl10 | $spl45)),
  inference(unit_resulting_resolution,[],[f69,f69,f121,f835,f98])).
fof(f3211,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl10 | $spl45)),
  inference(unit_resulting_resolution,[],[f68,f69,f121,f835,f98])).
fof(f3212,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl10 | $spl45)),
  inference(unit_resulting_resolution,[],[f103,f68,f121,f835,f98])).
fof(f3213,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl10 | $spl45)),
  inference(unit_resulting_resolution,[],[f69,f68,f121,f835,f98])).
fof(f3214,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl10 | $spl45)),
  inference(unit_resulting_resolution,[],[f68,f68,f121,f835,f98])).
fof(f3215,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl14 | $spl45)),
  inference(unit_resulting_resolution,[],[f103,f103,f126,f835,f98])).
fof(f3216,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl14 | $spl45)),
  inference(unit_resulting_resolution,[],[f69,f103,f126,f835,f98])).
fof(f3217,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl14 | $spl45)),
  inference(unit_resulting_resolution,[],[f68,f103,f126,f835,f98])).
fof(f3218,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl14 | $spl45)),
  inference(unit_resulting_resolution,[],[f103,f69,f126,f835,f98])).
fof(f3219,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl14 | $spl45)),
  inference(unit_resulting_resolution,[],[f69,f69,f126,f835,f98])).
fof(f3220,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl14 | $spl45)),
  inference(unit_resulting_resolution,[],[f68,f69,f126,f835,f98])).
fof(f3221,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl14 | $spl45)),
  inference(unit_resulting_resolution,[],[f103,f68,f126,f835,f98])).
fof(f3222,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl14 | $spl45)),
  inference(unit_resulting_resolution,[],[f69,f68,f126,f835,f98])).
fof(f3223,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl14 | $spl45)),
  inference(unit_resulting_resolution,[],[f68,f68,f126,f835,f98])).
fof(f3224,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl45),
  inference(unit_resulting_resolution,[],[f103,f103,f80,f835,f98])).
fof(f3225,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl45),
  inference(unit_resulting_resolution,[],[f69,f103,f80,f835,f98])).
fof(f3226,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl45),
  inference(unit_resulting_resolution,[],[f68,f103,f80,f835,f98])).
fof(f3227,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl45),
  inference(unit_resulting_resolution,[],[f103,f69,f80,f835,f98])).
fof(f3228,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl45),
  inference(unit_resulting_resolution,[],[f69,f69,f80,f835,f98])).
fof(f3229,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl45),
  inference(unit_resulting_resolution,[],[f68,f69,f80,f835,f98])).
fof(f3230,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl45),
  inference(unit_resulting_resolution,[],[f103,f68,f80,f835,f98])).
fof(f3231,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl45),
  inference(unit_resulting_resolution,[],[f69,f68,f80,f835,f98])).
fof(f3232,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl45),
  inference(unit_resulting_resolution,[],[f68,f68,f80,f835,f98])).
fof(f3238,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl12 | $spl45)),
  inference(unit_resulting_resolution,[],[f103,f70,f122,f835,f100])).
fof(f3239,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl12 | $spl45)),
  inference(unit_resulting_resolution,[],[f69,f70,f122,f835,f100])).
fof(f3240,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl12 | $spl45)),
  inference(unit_resulting_resolution,[],[f68,f70,f122,f835,f100])).
fof(f3241,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl16 | $spl45)),
  inference(unit_resulting_resolution,[],[f103,f70,f127,f835,f100])).
fof(f3242,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl16 | $spl45)),
  inference(unit_resulting_resolution,[],[f69,f70,f127,f835,f100])).
fof(f3243,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl16 | $spl45)),
  inference(unit_resulting_resolution,[],[f68,f70,f127,f835,f100])).
fof(f3244,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl45),
  inference(unit_resulting_resolution,[],[f103,f70,f80,f835,f100])).
fof(f3245,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl45),
  inference(unit_resulting_resolution,[],[f69,f70,f80,f835,f100])).
fof(f3246,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl45),
  inference(unit_resulting_resolution,[],[f68,f70,f80,f835,f100])).
fof(f3247,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl45),
  inference(unit_resulting_resolution,[],[f103,f835,f80,f835,f100])).
fof(f3248,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl45),
  inference(unit_resulting_resolution,[],[f69,f835,f80,f835,f100])).
fof(f3249,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl45),
  inference(unit_resulting_resolution,[],[f68,f835,f80,f835,f100])).
fof(f3255,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl45),
  inference(unit_resulting_resolution,[],[f835,f103,f80,f835,f100])).
fof(f3257,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl45),
  inference(unit_resulting_resolution,[],[f835,f69,f80,f835,f100])).
fof(f3259,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl45),
  inference(unit_resulting_resolution,[],[f835,f68,f80,f835,f100])).
fof(f3269,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl45),
  inference(unit_resulting_resolution,[],[f70,f103,f103,f835,f101])).
fof(f3270,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl45),
  inference(unit_resulting_resolution,[],[f70,f69,f103,f835,f101])).
fof(f3271,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl45),
  inference(unit_resulting_resolution,[],[f70,f68,f103,f835,f101])).
fof(f3272,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl45),
  inference(unit_resulting_resolution,[],[f70,f103,f69,f835,f101])).
fof(f3273,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl45),
  inference(unit_resulting_resolution,[],[f70,f69,f69,f835,f101])).
fof(f3274,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl45),
  inference(unit_resulting_resolution,[],[f70,f68,f69,f835,f101])).
fof(f3275,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl45),
  inference(unit_resulting_resolution,[],[f70,f103,f68,f835,f101])).
fof(f3276,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl45),
  inference(unit_resulting_resolution,[],[f70,f69,f68,f835,f101])).
fof(f3277,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl45),
  inference(unit_resulting_resolution,[],[f70,f68,f68,f835,f101])).
fof(f3287,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl45),
  inference(unit_resulting_resolution,[],[f103,f835,f103,f835,f101])).
fof(f3288,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl45),
  inference(unit_resulting_resolution,[],[f69,f835,f103,f835,f101])).
fof(f3289,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl45),
  inference(unit_resulting_resolution,[],[f68,f835,f103,f835,f101])).
fof(f3293,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl45),
  inference(unit_resulting_resolution,[],[f103,f835,f69,f835,f101])).
fof(f3294,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl45),
  inference(unit_resulting_resolution,[],[f69,f835,f69,f835,f101])).
fof(f3295,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl45),
  inference(unit_resulting_resolution,[],[f68,f835,f69,f835,f101])).
fof(f3299,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl45),
  inference(unit_resulting_resolution,[],[f103,f835,f68,f835,f101])).
fof(f3300,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl45),
  inference(unit_resulting_resolution,[],[f69,f835,f68,f835,f101])).
fof(f3301,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl45),
  inference(unit_resulting_resolution,[],[f68,f835,f68,f835,f101])).
fof(f3321,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl10 | $spl45)),
  inference(unit_resulting_resolution,[],[f103,f103,f121,f835,f102])).
fof(f3322,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl10 | $spl45)),
  inference(unit_resulting_resolution,[],[f69,f103,f121,f835,f102])).
fof(f3323,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl10 | $spl45)),
  inference(unit_resulting_resolution,[],[f68,f103,f121,f835,f102])).
fof(f3324,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl10 | $spl45)),
  inference(unit_resulting_resolution,[],[f103,f69,f121,f835,f102])).
fof(f3325,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl10 | $spl45)),
  inference(unit_resulting_resolution,[],[f69,f69,f121,f835,f102])).
fof(f3326,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl10 | $spl45)),
  inference(unit_resulting_resolution,[],[f68,f69,f121,f835,f102])).
fof(f3327,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl10 | $spl45)),
  inference(unit_resulting_resolution,[],[f103,f68,f121,f835,f102])).
fof(f3328,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl10 | $spl45)),
  inference(unit_resulting_resolution,[],[f69,f68,f121,f835,f102])).
fof(f3329,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl10 | $spl45)),
  inference(unit_resulting_resolution,[],[f68,f68,f121,f835,f102])).
fof(f3330,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl14 | $spl45)),
  inference(unit_resulting_resolution,[],[f103,f103,f126,f835,f102])).
fof(f3331,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl14 | $spl45)),
  inference(unit_resulting_resolution,[],[f69,f103,f126,f835,f102])).
fof(f3332,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl14 | $spl45)),
  inference(unit_resulting_resolution,[],[f68,f103,f126,f835,f102])).
fof(f3333,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl14 | $spl45)),
  inference(unit_resulting_resolution,[],[f103,f69,f126,f835,f102])).
fof(f3334,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl14 | $spl45)),
  inference(unit_resulting_resolution,[],[f69,f69,f126,f835,f102])).
fof(f3335,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl14 | $spl45)),
  inference(unit_resulting_resolution,[],[f68,f69,f126,f835,f102])).
fof(f3336,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl14 | $spl45)),
  inference(unit_resulting_resolution,[],[f103,f68,f126,f835,f102])).
fof(f3337,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl14 | $spl45)),
  inference(unit_resulting_resolution,[],[f69,f68,f126,f835,f102])).
fof(f3338,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl14 | $spl45)),
  inference(unit_resulting_resolution,[],[f68,f68,f126,f835,f102])).
fof(f3339,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl45),
  inference(unit_resulting_resolution,[],[f103,f103,f80,f835,f102])).
fof(f3340,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl45),
  inference(unit_resulting_resolution,[],[f69,f103,f80,f835,f102])).
fof(f3341,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl45),
  inference(unit_resulting_resolution,[],[f68,f103,f80,f835,f102])).
fof(f3342,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl45),
  inference(unit_resulting_resolution,[],[f103,f69,f80,f835,f102])).
fof(f3343,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl45),
  inference(unit_resulting_resolution,[],[f69,f69,f80,f835,f102])).
fof(f3344,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl45),
  inference(unit_resulting_resolution,[],[f68,f69,f80,f835,f102])).
fof(f3345,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl45),
  inference(unit_resulting_resolution,[],[f103,f68,f80,f835,f102])).
fof(f3346,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl45),
  inference(unit_resulting_resolution,[],[f69,f68,f80,f835,f102])).
fof(f3347,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl45),
  inference(unit_resulting_resolution,[],[f68,f68,f80,f835,f102])).
fof(f3353,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl45),
  inference(unit_resulting_resolution,[],[f70,f103,f835,f104])).
fof(f3354,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl45),
  inference(unit_resulting_resolution,[],[f70,f69,f835,f104])).
fof(f3355,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl45),
  inference(unit_resulting_resolution,[],[f70,f68,f835,f104])).
fof(f3361,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl45),
  inference(unit_resulting_resolution,[],[f103,f835,f835,f104])).
fof(f3362,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl45),
  inference(unit_resulting_resolution,[],[f69,f835,f835,f104])).
fof(f3363,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl45),
  inference(unit_resulting_resolution,[],[f68,f835,f835,f104])).
fof(f3168,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl58),
  inference(cnf_transformation,[],[f3168_D])).
fof(f3168_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f3109,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl44),
  inference(superposition,[],[f836,f71])).
fof(f3112,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl57),
  inference(cnf_transformation,[],[f3112_D])).
fof(f3112_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) <=> ~$spl57),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl57])])).
fof(f3015,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl12 | $spl44)),
  inference(unit_resulting_resolution,[],[f70,f70,f122,f836,f96])).
fof(f3016,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl16 | $spl44)),
  inference(unit_resulting_resolution,[],[f70,f70,f127,f836,f96])).
fof(f3017,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl38 | $spl44)),
  inference(unit_resulting_resolution,[],[f70,f70,f621,f836,f96])).
fof(f3018,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl44),
  inference(unit_resulting_resolution,[],[f70,f70,f80,f836,f96])).
fof(f3026,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl44),
  inference(unit_resulting_resolution,[],[f836,f70,f80,f836,f98])).
fof(f3037,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl10 | $spl44)),
  inference(unit_resulting_resolution,[],[f70,f103,f121,f836,f98])).
fof(f3038,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl10 | $spl44)),
  inference(unit_resulting_resolution,[],[f70,f69,f121,f836,f98])).
fof(f3039,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl10 | $spl44)),
  inference(unit_resulting_resolution,[],[f70,f68,f121,f836,f98])).
fof(f3040,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl14 | $spl44)),
  inference(unit_resulting_resolution,[],[f70,f103,f126,f836,f98])).
fof(f3041,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl14 | $spl44)),
  inference(unit_resulting_resolution,[],[f70,f69,f126,f836,f98])).
fof(f3042,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl14 | $spl44)),
  inference(unit_resulting_resolution,[],[f70,f68,f126,f836,f98])).
fof(f3043,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl44),
  inference(unit_resulting_resolution,[],[f103,f70,f80,f836,f98])).
fof(f3050,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl12 | $spl44)),
  inference(unit_resulting_resolution,[],[f70,f70,f122,f836,f100])).
fof(f3051,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl16 | $spl44)),
  inference(unit_resulting_resolution,[],[f70,f70,f127,f836,f100])).
fof(f3052,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl38 | $spl44)),
  inference(unit_resulting_resolution,[],[f70,f70,f621,f836,f100])).
fof(f3053,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl44),
  inference(unit_resulting_resolution,[],[f70,f70,f80,f836,f100])).
fof(f3058,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl44),
  inference(unit_resulting_resolution,[],[f70,f70,f103,f836,f101])).
fof(f3059,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl44),
  inference(unit_resulting_resolution,[],[f70,f70,f69,f836,f101])).
fof(f3060,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl44),
  inference(unit_resulting_resolution,[],[f70,f70,f68,f836,f101])).
fof(f3061,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl44),
  inference(unit_resulting_resolution,[],[f70,f70,f836,f836,f101])).
fof(f3069,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl44),
  inference(unit_resulting_resolution,[],[f70,f70,f836,f836,f101])).
fof(f3083,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl44),
  inference(unit_resulting_resolution,[],[f836,f70,f80,f836,f102])).
fof(f3094,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl10 | $spl44)),
  inference(unit_resulting_resolution,[],[f70,f103,f121,f836,f102])).
fof(f3095,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl10 | $spl44)),
  inference(unit_resulting_resolution,[],[f70,f69,f121,f836,f102])).
fof(f3096,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl10 | $spl44)),
  inference(unit_resulting_resolution,[],[f70,f68,f121,f836,f102])).
fof(f3097,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl14 | $spl44)),
  inference(unit_resulting_resolution,[],[f70,f103,f126,f836,f102])).
fof(f3098,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl14 | $spl44)),
  inference(unit_resulting_resolution,[],[f70,f69,f126,f836,f102])).
fof(f3099,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | ($spl14 | $spl44)),
  inference(unit_resulting_resolution,[],[f70,f68,f126,f836,f102])).
fof(f3100,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl44),
  inference(unit_resulting_resolution,[],[f103,f70,f80,f836,f102])).
fof(f3107,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2))) | $spl44),
  inference(unit_resulting_resolution,[],[f70,f70,f836,f104])).
fof(f836,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2)))) | $spl44),
  inference(cnf_transformation,[],[f836_D])).
fof(f836_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2)))) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f2084,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2)))) | $spl38),
  inference(superposition,[],[f87,f621])).
fof(f2791,plain,(
  $false | $spl42),
  inference(trivial_inequality_removal,[],[f2742])).
fof(f2742,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl42),
  inference(unit_resulting_resolution,[],[f70,f70,f723,f104])).
fof(f2790,plain,(
  $false | $spl42),
  inference(trivial_inequality_removal,[],[f2714])).
fof(f2714,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl42),
  inference(unit_resulting_resolution,[],[f103,f70,f80,f723,f102])).
fof(f2789,plain,(
  $false | ($spl38 | $spl42)),
  inference(trivial_inequality_removal,[],[f2713])).
fof(f2713,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl38 | $spl42)),
  inference(unit_resulting_resolution,[],[f70,f723,f621,f723,f102])).
fof(f2788,plain,(
  $false | ($spl16 | $spl42)),
  inference(trivial_inequality_removal,[],[f2712])).
fof(f2712,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl16 | $spl42)),
  inference(unit_resulting_resolution,[],[f70,f723,f127,f723,f102])).
fof(f2787,plain,(
  $false | ($spl12 | $spl42)),
  inference(trivial_inequality_removal,[],[f2711])).
fof(f2711,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl12 | $spl42)),
  inference(unit_resulting_resolution,[],[f70,f723,f122,f723,f102])).
fof(f2786,plain,(
  $false | ($spl14 | $spl42)),
  inference(trivial_inequality_removal,[],[f2710])).
fof(f2710,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl14 | $spl42)),
  inference(unit_resulting_resolution,[],[f70,f68,f126,f723,f102])).
fof(f2785,plain,(
  $false | ($spl14 | $spl42)),
  inference(trivial_inequality_removal,[],[f2709])).
fof(f2709,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl14 | $spl42)),
  inference(unit_resulting_resolution,[],[f70,f69,f126,f723,f102])).
fof(f2784,plain,(
  $false | ($spl14 | $spl42)),
  inference(trivial_inequality_removal,[],[f2708])).
fof(f2708,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl14 | $spl42)),
  inference(unit_resulting_resolution,[],[f70,f103,f126,f723,f102])).
fof(f2783,plain,(
  $false | ($spl10 | $spl42)),
  inference(trivial_inequality_removal,[],[f2707])).
fof(f2707,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl10 | $spl42)),
  inference(unit_resulting_resolution,[],[f70,f68,f121,f723,f102])).
fof(f2782,plain,(
  $false | ($spl10 | $spl42)),
  inference(trivial_inequality_removal,[],[f2706])).
fof(f2706,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl10 | $spl42)),
  inference(unit_resulting_resolution,[],[f70,f69,f121,f723,f102])).
fof(f2781,plain,(
  $false | ($spl10 | $spl42)),
  inference(trivial_inequality_removal,[],[f2705])).
fof(f2705,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl10 | $spl42)),
  inference(unit_resulting_resolution,[],[f70,f103,f121,f723,f102])).
fof(f2780,plain,(
  $false | $spl42),
  inference(trivial_inequality_removal,[],[f2691])).
fof(f2691,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl42),
  inference(unit_resulting_resolution,[],[f723,f70,f80,f723,f102])).
fof(f2779,plain,(
  $false | ($spl38 | $spl42)),
  inference(trivial_inequality_removal,[],[f2674])).
fof(f2674,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl38 | $spl42)),
  inference(unit_resulting_resolution,[],[f723,f70,f621,f723,f102])).
fof(f2778,plain,(
  $false | ($spl16 | $spl42)),
  inference(trivial_inequality_removal,[],[f2657])).
fof(f2657,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl16 | $spl42)),
  inference(unit_resulting_resolution,[],[f723,f70,f127,f723,f102])).
fof(f2777,plain,(
  $false | ($spl12 | $spl42)),
  inference(trivial_inequality_removal,[],[f2640])).
fof(f2640,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl12 | $spl42)),
  inference(unit_resulting_resolution,[],[f723,f70,f122,f723,f102])).
fof(f2776,plain,(
  $false | $spl42),
  inference(trivial_inequality_removal,[],[f2623])).
fof(f2623,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl42),
  inference(unit_resulting_resolution,[],[f70,f70,f723,f723,f101])).
fof(f2775,plain,(
  $false | $spl42),
  inference(trivial_inequality_removal,[],[f2603])).
fof(f2603,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl42),
  inference(unit_resulting_resolution,[],[f70,f70,f723,f723,f101])).
fof(f2774,plain,(
  $false | $spl42),
  inference(trivial_inequality_removal,[],[f2602])).
fof(f2602,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl42),
  inference(unit_resulting_resolution,[],[f70,f70,f68,f723,f101])).
fof(f2773,plain,(
  $false | $spl42),
  inference(trivial_inequality_removal,[],[f2601])).
fof(f2601,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl42),
  inference(unit_resulting_resolution,[],[f70,f70,f69,f723,f101])).
fof(f2772,plain,(
  $false | $spl42),
  inference(trivial_inequality_removal,[],[f2600])).
fof(f2600,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl42),
  inference(unit_resulting_resolution,[],[f70,f70,f103,f723,f101])).
fof(f2771,plain,(
  $false | $spl42),
  inference(trivial_inequality_removal,[],[f2583])).
fof(f2583,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl42),
  inference(unit_resulting_resolution,[],[f70,f70,f80,f723,f100])).
fof(f2770,plain,(
  $false | ($spl38 | $spl42)),
  inference(trivial_inequality_removal,[],[f2582])).
fof(f2582,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl38 | $spl42)),
  inference(unit_resulting_resolution,[],[f70,f70,f621,f723,f100])).
fof(f2769,plain,(
  $false | ($spl16 | $spl42)),
  inference(trivial_inequality_removal,[],[f2581])).
fof(f2581,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl16 | $spl42)),
  inference(unit_resulting_resolution,[],[f70,f70,f127,f723,f100])).
fof(f2768,plain,(
  $false | ($spl12 | $spl42)),
  inference(trivial_inequality_removal,[],[f2580])).
fof(f2580,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl12 | $spl42)),
  inference(unit_resulting_resolution,[],[f70,f70,f122,f723,f100])).
fof(f2767,plain,(
  $false | $spl42),
  inference(trivial_inequality_removal,[],[f2552])).
fof(f2552,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl42),
  inference(unit_resulting_resolution,[],[f103,f70,f80,f723,f98])).
fof(f2766,plain,(
  $false | ($spl38 | $spl42)),
  inference(trivial_inequality_removal,[],[f2551])).
fof(f2551,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl38 | $spl42)),
  inference(unit_resulting_resolution,[],[f70,f723,f621,f723,f98])).
fof(f2765,plain,(
  $false | ($spl16 | $spl42)),
  inference(trivial_inequality_removal,[],[f2550])).
fof(f2550,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl16 | $spl42)),
  inference(unit_resulting_resolution,[],[f70,f723,f127,f723,f98])).
fof(f2764,plain,(
  $false | ($spl12 | $spl42)),
  inference(trivial_inequality_removal,[],[f2549])).
fof(f2549,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl12 | $spl42)),
  inference(unit_resulting_resolution,[],[f70,f723,f122,f723,f98])).
fof(f2763,plain,(
  $false | ($spl14 | $spl42)),
  inference(trivial_inequality_removal,[],[f2548])).
fof(f2548,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl14 | $spl42)),
  inference(unit_resulting_resolution,[],[f70,f68,f126,f723,f98])).
fof(f2762,plain,(
  $false | ($spl14 | $spl42)),
  inference(trivial_inequality_removal,[],[f2547])).
fof(f2547,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl14 | $spl42)),
  inference(unit_resulting_resolution,[],[f70,f69,f126,f723,f98])).
fof(f2761,plain,(
  $false | ($spl14 | $spl42)),
  inference(trivial_inequality_removal,[],[f2546])).
fof(f2546,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl14 | $spl42)),
  inference(unit_resulting_resolution,[],[f70,f103,f126,f723,f98])).
fof(f2760,plain,(
  $false | ($spl10 | $spl42)),
  inference(trivial_inequality_removal,[],[f2545])).
fof(f2545,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl10 | $spl42)),
  inference(unit_resulting_resolution,[],[f70,f68,f121,f723,f98])).
fof(f2759,plain,(
  $false | ($spl10 | $spl42)),
  inference(trivial_inequality_removal,[],[f2544])).
fof(f2544,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl10 | $spl42)),
  inference(unit_resulting_resolution,[],[f70,f69,f121,f723,f98])).
fof(f2758,plain,(
  $false | ($spl10 | $spl42)),
  inference(trivial_inequality_removal,[],[f2543])).
fof(f2543,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl10 | $spl42)),
  inference(unit_resulting_resolution,[],[f70,f103,f121,f723,f98])).
fof(f2757,plain,(
  $false | $spl42),
  inference(trivial_inequality_removal,[],[f2529])).
fof(f2529,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl42),
  inference(unit_resulting_resolution,[],[f723,f70,f80,f723,f98])).
fof(f2756,plain,(
  $false | ($spl38 | $spl42)),
  inference(trivial_inequality_removal,[],[f2512])).
fof(f2512,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl38 | $spl42)),
  inference(unit_resulting_resolution,[],[f723,f70,f621,f723,f98])).
fof(f2755,plain,(
  $false | ($spl16 | $spl42)),
  inference(trivial_inequality_removal,[],[f2495])).
fof(f2495,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl16 | $spl42)),
  inference(unit_resulting_resolution,[],[f723,f70,f127,f723,f98])).
fof(f2754,plain,(
  $false | ($spl12 | $spl42)),
  inference(trivial_inequality_removal,[],[f2478])).
fof(f2478,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl12 | $spl42)),
  inference(unit_resulting_resolution,[],[f723,f70,f122,f723,f98])).
fof(f2753,plain,(
  $false | $spl42),
  inference(trivial_inequality_removal,[],[f2458])).
fof(f2458,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl42),
  inference(unit_resulting_resolution,[],[f70,f70,f80,f723,f96])).
fof(f2752,plain,(
  $false | ($spl38 | $spl42)),
  inference(trivial_inequality_removal,[],[f2457])).
fof(f2457,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl38 | $spl42)),
  inference(unit_resulting_resolution,[],[f70,f70,f621,f723,f96])).
fof(f2751,plain,(
  $false | ($spl16 | $spl42)),
  inference(trivial_inequality_removal,[],[f2456])).
fof(f2456,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl16 | $spl42)),
  inference(unit_resulting_resolution,[],[f70,f70,f127,f723,f96])).
fof(f2750,plain,(
  $false | ($spl12 | $spl42)),
  inference(trivial_inequality_removal,[],[f2455])).
fof(f2455,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl12 | $spl42)),
  inference(unit_resulting_resolution,[],[f70,f70,f122,f723,f96])).
fof(f2454,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f70,f723])).
fof(f2747,plain,(
  $false | $spl42),
  inference(resolution,[],[f723,f70])).
fof(f142,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))))) | $spl25),
  inference(cnf_transformation,[],[f142_D])).
fof(f142_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))))) <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f2450,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))))) | ($spl9 | $spl52)),
  inference(backward_demodulation,[],[f2153,f118])).
fof(f2440,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,X0))),s(fun(sK1,bool),sK3)))) | ~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK2))))) ) | $spl22),
  inference(forward_demodulation,[],[f2439,f80])).
fof(f2439,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,X0))),s(fun(sK1,bool),'const.combin.K_2'(s(fun(sK1,bool),sK3),s(sK0,X0)))))) | ~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK2))))) ) | $spl22),
  inference(forward_demodulation,[],[f2353,f81])).
fof(f2353,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK2)))) | p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(sK0,X0))),s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3))),s(sK0,X0))))))) ) | $spl22),
  inference(superposition,[],[f86,f140])).
fof(f2152,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3))))))) | $spl53),
  inference(cnf_transformation,[],[f2152_D])).
fof(f2152_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3))))))) <=> ~$spl53),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl53])])).
fof(f2154,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3))))))) | ($spl9 | $spl24)),
  inference(backward_demodulation,[],[f141,f118])).
fof(f141,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))))) | $spl24),
  inference(cnf_transformation,[],[f141_D])).
fof(f141_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f2153,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3))))))) | $spl52),
  inference(cnf_transformation,[],[f2153_D])).
fof(f2153_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3))))))) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f2150,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3))))))) | $spl23),
  inference(unit_resulting_resolution,[],[f139,f71])).
fof(f2039,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))))) | ($spl5 | $spl38)),
  inference(backward_demodulation,[],[f621,f114])).
fof(f1774,plain,(
  $false | $spl49),
  inference(trivial_inequality_removal,[],[f1589])).
fof(f1589,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl49),
  inference(unit_resulting_resolution,[],[f68,f1125,f1125,f104])).
fof(f1773,plain,(
  $false | $spl49),
  inference(trivial_inequality_removal,[],[f1588])).
fof(f1588,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl49),
  inference(unit_resulting_resolution,[],[f69,f1125,f1125,f104])).
fof(f1772,plain,(
  $false | $spl49),
  inference(trivial_inequality_removal,[],[f1587])).
fof(f1587,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl49),
  inference(unit_resulting_resolution,[],[f103,f1125,f1125,f104])).
fof(f1771,plain,(
  $false | $spl49),
  inference(trivial_inequality_removal,[],[f1578])).
fof(f1578,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl49),
  inference(unit_resulting_resolution,[],[f70,f68,f1125,f104])).
fof(f1770,plain,(
  $false | $spl49),
  inference(trivial_inequality_removal,[],[f1577])).
fof(f1577,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl49),
  inference(unit_resulting_resolution,[],[f70,f69,f1125,f104])).
fof(f1769,plain,(
  $false | $spl49),
  inference(trivial_inequality_removal,[],[f1576])).
fof(f1576,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl49),
  inference(unit_resulting_resolution,[],[f70,f103,f1125,f104])).
fof(f1768,plain,(
  $false | $spl49),
  inference(trivial_inequality_removal,[],[f1557])).
fof(f1557,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl49),
  inference(unit_resulting_resolution,[],[f68,f68,f80,f1125,f102])).
fof(f1767,plain,(
  $false | $spl49),
  inference(trivial_inequality_removal,[],[f1556])).
fof(f1556,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl49),
  inference(unit_resulting_resolution,[],[f69,f68,f80,f1125,f102])).
fof(f1766,plain,(
  $false | $spl49),
  inference(trivial_inequality_removal,[],[f1555])).
fof(f1555,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl49),
  inference(unit_resulting_resolution,[],[f103,f68,f80,f1125,f102])).
fof(f1765,plain,(
  $false | $spl49),
  inference(trivial_inequality_removal,[],[f1554])).
fof(f1554,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl49),
  inference(unit_resulting_resolution,[],[f68,f69,f80,f1125,f102])).
fof(f1764,plain,(
  $false | $spl49),
  inference(trivial_inequality_removal,[],[f1553])).
fof(f1553,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl49),
  inference(unit_resulting_resolution,[],[f69,f69,f80,f1125,f102])).
fof(f1763,plain,(
  $false | $spl49),
  inference(trivial_inequality_removal,[],[f1552])).
fof(f1552,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl49),
  inference(unit_resulting_resolution,[],[f103,f69,f80,f1125,f102])).
fof(f1762,plain,(
  $false | $spl49),
  inference(trivial_inequality_removal,[],[f1551])).
fof(f1551,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl49),
  inference(unit_resulting_resolution,[],[f68,f103,f80,f1125,f102])).
fof(f1761,plain,(
  $false | $spl49),
  inference(trivial_inequality_removal,[],[f1550])).
fof(f1550,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl49),
  inference(unit_resulting_resolution,[],[f69,f103,f80,f1125,f102])).
fof(f1760,plain,(
  $false | $spl49),
  inference(trivial_inequality_removal,[],[f1549])).
fof(f1549,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl49),
  inference(unit_resulting_resolution,[],[f103,f103,f80,f1125,f102])).
fof(f1759,plain,(
  $false | ($spl18 | $spl49)),
  inference(trivial_inequality_removal,[],[f1548])).
fof(f1548,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl18 | $spl49)),
  inference(unit_resulting_resolution,[],[f68,f68,f133,f1125,f102])).
fof(f1758,plain,(
  $false | ($spl18 | $spl49)),
  inference(trivial_inequality_removal,[],[f1547])).
fof(f1547,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl18 | $spl49)),
  inference(unit_resulting_resolution,[],[f69,f68,f133,f1125,f102])).
fof(f1757,plain,(
  $false | ($spl18 | $spl49)),
  inference(trivial_inequality_removal,[],[f1546])).
fof(f1546,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl18 | $spl49)),
  inference(unit_resulting_resolution,[],[f103,f68,f133,f1125,f102])).
fof(f1756,plain,(
  $false | ($spl18 | $spl49)),
  inference(trivial_inequality_removal,[],[f1545])).
fof(f1545,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl18 | $spl49)),
  inference(unit_resulting_resolution,[],[f68,f69,f133,f1125,f102])).
fof(f1755,plain,(
  $false | ($spl18 | $spl49)),
  inference(trivial_inequality_removal,[],[f1544])).
fof(f1544,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl18 | $spl49)),
  inference(unit_resulting_resolution,[],[f69,f69,f133,f1125,f102])).
fof(f1754,plain,(
  $false | ($spl18 | $spl49)),
  inference(trivial_inequality_removal,[],[f1543])).
fof(f1543,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl18 | $spl49)),
  inference(unit_resulting_resolution,[],[f103,f69,f133,f1125,f102])).
fof(f1753,plain,(
  $false | ($spl18 | $spl49)),
  inference(trivial_inequality_removal,[],[f1542])).
fof(f1542,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl18 | $spl49)),
  inference(unit_resulting_resolution,[],[f68,f103,f133,f1125,f102])).
fof(f1752,plain,(
  $false | ($spl18 | $spl49)),
  inference(trivial_inequality_removal,[],[f1541])).
fof(f1541,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl18 | $spl49)),
  inference(unit_resulting_resolution,[],[f69,f103,f133,f1125,f102])).
fof(f1751,plain,(
  $false | ($spl18 | $spl49)),
  inference(trivial_inequality_removal,[],[f1540])).
fof(f1540,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl18 | $spl49)),
  inference(unit_resulting_resolution,[],[f103,f103,f133,f1125,f102])).
fof(f1750,plain,(
  $false | ($spl14 | $spl49)),
  inference(trivial_inequality_removal,[],[f1539])).
fof(f1539,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl14 | $spl49)),
  inference(unit_resulting_resolution,[],[f68,f68,f126,f1125,f102])).
fof(f1749,plain,(
  $false | ($spl14 | $spl49)),
  inference(trivial_inequality_removal,[],[f1538])).
fof(f1538,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl14 | $spl49)),
  inference(unit_resulting_resolution,[],[f69,f68,f126,f1125,f102])).
fof(f1748,plain,(
  $false | ($spl14 | $spl49)),
  inference(trivial_inequality_removal,[],[f1537])).
fof(f1537,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl14 | $spl49)),
  inference(unit_resulting_resolution,[],[f103,f68,f126,f1125,f102])).
fof(f1747,plain,(
  $false | ($spl14 | $spl49)),
  inference(trivial_inequality_removal,[],[f1536])).
fof(f1536,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl14 | $spl49)),
  inference(unit_resulting_resolution,[],[f68,f69,f126,f1125,f102])).
fof(f1746,plain,(
  $false | ($spl14 | $spl49)),
  inference(trivial_inequality_removal,[],[f1535])).
fof(f1535,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl14 | $spl49)),
  inference(unit_resulting_resolution,[],[f69,f69,f126,f1125,f102])).
fof(f1745,plain,(
  $false | ($spl14 | $spl49)),
  inference(trivial_inequality_removal,[],[f1534])).
fof(f1534,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl14 | $spl49)),
  inference(unit_resulting_resolution,[],[f103,f69,f126,f1125,f102])).
fof(f1744,plain,(
  $false | ($spl14 | $spl49)),
  inference(trivial_inequality_removal,[],[f1533])).
fof(f1533,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl14 | $spl49)),
  inference(unit_resulting_resolution,[],[f68,f103,f126,f1125,f102])).
fof(f1743,plain,(
  $false | ($spl14 | $spl49)),
  inference(trivial_inequality_removal,[],[f1532])).
fof(f1532,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl14 | $spl49)),
  inference(unit_resulting_resolution,[],[f69,f103,f126,f1125,f102])).
fof(f1742,plain,(
  $false | ($spl14 | $spl49)),
  inference(trivial_inequality_removal,[],[f1531])).
fof(f1531,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl14 | $spl49)),
  inference(unit_resulting_resolution,[],[f103,f103,f126,f1125,f102])).
fof(f1741,plain,(
  $false | ($spl10 | $spl49)),
  inference(trivial_inequality_removal,[],[f1530])).
fof(f1530,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl10 | $spl49)),
  inference(unit_resulting_resolution,[],[f68,f68,f121,f1125,f102])).
fof(f1740,plain,(
  $false | ($spl10 | $spl49)),
  inference(trivial_inequality_removal,[],[f1529])).
fof(f1529,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl10 | $spl49)),
  inference(unit_resulting_resolution,[],[f69,f68,f121,f1125,f102])).
fof(f1739,plain,(
  $false | ($spl10 | $spl49)),
  inference(trivial_inequality_removal,[],[f1528])).
fof(f1528,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl10 | $spl49)),
  inference(unit_resulting_resolution,[],[f103,f68,f121,f1125,f102])).
fof(f1738,plain,(
  $false | ($spl10 | $spl49)),
  inference(trivial_inequality_removal,[],[f1527])).
fof(f1527,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl10 | $spl49)),
  inference(unit_resulting_resolution,[],[f68,f69,f121,f1125,f102])).
fof(f1737,plain,(
  $false | ($spl10 | $spl49)),
  inference(trivial_inequality_removal,[],[f1526])).
fof(f1526,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl10 | $spl49)),
  inference(unit_resulting_resolution,[],[f69,f69,f121,f1125,f102])).
fof(f1736,plain,(
  $false | ($spl10 | $spl49)),
  inference(trivial_inequality_removal,[],[f1525])).
fof(f1525,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl10 | $spl49)),
  inference(unit_resulting_resolution,[],[f103,f69,f121,f1125,f102])).
fof(f1735,plain,(
  $false | ($spl10 | $spl49)),
  inference(trivial_inequality_removal,[],[f1524])).
fof(f1524,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl10 | $spl49)),
  inference(unit_resulting_resolution,[],[f68,f103,f121,f1125,f102])).
fof(f1734,plain,(
  $false | ($spl10 | $spl49)),
  inference(trivial_inequality_removal,[],[f1523])).
fof(f1523,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl10 | $spl49)),
  inference(unit_resulting_resolution,[],[f69,f103,f121,f1125,f102])).
fof(f1733,plain,(
  $false | ($spl10 | $spl49)),
  inference(trivial_inequality_removal,[],[f1522])).
fof(f1522,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl10 | $spl49)),
  inference(unit_resulting_resolution,[],[f103,f103,f121,f1125,f102])).
fof(f1732,plain,(
  $false | $spl49),
  inference(trivial_inequality_removal,[],[f1497])).
fof(f1497,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl49),
  inference(unit_resulting_resolution,[],[f68,f1125,f68,f1125,f101])).
fof(f1731,plain,(
  $false | $spl49),
  inference(trivial_inequality_removal,[],[f1496])).
fof(f1496,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl49),
  inference(unit_resulting_resolution,[],[f69,f1125,f68,f1125,f101])).
fof(f1730,plain,(
  $false | $spl49),
  inference(trivial_inequality_removal,[],[f1495])).
fof(f1495,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl49),
  inference(unit_resulting_resolution,[],[f103,f1125,f68,f1125,f101])).
fof(f1729,plain,(
  $false | $spl49),
  inference(trivial_inequality_removal,[],[f1491])).
fof(f1491,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl49),
  inference(unit_resulting_resolution,[],[f68,f1125,f69,f1125,f101])).
fof(f1728,plain,(
  $false | $spl49),
  inference(trivial_inequality_removal,[],[f1490])).
fof(f1490,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl49),
  inference(unit_resulting_resolution,[],[f69,f1125,f69,f1125,f101])).
fof(f1727,plain,(
  $false | $spl49),
  inference(trivial_inequality_removal,[],[f1489])).
fof(f1489,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl49),
  inference(unit_resulting_resolution,[],[f103,f1125,f69,f1125,f101])).
fof(f1726,plain,(
  $false | $spl49),
  inference(trivial_inequality_removal,[],[f1485])).
fof(f1485,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl49),
  inference(unit_resulting_resolution,[],[f68,f1125,f103,f1125,f101])).
fof(f1725,plain,(
  $false | $spl49),
  inference(trivial_inequality_removal,[],[f1484])).
fof(f1484,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl49),
  inference(unit_resulting_resolution,[],[f69,f1125,f103,f1125,f101])).
fof(f1724,plain,(
  $false | $spl49),
  inference(trivial_inequality_removal,[],[f1483])).
fof(f1483,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl49),
  inference(unit_resulting_resolution,[],[f103,f1125,f103,f1125,f101])).
fof(f1723,plain,(
  $false | $spl49),
  inference(trivial_inequality_removal,[],[f1464])).
fof(f1464,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl49),
  inference(unit_resulting_resolution,[],[f70,f68,f68,f1125,f101])).
fof(f1722,plain,(
  $false | $spl49),
  inference(trivial_inequality_removal,[],[f1463])).
fof(f1463,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl49),
  inference(unit_resulting_resolution,[],[f70,f69,f68,f1125,f101])).
fof(f1721,plain,(
  $false | $spl49),
  inference(trivial_inequality_removal,[],[f1462])).
fof(f1462,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl49),
  inference(unit_resulting_resolution,[],[f70,f103,f68,f1125,f101])).
fof(f1720,plain,(
  $false | $spl49),
  inference(trivial_inequality_removal,[],[f1461])).
fof(f1461,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl49),
  inference(unit_resulting_resolution,[],[f70,f68,f69,f1125,f101])).
fof(f1719,plain,(
  $false | $spl49),
  inference(trivial_inequality_removal,[],[f1460])).
fof(f1460,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl49),
  inference(unit_resulting_resolution,[],[f70,f69,f69,f1125,f101])).
fof(f1718,plain,(
  $false | $spl49),
  inference(trivial_inequality_removal,[],[f1459])).
fof(f1459,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl49),
  inference(unit_resulting_resolution,[],[f70,f103,f69,f1125,f101])).
fof(f1717,plain,(
  $false | $spl49),
  inference(trivial_inequality_removal,[],[f1458])).
fof(f1458,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl49),
  inference(unit_resulting_resolution,[],[f70,f68,f103,f1125,f101])).
fof(f1716,plain,(
  $false | $spl49),
  inference(trivial_inequality_removal,[],[f1457])).
fof(f1457,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl49),
  inference(unit_resulting_resolution,[],[f70,f69,f103,f1125,f101])).
fof(f1715,plain,(
  $false | $spl49),
  inference(trivial_inequality_removal,[],[f1456])).
fof(f1456,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl49),
  inference(unit_resulting_resolution,[],[f70,f103,f103,f1125,f101])).
fof(f1714,plain,(
  $false | $spl49),
  inference(trivial_inequality_removal,[],[f1441])).
fof(f1441,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl49),
  inference(unit_resulting_resolution,[],[f1125,f68,f80,f1125,f100])).
fof(f1713,plain,(
  $false | $spl49),
  inference(trivial_inequality_removal,[],[f1439])).
fof(f1439,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl49),
  inference(unit_resulting_resolution,[],[f1125,f69,f80,f1125,f100])).
fof(f1712,plain,(
  $false | $spl49),
  inference(trivial_inequality_removal,[],[f1437])).
fof(f1437,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl49),
  inference(unit_resulting_resolution,[],[f1125,f103,f80,f1125,f100])).
fof(f1711,plain,(
  $false | ($spl18 | $spl49)),
  inference(trivial_inequality_removal,[],[f1421])).
fof(f1421,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl18 | $spl49)),
  inference(unit_resulting_resolution,[],[f1125,f68,f133,f1125,f100])).
fof(f1710,plain,(
  $false | ($spl18 | $spl49)),
  inference(trivial_inequality_removal,[],[f1419])).
fof(f1419,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl18 | $spl49)),
  inference(unit_resulting_resolution,[],[f1125,f69,f133,f1125,f100])).
fof(f1709,plain,(
  $false | ($spl18 | $spl49)),
  inference(trivial_inequality_removal,[],[f1417])).
fof(f1417,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl18 | $spl49)),
  inference(unit_resulting_resolution,[],[f1125,f103,f133,f1125,f100])).
fof(f1708,plain,(
  $false | ($spl14 | $spl49)),
  inference(trivial_inequality_removal,[],[f1401])).
fof(f1401,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl14 | $spl49)),
  inference(unit_resulting_resolution,[],[f1125,f68,f126,f1125,f100])).
fof(f1707,plain,(
  $false | ($spl14 | $spl49)),
  inference(trivial_inequality_removal,[],[f1399])).
fof(f1399,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl14 | $spl49)),
  inference(unit_resulting_resolution,[],[f1125,f69,f126,f1125,f100])).
fof(f1706,plain,(
  $false | ($spl14 | $spl49)),
  inference(trivial_inequality_removal,[],[f1397])).
fof(f1397,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl14 | $spl49)),
  inference(unit_resulting_resolution,[],[f1125,f103,f126,f1125,f100])).
fof(f1705,plain,(
  $false | ($spl10 | $spl49)),
  inference(trivial_inequality_removal,[],[f1381])).
fof(f1381,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl10 | $spl49)),
  inference(unit_resulting_resolution,[],[f1125,f68,f121,f1125,f100])).
fof(f1704,plain,(
  $false | ($spl10 | $spl49)),
  inference(trivial_inequality_removal,[],[f1379])).
fof(f1379,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl10 | $spl49)),
  inference(unit_resulting_resolution,[],[f1125,f69,f121,f1125,f100])).
fof(f1703,plain,(
  $false | ($spl10 | $spl49)),
  inference(trivial_inequality_removal,[],[f1377])).
fof(f1377,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl10 | $spl49)),
  inference(unit_resulting_resolution,[],[f1125,f103,f121,f1125,f100])).
fof(f1702,plain,(
  $false | $spl49),
  inference(trivial_inequality_removal,[],[f1346])).
fof(f1346,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl49),
  inference(unit_resulting_resolution,[],[f68,f1125,f80,f1125,f100])).
fof(f1701,plain,(
  $false | $spl49),
  inference(trivial_inequality_removal,[],[f1345])).
fof(f1345,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl49),
  inference(unit_resulting_resolution,[],[f69,f1125,f80,f1125,f100])).
fof(f1700,plain,(
  $false | $spl49),
  inference(trivial_inequality_removal,[],[f1344])).
fof(f1344,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl49),
  inference(unit_resulting_resolution,[],[f103,f1125,f80,f1125,f100])).
fof(f1699,plain,(
  $false | $spl49),
  inference(trivial_inequality_removal,[],[f1343])).
fof(f1343,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl49),
  inference(unit_resulting_resolution,[],[f68,f70,f80,f1125,f100])).
fof(f1698,plain,(
  $false | $spl49),
  inference(trivial_inequality_removal,[],[f1342])).
fof(f1342,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl49),
  inference(unit_resulting_resolution,[],[f69,f70,f80,f1125,f100])).
fof(f1697,plain,(
  $false | $spl49),
  inference(trivial_inequality_removal,[],[f1341])).
fof(f1341,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl49),
  inference(unit_resulting_resolution,[],[f103,f70,f80,f1125,f100])).
fof(f1696,plain,(
  $false | ($spl20 | $spl49)),
  inference(trivial_inequality_removal,[],[f1340])).
fof(f1340,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl20 | $spl49)),
  inference(unit_resulting_resolution,[],[f68,f70,f134,f1125,f100])).
fof(f1695,plain,(
  $false | ($spl20 | $spl49)),
  inference(trivial_inequality_removal,[],[f1339])).
fof(f1339,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl20 | $spl49)),
  inference(unit_resulting_resolution,[],[f69,f70,f134,f1125,f100])).
fof(f1694,plain,(
  $false | ($spl20 | $spl49)),
  inference(trivial_inequality_removal,[],[f1338])).
fof(f1338,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl20 | $spl49)),
  inference(unit_resulting_resolution,[],[f103,f70,f134,f1125,f100])).
fof(f1693,plain,(
  $false | ($spl16 | $spl49)),
  inference(trivial_inequality_removal,[],[f1337])).
fof(f1337,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl16 | $spl49)),
  inference(unit_resulting_resolution,[],[f68,f70,f127,f1125,f100])).
fof(f1692,plain,(
  $false | ($spl16 | $spl49)),
  inference(trivial_inequality_removal,[],[f1336])).
fof(f1336,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl16 | $spl49)),
  inference(unit_resulting_resolution,[],[f69,f70,f127,f1125,f100])).
fof(f1691,plain,(
  $false | ($spl16 | $spl49)),
  inference(trivial_inequality_removal,[],[f1335])).
fof(f1335,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl16 | $spl49)),
  inference(unit_resulting_resolution,[],[f103,f70,f127,f1125,f100])).
fof(f1690,plain,(
  $false | ($spl12 | $spl49)),
  inference(trivial_inequality_removal,[],[f1334])).
fof(f1334,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl12 | $spl49)),
  inference(unit_resulting_resolution,[],[f68,f70,f122,f1125,f100])).
fof(f1689,plain,(
  $false | ($spl12 | $spl49)),
  inference(trivial_inequality_removal,[],[f1333])).
fof(f1333,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl12 | $spl49)),
  inference(unit_resulting_resolution,[],[f69,f70,f122,f1125,f100])).
fof(f1688,plain,(
  $false | ($spl12 | $spl49)),
  inference(trivial_inequality_removal,[],[f1332])).
fof(f1332,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl12 | $spl49)),
  inference(unit_resulting_resolution,[],[f103,f70,f122,f1125,f100])).
fof(f1687,plain,(
  $false | ($spl18 | $spl49)),
  inference(trivial_inequality_removal,[],[f1331])).
fof(f1331,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl18 | $spl49)),
  inference(unit_resulting_resolution,[],[f68,f1125,f133,f1125,f100])).
fof(f1686,plain,(
  $false | ($spl18 | $spl49)),
  inference(trivial_inequality_removal,[],[f1330])).
fof(f1330,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl18 | $spl49)),
  inference(unit_resulting_resolution,[],[f69,f1125,f133,f1125,f100])).
fof(f1685,plain,(
  $false | ($spl18 | $spl49)),
  inference(trivial_inequality_removal,[],[f1329])).
fof(f1329,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl18 | $spl49)),
  inference(unit_resulting_resolution,[],[f103,f1125,f133,f1125,f100])).
fof(f1684,plain,(
  $false | ($spl14 | $spl49)),
  inference(trivial_inequality_removal,[],[f1328])).
fof(f1328,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl14 | $spl49)),
  inference(unit_resulting_resolution,[],[f68,f1125,f126,f1125,f100])).
fof(f1683,plain,(
  $false | ($spl14 | $spl49)),
  inference(trivial_inequality_removal,[],[f1327])).
fof(f1327,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl14 | $spl49)),
  inference(unit_resulting_resolution,[],[f69,f1125,f126,f1125,f100])).
fof(f1682,plain,(
  $false | ($spl14 | $spl49)),
  inference(trivial_inequality_removal,[],[f1326])).
fof(f1326,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl14 | $spl49)),
  inference(unit_resulting_resolution,[],[f103,f1125,f126,f1125,f100])).
fof(f1681,plain,(
  $false | ($spl10 | $spl49)),
  inference(trivial_inequality_removal,[],[f1325])).
fof(f1325,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl10 | $spl49)),
  inference(unit_resulting_resolution,[],[f68,f1125,f121,f1125,f100])).
fof(f1680,plain,(
  $false | ($spl10 | $spl49)),
  inference(trivial_inequality_removal,[],[f1324])).
fof(f1324,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl10 | $spl49)),
  inference(unit_resulting_resolution,[],[f69,f1125,f121,f1125,f100])).
fof(f1679,plain,(
  $false | ($spl10 | $spl49)),
  inference(trivial_inequality_removal,[],[f1323])).
fof(f1323,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl10 | $spl49)),
  inference(unit_resulting_resolution,[],[f103,f1125,f121,f1125,f100])).
fof(f1678,plain,(
  $false | $spl49),
  inference(trivial_inequality_removal,[],[f1304])).
fof(f1304,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl49),
  inference(unit_resulting_resolution,[],[f68,f68,f80,f1125,f98])).
fof(f1677,plain,(
  $false | $spl49),
  inference(trivial_inequality_removal,[],[f1303])).
fof(f1303,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl49),
  inference(unit_resulting_resolution,[],[f69,f68,f80,f1125,f98])).
fof(f1676,plain,(
  $false | $spl49),
  inference(trivial_inequality_removal,[],[f1302])).
fof(f1302,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl49),
  inference(unit_resulting_resolution,[],[f103,f68,f80,f1125,f98])).
fof(f1675,plain,(
  $false | $spl49),
  inference(trivial_inequality_removal,[],[f1301])).
fof(f1301,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl49),
  inference(unit_resulting_resolution,[],[f68,f69,f80,f1125,f98])).
fof(f1674,plain,(
  $false | $spl49),
  inference(trivial_inequality_removal,[],[f1300])).
fof(f1300,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl49),
  inference(unit_resulting_resolution,[],[f69,f69,f80,f1125,f98])).
fof(f1673,plain,(
  $false | $spl49),
  inference(trivial_inequality_removal,[],[f1299])).
fof(f1299,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl49),
  inference(unit_resulting_resolution,[],[f103,f69,f80,f1125,f98])).
fof(f1672,plain,(
  $false | $spl49),
  inference(trivial_inequality_removal,[],[f1298])).
fof(f1298,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl49),
  inference(unit_resulting_resolution,[],[f68,f103,f80,f1125,f98])).
fof(f1671,plain,(
  $false | $spl49),
  inference(trivial_inequality_removal,[],[f1297])).
fof(f1297,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl49),
  inference(unit_resulting_resolution,[],[f69,f103,f80,f1125,f98])).
fof(f1670,plain,(
  $false | $spl49),
  inference(trivial_inequality_removal,[],[f1296])).
fof(f1296,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl49),
  inference(unit_resulting_resolution,[],[f103,f103,f80,f1125,f98])).
fof(f1669,plain,(
  $false | ($spl18 | $spl49)),
  inference(trivial_inequality_removal,[],[f1295])).
fof(f1295,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl18 | $spl49)),
  inference(unit_resulting_resolution,[],[f68,f68,f133,f1125,f98])).
fof(f1668,plain,(
  $false | ($spl18 | $spl49)),
  inference(trivial_inequality_removal,[],[f1294])).
fof(f1294,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl18 | $spl49)),
  inference(unit_resulting_resolution,[],[f69,f68,f133,f1125,f98])).
fof(f1667,plain,(
  $false | ($spl18 | $spl49)),
  inference(trivial_inequality_removal,[],[f1293])).
fof(f1293,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl18 | $spl49)),
  inference(unit_resulting_resolution,[],[f103,f68,f133,f1125,f98])).
fof(f1666,plain,(
  $false | ($spl18 | $spl49)),
  inference(trivial_inequality_removal,[],[f1292])).
fof(f1292,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl18 | $spl49)),
  inference(unit_resulting_resolution,[],[f68,f69,f133,f1125,f98])).
fof(f1665,plain,(
  $false | ($spl18 | $spl49)),
  inference(trivial_inequality_removal,[],[f1291])).
fof(f1291,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl18 | $spl49)),
  inference(unit_resulting_resolution,[],[f69,f69,f133,f1125,f98])).
fof(f1664,plain,(
  $false | ($spl18 | $spl49)),
  inference(trivial_inequality_removal,[],[f1290])).
fof(f1290,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl18 | $spl49)),
  inference(unit_resulting_resolution,[],[f103,f69,f133,f1125,f98])).
fof(f1663,plain,(
  $false | ($spl18 | $spl49)),
  inference(trivial_inequality_removal,[],[f1289])).
fof(f1289,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl18 | $spl49)),
  inference(unit_resulting_resolution,[],[f68,f103,f133,f1125,f98])).
fof(f1662,plain,(
  $false | ($spl18 | $spl49)),
  inference(trivial_inequality_removal,[],[f1288])).
fof(f1288,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl18 | $spl49)),
  inference(unit_resulting_resolution,[],[f69,f103,f133,f1125,f98])).
fof(f1661,plain,(
  $false | ($spl18 | $spl49)),
  inference(trivial_inequality_removal,[],[f1287])).
fof(f1287,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl18 | $spl49)),
  inference(unit_resulting_resolution,[],[f103,f103,f133,f1125,f98])).
fof(f1660,plain,(
  $false | ($spl14 | $spl49)),
  inference(trivial_inequality_removal,[],[f1286])).
fof(f1286,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl14 | $spl49)),
  inference(unit_resulting_resolution,[],[f68,f68,f126,f1125,f98])).
fof(f1659,plain,(
  $false | ($spl14 | $spl49)),
  inference(trivial_inequality_removal,[],[f1285])).
fof(f1285,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl14 | $spl49)),
  inference(unit_resulting_resolution,[],[f69,f68,f126,f1125,f98])).
fof(f1658,plain,(
  $false | ($spl14 | $spl49)),
  inference(trivial_inequality_removal,[],[f1284])).
fof(f1284,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl14 | $spl49)),
  inference(unit_resulting_resolution,[],[f103,f68,f126,f1125,f98])).
fof(f1657,plain,(
  $false | ($spl14 | $spl49)),
  inference(trivial_inequality_removal,[],[f1283])).
fof(f1283,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl14 | $spl49)),
  inference(unit_resulting_resolution,[],[f68,f69,f126,f1125,f98])).
fof(f1656,plain,(
  $false | ($spl14 | $spl49)),
  inference(trivial_inequality_removal,[],[f1282])).
fof(f1282,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl14 | $spl49)),
  inference(unit_resulting_resolution,[],[f69,f69,f126,f1125,f98])).
fof(f1655,plain,(
  $false | ($spl14 | $spl49)),
  inference(trivial_inequality_removal,[],[f1281])).
fof(f1281,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl14 | $spl49)),
  inference(unit_resulting_resolution,[],[f103,f69,f126,f1125,f98])).
fof(f1654,plain,(
  $false | ($spl14 | $spl49)),
  inference(trivial_inequality_removal,[],[f1280])).
fof(f1280,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl14 | $spl49)),
  inference(unit_resulting_resolution,[],[f68,f103,f126,f1125,f98])).
fof(f1653,plain,(
  $false | ($spl14 | $spl49)),
  inference(trivial_inequality_removal,[],[f1279])).
fof(f1279,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl14 | $spl49)),
  inference(unit_resulting_resolution,[],[f69,f103,f126,f1125,f98])).
fof(f1652,plain,(
  $false | ($spl14 | $spl49)),
  inference(trivial_inequality_removal,[],[f1278])).
fof(f1278,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl14 | $spl49)),
  inference(unit_resulting_resolution,[],[f103,f103,f126,f1125,f98])).
fof(f1651,plain,(
  $false | ($spl10 | $spl49)),
  inference(trivial_inequality_removal,[],[f1277])).
fof(f1277,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl10 | $spl49)),
  inference(unit_resulting_resolution,[],[f68,f68,f121,f1125,f98])).
fof(f1650,plain,(
  $false | ($spl10 | $spl49)),
  inference(trivial_inequality_removal,[],[f1276])).
fof(f1276,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl10 | $spl49)),
  inference(unit_resulting_resolution,[],[f69,f68,f121,f1125,f98])).
fof(f1649,plain,(
  $false | ($spl10 | $spl49)),
  inference(trivial_inequality_removal,[],[f1275])).
fof(f1275,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl10 | $spl49)),
  inference(unit_resulting_resolution,[],[f103,f68,f121,f1125,f98])).
fof(f1648,plain,(
  $false | ($spl10 | $spl49)),
  inference(trivial_inequality_removal,[],[f1274])).
fof(f1274,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl10 | $spl49)),
  inference(unit_resulting_resolution,[],[f68,f69,f121,f1125,f98])).
fof(f1647,plain,(
  $false | ($spl10 | $spl49)),
  inference(trivial_inequality_removal,[],[f1273])).
fof(f1273,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl10 | $spl49)),
  inference(unit_resulting_resolution,[],[f69,f69,f121,f1125,f98])).
fof(f1646,plain,(
  $false | ($spl10 | $spl49)),
  inference(trivial_inequality_removal,[],[f1272])).
fof(f1272,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl10 | $spl49)),
  inference(unit_resulting_resolution,[],[f103,f69,f121,f1125,f98])).
fof(f1645,plain,(
  $false | ($spl10 | $spl49)),
  inference(trivial_inequality_removal,[],[f1271])).
fof(f1271,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl10 | $spl49)),
  inference(unit_resulting_resolution,[],[f68,f103,f121,f1125,f98])).
fof(f1644,plain,(
  $false | ($spl10 | $spl49)),
  inference(trivial_inequality_removal,[],[f1270])).
fof(f1270,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl10 | $spl49)),
  inference(unit_resulting_resolution,[],[f69,f103,f121,f1125,f98])).
fof(f1643,plain,(
  $false | ($spl10 | $spl49)),
  inference(trivial_inequality_removal,[],[f1269])).
fof(f1269,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl10 | $spl49)),
  inference(unit_resulting_resolution,[],[f103,f103,f121,f1125,f98])).
fof(f1642,plain,(
  $false | $spl49),
  inference(trivial_inequality_removal,[],[f1254])).
fof(f1254,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl49),
  inference(unit_resulting_resolution,[],[f1125,f68,f80,f1125,f96])).
fof(f1641,plain,(
  $false | $spl49),
  inference(trivial_inequality_removal,[],[f1252])).
fof(f1252,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl49),
  inference(unit_resulting_resolution,[],[f1125,f69,f80,f1125,f96])).
fof(f1640,plain,(
  $false | $spl49),
  inference(trivial_inequality_removal,[],[f1250])).
fof(f1250,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl49),
  inference(unit_resulting_resolution,[],[f1125,f103,f80,f1125,f96])).
fof(f1639,plain,(
  $false | ($spl18 | $spl49)),
  inference(trivial_inequality_removal,[],[f1234])).
fof(f1234,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl18 | $spl49)),
  inference(unit_resulting_resolution,[],[f1125,f68,f133,f1125,f96])).
fof(f1638,plain,(
  $false | ($spl18 | $spl49)),
  inference(trivial_inequality_removal,[],[f1232])).
fof(f1232,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl18 | $spl49)),
  inference(unit_resulting_resolution,[],[f1125,f69,f133,f1125,f96])).
fof(f1637,plain,(
  $false | ($spl18 | $spl49)),
  inference(trivial_inequality_removal,[],[f1230])).
fof(f1230,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl18 | $spl49)),
  inference(unit_resulting_resolution,[],[f1125,f103,f133,f1125,f96])).
fof(f1636,plain,(
  $false | ($spl14 | $spl49)),
  inference(trivial_inequality_removal,[],[f1214])).
fof(f1214,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl14 | $spl49)),
  inference(unit_resulting_resolution,[],[f1125,f68,f126,f1125,f96])).
fof(f1635,plain,(
  $false | ($spl14 | $spl49)),
  inference(trivial_inequality_removal,[],[f1212])).
fof(f1212,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl14 | $spl49)),
  inference(unit_resulting_resolution,[],[f1125,f69,f126,f1125,f96])).
fof(f1634,plain,(
  $false | ($spl14 | $spl49)),
  inference(trivial_inequality_removal,[],[f1210])).
fof(f1210,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl14 | $spl49)),
  inference(unit_resulting_resolution,[],[f1125,f103,f126,f1125,f96])).
fof(f1633,plain,(
  $false | ($spl10 | $spl49)),
  inference(trivial_inequality_removal,[],[f1194])).
fof(f1194,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl10 | $spl49)),
  inference(unit_resulting_resolution,[],[f1125,f68,f121,f1125,f96])).
fof(f1632,plain,(
  $false | ($spl10 | $spl49)),
  inference(trivial_inequality_removal,[],[f1192])).
fof(f1192,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl10 | $spl49)),
  inference(unit_resulting_resolution,[],[f1125,f69,f121,f1125,f96])).
fof(f1631,plain,(
  $false | ($spl10 | $spl49)),
  inference(trivial_inequality_removal,[],[f1190])).
fof(f1190,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl10 | $spl49)),
  inference(unit_resulting_resolution,[],[f1125,f103,f121,f1125,f96])).
fof(f1630,plain,(
  $false | $spl49),
  inference(trivial_inequality_removal,[],[f1159])).
fof(f1159,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl49),
  inference(unit_resulting_resolution,[],[f68,f1125,f80,f1125,f96])).
fof(f1629,plain,(
  $false | $spl49),
  inference(trivial_inequality_removal,[],[f1158])).
fof(f1158,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl49),
  inference(unit_resulting_resolution,[],[f69,f1125,f80,f1125,f96])).
fof(f1628,plain,(
  $false | $spl49),
  inference(trivial_inequality_removal,[],[f1157])).
fof(f1157,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl49),
  inference(unit_resulting_resolution,[],[f103,f1125,f80,f1125,f96])).
fof(f1627,plain,(
  $false | $spl49),
  inference(trivial_inequality_removal,[],[f1156])).
fof(f1156,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl49),
  inference(unit_resulting_resolution,[],[f68,f70,f80,f1125,f96])).
fof(f1626,plain,(
  $false | $spl49),
  inference(trivial_inequality_removal,[],[f1155])).
fof(f1155,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl49),
  inference(unit_resulting_resolution,[],[f69,f70,f80,f1125,f96])).
fof(f1625,plain,(
  $false | $spl49),
  inference(trivial_inequality_removal,[],[f1154])).
fof(f1154,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl49),
  inference(unit_resulting_resolution,[],[f103,f70,f80,f1125,f96])).
fof(f1624,plain,(
  $false | ($spl20 | $spl49)),
  inference(trivial_inequality_removal,[],[f1153])).
fof(f1153,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl20 | $spl49)),
  inference(unit_resulting_resolution,[],[f68,f70,f134,f1125,f96])).
fof(f1623,plain,(
  $false | ($spl20 | $spl49)),
  inference(trivial_inequality_removal,[],[f1152])).
fof(f1152,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl20 | $spl49)),
  inference(unit_resulting_resolution,[],[f69,f70,f134,f1125,f96])).
fof(f1622,plain,(
  $false | ($spl20 | $spl49)),
  inference(trivial_inequality_removal,[],[f1151])).
fof(f1151,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl20 | $spl49)),
  inference(unit_resulting_resolution,[],[f103,f70,f134,f1125,f96])).
fof(f1621,plain,(
  $false | ($spl16 | $spl49)),
  inference(trivial_inequality_removal,[],[f1150])).
fof(f1150,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl16 | $spl49)),
  inference(unit_resulting_resolution,[],[f68,f70,f127,f1125,f96])).
fof(f1620,plain,(
  $false | ($spl16 | $spl49)),
  inference(trivial_inequality_removal,[],[f1149])).
fof(f1149,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl16 | $spl49)),
  inference(unit_resulting_resolution,[],[f69,f70,f127,f1125,f96])).
fof(f1619,plain,(
  $false | ($spl16 | $spl49)),
  inference(trivial_inequality_removal,[],[f1148])).
fof(f1148,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl16 | $spl49)),
  inference(unit_resulting_resolution,[],[f103,f70,f127,f1125,f96])).
fof(f1618,plain,(
  $false | ($spl12 | $spl49)),
  inference(trivial_inequality_removal,[],[f1147])).
fof(f1147,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl12 | $spl49)),
  inference(unit_resulting_resolution,[],[f68,f70,f122,f1125,f96])).
fof(f1617,plain,(
  $false | ($spl12 | $spl49)),
  inference(trivial_inequality_removal,[],[f1146])).
fof(f1146,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl12 | $spl49)),
  inference(unit_resulting_resolution,[],[f69,f70,f122,f1125,f96])).
fof(f1616,plain,(
  $false | ($spl12 | $spl49)),
  inference(trivial_inequality_removal,[],[f1145])).
fof(f1145,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl12 | $spl49)),
  inference(unit_resulting_resolution,[],[f103,f70,f122,f1125,f96])).
fof(f1615,plain,(
  $false | ($spl18 | $spl49)),
  inference(trivial_inequality_removal,[],[f1144])).
fof(f1144,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl18 | $spl49)),
  inference(unit_resulting_resolution,[],[f68,f1125,f133,f1125,f96])).
fof(f1614,plain,(
  $false | ($spl18 | $spl49)),
  inference(trivial_inequality_removal,[],[f1143])).
fof(f1143,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl18 | $spl49)),
  inference(unit_resulting_resolution,[],[f69,f1125,f133,f1125,f96])).
fof(f1613,plain,(
  $false | ($spl18 | $spl49)),
  inference(trivial_inequality_removal,[],[f1142])).
fof(f1142,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl18 | $spl49)),
  inference(unit_resulting_resolution,[],[f103,f1125,f133,f1125,f96])).
fof(f1612,plain,(
  $false | ($spl14 | $spl49)),
  inference(trivial_inequality_removal,[],[f1141])).
fof(f1141,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl14 | $spl49)),
  inference(unit_resulting_resolution,[],[f68,f1125,f126,f1125,f96])).
fof(f1611,plain,(
  $false | ($spl14 | $spl49)),
  inference(trivial_inequality_removal,[],[f1140])).
fof(f1140,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl14 | $spl49)),
  inference(unit_resulting_resolution,[],[f69,f1125,f126,f1125,f96])).
fof(f1610,plain,(
  $false | ($spl14 | $spl49)),
  inference(trivial_inequality_removal,[],[f1139])).
fof(f1139,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl14 | $spl49)),
  inference(unit_resulting_resolution,[],[f103,f1125,f126,f1125,f96])).
fof(f1609,plain,(
  $false | ($spl10 | $spl49)),
  inference(trivial_inequality_removal,[],[f1138])).
fof(f1138,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl10 | $spl49)),
  inference(unit_resulting_resolution,[],[f68,f1125,f121,f1125,f96])).
fof(f1608,plain,(
  $false | ($spl10 | $spl49)),
  inference(trivial_inequality_removal,[],[f1137])).
fof(f1137,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl10 | $spl49)),
  inference(unit_resulting_resolution,[],[f69,f1125,f121,f1125,f96])).
fof(f1607,plain,(
  $false | ($spl10 | $spl49)),
  inference(trivial_inequality_removal,[],[f1136])).
fof(f1136,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl10 | $spl49)),
  inference(unit_resulting_resolution,[],[f103,f1125,f121,f1125,f96])).
fof(f1135,plain,(
  $false | $spl49),
  inference(unit_resulting_resolution,[],[f103,f1125])).
fof(f1604,plain,(
  $false | $spl49),
  inference(resolution,[],[f1125,f103])).
fof(f1605,plain,(
  $false | $spl49),
  inference(resolution,[],[f1125,f69])).
fof(f1606,plain,(
  $false | $spl49),
  inference(resolution,[],[f1125,f68])).
fof(f1123,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,X0))),s(fun(sK1,bool),sK3)))) | ~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK2))))) ) | $spl18),
  inference(forward_demodulation,[],[f1122,f80])).
fof(f1122,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,X0))),s(fun(sK1,bool),'const.combin.K_2'(s(fun(sK1,bool),sK3),s(sK0,X0)))))) | ~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK2))))) ) | $spl18),
  inference(forward_demodulation,[],[f1040,f81])).
fof(f1040,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK2)))) | p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(sK0,X0))),s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3))),s(sK0,X0))))))) ) | $spl18),
  inference(superposition,[],[f86,f133])).
fof(f139,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3))))))) | $spl23),
  inference(cnf_transformation,[],[f139_D])).
fof(f139_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3))))))) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f958,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3))))))) | ($spl9 | $spl46)),
  inference(backward_demodulation,[],[f956,f118])).
fof(f954,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))))) | $spl25),
  inference(trivial_inequality_removal,[],[f953])).
fof(f953,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))))) | $spl25),
  inference(superposition,[],[f142,f71])).
fof(f952,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))))) | $spl25),
  inference(unit_resulting_resolution,[],[f142,f71])).
fof(f781,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK7('const.combin.K_1'(s(fun(sK1,bool),sK3)),sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK0,sK1)),s(fun(sK0,bool),sK2)))) | $spl38),
  inference(superposition,[],[f87,f621])).
fof(f737,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3))))))) | $spl39),
  inference(trivial_inequality_removal,[],[f736])).
fof(f736,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3))))))) | $spl39),
  inference(superposition,[],[f622,f71])).
fof(f735,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3))))))) | $spl39),
  inference(unit_resulting_resolution,[],[f622,f71])).
fof(f132,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3))))))) | $spl19),
  inference(cnf_transformation,[],[f132_D])).
fof(f132_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3))))))) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f733,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3))))))) | ($spl5 | $spl36)),
  inference(backward_demodulation,[],[f619,f114])).
fof(f732,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3))))))) | s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))))) = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))))))),
  introduced(theory_axiom,[])).
fof(f721,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2)))) | $spl40),
  inference(cnf_transformation,[],[f721_D])).
fof(f721_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2)))) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f667,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK6(sK3,sK2,sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1)),sK1,sK0)),s(fun(sK0,bool),sK2)))) | $spl20),
  inference(superposition,[],[f84,f134])).
fof(f622,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3))))))) | $spl39),
  inference(cnf_transformation,[],[f622_D])).
fof(f622_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3))))))) <=> ~$spl39),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl39])])).
fof(f620,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3))))))) | ($spl5 | $spl20)),
  inference(backward_demodulation,[],[f134,f114])).
fof(f617,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))))) | $spl21),
  inference(trivial_inequality_removal,[],[f616])).
fof(f616,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))))) | $spl21),
  inference(superposition,[],[f135,f71])).
fof(f615,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))))) | $spl21),
  inference(unit_resulting_resolution,[],[f135,f71])).
fof(f369,plain,(
  s(bool,'T_0') != s(bool,f(s(fun(sK0,sK1),sK5('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),bool,fun(sK0,sK1))))) | $spl33),
  inference(cnf_transformation,[],[f369_D])).
fof(f369_D,plain,(
  s(bool,'T_0') != s(bool,f(s(fun(sK0,sK1),sK5('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),bool,fun(sK0,sK1))))) <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f520,plain,(
  s(bool,'T_0') != s(bool,f(s(fun(sK0,sK1),sK5('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),bool,fun(sK0,sK1))))) | $spl14),
  inference(global_subsumption,[],[f67,f126,f457])).
fof(f373,plain,(
  s(bool,'F_0') != s(bool,g(s(fun(sK0,sK1),sK5('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),bool,fun(sK0,sK1))))) | $spl35),
  inference(cnf_transformation,[],[f373_D])).
fof(f373_D,plain,(
  s(bool,'F_0') != s(bool,g(s(fun(sK0,sK1),sK5('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),bool,fun(sK0,sK1))))) <=> ~$spl35),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl35])])).
fof(f371,plain,(
  s(bool,'F_0') != s(bool,g(s(fun(sK0,sK1),sK5('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),bool,fun(sK0,sK1))))) | ($spl7 | $spl16)),
  inference(backward_demodulation,[],[f127,f116])).
fof(f370,plain,(
  s(bool,'T_0') = s(bool,f(s(fun(sK0,sK1),sK5('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),bool,fun(sK0,sK1))))) | $spl32),
  inference(cnf_transformation,[],[f370_D])).
fof(f370_D,plain,(
  s(bool,'T_0') = s(bool,f(s(fun(sK0,sK1),sK5('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),bool,fun(sK0,sK1))))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f368,plain,(
  s(bool,'T_0') = s(bool,f(s(fun(sK0,sK1),sK5('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),bool,fun(sK0,sK1))))) | $spl17),
  inference(trivial_inequality_removal,[],[f367])).
fof(f367,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(fun(sK0,sK1),sK5('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),bool,fun(sK0,sK1))))) | $spl17),
  inference(superposition,[],[f128,f71])).
fof(f366,plain,(
  s(bool,'T_0') = s(bool,f(s(fun(sK0,sK1),sK5('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),bool,fun(sK0,sK1))))) | $spl17),
  inference(unit_resulting_resolution,[],[f128,f71])).
fof(f146,plain,(
  s(bool,'T_0') != s(bool,f(s(fun(sK0,sK1),sK5('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),bool,fun(sK0,sK1))))) | $spl27),
  inference(cnf_transformation,[],[f146_D])).
fof(f146_D,plain,(
  s(bool,'T_0') != s(bool,f(s(fun(sK0,sK1),sK5('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),bool,fun(sK0,sK1))))) <=> ~$spl27),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl27])])).
fof(f287,plain,(
  s(bool,'T_0') != s(bool,f(s(fun(sK0,sK1),sK5('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),bool,fun(sK0,sK1))))) | $spl10),
  inference(global_subsumption,[],[f67,f121,f232])).
fof(f153,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl31),
  inference(cnf_transformation,[],[f153_D])).
fof(f153_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl31),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl31])])).
fof(f151,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ($spl10 | $spl29)),
  inference(forward_demodulation,[],[f150,f121])).
fof(f150,plain,(
  s(bool,'F_0') != s(bool,g(s(fun(sK0,sK1),sK5('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),bool,fun(sK0,sK1))))) | $spl29),
  inference(cnf_transformation,[],[f150_D])).
fof(f150_D,plain,(
  s(bool,'F_0') != s(bool,g(s(fun(sK0,sK1),sK5('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),bool,fun(sK0,sK1))))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f148,plain,(
  s(bool,'F_0') != s(bool,g(s(fun(sK0,sK1),sK5('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),bool,fun(sK0,sK1))))) | ($spl3 | $spl12)),
  inference(backward_demodulation,[],[f122,f112])).
fof(f147,plain,(
  s(bool,'T_0') = s(bool,f(s(fun(sK0,sK1),sK5('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),bool,fun(sK0,sK1))))) | $spl26),
  inference(cnf_transformation,[],[f147_D])).
fof(f147_D,plain,(
  s(bool,'T_0') = s(bool,f(s(fun(sK0,sK1),sK5('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),bool,fun(sK0,sK1))))) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f145,plain,(
  s(bool,'T_0') = s(bool,f(s(fun(sK0,sK1),sK5('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),bool,fun(sK0,sK1))))) | $spl13),
  inference(trivial_inequality_removal,[],[f144])).
fof(f144,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(fun(sK0,sK1),sK5('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),bool,fun(sK0,sK1))))) | $spl13),
  inference(superposition,[],[f123,f71])).
fof(f143,plain,(
  s(bool,'T_0') = s(bool,f(s(fun(sK0,sK1),sK5('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),bool,fun(sK0,sK1))))) | $spl13),
  inference(unit_resulting_resolution,[],[f123,f71])).
fof(f138,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3))))))) | $spl9),
  inference(superposition,[],[f118,f71])).
fof(f131,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3))))))) | $spl5),
  inference(superposition,[],[f114,f71])).
fof(f126,plain,(
  s(bool,'T_0') = s(bool,g(s(fun(sK0,sK1),sK5('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),bool,fun(sK0,sK1))))) | $spl14),
  inference(cnf_transformation,[],[f126_D])).
fof(f126_D,plain,(
  s(bool,'T_0') = s(bool,g(s(fun(sK0,sK1),sK5('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),bool,fun(sK0,sK1))))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f128,plain,(
  s(bool,'F_0') != s(bool,f(s(fun(sK0,sK1),sK5('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),bool,fun(sK0,sK1))))) | $spl17),
  inference(cnf_transformation,[],[f128_D])).
fof(f128_D,plain,(
  s(bool,'F_0') != s(bool,f(s(fun(sK0,sK1),sK5('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),bool,fun(sK0,sK1))))) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f124,plain,(
  s(bool,'F_0') != s(bool,f(s(fun(sK0,sK1),sK5('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),bool,fun(sK0,sK1))))) | s(bool,'T_0') = s(bool,g(s(fun(sK0,sK1),sK5('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),bool,fun(sK0,sK1))))) | $spl7),
  inference(superposition,[],[f116,f71])).
fof(f121,plain,(
  s(bool,'T_0') = s(bool,g(s(fun(sK0,sK1),sK5('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),bool,fun(sK0,sK1))))) | $spl10),
  inference(cnf_transformation,[],[f121_D])).
fof(f121_D,plain,(
  s(bool,'T_0') = s(bool,g(s(fun(sK0,sK1),sK5('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),bool,fun(sK0,sK1))))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f123,plain,(
  s(bool,'F_0') != s(bool,f(s(fun(sK0,sK1),sK5('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),bool,fun(sK0,sK1))))) | $spl13),
  inference(cnf_transformation,[],[f123_D])).
fof(f123_D,plain,(
  s(bool,'F_0') != s(bool,f(s(fun(sK0,sK1),sK5('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),bool,fun(sK0,sK1))))) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f119,plain,(
  s(bool,'F_0') != s(bool,f(s(fun(sK0,sK1),sK5('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),bool,fun(sK0,sK1))))) | s(bool,'T_0') = s(bool,g(s(fun(sK0,sK1),sK5('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),bool,fun(sK0,sK1))))) | $spl3),
  inference(superposition,[],[f112,f71])).
fof(f118,plain,(
  s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))))) != s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3))))))) | $spl9),
  inference(cnf_transformation,[],[f118_D])).
fof(f118_D,plain,(
  s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))))) != s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3))))))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f107,plain,(
  s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))))) != s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3))))))) | $spl1),
  inference(unit_resulting_resolution,[],[f106,f79])).
fof(f116,plain,(
  s(bool,f(s(fun(sK0,sK1),sK5('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),bool,fun(sK0,sK1))))) != s(bool,g(s(fun(sK0,sK1),sK5('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),bool,fun(sK0,sK1))))) | $spl7),
  inference(cnf_transformation,[],[f116_D])).
fof(f116_D,plain,(
  s(bool,f(s(fun(sK0,sK1),sK5('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),bool,fun(sK0,sK1))))) != s(bool,g(s(fun(sK0,sK1),sK5('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),bool,fun(sK0,sK1))))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f108,plain,(
  s(bool,f(s(fun(sK0,sK1),sK5('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),bool,fun(sK0,sK1))))) != s(bool,g(s(fun(sK0,sK1),sK5('const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),bool,fun(sK0,sK1))))) | $spl1),
  inference(unit_resulting_resolution,[],[f106,f82])).
fof(f114,plain,(
  s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))))) != s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3))))))) | $spl5),
  inference(cnf_transformation,[],[f114_D])).
fof(f114_D,plain,(
  s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))))) != s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3))))))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f109,plain,(
  s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))))) != s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK4('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),fun(sK0,sK1))),s(fun(fun(sK0,sK1),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3))))))) | $spl1),
  inference(unit_resulting_resolution,[],[f106,f79])).
fof(f112,plain,(
  s(bool,f(s(fun(sK0,sK1),sK5('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),bool,fun(sK0,sK1))))) != s(bool,g(s(fun(sK0,sK1),sK5('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),bool,fun(sK0,sK1))))) | $spl3),
  inference(cnf_transformation,[],[f112_D])).
fof(f112_D,plain,(
  s(bool,f(s(fun(sK0,sK1),sK5('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),bool,fun(sK0,sK1))))) != s(bool,g(s(fun(sK0,sK1),sK5('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),bool,fun(sK0,sK1))))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f110,plain,(
  s(bool,f(s(fun(sK0,sK1),sK5('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),bool,fun(sK0,sK1))))) != s(bool,g(s(fun(sK0,sK1),sK5('const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))),bool,fun(sK0,sK1))))) | $spl1),
  inference(unit_resulting_resolution,[],[f106,f82])).
fof(f106,plain,(
  s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))) != s(fun(fun(sK0,sK1),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3))))) | $spl1),
  inference(cnf_transformation,[],[f106_D])).
fof(f106_D,plain,(
  s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))) != s(fun(fun(sK0,sK1),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f67,plain,(
  s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))) != s(fun(fun(sK0,sK1),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))))),
  inference(cnf_transformation,[],[f48])).
fof(f48,plain,(
  s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),sK3))) != s(fun(fun(sK0,sK1),bool),'const.util_prob.DFUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,fun(sK1,bool)),'const.combin.K_1'(s(fun(sK1,bool),sK3)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f40])).
fof(f40,plain,(
  ? [X0,X1,X2,X3] : s(fun(fun(X0,X1),bool),'const.util_prob.FUNSET_2'(s(fun(X0,bool),X2),s(fun(X1,bool),X3))) != s(fun(fun(X0,X1),bool),'const.util_prob.DFUNSET_2'(s(fun(X0,bool),X2),s(fun(X0,fun(X1,bool)),'const.combin.K_1'(s(fun(X1,bool),X3)))))),
  inference(ennf_transformation,[],[f19])).
fof(f19,plain,(
  ~! [X0,X1,X2,X3] : s(fun(fun(X0,X1),bool),'const.util_prob.FUNSET_2'(s(fun(X0,bool),X2),s(fun(X1,bool),X3))) = s(fun(fun(X0,X1),bool),'const.util_prob.DFUNSET_2'(s(fun(X0,bool),X2),s(fun(X0,fun(X1,bool)),'const.combin.K_1'(s(fun(X1,bool),X3)))))),
  inference(flattening,[],[f18])).
fof(f18,plain,(
  ~! [X0,X1] : ! [X2,X3] : s(fun(fun(X0,X1),bool),'const.util_prob.FUNSET_2'(s(fun(X0,bool),X2),s(fun(X1,bool),X3))) = s(fun(fun(X0,X1),bool),'const.util_prob.DFUNSET_2'(s(fun(X0,bool),X2),s(fun(X0,fun(X1,bool)),'const.combin.K_1'(s(fun(X1,bool),X3)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9] : s(fun(fun(X6,X7),bool),'const.util_prob.FUNSET_2'(s(fun(X6,bool),X8),s(fun(X7,bool),X9))) = s(fun(fun(X6,X7),bool),'const.util_prob.DFUNSET_2'(s(fun(X6,bool),X8),s(fun(X6,fun(X7,bool)),'const.combin.K_1'(s(fun(X7,bool),X9)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9] : s(fun(fun(X6,X7),bool),'const.util_prob.FUNSET_2'(s(fun(X6,bool),X8),s(fun(X7,bool),X9))) = s(fun(fun(X6,X7),bool),'const.util_prob.DFUNSET_2'(s(fun(X6,bool),X8),s(fun(X6,fun(X7,bool)),'const.combin.K_1'(s(fun(X7,bool),X9)))))),
  file('thm.util_prob.FUNSET_DFUNSET','thm.util_prob.FUNSET_DFUNSET')).
% SZS output end Proof for thm.util_prob
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 137
% Passive clauses: 676
% Generated clauses: 9064
% Final active clauses: 69
% Final passive clauses: 505
% Input formulas: 16
% Initial clauses: 36
% 
% Duplicate literals: 10
% Trivial inequalities: 245
% Fw demodulations: 2876
% Bw demodulations: 22
% Global subsumptions: 101
% 
% Simple tautologies: 39
% Equational tautologies: 5
% Forward subsumptions: 4323
% Backward subsumptions: 54
% 
% Binary resolution: 335
% Unit resulting resolution: 3198
% Forward superposition: 1501
% Backward superposition: 553
% Equality resolution: 5
% 
% Split clauses: 96
% Split components: 223
% SAT solver clauses: 2098
% SAT solver unit clauses: 84
% SAT solver binary clauses: 717
% 
% Sat splits: 96
% Sat splitting refutations: 250
% 
% Memory used [KB]: 9850
% Time elapsed: 0.289 s
% ------------------------------
% ------------------------------
% Success in time 18.303 s
