Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.relation
% (139514)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 880
% Passive clauses: 8598
% Generated clauses: 20649
% Final active clauses: 870
% Final passive clauses: 7372
% Input formulas: 133
% Initial clauses: 408
% 
% Duplicate literals: 210
% Trivial inequalities: 25
% Fw subsumption resolutions: 1009
% Fw demodulations: 2781
% Bw demodulations: 9
% 
% Simple tautologies: 280
% Equational tautologies: 31
% Forward subsumptions: 7276
% 
% Binary resolution: 8165
% Factoring: 903
% Forward superposition: 2554
% Backward superposition: 4419
% Self superposition: 36
% Equality factoring: 2
% Equality resolution: 10
% 
% Split clauses: 225
% Split components: 452
% SAT solver clauses: 374
% SAT solver unit clauses: 9
% SAT solver binary clauses: 80
% 
% Sat splits: 225
% 
% TWLsolver clauses: 374
% TWLsolver calls for satisfiability: 331
% 
% Memory used [KB]: 14583
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 144 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.relation
% (139535)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 446
% Passive clauses: 3245
% Generated clauses: 12698
% Final active clauses: 429
% Final passive clauses: 2747
% Input formulas: 133
% Initial clauses: 408
% 
% Duplicate literals: 89
% Trivial inequalities: 5
% Fw subsumption resolutions: 378
% Fw demodulations: 2010
% Bw demodulations: 67
% Condensations: 72
% Global subsumptions: 23
% 
% Simple tautologies: 252
% Equational tautologies: 8
% Forward subsumptions: 3056
% 
% Binary resolution: 2830
% Factoring: 57
% Forward superposition: 1840
% Backward superposition: 1648
% Self superposition: 31
% Equality factoring: 2
% Equality resolution: 6
% 
% Split clauses: 24
% Split components: 48
% SAT solver clauses: 3955
% SAT solver unit clauses: 401
% SAT solver binary clauses: 109
% SAT solver learnt clauses: 46
% SAT solver learnt literals: 12
% 
% Sat splits: 24
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 3863
% TWLsolver calls for satisfiability: 3707
% 
% Memory used [KB]: 18166
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.relation
% (139564)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 2266
% Passive clauses: 17991
% Generated clauses: 91483
% Final active clauses: 1938
% Final passive clauses: 4462
% Input formulas: 133
% Initial clauses: 408
% Discarded non-redundant clauses: 8835
% 
% Duplicate literals: 341
% Trivial inequalities: 41
% Fw subsumption resolutions: 2679
% Fw demodulations: 17646
% Bw demodulations: 5117
% Condensations: 1417
% Global subsumptions: 613
% 
% Simple tautologies: 166
% Equational tautologies: 114
% Forward subsumptions: 48067
% Fw demodulations to eq. taut.: 14
% Bw demodulations to eq. taut.: 22
% 
% Unit resulting resolution: 31774
% Factoring: 1
% Forward superposition: 16171
% Backward superposition: 10675
% Self superposition: 93
% Equality factoring: 26
% Equality resolution: 9
% 
% Split clauses: 99
% Split components: 198
% SAT solver clauses: 20192
% SAT solver unit clauses: 4151
% SAT solver binary clauses: 7094
% 
% Sat splits: 99
% Sat splitting refutations: 1388
% 
% Memory used [KB]: 124475
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.relation
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.relation
% SZS output start Proof for thm.relation
fof(f65386,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f804,f1000,f805,f1002,f807,f1004,f808,f1006,f810,f998,f813,f829,f831,f832,f833,f836,f934,f1034,f1036,f1037,f1039,f1041,f1139,f1463,f1825,f2189,f2163,f2162,f2271,f2272,f2297,f2299,f2300,f2301,f2304,f2402,f2526,f2527,f2529,f2530,f2531,f2534,f2632,f3126,f3128,f3129,f3130,f3133,f3231,f4448,f4644,f4449,f4646,f4451,f4648,f4452,f4650,f4454,f4642,f4457,f4473,f4475,f4476,f4477,f4480,f4578,f6110,f6111,f6113,f6114,f6116,f6119,f6135,f6137,f6138,f6139,f6142,f6240,f6706,f7366,f7370,f7368,f7685,f1008,f1010,f7701,f7703,f7704,f7705,f7708,f7806,f1011,f7878,f7880,f7881,f7882,f7885,f7983,f2695,f8215,f8217,f8218,f8220,f8222,f8320,f1236,f8696,f8698,f8699,f8700,f8703,f8801,f3294,f9844,f10843,f10859,f10861,f10862,f10863,f10866,f10964,f11045,f11046,f11048,f11049,f11051,f11054,f11070,f11072,f11073,f11074,f11077,f11175,f12634,f12830,f12635,f12832,f12637,f12834,f12638,f12836,f12640,f12828,f12643,f12659,f12661,f12662,f12663,f12666,f12764,f14125,f14127,f14128,f14130,f14132,f14230,f14397,f14553,f14555,f14556,f14557,f14560,f14658,f14745,f14761,f14763,f14764,f14765,f14768,f14866,f14948,f15145,f14949,f15147,f14951,f15149,f14952,f15151,f14954,f15143,f14957,f14973,f14975,f14976,f14977,f14980,f15078,f15810,f15812,f15813,f15814,f15817,f15915,f16416,f16418,f16419,f16420,f16423,f16521,f10840,f10838,f6312,f10837,f6310,f10835,f6308,f10834,f6306,f6304,f17234,f17236,f17237,f17238,f17241,f17339,f17426,f6314,f6316,f17442,f17444,f17445,f17446,f17449,f17547,f6317,f17616,f17618,f17619,f17620,f17623,f17721,f15979,f18221,f18223,f18224,f18225,f18228,f18326,f16585,f18391,f18393,f18394,f18395,f18398,f18496,f18562,f18564,f18565,f18567,f18569,f18667,f18728,f18770,f18772,f18773,f18774,f18777,f18875,f18560,f19675,f19677,f19678,f19679,f19682,f19780,f20883,f20885,f20886,f20887,f20890,f20988,f21056,f11241,f21057,f11243,f21059,f11245,f21060,f11247,f21062,f11239,f21065,f21081,f21083,f21084,f21085,f21088,f21186,f21260,f21262,f21263,f21264,f21267,f21365,f21876,f21878,f21879,f21880,f21883,f21981,f22694,f22696,f22697,f22698,f22701,f22799,f23623,f23625,f23626,f23628,f23630,f23728,f23780,f23816,f23818,f23819,f23820,f23823,f23921,f24316,f24318,f24319,f24320,f24323,f24421,f24501,f24658,f24660,f24661,f24662,f24665,f24763,f24843,f24847,f24849,f24850,f24851,f24854,f24952,f25044,f25046,f25047,f25049,f25051,f25149,f25248,f25399,f25401,f25402,f25404,f25406,f25504,f25601,f26274,f26276,f26277,f26278,f26281,f26379,f27544,f27546,f27547,f27548,f27551,f27649,f28189,f28191,f28192,f28194,f28196,f28294,f28362,f28394,f28396,f28397,f28399,f28401,f28499,f28567,f28595,f28792,f28596,f28794,f28598,f28796,f28599,f28798,f28601,f28790,f28604,f28800,f28802,f28620,f28622,f28623,f28624,f28627,f28725,f28803,f29463,f29465,f29466,f29467,f29470,f29568,f29632,f29774,f29776,f29777,f29778,f29781,f29879,f29949,f32165,f32362,f32166,f32364,f32168,f32366,f32169,f32368,f32171,f32360,f32174,f32370,f32372,f32190,f32192,f32193,f32194,f32197,f32295,f32373,f36154,f36155,f36157,f36158,f36160,f36163,f36179,f36181,f36182,f36183,f36186,f36284,f36373,f36375,f36376,f36377,f36380,f36478,f17393,f36584,f4652,f4654,f36600,f36602,f36603,f36604,f36607,f36705,f4655,f36941,f15153,f15155,f36957,f36959,f36960,f36961,f36964,f37062,f15156,f37300,f37302,f37303,f37304,f37307,f37405,f27703,f38411,f38413,f38414,f38415,f38418,f38516,f38580,f39035,f39037,f39038,f39039,f39042,f39140,f39204,f39205,f39207,f39208,f39209,f39212,f39310,f39374,f39829,f39831,f39832,f39833,f39836,f39934,f39998,f41471,f41358,f41260,f41258,f41256,f41255,f41253,f41472,f41501,f41503,f41504,f41505,f41508,f41606,f21042,f41662,f41664,f41665,f41666,f41669,f41767,f14712,f41873,f41889,f41891,f41892,f41893,f41896,f41994,f42233,f42235,f42236,f42237,f42240,f42338,f42567,f42569,f42570,f42571,f42574,f42672,f25006,f43855,f43857,f43858,f43860,f43862,f43960,f44044,f44131,f44133,f44134,f44136,f44138,f44236,f44320,f44505,f44729,f44727,f44506,f44732,f44730,f44507,f44734,f44508,f44736,f44509,f44737,f44510,f44738,f44511,f44739,f44512,f44740,f44513,f44741,f44514,f44742,f44515,f44743,f44516,f44744,f44517,f44745,f44518,f44747,f44519,f44735,f44749,f44520,f44751,f44523,f44752,f44525,f44753,f44526,f44754,f44528,f44755,f44529,f44756,f44530,f44757,f44531,f44758,f44532,f44759,f44533,f44760,f44534,f44761,f44535,f44762,f44536,f44764,f44537,f44766,f44750,f44538,f44540,f44541,f44543,f44545,f44643,f44767,f44810,f44812,f44813,f44814,f44817,f44915,f44986,f45139,f45141,f45142,f45143,f45146,f45244,f45760,f45762,f45763,f45764,f45767,f45865,f45930,f45932,f45933,f45934,f45937,f46035,f46101,f46103,f46104,f46106,f46108,f46206,f46267,f46269,f46271,f46272,f46273,f46276,f46374,f46534,f46536,f46537,f46539,f46541,f46639,f46719,f46816,f46818,f46819,f46821,f46823,f46921,f47001,f47015,f47212,f47016,f47214,f47018,f47216,f47019,f47218,f47021,f47210,f47024,f47040,f47042,f47043,f47044,f47047,f47145,f47917,f47919,f47920,f47921,f47924,f48022,f48092,f48125,f11249,f11251,f48141,f48143,f48144,f48145,f48148,f48246,f11252,f48315,f48317,f48318,f48319,f48322,f48420,f21429,f48588,f48590,f48591,f48592,f48595,f48693,f42418,f49065,f49067,f49068,f49070,f49072,f49170,f49250,f49428,f49430,f49431,f49432,f49435,f49533,f22045,f49693,f49695,f49696,f49698,f49700,f49798,f49878,f49880,f49882,f49883,f49884,f49887,f49985,f46099,f50071,f47220,f47222,f50087,f50089,f50090,f50091,f50094,f50192,f47223,f50279,f50281,f50282,f50283,f50286,f50384,f50448,f50492,f50494,f50495,f50497,f50499,f50597,f50700,f50744,f50746,f50747,f50749,f50751,f50849,f50952,f50996,f50998,f50999,f51001,f51003,f51101,f51204,f51248,f51250,f51251,f51253,f51255,f51353,f51456,f52575,f52577,f52578,f52579,f52582,f52680,f52868,f56002,f56004,f56005,f56006,f56009,f56107,f56171,f59317,f36351,f59318,f36353,f59320,f36355,f59321,f36357,f59323,f36349,f59326,f36359,f36361,f59342,f59344,f59345,f59346,f59349,f59447,f36362,f59514,f59516,f59517,f59518,f59521,f59619,f59673,f59706,f12838,f12840,f59722,f59724,f59725,f59726,f59729,f59827,f12841,f59896,f59898,f59899,f59900,f59903,f60001,f60065,f60528,f60530,f60531,f60532,f60535,f60633,f60697,f60699,f60701,f60702,f60703,f60706,f60804,f60868,f60870,f60872,f60873,f60875,f60877,f60975,f61037,f62215,f62217,f62218,f62219,f62222,f62320,f62429,f63730,f63732,f63733,f63734,f63737,f63835,f63916,f64082,f64084,f64085,f64087,f64089,f64187,f64271,f64825,f64827,f64828,f1012,f64830,f1014,f64832,f1013,f64930,f65022,f1015,f65381])).
fof(f65381,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f65189,f1029])).
fof(f1029,plain,(
  ~p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17))))))),
  inference(resolution,[],[f500,f528])).
fof(f528,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f359])).
fof(f359,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | s(X0,X2) = s(X0,X3) | p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))) & ((s(X0,X2) != s(X0,X3) & ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))) | p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))))),
  inference(flattening,[],[f358])).
fof(f358,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | (s(X0,X2) = s(X0,X3) | p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))))) & ((s(X0,X2) != s(X0,X3) & ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))) | p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))))),
  inference(nnf_transformation,[],[f156])).
fof(f156,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) <=> (s(X0,X2) = s(X0,X3) | p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))))),
  inference(flattening,[],[f155])).
fof(f155,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) <=> (s(X0,X2) = s(X0,X3) | p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))))),
  inference(rectify,[],[f96])).
fof(f96,axiom,(
  ! [X6] : ! [X8,X9,X10] : (p(s(bool,'const.relation.RTC_3'(s(fun(X6,fun(X6,bool)),X8),s(X6,X9),s(X6,X10)))) <=> (s(X6,X9) = s(X6,X10) | p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X8),s(X6,X9),s(X6,X10))))))),
  file('thm.relation.RTC_lifts_monotonicities','thm.relation.RTC_CASES_TC')).
fof(f500,plain,(
  ~p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17))))))),
  inference(cnf_transformation,[],[f339])).
fof(f339,plain,(
  ! [X5,X6] : (~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X5))),s(sK13,X6)))) | p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,X5))))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,X6))))))) & (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)))) & ~p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13,sK14,sK15,sK16,sK17])],[f338])).
fof(f338,plain,(
  ? [X0,X1,X2] : (! [X5,X6] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X5))),s(X0,X6)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X0,X0),X1),s(X0,X5))))),s(X0,happ(s(fun(X0,X0),X1),s(X0,X6))))))) & ? [X3,X4] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X2),s(X0,X3),s(X0,X4)))) & ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X0,X0),X1),s(X0,X3))),s(X0,happ(s(fun(X0,X0),X1),s(X0,X4))))))))),
  inference(rectify,[],[f261])).
fof(f261,plain,(
  ? [X0,X1,X2] : (! [X3,X4] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X0,X0),X1),s(X0,X3))))),s(X0,happ(s(fun(X0,X0),X1),s(X0,X4))))))) & ? [X5,X6] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X2),s(X0,X5),s(X0,X6)))) & ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X0,X0),X1),s(X0,X5))),s(X0,happ(s(fun(X0,X0),X1),s(X0,X6))))))))),
  inference(ennf_transformation,[],[f136])).
fof(f136,plain,(
  ~! [X0,X1,X2] : (! [X3,X4] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X0,X0),X1),s(X0,X3))))),s(X0,happ(s(fun(X0,X0),X1),s(X0,X4))))))) => ! [X5,X6] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X2),s(X0,X5),s(X0,X6)))) => p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X0,X0),X1),s(X0,X5))),s(X0,happ(s(fun(X0,X0),X1),s(X0,X6))))))))),
  inference(flattening,[],[f135])).
fof(f135,plain,(
  ~! [X0] : ! [X1,X2] : (! [X3,X4] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X0,X0),X1),s(X0,X3))))),s(X0,happ(s(fun(X0,X0),X1),s(X0,X4))))))) => ! [X5,X6] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X2),s(X0,X5),s(X0,X6)))) => p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X0,X0),X1),s(X0,X5))),s(X0,happ(s(fun(X0,X0),X1),s(X0,X6))))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : (! [X9,X10] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,X9))),s(X6,X10)))) => p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,happ(s(fun(X6,X6),X7),s(X6,X9))))),s(X6,happ(s(fun(X6,X6),X7),s(X6,X10))))))) => ! [X9,X10] : (p(s(bool,'const.relation.RTC_3'(s(fun(X6,fun(X6,bool)),X8),s(X6,X9),s(X6,X10)))) => p(s(bool,'const.relation.RTC_3'(s(fun(X6,fun(X6,bool)),X8),s(X6,happ(s(fun(X6,X6),X7),s(X6,X9))),s(X6,happ(s(fun(X6,X6),X7),s(X6,X10))))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : (! [X9,X10] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,X9))),s(X6,X10)))) => p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,happ(s(fun(X6,X6),X7),s(X6,X9))))),s(X6,happ(s(fun(X6,X6),X7),s(X6,X10))))))) => ! [X9,X10] : (p(s(bool,'const.relation.RTC_3'(s(fun(X6,fun(X6,bool)),X8),s(X6,X9),s(X6,X10)))) => p(s(bool,'const.relation.RTC_3'(s(fun(X6,fun(X6,bool)),X8),s(X6,happ(s(fun(X6,X6),X7),s(X6,X9))),s(X6,happ(s(fun(X6,X6),X7),s(X6,X10))))))))),
  file('thm.relation.RTC_lifts_monotonicities','thm.relation.RTC_lifts_monotonicities')).
fof(f65189,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17)))))) | $spl10),
  inference(resolution,[],[f2525,f1008])).
fof(f2525,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X0),s(sK13,X1)))) | p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,X0))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,X1))))))) )),
  inference(subsumption_resolution,[],[f2273,f765])).
fof(f765,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X0,X0),X1),s(X0,X3))),s(X0,happ(s(fun(X0,X0),X1),s(X0,X4)))))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X2),s(X0,X3),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK75(X2,X1,X0)))),s(X0,sK76(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f490])).
fof(f490,plain,(
  ! [X0,X1,X2] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK75(X2,X1,X0)))),s(X0,sK76(X2,X1,X0))))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X0,X0),X1),s(X0,sK75(X2,X1,X0)))))),s(X0,happ(s(fun(X0,X0),X1),s(X0,sK76(X2,X1,X0)))))))) | ! [X3,X4] : (~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X2),s(X0,X3),s(X0,X4)))) | p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X0,X0),X1),s(X0,X3))),s(X0,happ(s(fun(X0,X0),X1),s(X0,X4))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK75,sK76])],[f489])).
fof(f489,plain,(
  ! [X0,X1,X2] : (? [X5,X6] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X5))),s(X0,X6)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X0,X0),X1),s(X0,X5))))),s(X0,happ(s(fun(X0,X0),X1),s(X0,X6))))))) | ! [X3,X4] : (~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X2),s(X0,X3),s(X0,X4)))) | p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X0,X0),X1),s(X0,X3))),s(X0,happ(s(fun(X0,X0),X1),s(X0,X4))))))))),
  inference(rectify,[],[f311])).
fof(f311,plain,(
  ! [X0,X1,X2] : (? [X3,X4] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X0,X0),X1),s(X0,X3))))),s(X0,happ(s(fun(X0,X0),X1),s(X0,X4))))))) | ! [X5,X6] : (~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X2),s(X0,X5),s(X0,X6)))) | p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X0,X0),X1),s(X0,X5))),s(X0,happ(s(fun(X0,X0),X1),s(X0,X6))))))))),
  inference(ennf_transformation,[],[f250])).
fof(f250,plain,(
  ! [X0,X1,X2] : (! [X3,X4] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X0,X0),X1),s(X0,X3))))),s(X0,happ(s(fun(X0,X0),X1),s(X0,X4))))))) => ! [X5,X6] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X2),s(X0,X5),s(X0,X6)))) => p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X0,X0),X1),s(X0,X5))),s(X0,happ(s(fun(X0,X0),X1),s(X0,X6))))))))),
  inference(flattening,[],[f249])).
fof(f249,plain,(
  ! [X0] : ! [X1,X2] : (! [X3,X4] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X0,X0),X1),s(X0,X3))))),s(X0,happ(s(fun(X0,X0),X1),s(X0,X4))))))) => ! [X5,X6] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X2),s(X0,X5),s(X0,X6)))) => p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X0,X0),X1),s(X0,X5))),s(X0,happ(s(fun(X0,X0),X1),s(X0,X6))))))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X6] : ! [X7,X8] : (! [X9,X10] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,X9))),s(X6,X10)))) => p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,happ(s(fun(X6,X6),X7),s(X6,X9))))),s(X6,happ(s(fun(X6,X6),X7),s(X6,X10))))))) => ! [X9,X10] : (p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X8),s(X6,X9),s(X6,X10)))) => p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X8),s(X6,happ(s(fun(X6,X6),X7),s(X6,X9))),s(X6,happ(s(fun(X6,X6),X7),s(X6,X10))))))))),
  file('thm.relation.RTC_lifts_monotonicities','thm.relation.TC_lifts_monotonicities')).
fof(f2273,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK75(sK15,sK14,sK13)))),s(sK13,sK76(sK15,sK14,sK13))))) | p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,X0))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,X1)))))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X0),s(sK13,X1))))) )),
  inference(resolution,[],[f498,f766])).
fof(f766,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X0,X0),X1),s(X0,X3))),s(X0,happ(s(fun(X0,X0),X1),s(X0,X4)))))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X2),s(X0,X3),s(X0,X4)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X0,X0),X1),s(X0,sK75(X2,X1,X0)))))),s(X0,happ(s(fun(X0,X0),X1),s(X0,sK76(X2,X1,X0)))))))) )),
  inference(cnf_transformation,[],[f490])).
fof(f498,plain,(
  ( ! [X6,X5] : (p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,X5))))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,X6)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X5))),s(sK13,X6))))) )),
  inference(cnf_transformation,[],[f339])).
fof(f1015,plain,(
  ( ! [X41] : (p(s(bool,X41))) ) | $spl22),
  inference(cnf_transformation,[],[f1015_D])).
fof(f1015_D,plain,(
  ( ! [X41] : (p(s(bool,X41))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f65022,plain,(
  ( ! [X34,X32] : (~p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X32),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17)))))) | ~p(s(bool,X34)) | s(bool,X34) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,X32)))) ) | $spl352),
  inference(cnf_transformation,[],[f65022_D])).
fof(f65022_D,plain,(
  ( ! [X34,X32] : (~p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X32),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17)))))) | ~p(s(bool,X34)) | s(bool,X34) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,X32)))) ) <=> ~$spl352),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl352])])).
fof(f64930,plain,(
  ( ! [X349,X347,X348] : (~p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X347),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17)))))) | p(s(bool,X348)) | ~p(s(bool,X349)) | s(bool,X349) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,X347)))) )),
  inference(resolution,[],[f1020,f783])).
fof(f783,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f555])).
fof(f555,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f369])).
fof(f369,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f368])).
fof(f368,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f288])).
fof(f288,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f287])).
fof(f287,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f181])).
fof(f181,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f120])).
fof(f120,axiom,(
  ! [X9,X45,X10,X46] : ((s(bool,X9) = s(bool,X45) & (p(s(bool,X45)) => s(bool,X10) = s(bool,X46))) => ((p(s(bool,X9)) => p(s(bool,X10))) <=> (p(s(bool,X45)) => p(s(bool,X46)))))),
  file('thm.relation.RTC_lifts_monotonicities','thm.bool.IMP_CONG')).
fof(f1020,plain,(
  ( ! [X4] : (~p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,X4)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X4),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17))))))) )),
  inference(resolution,[],[f500,f511])).
fof(f511,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X3)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4))))) )),
  inference(cnf_transformation,[],[f346])).
fof(f346,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,sK21(X3,X2,X1,X0))))) & p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,sK21(X3,X2,X1,X0)),s(X0,X3)))))) & (! [X4] : (~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X3))))) | p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK21])],[f345])).
fof(f345,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | ? [X5] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X5)))) & p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X5),s(X0,X3)))))) & (! [X4] : (~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X3))))) | p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))))),
  inference(rectify,[],[f344])).
fof(f344,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | ? [X4] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4)))) & p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X3)))))) & (! [X4] : (~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X3))))) | p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))))),
  inference(nnf_transformation,[],[f148])).
fof(f148,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) <=> ? [X4] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4)))) & p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X3))))))),
  inference(flattening,[],[f147])).
fof(f147,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) <=> ? [X4] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4)))) & p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X3))))))),
  inference(rectify,[],[f97])).
fof(f97,axiom,(
  ! [X6] : ! [X8,X9,X10] : (p(s(bool,'const.relation.RTC_3'(s(fun(X6,fun(X6,bool)),X8),s(X6,X9),s(X6,X10)))) <=> ? [X33] : (p(s(bool,'const.relation.RTC_3'(s(fun(X6,fun(X6,bool)),X8),s(X6,X9),s(X6,X33)))) & p(s(bool,'const.relation.RTC_3'(s(fun(X6,fun(X6,bool)),X8),s(X6,X33),s(X6,X10))))))),
  file('thm.relation.RTC_lifts_monotonicities','thm.relation.RTC_CASES_RTC_TWICE')).
fof(f1013,plain,(
  ( ! [X37,X36] : (s(bool,X36) != s(bool,X37) | ~p(s(bool,X36))) ) | $spl18),
  inference(cnf_transformation,[],[f1013_D])).
fof(f1013_D,plain,(
  ( ! [X37,X36] : (s(bool,X36) != s(bool,X37) | ~p(s(bool,X36))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f64832,plain,(
  ( ! [X61,X62,X60,X63] : (~p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X60),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17)))))) | ~p(s(bool,X61)) | ~p(s(bool,X62)) | s(bool,X62) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,X60))) | s(bool,X61) != s(bool,X63)) )),
  inference(resolution,[],[f1020,f558])).
fof(f558,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f369])).
fof(f1014,plain,(
  ( ! [X42] : (~p(s(bool,X42))) ) | $spl20),
  inference(cnf_transformation,[],[f1014_D])).
fof(f1014_D,plain,(
  ( ! [X42] : (~p(s(bool,X42))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f64830,plain,(
  ( ! [X54,X52,X55,X53] : (~p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X52),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17)))))) | p(s(bool,X53)) | ~p(s(bool,X54)) | ~p(s(bool,X55)) | s(bool,X54) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,X52)))) )),
  inference(resolution,[],[f1020,f557])).
fof(f557,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f369])).
fof(f1012,plain,(
  ( ! [X31,X29] : (s(bool,X29) != s(bool,X31) | p(s(bool,X29))) ) | $spl16),
  inference(cnf_transformation,[],[f1012_D])).
fof(f1012_D,plain,(
  ( ! [X31,X29] : (s(bool,X29) != s(bool,X31) | p(s(bool,X29))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f64828,plain,(
  ( ! [X47,X45,X46,X44] : (~p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X44),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17)))))) | p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X45) != s(bool,X47) | s(bool,X46) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,X44)))) )),
  inference(resolution,[],[f1020,f556])).
fof(f556,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f369])).
fof(f64827,plain,(
  ( ! [X43,X41,X42,X40] : (~p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X40),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17)))))) | ~p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X42) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,X40))) | s(bool,X41) != s(bool,X43)) )),
  inference(resolution,[],[f1020,f554])).
fof(f554,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f369])).
fof(f64825,plain,(
  ( ! [X35,X33,X34,X32] : (~p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X32),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17)))))) | p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X33) != s(bool,X35) | s(bool,X34) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,X32)))) )),
  inference(resolution,[],[f1020,f552])).
fof(f552,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f369])).
fof(f64271,plain,(
  ( ! [X87,X85] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),X85),s(sK13,sK16),s(sK13,sK17)))) | s(bool,X87) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),X85),s(sK13,sK18(X85,sK15,sK16,sK17,sK13)))),s(sK13,sK19(X85,sK15,sK16,sK17,sK13)))) | ~p(s(bool,X87))) ) | $spl350),
  inference(cnf_transformation,[],[f64271_D])).
fof(f64271_D,plain,(
  ( ! [X87,X85] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),X85),s(sK13,sK16),s(sK13,sK17)))) | s(bool,X87) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),X85),s(sK13,sK18(X85,sK15,sK16,sK17,sK13)))),s(sK13,sK19(X85,sK15,sK16,sK17,sK13)))) | ~p(s(bool,X87))) ) <=> ~$spl350),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl350])])).
fof(f64187,plain,(
  ( ! [X401,X400,X402] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),X400),s(sK13,sK16),s(sK13,sK17)))) | p(s(bool,X401)) | ~p(s(bool,X402)) | s(bool,X402) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),X400),s(sK13,sK18(X400,sK15,sK16,sK17,sK13)))),s(sK13,sK19(X400,sK15,sK16,sK17,sK13))))) )),
  inference(resolution,[],[f792,f783])).
fof(f792,plain,(
  ( ! [X1] : (~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),X1),s(sK13,sK18(X1,sK15,sK16,sK17,sK13)))),s(sK13,sK19(X1,sK15,sK16,sK17,sK13))))) | p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),X1),s(sK13,sK16),s(sK13,sK17))))) )),
  inference(resolution,[],[f499,f502])).
fof(f502,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X4),s(X0,X2),s(X0,X1)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X3),s(X0,X2),s(X0,X1)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,sK18(X4,X3,X2,X1,X0)))),s(X0,sK19(X4,X3,X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f340])).
fof(f340,plain,(
  ! [X0,X1,X2,X3,X4] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,sK18(X4,X3,X2,X1,X0)))),s(X0,sK19(X4,X3,X2,X1,X0))))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,sK18(X4,X3,X2,X1,X0)))),s(X0,sK19(X4,X3,X2,X1,X0)))))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X3),s(X0,X2),s(X0,X1)))) | p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X4),s(X0,X2),s(X0,X1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK18,sK19])],[f263])).
fof(f263,plain,(
  ! [X0,X1,X2,X3,X4] : (? [X5,X6] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,X5))),s(X0,X6)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X5))),s(X0,X6))))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X3),s(X0,X2),s(X0,X1)))) | p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X4),s(X0,X2),s(X0,X1)))))),
  inference(flattening,[],[f262])).
fof(f262,plain,(
  ! [X0,X1,X2,X3,X4] : (? [X5,X6] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,X5))),s(X0,X6)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X5))),s(X0,X6))))) | (~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X3),s(X0,X2),s(X0,X1)))) | p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X4),s(X0,X2),s(X0,X1))))))),
  inference(ennf_transformation,[],[f138])).
fof(f138,plain,(
  ! [X0,X1,X2,X3,X4] : (! [X5,X6] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,X5))),s(X0,X6)))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X5))),s(X0,X6))))) => (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X3),s(X0,X2),s(X0,X1)))) => p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X4),s(X0,X2),s(X0,X1))))))),
  inference(flattening,[],[f137])).
fof(f137,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : (! [X5,X6] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,X5))),s(X0,X6)))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X5))),s(X0,X6))))) => (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X3),s(X0,X2),s(X0,X1)))) => p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X4),s(X0,X2),s(X0,X1))))))),
  inference(rectify,[],[f90])).
fof(f90,axiom,(
  ! [X6] : ! [X10,X9,X8,X12] : (! [X42,X43] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,X42))),s(X6,X43)))) => p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X12),s(X6,X42))),s(X6,X43))))) => (p(s(bool,'const.relation.RTC_3'(s(fun(X6,fun(X6,bool)),X8),s(X6,X9),s(X6,X10)))) => p(s(bool,'const.relation.RTC_3'(s(fun(X6,fun(X6,bool)),X12),s(X6,X9),s(X6,X10))))))),
  file('thm.relation.RTC_lifts_monotonicities','thm.relation.RTC_MONOTONE')).
fof(f499,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17))))),
  inference(cnf_transformation,[],[f339])).
fof(f64089,plain,(
  ( ! [X116,X114,X115,X113] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),X113),s(sK13,sK16),s(sK13,sK17)))) | ~p(s(bool,X114)) | ~p(s(bool,X115)) | s(bool,X115) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),X113),s(sK13,sK18(X113,sK15,sK16,sK17,sK13)))),s(sK13,sK19(X113,sK15,sK16,sK17,sK13)))) | s(bool,X114) != s(bool,X116)) )),
  inference(resolution,[],[f792,f558])).
fof(f64087,plain,(
  ( ! [X107,X105,X108,X106] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),X105),s(sK13,sK16),s(sK13,sK17)))) | p(s(bool,X106)) | ~p(s(bool,X107)) | ~p(s(bool,X108)) | s(bool,X107) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),X105),s(sK13,sK18(X105,sK15,sK16,sK17,sK13)))),s(sK13,sK19(X105,sK15,sK16,sK17,sK13))))) )),
  inference(resolution,[],[f792,f557])).
fof(f64085,plain,(
  ( ! [X99,X97,X100,X98] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),X97),s(sK13,sK16),s(sK13,sK17)))) | p(s(bool,X98)) | ~p(s(bool,X99)) | s(bool,X98) != s(bool,X100) | s(bool,X99) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),X97),s(sK13,sK18(X97,sK15,sK16,sK17,sK13)))),s(sK13,sK19(X97,sK15,sK16,sK17,sK13))))) )),
  inference(resolution,[],[f792,f556])).
fof(f64084,plain,(
  ( ! [X94,X95,X93,X96] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),X93),s(sK13,sK16),s(sK13,sK17)))) | ~p(s(bool,X94)) | ~p(s(bool,X95)) | s(bool,X95) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),X93),s(sK13,sK18(X93,sK15,sK16,sK17,sK13)))),s(sK13,sK19(X93,sK15,sK16,sK17,sK13)))) | s(bool,X94) != s(bool,X96)) )),
  inference(resolution,[],[f792,f554])).
fof(f64082,plain,(
  ( ! [X88,X87,X85,X86] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),X85),s(sK13,sK16),s(sK13,sK17)))) | p(s(bool,X86)) | ~p(s(bool,X87)) | s(bool,X86) != s(bool,X88) | s(bool,X87) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),X85),s(sK13,sK18(X85,sK15,sK16,sK17,sK13)))),s(sK13,sK19(X85,sK15,sK16,sK17,sK13))))) )),
  inference(resolution,[],[f792,f552])).
fof(f63916,plain,(
  ( ! [X50,X48] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),X48),s(sK13,sK16),s(sK13,sK17)))) | s(bool,X50) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK18(X48,sK15,sK16,sK17,sK13)))),s(sK13,sK19(X48,sK15,sK16,sK17,sK13)))) | p(s(bool,X50))) ) | $spl348),
  inference(cnf_transformation,[],[f63916_D])).
fof(f63916_D,plain,(
  ( ! [X50,X48] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),X48),s(sK13,sK16),s(sK13,sK17)))) | s(bool,X50) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK18(X48,sK15,sK16,sK17,sK13)))),s(sK13,sK19(X48,sK15,sK16,sK17,sK13)))) | p(s(bool,X50))) ) <=> ~$spl348),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl348])])).
fof(f63835,plain,(
  ( ! [X366,X365,X367] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),X365),s(sK13,sK16),s(sK13,sK17)))) | p(s(bool,X366)) | p(s(bool,X367)) | s(bool,X367) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK18(X365,sK15,sK16,sK17,sK13)))),s(sK13,sK19(X365,sK15,sK16,sK17,sK13))))) )),
  inference(resolution,[],[f791,f783])).
fof(f791,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK18(X0,sK15,sK16,sK17,sK13)))),s(sK13,sK19(X0,sK15,sK16,sK17,sK13))))) | p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),X0),s(sK13,sK16),s(sK13,sK17))))) )),
  inference(resolution,[],[f499,f501])).
fof(f501,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X4),s(X0,X2),s(X0,X1)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X3),s(X0,X2),s(X0,X1)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,sK18(X4,X3,X2,X1,X0)))),s(X0,sK19(X4,X3,X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f340])).
fof(f63737,plain,(
  ( ! [X78,X76,X79,X77] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),X76),s(sK13,sK16),s(sK13,sK17)))) | p(s(bool,X77)) | ~p(s(bool,X78)) | s(bool,X77) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK18(X76,sK15,sK16,sK17,sK13)))),s(sK13,sK19(X76,sK15,sK16,sK17,sK13)))) | s(bool,X78) != s(bool,X79)) )),
  inference(resolution,[],[f791,f558])).
fof(f63734,plain,(
  ( ! [X66,X64,X67,X65] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),X64),s(sK13,sK16),s(sK13,sK17)))) | p(s(bool,X65)) | ~p(s(bool,X66)) | p(s(bool,X67)) | s(bool,X67) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK18(X64,sK15,sK16,sK17,sK13)))),s(sK13,sK19(X64,sK15,sK16,sK17,sK13))))) )),
  inference(resolution,[],[f791,f557])).
fof(f63733,plain,(
  ( ! [X61,X62,X60,X63] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),X60),s(sK13,sK16),s(sK13,sK17)))) | p(s(bool,X61)) | p(s(bool,X62)) | s(bool,X61) != s(bool,X63) | s(bool,X62) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK18(X60,sK15,sK16,sK17,sK13)))),s(sK13,sK19(X60,sK15,sK16,sK17,sK13))))) )),
  inference(resolution,[],[f791,f556])).
fof(f63732,plain,(
  ( ! [X59,X57,X58,X56] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),X56),s(sK13,sK16),s(sK13,sK17)))) | p(s(bool,X57)) | ~p(s(bool,X58)) | s(bool,X57) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK18(X56,sK15,sK16,sK17,sK13)))),s(sK13,sK19(X56,sK15,sK16,sK17,sK13)))) | s(bool,X58) != s(bool,X59)) )),
  inference(resolution,[],[f791,f554])).
fof(f63730,plain,(
  ( ! [X50,X48,X51,X49] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),X48),s(sK13,sK16),s(sK13,sK17)))) | p(s(bool,X49)) | p(s(bool,X50)) | s(bool,X49) != s(bool,X51) | s(bool,X50) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK18(X48,sK15,sK16,sK17,sK13)))),s(sK13,sK19(X48,sK15,sK16,sK17,sK13))))) )),
  inference(resolution,[],[f791,f552])).
fof(f62429,plain,(
  ( ! [X103,X105] : (p(s(bool,X103)) | s(bool,X103) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK26(happ(s(fun(sK13,sK13),sK14),s(sK13,sK17)),happ(s(fun(sK13,sK13),sK14),s(sK13,sK16)),sK15,sK13)),s(sK13,X105))),s(sK13,X105)))) ) | $spl346),
  inference(cnf_transformation,[],[f62429_D])).
fof(f62429_D,plain,(
  ( ! [X103,X105] : (p(s(bool,X103)) | s(bool,X103) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK26(happ(s(fun(sK13,sK13),sK14),s(sK13,sK17)),happ(s(fun(sK13,sK13),sK14),s(sK13,sK16)),sK15,sK13)),s(sK13,X105))),s(sK13,X105)))) ) <=> ~$spl346),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl346])])).
fof(f62320,plain,(
  ( ! [X420,X418,X419] : (p(s(bool,X418)) | p(s(bool,X419)) | s(bool,X419) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK26(happ(s(fun(sK13,sK13),sK14),s(sK13,sK17)),happ(s(fun(sK13,sK13),sK14),s(sK13,sK16)),sK15,sK13)),s(sK13,X420))),s(sK13,X420)))) )),
  inference(resolution,[],[f1021,f783])).
fof(f1021,plain,(
  ( ! [X5] : (p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK26(happ(s(fun(sK13,sK13),sK14),s(sK13,sK17)),happ(s(fun(sK13,sK13),sK14),s(sK13,sK16)),sK15,sK13)),s(sK13,X5))),s(sK13,X5))))) )),
  inference(resolution,[],[f500,f515])).
fof(f515,plain,(
  ( ! [X2,X0,X8,X3,X1] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),sK26(X3,X2,X1,X0)),s(X0,X8))),s(X0,X8))))) )),
  inference(cnf_transformation,[],[f349])).
fof(f349,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | ! [X9] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X9),s(X0,sK22(X9,X3,X2,X1,X0)))),s(X0,sK22(X9,X3,X2,X1,X0))))) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK23(X9,X3,X2,X1,X0)))),s(X0,sK24(X9,X3,X2,X1,X0))))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X9),s(X0,sK24(X9,X3,X2,X1,X0)))),s(X0,sK25(X9,X3,X2,X1,X0))))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X9),s(X0,sK23(X9,X3,X2,X1,X0)))),s(X0,sK25(X9,X3,X2,X1,X0)))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X9),s(X0,X2))),s(X0,X3)))))) & ((! [X8] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),sK26(X3,X2,X1,X0)),s(X0,X8))),s(X0,X8)))) & ! [X5,X6,X7] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X6)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),sK26(X3,X2,X1,X0)),s(X0,X6))),s(X0,X7)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),sK26(X3,X2,X1,X0)),s(X0,X5))),s(X0,X7))))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),sK26(X3,X2,X1,X0)),s(X0,X2))),s(X0,X3))))) | p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK22,sK23,sK24,sK25,sK26])],[f348])).
fof(f348,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | ! [X9] : (? [X13] : ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X9),s(X0,X13))),s(X0,X13)))) | ? [X10,X11,X12] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X10))),s(X0,X11)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X9),s(X0,X11))),s(X0,X12)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X9),s(X0,X10))),s(X0,X12))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X9),s(X0,X2))),s(X0,X3)))))) & (? [X4] : (! [X8] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X8))),s(X0,X8)))) & ! [X5,X6,X7] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X6)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X6))),s(X0,X7)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X5))),s(X0,X7))))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X2))),s(X0,X3))))) | p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))))),
  inference(rectify,[],[f347])).
fof(f347,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | ! [X4] : (? [X8] : ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X8))),s(X0,X8)))) | ? [X5,X6,X7] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X6)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X6))),s(X0,X7)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X5))),s(X0,X7))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X2))),s(X0,X3)))))) & (? [X4] : (! [X8] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X8))),s(X0,X8)))) & ! [X5,X6,X7] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X6)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X6))),s(X0,X7)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X5))),s(X0,X7))))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X2))),s(X0,X3))))) | p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))))),
  inference(nnf_transformation,[],[f271])).
fof(f271,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) <=> ! [X4] : (? [X8] : ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X8))),s(X0,X8)))) | ? [X5,X6,X7] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X6)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X6))),s(X0,X7)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X5))),s(X0,X7))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X2))),s(X0,X3))))))),
  inference(flattening,[],[f270])).
fof(f270,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) <=> ! [X4] : ((? [X8] : ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X8))),s(X0,X8)))) | ? [X5,X6,X7] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X6)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X6))),s(X0,X7))))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X5))),s(X0,X7)))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X2))),s(X0,X3))))))),
  inference(ennf_transformation,[],[f150])).
fof(f150,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) <=> ! [X4] : ((! [X8] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X8))),s(X0,X8)))) & ! [X5,X6,X7] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X6)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X6))),s(X0,X7))))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X5))),s(X0,X7)))))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X2))),s(X0,X3))))))),
  inference(flattening,[],[f149])).
fof(f149,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) <=> ! [X4] : ((! [X8] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X8))),s(X0,X8)))) & ! [X5,X6,X7] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X6)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X6))),s(X0,X7))))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X5))),s(X0,X7)))))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X2))),s(X0,X3))))))),
  inference(rectify,[],[f72])).
fof(f72,axiom,(
  ! [X6] : ! [X8,X40,X41] : (p(s(bool,'const.relation.RTC_3'(s(fun(X6,fun(X6,bool)),X8),s(X6,X40),s(X6,X41)))) <=> ! [X16] : ((! [X9] : p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X16),s(X6,X9))),s(X6,X9)))) & ! [X9,X10,X32] : ((p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,X9))),s(X6,X10)))) & p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X16),s(X6,X10))),s(X6,X32))))) => p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X16),s(X6,X9))),s(X6,X32)))))) => p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X16),s(X6,X40))),s(X6,X41))))))),
  file('thm.relation.RTC_lifts_monotonicities','thm.relation.RTC_DEF')).
fof(f62222,plain,(
  ( ! [X132,X130,X131,X129] : (p(s(bool,X129)) | ~p(s(bool,X130)) | s(bool,X129) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK26(happ(s(fun(sK13,sK13),sK14),s(sK13,sK17)),happ(s(fun(sK13,sK13),sK14),s(sK13,sK16)),sK15,sK13)),s(sK13,X131))),s(sK13,X131))) | s(bool,X130) != s(bool,X132)) )),
  inference(resolution,[],[f1021,f558])).
fof(f62219,plain,(
  ( ! [X121,X118,X120,X119] : (p(s(bool,X118)) | ~p(s(bool,X119)) | p(s(bool,X120)) | s(bool,X120) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK26(happ(s(fun(sK13,sK13),sK14),s(sK13,sK17)),happ(s(fun(sK13,sK13),sK14),s(sK13,sK16)),sK15,sK13)),s(sK13,X121))),s(sK13,X121)))) )),
  inference(resolution,[],[f1021,f557])).
fof(f62218,plain,(
  ( ! [X116,X114,X117,X115] : (p(s(bool,X114)) | p(s(bool,X115)) | s(bool,X114) != s(bool,X116) | s(bool,X115) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK26(happ(s(fun(sK13,sK13),sK14),s(sK13,sK17)),happ(s(fun(sK13,sK13),sK14),s(sK13,sK16)),sK15,sK13)),s(sK13,X117))),s(sK13,X117)))) )),
  inference(resolution,[],[f1021,f556])).
fof(f62217,plain,(
  ( ! [X111,X112,X110,X113] : (p(s(bool,X110)) | ~p(s(bool,X111)) | s(bool,X110) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK26(happ(s(fun(sK13,sK13),sK14),s(sK13,sK17)),happ(s(fun(sK13,sK13),sK14),s(sK13,sK16)),sK15,sK13)),s(sK13,X112))),s(sK13,X112))) | s(bool,X111) != s(bool,X113)) )),
  inference(resolution,[],[f1021,f554])).
fof(f62215,plain,(
  ( ! [X103,X105,X102,X104] : (p(s(bool,X102)) | p(s(bool,X103)) | s(bool,X102) != s(bool,X104) | s(bool,X103) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK26(happ(s(fun(sK13,sK13),sK14),s(sK13,sK17)),happ(s(fun(sK13,sK13),sK14),s(sK13,sK16)),sK15,sK13)),s(sK13,X105))),s(sK13,X105)))) )),
  inference(resolution,[],[f1021,f552])).
fof(f61037,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl344),
  inference(cnf_transformation,[],[f61037_D])).
fof(f61037_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) <=> ~$spl344),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl344])])).
fof(f60975,plain,(
  ( ! [X210,X211] : (p(s(bool,X210)) | ~p(s(bool,X211)) | s(bool,X211) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl97),
  inference(resolution,[],[f12839,f783])).
fof(f12839,plain,(
  ~p(s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) | $spl97),
  inference(cnf_transformation,[],[f12839_D])).
fof(f12839_D,plain,(
  ~p(s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) <=> ~$spl97),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl97])])).
fof(f60877,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK16,sK15,sK13)),s(sK13,sK17))) | s(bool,X21) != s(bool,X23)) ) | $spl97),
  inference(resolution,[],[f12839,f558])).
fof(f60875,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl97),
  inference(resolution,[],[f12839,f557])).
fof(f60873,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl97),
  inference(resolution,[],[f12839,f556])).
fof(f60872,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK16,sK15,sK13)),s(sK13,sK17))) | s(bool,X6) != s(bool,X8)) ) | $spl97),
  inference(resolution,[],[f12839,f554])).
fof(f60870,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl97),
  inference(resolution,[],[f12839,f552])).
fof(f60868,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl342),
  inference(cnf_transformation,[],[f60868_D])).
fof(f60868_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) <=> ~$spl342),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl342])])).
fof(f60804,plain,(
  ( ! [X212,X213] : (p(s(bool,X212)) | p(s(bool,X213)) | s(bool,X213) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl96),
  inference(resolution,[],[f12840,f783])).
fof(f60706,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK16,sK15,sK13)),s(sK13,sK17))) | s(bool,X22) != s(bool,X23)) ) | $spl96),
  inference(resolution,[],[f12840,f558])).
fof(f60703,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl96),
  inference(resolution,[],[f12840,f557])).
fof(f60702,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl96),
  inference(resolution,[],[f12840,f556])).
fof(f60701,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK16,sK15,sK13)),s(sK13,sK17))) | s(bool,X7) != s(bool,X8)) ) | $spl96),
  inference(resolution,[],[f12840,f554])).
fof(f60699,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl96),
  inference(resolution,[],[f12840,f552])).
fof(f60697,plain,(
  ( ! [X15] : (p(s(bool,X15)) | s(bool,X15) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl340),
  inference(cnf_transformation,[],[f60697_D])).
fof(f60697_D,plain,(
  ( ! [X15] : (p(s(bool,X15)) | s(bool,X15) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) <=> ~$spl340),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl340])])).
fof(f60633,plain,(
  ( ! [X227,X226] : (p(s(bool,X226)) | p(s(bool,X227)) | s(bool,X227) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl94),
  inference(resolution,[],[f12838,f783])).
fof(f60535,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK16,sK15,sK13)),s(sK13,sK17))) | s(bool,X36) != s(bool,X37)) ) | $spl94),
  inference(resolution,[],[f12838,f558])).
fof(f60532,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | p(s(bool,X28)) | s(bool,X28) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl94),
  inference(resolution,[],[f12838,f557])).
fof(f60531,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X23) != s(bool,X25) | s(bool,X24) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl94),
  inference(resolution,[],[f12838,f556])).
fof(f60530,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X20) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK16,sK15,sK13)),s(sK13,sK17))) | s(bool,X21) != s(bool,X22)) ) | $spl94),
  inference(resolution,[],[f12838,f554])).
fof(f60528,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl94),
  inference(resolution,[],[f12838,f552])).
fof(f60065,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl338),
  inference(cnf_transformation,[],[f60065_D])).
fof(f60065_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) <=> ~$spl338),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl338])])).
fof(f60001,plain,(
  ( ! [X212,X213] : (p(s(bool,X212)) | p(s(bool,X213)) | s(bool,X213) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl10),
  inference(resolution,[],[f12644,f783])).
fof(f12644,plain,(
  p(s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) | $spl10),
  inference(resolution,[],[f3109,f534])).
fof(f534,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.relation.EQC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f278])).
fof(f278,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | p(s(bool,'const.relation.EQC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))))),
  inference(ennf_transformation,[],[f168])).
fof(f168,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) => p(s(bool,'const.relation.EQC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))))),
  inference(flattening,[],[f167])).
fof(f167,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) => p(s(bool,'const.relation.EQC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))))),
  inference(rectify,[],[f98])).
fof(f98,axiom,(
  ! [X6] : ! [X8,X9,X10] : (p(s(bool,'const.relation.RTC_3'(s(fun(X6,fun(X6,bool)),X8),s(X6,X9),s(X6,X10)))) => p(s(bool,'const.relation.EQC_3'(s(fun(X6,fun(X6,bool)),X8),s(X6,X9),s(X6,X10)))))),
  file('thm.relation.RTC_lifts_monotonicities','thm.relation.RTC_EQC')).
fof(f3109,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) | $spl10),
  inference(resolution,[],[f1008,f507])).
fof(f507,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,sK20(X3,X2,X1,X0)),s(X0,X3)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f343])).
fof(f343,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,sK20(X3,X2,X1,X0))))) & p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,sK20(X3,X2,X1,X0)),s(X0,X3)))))) & (! [X4] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X4)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X3))))) | p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK20])],[f342])).
fof(f342,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | ? [X5] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X5)))) & p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X5),s(X0,X3)))))) & (! [X4] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X4)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X3))))) | p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))))),
  inference(rectify,[],[f341])).
fof(f341,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | ? [X4] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X4)))) & p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X3)))))) & (! [X4] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X4)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X3))))) | p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))))),
  inference(nnf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) <=> ? [X4] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X4)))) & p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X3))))))),
  inference(flattening,[],[f145])).
fof(f145,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) <=> ? [X4] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X4)))) & p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X3))))))),
  inference(rectify,[],[f92])).
fof(f92,axiom,(
  ! [X6] : ! [X8,X9,X32] : (p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X8),s(X6,X9),s(X6,X32)))) <=> ? [X10] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,X9))),s(X6,X10)))) & p(s(bool,'const.relation.RTC_3'(s(fun(X6,fun(X6,bool)),X8),s(X6,X10),s(X6,X32))))))),
  file('thm.relation.RTC_lifts_monotonicities','thm.relation.EXTEND_RTC_TC_EQN')).
fof(f59903,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK16,sK15,sK13)),s(sK13,sK17))) | s(bool,X22) != s(bool,X23)) ) | $spl10),
  inference(resolution,[],[f12644,f558])).
fof(f59900,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl10),
  inference(resolution,[],[f12644,f557])).
fof(f59899,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl10),
  inference(resolution,[],[f12644,f556])).
fof(f59898,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK16,sK15,sK13)),s(sK13,sK17))) | s(bool,X7) != s(bool,X8)) ) | $spl10),
  inference(resolution,[],[f12644,f554])).
fof(f59896,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl10),
  inference(resolution,[],[f12644,f552])).
fof(f12841,plain,(
  ( ! [X30] : (p(s(bool,X30)) | s(bool,X30) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl98),
  inference(cnf_transformation,[],[f12841_D])).
fof(f12841_D,plain,(
  ( ! [X30] : (p(s(bool,X30)) | s(bool,X30) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f59827,plain,(
  ( ! [X241,X242] : (p(s(bool,X241)) | p(s(bool,X242)) | s(bool,X242) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl10),
  inference(resolution,[],[f3109,f783])).
fof(f59729,plain,(
  ( ! [X52,X50,X51] : (p(s(bool,X50)) | ~p(s(bool,X51)) | s(bool,X50) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK16,sK15,sK13)),s(sK13,sK17))) | s(bool,X51) != s(bool,X52)) ) | $spl10),
  inference(resolution,[],[f3109,f558])).
fof(f59726,plain,(
  ( ! [X43,X41,X42] : (p(s(bool,X41)) | ~p(s(bool,X42)) | p(s(bool,X43)) | s(bool,X43) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl10),
  inference(resolution,[],[f3109,f557])).
fof(f59725,plain,(
  ( ! [X39,X38,X40] : (p(s(bool,X38)) | p(s(bool,X39)) | s(bool,X38) != s(bool,X40) | s(bool,X39) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl10),
  inference(resolution,[],[f3109,f556])).
fof(f59724,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK16,sK15,sK13)),s(sK13,sK17))) | s(bool,X36) != s(bool,X37)) ) | $spl10),
  inference(resolution,[],[f3109,f554])).
fof(f59722,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl10),
  inference(resolution,[],[f3109,f552])).
fof(f12840,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) | $spl96),
  inference(cnf_transformation,[],[f12840_D])).
fof(f12840_D,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f12838,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) | $spl94),
  inference(cnf_transformation,[],[f12838_D])).
fof(f12838_D,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f59706,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) | p(s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) | $spl10),
  inference(resolution,[],[f3109,f531])).
fof(f531,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f275])).
fof(f275,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))))),
  inference(flattening,[],[f274])).
fof(f274,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | (p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))))),
  inference(ennf_transformation,[],[f162])).
fof(f162,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) => (p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))))),
  inference(flattening,[],[f161])).
fof(f161,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) => (p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))))),
  inference(rectify,[],[f102])).
fof(f102,axiom,(
  ! [X6] : ! [X8,X9,X10] : (p(s(bool,'const.relation.RTC_3'(s(fun(X6,fun(X6,bool)),X8),s(X6,X9),s(X6,X10)))) => (p(s(bool,'const.relation.RC_3'(s(fun(X6,fun(X6,bool)),X8),s(X6,X9),s(X6,X10)))) | p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X8),s(X6,X9),s(X6,X10))))))),
  file('thm.relation.RTC_lifts_monotonicities','thm.relation.RTC_TC_RC')).
fof(f59673,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17),s(sK13,sK17)))) ) | $spl336),
  inference(cnf_transformation,[],[f59673_D])).
fof(f59673_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17),s(sK13,sK17)))) ) <=> ~$spl336),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl336])])).
fof(f59619,plain,(
  ( ! [X212,X213] : (p(s(bool,X212)) | p(s(bool,X213)) | s(bool,X213) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17),s(sK13,sK17)))) ) | ($spl84 | $spl96)),
  inference(resolution,[],[f59511,f783])).
fof(f59511,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17),s(sK13,sK17)))) | ($spl84 | $spl96)),
  inference(forward_demodulation,[],[f12840,f12828])).
fof(f59521,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17),s(sK13,sK17))) | s(bool,X22) != s(bool,X23)) ) | ($spl84 | $spl96)),
  inference(resolution,[],[f59511,f558])).
fof(f59518,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17),s(sK13,sK17)))) ) | ($spl84 | $spl96)),
  inference(resolution,[],[f59511,f557])).
fof(f59517,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17),s(sK13,sK17)))) ) | ($spl84 | $spl96)),
  inference(resolution,[],[f59511,f556])).
fof(f59516,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17),s(sK13,sK17))) | s(bool,X7) != s(bool,X8)) ) | ($spl84 | $spl96)),
  inference(resolution,[],[f59511,f554])).
fof(f59514,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17),s(sK13,sK17)))) ) | ($spl84 | $spl96)),
  inference(resolution,[],[f59511,f552])).
fof(f36362,plain,(
  ( ! [X30] : (p(s(bool,X30)) | s(bool,X30) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK17,sK15,sK13)),s(sK13,sK17)))) ) | $spl204),
  inference(cnf_transformation,[],[f36362_D])).
fof(f36362_D,plain,(
  ( ! [X30] : (p(s(bool,X30)) | s(bool,X30) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK17,sK15,sK13)),s(sK13,sK17)))) ) <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f59447,plain,(
  ( ! [X241,X242] : (p(s(bool,X241)) | p(s(bool,X242)) | s(bool,X242) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK17,sK15,sK13)),s(sK13,sK17)))) ) | ($spl84 | $spl94)),
  inference(resolution,[],[f37283,f783])).
fof(f37283,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK17,sK15,sK13)),s(sK13,sK17)))) | ($spl84 | $spl94)),
  inference(resolution,[],[f30808,f507])).
fof(f30808,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17),s(sK13,sK17)))) | ($spl84 | $spl94)),
  inference(forward_demodulation,[],[f12838,f12828])).
fof(f59349,plain,(
  ( ! [X52,X50,X51] : (p(s(bool,X50)) | ~p(s(bool,X51)) | s(bool,X50) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK17,sK15,sK13)),s(sK13,sK17))) | s(bool,X51) != s(bool,X52)) ) | ($spl84 | $spl94)),
  inference(resolution,[],[f37283,f558])).
fof(f59346,plain,(
  ( ! [X43,X41,X42] : (p(s(bool,X41)) | ~p(s(bool,X42)) | p(s(bool,X43)) | s(bool,X43) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK17,sK15,sK13)),s(sK13,sK17)))) ) | ($spl84 | $spl94)),
  inference(resolution,[],[f37283,f557])).
fof(f59345,plain,(
  ( ! [X39,X38,X40] : (p(s(bool,X38)) | p(s(bool,X39)) | s(bool,X38) != s(bool,X40) | s(bool,X39) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK17,sK15,sK13)),s(sK13,sK17)))) ) | ($spl84 | $spl94)),
  inference(resolution,[],[f37283,f556])).
fof(f59344,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK17,sK15,sK13)),s(sK13,sK17))) | s(bool,X36) != s(bool,X37)) ) | ($spl84 | $spl94)),
  inference(resolution,[],[f37283,f554])).
fof(f59342,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK17,sK15,sK13)),s(sK13,sK17)))) ) | ($spl84 | $spl94)),
  inference(resolution,[],[f37283,f552])).
fof(f36361,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK17,sK15,sK13)),s(sK13,sK17)))) | $spl202),
  inference(cnf_transformation,[],[f36361_D])).
fof(f36361_D,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK17,sK15,sK13)),s(sK13,sK17)))) <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f36359,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK17,sK15,sK13)),s(sK13,sK17)))) | $spl200),
  inference(cnf_transformation,[],[f36359_D])).
fof(f36359_D,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK17,sK15,sK13)),s(sK13,sK17)))) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f59326,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK17,sK15,sK13)),s(sK13,sK17)))) | p(s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK17,sK15,sK13)),s(sK13,sK17)))) | ($spl84 | $spl94)),
  inference(resolution,[],[f37283,f531])).
fof(f36349,plain,(
  s(sK13,sK17) = s(sK13,sK20(sK17,sK17,sK15,sK13)) | $spl190),
  inference(cnf_transformation,[],[f36349_D])).
fof(f36349_D,plain,(
  s(sK13,sK17) = s(sK13,sK20(sK17,sK17,sK15,sK13)) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f59323,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK17,sK15,sK13)),s(sK13,sK17)))) | s(sK13,sK17) = s(sK13,sK20(sK17,sK17,sK15,sK13)) | ($spl84 | $spl94)),
  inference(resolution,[],[f37283,f526])).
fof(f526,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | s(X0,X2) = s(X0,X3) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f359])).
fof(f36357,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK20(sK17,sK17,sK15,sK13),sK15,sK13)))),s(sK13,sK17)))) | $spl198),
  inference(cnf_transformation,[],[f36357_D])).
fof(f36357_D,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK20(sK17,sK17,sK15,sK13),sK15,sK13)))),s(sK13,sK17)))) <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f59321,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK20(sK17,sK17,sK15,sK13),sK15,sK13)))),s(sK13,sK17)))) | s(sK13,sK17) = s(sK13,sK20(sK17,sK17,sK15,sK13)) | ($spl84 | $spl94)),
  inference(resolution,[],[f37283,f523])).
fof(f523,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK28(X3,X2,X1,X0)))),s(X0,X3)))) | s(X0,X2) = s(X0,X3) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f357])).
fof(f357,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | s(X0,X2) = s(X0,X3) | (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,sK28(X3,X2,X1,X0))))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK28(X3,X2,X1,X0)))),s(X0,X3)))))) & ((s(X0,X2) != s(X0,X3) & ! [X4] : (~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X3)))))) | p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK28])],[f356])).
fof(f356,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | s(X0,X2) = s(X0,X3) | ? [X5] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X5)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X3)))))) & ((s(X0,X2) != s(X0,X3) & ! [X4] : (~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X3)))))) | p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))))),
  inference(rectify,[],[f355])).
fof(f355,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | s(X0,X2) = s(X0,X3) | ? [X4] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X3)))))) & ((s(X0,X2) != s(X0,X3) & ! [X4] : (~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X3)))))) | p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))))),
  inference(flattening,[],[f354])).
fof(f354,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | (s(X0,X2) = s(X0,X3) | ? [X4] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X3))))))) & ((s(X0,X2) != s(X0,X3) & ! [X4] : (~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X3)))))) | p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))))),
  inference(nnf_transformation,[],[f154])).
fof(f154,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) <=> (s(X0,X2) = s(X0,X3) | ? [X4] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X3)))))))),
  inference(flattening,[],[f153])).
fof(f153,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) <=> (s(X0,X2) = s(X0,X3) | ? [X4] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X3)))))))),
  inference(rectify,[],[f87])).
fof(f87,axiom,(
  ! [X6] : ! [X8,X9,X10] : (p(s(bool,'const.relation.RTC_3'(s(fun(X6,fun(X6,bool)),X8),s(X6,X9),s(X6,X10)))) <=> (s(X6,X9) = s(X6,X10) | ? [X33] : (p(s(bool,'const.relation.RTC_3'(s(fun(X6,fun(X6,bool)),X8),s(X6,X9),s(X6,X33)))) & p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,X33))),s(X6,X10)))))))),
  file('thm.relation.RTC_lifts_monotonicities','thm.relation.RTC_CASES2')).
fof(f36355,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK17,sK15,sK13)),s(sK13,sK28(sK17,sK20(sK17,sK17,sK15,sK13),sK15,sK13))))) | $spl196),
  inference(cnf_transformation,[],[f36355_D])).
fof(f36355_D,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK17,sK15,sK13)),s(sK13,sK28(sK17,sK20(sK17,sK17,sK15,sK13),sK15,sK13))))) <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f59320,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK17,sK15,sK13)),s(sK13,sK28(sK17,sK20(sK17,sK17,sK15,sK13),sK15,sK13))))) | s(sK13,sK17) = s(sK13,sK20(sK17,sK17,sK15,sK13)) | ($spl84 | $spl94)),
  inference(resolution,[],[f37283,f522])).
fof(f522,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,sK28(X3,X2,X1,X0))))) | s(X0,X2) = s(X0,X3) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f357])).
fof(f36353,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK20(sK17,sK17,sK15,sK13),sK15,sK13)),s(sK13,sK17)))) | $spl194),
  inference(cnf_transformation,[],[f36353_D])).
fof(f36353_D,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK20(sK17,sK17,sK15,sK13),sK15,sK13)),s(sK13,sK17)))) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f59318,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK20(sK17,sK17,sK15,sK13),sK15,sK13)),s(sK13,sK17)))) | s(sK13,sK17) = s(sK13,sK20(sK17,sK17,sK15,sK13)) | ($spl84 | $spl94)),
  inference(resolution,[],[f37283,f519])).
fof(f519,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,sK27(X3,X2,X1,X0)),s(X0,X3)))) | s(X0,X2) = s(X0,X3) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f353])).
fof(f353,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | s(X0,X2) = s(X0,X3) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,sK27(X3,X2,X1,X0))))) & p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,sK27(X3,X2,X1,X0)),s(X0,X3)))))) & ((s(X0,X2) != s(X0,X3) & ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X4)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X3)))))) | p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK27])],[f352])).
fof(f352,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | s(X0,X2) = s(X0,X3) | ? [X5] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X5)))) & p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X5),s(X0,X3)))))) & ((s(X0,X2) != s(X0,X3) & ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X4)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X3)))))) | p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))))),
  inference(rectify,[],[f351])).
fof(f351,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | s(X0,X2) = s(X0,X3) | ? [X4] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X4)))) & p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X3)))))) & ((s(X0,X2) != s(X0,X3) & ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X4)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X3)))))) | p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))))),
  inference(flattening,[],[f350])).
fof(f350,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | (s(X0,X2) = s(X0,X3) | ? [X4] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X4)))) & p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X3))))))) & ((s(X0,X2) != s(X0,X3) & ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X4)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X3)))))) | p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))))),
  inference(nnf_transformation,[],[f152])).
fof(f152,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) <=> (s(X0,X2) = s(X0,X3) | ? [X4] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X4)))) & p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X3)))))))),
  inference(flattening,[],[f151])).
fof(f151,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) <=> (s(X0,X2) = s(X0,X3) | ? [X4] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X4)))) & p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X3)))))))),
  inference(rectify,[],[f75])).
fof(f75,axiom,(
  ! [X6] : ! [X8,X9,X10] : (p(s(bool,'const.relation.RTC_3'(s(fun(X6,fun(X6,bool)),X8),s(X6,X9),s(X6,X10)))) <=> (s(X6,X9) = s(X6,X10) | ? [X33] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,X9))),s(X6,X33)))) & p(s(bool,'const.relation.RTC_3'(s(fun(X6,fun(X6,bool)),X8),s(X6,X33),s(X6,X10)))))))),
  file('thm.relation.RTC_lifts_monotonicities','thm.relation.RTC_CASES|49|')).
fof(f36351,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK17,sK15,sK13)))),s(sK13,sK27(sK17,sK20(sK17,sK17,sK15,sK13),sK15,sK13))))) | $spl192),
  inference(cnf_transformation,[],[f36351_D])).
fof(f36351_D,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK17,sK15,sK13)))),s(sK13,sK27(sK17,sK20(sK17,sK17,sK15,sK13),sK15,sK13))))) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f59317,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK17,sK15,sK13)))),s(sK13,sK27(sK17,sK20(sK17,sK17,sK15,sK13),sK15,sK13))))) | s(sK13,sK17) = s(sK13,sK20(sK17,sK17,sK15,sK13)) | ($spl84 | $spl94)),
  inference(resolution,[],[f37283,f518])).
fof(f518,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,sK27(X3,X2,X1,X0))))) | s(X0,X2) = s(X0,X3) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f353])).
fof(f56171,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl334),
  inference(cnf_transformation,[],[f56171_D])).
fof(f56171_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) <=> ~$spl334),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl334])])).
fof(f56107,plain,(
  ( ! [X212,X213] : (p(s(bool,X212)) | p(s(bool,X213)) | s(bool,X213) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl8),
  inference(resolution,[],[f47025,f783])).
fof(f47025,plain,(
  p(s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) | $spl8),
  inference(resolution,[],[f42217,f534])).
fof(f42217,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) | $spl8),
  inference(resolution,[],[f1006,f529])).
fof(f529,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f272])).
fof(f272,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))))),
  inference(ennf_transformation,[],[f158])).
fof(f158,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) => p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))))),
  inference(flattening,[],[f157])).
fof(f157,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) => p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))))),
  inference(rectify,[],[f73])).
fof(f73,axiom,(
  ! [X6] : ! [X8,X9,X10] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,X9))),s(X6,X10)))) => p(s(bool,'const.relation.RTC_3'(s(fun(X6,fun(X6,bool)),X8),s(X6,X9),s(X6,X10)))))),
  file('thm.relation.RTC_lifts_monotonicities','thm.relation.RTC_SINGLE')).
fof(f56009,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK16,sK15,sK13)),s(sK13,sK17))) | s(bool,X22) != s(bool,X23)) ) | $spl8),
  inference(resolution,[],[f47025,f558])).
fof(f56006,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl8),
  inference(resolution,[],[f47025,f557])).
fof(f56005,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl8),
  inference(resolution,[],[f47025,f556])).
fof(f56004,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK16,sK15,sK13)),s(sK13,sK17))) | s(bool,X7) != s(bool,X8)) ) | $spl8),
  inference(resolution,[],[f47025,f554])).
fof(f56002,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl8),
  inference(resolution,[],[f47025,f552])).
fof(f52868,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17))))) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) | $spl332),
  inference(cnf_transformation,[],[f52868_D])).
fof(f52868_D,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17))))) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) <=> ~$spl332),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl332])])).
fof(f52680,plain,(
  ( ! [X318,X317,X319] : (s(bool,X317) = s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17))))) | p(s(bool,X318)) | p(s(bool,X319)) | s(bool,X317) != s(bool,X319)) )),
  inference(resolution,[],[f1129,f783])).
fof(f1129,plain,(
  ( ! [X200] : (p(s(bool,X200)) | s(bool,X200) = s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17)))))) )),
  inference(resolution,[],[f500,f697])).
fof(f697,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f435])).
fof(f435,plain,(
  ! [X0,X1] : ((s(bool,X0) != s(bool,X1) | ((~p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | p(s(bool,X0))))) & ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(flattening,[],[f434])).
fof(f434,plain,(
  ! [X0,X1] : ((s(bool,X0) != s(bool,X1) | ((~p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | p(s(bool,X0))))) & (((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0)))) | s(bool,X0) = s(bool,X1)))),
  inference(nnf_transformation,[],[f295])).
fof(f295,plain,(
  ! [X0,X1] : (s(bool,X0) = s(bool,X1) <=> ((~p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | p(s(bool,X0)))))),
  inference(ennf_transformation,[],[f204])).
fof(f204,plain,(
  ! [X0,X1] : (s(bool,X0) = s(bool,X1) <=> ((p(s(bool,X0)) => p(s(bool,X1))) & (p(s(bool,X1)) => p(s(bool,X0)))))),
  inference(rectify,[],[f118])).
fof(f118,axiom,(
  ! [X14,X15] : (s(bool,X14) = s(bool,X15) <=> ((p(s(bool,X14)) => p(s(bool,X15))) & (p(s(bool,X15)) => p(s(bool,X14)))))),
  file('thm.relation.RTC_lifts_monotonicities','thm.bool.EQ_IMP_THM')).
fof(f52582,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17))))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) )),
  inference(resolution,[],[f1129,f558])).
fof(f52579,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,X16) = s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17))))) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) )),
  inference(resolution,[],[f1129,f557])).
fof(f52578,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17))))) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) )),
  inference(resolution,[],[f1129,f556])).
fof(f52577,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17))))) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) )),
  inference(resolution,[],[f1129,f554])).
fof(f52575,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17))))) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) )),
  inference(resolution,[],[f1129,f552])).
fof(f51456,plain,(
  ( ! [X202,X205,X203,X201] : (p(s(bool,happ(s(fun(X201,bool),X202),s(X201,X203)))) | s(bool,X205) != s(bool,happ(s(fun(X201,bool),X202),s(X201,sK69(X202,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X201)))) | ~p(s(bool,X205))) ) | $spl330),
  inference(cnf_transformation,[],[f51456_D])).
fof(f51456_D,plain,(
  ( ! [X202,X205,X203,X201] : (p(s(bool,happ(s(fun(X201,bool),X202),s(X201,X203)))) | s(bool,X205) != s(bool,happ(s(fun(X201,bool),X202),s(X201,sK69(X202,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X201)))) | ~p(s(bool,X205))) ) <=> ~$spl330),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl330])])).
fof(f51353,plain,(
  ( ! [X728,X730,X726,X729,X727] : (p(s(bool,happ(s(fun(X726,bool),X727),s(X726,X728)))) | p(s(bool,X729)) | ~p(s(bool,X730)) | s(bool,X730) != s(bool,happ(s(fun(X726,bool),X727),s(X726,sK69(X727,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X726))))) )),
  inference(resolution,[],[f933,f783])).
fof(f933,plain,(
  ( ! [X239,X240,X238] : (~p(s(bool,happ(s(fun(X238,bool),X239),s(X238,sK69(X239,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X238))))) | p(s(bool,happ(s(fun(X238,bool),X239),s(X238,X240))))) )),
  inference(resolution,[],[f499,f745])).
fof(f745,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK69(X2,X1,X0))))) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f472])).
fof(f472,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK68(X2,X1,X0))))) | ! [X5] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))))) & ((~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK69(X2,X1,X0)))))) | (p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK68,sK69])],[f471])).
fof(f471,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,X1)) | ? [X6] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))) | ! [X5] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))))) & (? [X4] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(rectify,[],[f470])).
fof(f470,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,X1)) | ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | ! [X4] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (? [X4] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(flattening,[],[f469])).
fof(f469,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) | ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ! [X4] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (? [X4] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f231])).
fof(f231,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> ! [X4] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f230])).
fof(f230,plain,(
  ! [X0] : ! [X1,X2] : ((p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> ! [X4] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X6] : ! [X16,X12] : ((p(s(bool,X16)) & ! [X9] : p(s(bool,happ(s(fun(X6,bool),X12),s(X6,X9))))) <=> ! [X9] : (p(s(bool,X16)) & p(s(bool,happ(s(fun(X6,bool),X12),s(X6,X9))))))),
  file('thm.relation.RTC_lifts_monotonicities','thm.bool.RIGHT_AND_FORALL_THM')).
fof(f51255,plain,(
  ( ! [X243,X246,X244,X248,X247,X245] : (p(s(bool,happ(s(fun(X243,bool),X244),s(X243,X245)))) | ~p(s(bool,X246)) | ~p(s(bool,X247)) | s(bool,X247) != s(bool,happ(s(fun(X243,bool),X244),s(X243,sK69(X244,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X243)))) | s(bool,X246) != s(bool,X248)) )),
  inference(resolution,[],[f933,f558])).
fof(f51253,plain,(
  ( ! [X231,X235,X233,X236,X234,X232] : (p(s(bool,happ(s(fun(X231,bool),X232),s(X231,X233)))) | p(s(bool,X234)) | ~p(s(bool,X235)) | ~p(s(bool,X236)) | s(bool,X235) != s(bool,happ(s(fun(X231,bool),X232),s(X231,sK69(X232,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X231))))) )),
  inference(resolution,[],[f933,f557])).
fof(f51251,plain,(
  ( ! [X222,X220,X219,X223,X221,X224] : (p(s(bool,happ(s(fun(X219,bool),X220),s(X219,X221)))) | p(s(bool,X222)) | ~p(s(bool,X223)) | s(bool,X222) != s(bool,X224) | s(bool,X223) != s(bool,happ(s(fun(X219,bool),X220),s(X219,sK69(X220,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X219))))) )),
  inference(resolution,[],[f933,f556])).
fof(f51250,plain,(
  ( ! [X218,X216,X215,X213,X217,X214] : (p(s(bool,happ(s(fun(X213,bool),X214),s(X213,X215)))) | ~p(s(bool,X216)) | ~p(s(bool,X217)) | s(bool,X217) != s(bool,happ(s(fun(X213,bool),X214),s(X213,sK69(X214,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X213)))) | s(bool,X216) != s(bool,X218)) )),
  inference(resolution,[],[f933,f554])).
fof(f51248,plain,(
  ( ! [X206,X204,X202,X205,X203,X201] : (p(s(bool,happ(s(fun(X201,bool),X202),s(X201,X203)))) | p(s(bool,X204)) | ~p(s(bool,X205)) | s(bool,X204) != s(bool,X206) | s(bool,X205) != s(bool,happ(s(fun(X201,bool),X202),s(X201,sK69(X202,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X201))))) )),
  inference(resolution,[],[f933,f552])).
fof(f51204,plain,(
  ( ! [X202,X205,X203,X201] : (p(s(bool,happ(s(fun(X201,bool),X202),s(X201,X203)))) | s(bool,X205) != s(bool,happ(s(fun(X201,bool),X202),s(X201,sK68(X202,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X201)))) | ~p(s(bool,X205))) ) | $spl328),
  inference(cnf_transformation,[],[f51204_D])).
fof(f51204_D,plain,(
  ( ! [X202,X205,X203,X201] : (p(s(bool,happ(s(fun(X201,bool),X202),s(X201,X203)))) | s(bool,X205) != s(bool,happ(s(fun(X201,bool),X202),s(X201,sK68(X202,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X201)))) | ~p(s(bool,X205))) ) <=> ~$spl328),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl328])])).
fof(f51101,plain,(
  ( ! [X728,X730,X726,X729,X727] : (p(s(bool,happ(s(fun(X726,bool),X727),s(X726,X728)))) | p(s(bool,X729)) | ~p(s(bool,X730)) | s(bool,X730) != s(bool,happ(s(fun(X726,bool),X727),s(X726,sK68(X727,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X726))))) )),
  inference(resolution,[],[f932,f783])).
fof(f932,plain,(
  ( ! [X237,X235,X236] : (~p(s(bool,happ(s(fun(X235,bool),X236),s(X235,sK68(X236,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X235))))) | p(s(bool,happ(s(fun(X235,bool),X236),s(X235,X237))))) )),
  inference(resolution,[],[f499,f743])).
fof(f743,plain,(
  ( ! [X2,X0,X5,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK68(X2,X1,X0))))) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f472])).
fof(f51003,plain,(
  ( ! [X243,X246,X244,X248,X247,X245] : (p(s(bool,happ(s(fun(X243,bool),X244),s(X243,X245)))) | ~p(s(bool,X246)) | ~p(s(bool,X247)) | s(bool,X247) != s(bool,happ(s(fun(X243,bool),X244),s(X243,sK68(X244,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X243)))) | s(bool,X246) != s(bool,X248)) )),
  inference(resolution,[],[f932,f558])).
fof(f51001,plain,(
  ( ! [X231,X235,X233,X236,X234,X232] : (p(s(bool,happ(s(fun(X231,bool),X232),s(X231,X233)))) | p(s(bool,X234)) | ~p(s(bool,X235)) | ~p(s(bool,X236)) | s(bool,X235) != s(bool,happ(s(fun(X231,bool),X232),s(X231,sK68(X232,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X231))))) )),
  inference(resolution,[],[f932,f557])).
fof(f50999,plain,(
  ( ! [X222,X220,X219,X223,X221,X224] : (p(s(bool,happ(s(fun(X219,bool),X220),s(X219,X221)))) | p(s(bool,X222)) | ~p(s(bool,X223)) | s(bool,X222) != s(bool,X224) | s(bool,X223) != s(bool,happ(s(fun(X219,bool),X220),s(X219,sK68(X220,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X219))))) )),
  inference(resolution,[],[f932,f556])).
fof(f50998,plain,(
  ( ! [X218,X216,X215,X213,X217,X214] : (p(s(bool,happ(s(fun(X213,bool),X214),s(X213,X215)))) | ~p(s(bool,X216)) | ~p(s(bool,X217)) | s(bool,X217) != s(bool,happ(s(fun(X213,bool),X214),s(X213,sK68(X214,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X213)))) | s(bool,X216) != s(bool,X218)) )),
  inference(resolution,[],[f932,f554])).
fof(f50996,plain,(
  ( ! [X206,X204,X202,X205,X203,X201] : (p(s(bool,happ(s(fun(X201,bool),X202),s(X201,X203)))) | p(s(bool,X204)) | ~p(s(bool,X205)) | s(bool,X204) != s(bool,X206) | s(bool,X205) != s(bool,happ(s(fun(X201,bool),X202),s(X201,sK68(X202,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X201))))) )),
  inference(resolution,[],[f932,f552])).
fof(f50952,plain,(
  ( ! [X202,X205,X203,X201] : (p(s(bool,happ(s(fun(X201,bool),X202),s(X201,X203)))) | s(bool,X205) != s(bool,happ(s(fun(X201,bool),X202),s(X201,sK67('const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X202,X201)))) | ~p(s(bool,X205))) ) | $spl326),
  inference(cnf_transformation,[],[f50952_D])).
fof(f50952_D,plain,(
  ( ! [X202,X205,X203,X201] : (p(s(bool,happ(s(fun(X201,bool),X202),s(X201,X203)))) | s(bool,X205) != s(bool,happ(s(fun(X201,bool),X202),s(X201,sK67('const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X202,X201)))) | ~p(s(bool,X205))) ) <=> ~$spl326),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl326])])).
fof(f50849,plain,(
  ( ! [X728,X730,X726,X729,X727] : (p(s(bool,happ(s(fun(X726,bool),X727),s(X726,X728)))) | p(s(bool,X729)) | ~p(s(bool,X730)) | s(bool,X730) != s(bool,happ(s(fun(X726,bool),X727),s(X726,sK67('const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X727,X726))))) )),
  inference(resolution,[],[f931,f783])).
fof(f931,plain,(
  ( ! [X233,X234,X232] : (~p(s(bool,happ(s(fun(X232,bool),X233),s(X232,sK67('const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X233,X232))))) | p(s(bool,happ(s(fun(X232,bool),X233),s(X232,X234))))) )),
  inference(resolution,[],[f499,f740])).
fof(f740,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,X2)) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK67(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f468])).
fof(f468,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK66(X2,X1,X0))))) | ~p(s(bool,X2)) | ! [X5] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & p(s(bool,X2)))) & ((~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK67(X2,X1,X0))))) | ~p(s(bool,X2))) | (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK66,sK67])],[f467])).
fof(f467,plain,(
  ! [X0,X1,X2] : ((? [X6] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) | ~p(s(bool,X2)) | ! [X5] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & p(s(bool,X2)))) & (? [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,X2))) | (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2)))))),
  inference(rectify,[],[f466])).
fof(f466,plain,(
  ! [X0,X1,X2] : ((? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,X2)) | ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & p(s(bool,X2)))) & (? [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,X2))) | (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2)))))),
  inference(flattening,[],[f465])).
fof(f465,plain,(
  ! [X0,X1,X2] : (((? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,X2))) | ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & p(s(bool,X2)))) & (? [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,X2))) | (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2)))))),
  inference(nnf_transformation,[],[f229])).
fof(f229,plain,(
  ! [X0,X1,X2] : ((! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2))) <=> ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & p(s(bool,X2))))),
  inference(flattening,[],[f228])).
fof(f228,plain,(
  ! [X0] : ! [X1,X2] : ((! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2))) <=> ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & p(s(bool,X2))))),
  inference(rectify,[],[f130])).
fof(f130,axiom,(
  ! [X6] : ! [X16,X12] : ((! [X9] : p(s(bool,happ(s(fun(X6,bool),X16),s(X6,X9)))) & p(s(bool,X12))) <=> ! [X9] : (p(s(bool,happ(s(fun(X6,bool),X16),s(X6,X9)))) & p(s(bool,X12))))),
  file('thm.relation.RTC_lifts_monotonicities','thm.bool.LEFT_AND_FORALL_THM')).
fof(f50751,plain,(
  ( ! [X243,X246,X244,X248,X247,X245] : (p(s(bool,happ(s(fun(X243,bool),X244),s(X243,X245)))) | ~p(s(bool,X246)) | ~p(s(bool,X247)) | s(bool,X247) != s(bool,happ(s(fun(X243,bool),X244),s(X243,sK67('const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X244,X243)))) | s(bool,X246) != s(bool,X248)) )),
  inference(resolution,[],[f931,f558])).
fof(f50749,plain,(
  ( ! [X231,X235,X233,X236,X234,X232] : (p(s(bool,happ(s(fun(X231,bool),X232),s(X231,X233)))) | p(s(bool,X234)) | ~p(s(bool,X235)) | ~p(s(bool,X236)) | s(bool,X235) != s(bool,happ(s(fun(X231,bool),X232),s(X231,sK67('const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X232,X231))))) )),
  inference(resolution,[],[f931,f557])).
fof(f50747,plain,(
  ( ! [X222,X220,X219,X223,X221,X224] : (p(s(bool,happ(s(fun(X219,bool),X220),s(X219,X221)))) | p(s(bool,X222)) | ~p(s(bool,X223)) | s(bool,X222) != s(bool,X224) | s(bool,X223) != s(bool,happ(s(fun(X219,bool),X220),s(X219,sK67('const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X220,X219))))) )),
  inference(resolution,[],[f931,f556])).
fof(f50746,plain,(
  ( ! [X218,X216,X215,X213,X217,X214] : (p(s(bool,happ(s(fun(X213,bool),X214),s(X213,X215)))) | ~p(s(bool,X216)) | ~p(s(bool,X217)) | s(bool,X217) != s(bool,happ(s(fun(X213,bool),X214),s(X213,sK67('const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X214,X213)))) | s(bool,X216) != s(bool,X218)) )),
  inference(resolution,[],[f931,f554])).
fof(f50744,plain,(
  ( ! [X206,X204,X202,X205,X203,X201] : (p(s(bool,happ(s(fun(X201,bool),X202),s(X201,X203)))) | p(s(bool,X204)) | ~p(s(bool,X205)) | s(bool,X204) != s(bool,X206) | s(bool,X205) != s(bool,happ(s(fun(X201,bool),X202),s(X201,sK67('const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X202,X201))))) )),
  inference(resolution,[],[f931,f552])).
fof(f50700,plain,(
  ( ! [X202,X205,X203,X201] : (p(s(bool,happ(s(fun(X201,bool),X202),s(X201,X203)))) | s(bool,X205) != s(bool,happ(s(fun(X201,bool),X202),s(X201,sK66('const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X202,X201)))) | ~p(s(bool,X205))) ) | $spl324),
  inference(cnf_transformation,[],[f50700_D])).
fof(f50700_D,plain,(
  ( ! [X202,X205,X203,X201] : (p(s(bool,happ(s(fun(X201,bool),X202),s(X201,X203)))) | s(bool,X205) != s(bool,happ(s(fun(X201,bool),X202),s(X201,sK66('const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X202,X201)))) | ~p(s(bool,X205))) ) <=> ~$spl324),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl324])])).
fof(f50597,plain,(
  ( ! [X728,X730,X726,X729,X727] : (p(s(bool,happ(s(fun(X726,bool),X727),s(X726,X728)))) | p(s(bool,X729)) | ~p(s(bool,X730)) | s(bool,X730) != s(bool,happ(s(fun(X726,bool),X727),s(X726,sK66('const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X727,X726))))) )),
  inference(resolution,[],[f930,f783])).
fof(f930,plain,(
  ( ! [X231,X229,X230] : (~p(s(bool,happ(s(fun(X229,bool),X230),s(X229,sK66('const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X230,X229))))) | p(s(bool,happ(s(fun(X229,bool),X230),s(X229,X231))))) )),
  inference(resolution,[],[f499,f738])).
fof(f738,plain,(
  ( ! [X2,X0,X5,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ~p(s(bool,X2)) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK66(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f468])).
fof(f50499,plain,(
  ( ! [X243,X246,X244,X248,X247,X245] : (p(s(bool,happ(s(fun(X243,bool),X244),s(X243,X245)))) | ~p(s(bool,X246)) | ~p(s(bool,X247)) | s(bool,X247) != s(bool,happ(s(fun(X243,bool),X244),s(X243,sK66('const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X244,X243)))) | s(bool,X246) != s(bool,X248)) )),
  inference(resolution,[],[f930,f558])).
fof(f50497,plain,(
  ( ! [X231,X235,X233,X236,X234,X232] : (p(s(bool,happ(s(fun(X231,bool),X232),s(X231,X233)))) | p(s(bool,X234)) | ~p(s(bool,X235)) | ~p(s(bool,X236)) | s(bool,X235) != s(bool,happ(s(fun(X231,bool),X232),s(X231,sK66('const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X232,X231))))) )),
  inference(resolution,[],[f930,f557])).
fof(f50495,plain,(
  ( ! [X222,X220,X219,X223,X221,X224] : (p(s(bool,happ(s(fun(X219,bool),X220),s(X219,X221)))) | p(s(bool,X222)) | ~p(s(bool,X223)) | s(bool,X222) != s(bool,X224) | s(bool,X223) != s(bool,happ(s(fun(X219,bool),X220),s(X219,sK66('const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X220,X219))))) )),
  inference(resolution,[],[f930,f556])).
fof(f50494,plain,(
  ( ! [X218,X216,X215,X213,X217,X214] : (p(s(bool,happ(s(fun(X213,bool),X214),s(X213,X215)))) | ~p(s(bool,X216)) | ~p(s(bool,X217)) | s(bool,X217) != s(bool,happ(s(fun(X213,bool),X214),s(X213,sK66('const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X214,X213)))) | s(bool,X216) != s(bool,X218)) )),
  inference(resolution,[],[f930,f554])).
fof(f50492,plain,(
  ( ! [X206,X204,X202,X205,X203,X201] : (p(s(bool,happ(s(fun(X201,bool),X202),s(X201,X203)))) | p(s(bool,X204)) | ~p(s(bool,X205)) | s(bool,X204) != s(bool,X206) | s(bool,X205) != s(bool,happ(s(fun(X201,bool),X202),s(X201,sK66('const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X202,X201))))) )),
  inference(resolution,[],[f930,f552])).
fof(f50448,plain,(
  ( ! [X15] : (p(s(bool,X15)) | s(bool,X15) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl322),
  inference(cnf_transformation,[],[f50448_D])).
fof(f50448_D,plain,(
  ( ! [X15] : (p(s(bool,X15)) | s(bool,X15) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) <=> ~$spl322),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl322])])).
fof(f50384,plain,(
  ( ! [X227,X226] : (p(s(bool,X226)) | p(s(bool,X227)) | s(bool,X227) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl8),
  inference(resolution,[],[f42218,f783])).
fof(f42218,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) | $spl8),
  inference(resolution,[],[f1006,f764])).
fof(f764,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f310])).
fof(f310,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))))),
  inference(ennf_transformation,[],[f248])).
fof(f248,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) => p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))))),
  inference(flattening,[],[f247])).
fof(f247,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) => p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))))),
  inference(rectify,[],[f55])).
fof(f55,axiom,(
  ! [X6] : ! [X8,X9,X10] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,X9))),s(X6,X10)))) => p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X8),s(X6,X9),s(X6,X10)))))),
  file('thm.relation.RTC_lifts_monotonicities','thm.relation.TC_RULES|split|0')).
fof(f50286,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK16,sK15,sK13)),s(sK13,sK17))) | s(bool,X36) != s(bool,X37)) ) | $spl8),
  inference(resolution,[],[f42218,f558])).
fof(f50283,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | p(s(bool,X28)) | s(bool,X28) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl8),
  inference(resolution,[],[f42218,f557])).
fof(f50282,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X23) != s(bool,X25) | s(bool,X24) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl8),
  inference(resolution,[],[f42218,f556])).
fof(f50281,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X20) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK16,sK15,sK13)),s(sK13,sK17))) | s(bool,X21) != s(bool,X22)) ) | $spl8),
  inference(resolution,[],[f42218,f554])).
fof(f50279,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl8),
  inference(resolution,[],[f42218,f552])).
fof(f47223,plain,(
  ( ! [X30] : (p(s(bool,X30)) | s(bool,X30) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl314),
  inference(cnf_transformation,[],[f47223_D])).
fof(f47223_D,plain,(
  ( ! [X30] : (p(s(bool,X30)) | s(bool,X30) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) <=> ~$spl314),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl314])])).
fof(f50192,plain,(
  ( ! [X241,X242] : (p(s(bool,X241)) | p(s(bool,X242)) | s(bool,X242) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl8),
  inference(resolution,[],[f42217,f783])).
fof(f50094,plain,(
  ( ! [X52,X50,X51] : (p(s(bool,X50)) | ~p(s(bool,X51)) | s(bool,X50) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK16,sK15,sK13)),s(sK13,sK17))) | s(bool,X51) != s(bool,X52)) ) | $spl8),
  inference(resolution,[],[f42217,f558])).
fof(f50091,plain,(
  ( ! [X43,X41,X42] : (p(s(bool,X41)) | ~p(s(bool,X42)) | p(s(bool,X43)) | s(bool,X43) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl8),
  inference(resolution,[],[f42217,f557])).
fof(f50090,plain,(
  ( ! [X39,X38,X40] : (p(s(bool,X38)) | p(s(bool,X39)) | s(bool,X38) != s(bool,X40) | s(bool,X39) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl8),
  inference(resolution,[],[f42217,f556])).
fof(f50089,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK16,sK15,sK13)),s(sK13,sK17))) | s(bool,X36) != s(bool,X37)) ) | $spl8),
  inference(resolution,[],[f42217,f554])).
fof(f50087,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl8),
  inference(resolution,[],[f42217,f552])).
fof(f47222,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) | $spl312),
  inference(cnf_transformation,[],[f47222_D])).
fof(f47222_D,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) <=> ~$spl312),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl312])])).
fof(f47220,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) | $spl310),
  inference(cnf_transformation,[],[f47220_D])).
fof(f47220_D,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) <=> ~$spl310),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl310])])).
fof(f50071,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) | p(s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) | $spl8),
  inference(resolution,[],[f42217,f531])).
fof(f46099,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) | $spl292),
  inference(cnf_transformation,[],[f46099_D])).
fof(f46099_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) <=> ~$spl292),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl292])])).
fof(f49985,plain,(
  ( ! [X212,X213] : (p(s(bool,X212)) | p(s(bool,X213)) | s(bool,X213) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) | $spl80),
  inference(resolution,[],[f11251,f783])).
fof(f49887,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13)))) | s(bool,X22) != s(bool,X23)) ) | $spl80),
  inference(resolution,[],[f11251,f558])).
fof(f49884,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) | $spl80),
  inference(resolution,[],[f11251,f557])).
fof(f49883,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) | $spl80),
  inference(resolution,[],[f11251,f556])).
fof(f49882,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13)))) | s(bool,X7) != s(bool,X8)) ) | $spl80),
  inference(resolution,[],[f11251,f554])).
fof(f49880,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) | $spl80),
  inference(resolution,[],[f11251,f552])).
fof(f49878,plain,(
  ( ! [X420,X422,X418,X417] : (p(s(bool,happ(s(fun(X417,bool),X418),s(X417,sK56(X418,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X417))))) | s(bool,X420) != s(bool,happ(s(fun(X417,bool),X418),s(X417,X422))) | ~p(s(bool,X420))) ) | $spl320),
  inference(cnf_transformation,[],[f49878_D])).
fof(f49878_D,plain,(
  ( ! [X420,X422,X418,X417] : (p(s(bool,happ(s(fun(X417,bool),X418),s(X417,sK56(X418,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X417))))) | s(bool,X420) != s(bool,happ(s(fun(X417,bool),X418),s(X417,X422))) | ~p(s(bool,X420))) ) <=> ~$spl320),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl320])])).
fof(f49798,plain,(
  ( ! [X945,X942,X944,X943,X941] : (p(s(bool,happ(s(fun(X941,bool),X942),s(X941,sK56(X942,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X941))))) | p(s(bool,X943)) | ~p(s(bool,X944)) | s(bool,X944) != s(bool,happ(s(fun(X941,bool),X942),s(X941,X945)))) )),
  inference(resolution,[],[f929,f783])).
fof(f929,plain,(
  ( ! [X227,X228,X226] : (~p(s(bool,happ(s(fun(X226,bool),X227),s(X226,X228)))) | p(s(bool,happ(s(fun(X226,bool),X227),s(X226,sK56(X227,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X226)))))) )),
  inference(resolution,[],[f499,f721])).
fof(f721,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK56(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f448])).
fof(f448,plain,(
  ! [X0,X1,X2] : ((! [X6] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK55(X2,X1,X0))))))) & (~p(s(bool,X1)) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK56(X2,X1,X0))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK55,sK56])],[f447])).
fof(f447,plain,(
  ! [X0,X1,X2] : ((! [X6] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | (p(s(bool,X1)) & ? [X5] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))))) & (~p(s(bool,X1)) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X3] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(rectify,[],[f446])).
fof(f446,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (p(s(bool,X1)) & ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (~p(s(bool,X1)) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X3] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(flattening,[],[f445])).
fof(f445,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (p(s(bool,X1)) & ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & ((~p(s(bool,X1)) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ? [X3] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f219])).
fof(f219,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (p(s(bool,X1)) & ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f218])).
fof(f218,plain,(
  ! [X0] : ! [X1,X2] : (? [X3] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (p(s(bool,X1)) & ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f99])).
fof(f99,axiom,(
  ! [X6] : ! [X16,X12] : (? [X9] : (p(s(bool,X16)) & p(s(bool,happ(s(fun(X6,bool),X12),s(X6,X9))))) <=> (p(s(bool,X16)) & ? [X9] : p(s(bool,happ(s(fun(X6,bool),X12),s(X6,X9))))))),
  file('thm.relation.RTC_lifts_monotonicities','thm.bool.RIGHT_EXISTS_AND_THM')).
fof(f49700,plain,(
  ( ! [X461,X463,X459,X460,X462,X458] : (p(s(bool,happ(s(fun(X458,bool),X459),s(X458,sK56(X459,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X458))))) | ~p(s(bool,X460)) | ~p(s(bool,X461)) | s(bool,X461) != s(bool,happ(s(fun(X458,bool),X459),s(X458,X462))) | s(bool,X460) != s(bool,X463)) )),
  inference(resolution,[],[f929,f558])).
fof(f49698,plain,(
  ( ! [X447,X449,X451,X452,X448,X450] : (p(s(bool,happ(s(fun(X447,bool),X448),s(X447,sK56(X448,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X447))))) | p(s(bool,X449)) | ~p(s(bool,X450)) | ~p(s(bool,X451)) | s(bool,X450) != s(bool,happ(s(fun(X447,bool),X448),s(X447,X452)))) )),
  inference(resolution,[],[f929,f557])).
fof(f49696,plain,(
  ( ! [X436,X438,X440,X437,X439,X435] : (p(s(bool,happ(s(fun(X435,bool),X436),s(X435,sK56(X436,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X435))))) | p(s(bool,X437)) | ~p(s(bool,X438)) | s(bool,X437) != s(bool,X439) | s(bool,X438) != s(bool,happ(s(fun(X435,bool),X436),s(X435,X440)))) )),
  inference(resolution,[],[f929,f556])).
fof(f49695,plain,(
  ( ! [X430,X432,X434,X429,X431,X433] : (p(s(bool,happ(s(fun(X429,bool),X430),s(X429,sK56(X430,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X429))))) | ~p(s(bool,X431)) | ~p(s(bool,X432)) | s(bool,X432) != s(bool,happ(s(fun(X429,bool),X430),s(X429,X433))) | s(bool,X431) != s(bool,X434)) )),
  inference(resolution,[],[f929,f554])).
fof(f49693,plain,(
  ( ! [X420,X422,X418,X421,X417,X419] : (p(s(bool,happ(s(fun(X417,bool),X418),s(X417,sK56(X418,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X417))))) | p(s(bool,X419)) | ~p(s(bool,X420)) | s(bool,X419) != s(bool,X421) | s(bool,X420) != s(bool,happ(s(fun(X417,bool),X418),s(X417,X422)))) )),
  inference(resolution,[],[f929,f552])).
fof(f22045,plain,(
  ( ! [X15] : (p(s(bool,X15)) | s(bool,X15) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) | $spl134),
  inference(cnf_transformation,[],[f22045_D])).
fof(f22045_D,plain,(
  ( ! [X15] : (p(s(bool,X15)) | s(bool,X15) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f49533,plain,(
  ( ! [X227,X226] : (p(s(bool,X226)) | p(s(bool,X227)) | s(bool,X227) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) | $spl78),
  inference(resolution,[],[f11249,f783])).
fof(f49435,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13)))) | s(bool,X36) != s(bool,X37)) ) | $spl78),
  inference(resolution,[],[f11249,f558])).
fof(f49432,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | p(s(bool,X28)) | s(bool,X28) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) | $spl78),
  inference(resolution,[],[f11249,f557])).
fof(f49431,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X23) != s(bool,X25) | s(bool,X24) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) | $spl78),
  inference(resolution,[],[f11249,f556])).
fof(f49430,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X20) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13)))) | s(bool,X21) != s(bool,X22)) ) | $spl78),
  inference(resolution,[],[f11249,f554])).
fof(f49428,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) | $spl78),
  inference(resolution,[],[f11249,f552])).
fof(f49250,plain,(
  ( ! [X420,X422,X418,X417] : (p(s(bool,happ(s(fun(X417,bool),X418),s(X417,sK55(X418,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X417))))) | s(bool,X420) != s(bool,happ(s(fun(X417,bool),X418),s(X417,X422))) | ~p(s(bool,X420))) ) | $spl318),
  inference(cnf_transformation,[],[f49250_D])).
fof(f49250_D,plain,(
  ( ! [X420,X422,X418,X417] : (p(s(bool,happ(s(fun(X417,bool),X418),s(X417,sK55(X418,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X417))))) | s(bool,X420) != s(bool,happ(s(fun(X417,bool),X418),s(X417,X422))) | ~p(s(bool,X420))) ) <=> ~$spl318),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl318])])).
fof(f49170,plain,(
  ( ! [X945,X942,X944,X943,X941] : (p(s(bool,happ(s(fun(X941,bool),X942),s(X941,sK55(X942,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X941))))) | p(s(bool,X943)) | ~p(s(bool,X944)) | s(bool,X944) != s(bool,happ(s(fun(X941,bool),X942),s(X941,X945)))) )),
  inference(resolution,[],[f928,f783])).
fof(f928,plain,(
  ( ! [X225,X223,X224] : (~p(s(bool,happ(s(fun(X223,bool),X224),s(X223,X225)))) | p(s(bool,happ(s(fun(X223,bool),X224),s(X223,sK55(X224,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X223)))))) )),
  inference(resolution,[],[f499,f719])).
fof(f719,plain,(
  ( ! [X6,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK55(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f448])).
fof(f49072,plain,(
  ( ! [X461,X463,X459,X460,X462,X458] : (p(s(bool,happ(s(fun(X458,bool),X459),s(X458,sK55(X459,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X458))))) | ~p(s(bool,X460)) | ~p(s(bool,X461)) | s(bool,X461) != s(bool,happ(s(fun(X458,bool),X459),s(X458,X462))) | s(bool,X460) != s(bool,X463)) )),
  inference(resolution,[],[f928,f558])).
fof(f49070,plain,(
  ( ! [X447,X449,X451,X452,X448,X450] : (p(s(bool,happ(s(fun(X447,bool),X448),s(X447,sK55(X448,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X447))))) | p(s(bool,X449)) | ~p(s(bool,X450)) | ~p(s(bool,X451)) | s(bool,X450) != s(bool,happ(s(fun(X447,bool),X448),s(X447,X452)))) )),
  inference(resolution,[],[f928,f557])).
fof(f49068,plain,(
  ( ! [X436,X438,X440,X437,X439,X435] : (p(s(bool,happ(s(fun(X435,bool),X436),s(X435,sK55(X436,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X435))))) | p(s(bool,X437)) | ~p(s(bool,X438)) | s(bool,X437) != s(bool,X439) | s(bool,X438) != s(bool,happ(s(fun(X435,bool),X436),s(X435,X440)))) )),
  inference(resolution,[],[f928,f556])).
fof(f49067,plain,(
  ( ! [X430,X432,X434,X429,X431,X433] : (p(s(bool,happ(s(fun(X429,bool),X430),s(X429,sK55(X430,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X429))))) | ~p(s(bool,X431)) | ~p(s(bool,X432)) | s(bool,X432) != s(bool,happ(s(fun(X429,bool),X430),s(X429,X433))) | s(bool,X431) != s(bool,X434)) )),
  inference(resolution,[],[f928,f554])).
fof(f49065,plain,(
  ( ! [X420,X422,X418,X421,X417,X419] : (p(s(bool,happ(s(fun(X417,bool),X418),s(X417,sK55(X418,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X417))))) | p(s(bool,X419)) | ~p(s(bool,X420)) | s(bool,X419) != s(bool,X421) | s(bool,X420) != s(bool,happ(s(fun(X417,bool),X418),s(X417,X422)))) )),
  inference(resolution,[],[f928,f552])).
fof(f42418,plain,(
  ( ! [X22] : (p(s(bool,X22)) | s(bool,X22) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK16,sK15,sK13)))),s(sK13,sK17)))) ) | $spl216),
  inference(cnf_transformation,[],[f42418_D])).
fof(f42418_D,plain,(
  ( ! [X22] : (p(s(bool,X22)) | s(bool,X22) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK16,sK15,sK13)))),s(sK13,sK17)))) ) <=> ~$spl216),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl216])])).
fof(f48693,plain,(
  ( ! [X233,X234] : (p(s(bool,X233)) | p(s(bool,X234)) | s(bool,X234) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK16,sK15,sK13)))),s(sK13,sK17)))) ) | $spl8),
  inference(resolution,[],[f1006,f783])).
fof(f48595,plain,(
  ( ! [X43,X44,X42] : (p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X42) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK16,sK15,sK13)))),s(sK13,sK17))) | s(bool,X43) != s(bool,X44)) ) | $spl8),
  inference(resolution,[],[f1006,f558])).
fof(f48592,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X33)) | ~p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X35) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK16,sK15,sK13)))),s(sK13,sK17)))) ) | $spl8),
  inference(resolution,[],[f1006,f557])).
fof(f48591,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | p(s(bool,X31)) | s(bool,X30) != s(bool,X32) | s(bool,X31) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK16,sK15,sK13)))),s(sK13,sK17)))) ) | $spl8),
  inference(resolution,[],[f1006,f556])).
fof(f48590,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK16,sK15,sK13)))),s(sK13,sK17))) | s(bool,X28) != s(bool,X29)) ) | $spl8),
  inference(resolution,[],[f1006,f554])).
fof(f48588,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK16,sK15,sK13)))),s(sK13,sK17)))) ) | $spl8),
  inference(resolution,[],[f1006,f552])).
fof(f21429,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) | $spl132),
  inference(cnf_transformation,[],[f21429_D])).
fof(f21429_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f48420,plain,(
  ( ! [X212,X213] : (p(s(bool,X212)) | p(s(bool,X213)) | s(bool,X213) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) | $spl6),
  inference(resolution,[],[f11055,f783])).
fof(f11055,plain,(
  p(s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) | $spl6),
  inference(resolution,[],[f1004,f534])).
fof(f48322,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13)))) | s(bool,X22) != s(bool,X23)) ) | $spl6),
  inference(resolution,[],[f11055,f558])).
fof(f48319,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) | $spl6),
  inference(resolution,[],[f11055,f557])).
fof(f48318,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) | $spl6),
  inference(resolution,[],[f11055,f556])).
fof(f48317,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13)))) | s(bool,X7) != s(bool,X8)) ) | $spl6),
  inference(resolution,[],[f11055,f554])).
fof(f48315,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) | $spl6),
  inference(resolution,[],[f11055,f552])).
fof(f11252,plain,(
  ( ! [X30] : (p(s(bool,X30)) | s(bool,X30) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) | $spl82),
  inference(cnf_transformation,[],[f11252_D])).
fof(f11252_D,plain,(
  ( ! [X30] : (p(s(bool,X30)) | s(bool,X30) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f48246,plain,(
  ( ! [X241,X242] : (p(s(bool,X241)) | p(s(bool,X242)) | s(bool,X242) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) | $spl6),
  inference(resolution,[],[f1004,f783])).
fof(f48148,plain,(
  ( ! [X52,X50,X51] : (p(s(bool,X50)) | ~p(s(bool,X51)) | s(bool,X50) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13)))) | s(bool,X51) != s(bool,X52)) ) | $spl6),
  inference(resolution,[],[f1004,f558])).
fof(f48145,plain,(
  ( ! [X43,X41,X42] : (p(s(bool,X41)) | ~p(s(bool,X42)) | p(s(bool,X43)) | s(bool,X43) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) | $spl6),
  inference(resolution,[],[f1004,f557])).
fof(f48144,plain,(
  ( ! [X39,X38,X40] : (p(s(bool,X38)) | p(s(bool,X39)) | s(bool,X38) != s(bool,X40) | s(bool,X39) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) | $spl6),
  inference(resolution,[],[f1004,f556])).
fof(f48143,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13)))) | s(bool,X36) != s(bool,X37)) ) | $spl6),
  inference(resolution,[],[f1004,f554])).
fof(f48141,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) | $spl6),
  inference(resolution,[],[f1004,f552])).
fof(f11251,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) | $spl80),
  inference(cnf_transformation,[],[f11251_D])).
fof(f11251_D,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f11249,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) | $spl78),
  inference(cnf_transformation,[],[f11249_D])).
fof(f11249_D,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f48125,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) | p(s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) | $spl6),
  inference(resolution,[],[f1004,f531])).
fof(f48092,plain,(
  ( ! [X22] : (p(s(bool,X22)) | s(bool,X22) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,sK17)))) ) | $spl316),
  inference(cnf_transformation,[],[f48092_D])).
fof(f48092_D,plain,(
  ( ! [X22] : (p(s(bool,X22)) | s(bool,X22) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,sK17)))) ) <=> ~$spl316),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl316])])).
fof(f48022,plain,(
  ( ! [X233,X234] : (p(s(bool,X233)) | p(s(bool,X234)) | s(bool,X234) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,sK17)))) ) | ($spl8 | $spl300)),
  inference(resolution,[],[f47225,f783])).
fof(f47225,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,sK17)))) | ($spl8 | $spl300)),
  inference(backward_demodulation,[],[f47210,f1006])).
fof(f47924,plain,(
  ( ! [X43,X44,X42] : (p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X42) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,sK17))) | s(bool,X43) != s(bool,X44)) ) | ($spl8 | $spl300)),
  inference(resolution,[],[f47225,f558])).
fof(f47921,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X33)) | ~p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X35) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,sK17)))) ) | ($spl8 | $spl300)),
  inference(resolution,[],[f47225,f557])).
fof(f47920,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | p(s(bool,X31)) | s(bool,X30) != s(bool,X32) | s(bool,X31) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,sK17)))) ) | ($spl8 | $spl300)),
  inference(resolution,[],[f47225,f556])).
fof(f47919,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,sK17))) | s(bool,X28) != s(bool,X29)) ) | ($spl8 | $spl300)),
  inference(resolution,[],[f47225,f554])).
fof(f47917,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,sK17)))) ) | ($spl8 | $spl300)),
  inference(resolution,[],[f47225,f552])).
fof(f47145,plain,(
  ( ! [X241,X242] : (p(s(bool,X241)) | p(s(bool,X242)) | s(bool,X242) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl8),
  inference(resolution,[],[f42217,f783])).
fof(f47047,plain,(
  ( ! [X52,X50,X51] : (p(s(bool,X50)) | ~p(s(bool,X51)) | s(bool,X50) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK16,sK15,sK13)),s(sK13,sK17))) | s(bool,X51) != s(bool,X52)) ) | $spl8),
  inference(resolution,[],[f42217,f558])).
fof(f47044,plain,(
  ( ! [X43,X41,X42] : (p(s(bool,X41)) | ~p(s(bool,X42)) | p(s(bool,X43)) | s(bool,X43) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl8),
  inference(resolution,[],[f42217,f557])).
fof(f47043,plain,(
  ( ! [X39,X38,X40] : (p(s(bool,X38)) | p(s(bool,X39)) | s(bool,X38) != s(bool,X40) | s(bool,X39) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl8),
  inference(resolution,[],[f42217,f556])).
fof(f47042,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK16,sK15,sK13)),s(sK13,sK17))) | s(bool,X36) != s(bool,X37)) ) | $spl8),
  inference(resolution,[],[f42217,f554])).
fof(f47040,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl8),
  inference(resolution,[],[f42217,f552])).
fof(f47024,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) | p(s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) | $spl8),
  inference(resolution,[],[f42217,f531])).
fof(f47210,plain,(
  s(sK13,sK17) = s(sK13,sK28(sK17,sK16,sK15,sK13)) | $spl300),
  inference(cnf_transformation,[],[f47210_D])).
fof(f47210_D,plain,(
  s(sK13,sK17) = s(sK13,sK28(sK17,sK16,sK15,sK13)) <=> ~$spl300),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl300])])).
fof(f47021,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) | s(sK13,sK17) = s(sK13,sK28(sK17,sK16,sK15,sK13)) | $spl8),
  inference(resolution,[],[f42217,f526])).
fof(f47218,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK28(sK17,sK16,sK15,sK13),sK15,sK13)))),s(sK13,sK17)))) | $spl308),
  inference(cnf_transformation,[],[f47218_D])).
fof(f47218_D,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK28(sK17,sK16,sK15,sK13),sK15,sK13)))),s(sK13,sK17)))) <=> ~$spl308),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl308])])).
fof(f47019,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK28(sK17,sK16,sK15,sK13),sK15,sK13)))),s(sK13,sK17)))) | s(sK13,sK17) = s(sK13,sK28(sK17,sK16,sK15,sK13)) | $spl8),
  inference(resolution,[],[f42217,f523])).
fof(f47216,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK16,sK15,sK13)),s(sK13,sK28(sK17,sK28(sK17,sK16,sK15,sK13),sK15,sK13))))) | $spl306),
  inference(cnf_transformation,[],[f47216_D])).
fof(f47216_D,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK16,sK15,sK13)),s(sK13,sK28(sK17,sK28(sK17,sK16,sK15,sK13),sK15,sK13))))) <=> ~$spl306),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl306])])).
fof(f47018,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK16,sK15,sK13)),s(sK13,sK28(sK17,sK28(sK17,sK16,sK15,sK13),sK15,sK13))))) | s(sK13,sK17) = s(sK13,sK28(sK17,sK16,sK15,sK13)) | $spl8),
  inference(resolution,[],[f42217,f522])).
fof(f47214,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK28(sK17,sK16,sK15,sK13),sK15,sK13)),s(sK13,sK17)))) | $spl304),
  inference(cnf_transformation,[],[f47214_D])).
fof(f47214_D,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK28(sK17,sK16,sK15,sK13),sK15,sK13)),s(sK13,sK17)))) <=> ~$spl304),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl304])])).
fof(f47016,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK28(sK17,sK16,sK15,sK13),sK15,sK13)),s(sK13,sK17)))) | s(sK13,sK17) = s(sK13,sK28(sK17,sK16,sK15,sK13)) | $spl8),
  inference(resolution,[],[f42217,f519])).
fof(f47212,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK16,sK15,sK13)))),s(sK13,sK27(sK17,sK28(sK17,sK16,sK15,sK13),sK15,sK13))))) | $spl302),
  inference(cnf_transformation,[],[f47212_D])).
fof(f47212_D,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK16,sK15,sK13)))),s(sK13,sK27(sK17,sK28(sK17,sK16,sK15,sK13),sK15,sK13))))) <=> ~$spl302),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl302])])).
fof(f47015,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK16,sK15,sK13)))),s(sK13,sK27(sK17,sK28(sK17,sK16,sK15,sK13),sK15,sK13))))) | s(sK13,sK17) = s(sK13,sK28(sK17,sK16,sK15,sK13)) | $spl8),
  inference(resolution,[],[f42217,f518])).
fof(f47001,plain,(
  ( ! [X420,X422,X418,X417] : (p(s(bool,happ(s(fun(X417,bool),X418),s(X417,sK54('const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X418,X417))))) | s(bool,X420) != s(bool,happ(s(fun(X417,bool),X418),s(X417,X422))) | ~p(s(bool,X420))) ) | $spl298),
  inference(cnf_transformation,[],[f47001_D])).
fof(f47001_D,plain,(
  ( ! [X420,X422,X418,X417] : (p(s(bool,happ(s(fun(X417,bool),X418),s(X417,sK54('const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X418,X417))))) | s(bool,X420) != s(bool,happ(s(fun(X417,bool),X418),s(X417,X422))) | ~p(s(bool,X420))) ) <=> ~$spl298),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl298])])).
fof(f46921,plain,(
  ( ! [X945,X942,X944,X943,X941] : (p(s(bool,happ(s(fun(X941,bool),X942),s(X941,sK54('const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X942,X941))))) | p(s(bool,X943)) | ~p(s(bool,X944)) | s(bool,X944) != s(bool,happ(s(fun(X941,bool),X942),s(X941,X945)))) )),
  inference(resolution,[],[f927,f783])).
fof(f927,plain,(
  ( ! [X222,X220,X221] : (~p(s(bool,happ(s(fun(X220,bool),X221),s(X220,X222)))) | p(s(bool,happ(s(fun(X220,bool),X221),s(X220,sK54('const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X221,X220)))))) )),
  inference(resolution,[],[f499,f716])).
fof(f716,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK54(X2,X1,X0))))) | ~p(s(bool,X2)) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) )),
  inference(cnf_transformation,[],[f444])).
fof(f444,plain,(
  ! [X0,X1,X2] : ((! [X6] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) | ~p(s(bool,X2))) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK53(X2,X1,X0))))) & p(s(bool,X2)))) & (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,X2)) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK54(X2,X1,X0))))) & p(s(bool,X2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK53,sK54])],[f443])).
fof(f443,plain,(
  ! [X0,X1,X2] : ((! [X6] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) | ~p(s(bool,X2))) | (? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & p(s(bool,X2)))) & (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,X2)) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2)))))),
  inference(rectify,[],[f442])).
fof(f442,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,X2))) | (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & p(s(bool,X2)))) & (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,X2)) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2)))))),
  inference(flattening,[],[f441])).
fof(f441,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,X2))) | (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & p(s(bool,X2)))) & ((! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,X2))) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2)))))),
  inference(nnf_transformation,[],[f217])).
fof(f217,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2))) <=> (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & p(s(bool,X2))))),
  inference(flattening,[],[f216])).
fof(f216,plain,(
  ! [X0] : ! [X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2))) <=> (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & p(s(bool,X2))))),
  inference(rectify,[],[f108])).
fof(f108,axiom,(
  ! [X6] : ! [X16,X12] : (? [X9] : (p(s(bool,happ(s(fun(X6,bool),X16),s(X6,X9)))) & p(s(bool,X12))) <=> (? [X9] : p(s(bool,happ(s(fun(X6,bool),X16),s(X6,X9)))) & p(s(bool,X12))))),
  file('thm.relation.RTC_lifts_monotonicities','thm.bool.LEFT_EXISTS_AND_THM')).
fof(f46823,plain,(
  ( ! [X461,X463,X459,X460,X462,X458] : (p(s(bool,happ(s(fun(X458,bool),X459),s(X458,sK54('const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X459,X458))))) | ~p(s(bool,X460)) | ~p(s(bool,X461)) | s(bool,X461) != s(bool,happ(s(fun(X458,bool),X459),s(X458,X462))) | s(bool,X460) != s(bool,X463)) )),
  inference(resolution,[],[f927,f558])).
fof(f46821,plain,(
  ( ! [X447,X449,X451,X452,X448,X450] : (p(s(bool,happ(s(fun(X447,bool),X448),s(X447,sK54('const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X448,X447))))) | p(s(bool,X449)) | ~p(s(bool,X450)) | ~p(s(bool,X451)) | s(bool,X450) != s(bool,happ(s(fun(X447,bool),X448),s(X447,X452)))) )),
  inference(resolution,[],[f927,f557])).
fof(f46819,plain,(
  ( ! [X436,X438,X440,X437,X439,X435] : (p(s(bool,happ(s(fun(X435,bool),X436),s(X435,sK54('const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X436,X435))))) | p(s(bool,X437)) | ~p(s(bool,X438)) | s(bool,X437) != s(bool,X439) | s(bool,X438) != s(bool,happ(s(fun(X435,bool),X436),s(X435,X440)))) )),
  inference(resolution,[],[f927,f556])).
fof(f46818,plain,(
  ( ! [X430,X432,X434,X429,X431,X433] : (p(s(bool,happ(s(fun(X429,bool),X430),s(X429,sK54('const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X430,X429))))) | ~p(s(bool,X431)) | ~p(s(bool,X432)) | s(bool,X432) != s(bool,happ(s(fun(X429,bool),X430),s(X429,X433))) | s(bool,X431) != s(bool,X434)) )),
  inference(resolution,[],[f927,f554])).
fof(f46816,plain,(
  ( ! [X420,X422,X418,X421,X417,X419] : (p(s(bool,happ(s(fun(X417,bool),X418),s(X417,sK54('const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X418,X417))))) | p(s(bool,X419)) | ~p(s(bool,X420)) | s(bool,X419) != s(bool,X421) | s(bool,X420) != s(bool,happ(s(fun(X417,bool),X418),s(X417,X422)))) )),
  inference(resolution,[],[f927,f552])).
fof(f46719,plain,(
  ( ! [X420,X422,X418,X417] : (p(s(bool,happ(s(fun(X417,bool),X418),s(X417,sK53('const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X418,X417))))) | s(bool,X420) != s(bool,happ(s(fun(X417,bool),X418),s(X417,X422))) | ~p(s(bool,X420))) ) | $spl296),
  inference(cnf_transformation,[],[f46719_D])).
fof(f46719_D,plain,(
  ( ! [X420,X422,X418,X417] : (p(s(bool,happ(s(fun(X417,bool),X418),s(X417,sK53('const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X418,X417))))) | s(bool,X420) != s(bool,happ(s(fun(X417,bool),X418),s(X417,X422))) | ~p(s(bool,X420))) ) <=> ~$spl296),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl296])])).
fof(f46639,plain,(
  ( ! [X945,X942,X944,X943,X941] : (p(s(bool,happ(s(fun(X941,bool),X942),s(X941,sK53('const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X942,X941))))) | p(s(bool,X943)) | ~p(s(bool,X944)) | s(bool,X944) != s(bool,happ(s(fun(X941,bool),X942),s(X941,X945)))) )),
  inference(resolution,[],[f926,f783])).
fof(f926,plain,(
  ( ! [X218,X219,X217] : (~p(s(bool,happ(s(fun(X217,bool),X218),s(X217,X219)))) | p(s(bool,happ(s(fun(X217,bool),X218),s(X217,sK53('const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X218,X217)))))) )),
  inference(resolution,[],[f499,f714])).
fof(f714,plain,(
  ( ! [X6,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK53(X2,X1,X0))))) | ~p(s(bool,X2)) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6))))) )),
  inference(cnf_transformation,[],[f444])).
fof(f46541,plain,(
  ( ! [X461,X463,X459,X460,X462,X458] : (p(s(bool,happ(s(fun(X458,bool),X459),s(X458,sK53('const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X459,X458))))) | ~p(s(bool,X460)) | ~p(s(bool,X461)) | s(bool,X461) != s(bool,happ(s(fun(X458,bool),X459),s(X458,X462))) | s(bool,X460) != s(bool,X463)) )),
  inference(resolution,[],[f926,f558])).
fof(f46539,plain,(
  ( ! [X447,X449,X451,X452,X448,X450] : (p(s(bool,happ(s(fun(X447,bool),X448),s(X447,sK53('const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X448,X447))))) | p(s(bool,X449)) | ~p(s(bool,X450)) | ~p(s(bool,X451)) | s(bool,X450) != s(bool,happ(s(fun(X447,bool),X448),s(X447,X452)))) )),
  inference(resolution,[],[f926,f557])).
fof(f46537,plain,(
  ( ! [X436,X438,X440,X437,X439,X435] : (p(s(bool,happ(s(fun(X435,bool),X436),s(X435,sK53('const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X436,X435))))) | p(s(bool,X437)) | ~p(s(bool,X438)) | s(bool,X437) != s(bool,X439) | s(bool,X438) != s(bool,happ(s(fun(X435,bool),X436),s(X435,X440)))) )),
  inference(resolution,[],[f926,f556])).
fof(f46536,plain,(
  ( ! [X430,X432,X434,X429,X431,X433] : (p(s(bool,happ(s(fun(X429,bool),X430),s(X429,sK53('const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X430,X429))))) | ~p(s(bool,X431)) | ~p(s(bool,X432)) | s(bool,X432) != s(bool,happ(s(fun(X429,bool),X430),s(X429,X433))) | s(bool,X431) != s(bool,X434)) )),
  inference(resolution,[],[f926,f554])).
fof(f46534,plain,(
  ( ! [X420,X422,X418,X421,X417,X419] : (p(s(bool,happ(s(fun(X417,bool),X418),s(X417,sK53('const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)),X418,X417))))) | p(s(bool,X419)) | ~p(s(bool,X420)) | s(bool,X419) != s(bool,X421) | s(bool,X420) != s(bool,happ(s(fun(X417,bool),X418),s(X417,X422)))) )),
  inference(resolution,[],[f926,f552])).
fof(f46374,plain,(
  ( ! [X212,X213] : (p(s(bool,X212)) | p(s(bool,X213)) | s(bool,X213) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) | $spl80),
  inference(resolution,[],[f11251,f783])).
fof(f46276,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13)))) | s(bool,X22) != s(bool,X23)) ) | $spl80),
  inference(resolution,[],[f11251,f558])).
fof(f46273,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) | $spl80),
  inference(resolution,[],[f11251,f557])).
fof(f46272,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) | $spl80),
  inference(resolution,[],[f11251,f556])).
fof(f46271,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13)))) | s(bool,X7) != s(bool,X8)) ) | $spl80),
  inference(resolution,[],[f11251,f554])).
fof(f46269,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) | $spl80),
  inference(resolution,[],[f11251,f552])).
fof(f46267,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) | $spl294),
  inference(cnf_transformation,[],[f46267_D])).
fof(f46267_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) <=> ~$spl294),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl294])])).
fof(f46206,plain,(
  ( ! [X210,X211] : (p(s(bool,X210)) | ~p(s(bool,X211)) | s(bool,X211) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) | $spl81),
  inference(resolution,[],[f11250,f783])).
fof(f11250,plain,(
  ~p(s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) | $spl81),
  inference(cnf_transformation,[],[f11250_D])).
fof(f11250_D,plain,(
  ~p(s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) <=> ~$spl81),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl81])])).
fof(f46108,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13)))) | s(bool,X21) != s(bool,X23)) ) | $spl81),
  inference(resolution,[],[f11250,f558])).
fof(f46106,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) | $spl81),
  inference(resolution,[],[f11250,f557])).
fof(f46104,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) | $spl81),
  inference(resolution,[],[f11250,f556])).
fof(f46103,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13)))) | s(bool,X6) != s(bool,X8)) ) | $spl81),
  inference(resolution,[],[f11250,f554])).
fof(f46101,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) | $spl81),
  inference(resolution,[],[f11250,f552])).
fof(f46035,plain,(
  ( ! [X212,X213] : (p(s(bool,X212)) | p(s(bool,X213)) | s(bool,X213) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) | $spl80),
  inference(resolution,[],[f11251,f783])).
fof(f45937,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13)))) | s(bool,X22) != s(bool,X23)) ) | $spl80),
  inference(resolution,[],[f11251,f558])).
fof(f45934,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) | $spl80),
  inference(resolution,[],[f11251,f557])).
fof(f45933,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) | $spl80),
  inference(resolution,[],[f11251,f556])).
fof(f45932,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13)))) | s(bool,X7) != s(bool,X8)) ) | $spl80),
  inference(resolution,[],[f11251,f554])).
fof(f45930,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) | $spl80),
  inference(resolution,[],[f11251,f552])).
fof(f45865,plain,(
  ( ! [X227,X226] : (p(s(bool,X226)) | p(s(bool,X227)) | s(bool,X227) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) | $spl78),
  inference(resolution,[],[f11249,f783])).
fof(f45767,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13)))) | s(bool,X36) != s(bool,X37)) ) | $spl78),
  inference(resolution,[],[f11249,f558])).
fof(f45764,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | p(s(bool,X28)) | s(bool,X28) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) | $spl78),
  inference(resolution,[],[f11249,f557])).
fof(f45763,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X23) != s(bool,X25) | s(bool,X24) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) | $spl78),
  inference(resolution,[],[f11249,f556])).
fof(f45762,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X20) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13)))) | s(bool,X21) != s(bool,X22)) ) | $spl78),
  inference(resolution,[],[f11249,f554])).
fof(f45760,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) | $spl78),
  inference(resolution,[],[f11249,f552])).
fof(f45244,plain,(
  ( ! [X212,X213] : (p(s(bool,X212)) | p(s(bool,X213)) | s(bool,X213) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) | $spl6),
  inference(resolution,[],[f11055,f783])).
fof(f45146,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13)))) | s(bool,X22) != s(bool,X23)) ) | $spl6),
  inference(resolution,[],[f11055,f558])).
fof(f45143,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) | $spl6),
  inference(resolution,[],[f11055,f557])).
fof(f45142,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) | $spl6),
  inference(resolution,[],[f11055,f556])).
fof(f45141,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13)))) | s(bool,X7) != s(bool,X8)) ) | $spl6),
  inference(resolution,[],[f11055,f554])).
fof(f45139,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) | $spl6),
  inference(resolution,[],[f11055,f552])).
fof(f44986,plain,(
  ( ! [X70,X68] : (p(s(bool,X68)) | s(bool,X68) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,X70)))) ) | $spl290),
  inference(cnf_transformation,[],[f44986_D])).
fof(f44986_D,plain,(
  ( ! [X70,X68] : (p(s(bool,X68)) | s(bool,X68) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,X70)))) ) <=> ~$spl290),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl290])])).
fof(f44915,plain,(
  ( ! [X383,X385,X384] : (p(s(bool,X383)) | p(s(bool,X384)) | s(bool,X384) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,X385)))) ) | $spl232),
  inference(resolution,[],[f44735,f783])).
fof(f44817,plain,(
  ( ! [X94,X97,X95,X96] : (p(s(bool,X94)) | ~p(s(bool,X95)) | s(bool,X94) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,X96))) | s(bool,X95) != s(bool,X97)) ) | $spl232),
  inference(resolution,[],[f44735,f558])).
fof(f44814,plain,(
  ( ! [X85,X83,X86,X84] : (p(s(bool,X83)) | ~p(s(bool,X84)) | p(s(bool,X85)) | s(bool,X85) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,X86)))) ) | $spl232),
  inference(resolution,[],[f44735,f557])).
fof(f44813,plain,(
  ( ! [X80,X81,X79,X82] : (p(s(bool,X79)) | p(s(bool,X80)) | s(bool,X79) != s(bool,X81) | s(bool,X80) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,X82)))) ) | $spl232),
  inference(resolution,[],[f44735,f556])).
fof(f44812,plain,(
  ( ! [X78,X76,X77,X75] : (p(s(bool,X75)) | ~p(s(bool,X76)) | s(bool,X75) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,X77))) | s(bool,X76) != s(bool,X78)) ) | $spl232),
  inference(resolution,[],[f44735,f554])).
fof(f44810,plain,(
  ( ! [X70,X68,X69,X67] : (p(s(bool,X67)) | p(s(bool,X68)) | s(bool,X67) != s(bool,X69) | s(bool,X68) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,X70)))) ) | $spl232),
  inference(resolution,[],[f44735,f552])).
fof(f44767,plain,(
  ( ! [X121,X119] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,X119)))) | ~p(s(bool,X121)) | s(bool,X121) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,X119)))) ) | $spl288),
  inference(cnf_transformation,[],[f44767_D])).
fof(f44767_D,plain,(
  ( ! [X121,X119] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,X119)))) | ~p(s(bool,X121)) | s(bool,X121) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,X119)))) ) <=> ~$spl288),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl288])])).
fof(f44643,plain,(
  ( ! [X436,X434,X435] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,X434)))) | p(s(bool,X435)) | ~p(s(bool,X436)) | s(bool,X436) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,X434)))) )),
  inference(resolution,[],[f795,f783])).
fof(f795,plain,(
  ( ! [X4] : (~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,X4)))) | p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,X4))))) )),
  inference(resolution,[],[f499,f505])).
fof(f505,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X4)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f269])).
fof(f269,plain,(
  ! [X0,X1,X2,X3,X4] : (~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X4)))) | p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4)))))),
  inference(flattening,[],[f268])).
fof(f268,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X4))))) | p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4)))))),
  inference(ennf_transformation,[],[f144])).
fof(f144,plain,(
  ! [X0,X1,X2,X3,X4] : ((p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X4))))) => p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4)))))),
  inference(flattening,[],[f143])).
fof(f143,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : ((p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X4))))) => p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4)))))),
  inference(rectify,[],[f74])).
fof(f74,axiom,(
  ! [X6] : ! [X8,X9,X10,X32] : ((p(s(bool,'const.relation.RTC_3'(s(fun(X6,fun(X6,bool)),X8),s(X6,X9),s(X6,X10)))) & p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,X10))),s(X6,X32))))) => p(s(bool,'const.relation.RTC_3'(s(fun(X6,fun(X6,bool)),X8),s(X6,X9),s(X6,X32)))))),
  file('thm.relation.RTC_lifts_monotonicities','thm.relation.RTC_RULES_RIGHT|49||split|1')).
fof(f44545,plain,(
  ( ! [X149,X147,X150,X148] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,X147)))) | ~p(s(bool,X148)) | ~p(s(bool,X149)) | s(bool,X149) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,X147))) | s(bool,X148) != s(bool,X150)) )),
  inference(resolution,[],[f795,f558])).
fof(f44543,plain,(
  ( ! [X142,X140,X141,X139] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,X139)))) | p(s(bool,X140)) | ~p(s(bool,X141)) | ~p(s(bool,X142)) | s(bool,X141) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,X139)))) )),
  inference(resolution,[],[f795,f557])).
fof(f44541,plain,(
  ( ! [X134,X132,X133,X131] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,X131)))) | p(s(bool,X132)) | ~p(s(bool,X133)) | s(bool,X132) != s(bool,X134) | s(bool,X133) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,X131)))) )),
  inference(resolution,[],[f795,f556])).
fof(f44540,plain,(
  ( ! [X127,X130,X128,X129] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,X127)))) | ~p(s(bool,X128)) | ~p(s(bool,X129)) | s(bool,X129) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,X127))) | s(bool,X128) != s(bool,X130)) )),
  inference(resolution,[],[f795,f554])).
fof(f44538,plain,(
  ( ! [X121,X122,X120,X119] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,X119)))) | p(s(bool,X120)) | ~p(s(bool,X121)) | s(bool,X120) != s(bool,X122) | s(bool,X121) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,X119)))) )),
  inference(resolution,[],[f795,f552])).
fof(f44750,plain,(
  ( ! [X86] : (~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,X86))))) ) | $spl258),
  inference(cnf_transformation,[],[f44750_D])).
fof(f44750_D,plain,(
  ( ! [X86] : (~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,X86))))) ) <=> ~$spl258),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl258])])).
fof(f44766,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK71(happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),sK13))))) | $spl286),
  inference(cnf_transformation,[],[f44766_D])).
fof(f44766_D,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK71(happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),sK13))))) <=> ~$spl286),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl286])])).
fof(f44537,plain,(
  ( ! [X118] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK71(happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,X118))))) )),
  inference(resolution,[],[f795,f749])).
fof(f749,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK71(X1,X0)))))) )),
  inference(cnf_transformation,[],[f477])).
fof(f477,plain,(
  ! [X0,X1] : ((p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK70(X1,X0))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) & (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK71(X1,X0))))) | ! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK70,sK71])],[f476])).
fof(f476,plain,(
  ! [X0,X1] : ((? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) & (? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(rectify,[],[f475])).
fof(f475,plain,(
  ! [X0,X1] : ((? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(nnf_transformation,[],[f307])).
fof(f307,plain,(
  ! [X0,X1] : (! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(ennf_transformation,[],[f234])).
fof(f234,plain,(
  ! [X0,X1] : (~? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(flattening,[],[f233])).
fof(f233,plain,(
  ! [X0] : ! [X1] : (~? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(rectify,[],[f109])).
fof(f109,axiom,(
  ! [X6] : ! [X16] : (~? [X9] : p(s(bool,happ(s(fun(X6,bool),X16),s(X6,X9)))) <=> ! [X9] : ~p(s(bool,happ(s(fun(X6,bool),X16),s(X6,X9)))))),
  file('thm.relation.RTC_lifts_monotonicities','thm.bool.NOT_EXISTS_THM')).
fof(f44764,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK70(happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),sK13))))) | $spl284),
  inference(cnf_transformation,[],[f44764_D])).
fof(f44764_D,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK70(happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),sK13))))) <=> ~$spl284),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl284])])).
fof(f44536,plain,(
  ( ! [X117] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK70(happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,X117))))) )),
  inference(resolution,[],[f795,f748])).
fof(f748,plain,(
  ( ! [X4,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK70(X1,X0)))))) )),
  inference(cnf_transformation,[],[f477])).
fof(f44762,plain,(
  ( ! [X115] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK65(happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),X115,sK13))))) | p(s(bool,X115))) ) | $spl282),
  inference(cnf_transformation,[],[f44762_D])).
fof(f44762_D,plain,(
  ( ! [X115] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK65(happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),X115,sK13))))) | p(s(bool,X115))) ) <=> ~$spl282),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl282])])).
fof(f44535,plain,(
  ( ! [X116,X115] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK65(happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),X115,sK13))))) | p(s(bool,X115)) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,X116))))) )),
  inference(resolution,[],[f795,f737])).
fof(f737,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK65(X2,X1,X0))))) | p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) )),
  inference(cnf_transformation,[],[f464])).
fof(f464,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X6] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK64(X2,X1,X0))))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK65(X2,X1,X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK64,sK65])],[f463])).
fof(f463,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X6] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | ? [X5] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(rectify,[],[f462])).
fof(f462,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(flattening,[],[f461])).
fof(f461,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f227])).
fof(f227,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f226])).
fof(f226,plain,(
  ! [X0] : ! [X1,X2] : ((p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f94])).
fof(f94,axiom,(
  ! [X6] : ! [X16,X12] : ((p(s(bool,X16)) | ? [X9] : p(s(bool,happ(s(fun(X6,bool),X12),s(X6,X9))))) <=> ? [X9] : (p(s(bool,X16)) | p(s(bool,happ(s(fun(X6,bool),X12),s(X6,X9))))))),
  file('thm.relation.RTC_lifts_monotonicities','thm.bool.RIGHT_OR_EXISTS_THM')).
fof(f44761,plain,(
  ( ! [X113] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK64(happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),X113,sK13))))) | p(s(bool,X113))) ) | $spl280),
  inference(cnf_transformation,[],[f44761_D])).
fof(f44761_D,plain,(
  ( ! [X113] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK64(happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),X113,sK13))))) | p(s(bool,X113))) ) <=> ~$spl280),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl280])])).
fof(f44534,plain,(
  ( ! [X114,X113] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK64(happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),X113,sK13))))) | p(s(bool,X113)) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,X114))))) )),
  inference(resolution,[],[f795,f735])).
fof(f735,plain,(
  ( ! [X6,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK64(X2,X1,X0))))) | p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) )),
  inference(cnf_transformation,[],[f464])).
fof(f44760,plain,(
  ( ! [X111] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK63(X111,happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),sK13))))) | p(s(bool,X111))) ) | $spl278),
  inference(cnf_transformation,[],[f44760_D])).
fof(f44760_D,plain,(
  ( ! [X111] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK63(X111,happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),sK13))))) | p(s(bool,X111))) ) <=> ~$spl278),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl278])])).
fof(f44533,plain,(
  ( ! [X111,X112] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK63(X111,happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),sK13))))) | p(s(bool,X111)) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,X112))))) )),
  inference(resolution,[],[f795,f732])).
fof(f732,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,X2)) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK63(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) )),
  inference(cnf_transformation,[],[f460])).
fof(f460,plain,(
  ! [X0,X1,X2] : (((! [X6] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) & ~p(s(bool,X2))) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK62(X2,X1,X0))))) | p(s(bool,X2)))) & (! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ~p(s(bool,X2))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK63(X2,X1,X0))))) | p(s(bool,X2))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK62,sK63])],[f459])).
fof(f459,plain,(
  ! [X0,X1,X2] : (((! [X6] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) & ~p(s(bool,X2))) | ? [X5] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | p(s(bool,X2)))) & (! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ~p(s(bool,X2))) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2))))),
  inference(rectify,[],[f458])).
fof(f458,plain,(
  ! [X0,X1,X2] : (((! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & ~p(s(bool,X2))) | ? [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,X2)))) & (! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ~p(s(bool,X2))) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2))))),
  inference(flattening,[],[f457])).
fof(f457,plain,(
  ! [X0,X1,X2] : (((! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & ~p(s(bool,X2))) | ? [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,X2)))) & (! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ~p(s(bool,X2))) | (? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2)))))),
  inference(nnf_transformation,[],[f225])).
fof(f225,plain,(
  ! [X0,X1,X2] : ((? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2))) <=> ? [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,X2))))),
  inference(flattening,[],[f224])).
fof(f224,plain,(
  ! [X0] : ! [X1,X2] : ((? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2))) <=> ? [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,X2))))),
  inference(rectify,[],[f122])).
fof(f122,axiom,(
  ! [X6] : ! [X16,X12] : ((? [X9] : p(s(bool,happ(s(fun(X6,bool),X16),s(X6,X9)))) | p(s(bool,X12))) <=> ? [X9] : (p(s(bool,happ(s(fun(X6,bool),X16),s(X6,X9)))) | p(s(bool,X12))))),
  file('thm.relation.RTC_lifts_monotonicities','thm.bool.LEFT_OR_EXISTS_THM')).
fof(f44759,plain,(
  ( ! [X109] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK62(X109,happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),sK13))))) | p(s(bool,X109))) ) | $spl276),
  inference(cnf_transformation,[],[f44759_D])).
fof(f44759_D,plain,(
  ( ! [X109] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK62(X109,happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),sK13))))) | p(s(bool,X109))) ) <=> ~$spl276),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl276])])).
fof(f44532,plain,(
  ( ! [X109,X110] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK62(X109,happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),sK13))))) | p(s(bool,X109)) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,X110))))) )),
  inference(resolution,[],[f795,f730])).
fof(f730,plain,(
  ( ! [X6,X2,X0,X1] : (p(s(bool,X2)) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK62(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6))))) )),
  inference(cnf_transformation,[],[f460])).
fof(f44758,plain,(
  ( ! [X107] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK56(happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),X107,sK13))))) | ~p(s(bool,X107))) ) | $spl274),
  inference(cnf_transformation,[],[f44758_D])).
fof(f44758_D,plain,(
  ( ! [X107] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK56(happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),X107,sK13))))) | ~p(s(bool,X107))) ) <=> ~$spl274),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl274])])).
fof(f44531,plain,(
  ( ! [X107,X108] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK56(happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),X107,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,X108)))) | ~p(s(bool,X107))) )),
  inference(resolution,[],[f795,f721])).
fof(f44757,plain,(
  ( ! [X105] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK55(happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),X105,sK13))))) | ~p(s(bool,X105))) ) | $spl272),
  inference(cnf_transformation,[],[f44757_D])).
fof(f44757_D,plain,(
  ( ! [X105] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK55(happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),X105,sK13))))) | ~p(s(bool,X105))) ) <=> ~$spl272),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl272])])).
fof(f44530,plain,(
  ( ! [X105,X106] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK55(happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),X105,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,X106)))) | ~p(s(bool,X105))) )),
  inference(resolution,[],[f795,f719])).
fof(f44756,plain,(
  ( ! [X103] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK54(X103,happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),sK13))))) | ~p(s(bool,X103))) ) | $spl270),
  inference(cnf_transformation,[],[f44756_D])).
fof(f44756_D,plain,(
  ( ! [X103] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK54(X103,happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),sK13))))) | ~p(s(bool,X103))) ) <=> ~$spl270),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl270])])).
fof(f44529,plain,(
  ( ! [X103,X104] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK54(X103,happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),sK13))))) | ~p(s(bool,X103)) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,X104))))) )),
  inference(resolution,[],[f795,f716])).
fof(f44755,plain,(
  ( ! [X101] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK53(X101,happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),sK13))))) | ~p(s(bool,X101))) ) | $spl268),
  inference(cnf_transformation,[],[f44755_D])).
fof(f44755_D,plain,(
  ( ! [X101] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK53(X101,happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),sK13))))) | ~p(s(bool,X101))) ) <=> ~$spl268),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl268])])).
fof(f44528,plain,(
  ( ! [X101,X102] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK53(X101,happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),sK13))))) | ~p(s(bool,X101)) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,X102))))) )),
  inference(resolution,[],[f795,f714])).
fof(f44754,plain,(
  ( ! [X97] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK52(X97,happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X97),s(sK13,sK52(X97,happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),sK13)))))) ) | $spl266),
  inference(cnf_transformation,[],[f44754_D])).
fof(f44754_D,plain,(
  ( ! [X97] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK52(X97,happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X97),s(sK13,sK52(X97,happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),sK13)))))) ) <=> ~$spl266),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl266])])).
fof(f44526,plain,(
  ( ! [X97,X98] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK52(X97,happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X97),s(sK13,sK52(X97,happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,X98))))) )),
  inference(resolution,[],[f795,f712])).
fof(f712,plain,(
  ( ! [X2,X0,X5,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK52(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK52(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))) )),
  inference(cnf_transformation,[],[f440])).
fof(f440,plain,(
  ! [X0,X1,X2] : ((! [X8] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X8)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X8))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK50(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK51(X2,X1,X0)))))) & ((! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK52(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK52(X2,X1,X0))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK50,sK51,sK52])],[f439])).
fof(f439,plain,(
  ! [X0,X1,X2] : ((! [X8] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X8)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X8))))) | ? [X7] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X7)))) | ? [X6] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) & ((! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(rectify,[],[f438])).
fof(f438,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) & ((! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(flattening,[],[f437])).
fof(f437,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & ((! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f215])).
fof(f215,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f214])).
fof(f214,plain,(
  ! [X0] : ! [X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f93])).
fof(f93,axiom,(
  ! [X6] : ! [X16,X12] : (? [X9] : (p(s(bool,happ(s(fun(X6,bool),X16),s(X6,X9)))) | p(s(bool,happ(s(fun(X6,bool),X12),s(X6,X9))))) <=> (? [X9] : p(s(bool,happ(s(fun(X6,bool),X16),s(X6,X9)))) | ? [X9] : p(s(bool,happ(s(fun(X6,bool),X12),s(X6,X9))))))),
  file('thm.relation.RTC_lifts_monotonicities','thm.bool.EXISTS_OR_THM')).
fof(f44753,plain,(
  ( ! [X95] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK52(happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),X95,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X95),s(sK13,sK52(happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),X95,sK13)))))) ) | $spl264),
  inference(cnf_transformation,[],[f44753_D])).
fof(f44753_D,plain,(
  ( ! [X95] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK52(happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),X95,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X95),s(sK13,sK52(happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),X95,sK13)))))) ) <=> ~$spl264),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl264])])).
fof(f44525,plain,(
  ( ! [X95,X96] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK52(happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),X95,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X95),s(sK13,sK52(happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),X95,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,X96))))) )),
  inference(resolution,[],[f795,f713])).
fof(f713,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK52(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK52(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) )),
  inference(cnf_transformation,[],[f440])).
fof(f44752,plain,(
  ( ! [X91] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK51(happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),X91,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X91),s(sK13,sK50(happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),X91,sK13)))))) ) | $spl262),
  inference(cnf_transformation,[],[f44752_D])).
fof(f44752_D,plain,(
  ( ! [X91] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK51(happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),X91,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X91),s(sK13,sK50(happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),X91,sK13)))))) ) <=> ~$spl262),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl262])])).
fof(f44523,plain,(
  ( ! [X92,X91] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK51(happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),X91,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X91),s(sK13,sK50(happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),X91,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,X92))))) )),
  inference(resolution,[],[f795,f711])).
fof(f711,plain,(
  ( ! [X2,X0,X8,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK51(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK50(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X8))))) )),
  inference(cnf_transformation,[],[f440])).
fof(f44751,plain,(
  ( ! [X85] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK50(X85,happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X85),s(sK13,sK51(X85,happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),sK13)))))) ) | $spl260),
  inference(cnf_transformation,[],[f44751_D])).
fof(f44751_D,plain,(
  ( ! [X85] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK50(X85,happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X85),s(sK13,sK51(X85,happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),sK13)))))) ) <=> ~$spl260),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl260])])).
fof(f44520,plain,(
  ( ! [X85,X86] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK50(X85,happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X85),s(sK13,sK51(X85,happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,X86))))) )),
  inference(resolution,[],[f795,f710])).
fof(f710,plain,(
  ( ! [X2,X0,X8,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK51(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK50(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X8))))) )),
  inference(cnf_transformation,[],[f440])).
fof(f44749,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,sK73(happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),sK13))))) | $spl257),
  inference(cnf_transformation,[],[f44749_D])).
fof(f44749_D,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,sK73(happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),sK13))))) <=> ~$spl257),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl257])])).
fof(f44735,plain,(
  ( ! [X62] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,X62))))) ) | $spl232),
  inference(cnf_transformation,[],[f44735_D])).
fof(f44735_D,plain,(
  ( ! [X62] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,X62))))) ) <=> ~$spl232),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl232])])).
fof(f44519,plain,(
  ( ! [X84] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,X84)))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,sK73(happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),sK13)))))) )),
  inference(resolution,[],[f795,f751])).
fof(f751,plain,(
  ( ! [X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK73(X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f480])).
fof(f480,plain,(
  ! [X0,X1] : ((! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK72(X1,X0)))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK73(X1,X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK72,sK73])],[f479])).
fof(f479,plain,(
  ! [X0,X1] : ((! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(rectify,[],[f478])).
fof(f478,plain,(
  ! [X0,X1] : ((! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(nnf_transformation,[],[f308])).
fof(f308,plain,(
  ! [X0,X1] : (? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(ennf_transformation,[],[f236])).
fof(f236,plain,(
  ! [X0,X1] : (~! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(flattening,[],[f235])).
fof(f235,plain,(
  ! [X0] : ! [X1] : (~! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(rectify,[],[f88])).
fof(f88,axiom,(
  ! [X6] : ! [X16] : (~! [X9] : p(s(bool,happ(s(fun(X6,bool),X16),s(X6,X9)))) <=> ? [X9] : ~p(s(bool,happ(s(fun(X6,bool),X16),s(X6,X9)))))),
  file('thm.relation.RTC_lifts_monotonicities','thm.bool.NOT_FORALL_THM')).
fof(f44747,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,sK72(happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),sK13))))) | $spl255),
  inference(cnf_transformation,[],[f44747_D])).
fof(f44747_D,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,sK72(happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),sK13))))) <=> ~$spl255),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl255])])).
fof(f44518,plain,(
  ( ! [X83] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,X83)))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,sK72(happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),sK13)))))) )),
  inference(resolution,[],[f795,f750])).
fof(f750,plain,(
  ( ! [X0,X5,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK72(X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))) )),
  inference(cnf_transformation,[],[f480])).
fof(f44745,plain,(
  ( ! [X82] : (~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,sK69(happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),X82,sK13))))) | ~p(s(bool,X82))) ) | $spl252),
  inference(cnf_transformation,[],[f44745_D])).
fof(f44745_D,plain,(
  ( ! [X82] : (~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,sK69(happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),X82,sK13))))) | ~p(s(bool,X82))) ) <=> ~$spl252),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl252])])).
fof(f44517,plain,(
  ( ! [X81,X82] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,X81)))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,sK69(happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),X82,sK13))))) | ~p(s(bool,X82))) )),
  inference(resolution,[],[f795,f745])).
fof(f44744,plain,(
  ( ! [X80] : (~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,sK68(happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),X80,sK13))))) | ~p(s(bool,X80))) ) | $spl250),
  inference(cnf_transformation,[],[f44744_D])).
fof(f44744_D,plain,(
  ( ! [X80] : (~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,sK68(happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),X80,sK13))))) | ~p(s(bool,X80))) ) <=> ~$spl250),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl250])])).
fof(f44516,plain,(
  ( ! [X80,X79] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,X79)))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,sK68(happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),X80,sK13))))) | ~p(s(bool,X80))) )),
  inference(resolution,[],[f795,f743])).
fof(f44743,plain,(
  ( ! [X78] : (~p(s(bool,X78)) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,sK67(X78,happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),sK13)))))) ) | $spl248),
  inference(cnf_transformation,[],[f44743_D])).
fof(f44743_D,plain,(
  ( ! [X78] : (~p(s(bool,X78)) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,sK67(X78,happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),sK13)))))) ) <=> ~$spl248),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl248])])).
fof(f44515,plain,(
  ( ! [X78,X77] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,X77)))) | ~p(s(bool,X78)) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,sK67(X78,happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),sK13)))))) )),
  inference(resolution,[],[f795,f740])).
fof(f44742,plain,(
  ( ! [X76] : (~p(s(bool,X76)) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,sK66(X76,happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),sK13)))))) ) | $spl246),
  inference(cnf_transformation,[],[f44742_D])).
fof(f44742_D,plain,(
  ( ! [X76] : (~p(s(bool,X76)) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,sK66(X76,happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),sK13)))))) ) <=> ~$spl246),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl246])])).
fof(f44514,plain,(
  ( ! [X76,X75] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,X75)))) | ~p(s(bool,X76)) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,sK66(X76,happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),sK13)))))) )),
  inference(resolution,[],[f795,f738])).
fof(f44741,plain,(
  ( ! [X74] : (~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,sK61(happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),X74,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),X74),s(sK13,sK60(happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),X74,sK13)))))) ) | $spl244),
  inference(cnf_transformation,[],[f44741_D])).
fof(f44741_D,plain,(
  ( ! [X74] : (~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,sK61(happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),X74,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),X74),s(sK13,sK60(happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),X74,sK13)))))) ) <=> ~$spl244),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl244])])).
fof(f44513,plain,(
  ( ! [X74,X73] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,X73)))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,sK61(happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),X74,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),X74),s(sK13,sK60(happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),X74,sK13)))))) )),
  inference(resolution,[],[f795,f729])).
fof(f729,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK61(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK60(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f456])).
fof(f456,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK59(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK59(X2,X1,X0)))))) | (! [X7] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X7)))) & ! [X6] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK60(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK61(X2,X1,X0))))) | ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK59,sK60,sK61])],[f455])).
fof(f455,plain,(
  ! [X0,X1,X2] : ((? [X8] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X8)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X8))))) | (! [X7] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X7)))) & ! [X6] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))))) & (? [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(rectify,[],[f454])).
fof(f454,plain,(
  ! [X0,X1,X2] : ((? [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (? [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(flattening,[],[f453])).
fof(f453,plain,(
  ! [X0,X1,X2] : ((? [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & ((? [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f223])).
fof(f223,plain,(
  ! [X0,X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f222])).
fof(f222,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f48])).
fof(f48,axiom,(
  ! [X6] : ! [X16,X12] : (! [X9] : (p(s(bool,happ(s(fun(X6,bool),X16),s(X6,X9)))) & p(s(bool,happ(s(fun(X6,bool),X12),s(X6,X9))))) <=> (! [X9] : p(s(bool,happ(s(fun(X6,bool),X16),s(X6,X9)))) & ! [X9] : p(s(bool,happ(s(fun(X6,bool),X12),s(X6,X9))))))),
  file('thm.relation.RTC_lifts_monotonicities','thm.bool.FORALL_AND_THM')).
fof(f44740,plain,(
  ( ! [X72] : (~p(s(bool,happ(s(fun(sK13,bool),X72),s(sK13,sK61(X72,happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,sK60(X72,happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),sK13)))))) ) | $spl242),
  inference(cnf_transformation,[],[f44740_D])).
fof(f44740_D,plain,(
  ( ! [X72] : (~p(s(bool,happ(s(fun(sK13,bool),X72),s(sK13,sK61(X72,happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,sK60(X72,happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),sK13)))))) ) <=> ~$spl242),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl242])])).
fof(f44512,plain,(
  ( ! [X72,X71] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,X71)))) | ~p(s(bool,happ(s(fun(sK13,bool),X72),s(sK13,sK61(X72,happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,sK60(X72,happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),sK13)))))) )),
  inference(resolution,[],[f795,f728])).
fof(f728,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK61(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK60(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f456])).
fof(f44739,plain,(
  ( ! [X70] : (~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,sK59(happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),X70,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),X70),s(sK13,sK59(happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),X70,sK13)))))) ) | $spl240),
  inference(cnf_transformation,[],[f44739_D])).
fof(f44739_D,plain,(
  ( ! [X70] : (~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,sK59(happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),X70,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),X70),s(sK13,sK59(happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),X70,sK13)))))) ) <=> ~$spl240),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl240])])).
fof(f44511,plain,(
  ( ! [X70,X69] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,X69)))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,sK59(happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),X70,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),X70),s(sK13,sK59(happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),X70,sK13)))))) )),
  inference(resolution,[],[f795,f727])).
fof(f727,plain,(
  ( ! [X6,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK59(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK59(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f456])).
fof(f44738,plain,(
  ( ! [X68] : (~p(s(bool,happ(s(fun(sK13,bool),X68),s(sK13,sK59(X68,happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,sK59(X68,happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),sK13)))))) ) | $spl238),
  inference(cnf_transformation,[],[f44738_D])).
fof(f44738_D,plain,(
  ( ! [X68] : (~p(s(bool,happ(s(fun(sK13,bool),X68),s(sK13,sK59(X68,happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,sK59(X68,happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),sK13)))))) ) <=> ~$spl238),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl238])])).
fof(f44510,plain,(
  ( ! [X68,X67] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,X67)))) | ~p(s(bool,happ(s(fun(sK13,bool),X68),s(sK13,sK59(X68,happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,sK59(X68,happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),sK13)))))) )),
  inference(resolution,[],[f795,f726])).
fof(f726,plain,(
  ( ! [X2,X0,X7,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X7)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK59(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK59(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f456])).
fof(f44737,plain,(
  ( ! [X66] : (p(s(bool,X66)) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,sK58(happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),X66,sK13)))))) ) | $spl236),
  inference(cnf_transformation,[],[f44737_D])).
fof(f44737_D,plain,(
  ( ! [X66] : (p(s(bool,X66)) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,sK58(happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),X66,sK13)))))) ) <=> ~$spl236),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl236])])).
fof(f44509,plain,(
  ( ! [X66,X65] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,X65)))) | p(s(bool,X66)) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,sK58(happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),X66,sK13)))))) )),
  inference(resolution,[],[f795,f725])).
fof(f725,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK58(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f452])).
fof(f452,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK57(X2,X1,X0)))))) | p(s(bool,X1)) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5))))) & ((~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK58(X2,X1,X0)))))) | ! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK57,sK58])],[f451])).
fof(f451,plain,(
  ! [X0,X1,X2] : ((? [X6] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | p(s(bool,X1)) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5))))) & ((~p(s(bool,X1)) & ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(rectify,[],[f450])).
fof(f450,plain,(
  ! [X0,X1,X2] : ((? [X3] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | p(s(bool,X1)) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) & ((~p(s(bool,X1)) & ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(flattening,[],[f449])).
fof(f449,plain,(
  ! [X0,X1,X2] : ((? [X3] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (p(s(bool,X1)) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & ((~p(s(bool,X1)) & ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f221])).
fof(f221,plain,(
  ! [X0,X1,X2] : (! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (p(s(bool,X1)) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f220])).
fof(f220,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (p(s(bool,X1)) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f58])).
fof(f58,axiom,(
  ! [X6] : ! [X16,X12] : (! [X9] : (p(s(bool,X16)) | p(s(bool,happ(s(fun(X6,bool),X12),s(X6,X9))))) <=> (p(s(bool,X16)) | ! [X9] : p(s(bool,happ(s(fun(X6,bool),X12),s(X6,X9))))))),
  file('thm.relation.RTC_lifts_monotonicities','thm.bool.RIGHT_FORALL_OR_THM')).
fof(f44736,plain,(
  ( ! [X64] : (p(s(bool,X64)) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,sK57(happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),X64,sK13)))))) ) | $spl234),
  inference(cnf_transformation,[],[f44736_D])).
fof(f44736_D,plain,(
  ( ! [X64] : (p(s(bool,X64)) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,sK57(happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),X64,sK13)))))) ) <=> ~$spl234),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl234])])).
fof(f44508,plain,(
  ( ! [X64,X63] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,X63)))) | p(s(bool,X64)) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,sK57(happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)),X64,sK13)))))) )),
  inference(resolution,[],[f795,f723])).
fof(f723,plain,(
  ( ! [X2,X0,X5,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) | p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK57(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f452])).
fof(f44734,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,happ(s(fun(fun(sK13,bool),sK13),sK49(sK13)),s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)))))))) | $spl231),
  inference(cnf_transformation,[],[f44734_D])).
fof(f44734_D,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,happ(s(fun(fun(sK13,bool),sK13),sK49(sK13)),s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17)))))))) <=> ~$spl231),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl231])])).
fof(f44507,plain,(
  ( ! [X62] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,X62)))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,happ(s(fun(fun(sK13,bool),sK13),sK49(sK13)),s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))))))))) )),
  inference(resolution,[],[f795,f709])).
fof(f709,plain,(
  ( ! [X2,X0,X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,happ(s(fun(fun(X0,bool),X0),sK49(X0)),s(fun(X0,bool),X2))))))) )),
  inference(cnf_transformation,[],[f436])).
fof(f436,plain,(
  ! [X0] : ! [X2] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,happ(s(fun(fun(X0,bool),X0),sK49(X0)),s(fun(X0,bool),X2)))))) | ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK49])],[f306])).
fof(f306,plain,(
  ! [X0] : ? [X1] : ! [X2] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,happ(s(fun(fun(X0,bool),X0),X1),s(fun(X0,bool),X2)))))) | ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))),
  inference(ennf_transformation,[],[f213])).
fof(f213,plain,(
  ! [X0] : ? [X1] : ! [X2] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,happ(s(fun(fun(X0,bool),X0),X1),s(fun(X0,bool),X2)))))) => ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X6] : ? [X7] : ! [X19] : (p(s(bool,happ(s(fun(X6,bool),X19),s(X6,happ(s(fun(fun(X6,bool),X6),X7),s(fun(X6,bool),X19)))))) => ! [X9] : p(s(bool,happ(s(fun(X6,bool),X19),s(X6,X9)))))),
  file('thm.relation.RTC_lifts_monotonicities','thm.normalForms.normalForms_2_prove_ltb')).
fof(f44730,plain,(
  ( ! [X61] : (~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK45(sK15,sK13,sK13)))),s(sK13,X61))))) ) | $spl226),
  inference(cnf_transformation,[],[f44730_D])).
fof(f44730_D,plain,(
  ( ! [X61] : (~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK45(sK15,sK13,sK13)))),s(sK13,X61))))) ) <=> ~$spl226),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl226])])).
fof(f44732,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,happ(s(fun(sK13,sK13),sK46(sK15,sK13,sK13)),s(sK13,sK17)))))) | $spl228),
  inference(cnf_transformation,[],[f44732_D])).
fof(f44732_D,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,happ(s(fun(sK13,sK13),sK46(sK15,sK13,sK13)),s(sK13,sK17)))))) <=> ~$spl228),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl228])])).
fof(f44506,plain,(
  ( ! [X61] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,happ(s(fun(sK13,sK13),sK46(sK15,sK13,sK13)),s(sK13,sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK45(sK15,sK13,sK13)))),s(sK13,X61))))) )),
  inference(resolution,[],[f795,f671])).
fof(f671,plain,(
  ( ! [X2,X0,X10,X8,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X8))),s(X0,happ(s(fun(X1,X0),sK46(X2,X1,X0)),s(X1,X8)))))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,sK45(X2,X1,X0)))),s(X0,X10))))) )),
  inference(cnf_transformation,[],[f423])).
fof(f423,plain,(
  ! [X0,X1,X2] : ((! [X10] : ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,sK45(X2,X1,X0)))),s(X0,X10)))) | ! [X8] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X8))),s(X0,happ(s(fun(X1,X0),sK46(X2,X1,X0)),s(X1,X8))))))) & (! [X5] : ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,sK47(X5,X2,X1,X0)))),s(X0,happ(s(fun(X1,X0),X5),s(X1,sK47(X5,X2,X1,X0))))))) | ! [X3] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X3))),s(X0,sK48(X3,X2,X1,X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK45,sK46,sK47,sK48])],[f422])).
fof(f422,plain,(
  ! [X0,X1,X2] : ((? [X9] : ! [X10] : ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X9))),s(X0,X10)))) | ? [X7] : ! [X8] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X8))),s(X0,happ(s(fun(X1,X0),X7),s(X1,X8))))))) & (! [X5] : ? [X6] : ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X6))),s(X0,happ(s(fun(X1,X0),X5),s(X1,X6)))))) | ! [X3] : ? [X4] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X3))),s(X0,X4))))))),
  inference(rectify,[],[f421])).
fof(f421,plain,(
  ! [X0,X1,X2] : ((? [X3] : ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X3))),s(X0,X4)))) | ? [X5] : ! [X6] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X6))),s(X0,happ(s(fun(X1,X0),X5),s(X1,X6))))))) & (! [X5] : ? [X6] : ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X6))),s(X0,happ(s(fun(X1,X0),X5),s(X1,X6)))))) | ! [X3] : ? [X4] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X3))),s(X0,X4))))))),
  inference(nnf_transformation,[],[f196])).
fof(f196,plain,(
  ! [X0,X1,X2] : (! [X3] : ? [X4] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X3))),s(X0,X4)))) <=> ? [X5] : ! [X6] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X6))),s(X0,happ(s(fun(X1,X0),X5),s(X1,X6)))))))),
  inference(flattening,[],[f195])).
fof(f195,plain,(
  ! [X0,X1] : ! [X2] : (! [X3] : ? [X4] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X3))),s(X0,X4)))) <=> ? [X5] : ! [X6] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X6))),s(X0,happ(s(fun(X1,X0),X5),s(X1,X6)))))))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X11,X6] : ! [X16] : (! [X9] : ? [X10] : p(s(bool,happ(s(fun(X11,bool),happ(s(fun(X6,fun(X11,bool)),X16),s(X6,X9))),s(X11,X10)))) <=> ? [X7] : ! [X9] : p(s(bool,happ(s(fun(X11,bool),happ(s(fun(X6,fun(X11,bool)),X16),s(X6,X9))),s(X11,happ(s(fun(X6,X11),X7),s(X6,X9)))))))),
  file('thm.relation.RTC_lifts_monotonicities','thm.bool.SKOLEM_THM')).
fof(f44727,plain,(
  ( ! [X60] : (~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK47(X60,sK15,sK13,sK13)))),s(sK13,happ(s(fun(sK13,sK13),X60),s(sK13,sK47(X60,sK15,sK13,sK13)))))))) ) | $spl222),
  inference(cnf_transformation,[],[f44727_D])).
fof(f44727_D,plain,(
  ( ! [X60] : (~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK47(X60,sK15,sK13,sK13)))),s(sK13,happ(s(fun(sK13,sK13),X60),s(sK13,sK47(X60,sK15,sK13,sK13)))))))) ) <=> ~$spl222),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl222])])).
fof(f44729,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK48(sK17,sK15,sK13,sK13))))) | $spl224),
  inference(cnf_transformation,[],[f44729_D])).
fof(f44729_D,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK48(sK17,sK15,sK13,sK13))))) <=> ~$spl224),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl224])])).
fof(f44505,plain,(
  ( ! [X60] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK48(sK17,sK15,sK13,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK47(X60,sK15,sK13,sK13)))),s(sK13,happ(s(fun(sK13,sK13),X60),s(sK13,sK47(X60,sK15,sK13,sK13)))))))) )),
  inference(resolution,[],[f795,f672])).
fof(f672,plain,(
  ( ! [X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X3))),s(X0,sK48(X3,X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,sK47(X5,X2,X1,X0)))),s(X0,happ(s(fun(X1,X0),X5),s(X1,sK47(X5,X2,X1,X0)))))))) )),
  inference(cnf_transformation,[],[f423])).
fof(f44320,plain,(
  ( ! [X84,X82] : (p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X82),s(sK13,sK17)))) | ~p(s(bool,X84)) | s(bool,X84) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X82))),s(sK13,sK16)))) ) | $spl220),
  inference(cnf_transformation,[],[f44320_D])).
fof(f44320_D,plain,(
  ( ! [X84,X82] : (p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X82),s(sK13,sK17)))) | ~p(s(bool,X84)) | s(bool,X84) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X82))),s(sK13,sK16)))) ) <=> ~$spl220),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl220])])).
fof(f44236,plain,(
  ( ! [X397,X399,X398] : (p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X397),s(sK13,sK17)))) | p(s(bool,X398)) | ~p(s(bool,X399)) | s(bool,X399) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X397))),s(sK13,sK16)))) )),
  inference(resolution,[],[f794,f783])).
fof(f794,plain,(
  ( ! [X3] : (~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X3))),s(sK13,sK16)))) | p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X3),s(sK13,sK17))))) )),
  inference(resolution,[],[f499,f504])).
fof(f504,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f267])).
fof(f267,plain,(
  ! [X0,X1,X2,X3,X4] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4)))))),
  inference(flattening,[],[f266])).
fof(f266,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4))))) | p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4)))))),
  inference(ennf_transformation,[],[f142])).
fof(f142,plain,(
  ! [X0,X1,X2,X3,X4] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4))))) => p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4)))))),
  inference(flattening,[],[f141])).
fof(f141,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4))))) => p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4)))))),
  inference(rectify,[],[f86])).
fof(f86,axiom,(
  ! [X6] : ! [X8,X9,X10,X32] : ((p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,X9))),s(X6,X10)))) & p(s(bool,'const.relation.RTC_3'(s(fun(X6,fun(X6,bool)),X8),s(X6,X10),s(X6,X32))))) => p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X8),s(X6,X9),s(X6,X32)))))),
  file('thm.relation.RTC_lifts_monotonicities','thm.relation.EXTEND_RTC_TC')).
fof(f44138,plain,(
  ( ! [X111,X112,X110,X113] : (p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X110),s(sK13,sK17)))) | ~p(s(bool,X111)) | ~p(s(bool,X112)) | s(bool,X112) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X110))),s(sK13,sK16))) | s(bool,X111) != s(bool,X113)) )),
  inference(resolution,[],[f794,f558])).
fof(f44136,plain,(
  ( ! [X103,X105,X102,X104] : (p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X102),s(sK13,sK17)))) | p(s(bool,X103)) | ~p(s(bool,X104)) | ~p(s(bool,X105)) | s(bool,X104) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X102))),s(sK13,sK16)))) )),
  inference(resolution,[],[f794,f557])).
fof(f44134,plain,(
  ( ! [X94,X97,X95,X96] : (p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X94),s(sK13,sK17)))) | p(s(bool,X95)) | ~p(s(bool,X96)) | s(bool,X95) != s(bool,X97) | s(bool,X96) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X94))),s(sK13,sK16)))) )),
  inference(resolution,[],[f794,f556])).
fof(f44133,plain,(
  ( ! [X92,X90,X93,X91] : (p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X90),s(sK13,sK17)))) | ~p(s(bool,X91)) | ~p(s(bool,X92)) | s(bool,X92) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X90))),s(sK13,sK16))) | s(bool,X91) != s(bool,X93)) )),
  inference(resolution,[],[f794,f554])).
fof(f44131,plain,(
  ( ! [X85,X83,X84,X82] : (p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X82),s(sK13,sK17)))) | p(s(bool,X83)) | ~p(s(bool,X84)) | s(bool,X83) != s(bool,X85) | s(bool,X84) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X82))),s(sK13,sK16)))) )),
  inference(resolution,[],[f794,f552])).
fof(f44044,plain,(
  ( ! [X84,X82] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X82),s(sK13,sK17)))) | ~p(s(bool,X84)) | s(bool,X84) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X82))),s(sK13,sK16)))) ) | $spl218),
  inference(cnf_transformation,[],[f44044_D])).
fof(f44044_D,plain,(
  ( ! [X84,X82] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X82),s(sK13,sK17)))) | ~p(s(bool,X84)) | s(bool,X84) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X82))),s(sK13,sK16)))) ) <=> ~$spl218),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl218])])).
fof(f43960,plain,(
  ( ! [X397,X399,X398] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X397),s(sK13,sK17)))) | p(s(bool,X398)) | ~p(s(bool,X399)) | s(bool,X399) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X397))),s(sK13,sK16)))) )),
  inference(resolution,[],[f793,f783])).
fof(f793,plain,(
  ( ! [X2] : (~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X2))),s(sK13,sK16)))) | p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X2),s(sK13,sK17))))) )),
  inference(resolution,[],[f499,f503])).
fof(f503,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f265])).
fof(f265,plain,(
  ! [X0,X1,X2,X3,X4] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4)))))),
  inference(flattening,[],[f264])).
fof(f264,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4))))) | p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4)))))),
  inference(ennf_transformation,[],[f140])).
fof(f140,plain,(
  ! [X0,X1,X2,X3,X4] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4))))) => p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4)))))),
  inference(flattening,[],[f139])).
fof(f139,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4))))) => p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4)))))),
  inference(rectify,[],[f66])).
fof(f66,axiom,(
  ! [X6] : ! [X8,X9,X10,X32] : ((p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,X9))),s(X6,X10)))) & p(s(bool,'const.relation.RTC_3'(s(fun(X6,fun(X6,bool)),X8),s(X6,X10),s(X6,X32))))) => p(s(bool,'const.relation.RTC_3'(s(fun(X6,fun(X6,bool)),X8),s(X6,X9),s(X6,X32)))))),
  file('thm.relation.RTC_lifts_monotonicities','thm.relation.RTC_RULES|split|1')).
fof(f43862,plain,(
  ( ! [X111,X112,X110,X113] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X110),s(sK13,sK17)))) | ~p(s(bool,X111)) | ~p(s(bool,X112)) | s(bool,X112) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X110))),s(sK13,sK16))) | s(bool,X111) != s(bool,X113)) )),
  inference(resolution,[],[f793,f558])).
fof(f43860,plain,(
  ( ! [X103,X105,X102,X104] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X102),s(sK13,sK17)))) | p(s(bool,X103)) | ~p(s(bool,X104)) | ~p(s(bool,X105)) | s(bool,X104) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X102))),s(sK13,sK16)))) )),
  inference(resolution,[],[f793,f557])).
fof(f43858,plain,(
  ( ! [X94,X97,X95,X96] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X94),s(sK13,sK17)))) | p(s(bool,X95)) | ~p(s(bool,X96)) | s(bool,X95) != s(bool,X97) | s(bool,X96) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X94))),s(sK13,sK16)))) )),
  inference(resolution,[],[f793,f556])).
fof(f43857,plain,(
  ( ! [X92,X90,X93,X91] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X90),s(sK13,sK17)))) | ~p(s(bool,X91)) | ~p(s(bool,X92)) | s(bool,X92) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X90))),s(sK13,sK16))) | s(bool,X91) != s(bool,X93)) )),
  inference(resolution,[],[f793,f554])).
fof(f43855,plain,(
  ( ! [X85,X83,X84,X82] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X82),s(sK13,sK17)))) | p(s(bool,X83)) | ~p(s(bool,X84)) | s(bool,X83) != s(bool,X85) | s(bool,X84) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X82))),s(sK13,sK16)))) )),
  inference(resolution,[],[f793,f552])).
fof(f25006,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16)))) ) | $spl142),
  inference(cnf_transformation,[],[f25006_D])).
fof(f25006_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16)))) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f42672,plain,(
  ( ! [X212,X213] : (p(s(bool,X212)) | p(s(bool,X213)) | s(bool,X213) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16)))) ) | ($spl154 | $spl166)),
  inference(resolution,[],[f40013,f783])).
fof(f40013,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16)))) | ($spl154 | $spl166)),
  inference(forward_demodulation,[],[f28802,f28790])).
fof(f42574,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16))) | s(bool,X22) != s(bool,X23)) ) | ($spl154 | $spl166)),
  inference(resolution,[],[f40013,f558])).
fof(f42571,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16)))) ) | ($spl154 | $spl166)),
  inference(resolution,[],[f40013,f557])).
fof(f42570,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16)))) ) | ($spl154 | $spl166)),
  inference(resolution,[],[f40013,f556])).
fof(f42569,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16))) | s(bool,X7) != s(bool,X8)) ) | ($spl154 | $spl166)),
  inference(resolution,[],[f40013,f554])).
fof(f42567,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16)))) ) | ($spl154 | $spl166)),
  inference(resolution,[],[f40013,f552])).
fof(f42338,plain,(
  ( ! [X233,X234] : (p(s(bool,X233)) | p(s(bool,X234)) | s(bool,X234) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK16,sK15,sK13)))),s(sK13,sK17)))) ) | $spl8),
  inference(resolution,[],[f1006,f783])).
fof(f42240,plain,(
  ( ! [X43,X44,X42] : (p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X42) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK16,sK15,sK13)))),s(sK13,sK17))) | s(bool,X43) != s(bool,X44)) ) | $spl8),
  inference(resolution,[],[f1006,f558])).
fof(f42237,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X33)) | ~p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X35) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK16,sK15,sK13)))),s(sK13,sK17)))) ) | $spl8),
  inference(resolution,[],[f1006,f557])).
fof(f42236,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | p(s(bool,X31)) | s(bool,X30) != s(bool,X32) | s(bool,X31) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK16,sK15,sK13)))),s(sK13,sK17)))) ) | $spl8),
  inference(resolution,[],[f1006,f556])).
fof(f42235,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK16,sK15,sK13)))),s(sK13,sK17))) | s(bool,X28) != s(bool,X29)) ) | $spl8),
  inference(resolution,[],[f1006,f554])).
fof(f42233,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK16,sK15,sK13)))),s(sK13,sK17)))) ) | $spl8),
  inference(resolution,[],[f1006,f552])).
fof(f41994,plain,(
  ( ! [X241,X242] : (p(s(bool,X241)) | p(s(bool,X242)) | s(bool,X242) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) | $spl6),
  inference(resolution,[],[f1004,f783])).
fof(f41896,plain,(
  ( ! [X52,X50,X51] : (p(s(bool,X50)) | ~p(s(bool,X51)) | s(bool,X50) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13)))) | s(bool,X51) != s(bool,X52)) ) | $spl6),
  inference(resolution,[],[f1004,f558])).
fof(f41893,plain,(
  ( ! [X43,X41,X42] : (p(s(bool,X41)) | ~p(s(bool,X42)) | p(s(bool,X43)) | s(bool,X43) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) | $spl6),
  inference(resolution,[],[f1004,f557])).
fof(f41892,plain,(
  ( ! [X39,X38,X40] : (p(s(bool,X38)) | p(s(bool,X39)) | s(bool,X38) != s(bool,X40) | s(bool,X39) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) | $spl6),
  inference(resolution,[],[f1004,f556])).
fof(f41891,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13)))) | s(bool,X36) != s(bool,X37)) ) | $spl6),
  inference(resolution,[],[f1004,f554])).
fof(f41889,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) | $spl6),
  inference(resolution,[],[f1004,f552])).
fof(f41873,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) | p(s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) | $spl6),
  inference(resolution,[],[f1004,f531])).
fof(f14712,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16)))) ) | $spl102),
  inference(cnf_transformation,[],[f14712_D])).
fof(f14712_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16)))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f41767,plain,(
  ( ! [X212,X213] : (p(s(bool,X212)) | p(s(bool,X213)) | s(bool,X213) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16)))) ) | ($spl2 | $spl154)),
  inference(resolution,[],[f39999,f783])).
fof(f39999,plain,(
  p(s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16)))) | ($spl2 | $spl154)),
  inference(forward_demodulation,[],[f28605,f28790])).
fof(f28605,plain,(
  p(s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK27(sK17,sK16,sK15,sK13))))) | $spl2),
  inference(resolution,[],[f24300,f534])).
fof(f24300,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK27(sK17,sK16,sK15,sK13))))) | $spl2),
  inference(resolution,[],[f1000,f529])).
fof(f41669,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16))) | s(bool,X22) != s(bool,X23)) ) | ($spl2 | $spl154)),
  inference(resolution,[],[f39999,f558])).
fof(f41666,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16)))) ) | ($spl2 | $spl154)),
  inference(resolution,[],[f39999,f557])).
fof(f41665,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16)))) ) | ($spl2 | $spl154)),
  inference(resolution,[],[f39999,f556])).
fof(f41664,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16))) | s(bool,X7) != s(bool,X8)) ) | ($spl2 | $spl154)),
  inference(resolution,[],[f39999,f554])).
fof(f41662,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16)))) ) | ($spl2 | $spl154)),
  inference(resolution,[],[f39999,f552])).
fof(f21042,plain,(
  ( ! [X15] : (p(s(bool,X15)) | s(bool,X15) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16)))) ) | $spl130),
  inference(cnf_transformation,[],[f21042_D])).
fof(f21042_D,plain,(
  ( ! [X15] : (p(s(bool,X15)) | s(bool,X15) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16)))) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f41606,plain,(
  ( ! [X227,X226] : (p(s(bool,X226)) | p(s(bool,X227)) | s(bool,X227) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16)))) ) | ($spl2 | $spl154)),
  inference(resolution,[],[f29759,f783])).
fof(f29759,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16)))) | ($spl2 | $spl154)),
  inference(resolution,[],[f28804,f764])).
fof(f28804,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16))),s(sK13,sK16)))) | ($spl2 | $spl154)),
  inference(backward_demodulation,[],[f28790,f1000])).
fof(f41508,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16))) | s(bool,X36) != s(bool,X37)) ) | ($spl2 | $spl154)),
  inference(resolution,[],[f29759,f558])).
fof(f41505,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | p(s(bool,X28)) | s(bool,X28) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16)))) ) | ($spl2 | $spl154)),
  inference(resolution,[],[f29759,f557])).
fof(f41504,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X23) != s(bool,X25) | s(bool,X24) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16)))) ) | ($spl2 | $spl154)),
  inference(resolution,[],[f29759,f556])).
fof(f41503,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X20) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16))) | s(bool,X21) != s(bool,X22)) ) | ($spl2 | $spl154)),
  inference(resolution,[],[f29759,f554])).
fof(f41501,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16)))) ) | ($spl2 | $spl154)),
  inference(resolution,[],[f29759,f552])).
fof(f41472,plain,(
  ( ! [X250] : (~p(s(bool,X250)) | s(bool,X250) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17)))))) ) | $spl214),
  inference(cnf_transformation,[],[f41472_D])).
fof(f41472_D,plain,(
  ( ! [X250] : (~p(s(bool,X250)) | s(bool,X250) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17)))))) ) <=> ~$spl214),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl214])])).
fof(f41253,plain,(
  ( ! [X39,X41,X40] : (p(s(bool,X39)) | ~p(s(bool,X40)) | s(bool,X39) != s(bool,X41) | s(bool,X40) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17)))))) )),
  inference(resolution,[],[f1030,f552])).
fof(f1030,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17))))))),
  inference(resolution,[],[f500,f529])).
fof(f41255,plain,(
  ( ! [X47,X45,X46] : (~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X46) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17))))) | s(bool,X45) != s(bool,X47)) )),
  inference(resolution,[],[f1030,f554])).
fof(f41256,plain,(
  ( ! [X50,X48,X49] : (p(s(bool,X48)) | ~p(s(bool,X49)) | s(bool,X48) != s(bool,X50) | s(bool,X49) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17)))))) )),
  inference(resolution,[],[f1030,f556])).
fof(f41258,plain,(
  ( ! [X54,X56,X55] : (p(s(bool,X54)) | ~p(s(bool,X55)) | ~p(s(bool,X56)) | s(bool,X55) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17)))))) )),
  inference(resolution,[],[f1030,f557])).
fof(f41260,plain,(
  ( ! [X61,X62,X60] : (~p(s(bool,X60)) | ~p(s(bool,X61)) | s(bool,X61) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17))))) | s(bool,X60) != s(bool,X62)) )),
  inference(resolution,[],[f1030,f558])).
fof(f41358,plain,(
  ( ! [X249,X250] : (p(s(bool,X249)) | ~p(s(bool,X250)) | s(bool,X250) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17)))))) )),
  inference(resolution,[],[f1030,f783])).
fof(f41471,plain,(
  $false | ($spl8 | $spl68)),
  inference(subsumption_resolution,[],[f41210,f24502])).
fof(f24502,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16))),s(sK13,sK17)))) | ($spl8 | $spl68)),
  inference(forward_demodulation,[],[f1006,f11239])).
fof(f41210,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16))),s(sK13,sK17))))),
  inference(resolution,[],[f1030,f498])).
fof(f39998,plain,(
  ( ! [X15] : (p(s(bool,X15)) | s(bool,X15) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK21(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl212),
  inference(cnf_transformation,[],[f39998_D])).
fof(f39998_D,plain,(
  ( ! [X15] : (p(s(bool,X15)) | s(bool,X15) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK21(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) <=> ~$spl212),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl212])])).
fof(f39934,plain,(
  ( ! [X227,X226] : (p(s(bool,X226)) | p(s(bool,X227)) | s(bool,X227) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK21(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl114),
  inference(resolution,[],[f15153,f783])).
fof(f39836,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK21(sK17,sK16,sK15,sK13)),s(sK13,sK17))) | s(bool,X36) != s(bool,X37)) ) | $spl114),
  inference(resolution,[],[f15153,f558])).
fof(f39833,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | p(s(bool,X28)) | s(bool,X28) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK21(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl114),
  inference(resolution,[],[f15153,f557])).
fof(f39832,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X23) != s(bool,X25) | s(bool,X24) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK21(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl114),
  inference(resolution,[],[f15153,f556])).
fof(f39831,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X20) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK21(sK17,sK16,sK15,sK13)),s(sK13,sK17))) | s(bool,X21) != s(bool,X22)) ) | $spl114),
  inference(resolution,[],[f15153,f554])).
fof(f39829,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK21(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl114),
  inference(resolution,[],[f15153,f552])).
fof(f39374,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK21(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl210),
  inference(cnf_transformation,[],[f39374_D])).
fof(f39374_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK21(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) <=> ~$spl210),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl210])])).
fof(f39310,plain,(
  ( ! [X212,X213] : (p(s(bool,X212)) | p(s(bool,X213)) | s(bool,X213) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK21(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) )),
  inference(resolution,[],[f14958,f783])).
fof(f14958,plain,(
  p(s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK21(sK17,sK16,sK15,sK13)),s(sK13,sK17))))),
  inference(resolution,[],[f798,f534])).
fof(f798,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK21(sK17,sK16,sK15,sK13)),s(sK13,sK17))))),
  inference(resolution,[],[f499,f510])).
fof(f510,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,sK21(X3,X2,X1,X0)),s(X0,X3)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f346])).
fof(f39212,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK21(sK17,sK16,sK15,sK13)),s(sK13,sK17))) | s(bool,X22) != s(bool,X23)) )),
  inference(resolution,[],[f14958,f558])).
fof(f39209,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK21(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) )),
  inference(resolution,[],[f14958,f557])).
fof(f39208,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK21(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) )),
  inference(resolution,[],[f14958,f556])).
fof(f39207,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK21(sK17,sK16,sK15,sK13)),s(sK13,sK17))) | s(bool,X7) != s(bool,X8)) )),
  inference(resolution,[],[f14958,f554])).
fof(f39205,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK21(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) )),
  inference(resolution,[],[f14958,f552])).
fof(f39204,plain,(
  ( ! [X15] : (p(s(bool,X15)) | s(bool,X15) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK21(sK17,sK16,sK15,sK13))))) ) | $spl208),
  inference(cnf_transformation,[],[f39204_D])).
fof(f39204_D,plain,(
  ( ! [X15] : (p(s(bool,X15)) | s(bool,X15) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK21(sK17,sK16,sK15,sK13))))) ) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f39140,plain,(
  ( ! [X227,X226] : (p(s(bool,X226)) | p(s(bool,X227)) | s(bool,X227) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK21(sK17,sK16,sK15,sK13))))) ) | $spl42),
  inference(resolution,[],[f4652,f783])).
fof(f39042,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK21(sK17,sK16,sK15,sK13)))) | s(bool,X36) != s(bool,X37)) ) | $spl42),
  inference(resolution,[],[f4652,f558])).
fof(f39039,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | p(s(bool,X28)) | s(bool,X28) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK21(sK17,sK16,sK15,sK13))))) ) | $spl42),
  inference(resolution,[],[f4652,f557])).
fof(f39038,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X23) != s(bool,X25) | s(bool,X24) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK21(sK17,sK16,sK15,sK13))))) ) | $spl42),
  inference(resolution,[],[f4652,f556])).
fof(f39037,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X20) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK21(sK17,sK16,sK15,sK13)))) | s(bool,X21) != s(bool,X22)) ) | $spl42),
  inference(resolution,[],[f4652,f554])).
fof(f39035,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK21(sK17,sK16,sK15,sK13))))) ) | $spl42),
  inference(resolution,[],[f4652,f552])).
fof(f38580,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK21(sK17,sK16,sK15,sK13))))) ) | $spl206),
  inference(cnf_transformation,[],[f38580_D])).
fof(f38580_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK21(sK17,sK16,sK15,sK13))))) ) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f38516,plain,(
  ( ! [X212,X213] : (p(s(bool,X212)) | p(s(bool,X213)) | s(bool,X213) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK21(sK17,sK16,sK15,sK13))))) )),
  inference(resolution,[],[f4458,f783])).
fof(f4458,plain,(
  p(s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK21(sK17,sK16,sK15,sK13)))))),
  inference(resolution,[],[f797,f534])).
fof(f797,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK21(sK17,sK16,sK15,sK13)))))),
  inference(resolution,[],[f499,f509])).
fof(f509,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,sK21(X3,X2,X1,X0))))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f346])).
fof(f38418,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK21(sK17,sK16,sK15,sK13)))) | s(bool,X22) != s(bool,X23)) )),
  inference(resolution,[],[f4458,f558])).
fof(f38415,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK21(sK17,sK16,sK15,sK13))))) )),
  inference(resolution,[],[f4458,f557])).
fof(f38414,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK21(sK17,sK16,sK15,sK13))))) )),
  inference(resolution,[],[f4458,f556])).
fof(f38413,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK21(sK17,sK16,sK15,sK13)))) | s(bool,X7) != s(bool,X8)) )),
  inference(resolution,[],[f4458,f554])).
fof(f38411,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK21(sK17,sK16,sK15,sK13))))) )),
  inference(resolution,[],[f4458,f552])).
fof(f27703,plain,(
  ( ! [X15] : (p(s(bool,X15)) | s(bool,X15) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17),s(sK13,sK17)))) ) | $spl148),
  inference(cnf_transformation,[],[f27703_D])).
fof(f27703_D,plain,(
  ( ! [X15] : (p(s(bool,X15)) | s(bool,X15) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17),s(sK13,sK17)))) ) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f37405,plain,(
  ( ! [X227,X226] : (p(s(bool,X226)) | p(s(bool,X227)) | s(bool,X227) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17),s(sK13,sK17)))) ) | ($spl84 | $spl94)),
  inference(resolution,[],[f30808,f783])).
fof(f37307,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17),s(sK13,sK17))) | s(bool,X36) != s(bool,X37)) ) | ($spl84 | $spl94)),
  inference(resolution,[],[f30808,f558])).
fof(f37304,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | p(s(bool,X28)) | s(bool,X28) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17),s(sK13,sK17)))) ) | ($spl84 | $spl94)),
  inference(resolution,[],[f30808,f557])).
fof(f37303,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X23) != s(bool,X25) | s(bool,X24) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17),s(sK13,sK17)))) ) | ($spl84 | $spl94)),
  inference(resolution,[],[f30808,f556])).
fof(f37302,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X20) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17),s(sK13,sK17))) | s(bool,X21) != s(bool,X22)) ) | ($spl84 | $spl94)),
  inference(resolution,[],[f30808,f554])).
fof(f37300,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17),s(sK13,sK17)))) ) | ($spl84 | $spl94)),
  inference(resolution,[],[f30808,f552])).
fof(f15156,plain,(
  ( ! [X30] : (p(s(bool,X30)) | s(bool,X30) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK21(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl118),
  inference(cnf_transformation,[],[f15156_D])).
fof(f15156_D,plain,(
  ( ! [X30] : (p(s(bool,X30)) | s(bool,X30) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK21(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f37062,plain,(
  ( ! [X241,X242] : (p(s(bool,X241)) | p(s(bool,X242)) | s(bool,X242) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK21(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) )),
  inference(resolution,[],[f798,f783])).
fof(f36964,plain,(
  ( ! [X52,X50,X51] : (p(s(bool,X50)) | ~p(s(bool,X51)) | s(bool,X50) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK21(sK17,sK16,sK15,sK13)),s(sK13,sK17))) | s(bool,X51) != s(bool,X52)) )),
  inference(resolution,[],[f798,f558])).
fof(f36961,plain,(
  ( ! [X43,X41,X42] : (p(s(bool,X41)) | ~p(s(bool,X42)) | p(s(bool,X43)) | s(bool,X43) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK21(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) )),
  inference(resolution,[],[f798,f557])).
fof(f36960,plain,(
  ( ! [X39,X38,X40] : (p(s(bool,X38)) | p(s(bool,X39)) | s(bool,X38) != s(bool,X40) | s(bool,X39) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK21(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) )),
  inference(resolution,[],[f798,f556])).
fof(f36959,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK21(sK17,sK16,sK15,sK13)),s(sK13,sK17))) | s(bool,X36) != s(bool,X37)) )),
  inference(resolution,[],[f798,f554])).
fof(f36957,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK21(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) )),
  inference(resolution,[],[f798,f552])).
fof(f15155,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK21(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) | $spl116),
  inference(cnf_transformation,[],[f15155_D])).
fof(f15155_D,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK21(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f15153,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK21(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) | $spl114),
  inference(cnf_transformation,[],[f15153_D])).
fof(f15153_D,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK21(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f36941,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK21(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) | p(s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK21(sK17,sK16,sK15,sK13)),s(sK13,sK17))))),
  inference(resolution,[],[f798,f531])).
fof(f4655,plain,(
  ( ! [X30] : (p(s(bool,X30)) | s(bool,X30) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK21(sK17,sK16,sK15,sK13))))) ) | $spl46),
  inference(cnf_transformation,[],[f4655_D])).
fof(f4655_D,plain,(
  ( ! [X30] : (p(s(bool,X30)) | s(bool,X30) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK21(sK17,sK16,sK15,sK13))))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f36705,plain,(
  ( ! [X241,X242] : (p(s(bool,X241)) | p(s(bool,X242)) | s(bool,X242) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK21(sK17,sK16,sK15,sK13))))) )),
  inference(resolution,[],[f797,f783])).
fof(f36607,plain,(
  ( ! [X52,X50,X51] : (p(s(bool,X50)) | ~p(s(bool,X51)) | s(bool,X50) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK21(sK17,sK16,sK15,sK13)))) | s(bool,X51) != s(bool,X52)) )),
  inference(resolution,[],[f797,f558])).
fof(f36604,plain,(
  ( ! [X43,X41,X42] : (p(s(bool,X41)) | ~p(s(bool,X42)) | p(s(bool,X43)) | s(bool,X43) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK21(sK17,sK16,sK15,sK13))))) )),
  inference(resolution,[],[f797,f557])).
fof(f36603,plain,(
  ( ! [X39,X38,X40] : (p(s(bool,X38)) | p(s(bool,X39)) | s(bool,X38) != s(bool,X40) | s(bool,X39) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK21(sK17,sK16,sK15,sK13))))) )),
  inference(resolution,[],[f797,f556])).
fof(f36602,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK21(sK17,sK16,sK15,sK13)))) | s(bool,X36) != s(bool,X37)) )),
  inference(resolution,[],[f797,f554])).
fof(f36600,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK21(sK17,sK16,sK15,sK13))))) )),
  inference(resolution,[],[f797,f552])).
fof(f4654,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK21(sK17,sK16,sK15,sK13))))) | $spl44),
  inference(cnf_transformation,[],[f4654_D])).
fof(f4654_D,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK21(sK17,sK16,sK15,sK13))))) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f4652,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK21(sK17,sK16,sK15,sK13))))) | $spl42),
  inference(cnf_transformation,[],[f4652_D])).
fof(f4652_D,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK21(sK17,sK16,sK15,sK13))))) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f36584,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK21(sK17,sK16,sK15,sK13))))) | p(s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK21(sK17,sK16,sK15,sK13)))))),
  inference(resolution,[],[f797,f531])).
fof(f17393,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17),s(sK13,sK17)))) ) | $spl124),
  inference(cnf_transformation,[],[f17393_D])).
fof(f17393_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17),s(sK13,sK17)))) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f36478,plain,(
  ( ! [X212,X213] : (p(s(bool,X212)) | p(s(bool,X213)) | s(bool,X213) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17),s(sK13,sK17)))) ) | ($spl10 | $spl84)),
  inference(resolution,[],[f30543,f783])).
fof(f30543,plain,(
  p(s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17),s(sK13,sK17)))) | ($spl10 | $spl84)),
  inference(forward_demodulation,[],[f12644,f12828])).
fof(f36380,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17),s(sK13,sK17))) | s(bool,X22) != s(bool,X23)) ) | ($spl10 | $spl84)),
  inference(resolution,[],[f30543,f558])).
fof(f36377,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17),s(sK13,sK17)))) ) | ($spl10 | $spl84)),
  inference(resolution,[],[f30543,f557])).
fof(f36376,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17),s(sK13,sK17)))) ) | ($spl10 | $spl84)),
  inference(resolution,[],[f30543,f556])).
fof(f36375,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17),s(sK13,sK17))) | s(bool,X7) != s(bool,X8)) ) | ($spl10 | $spl84)),
  inference(resolution,[],[f30543,f554])).
fof(f36373,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17),s(sK13,sK17)))) ) | ($spl10 | $spl84)),
  inference(resolution,[],[f30543,f552])).
fof(f36284,plain,(
  ( ! [X241,X242] : (p(s(bool,X241)) | p(s(bool,X242)) | s(bool,X242) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK17,sK15,sK13)),s(sK13,sK17)))) ) | ($spl104 | $spl114)),
  inference(resolution,[],[f27527,f783])).
fof(f27527,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK17,sK15,sK13)),s(sK13,sK17)))) | ($spl104 | $spl114)),
  inference(resolution,[],[f27525,f507])).
fof(f27525,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17),s(sK13,sK17)))) | ($spl104 | $spl114)),
  inference(forward_demodulation,[],[f15153,f15143])).
fof(f36186,plain,(
  ( ! [X52,X50,X51] : (p(s(bool,X50)) | ~p(s(bool,X51)) | s(bool,X50) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK17,sK15,sK13)),s(sK13,sK17))) | s(bool,X51) != s(bool,X52)) ) | ($spl104 | $spl114)),
  inference(resolution,[],[f27527,f558])).
fof(f36183,plain,(
  ( ! [X43,X41,X42] : (p(s(bool,X41)) | ~p(s(bool,X42)) | p(s(bool,X43)) | s(bool,X43) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK17,sK15,sK13)),s(sK13,sK17)))) ) | ($spl104 | $spl114)),
  inference(resolution,[],[f27527,f557])).
fof(f36182,plain,(
  ( ! [X39,X38,X40] : (p(s(bool,X38)) | p(s(bool,X39)) | s(bool,X38) != s(bool,X40) | s(bool,X39) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK17,sK15,sK13)),s(sK13,sK17)))) ) | ($spl104 | $spl114)),
  inference(resolution,[],[f27527,f556])).
fof(f36181,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK17,sK15,sK13)),s(sK13,sK17))) | s(bool,X36) != s(bool,X37)) ) | ($spl104 | $spl114)),
  inference(resolution,[],[f27527,f554])).
fof(f36179,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK17,sK15,sK13)),s(sK13,sK17)))) ) | ($spl104 | $spl114)),
  inference(resolution,[],[f27527,f552])).
fof(f36163,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK17,sK15,sK13)),s(sK13,sK17)))) | p(s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK17,sK15,sK13)),s(sK13,sK17)))) | ($spl104 | $spl114)),
  inference(resolution,[],[f27527,f531])).
fof(f36160,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK17,sK15,sK13)),s(sK13,sK17)))) | s(sK13,sK17) = s(sK13,sK20(sK17,sK17,sK15,sK13)) | ($spl104 | $spl114)),
  inference(resolution,[],[f27527,f526])).
fof(f36158,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK20(sK17,sK17,sK15,sK13),sK15,sK13)))),s(sK13,sK17)))) | s(sK13,sK17) = s(sK13,sK20(sK17,sK17,sK15,sK13)) | ($spl104 | $spl114)),
  inference(resolution,[],[f27527,f523])).
fof(f36157,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK17,sK15,sK13)),s(sK13,sK28(sK17,sK20(sK17,sK17,sK15,sK13),sK15,sK13))))) | s(sK13,sK17) = s(sK13,sK20(sK17,sK17,sK15,sK13)) | ($spl104 | $spl114)),
  inference(resolution,[],[f27527,f522])).
fof(f36155,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK20(sK17,sK17,sK15,sK13),sK15,sK13)),s(sK13,sK17)))) | s(sK13,sK17) = s(sK13,sK20(sK17,sK17,sK15,sK13)) | ($spl104 | $spl114)),
  inference(resolution,[],[f27527,f519])).
fof(f36154,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK17,sK15,sK13)))),s(sK13,sK27(sK17,sK20(sK17,sK17,sK15,sK13),sK15,sK13))))) | s(sK13,sK17) = s(sK13,sK20(sK17,sK17,sK15,sK13)) | ($spl104 | $spl114)),
  inference(resolution,[],[f27527,f518])).
fof(f32373,plain,(
  ( ! [X30] : (p(s(bool,X30)) | s(bool,X30) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK16,sK16,sK15,sK13)),s(sK13,sK16)))) ) | $spl188),
  inference(cnf_transformation,[],[f32373_D])).
fof(f32373_D,plain,(
  ( ! [X30] : (p(s(bool,X30)) | s(bool,X30) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK16,sK16,sK15,sK13)),s(sK13,sK16)))) ) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f32295,plain,(
  ( ! [X241,X242] : (p(s(bool,X241)) | p(s(bool,X242)) | s(bool,X242) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK16,sK16,sK15,sK13)),s(sK13,sK16)))) ) | ($spl68 | $spl78)),
  inference(resolution,[],[f20866,f783])).
fof(f20866,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK16,sK16,sK15,sK13)),s(sK13,sK16)))) | ($spl68 | $spl78)),
  inference(resolution,[],[f20864,f507])).
fof(f20864,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16)))) | ($spl68 | $spl78)),
  inference(forward_demodulation,[],[f11249,f11239])).
fof(f32197,plain,(
  ( ! [X52,X50,X51] : (p(s(bool,X50)) | ~p(s(bool,X51)) | s(bool,X50) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK16,sK16,sK15,sK13)),s(sK13,sK16))) | s(bool,X51) != s(bool,X52)) ) | ($spl68 | $spl78)),
  inference(resolution,[],[f20866,f558])).
fof(f32194,plain,(
  ( ! [X43,X41,X42] : (p(s(bool,X41)) | ~p(s(bool,X42)) | p(s(bool,X43)) | s(bool,X43) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK16,sK16,sK15,sK13)),s(sK13,sK16)))) ) | ($spl68 | $spl78)),
  inference(resolution,[],[f20866,f557])).
fof(f32193,plain,(
  ( ! [X39,X38,X40] : (p(s(bool,X38)) | p(s(bool,X39)) | s(bool,X38) != s(bool,X40) | s(bool,X39) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK16,sK16,sK15,sK13)),s(sK13,sK16)))) ) | ($spl68 | $spl78)),
  inference(resolution,[],[f20866,f556])).
fof(f32192,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK16,sK16,sK15,sK13)),s(sK13,sK16))) | s(bool,X36) != s(bool,X37)) ) | ($spl68 | $spl78)),
  inference(resolution,[],[f20866,f554])).
fof(f32190,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK16,sK16,sK15,sK13)),s(sK13,sK16)))) ) | ($spl68 | $spl78)),
  inference(resolution,[],[f20866,f552])).
fof(f32372,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK16,sK16,sK15,sK13)),s(sK13,sK16)))) | $spl186),
  inference(cnf_transformation,[],[f32372_D])).
fof(f32372_D,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK16,sK16,sK15,sK13)),s(sK13,sK16)))) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f32370,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK16,sK16,sK15,sK13)),s(sK13,sK16)))) | $spl184),
  inference(cnf_transformation,[],[f32370_D])).
fof(f32370_D,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK16,sK16,sK15,sK13)),s(sK13,sK16)))) <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f32174,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK16,sK16,sK15,sK13)),s(sK13,sK16)))) | p(s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK16,sK16,sK15,sK13)),s(sK13,sK16)))) | ($spl68 | $spl78)),
  inference(resolution,[],[f20866,f531])).
fof(f32360,plain,(
  s(sK13,sK16) = s(sK13,sK20(sK16,sK16,sK15,sK13)) | $spl174),
  inference(cnf_transformation,[],[f32360_D])).
fof(f32360_D,plain,(
  s(sK13,sK16) = s(sK13,sK20(sK16,sK16,sK15,sK13)) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f32171,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK16,sK16,sK15,sK13)),s(sK13,sK16)))) | s(sK13,sK16) = s(sK13,sK20(sK16,sK16,sK15,sK13)) | ($spl68 | $spl78)),
  inference(resolution,[],[f20866,f526])).
fof(f32368,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK16,sK20(sK16,sK16,sK15,sK13),sK15,sK13)))),s(sK13,sK16)))) | $spl182),
  inference(cnf_transformation,[],[f32368_D])).
fof(f32368_D,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK16,sK20(sK16,sK16,sK15,sK13),sK15,sK13)))),s(sK13,sK16)))) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f32169,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK16,sK20(sK16,sK16,sK15,sK13),sK15,sK13)))),s(sK13,sK16)))) | s(sK13,sK16) = s(sK13,sK20(sK16,sK16,sK15,sK13)) | ($spl68 | $spl78)),
  inference(resolution,[],[f20866,f523])).
fof(f32366,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK16,sK16,sK15,sK13)),s(sK13,sK28(sK16,sK20(sK16,sK16,sK15,sK13),sK15,sK13))))) | $spl180),
  inference(cnf_transformation,[],[f32366_D])).
fof(f32366_D,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK16,sK16,sK15,sK13)),s(sK13,sK28(sK16,sK20(sK16,sK16,sK15,sK13),sK15,sK13))))) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f32168,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK16,sK16,sK15,sK13)),s(sK13,sK28(sK16,sK20(sK16,sK16,sK15,sK13),sK15,sK13))))) | s(sK13,sK16) = s(sK13,sK20(sK16,sK16,sK15,sK13)) | ($spl68 | $spl78)),
  inference(resolution,[],[f20866,f522])).
fof(f32364,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK16,sK20(sK16,sK16,sK15,sK13),sK15,sK13)),s(sK13,sK16)))) | $spl178),
  inference(cnf_transformation,[],[f32364_D])).
fof(f32364_D,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK16,sK20(sK16,sK16,sK15,sK13),sK15,sK13)),s(sK13,sK16)))) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f32166,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK16,sK20(sK16,sK16,sK15,sK13),sK15,sK13)),s(sK13,sK16)))) | s(sK13,sK16) = s(sK13,sK20(sK16,sK16,sK15,sK13)) | ($spl68 | $spl78)),
  inference(resolution,[],[f20866,f519])).
fof(f32362,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK16,sK16,sK15,sK13)))),s(sK13,sK27(sK16,sK20(sK16,sK16,sK15,sK13),sK15,sK13))))) | $spl176),
  inference(cnf_transformation,[],[f32362_D])).
fof(f32362_D,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK16,sK16,sK15,sK13)))),s(sK13,sK27(sK16,sK20(sK16,sK16,sK15,sK13),sK15,sK13))))) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f32165,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK16,sK16,sK15,sK13)))),s(sK13,sK27(sK16,sK20(sK16,sK16,sK15,sK13),sK15,sK13))))) | s(sK13,sK16) = s(sK13,sK20(sK16,sK16,sK15,sK13)) | ($spl68 | $spl78)),
  inference(resolution,[],[f20866,f518])).
fof(f29949,plain,(
  ( ! [X22] : (p(s(bool,X22)) | s(bool,X22) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16))),s(sK13,sK16)))) ) | $spl172),
  inference(cnf_transformation,[],[f29949_D])).
fof(f29949_D,plain,(
  ( ! [X22] : (p(s(bool,X22)) | s(bool,X22) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16))),s(sK13,sK16)))) ) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f29879,plain,(
  ( ! [X233,X234] : (p(s(bool,X233)) | p(s(bool,X234)) | s(bool,X234) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16))),s(sK13,sK16)))) ) | ($spl2 | $spl154)),
  inference(resolution,[],[f28804,f783])).
fof(f29781,plain,(
  ( ! [X43,X44,X42] : (p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X42) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16))),s(sK13,sK16))) | s(bool,X43) != s(bool,X44)) ) | ($spl2 | $spl154)),
  inference(resolution,[],[f28804,f558])).
fof(f29778,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X33)) | ~p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X35) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16))),s(sK13,sK16)))) ) | ($spl2 | $spl154)),
  inference(resolution,[],[f28804,f557])).
fof(f29777,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | p(s(bool,X31)) | s(bool,X30) != s(bool,X32) | s(bool,X31) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16))),s(sK13,sK16)))) ) | ($spl2 | $spl154)),
  inference(resolution,[],[f28804,f556])).
fof(f29776,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16))),s(sK13,sK16))) | s(bool,X28) != s(bool,X29)) ) | ($spl2 | $spl154)),
  inference(resolution,[],[f28804,f554])).
fof(f29774,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16))),s(sK13,sK16)))) ) | ($spl2 | $spl154)),
  inference(resolution,[],[f28804,f552])).
fof(f29632,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)))) ) | $spl170),
  inference(cnf_transformation,[],[f29632_D])).
fof(f29632_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)))) ) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f29568,plain,(
  ( ! [X212,X213] : (p(s(bool,X212)) | p(s(bool,X213)) | s(bool,X213) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)))) ) | ($spl60 | $spl154)),
  inference(resolution,[],[f28820,f783])).
fof(f28820,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)))) | ($spl60 | $spl154)),
  inference(backward_demodulation,[],[f28790,f6316])).
fof(f29470,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17))) | s(bool,X22) != s(bool,X23)) ) | ($spl60 | $spl154)),
  inference(resolution,[],[f28820,f558])).
fof(f29467,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)))) ) | ($spl60 | $spl154)),
  inference(resolution,[],[f28820,f557])).
fof(f29466,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)))) ) | ($spl60 | $spl154)),
  inference(resolution,[],[f28820,f556])).
fof(f29465,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17))) | s(bool,X7) != s(bool,X8)) ) | ($spl60 | $spl154)),
  inference(resolution,[],[f28820,f554])).
fof(f29463,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)))) ) | ($spl60 | $spl154)),
  inference(resolution,[],[f28820,f552])).
fof(f28803,plain,(
  ( ! [X30] : (p(s(bool,X30)) | s(bool,X30) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK27(sK17,sK16,sK15,sK13))))) ) | $spl168),
  inference(cnf_transformation,[],[f28803_D])).
fof(f28803_D,plain,(
  ( ! [X30] : (p(s(bool,X30)) | s(bool,X30) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK27(sK17,sK16,sK15,sK13))))) ) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f28725,plain,(
  ( ! [X241,X242] : (p(s(bool,X241)) | p(s(bool,X242)) | s(bool,X242) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK27(sK17,sK16,sK15,sK13))))) ) | $spl2),
  inference(resolution,[],[f24300,f783])).
fof(f28627,plain,(
  ( ! [X52,X50,X51] : (p(s(bool,X50)) | ~p(s(bool,X51)) | s(bool,X50) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK27(sK17,sK16,sK15,sK13)))) | s(bool,X51) != s(bool,X52)) ) | $spl2),
  inference(resolution,[],[f24300,f558])).
fof(f28624,plain,(
  ( ! [X43,X41,X42] : (p(s(bool,X41)) | ~p(s(bool,X42)) | p(s(bool,X43)) | s(bool,X43) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK27(sK17,sK16,sK15,sK13))))) ) | $spl2),
  inference(resolution,[],[f24300,f557])).
fof(f28623,plain,(
  ( ! [X39,X38,X40] : (p(s(bool,X38)) | p(s(bool,X39)) | s(bool,X38) != s(bool,X40) | s(bool,X39) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK27(sK17,sK16,sK15,sK13))))) ) | $spl2),
  inference(resolution,[],[f24300,f556])).
fof(f28622,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK27(sK17,sK16,sK15,sK13)))) | s(bool,X36) != s(bool,X37)) ) | $spl2),
  inference(resolution,[],[f24300,f554])).
fof(f28620,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK27(sK17,sK16,sK15,sK13))))) ) | $spl2),
  inference(resolution,[],[f24300,f552])).
fof(f28802,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK27(sK17,sK16,sK15,sK13))))) | $spl166),
  inference(cnf_transformation,[],[f28802_D])).
fof(f28802_D,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK27(sK17,sK16,sK15,sK13))))) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f28800,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK27(sK17,sK16,sK15,sK13))))) | $spl164),
  inference(cnf_transformation,[],[f28800_D])).
fof(f28800_D,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK27(sK17,sK16,sK15,sK13))))) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f28604,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK27(sK17,sK16,sK15,sK13))))) | p(s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK27(sK17,sK16,sK15,sK13))))) | $spl2),
  inference(resolution,[],[f24300,f531])).
fof(f28790,plain,(
  s(sK13,sK16) = s(sK13,sK27(sK17,sK16,sK15,sK13)) | $spl154),
  inference(cnf_transformation,[],[f28790_D])).
fof(f28790_D,plain,(
  s(sK13,sK16) = s(sK13,sK27(sK17,sK16,sK15,sK13)) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f28601,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK27(sK17,sK16,sK15,sK13))))) | s(sK13,sK16) = s(sK13,sK27(sK17,sK16,sK15,sK13)) | $spl2),
  inference(resolution,[],[f24300,f526])).
fof(f28798,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK27(sK17,sK16,sK15,sK13),sK16,sK15,sK13)))),s(sK13,sK27(sK17,sK16,sK15,sK13))))) | $spl162),
  inference(cnf_transformation,[],[f28798_D])).
fof(f28798_D,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK27(sK17,sK16,sK15,sK13),sK16,sK15,sK13)))),s(sK13,sK27(sK17,sK16,sK15,sK13))))) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f28599,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK27(sK17,sK16,sK15,sK13),sK16,sK15,sK13)))),s(sK13,sK27(sK17,sK16,sK15,sK13))))) | s(sK13,sK16) = s(sK13,sK27(sK17,sK16,sK15,sK13)) | $spl2),
  inference(resolution,[],[f24300,f523])).
fof(f28796,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK27(sK17,sK16,sK15,sK13),sK16,sK15,sK13))))) | $spl160),
  inference(cnf_transformation,[],[f28796_D])).
fof(f28796_D,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK27(sK17,sK16,sK15,sK13),sK16,sK15,sK13))))) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f28598,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK27(sK17,sK16,sK15,sK13),sK16,sK15,sK13))))) | s(sK13,sK16) = s(sK13,sK27(sK17,sK16,sK15,sK13)) | $spl2),
  inference(resolution,[],[f24300,f522])).
fof(f28794,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK27(sK17,sK16,sK15,sK13),sK16,sK15,sK13)),s(sK13,sK27(sK17,sK16,sK15,sK13))))) | $spl158),
  inference(cnf_transformation,[],[f28794_D])).
fof(f28794_D,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK27(sK17,sK16,sK15,sK13),sK16,sK15,sK13)),s(sK13,sK27(sK17,sK16,sK15,sK13))))) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f28596,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK27(sK17,sK16,sK15,sK13),sK16,sK15,sK13)),s(sK13,sK27(sK17,sK16,sK15,sK13))))) | s(sK13,sK16) = s(sK13,sK27(sK17,sK16,sK15,sK13)) | $spl2),
  inference(resolution,[],[f24300,f519])).
fof(f28792,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16))),s(sK13,sK27(sK27(sK17,sK16,sK15,sK13),sK16,sK15,sK13))))) | $spl156),
  inference(cnf_transformation,[],[f28792_D])).
fof(f28792_D,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16))),s(sK13,sK27(sK27(sK17,sK16,sK15,sK13),sK16,sK15,sK13))))) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f28595,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16))),s(sK13,sK27(sK27(sK17,sK16,sK15,sK13),sK16,sK15,sK13))))) | s(sK13,sK16) = s(sK13,sK27(sK17,sK16,sK15,sK13)) | $spl2),
  inference(resolution,[],[f24300,f518])).
fof(f28567,plain,(
  ( ! [X34,X32] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,X32)))) | ~p(s(bool,X34)) | s(bool,X34) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17),s(sK13,X32)))) ) | $spl152),
  inference(cnf_transformation,[],[f28567_D])).
fof(f28567_D,plain,(
  ( ! [X34,X32] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,X32)))) | ~p(s(bool,X34)) | s(bool,X34) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17),s(sK13,X32)))) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f28499,plain,(
  ( ! [X349,X347,X348] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,X347)))) | p(s(bool,X348)) | ~p(s(bool,X349)) | s(bool,X349) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17),s(sK13,X347)))) )),
  inference(resolution,[],[f800,f783])).
fof(f800,plain,(
  ( ! [X7] : (~p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17),s(sK13,X7)))) | p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,X7))))) )),
  inference(resolution,[],[f499,f511])).
fof(f28401,plain,(
  ( ! [X61,X62,X60,X63] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,X60)))) | ~p(s(bool,X61)) | ~p(s(bool,X62)) | s(bool,X62) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17),s(sK13,X60))) | s(bool,X61) != s(bool,X63)) )),
  inference(resolution,[],[f800,f558])).
fof(f28399,plain,(
  ( ! [X54,X52,X55,X53] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,X52)))) | p(s(bool,X53)) | ~p(s(bool,X54)) | ~p(s(bool,X55)) | s(bool,X54) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17),s(sK13,X52)))) )),
  inference(resolution,[],[f800,f557])).
fof(f28397,plain,(
  ( ! [X47,X45,X46,X44] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,X44)))) | p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X45) != s(bool,X47) | s(bool,X46) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17),s(sK13,X44)))) )),
  inference(resolution,[],[f800,f556])).
fof(f28396,plain,(
  ( ! [X43,X41,X42,X40] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,X40)))) | ~p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X42) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17),s(sK13,X40))) | s(bool,X41) != s(bool,X43)) )),
  inference(resolution,[],[f800,f554])).
fof(f28394,plain,(
  ( ! [X35,X33,X34,X32] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,X32)))) | p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X33) != s(bool,X35) | s(bool,X34) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17),s(sK13,X32)))) )),
  inference(resolution,[],[f800,f552])).
fof(f28362,plain,(
  ( ! [X35,X33] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X33),s(sK13,sK17)))) | ~p(s(bool,X35)) | s(bool,X35) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X33),s(sK13,sK16)))) ) | $spl150),
  inference(cnf_transformation,[],[f28362_D])).
fof(f28362_D,plain,(
  ( ! [X35,X33] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X33),s(sK13,sK17)))) | ~p(s(bool,X35)) | s(bool,X35) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X33),s(sK13,sK16)))) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f28294,plain,(
  ( ! [X350,X349,X348] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X348),s(sK13,sK17)))) | p(s(bool,X349)) | ~p(s(bool,X350)) | s(bool,X350) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X348),s(sK13,sK16)))) )),
  inference(resolution,[],[f799,f783])).
fof(f799,plain,(
  ( ! [X6] : (~p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X6),s(sK13,sK16)))) | p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X6),s(sK13,sK17))))) )),
  inference(resolution,[],[f499,f511])).
fof(f28196,plain,(
  ( ! [X61,X64,X62,X63] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X61),s(sK13,sK17)))) | ~p(s(bool,X62)) | ~p(s(bool,X63)) | s(bool,X63) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X61),s(sK13,sK16))) | s(bool,X62) != s(bool,X64)) )),
  inference(resolution,[],[f799,f558])).
fof(f28194,plain,(
  ( ! [X54,X56,X55,X53] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X53),s(sK13,sK17)))) | p(s(bool,X54)) | ~p(s(bool,X55)) | ~p(s(bool,X56)) | s(bool,X55) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X53),s(sK13,sK16)))) )),
  inference(resolution,[],[f799,f557])).
fof(f28192,plain,(
  ( ! [X47,X45,X48,X46] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X45),s(sK13,sK17)))) | p(s(bool,X46)) | ~p(s(bool,X47)) | s(bool,X46) != s(bool,X48) | s(bool,X47) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X45),s(sK13,sK16)))) )),
  inference(resolution,[],[f799,f556])).
fof(f28191,plain,(
  ( ! [X43,X41,X44,X42] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X41),s(sK13,sK17)))) | ~p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X43) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X41),s(sK13,sK16))) | s(bool,X42) != s(bool,X44)) )),
  inference(resolution,[],[f799,f554])).
fof(f28189,plain,(
  ( ! [X35,X33,X36,X34] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X33),s(sK13,sK17)))) | p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X34) != s(bool,X36) | s(bool,X35) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X33),s(sK13,sK16)))) )),
  inference(resolution,[],[f799,f552])).
fof(f27649,plain,(
  ( ! [X227,X226] : (p(s(bool,X226)) | p(s(bool,X227)) | s(bool,X227) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17),s(sK13,sK17)))) ) | ($spl104 | $spl114)),
  inference(resolution,[],[f27525,f783])).
fof(f27551,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17),s(sK13,sK17))) | s(bool,X36) != s(bool,X37)) ) | ($spl104 | $spl114)),
  inference(resolution,[],[f27525,f558])).
fof(f27548,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | p(s(bool,X28)) | s(bool,X28) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17),s(sK13,sK17)))) ) | ($spl104 | $spl114)),
  inference(resolution,[],[f27525,f557])).
fof(f27547,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X23) != s(bool,X25) | s(bool,X24) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17),s(sK13,sK17)))) ) | ($spl104 | $spl114)),
  inference(resolution,[],[f27525,f556])).
fof(f27546,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X20) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17),s(sK13,sK17))) | s(bool,X21) != s(bool,X22)) ) | ($spl104 | $spl114)),
  inference(resolution,[],[f27525,f554])).
fof(f27544,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17),s(sK13,sK17)))) ) | ($spl104 | $spl114)),
  inference(resolution,[],[f27525,f552])).
fof(f26379,plain,(
  ( ! [X212,X213] : (p(s(bool,X212)) | p(s(bool,X213)) | s(bool,X213) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17),s(sK13,sK17)))) ) | $spl104),
  inference(resolution,[],[f26273,f783])).
fof(f26273,plain,(
  p(s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17),s(sK13,sK17)))) | $spl104),
  inference(forward_demodulation,[],[f14958,f15143])).
fof(f26281,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17),s(sK13,sK17))) | s(bool,X22) != s(bool,X23)) ) | $spl104),
  inference(resolution,[],[f26273,f558])).
fof(f26278,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17),s(sK13,sK17)))) ) | $spl104),
  inference(resolution,[],[f26273,f557])).
fof(f26277,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17),s(sK13,sK17)))) ) | $spl104),
  inference(resolution,[],[f26273,f556])).
fof(f26276,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17),s(sK13,sK17))) | s(bool,X7) != s(bool,X8)) ) | $spl104),
  inference(resolution,[],[f26273,f554])).
fof(f26274,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17),s(sK13,sK17)))) ) | $spl104),
  inference(resolution,[],[f26273,f552])).
fof(f25601,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17)))))) ) | $spl146),
  inference(cnf_transformation,[],[f25601_D])).
fof(f25601_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17)))))) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f25504,plain,(
  ( ! [X210,X211] : (p(s(bool,X210)) | ~p(s(bool,X211)) | s(bool,X211) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17)))))) )),
  inference(resolution,[],[f1032,f783])).
fof(f1032,plain,(
  ~p(s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17))))))),
  inference(resolution,[],[f500,f532])).
fof(f532,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | ~p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f276])).
fof(f276,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))))),
  inference(ennf_transformation,[],[f164])).
fof(f164,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) => p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))))),
  inference(flattening,[],[f163])).
fof(f163,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) => p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))))),
  inference(rectify,[],[f100])).
fof(f100,axiom,(
  ! [X6] : ! [X8,X9,X10] : (p(s(bool,'const.relation.RC_3'(s(fun(X6,fun(X6,bool)),X8),s(X6,X9),s(X6,X10)))) => p(s(bool,'const.relation.RTC_3'(s(fun(X6,fun(X6,bool)),X8),s(X6,X9),s(X6,X10)))))),
  file('thm.relation.RTC_lifts_monotonicities','thm.relation.RC_RTC')).
fof(f25406,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17))))) | s(bool,X21) != s(bool,X23)) )),
  inference(resolution,[],[f1032,f558])).
fof(f25404,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17)))))) )),
  inference(resolution,[],[f1032,f557])).
fof(f25402,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17)))))) )),
  inference(resolution,[],[f1032,f556])).
fof(f25401,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17))))) | s(bool,X6) != s(bool,X8)) )),
  inference(resolution,[],[f1032,f554])).
fof(f25399,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17)))))) )),
  inference(resolution,[],[f1032,f552])).
fof(f25248,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | s(bool,X3) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17)))))) ) | $spl144),
  inference(cnf_transformation,[],[f25248_D])).
fof(f25248_D,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | s(bool,X3) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17)))))) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f25149,plain,(
  ( ! [X212,X213] : (p(s(bool,X212)) | ~p(s(bool,X213)) | s(bool,X213) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17)))))) )),
  inference(resolution,[],[f1029,f783])).
fof(f25051,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17))))) | s(bool,X23) != s(bool,X25)) )),
  inference(resolution,[],[f1029,f558])).
fof(f25049,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17)))))) )),
  inference(resolution,[],[f1029,f557])).
fof(f25047,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17)))))) )),
  inference(resolution,[],[f1029,f556])).
fof(f25046,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X9) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17))))) | s(bool,X8) != s(bool,X10)) )),
  inference(resolution,[],[f1029,f554])).
fof(f25044,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17)))))) )),
  inference(resolution,[],[f1029,f552])).
fof(f24952,plain,(
  ( ! [X212,X213] : (p(s(bool,X212)) | p(s(bool,X213)) | s(bool,X213) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16)))) ) | ($spl68 | $spl80)),
  inference(resolution,[],[f24845,f783])).
fof(f24845,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16)))) | ($spl68 | $spl80)),
  inference(forward_demodulation,[],[f11251,f11239])).
fof(f24854,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16))) | s(bool,X22) != s(bool,X23)) ) | ($spl68 | $spl80)),
  inference(resolution,[],[f24845,f558])).
fof(f24851,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16)))) ) | ($spl68 | $spl80)),
  inference(resolution,[],[f24845,f557])).
fof(f24850,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16)))) ) | ($spl68 | $spl80)),
  inference(resolution,[],[f24845,f556])).
fof(f24849,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16))) | s(bool,X7) != s(bool,X8)) ) | ($spl68 | $spl80)),
  inference(resolution,[],[f24845,f554])).
fof(f24847,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16)))) ) | ($spl68 | $spl80)),
  inference(resolution,[],[f24845,f552])).
fof(f24843,plain,(
  ( ! [X22] : (p(s(bool,X22)) | s(bool,X22) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16))),s(sK13,sK17)))) ) | $spl140),
  inference(cnf_transformation,[],[f24843_D])).
fof(f24843_D,plain,(
  ( ! [X22] : (p(s(bool,X22)) | s(bool,X22) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16))),s(sK13,sK17)))) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f24763,plain,(
  ( ! [X233,X234] : (p(s(bool,X233)) | p(s(bool,X234)) | s(bool,X234) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16))),s(sK13,sK17)))) ) | ($spl8 | $spl68)),
  inference(resolution,[],[f24502,f783])).
fof(f24665,plain,(
  ( ! [X43,X44,X42] : (p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X42) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16))),s(sK13,sK17))) | s(bool,X43) != s(bool,X44)) ) | ($spl8 | $spl68)),
  inference(resolution,[],[f24502,f558])).
fof(f24662,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X33)) | ~p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X35) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16))),s(sK13,sK17)))) ) | ($spl8 | $spl68)),
  inference(resolution,[],[f24502,f557])).
fof(f24661,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | p(s(bool,X31)) | s(bool,X30) != s(bool,X32) | s(bool,X31) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16))),s(sK13,sK17)))) ) | ($spl8 | $spl68)),
  inference(resolution,[],[f24502,f556])).
fof(f24660,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16))),s(sK13,sK17))) | s(bool,X28) != s(bool,X29)) ) | ($spl8 | $spl68)),
  inference(resolution,[],[f24502,f554])).
fof(f24658,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16))),s(sK13,sK17)))) ) | ($spl8 | $spl68)),
  inference(resolution,[],[f24502,f552])).
fof(f24501,plain,(
  ( ! [X22] : (p(s(bool,X22)) | s(bool,X22) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16))),s(sK13,sK27(sK17,sK16,sK15,sK13))))) ) | $spl138),
  inference(cnf_transformation,[],[f24501_D])).
fof(f24501_D,plain,(
  ( ! [X22] : (p(s(bool,X22)) | s(bool,X22) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16))),s(sK13,sK27(sK17,sK16,sK15,sK13))))) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f24421,plain,(
  ( ! [X233,X234] : (p(s(bool,X233)) | p(s(bool,X234)) | s(bool,X234) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16))),s(sK13,sK27(sK17,sK16,sK15,sK13))))) ) | $spl2),
  inference(resolution,[],[f1000,f783])).
fof(f24323,plain,(
  ( ! [X43,X44,X42] : (p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X42) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16))),s(sK13,sK27(sK17,sK16,sK15,sK13)))) | s(bool,X43) != s(bool,X44)) ) | $spl2),
  inference(resolution,[],[f1000,f558])).
fof(f24320,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X33)) | ~p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X35) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16))),s(sK13,sK27(sK17,sK16,sK15,sK13))))) ) | $spl2),
  inference(resolution,[],[f1000,f557])).
fof(f24319,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | p(s(bool,X31)) | s(bool,X30) != s(bool,X32) | s(bool,X31) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16))),s(sK13,sK27(sK17,sK16,sK15,sK13))))) ) | $spl2),
  inference(resolution,[],[f1000,f556])).
fof(f24318,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16))),s(sK13,sK27(sK17,sK16,sK15,sK13)))) | s(bool,X28) != s(bool,X29)) ) | $spl2),
  inference(resolution,[],[f1000,f554])).
fof(f24316,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16))),s(sK13,sK27(sK17,sK16,sK15,sK13))))) ) | $spl2),
  inference(resolution,[],[f1000,f552])).
fof(f23921,plain,(
  ( ! [X227,X226] : (p(s(bool,X226)) | p(s(bool,X227)) | s(bool,X227) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16)))) ) | ($spl68 | $spl78)),
  inference(resolution,[],[f20864,f783])).
fof(f23823,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16))) | s(bool,X36) != s(bool,X37)) ) | ($spl68 | $spl78)),
  inference(resolution,[],[f20864,f558])).
fof(f23820,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | p(s(bool,X28)) | s(bool,X28) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16)))) ) | ($spl68 | $spl78)),
  inference(resolution,[],[f20864,f557])).
fof(f23819,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X23) != s(bool,X25) | s(bool,X24) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16)))) ) | ($spl68 | $spl78)),
  inference(resolution,[],[f20864,f556])).
fof(f23818,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X20) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16))) | s(bool,X21) != s(bool,X22)) ) | ($spl68 | $spl78)),
  inference(resolution,[],[f20864,f554])).
fof(f23816,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16)))) ) | ($spl68 | $spl78)),
  inference(resolution,[],[f20864,f552])).
fof(f23780,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | s(bool,X3) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16)))) ) | $spl136),
  inference(cnf_transformation,[],[f23780_D])).
fof(f23780_D,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | s(bool,X3) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16)))) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f23728,plain,(
  ( ! [X212,X213] : (p(s(bool,X212)) | ~p(s(bool,X213)) | s(bool,X213) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16)))) ) | ($spl68 | $spl79)),
  inference(resolution,[],[f23617,f783])).
fof(f23617,plain,(
  ~p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16)))) | ($spl68 | $spl79)),
  inference(forward_demodulation,[],[f11248,f11239])).
fof(f11248,plain,(
  ~p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) | $spl79),
  inference(cnf_transformation,[],[f11248_D])).
fof(f11248_D,plain,(
  ~p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) <=> ~$spl79),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl79])])).
fof(f23630,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16))) | s(bool,X23) != s(bool,X25)) ) | ($spl68 | $spl79)),
  inference(resolution,[],[f23617,f558])).
fof(f23628,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16)))) ) | ($spl68 | $spl79)),
  inference(resolution,[],[f23617,f557])).
fof(f23626,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16)))) ) | ($spl68 | $spl79)),
  inference(resolution,[],[f23617,f556])).
fof(f23625,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X9) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16))) | s(bool,X8) != s(bool,X10)) ) | ($spl68 | $spl79)),
  inference(resolution,[],[f23617,f554])).
fof(f23623,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16)))) ) | ($spl68 | $spl79)),
  inference(resolution,[],[f23617,f552])).
fof(f22799,plain,(
  ( ! [X212,X213] : (p(s(bool,X212)) | p(s(bool,X213)) | s(bool,X213) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16)))) ) | ($spl6 | $spl68)),
  inference(resolution,[],[f19674,f783])).
fof(f19674,plain,(
  p(s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16)))) | ($spl6 | $spl68)),
  inference(forward_demodulation,[],[f11055,f11239])).
fof(f22701,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16))) | s(bool,X22) != s(bool,X23)) ) | ($spl6 | $spl68)),
  inference(resolution,[],[f19674,f558])).
fof(f22698,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16)))) ) | ($spl6 | $spl68)),
  inference(resolution,[],[f19674,f557])).
fof(f22697,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16)))) ) | ($spl6 | $spl68)),
  inference(resolution,[],[f19674,f556])).
fof(f22696,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16))) | s(bool,X7) != s(bool,X8)) ) | ($spl6 | $spl68)),
  inference(resolution,[],[f19674,f554])).
fof(f22694,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16)))) ) | ($spl6 | $spl68)),
  inference(resolution,[],[f19674,f552])).
fof(f21981,plain,(
  ( ! [X227,X226] : (p(s(bool,X226)) | p(s(bool,X227)) | s(bool,X227) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) | $spl78),
  inference(resolution,[],[f11249,f783])).
fof(f21883,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13)))) | s(bool,X36) != s(bool,X37)) ) | $spl78),
  inference(resolution,[],[f11249,f558])).
fof(f21880,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | p(s(bool,X28)) | s(bool,X28) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) | $spl78),
  inference(resolution,[],[f11249,f557])).
fof(f21879,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X23) != s(bool,X25) | s(bool,X24) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) | $spl78),
  inference(resolution,[],[f11249,f556])).
fof(f21878,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X20) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13)))) | s(bool,X21) != s(bool,X22)) ) | $spl78),
  inference(resolution,[],[f11249,f554])).
fof(f21876,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) | $spl78),
  inference(resolution,[],[f11249,f552])).
fof(f21365,plain,(
  ( ! [X212,X213] : (p(s(bool,X212)) | p(s(bool,X213)) | s(bool,X213) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) | $spl6),
  inference(resolution,[],[f11055,f783])).
fof(f21267,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13)))) | s(bool,X22) != s(bool,X23)) ) | $spl6),
  inference(resolution,[],[f11055,f558])).
fof(f21264,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) | $spl6),
  inference(resolution,[],[f11055,f557])).
fof(f21263,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) | $spl6),
  inference(resolution,[],[f11055,f556])).
fof(f21262,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13)))) | s(bool,X7) != s(bool,X8)) ) | $spl6),
  inference(resolution,[],[f11055,f554])).
fof(f21260,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) | $spl6),
  inference(resolution,[],[f11055,f552])).
fof(f21186,plain,(
  ( ! [X241,X242] : (p(s(bool,X241)) | p(s(bool,X242)) | s(bool,X242) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) | $spl6),
  inference(resolution,[],[f1004,f783])).
fof(f21088,plain,(
  ( ! [X52,X50,X51] : (p(s(bool,X50)) | ~p(s(bool,X51)) | s(bool,X50) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13)))) | s(bool,X51) != s(bool,X52)) ) | $spl6),
  inference(resolution,[],[f1004,f558])).
fof(f21085,plain,(
  ( ! [X43,X41,X42] : (p(s(bool,X41)) | ~p(s(bool,X42)) | p(s(bool,X43)) | s(bool,X43) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) | $spl6),
  inference(resolution,[],[f1004,f557])).
fof(f21084,plain,(
  ( ! [X39,X38,X40] : (p(s(bool,X38)) | p(s(bool,X39)) | s(bool,X38) != s(bool,X40) | s(bool,X39) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) | $spl6),
  inference(resolution,[],[f1004,f556])).
fof(f21083,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13)))) | s(bool,X36) != s(bool,X37)) ) | $spl6),
  inference(resolution,[],[f1004,f554])).
fof(f21081,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) | $spl6),
  inference(resolution,[],[f1004,f552])).
fof(f21065,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) | p(s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) | $spl6),
  inference(resolution,[],[f1004,f531])).
fof(f11239,plain,(
  s(sK13,sK16) = s(sK13,sK28(sK17,sK16,sK15,sK13)) | $spl68),
  inference(cnf_transformation,[],[f11239_D])).
fof(f11239_D,plain,(
  s(sK13,sK16) = s(sK13,sK28(sK17,sK16,sK15,sK13)) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f21062,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) | s(sK13,sK16) = s(sK13,sK28(sK17,sK16,sK15,sK13)) | $spl6),
  inference(resolution,[],[f1004,f526])).
fof(f11247,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK28(sK17,sK16,sK15,sK13),sK16,sK15,sK13)))),s(sK13,sK28(sK17,sK16,sK15,sK13))))) | $spl76),
  inference(cnf_transformation,[],[f11247_D])).
fof(f11247_D,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK28(sK17,sK16,sK15,sK13),sK16,sK15,sK13)))),s(sK13,sK28(sK17,sK16,sK15,sK13))))) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f21060,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK28(sK17,sK16,sK15,sK13),sK16,sK15,sK13)))),s(sK13,sK28(sK17,sK16,sK15,sK13))))) | s(sK13,sK16) = s(sK13,sK28(sK17,sK16,sK15,sK13)) | $spl6),
  inference(resolution,[],[f1004,f523])).
fof(f11245,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK28(sK17,sK16,sK15,sK13),sK16,sK15,sK13))))) | $spl74),
  inference(cnf_transformation,[],[f11245_D])).
fof(f11245_D,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK28(sK17,sK16,sK15,sK13),sK16,sK15,sK13))))) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f21059,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK28(sK17,sK16,sK15,sK13),sK16,sK15,sK13))))) | s(sK13,sK16) = s(sK13,sK28(sK17,sK16,sK15,sK13)) | $spl6),
  inference(resolution,[],[f1004,f522])).
fof(f11243,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK28(sK17,sK16,sK15,sK13),sK16,sK15,sK13)),s(sK13,sK28(sK17,sK16,sK15,sK13))))) | $spl72),
  inference(cnf_transformation,[],[f11243_D])).
fof(f11243_D,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK28(sK17,sK16,sK15,sK13),sK16,sK15,sK13)),s(sK13,sK28(sK17,sK16,sK15,sK13))))) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f21057,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK28(sK17,sK16,sK15,sK13),sK16,sK15,sK13)),s(sK13,sK28(sK17,sK16,sK15,sK13))))) | s(sK13,sK16) = s(sK13,sK28(sK17,sK16,sK15,sK13)) | $spl6),
  inference(resolution,[],[f1004,f519])).
fof(f11241,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16))),s(sK13,sK27(sK28(sK17,sK16,sK15,sK13),sK16,sK15,sK13))))) | $spl70),
  inference(cnf_transformation,[],[f11241_D])).
fof(f11241_D,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16))),s(sK13,sK27(sK28(sK17,sK16,sK15,sK13),sK16,sK15,sK13))))) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f21056,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16))),s(sK13,sK27(sK28(sK17,sK16,sK15,sK13),sK16,sK15,sK13))))) | s(sK13,sK16) = s(sK13,sK28(sK17,sK16,sK15,sK13)) | $spl6),
  inference(resolution,[],[f1004,f518])).
fof(f20988,plain,(
  ( ! [X227,X226] : (p(s(bool,X226)) | p(s(bool,X227)) | s(bool,X227) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16)))) ) | ($spl68 | $spl78)),
  inference(resolution,[],[f20864,f783])).
fof(f20890,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16))) | s(bool,X36) != s(bool,X37)) ) | ($spl68 | $spl78)),
  inference(resolution,[],[f20864,f558])).
fof(f20887,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | p(s(bool,X28)) | s(bool,X28) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16)))) ) | ($spl68 | $spl78)),
  inference(resolution,[],[f20864,f557])).
fof(f20886,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X23) != s(bool,X25) | s(bool,X24) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16)))) ) | ($spl68 | $spl78)),
  inference(resolution,[],[f20864,f556])).
fof(f20885,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X20) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16))) | s(bool,X21) != s(bool,X22)) ) | ($spl68 | $spl78)),
  inference(resolution,[],[f20864,f554])).
fof(f20883,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16)))) ) | ($spl68 | $spl78)),
  inference(resolution,[],[f20864,f552])).
fof(f19780,plain,(
  ( ! [X212,X213] : (p(s(bool,X212)) | p(s(bool,X213)) | s(bool,X213) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16)))) ) | ($spl6 | $spl68)),
  inference(resolution,[],[f19674,f783])).
fof(f19682,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16))) | s(bool,X22) != s(bool,X23)) ) | ($spl6 | $spl68)),
  inference(resolution,[],[f19674,f558])).
fof(f19679,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16)))) ) | ($spl6 | $spl68)),
  inference(resolution,[],[f19674,f557])).
fof(f19678,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16)))) ) | ($spl6 | $spl68)),
  inference(resolution,[],[f19674,f556])).
fof(f19677,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16))) | s(bool,X7) != s(bool,X8)) ) | ($spl6 | $spl68)),
  inference(resolution,[],[f19674,f554])).
fof(f19675,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16)))) ) | ($spl6 | $spl68)),
  inference(resolution,[],[f19674,f552])).
fof(f18560,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl126),
  inference(cnf_transformation,[],[f18560_D])).
fof(f18560_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f18875,plain,(
  ( ! [X212,X213] : (p(s(bool,X212)) | p(s(bool,X213)) | s(bool,X213) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl60),
  inference(resolution,[],[f6316,f783])).
fof(f18777,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17))) | s(bool,X22) != s(bool,X23)) ) | $spl60),
  inference(resolution,[],[f6316,f558])).
fof(f18774,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl60),
  inference(resolution,[],[f6316,f557])).
fof(f18773,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl60),
  inference(resolution,[],[f6316,f556])).
fof(f18772,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17))) | s(bool,X7) != s(bool,X8)) ) | $spl60),
  inference(resolution,[],[f6316,f554])).
fof(f18770,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl60),
  inference(resolution,[],[f6316,f552])).
fof(f18728,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl128),
  inference(cnf_transformation,[],[f18728_D])).
fof(f18728_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f18667,plain,(
  ( ! [X210,X211] : (p(s(bool,X210)) | ~p(s(bool,X211)) | s(bool,X211) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl61),
  inference(resolution,[],[f6315,f783])).
fof(f6315,plain,(
  ~p(s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) | $spl61),
  inference(cnf_transformation,[],[f6315_D])).
fof(f6315_D,plain,(
  ~p(s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) <=> ~$spl61),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl61])])).
fof(f18569,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17))) | s(bool,X21) != s(bool,X23)) ) | $spl61),
  inference(resolution,[],[f6315,f558])).
fof(f18567,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl61),
  inference(resolution,[],[f6315,f557])).
fof(f18565,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl61),
  inference(resolution,[],[f6315,f556])).
fof(f18564,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17))) | s(bool,X6) != s(bool,X8)) ) | $spl61),
  inference(resolution,[],[f6315,f554])).
fof(f18562,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl61),
  inference(resolution,[],[f6315,f552])).
fof(f18496,plain,(
  ( ! [X212,X213] : (p(s(bool,X212)) | p(s(bool,X213)) | s(bool,X213) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl60),
  inference(resolution,[],[f6316,f783])).
fof(f18398,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17))) | s(bool,X22) != s(bool,X23)) ) | $spl60),
  inference(resolution,[],[f6316,f558])).
fof(f18395,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl60),
  inference(resolution,[],[f6316,f557])).
fof(f18394,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl60),
  inference(resolution,[],[f6316,f556])).
fof(f18393,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17))) | s(bool,X7) != s(bool,X8)) ) | $spl60),
  inference(resolution,[],[f6316,f554])).
fof(f18391,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl60),
  inference(resolution,[],[f6316,f552])).
fof(f16585,plain,(
  ( ! [X15] : (p(s(bool,X15)) | s(bool,X15) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl122),
  inference(cnf_transformation,[],[f16585_D])).
fof(f16585_D,plain,(
  ( ! [X15] : (p(s(bool,X15)) | s(bool,X15) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f18326,plain,(
  ( ! [X227,X226] : (p(s(bool,X226)) | p(s(bool,X227)) | s(bool,X227) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl58),
  inference(resolution,[],[f6314,f783])).
fof(f18228,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17))) | s(bool,X36) != s(bool,X37)) ) | $spl58),
  inference(resolution,[],[f6314,f558])).
fof(f18225,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | p(s(bool,X28)) | s(bool,X28) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl58),
  inference(resolution,[],[f6314,f557])).
fof(f18224,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X23) != s(bool,X25) | s(bool,X24) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl58),
  inference(resolution,[],[f6314,f556])).
fof(f18223,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X20) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17))) | s(bool,X21) != s(bool,X22)) ) | $spl58),
  inference(resolution,[],[f6314,f554])).
fof(f18221,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl58),
  inference(resolution,[],[f6314,f552])).
fof(f15979,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl120),
  inference(cnf_transformation,[],[f15979_D])).
fof(f15979_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f17721,plain,(
  ( ! [X212,X213] : (p(s(bool,X212)) | p(s(bool,X213)) | s(bool,X213) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl4),
  inference(resolution,[],[f6120,f783])).
fof(f6120,plain,(
  p(s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) | $spl4),
  inference(resolution,[],[f1002,f534])).
fof(f17623,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17))) | s(bool,X22) != s(bool,X23)) ) | $spl4),
  inference(resolution,[],[f6120,f558])).
fof(f17620,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl4),
  inference(resolution,[],[f6120,f557])).
fof(f17619,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl4),
  inference(resolution,[],[f6120,f556])).
fof(f17618,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17))) | s(bool,X7) != s(bool,X8)) ) | $spl4),
  inference(resolution,[],[f6120,f554])).
fof(f17616,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl4),
  inference(resolution,[],[f6120,f552])).
fof(f6317,plain,(
  ( ! [X30] : (p(s(bool,X30)) | s(bool,X30) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl62),
  inference(cnf_transformation,[],[f6317_D])).
fof(f6317_D,plain,(
  ( ! [X30] : (p(s(bool,X30)) | s(bool,X30) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f17547,plain,(
  ( ! [X241,X242] : (p(s(bool,X241)) | p(s(bool,X242)) | s(bool,X242) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl4),
  inference(resolution,[],[f1002,f783])).
fof(f17449,plain,(
  ( ! [X52,X50,X51] : (p(s(bool,X50)) | ~p(s(bool,X51)) | s(bool,X50) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17))) | s(bool,X51) != s(bool,X52)) ) | $spl4),
  inference(resolution,[],[f1002,f558])).
fof(f17446,plain,(
  ( ! [X43,X41,X42] : (p(s(bool,X41)) | ~p(s(bool,X42)) | p(s(bool,X43)) | s(bool,X43) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl4),
  inference(resolution,[],[f1002,f557])).
fof(f17445,plain,(
  ( ! [X39,X38,X40] : (p(s(bool,X38)) | p(s(bool,X39)) | s(bool,X38) != s(bool,X40) | s(bool,X39) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl4),
  inference(resolution,[],[f1002,f556])).
fof(f17444,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17))) | s(bool,X36) != s(bool,X37)) ) | $spl4),
  inference(resolution,[],[f1002,f554])).
fof(f17442,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl4),
  inference(resolution,[],[f1002,f552])).
fof(f6316,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) | $spl60),
  inference(cnf_transformation,[],[f6316_D])).
fof(f6316_D,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f6314,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) | $spl58),
  inference(cnf_transformation,[],[f6314_D])).
fof(f6314_D,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f17426,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) | p(s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) | $spl4),
  inference(resolution,[],[f1002,f531])).
fof(f17339,plain,(
  ( ! [X212,X213] : (p(s(bool,X212)) | p(s(bool,X213)) | s(bool,X213) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17),s(sK13,sK17)))) ) | ($spl4 | $spl48)),
  inference(resolution,[],[f16587,f783])).
fof(f16587,plain,(
  p(s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17),s(sK13,sK17)))) | ($spl4 | $spl48)),
  inference(backward_demodulation,[],[f6304,f6120])).
fof(f17241,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17),s(sK13,sK17))) | s(bool,X22) != s(bool,X23)) ) | ($spl4 | $spl48)),
  inference(resolution,[],[f16587,f558])).
fof(f17238,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17),s(sK13,sK17)))) ) | ($spl4 | $spl48)),
  inference(resolution,[],[f16587,f557])).
fof(f17237,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17),s(sK13,sK17)))) ) | ($spl4 | $spl48)),
  inference(resolution,[],[f16587,f556])).
fof(f17236,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17),s(sK13,sK17))) | s(bool,X7) != s(bool,X8)) ) | ($spl4 | $spl48)),
  inference(resolution,[],[f16587,f554])).
fof(f17234,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17),s(sK13,sK17)))) ) | ($spl4 | $spl48)),
  inference(resolution,[],[f16587,f552])).
fof(f6304,plain,(
  s(sK13,sK17) = s(sK13,sK27(sK17,sK16,sK15,sK13)) | $spl48),
  inference(cnf_transformation,[],[f6304_D])).
fof(f6304_D,plain,(
  s(sK13,sK17) = s(sK13,sK27(sK17,sK16,sK15,sK13)) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f6306,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)))),s(sK13,sK27(sK17,sK27(sK17,sK16,sK15,sK13),sK15,sK13))))) | $spl50),
  inference(cnf_transformation,[],[f6306_D])).
fof(f6306_D,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)))),s(sK13,sK27(sK17,sK27(sK17,sK16,sK15,sK13),sK15,sK13))))) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f10834,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)))),s(sK13,sK27(sK17,sK27(sK17,sK16,sK15,sK13),sK15,sK13))))) | s(sK13,sK17) = s(sK13,sK27(sK17,sK16,sK15,sK13)) | $spl4),
  inference(resolution,[],[f1002,f518])).
fof(f6308,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK27(sK17,sK16,sK15,sK13),sK15,sK13)),s(sK13,sK17)))) | $spl52),
  inference(cnf_transformation,[],[f6308_D])).
fof(f6308_D,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK27(sK17,sK16,sK15,sK13),sK15,sK13)),s(sK13,sK17)))) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f10835,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK27(sK17,sK16,sK15,sK13),sK15,sK13)),s(sK13,sK17)))) | s(sK13,sK17) = s(sK13,sK27(sK17,sK16,sK15,sK13)) | $spl4),
  inference(resolution,[],[f1002,f519])).
fof(f6310,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK28(sK17,sK27(sK17,sK16,sK15,sK13),sK15,sK13))))) | $spl54),
  inference(cnf_transformation,[],[f6310_D])).
fof(f6310_D,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK28(sK17,sK27(sK17,sK16,sK15,sK13),sK15,sK13))))) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f10837,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK28(sK17,sK27(sK17,sK16,sK15,sK13),sK15,sK13))))) | s(sK13,sK17) = s(sK13,sK27(sK17,sK16,sK15,sK13)) | $spl4),
  inference(resolution,[],[f1002,f522])).
fof(f6312,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK27(sK17,sK16,sK15,sK13),sK15,sK13)))),s(sK13,sK17)))) | $spl56),
  inference(cnf_transformation,[],[f6312_D])).
fof(f6312_D,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK27(sK17,sK16,sK15,sK13),sK15,sK13)))),s(sK13,sK17)))) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f10838,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK27(sK17,sK16,sK15,sK13),sK15,sK13)))),s(sK13,sK17)))) | s(sK13,sK17) = s(sK13,sK27(sK17,sK16,sK15,sK13)) | $spl4),
  inference(resolution,[],[f1002,f523])).
fof(f10840,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) | s(sK13,sK17) = s(sK13,sK27(sK17,sK16,sK15,sK13)) | $spl4),
  inference(resolution,[],[f1002,f526])).
fof(f16521,plain,(
  ( ! [X227,X226] : (p(s(bool,X226)) | p(s(bool,X227)) | s(bool,X227) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl58),
  inference(resolution,[],[f6314,f783])).
fof(f16423,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17))) | s(bool,X36) != s(bool,X37)) ) | $spl58),
  inference(resolution,[],[f6314,f558])).
fof(f16420,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | p(s(bool,X28)) | s(bool,X28) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl58),
  inference(resolution,[],[f6314,f557])).
fof(f16419,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X23) != s(bool,X25) | s(bool,X24) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl58),
  inference(resolution,[],[f6314,f556])).
fof(f16418,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X20) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17))) | s(bool,X21) != s(bool,X22)) ) | $spl58),
  inference(resolution,[],[f6314,f554])).
fof(f16416,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl58),
  inference(resolution,[],[f6314,f552])).
fof(f15915,plain,(
  ( ! [X212,X213] : (p(s(bool,X212)) | p(s(bool,X213)) | s(bool,X213) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl4),
  inference(resolution,[],[f6120,f783])).
fof(f15817,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17))) | s(bool,X22) != s(bool,X23)) ) | $spl4),
  inference(resolution,[],[f6120,f558])).
fof(f15814,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl4),
  inference(resolution,[],[f6120,f557])).
fof(f15813,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl4),
  inference(resolution,[],[f6120,f556])).
fof(f15812,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17))) | s(bool,X7) != s(bool,X8)) ) | $spl4),
  inference(resolution,[],[f6120,f554])).
fof(f15810,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl4),
  inference(resolution,[],[f6120,f552])).
fof(f15078,plain,(
  ( ! [X241,X242] : (p(s(bool,X241)) | p(s(bool,X242)) | s(bool,X242) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK21(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) )),
  inference(resolution,[],[f798,f783])).
fof(f14980,plain,(
  ( ! [X52,X50,X51] : (p(s(bool,X50)) | ~p(s(bool,X51)) | s(bool,X50) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK21(sK17,sK16,sK15,sK13)),s(sK13,sK17))) | s(bool,X51) != s(bool,X52)) )),
  inference(resolution,[],[f798,f558])).
fof(f14977,plain,(
  ( ! [X43,X41,X42] : (p(s(bool,X41)) | ~p(s(bool,X42)) | p(s(bool,X43)) | s(bool,X43) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK21(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) )),
  inference(resolution,[],[f798,f557])).
fof(f14976,plain,(
  ( ! [X39,X38,X40] : (p(s(bool,X38)) | p(s(bool,X39)) | s(bool,X38) != s(bool,X40) | s(bool,X39) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK21(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) )),
  inference(resolution,[],[f798,f556])).
fof(f14975,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK21(sK17,sK16,sK15,sK13)),s(sK13,sK17))) | s(bool,X36) != s(bool,X37)) )),
  inference(resolution,[],[f798,f554])).
fof(f14973,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK21(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) )),
  inference(resolution,[],[f798,f552])).
fof(f14957,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK21(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) | p(s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK21(sK17,sK16,sK15,sK13)),s(sK13,sK17))))),
  inference(resolution,[],[f798,f531])).
fof(f15143,plain,(
  s(sK13,sK17) = s(sK13,sK21(sK17,sK16,sK15,sK13)) | $spl104),
  inference(cnf_transformation,[],[f15143_D])).
fof(f15143_D,plain,(
  s(sK13,sK17) = s(sK13,sK21(sK17,sK16,sK15,sK13)) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f14954,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK21(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) | s(sK13,sK17) = s(sK13,sK21(sK17,sK16,sK15,sK13))),
  inference(resolution,[],[f798,f526])).
fof(f15151,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK21(sK17,sK16,sK15,sK13),sK15,sK13)))),s(sK13,sK17)))) | $spl112),
  inference(cnf_transformation,[],[f15151_D])).
fof(f15151_D,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK21(sK17,sK16,sK15,sK13),sK15,sK13)))),s(sK13,sK17)))) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f14952,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK21(sK17,sK16,sK15,sK13),sK15,sK13)))),s(sK13,sK17)))) | s(sK13,sK17) = s(sK13,sK21(sK17,sK16,sK15,sK13))),
  inference(resolution,[],[f798,f523])).
fof(f15149,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK21(sK17,sK16,sK15,sK13)),s(sK13,sK28(sK17,sK21(sK17,sK16,sK15,sK13),sK15,sK13))))) | $spl110),
  inference(cnf_transformation,[],[f15149_D])).
fof(f15149_D,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK21(sK17,sK16,sK15,sK13)),s(sK13,sK28(sK17,sK21(sK17,sK16,sK15,sK13),sK15,sK13))))) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f14951,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK21(sK17,sK16,sK15,sK13)),s(sK13,sK28(sK17,sK21(sK17,sK16,sK15,sK13),sK15,sK13))))) | s(sK13,sK17) = s(sK13,sK21(sK17,sK16,sK15,sK13))),
  inference(resolution,[],[f798,f522])).
fof(f15147,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK21(sK17,sK16,sK15,sK13),sK15,sK13)),s(sK13,sK17)))) | $spl108),
  inference(cnf_transformation,[],[f15147_D])).
fof(f15147_D,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK21(sK17,sK16,sK15,sK13),sK15,sK13)),s(sK13,sK17)))) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f14949,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK21(sK17,sK16,sK15,sK13),sK15,sK13)),s(sK13,sK17)))) | s(sK13,sK17) = s(sK13,sK21(sK17,sK16,sK15,sK13))),
  inference(resolution,[],[f798,f519])).
fof(f15145,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK21(sK17,sK16,sK15,sK13)))),s(sK13,sK27(sK17,sK21(sK17,sK16,sK15,sK13),sK15,sK13))))) | $spl106),
  inference(cnf_transformation,[],[f15145_D])).
fof(f15145_D,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK21(sK17,sK16,sK15,sK13)))),s(sK13,sK27(sK17,sK21(sK17,sK16,sK15,sK13),sK15,sK13))))) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f14948,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK21(sK17,sK16,sK15,sK13)))),s(sK13,sK27(sK17,sK21(sK17,sK16,sK15,sK13),sK15,sK13))))) | s(sK13,sK17) = s(sK13,sK21(sK17,sK16,sK15,sK13))),
  inference(resolution,[],[f798,f518])).
fof(f14866,plain,(
  ( ! [X241,X242] : (p(s(bool,X241)) | p(s(bool,X242)) | s(bool,X242) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK21(sK17,sK16,sK15,sK13))))) )),
  inference(resolution,[],[f797,f783])).
fof(f14768,plain,(
  ( ! [X52,X50,X51] : (p(s(bool,X50)) | ~p(s(bool,X51)) | s(bool,X50) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK21(sK17,sK16,sK15,sK13)))) | s(bool,X51) != s(bool,X52)) )),
  inference(resolution,[],[f797,f558])).
fof(f14765,plain,(
  ( ! [X43,X41,X42] : (p(s(bool,X41)) | ~p(s(bool,X42)) | p(s(bool,X43)) | s(bool,X43) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK21(sK17,sK16,sK15,sK13))))) )),
  inference(resolution,[],[f797,f557])).
fof(f14764,plain,(
  ( ! [X39,X38,X40] : (p(s(bool,X38)) | p(s(bool,X39)) | s(bool,X38) != s(bool,X40) | s(bool,X39) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK21(sK17,sK16,sK15,sK13))))) )),
  inference(resolution,[],[f797,f556])).
fof(f14763,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK21(sK17,sK16,sK15,sK13)))) | s(bool,X36) != s(bool,X37)) )),
  inference(resolution,[],[f797,f554])).
fof(f14761,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK21(sK17,sK16,sK15,sK13))))) )),
  inference(resolution,[],[f797,f552])).
fof(f14745,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK21(sK17,sK16,sK15,sK13))))) | p(s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK21(sK17,sK16,sK15,sK13)))))),
  inference(resolution,[],[f797,f531])).
fof(f14658,plain,(
  ( ! [X212,X213] : (p(s(bool,X212)) | p(s(bool,X213)) | s(bool,X213) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16)))) ) | $spl32),
  inference(resolution,[],[f14552,f783])).
fof(f14552,plain,(
  p(s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16)))) | $spl32),
  inference(forward_demodulation,[],[f4458,f4642])).
fof(f14560,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16))) | s(bool,X22) != s(bool,X23)) ) | $spl32),
  inference(resolution,[],[f14552,f558])).
fof(f14557,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16)))) ) | $spl32),
  inference(resolution,[],[f14552,f557])).
fof(f14556,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16)))) ) | $spl32),
  inference(resolution,[],[f14552,f556])).
fof(f14555,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16))) | s(bool,X7) != s(bool,X8)) ) | $spl32),
  inference(resolution,[],[f14552,f554])).
fof(f14553,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK16)))) ) | $spl32),
  inference(resolution,[],[f14552,f552])).
fof(f14397,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17))) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) | $spl100),
  inference(cnf_transformation,[],[f14397_D])).
fof(f14397_D,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17))) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f14230,plain,(
  ( ! [X315,X316,X317] : (s(bool,X315) = s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17))) | p(s(bool,X316)) | ~p(s(bool,X317)) | s(bool,X315) != s(bool,X317)) )),
  inference(resolution,[],[f922,f783])).
fof(f922,plain,(
  ( ! [X213] : (~p(s(bool,X213)) | s(bool,X213) = s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)))) )),
  inference(resolution,[],[f499,f700])).
fof(f700,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f435])).
fof(f14132,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17))) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) )),
  inference(resolution,[],[f922,f558])).
fof(f14130,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,X20) = s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17))) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X20) != s(bool,X22)) )),
  inference(resolution,[],[f922,f557])).
fof(f14128,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17))) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) )),
  inference(resolution,[],[f922,f556])).
fof(f14127,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17))) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) )),
  inference(resolution,[],[f922,f554])).
fof(f14125,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17))) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) )),
  inference(resolution,[],[f922,f552])).
fof(f12764,plain,(
  ( ! [X241,X242] : (p(s(bool,X241)) | p(s(bool,X242)) | s(bool,X242) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl10),
  inference(resolution,[],[f3109,f783])).
fof(f12666,plain,(
  ( ! [X52,X50,X51] : (p(s(bool,X50)) | ~p(s(bool,X51)) | s(bool,X50) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK16,sK15,sK13)),s(sK13,sK17))) | s(bool,X51) != s(bool,X52)) ) | $spl10),
  inference(resolution,[],[f3109,f558])).
fof(f12663,plain,(
  ( ! [X43,X41,X42] : (p(s(bool,X41)) | ~p(s(bool,X42)) | p(s(bool,X43)) | s(bool,X43) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl10),
  inference(resolution,[],[f3109,f557])).
fof(f12662,plain,(
  ( ! [X39,X38,X40] : (p(s(bool,X38)) | p(s(bool,X39)) | s(bool,X38) != s(bool,X40) | s(bool,X39) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl10),
  inference(resolution,[],[f3109,f556])).
fof(f12661,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK16,sK15,sK13)),s(sK13,sK17))) | s(bool,X36) != s(bool,X37)) ) | $spl10),
  inference(resolution,[],[f3109,f554])).
fof(f12659,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl10),
  inference(resolution,[],[f3109,f552])).
fof(f12643,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) | p(s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) | $spl10),
  inference(resolution,[],[f3109,f531])).
fof(f12828,plain,(
  s(sK13,sK17) = s(sK13,sK20(sK17,sK16,sK15,sK13)) | $spl84),
  inference(cnf_transformation,[],[f12828_D])).
fof(f12828_D,plain,(
  s(sK13,sK17) = s(sK13,sK20(sK17,sK16,sK15,sK13)) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f12640,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) | s(sK13,sK17) = s(sK13,sK20(sK17,sK16,sK15,sK13)) | $spl10),
  inference(resolution,[],[f3109,f526])).
fof(f12836,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK20(sK17,sK16,sK15,sK13),sK15,sK13)))),s(sK13,sK17)))) | $spl92),
  inference(cnf_transformation,[],[f12836_D])).
fof(f12836_D,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK20(sK17,sK16,sK15,sK13),sK15,sK13)))),s(sK13,sK17)))) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f12638,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK20(sK17,sK16,sK15,sK13),sK15,sK13)))),s(sK13,sK17)))) | s(sK13,sK17) = s(sK13,sK20(sK17,sK16,sK15,sK13)) | $spl10),
  inference(resolution,[],[f3109,f523])).
fof(f12834,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK16,sK15,sK13)),s(sK13,sK28(sK17,sK20(sK17,sK16,sK15,sK13),sK15,sK13))))) | $spl90),
  inference(cnf_transformation,[],[f12834_D])).
fof(f12834_D,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK16,sK15,sK13)),s(sK13,sK28(sK17,sK20(sK17,sK16,sK15,sK13),sK15,sK13))))) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f12637,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK16,sK15,sK13)),s(sK13,sK28(sK17,sK20(sK17,sK16,sK15,sK13),sK15,sK13))))) | s(sK13,sK17) = s(sK13,sK20(sK17,sK16,sK15,sK13)) | $spl10),
  inference(resolution,[],[f3109,f522])).
fof(f12832,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK20(sK17,sK16,sK15,sK13),sK15,sK13)),s(sK13,sK17)))) | $spl88),
  inference(cnf_transformation,[],[f12832_D])).
fof(f12832_D,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK20(sK17,sK16,sK15,sK13),sK15,sK13)),s(sK13,sK17)))) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f12635,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK20(sK17,sK16,sK15,sK13),sK15,sK13)),s(sK13,sK17)))) | s(sK13,sK17) = s(sK13,sK20(sK17,sK16,sK15,sK13)) | $spl10),
  inference(resolution,[],[f3109,f519])).
fof(f12830,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK16,sK15,sK13)))),s(sK13,sK27(sK17,sK20(sK17,sK16,sK15,sK13),sK15,sK13))))) | $spl86),
  inference(cnf_transformation,[],[f12830_D])).
fof(f12830_D,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK16,sK15,sK13)))),s(sK13,sK27(sK17,sK20(sK17,sK16,sK15,sK13),sK15,sK13))))) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f12634,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK20(sK17,sK16,sK15,sK13)))),s(sK13,sK27(sK17,sK20(sK17,sK16,sK15,sK13),sK15,sK13))))) | s(sK13,sK17) = s(sK13,sK20(sK17,sK16,sK15,sK13)) | $spl10),
  inference(resolution,[],[f3109,f518])).
fof(f11175,plain,(
  ( ! [X241,X242] : (p(s(bool,X241)) | p(s(bool,X242)) | s(bool,X242) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) | $spl6),
  inference(resolution,[],[f1004,f783])).
fof(f11077,plain,(
  ( ! [X52,X50,X51] : (p(s(bool,X50)) | ~p(s(bool,X51)) | s(bool,X50) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13)))) | s(bool,X51) != s(bool,X52)) ) | $spl6),
  inference(resolution,[],[f1004,f558])).
fof(f11074,plain,(
  ( ! [X43,X41,X42] : (p(s(bool,X41)) | ~p(s(bool,X42)) | p(s(bool,X43)) | s(bool,X43) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) | $spl6),
  inference(resolution,[],[f1004,f557])).
fof(f11073,plain,(
  ( ! [X39,X38,X40] : (p(s(bool,X38)) | p(s(bool,X39)) | s(bool,X38) != s(bool,X40) | s(bool,X39) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) | $spl6),
  inference(resolution,[],[f1004,f556])).
fof(f11072,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13)))) | s(bool,X36) != s(bool,X37)) ) | $spl6),
  inference(resolution,[],[f1004,f554])).
fof(f11070,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) ) | $spl6),
  inference(resolution,[],[f1004,f552])).
fof(f11054,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) | p(s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) | $spl6),
  inference(resolution,[],[f1004,f531])).
fof(f11051,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) | s(sK13,sK16) = s(sK13,sK28(sK17,sK16,sK15,sK13)) | $spl6),
  inference(resolution,[],[f1004,f526])).
fof(f11049,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK28(sK17,sK16,sK15,sK13),sK16,sK15,sK13)))),s(sK13,sK28(sK17,sK16,sK15,sK13))))) | s(sK13,sK16) = s(sK13,sK28(sK17,sK16,sK15,sK13)) | $spl6),
  inference(resolution,[],[f1004,f523])).
fof(f11048,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK28(sK17,sK16,sK15,sK13),sK16,sK15,sK13))))) | s(sK13,sK16) = s(sK13,sK28(sK17,sK16,sK15,sK13)) | $spl6),
  inference(resolution,[],[f1004,f522])).
fof(f11046,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK28(sK17,sK16,sK15,sK13),sK16,sK15,sK13)),s(sK13,sK28(sK17,sK16,sK15,sK13))))) | s(sK13,sK16) = s(sK13,sK28(sK17,sK16,sK15,sK13)) | $spl6),
  inference(resolution,[],[f1004,f519])).
fof(f11045,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16))),s(sK13,sK27(sK28(sK17,sK16,sK15,sK13),sK16,sK15,sK13))))) | s(sK13,sK16) = s(sK13,sK28(sK17,sK16,sK15,sK13)) | $spl6),
  inference(resolution,[],[f1004,f518])).
fof(f10964,plain,(
  ( ! [X241,X242] : (p(s(bool,X241)) | p(s(bool,X242)) | s(bool,X242) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl4),
  inference(resolution,[],[f1002,f783])).
fof(f10866,plain,(
  ( ! [X52,X50,X51] : (p(s(bool,X50)) | ~p(s(bool,X51)) | s(bool,X50) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17))) | s(bool,X51) != s(bool,X52)) ) | $spl4),
  inference(resolution,[],[f1002,f558])).
fof(f10863,plain,(
  ( ! [X43,X41,X42] : (p(s(bool,X41)) | ~p(s(bool,X42)) | p(s(bool,X43)) | s(bool,X43) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl4),
  inference(resolution,[],[f1002,f557])).
fof(f10862,plain,(
  ( ! [X39,X38,X40] : (p(s(bool,X38)) | p(s(bool,X39)) | s(bool,X38) != s(bool,X40) | s(bool,X39) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl4),
  inference(resolution,[],[f1002,f556])).
fof(f10861,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17))) | s(bool,X36) != s(bool,X37)) ) | $spl4),
  inference(resolution,[],[f1002,f554])).
fof(f10859,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl4),
  inference(resolution,[],[f1002,f552])).
fof(f10843,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) | p(s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) | $spl4),
  inference(resolution,[],[f1002,f531])).
fof(f9844,plain,(
  $false | $spl65),
  inference(subsumption_resolution,[],[f7368,f499])).
fof(f3294,plain,(
  ( ! [X15] : (p(s(bool,X15)) | s(bool,X15) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)))) ) | $spl30),
  inference(cnf_transformation,[],[f3294_D])).
fof(f3294_D,plain,(
  ( ! [X15] : (p(s(bool,X15)) | s(bool,X15) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f8801,plain,(
  ( ! [X227,X226] : (p(s(bool,X226)) | p(s(bool,X227)) | s(bool,X227) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)))) ) | $spl10),
  inference(resolution,[],[f1008,f783])).
fof(f8703,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17))) | s(bool,X36) != s(bool,X37)) ) | $spl10),
  inference(resolution,[],[f1008,f558])).
fof(f8700,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | p(s(bool,X28)) | s(bool,X28) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)))) ) | $spl10),
  inference(resolution,[],[f1008,f557])).
fof(f8699,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X23) != s(bool,X25) | s(bool,X24) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)))) ) | $spl10),
  inference(resolution,[],[f1008,f556])).
fof(f8698,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X20) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17))) | s(bool,X21) != s(bool,X22)) ) | $spl10),
  inference(resolution,[],[f1008,f554])).
fof(f8696,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)))) ) | $spl10),
  inference(resolution,[],[f1008,f552])).
fof(f1236,plain,(
  ( ! [X13] : (~p(s(bool,X13)) | s(bool,X13) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17)))))) ) | $spl24),
  inference(cnf_transformation,[],[f1236_D])).
fof(f1236_D,plain,(
  ( ! [X13] : (~p(s(bool,X13)) | s(bool,X13) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17)))))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f8320,plain,(
  ( ! [X222,X223] : (p(s(bool,X222)) | ~p(s(bool,X223)) | s(bool,X223) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17)))))) )),
  inference(resolution,[],[f500,f783])).
fof(f8222,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X34) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17))))) | s(bool,X33) != s(bool,X35)) )),
  inference(resolution,[],[f500,f558])).
fof(f8220,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | ~p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X28) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17)))))) )),
  inference(resolution,[],[f500,f557])).
fof(f8218,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17)))))) )),
  inference(resolution,[],[f500,f556])).
fof(f8217,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17))))) | s(bool,X18) != s(bool,X20)) )),
  inference(resolution,[],[f500,f554])).
fof(f8215,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17)))))) )),
  inference(resolution,[],[f500,f552])).
fof(f2695,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)))) ) | $spl28),
  inference(cnf_transformation,[],[f2695_D])).
fof(f2695_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f7983,plain,(
  ( ! [X212,X213] : (p(s(bool,X212)) | p(s(bool,X213)) | s(bool,X213) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)))) )),
  inference(resolution,[],[f814,f783])).
fof(f814,plain,(
  p(s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17))))),
  inference(resolution,[],[f499,f534])).
fof(f7885,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17))) | s(bool,X22) != s(bool,X23)) )),
  inference(resolution,[],[f814,f558])).
fof(f7882,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)))) )),
  inference(resolution,[],[f814,f557])).
fof(f7881,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)))) )),
  inference(resolution,[],[f814,f556])).
fof(f7880,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17))) | s(bool,X7) != s(bool,X8)) )),
  inference(resolution,[],[f814,f554])).
fof(f7878,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)))) )),
  inference(resolution,[],[f814,f552])).
fof(f1011,plain,(
  ( ! [X30] : (p(s(bool,X30)) | s(bool,X30) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)))) ) | $spl14),
  inference(cnf_transformation,[],[f1011_D])).
fof(f1011_D,plain,(
  ( ! [X30] : (p(s(bool,X30)) | s(bool,X30) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f7806,plain,(
  ( ! [X241,X242] : (p(s(bool,X241)) | p(s(bool,X242)) | s(bool,X242) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)))) )),
  inference(resolution,[],[f499,f783])).
fof(f7708,plain,(
  ( ! [X52,X50,X51] : (p(s(bool,X50)) | ~p(s(bool,X51)) | s(bool,X50) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17))) | s(bool,X51) != s(bool,X52)) )),
  inference(resolution,[],[f499,f558])).
fof(f7705,plain,(
  ( ! [X43,X41,X42] : (p(s(bool,X41)) | ~p(s(bool,X42)) | p(s(bool,X43)) | s(bool,X43) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)))) )),
  inference(resolution,[],[f499,f557])).
fof(f7704,plain,(
  ( ! [X39,X38,X40] : (p(s(bool,X38)) | p(s(bool,X39)) | s(bool,X38) != s(bool,X40) | s(bool,X39) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)))) )),
  inference(resolution,[],[f499,f556])).
fof(f7703,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17))) | s(bool,X36) != s(bool,X37)) )),
  inference(resolution,[],[f499,f554])).
fof(f7701,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)))) )),
  inference(resolution,[],[f499,f552])).
fof(f1010,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)))) | $spl12),
  inference(cnf_transformation,[],[f1010_D])).
fof(f1010_D,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f1008,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)))) | $spl10),
  inference(cnf_transformation,[],[f1008_D])).
fof(f1008_D,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f7685,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)))) | p(s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17))))),
  inference(resolution,[],[f499,f531])).
fof(f7368,plain,(
  ~p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)))) | $spl65),
  inference(cnf_transformation,[],[f7368_D])).
fof(f7368_D,plain,(
  ~p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)))) <=> ~$spl65),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl65])])).
fof(f7370,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16))),s(sK13,sK17)))) | $spl66),
  inference(cnf_transformation,[],[f7370_D])).
fof(f7370_D,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16))),s(sK13,sK17)))) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f7366,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16))),s(sK13,sK17)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)))) | ($spl1 | $spl48)),
  inference(subsumption_resolution,[],[f7035,f997])).
fof(f997,plain,(
  s(sK13,sK16) != s(sK13,sK17) | $spl1),
  inference(cnf_transformation,[],[f997_D])).
fof(f997_D,plain,(
  s(sK13,sK16) != s(sK13,sK17) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f7035,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16))),s(sK13,sK17)))) | s(sK13,sK16) = s(sK13,sK17) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)))) | $spl48),
  inference(superposition,[],[f518,f6304])).
fof(f6706,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f6319,f550])).
fof(f550,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X2),s(X0,X1),s(X0,X1))))) )),
  inference(cnf_transformation,[],[f180])).
fof(f180,plain,(
  ! [X0,X1,X2] : p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X2),s(X0,X1),s(X0,X1))))),
  inference(flattening,[],[f179])).
fof(f179,plain,(
  ! [X0] : ! [X1,X2] : p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X2),s(X0,X1),s(X0,X1))))),
  inference(rectify,[],[f95])).
fof(f95,axiom,(
  ! [X6] : ! [X9,X8] : p(s(bool,'const.relation.RTC_3'(s(fun(X6,fun(X6,bool)),X8),s(X6,X9),s(X6,X9))))),
  file('thm.relation.RTC_lifts_monotonicities','thm.relation.RTC_REFL')).
fof(f6319,plain,(
  ~p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16)))))) | $spl0),
  inference(backward_demodulation,[],[f998,f500])).
fof(f6240,plain,(
  ( ! [X241,X242] : (p(s(bool,X241)) | p(s(bool,X242)) | s(bool,X242) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl4),
  inference(resolution,[],[f1002,f783])).
fof(f6142,plain,(
  ( ! [X52,X50,X51] : (p(s(bool,X50)) | ~p(s(bool,X51)) | s(bool,X50) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17))) | s(bool,X51) != s(bool,X52)) ) | $spl4),
  inference(resolution,[],[f1002,f558])).
fof(f6139,plain,(
  ( ! [X43,X41,X42] : (p(s(bool,X41)) | ~p(s(bool,X42)) | p(s(bool,X43)) | s(bool,X43) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl4),
  inference(resolution,[],[f1002,f557])).
fof(f6138,plain,(
  ( ! [X39,X38,X40] : (p(s(bool,X38)) | p(s(bool,X39)) | s(bool,X38) != s(bool,X40) | s(bool,X39) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl4),
  inference(resolution,[],[f1002,f556])).
fof(f6137,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17))) | s(bool,X36) != s(bool,X37)) ) | $spl4),
  inference(resolution,[],[f1002,f554])).
fof(f6135,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) ) | $spl4),
  inference(resolution,[],[f1002,f552])).
fof(f6119,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) | p(s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) | $spl4),
  inference(resolution,[],[f1002,f531])).
fof(f6116,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) | s(sK13,sK17) = s(sK13,sK27(sK17,sK16,sK15,sK13)) | $spl4),
  inference(resolution,[],[f1002,f526])).
fof(f6114,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK27(sK17,sK16,sK15,sK13),sK15,sK13)))),s(sK13,sK17)))) | s(sK13,sK17) = s(sK13,sK27(sK17,sK16,sK15,sK13)) | $spl4),
  inference(resolution,[],[f1002,f523])).
fof(f6113,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK28(sK17,sK27(sK17,sK16,sK15,sK13),sK15,sK13))))) | s(sK13,sK17) = s(sK13,sK27(sK17,sK16,sK15,sK13)) | $spl4),
  inference(resolution,[],[f1002,f522])).
fof(f6111,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK27(sK17,sK16,sK15,sK13),sK15,sK13)),s(sK13,sK17)))) | s(sK13,sK17) = s(sK13,sK27(sK17,sK16,sK15,sK13)) | $spl4),
  inference(resolution,[],[f1002,f519])).
fof(f6110,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)))),s(sK13,sK27(sK17,sK27(sK17,sK16,sK15,sK13),sK15,sK13))))) | s(sK13,sK17) = s(sK13,sK27(sK17,sK16,sK15,sK13)) | $spl4),
  inference(resolution,[],[f1002,f518])).
fof(f4578,plain,(
  ( ! [X241,X242] : (p(s(bool,X241)) | p(s(bool,X242)) | s(bool,X242) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK21(sK17,sK16,sK15,sK13))))) )),
  inference(resolution,[],[f797,f783])).
fof(f4480,plain,(
  ( ! [X52,X50,X51] : (p(s(bool,X50)) | ~p(s(bool,X51)) | s(bool,X50) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK21(sK17,sK16,sK15,sK13)))) | s(bool,X51) != s(bool,X52)) )),
  inference(resolution,[],[f797,f558])).
fof(f4477,plain,(
  ( ! [X43,X41,X42] : (p(s(bool,X41)) | ~p(s(bool,X42)) | p(s(bool,X43)) | s(bool,X43) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK21(sK17,sK16,sK15,sK13))))) )),
  inference(resolution,[],[f797,f557])).
fof(f4476,plain,(
  ( ! [X39,X38,X40] : (p(s(bool,X38)) | p(s(bool,X39)) | s(bool,X38) != s(bool,X40) | s(bool,X39) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK21(sK17,sK16,sK15,sK13))))) )),
  inference(resolution,[],[f797,f556])).
fof(f4475,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK21(sK17,sK16,sK15,sK13)))) | s(bool,X36) != s(bool,X37)) )),
  inference(resolution,[],[f797,f554])).
fof(f4473,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK21(sK17,sK16,sK15,sK13))))) )),
  inference(resolution,[],[f797,f552])).
fof(f4457,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK21(sK17,sK16,sK15,sK13))))) | p(s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK21(sK17,sK16,sK15,sK13)))))),
  inference(resolution,[],[f797,f531])).
fof(f4642,plain,(
  s(sK13,sK16) = s(sK13,sK21(sK17,sK16,sK15,sK13)) | $spl32),
  inference(cnf_transformation,[],[f4642_D])).
fof(f4642_D,plain,(
  s(sK13,sK16) = s(sK13,sK21(sK17,sK16,sK15,sK13)) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f4454,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK21(sK17,sK16,sK15,sK13))))) | s(sK13,sK16) = s(sK13,sK21(sK17,sK16,sK15,sK13))),
  inference(resolution,[],[f797,f526])).
fof(f4650,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK21(sK17,sK16,sK15,sK13),sK16,sK15,sK13)))),s(sK13,sK21(sK17,sK16,sK15,sK13))))) | $spl40),
  inference(cnf_transformation,[],[f4650_D])).
fof(f4650_D,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK21(sK17,sK16,sK15,sK13),sK16,sK15,sK13)))),s(sK13,sK21(sK17,sK16,sK15,sK13))))) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f4452,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK21(sK17,sK16,sK15,sK13),sK16,sK15,sK13)))),s(sK13,sK21(sK17,sK16,sK15,sK13))))) | s(sK13,sK16) = s(sK13,sK21(sK17,sK16,sK15,sK13))),
  inference(resolution,[],[f797,f523])).
fof(f4648,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK21(sK17,sK16,sK15,sK13),sK16,sK15,sK13))))) | $spl38),
  inference(cnf_transformation,[],[f4648_D])).
fof(f4648_D,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK21(sK17,sK16,sK15,sK13),sK16,sK15,sK13))))) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f4451,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK21(sK17,sK16,sK15,sK13),sK16,sK15,sK13))))) | s(sK13,sK16) = s(sK13,sK21(sK17,sK16,sK15,sK13))),
  inference(resolution,[],[f797,f522])).
fof(f4646,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK21(sK17,sK16,sK15,sK13),sK16,sK15,sK13)),s(sK13,sK21(sK17,sK16,sK15,sK13))))) | $spl36),
  inference(cnf_transformation,[],[f4646_D])).
fof(f4646_D,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK21(sK17,sK16,sK15,sK13),sK16,sK15,sK13)),s(sK13,sK21(sK17,sK16,sK15,sK13))))) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f4449,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK21(sK17,sK16,sK15,sK13),sK16,sK15,sK13)),s(sK13,sK21(sK17,sK16,sK15,sK13))))) | s(sK13,sK16) = s(sK13,sK21(sK17,sK16,sK15,sK13))),
  inference(resolution,[],[f797,f519])).
fof(f4644,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16))),s(sK13,sK27(sK21(sK17,sK16,sK15,sK13),sK16,sK15,sK13))))) | $spl34),
  inference(cnf_transformation,[],[f4644_D])).
fof(f4644_D,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16))),s(sK13,sK27(sK21(sK17,sK16,sK15,sK13),sK16,sK15,sK13))))) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f4448,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16))),s(sK13,sK27(sK21(sK17,sK16,sK15,sK13),sK16,sK15,sK13))))) | s(sK13,sK16) = s(sK13,sK21(sK17,sK16,sK15,sK13))),
  inference(resolution,[],[f797,f518])).
fof(f3231,plain,(
  ( ! [X227,X226] : (p(s(bool,X226)) | p(s(bool,X227)) | s(bool,X227) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)))) ) | $spl10),
  inference(resolution,[],[f1008,f783])).
fof(f3133,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17))) | s(bool,X36) != s(bool,X37)) ) | $spl10),
  inference(resolution,[],[f1008,f558])).
fof(f3130,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | p(s(bool,X28)) | s(bool,X28) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)))) ) | $spl10),
  inference(resolution,[],[f1008,f557])).
fof(f3129,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X23) != s(bool,X25) | s(bool,X24) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)))) ) | $spl10),
  inference(resolution,[],[f1008,f556])).
fof(f3128,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X20) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17))) | s(bool,X21) != s(bool,X22)) ) | $spl10),
  inference(resolution,[],[f1008,f554])).
fof(f3126,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)))) ) | $spl10),
  inference(resolution,[],[f1008,f552])).
fof(f2632,plain,(
  ( ! [X212,X213] : (p(s(bool,X212)) | p(s(bool,X213)) | s(bool,X213) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)))) )),
  inference(resolution,[],[f814,f783])).
fof(f2534,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17))) | s(bool,X22) != s(bool,X23)) )),
  inference(resolution,[],[f814,f558])).
fof(f2531,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)))) )),
  inference(resolution,[],[f814,f557])).
fof(f2530,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)))) )),
  inference(resolution,[],[f814,f556])).
fof(f2529,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17))) | s(bool,X7) != s(bool,X8)) )),
  inference(resolution,[],[f814,f554])).
fof(f2527,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.relation.EQC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)))) )),
  inference(resolution,[],[f814,f552])).
fof(f2526,plain,(
  ( ! [X70,X72,X69] : (~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X69))),s(sK13,X70)))) | s(bool,X72) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,X69))))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,X70))))) | p(s(bool,X72))) ) | $spl26),
  inference(cnf_transformation,[],[f2526_D])).
fof(f2526_D,plain,(
  ( ! [X70,X72,X69] : (~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X69))),s(sK13,X70)))) | s(bool,X72) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,X69))))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,X70))))) | p(s(bool,X72))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f2402,plain,(
  ( ! [X492,X494,X493,X491] : (~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X491))),s(sK13,X492)))) | p(s(bool,X493)) | p(s(bool,X494)) | s(bool,X494) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,X491))))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,X492)))))) )),
  inference(resolution,[],[f498,f783])).
fof(f2304,plain,(
  ( ! [X107,X105,X108,X106,X104] : (~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X104))),s(sK13,X105)))) | p(s(bool,X106)) | ~p(s(bool,X107)) | s(bool,X106) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,X104))))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,X105))))) | s(bool,X107) != s(bool,X108)) )),
  inference(resolution,[],[f498,f558])).
fof(f2301,plain,(
  ( ! [X92,X90,X93,X91,X89] : (~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X89))),s(sK13,X90)))) | p(s(bool,X91)) | ~p(s(bool,X92)) | p(s(bool,X93)) | s(bool,X93) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,X89))))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,X90)))))) )),
  inference(resolution,[],[f498,f557])).
fof(f2300,plain,(
  ( ! [X88,X87,X85,X86,X84] : (~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X84))),s(sK13,X85)))) | p(s(bool,X86)) | p(s(bool,X87)) | s(bool,X86) != s(bool,X88) | s(bool,X87) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,X84))))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,X85)))))) )),
  inference(resolution,[],[f498,f556])).
fof(f2299,plain,(
  ( ! [X80,X83,X81,X79,X82] : (~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X79))),s(sK13,X80)))) | p(s(bool,X81)) | ~p(s(bool,X82)) | s(bool,X81) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,X79))))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,X80))))) | s(bool,X82) != s(bool,X83)) )),
  inference(resolution,[],[f498,f554])).
fof(f2297,plain,(
  ( ! [X70,X72,X71,X69,X73] : (~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X69))),s(sK13,X70)))) | p(s(bool,X71)) | p(s(bool,X72)) | s(bool,X71) != s(bool,X73) | s(bool,X72) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,X69))))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,X70)))))) )),
  inference(resolution,[],[f498,f552])).
fof(f2272,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2169,f1956])).
fof(f1956,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl18),
  inference(equality_resolution,[],[f1013])).
fof(f2169,plain,(
  ( ! [X761,X763,X759,X760,X762] : (p(s(bool,happ(s(fun(X759,bool),happ(s(fun(X759,fun(X759,bool)),sK26(X760,X761,X762,X759)),s(X759,X763))),s(X759,X763))))) ) | $spl18),
  inference(resolution,[],[f1956,f515])).
fof(f2271,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2128,f1956])).
fof(f2128,plain,(
  ( ! [X617,X619,X616,X618] : (p(s(bool,'const.relation.RTC_3'(s(fun(X616,fun(X616,bool)),X617),s(X616,X618),s(X616,X619))))) ) | $spl18),
  inference(resolution,[],[f1956,f515])).
fof(f2162,plain,(
  $false | $spl18),
  inference(resolution,[],[f1956,f550])).
fof(f2163,plain,(
  $false | $spl18),
  inference(resolution,[],[f1956,f549])).
fof(f549,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2))))) )),
  inference(cnf_transformation,[],[f178])).
fof(f178,plain,(
  ! [X0,X1,X2] : p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2))))),
  inference(flattening,[],[f177])).
fof(f177,plain,(
  ! [X0] : ! [X1,X2] : p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2))))),
  inference(rectify,[],[f65])).
fof(f65,axiom,(
  ! [X6] : ! [X8,X9] : p(s(bool,'const.relation.RTC_3'(s(fun(X6,fun(X6,bool)),X8),s(X6,X9),s(X6,X9))))),
  file('thm.relation.RTC_lifts_monotonicities','thm.relation.RTC_RULES|split|0')).
fof(f2189,plain,(
  $false | $spl18),
  inference(resolution,[],[f1956,f499])).
fof(f1825,plain,(
  $false | $spl16),
  inference(resolution,[],[f1594,f500])).
fof(f1594,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl16),
  inference(equality_resolution,[],[f1012])).
fof(f1463,plain,(
  $false | $spl22),
  inference(resolution,[],[f1015,f500])).
fof(f1139,plain,(
  ( ! [X222,X223] : (p(s(bool,X222)) | ~p(s(bool,X223)) | s(bool,X223) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17)))))) )),
  inference(resolution,[],[f500,f783])).
fof(f1041,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X34) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17))))) | s(bool,X33) != s(bool,X35)) )),
  inference(resolution,[],[f500,f558])).
fof(f1039,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | ~p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X28) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17)))))) )),
  inference(resolution,[],[f500,f557])).
fof(f1037,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17)))))) )),
  inference(resolution,[],[f500,f556])).
fof(f1036,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17))))) | s(bool,X18) != s(bool,X20)) )),
  inference(resolution,[],[f500,f554])).
fof(f1034,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17)))))) )),
  inference(resolution,[],[f500,f552])).
fof(f934,plain,(
  ( ! [X241,X242] : (p(s(bool,X241)) | p(s(bool,X242)) | s(bool,X242) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)))) )),
  inference(resolution,[],[f499,f783])).
fof(f836,plain,(
  ( ! [X52,X50,X51] : (p(s(bool,X50)) | ~p(s(bool,X51)) | s(bool,X50) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17))) | s(bool,X51) != s(bool,X52)) )),
  inference(resolution,[],[f499,f558])).
fof(f833,plain,(
  ( ! [X43,X41,X42] : (p(s(bool,X41)) | ~p(s(bool,X42)) | p(s(bool,X43)) | s(bool,X43) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)))) )),
  inference(resolution,[],[f499,f557])).
fof(f832,plain,(
  ( ! [X39,X38,X40] : (p(s(bool,X38)) | p(s(bool,X39)) | s(bool,X38) != s(bool,X40) | s(bool,X39) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)))) )),
  inference(resolution,[],[f499,f556])).
fof(f831,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17))) | s(bool,X36) != s(bool,X37)) )),
  inference(resolution,[],[f499,f554])).
fof(f829,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)))) )),
  inference(resolution,[],[f499,f552])).
fof(f813,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)))) | p(s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17))))),
  inference(resolution,[],[f499,f531])).
fof(f998,plain,(
  s(sK13,sK16) = s(sK13,sK17) | $spl0),
  inference(cnf_transformation,[],[f998_D])).
fof(f998_D,plain,(
  s(sK13,sK16) = s(sK13,sK17) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f810,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)))) | s(sK13,sK16) = s(sK13,sK17)),
  inference(resolution,[],[f499,f526])).
fof(f1006,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK16,sK15,sK13)))),s(sK13,sK17)))) | $spl8),
  inference(cnf_transformation,[],[f1006_D])).
fof(f1006_D,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK16,sK15,sK13)))),s(sK13,sK17)))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f808,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK28(sK17,sK16,sK15,sK13)))),s(sK13,sK17)))) | s(sK13,sK16) = s(sK13,sK17)),
  inference(resolution,[],[f499,f523])).
fof(f1004,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) | $spl6),
  inference(cnf_transformation,[],[f1004_D])).
fof(f1004_D,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f807,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK28(sK17,sK16,sK15,sK13))))) | s(sK13,sK16) = s(sK13,sK17)),
  inference(resolution,[],[f499,f522])).
fof(f1002,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) | $spl4),
  inference(cnf_transformation,[],[f1002_D])).
fof(f1002_D,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f805,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK27(sK17,sK16,sK15,sK13)),s(sK13,sK17)))) | s(sK13,sK16) = s(sK13,sK17)),
  inference(resolution,[],[f499,f519])).
fof(f1000,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16))),s(sK13,sK27(sK17,sK16,sK15,sK13))))) | $spl2),
  inference(cnf_transformation,[],[f1000_D])).
fof(f1000_D,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16))),s(sK13,sK27(sK17,sK16,sK15,sK13))))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f804,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16))),s(sK13,sK27(sK17,sK16,sK15,sK13))))) | s(sK13,sK16) = s(sK13,sK17)),
  inference(resolution,[],[f499,f518])).
% SZS output end Proof for thm.relation
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 1205
% Passive clauses: 45776
% Generated clauses: 65217
% Final active clauses: 778
% Final passive clauses: 33368
% Input formulas: 133
% Initial clauses: 285
% 
% Selected by SInE selection: 70
% SInE iterations: 1
% 
% Duplicate literals: 1061
% Fw subsumption resolutions: 812
% Fw demodulations: 3654
% Bw demodulations: 132
% 
% Simple tautologies: 2164
% Equational tautologies: 4
% Forward subsumptions: 11704
% 
% Binary resolution: 22415
% Forward superposition: 32806
% Backward superposition: 3476
% Equality resolution: 6
% 
% Split clauses: 729
% Split components: 1560
% SAT solver clauses: 741
% SAT solver unit clauses: 10
% SAT solver binary clauses: 232
% 
% Sat splits: 729
% Sat splitting refutations: 11
% 
% Memory used [KB]: 28016
% Time elapsed: 2.123 s
% ------------------------------
% ------------------------------
% Success in time 13.85 s
