Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.extreal
% (72446)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 227
% Passive clauses: 2628
% Generated clauses: 18512
% Final active clauses: 198
% Final passive clauses: 2167
% Input formulas: 27
% Initial clauses: 76
% 
% Duplicate literals: 890
% Trivial inequalities: 328
% Fw subsumption resolutions: 3039
% Fw demodulations: 1347
% Bw demodulations: 22
% 
% Simple tautologies: 392
% Equational tautologies: 43
% Forward subsumptions: 6801
% Bw demodulations to eq. taut.: 2
% 
% Binary resolution: 1091
% Factoring: 1
% Forward superposition: 4996
% Backward superposition: 6009
% Self superposition: 157
% Equality factoring: 130
% Equality resolution: 51
% 
% Split clauses: 1541
% Split components: 3223
% SAT solver clauses: 2216
% SAT solver unit clauses: 32
% SAT solver binary clauses: 288
% 
% Sat splits: 1541
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 2215
% TWLsolver calls for satisfiability: 306
% 
% Memory used [KB]: 9594
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.extreal
% (72465)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 285
% Passive clauses: 2496
% Generated clauses: 22823
% Final active clauses: 64
% Final passive clauses: 644
% Input formulas: 27
% Initial clauses: 76
% 
% Duplicate literals: 134
% Trivial inequalities: 444
% Fw subsumption resolutions: 1651
% Fw demodulations: 4222
% Bw demodulations: 1476
% Condensations: 224
% Global subsumptions: 8
% 
% Simple tautologies: 93
% Equational tautologies: 20
% Forward subsumptions: 8571
% Bw demodulations to eq. taut.: 168
% 
% Binary resolution: 490
% Forward superposition: 4339
% Backward superposition: 2326
% Self superposition: 48
% Equality factoring: 26
% Equality resolution: 33
% 
% Split clauses: 357
% Split components: 726
% SAT solver clauses: 2787
% SAT solver unit clauses: 79
% SAT solver binary clauses: 69
% SAT solver learnt clauses: 23
% SAT solver learnt literals: 23
% 
% Sat splits: 357
% Sat splitting refutations: 220
% 
% TWLsolver clauses: 2748
% TWLsolver calls for satisfiability: 569
% 
% Memory used [KB]: 12409
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.extreal
% (72490)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 936
% Passive clauses: 8397
% Generated clauses: 119793
% Final active clauses: 126
% Final passive clauses: 686
% Input formulas: 27
% Initial clauses: 76
% 
% Duplicate literals: 6653
% Trivial inequalities: 938
% Fw subsumption resolutions: 8594
% Fw demodulations: 12512
% Bw demodulations: 6476
% Condensations: 12133
% Global subsumptions: 1919
% 
% Simple tautologies: 1047
% Equational tautologies: 1960
% Forward subsumptions: 63000
% Bw demodulations to eq. taut.: 491
% 
% Unit resulting resolution: 6719
% Forward superposition: 30743
% Backward superposition: 17725
% Self superposition: 459
% Equality factoring: 344
% Equality resolution: 82
% 
% Split clauses: 778
% Split components: 1575
% SAT solver clauses: 13779
% SAT solver unit clauses: 216
% SAT solver binary clauses: 1762
% 
% Sat splits: 778
% Sat splitting refutations: 1905
% 
% Memory used [KB]: 54753
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 183 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.extreal
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.extreal
% SZS output start Proof for thm.extreal
fof(f2884,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f197,f198,f199,f200,f201,f203,f225,f204,f214,f223,f264,f294,f327,f329,f325,f324,f362,f363,f330,f359,f358,f328,f394,f435,f395,f436,f434,f396,f397,f398,f399,f400,f438,f401,f439,f403,f416,f417,f440,f491,f492,f493,f494,f495,f497,f518,f498,f508,f517,f524,f575,f576,f577,f578,f579,f581,f625,f582,f592,f624,f595,f596,f629,f598,f630,f623,f622,f628,f672,f673,f674,f675,f676,f678,f722,f226,f679,f689,f721,f692,f693,f726,f695,f727,f720,f719,f725,f997,f998,f1012,f1044,f1013,f1015,f1016,f1017,f1020,f1030,f1045,f1116,f1118,f1119,f1120,f1123,f1133,f1148,f1508,f1512,f1515,f1519,f1521,f1531,f437,f1537,f1538,f1539,f1540,f1628,f1630,f1631,f1632,f1635,f1645,f1662,f1920,f1922,f1923,f1924,f1927,f1937,f1972,f2008,f2012,f2010,f2014,f2018,f2016,f2020,f2024,f2022,f2001,f2028,f2026,f2034,f2036,f2037,f2039,f2041,f2054,f2080,f2082,f2083,f2084,f2087,f2097,f2115,f2117,f2118,f2120,f2122,f2135,f2074,f2163,f2165,f2166,f2168,f2170,f2183,f2200,f2202,f2203,f2204,f2207,f2217,f2224,f2254,f2256,f2257,f2259,f2261,f2274,f2194,f2295,f2297,f2298,f2300,f2302,f2315,f2341,f2343,f2344,f2346,f2348,f2361,f2380,f2382,f2383,f2384,f2387,f2397,f2406,f2415,f2417,f2418,f2420,f2422,f2435,f2457,f2459,f2460,f2462,f2464,f2477,f2332,f2502,f2504,f2505,f2506,f2509,f2519,f2529,f2537,f2539,f2540,f2541,f2544,f2554,f2104,f2572,f2574,f2575,f2576,f2579,f2589,f2608,f2610,f2611,f2612,f2615,f2625,f2669,f2699,f2705,f2707,f2708,f2710,f2712,f2725,f2372,f2753,f2755,f2756,f2757,f2760,f2770,f2633,f2789,f2791,f2792,f224,f2793,f228,f2796,f227,f2806,f229,f2599,f2854,f2864])).
fof(f2864,plain,(
  $false | $spl74),
  inference(subsumption_resolution,[],[f2863,f2027])).
fof(f2027,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK1,sK0)),s('type.extreal.extreal',sK1)))) | $spl74),
  inference(cnf_transformation,[],[f2027_D])).
fof(f2027_D,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK1,sK0)),s('type.extreal.extreal',sK1)))) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f2863,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK1,sK0)),s('type.extreal.extreal',sK1)))) | $spl74),
  inference(subsumption_resolution,[],[f2862,f116])).
fof(f116,plain,(
  s('type.extreal.extreal',sK1) != s('type.extreal.extreal','const.extreal.extreal_inf_1'(s(fun('type.extreal.extreal',bool),sK0)))),
  inference(cnf_transformation,[],[f73])).
fof(f73,plain,(
  p(s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK2)))) & ! [X2] : (~p(s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',X2)))) | s('type.extreal.extreal',sK1) = s('type.extreal.extreal',X2)) & s('type.extreal.extreal',sK1) != s('type.extreal.extreal','const.extreal.extreal_inf_1'(s(fun('type.extreal.extreal',bool),sK0)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f60])).
fof(f60,plain,(
  ? [X0,X1] : (? [X3] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal',X3)))) & ! [X2] : (~p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal',X2)))) | s('type.extreal.extreal',X1) = s('type.extreal.extreal',X2)) & s('type.extreal.extreal',X1) != s('type.extreal.extreal','const.extreal.extreal_inf_1'(s(fun('type.extreal.extreal',bool),X0))))),
  inference(flattening,[],[f59])).
fof(f59,plain,(
  ? [X0,X1] : ((? [X3] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal',X3)))) & ! [X2] : (~p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal',X2)))) | s('type.extreal.extreal',X1) = s('type.extreal.extreal',X2))) & s('type.extreal.extreal',X1) != s('type.extreal.extreal','const.extreal.extreal_inf_1'(s(fun('type.extreal.extreal',bool),X0))))),
  inference(ennf_transformation,[],[f29])).
fof(f29,plain,(
  ~! [X0,X1] : ((? [X3] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal',X3)))) & ! [X2] : (p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal',X2)))) => s('type.extreal.extreal',X1) = s('type.extreal.extreal',X2))) => s('type.extreal.extreal',X1) = s('type.extreal.extreal','const.extreal.extreal_inf_1'(s(fun('type.extreal.extreal',bool),X0))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ((? [X8] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),X6),s('type.extreal.extreal',X8)))) & ! [X8] : (p(s(bool,happ(s(fun('type.extreal.extreal',bool),X6),s('type.extreal.extreal',X8)))) => s('type.extreal.extreal',X7) = s('type.extreal.extreal',X8))) => s('type.extreal.extreal',X7) = s('type.extreal.extreal','const.extreal.extreal_inf_1'(s(fun('type.extreal.extreal',bool),X6))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ((? [X8] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),X6),s('type.extreal.extreal',X8)))) & ! [X8] : (p(s(bool,happ(s(fun('type.extreal.extreal',bool),X6),s('type.extreal.extreal',X8)))) => s('type.extreal.extreal',X7) = s('type.extreal.extreal',X8))) => s('type.extreal.extreal',X7) = s('type.extreal.extreal','const.extreal.extreal_inf_1'(s(fun('type.extreal.extreal',bool),X6))))),
  file('thm.extreal.inf_const_alt','thm.extreal.inf_const_alt')).
fof(f2862,plain,(
  s('type.extreal.extreal',sK1) = s('type.extreal.extreal','const.extreal.extreal_inf_1'(s(fun('type.extreal.extreal',bool),sK0))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK1,sK0)),s('type.extreal.extreal',sK1)))) | $spl74),
  inference(subsumption_resolution,[],[f2827,f189])).
fof(f189,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0))))) )),
  inference(cnf_transformation,[],[f58])).
fof(f58,plain,(
  ! [X0] : p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0))))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X8] : p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X8),s('type.extreal.extreal',X8))))),
  file('thm.extreal.inf_const_alt','thm.extreal.le_refl')).
fof(f2827,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK1)))) | s('type.extreal.extreal',sK1) = s('type.extreal.extreal','const.extreal.extreal_inf_1'(s(fun('type.extreal.extreal',bool),sK0))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK1,sK0)),s('type.extreal.extreal',sK1)))) | $spl74),
  inference(superposition,[],[f123,f2783])).
fof(f2783,plain,(
  s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK4(sK1,sK0)) | $spl74),
  inference(resolution,[],[f2777,f115])).
fof(f115,plain,(
  ( ! [X2] : (~p(s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',X2)))) | s('type.extreal.extreal',sK1) = s('type.extreal.extreal',X2)) )),
  inference(cnf_transformation,[],[f73])).
fof(f2777,plain,(
  p(s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK4(sK1,sK0))))) | $spl74),
  inference(subsumption_resolution,[],[f2751,f116])).
fof(f2751,plain,(
  s('type.extreal.extreal',sK1) = s('type.extreal.extreal','const.extreal.extreal_inf_1'(s(fun('type.extreal.extreal',bool),sK0))) | p(s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK4(sK1,sK0))))) | $spl74),
  inference(resolution,[],[f2027,f122])).
fof(f122,plain,(
  ( ! [X0,X1] : (s('type.extreal.extreal',X1) = s('type.extreal.extreal','const.extreal.extreal_inf_1'(s(fun('type.extreal.extreal',bool),X0))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(X1,X0)),s('type.extreal.extreal',X1)))) | p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal',sK4(X1,X0)))))) )),
  inference(cnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0,X1] : ((s('type.extreal.extreal',X1) != s('type.extreal.extreal','const.extreal.extreal_inf_1'(s(fun('type.extreal.extreal',bool),X0))) | (! [X7] : (~p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal',X7)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X7))))) & ! [X5] : ((p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal',sK3(X5,X1,X0))))) & ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X5),s('type.extreal.extreal',sK3(X5,X1,X0)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X5),s('type.extreal.extreal',X1))))))) & ((p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal',sK4(X1,X0))))) & ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',sK4(X1,X0)))))) | (! [X3] : (~p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal',X3)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(X1,X0)),s('type.extreal.extreal',X3))))) & ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(X1,X0)),s('type.extreal.extreal',X1))))) | s('type.extreal.extreal',X1) = s('type.extreal.extreal','const.extreal.extreal_inf_1'(s(fun('type.extreal.extreal',bool),X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3,sK4,sK5])],[f76])).
fof(f76,plain,(
  ! [X0,X1] : ((s('type.extreal.extreal',X1) != s('type.extreal.extreal','const.extreal.extreal_inf_1'(s(fun('type.extreal.extreal',bool),X0))) | (! [X7] : (~p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal',X7)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X7))))) & ! [X5] : (? [X6] : (p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal',X6)))) & ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X5),s('type.extreal.extreal',X6))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X5),s('type.extreal.extreal',X1))))))) & (? [X4] : (p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal',X4)))) & ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X4))))) | ? [X2] : (! [X3] : (~p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal',X3)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X3))))) & ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X1))))) | s('type.extreal.extreal',X1) = s('type.extreal.extreal','const.extreal.extreal_inf_1'(s(fun('type.extreal.extreal',bool),X0)))))),
  inference(rectify,[],[f75])).
fof(f75,plain,(
  ! [X0,X1] : ((s('type.extreal.extreal',X1) != s('type.extreal.extreal','const.extreal.extreal_inf_1'(s(fun('type.extreal.extreal',bool),X0))) | (! [X4] : (~p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal',X4)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X4))))) & ! [X2] : (? [X3] : (p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal',X3)))) & ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X3))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X1))))))) & (? [X4] : (p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal',X4)))) & ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X4))))) | ? [X2] : (! [X3] : (~p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal',X3)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X3))))) & ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X1))))) | s('type.extreal.extreal',X1) = s('type.extreal.extreal','const.extreal.extreal_inf_1'(s(fun('type.extreal.extreal',bool),X0)))))),
  inference(flattening,[],[f74])).
fof(f74,plain,(
  ! [X0,X1] : ((s('type.extreal.extreal',X1) != s('type.extreal.extreal','const.extreal.extreal_inf_1'(s(fun('type.extreal.extreal',bool),X0))) | (! [X4] : (~p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal',X4)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X4))))) & ! [X2] : (? [X3] : (p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal',X3)))) & ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X3))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X1))))))) & ((? [X4] : (p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal',X4)))) & ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X4))))) | ? [X2] : (! [X3] : (~p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal',X3)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X3))))) & ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X1)))))) | s('type.extreal.extreal',X1) = s('type.extreal.extreal','const.extreal.extreal_inf_1'(s(fun('type.extreal.extreal',bool),X0)))))),
  inference(nnf_transformation,[],[f61])).
fof(f61,plain,(
  ! [X0,X1] : (s('type.extreal.extreal',X1) = s('type.extreal.extreal','const.extreal.extreal_inf_1'(s(fun('type.extreal.extreal',bool),X0))) <=> (! [X4] : (~p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal',X4)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X4))))) & ! [X2] : (? [X3] : (p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal',X3)))) & ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X3))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X1)))))))),
  inference(ennf_transformation,[],[f30])).
fof(f30,plain,(
  ! [X0,X1] : (s('type.extreal.extreal',X1) = s('type.extreal.extreal','const.extreal.extreal_inf_1'(s(fun('type.extreal.extreal',bool),X0))) <=> (! [X4] : (p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal',X4)))) => p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X4))))) & ! [X2] : (! [X3] : (p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal',X3)))) => p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X3))))) => p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X1)))))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X6,X8] : (s('type.extreal.extreal',X8) = s('type.extreal.extreal','const.extreal.extreal_inf_1'(s(fun('type.extreal.extreal',bool),X6))) <=> (! [X11] : (p(s(bool,happ(s(fun('type.extreal.extreal',bool),X6),s('type.extreal.extreal',X11)))) => p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X8),s('type.extreal.extreal',X11))))) & ! [X11] : (! [X7] : (p(s(bool,happ(s(fun('type.extreal.extreal',bool),X6),s('type.extreal.extreal',X7)))) => p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X11),s('type.extreal.extreal',X7))))) => p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X11),s('type.extreal.extreal',X8)))))))),
  file('thm.extreal.inf_const_alt','thm.extreal.inf_eq')).
fof(f123,plain,(
  ( ! [X0,X1] : (s('type.extreal.extreal',X1) = s('type.extreal.extreal','const.extreal.extreal_inf_1'(s(fun('type.extreal.extreal',bool),X0))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(X1,X0)),s('type.extreal.extreal',X1)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',sK4(X1,X0)))))) )),
  inference(cnf_transformation,[],[f77])).
fof(f2854,plain,(
  $false | ($spl73 | $spl74)),
  inference(subsumption_resolution,[],[f2823,f189])).
fof(f2823,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK1)))) | ($spl73 | $spl74)),
  inference(backward_demodulation,[],[f2783,f2026])).
fof(f2599,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,X6) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK4(sK1,sK0))))) ) | $spl92),
  inference(cnf_transformation,[],[f2599_D])).
fof(f2599_D,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,X6) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK4(sK1,sK0))))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f229,plain,(
  ( ! [X17] : (p(s(bool,X17))) ) | $spl12),
  inference(cnf_transformation,[],[f229_D])).
fof(f229_D,plain,(
  ( ! [X17] : (p(s(bool,X17))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f2806,plain,(
  ( ! [X30,X29] : (p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X30) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK4(sK1,sK0))))) ) | $spl74),
  inference(resolution,[],[f2777,f190])).
fof(f190,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f128])).
fof(f128,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f78])).
fof(f78,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f63])).
fof(f63,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f62])).
fof(f62,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f31])).
fof(f31,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X8,X17,X11,X18] : ((s(bool,X8) = s(bool,X17) & (p(s(bool,X17)) => s(bool,X11) = s(bool,X18))) => ((p(s(bool,X8)) => p(s(bool,X11))) <=> (p(s(bool,X17)) => p(s(bool,X18)))))),
  file('thm.extreal.inf_const_alt','thm.bool.IMP_CONG')).
fof(f227,plain,(
  ( ! [X12,X13] : (s(bool,X12) != s(bool,X13) | ~p(s(bool,X12))) ) | $spl8),
  inference(cnf_transformation,[],[f227_D])).
fof(f227_D,plain,(
  ( ! [X12,X13] : (s(bool,X12) != s(bool,X13) | ~p(s(bool,X12))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f2796,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK4(sK1,sK0)))) | s(bool,X27) != s(bool,X28)) ) | $spl74),
  inference(resolution,[],[f2777,f131])).
fof(f131,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f79])).
fof(f228,plain,(
  ( ! [X18] : (~p(s(bool,X18))) ) | $spl10),
  inference(cnf_transformation,[],[f228_D])).
fof(f228_D,plain,(
  ( ! [X18] : (~p(s(bool,X18))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f2793,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X19) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK4(sK1,sK0))))) ) | $spl74),
  inference(resolution,[],[f2777,f130])).
fof(f130,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f79])).
fof(f224,plain,(
  ( ! [X7,X5] : (s(bool,X5) != s(bool,X7) | p(s(bool,X5))) ) | $spl2),
  inference(cnf_transformation,[],[f224_D])).
fof(f224_D,plain,(
  ( ! [X7,X5] : (s(bool,X5) != s(bool,X7) | p(s(bool,X5))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f2792,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK4(sK1,sK0))))) ) | $spl74),
  inference(resolution,[],[f2777,f129])).
fof(f129,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f79])).
fof(f2791,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK4(sK1,sK0)))) | s(bool,X12) != s(bool,X13)) ) | $spl74),
  inference(resolution,[],[f2777,f127])).
fof(f127,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f79])).
fof(f2789,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK4(sK1,sK0))))) ) | $spl74),
  inference(resolution,[],[f2777,f125])).
fof(f125,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f79])).
fof(f2633,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK1,sK0)),s('type.extreal.extreal',sK1)))) ) | $spl94),
  inference(cnf_transformation,[],[f2633_D])).
fof(f2633_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK1,sK0)),s('type.extreal.extreal',sK1)))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f2770,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | p(s(bool,X25)) | s(bool,X25) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK1,sK0)),s('type.extreal.extreal',sK1)))) ) | $spl74),
  inference(resolution,[],[f2027,f190])).
fof(f2760,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK1,sK0)),s('type.extreal.extreal',sK1))) | s(bool,X22) != s(bool,X23)) ) | $spl74),
  inference(resolution,[],[f2027,f131])).
fof(f2757,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK1,sK0)),s('type.extreal.extreal',sK1)))) ) | $spl74),
  inference(resolution,[],[f2027,f130])).
fof(f2756,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK1,sK0)),s('type.extreal.extreal',sK1)))) ) | $spl74),
  inference(resolution,[],[f2027,f129])).
fof(f2755,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK1,sK0)),s('type.extreal.extreal',sK1))) | s(bool,X7) != s(bool,X8)) ) | $spl74),
  inference(resolution,[],[f2027,f127])).
fof(f2753,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK1,sK0)),s('type.extreal.extreal',sK1)))) ) | $spl74),
  inference(resolution,[],[f2027,f125])).
fof(f2372,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | s(bool,X4) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK1,sK0))))) ) | $spl86),
  inference(cnf_transformation,[],[f2372_D])).
fof(f2372_D,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | s(bool,X4) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK1,sK0))))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f2725,plain,(
  ( ! [X33,X34] : (p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X34) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK1,sK0))))) ) | $spl73),
  inference(resolution,[],[f2026,f190])).
fof(f2712,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK1,sK0)))) | s(bool,X24) != s(bool,X26)) ) | $spl73),
  inference(resolution,[],[f2026,f131])).
fof(f2710,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK1,sK0))))) ) | $spl73),
  inference(resolution,[],[f2026,f130])).
fof(f2708,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK1,sK0))))) ) | $spl73),
  inference(resolution,[],[f2026,f129])).
fof(f2707,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK1,sK0)))) | s(bool,X9) != s(bool,X11)) ) | $spl73),
  inference(resolution,[],[f2026,f127])).
fof(f2705,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK1,sK0))))) ) | $spl73),
  inference(resolution,[],[f2026,f125])).
fof(f2699,plain,(
  $false | $spl75),
  inference(subsumption_resolution,[],[f2698,f189])).
fof(f2698,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK1)))) | $spl75),
  inference(forward_demodulation,[],[f2331,f2566])).
fof(f2566,plain,(
  s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK4(sK1,sK0)) | $spl75),
  inference(resolution,[],[f2490,f115])).
fof(f2490,plain,(
  p(s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK4(sK1,sK0))))) | $spl75),
  inference(subsumption_resolution,[],[f2489,f154])).
fof(f154,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f92])).
fof(f92,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))))),
  inference(flattening,[],[f91])).
fof(f91,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f67])).
fof(f67,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f44])).
fof(f44,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X10] : ((p(s(bool,X10)) => p(s(bool,X10))) <=> p(s(bool,'T_0')))),
  file('thm.extreal.inf_const_alt','thm.bool.IMP_CLAUSES|split|3')).
fof(f2489,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK4(sK1,sK0))))) | $spl75),
  inference(forward_demodulation,[],[f2488,f501])).
fof(f501,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK1)))),
  inference(resolution,[],[f447,f139])).
fof(f139,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X10] : (s(bool,'T_0') = s(bool,X10) <=> p(s(bool,X10)))),
  file('thm.extreal.inf_const_alt','thm.bool.EQ_CLAUSES|split|1')).
fof(f447,plain,(
  p(s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK1))))),
  inference(backward_demodulation,[],[f390,f114])).
fof(f114,plain,(
  p(s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK2))))),
  inference(cnf_transformation,[],[f73])).
fof(f390,plain,(
  s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK2)),
  inference(resolution,[],[f115,f114])).
fof(f2488,plain,(
  ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK1)))) | p(s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK4(sK1,sK0))))) | $spl75),
  inference(subsumption_resolution,[],[f2452,f116])).
fof(f2452,plain,(
  s('type.extreal.extreal',sK1) = s('type.extreal.extreal','const.extreal.extreal_inf_1'(s(fun('type.extreal.extreal',bool),sK0))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK1)))) | p(s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK4(sK1,sK0))))) | $spl75),
  inference(resolution,[],[f2028,f120])).
fof(f120,plain,(
  ( ! [X0,X3,X1] : (s('type.extreal.extreal',X1) = s('type.extreal.extreal','const.extreal.extreal_inf_1'(s(fun('type.extreal.extreal',bool),X0))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(X1,X0)),s('type.extreal.extreal',X3)))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal',X3)))) | p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal',sK4(X1,X0)))))) )),
  inference(cnf_transformation,[],[f77])).
fof(f2331,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK1,sK0))))) | $spl75),
  inference(subsumption_resolution,[],[f2330,f158])).
fof(f158,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f94])).
fof(f94,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f93])).
fof(f93,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f68])).
fof(f68,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f45])).
fof(f45,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X10] : ((p(s(bool,'F_0')) => p(s(bool,X10))) <=> p(s(bool,'T_0')))),
  file('thm.extreal.inf_const_alt','thm.bool.IMP_CLAUSES|split|2')).
fof(f2330,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK1,sK0))))) | $spl75),
  inference(forward_demodulation,[],[f2329,f501])).
fof(f2329,plain,(
  ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK1)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK1,sK0))))) | $spl75),
  inference(subsumption_resolution,[],[f2291,f116])).
fof(f2291,plain,(
  s('type.extreal.extreal',sK1) = s('type.extreal.extreal','const.extreal.extreal_inf_1'(s(fun('type.extreal.extreal',bool),sK0))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK1)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK1,sK0))))) | $spl75),
  inference(resolution,[],[f2028,f121])).
fof(f121,plain,(
  ( ! [X0,X3,X1] : (s('type.extreal.extreal',X1) = s('type.extreal.extreal','const.extreal.extreal_inf_1'(s(fun('type.extreal.extreal',bool),X0))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(X1,X0)),s('type.extreal.extreal',X3)))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal',X3)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',sK4(X1,X0)))))) )),
  inference(cnf_transformation,[],[f77])).
fof(f2669,plain,(
  $false | ($spl73 | $spl75)),
  inference(subsumption_resolution,[],[f2639,f189])).
fof(f2639,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK1)))) | ($spl73 | $spl75)),
  inference(backward_demodulation,[],[f2566,f2026])).
fof(f2625,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | p(s(bool,X25)) | s(bool,X25) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK1,sK0)),s('type.extreal.extreal',sK1)))) ) | $spl74),
  inference(resolution,[],[f2027,f190])).
fof(f2615,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK1,sK0)),s('type.extreal.extreal',sK1))) | s(bool,X22) != s(bool,X23)) ) | $spl74),
  inference(resolution,[],[f2027,f131])).
fof(f2612,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK1,sK0)),s('type.extreal.extreal',sK1)))) ) | $spl74),
  inference(resolution,[],[f2027,f130])).
fof(f2611,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK1,sK0)),s('type.extreal.extreal',sK1)))) ) | $spl74),
  inference(resolution,[],[f2027,f129])).
fof(f2610,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK1,sK0)),s('type.extreal.extreal',sK1))) | s(bool,X7) != s(bool,X8)) ) | $spl74),
  inference(resolution,[],[f2027,f127])).
fof(f2608,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK1,sK0)),s('type.extreal.extreal',sK1)))) ) | $spl74),
  inference(resolution,[],[f2027,f125])).
fof(f2589,plain,(
  ( ! [X30,X29] : (p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X30) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK4(sK1,sK0))))) ) | $spl75),
  inference(resolution,[],[f2490,f190])).
fof(f2579,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK4(sK1,sK0)))) | s(bool,X27) != s(bool,X28)) ) | $spl75),
  inference(resolution,[],[f2490,f131])).
fof(f2576,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X19) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK4(sK1,sK0))))) ) | $spl75),
  inference(resolution,[],[f2490,f130])).
fof(f2575,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK4(sK1,sK0))))) ) | $spl75),
  inference(resolution,[],[f2490,f129])).
fof(f2574,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK4(sK1,sK0)))) | s(bool,X12) != s(bool,X13)) ) | $spl75),
  inference(resolution,[],[f2490,f127])).
fof(f2572,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK4(sK1,sK0))))) ) | $spl75),
  inference(resolution,[],[f2490,f125])).
fof(f2104,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK2,sK0)),s('type.extreal.extreal',sK1)))) ) | $spl78),
  inference(cnf_transformation,[],[f2104_D])).
fof(f2104_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK2,sK0)),s('type.extreal.extreal',sK1)))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f2554,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | p(s(bool,X25)) | s(bool,X25) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK2,sK0)),s('type.extreal.extreal',sK1)))) ) | $spl60),
  inference(resolution,[],[f2009,f190])).
fof(f2009,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK2,sK0)),s('type.extreal.extreal',sK1)))) | $spl60),
  inference(cnf_transformation,[],[f2009_D])).
fof(f2009_D,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK2,sK0)),s('type.extreal.extreal',sK1)))) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f2544,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK2,sK0)),s('type.extreal.extreal',sK1))) | s(bool,X22) != s(bool,X23)) ) | $spl60),
  inference(resolution,[],[f2009,f131])).
fof(f2541,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK2,sK0)),s('type.extreal.extreal',sK1)))) ) | $spl60),
  inference(resolution,[],[f2009,f130])).
fof(f2540,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK2,sK0)),s('type.extreal.extreal',sK1)))) ) | $spl60),
  inference(resolution,[],[f2009,f129])).
fof(f2539,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK2,sK0)),s('type.extreal.extreal',sK1))) | s(bool,X7) != s(bool,X8)) ) | $spl60),
  inference(resolution,[],[f2009,f127])).
fof(f2537,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK2,sK0)),s('type.extreal.extreal',sK1)))) ) | $spl60),
  inference(resolution,[],[f2009,f125])).
fof(f2529,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,X6) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK4(sK2,sK0))))) ) | $spl90),
  inference(cnf_transformation,[],[f2529_D])).
fof(f2529_D,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,X6) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK4(sK2,sK0))))) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f2519,plain,(
  ( ! [X30,X29] : (p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X30) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK4(sK2,sK0))))) ) | $spl61),
  inference(resolution,[],[f2149,f190])).
fof(f2149,plain,(
  p(s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK4(sK2,sK0))))) | $spl61),
  inference(subsumption_resolution,[],[f2148,f154])).
fof(f2148,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK4(sK2,sK0))))) | $spl61),
  inference(forward_demodulation,[],[f2147,f501])).
fof(f2147,plain,(
  ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK1)))) | p(s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK4(sK2,sK0))))) | $spl61),
  inference(subsumption_resolution,[],[f2146,f116])).
fof(f2146,plain,(
  s('type.extreal.extreal',sK1) = s('type.extreal.extreal','const.extreal.extreal_inf_1'(s(fun('type.extreal.extreal',bool),sK0))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK1)))) | p(s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK4(sK2,sK0))))) | $spl61),
  inference(forward_demodulation,[],[f2110,f390])).
fof(f2110,plain,(
  s('type.extreal.extreal',sK2) = s('type.extreal.extreal','const.extreal.extreal_inf_1'(s(fun('type.extreal.extreal',bool),sK0))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK1)))) | p(s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK4(sK2,sK0))))) | $spl61),
  inference(resolution,[],[f2010,f120])).
fof(f2509,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK4(sK2,sK0)))) | s(bool,X27) != s(bool,X28)) ) | $spl61),
  inference(resolution,[],[f2149,f131])).
fof(f2506,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X19) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK4(sK2,sK0))))) ) | $spl61),
  inference(resolution,[],[f2149,f130])).
fof(f2505,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK4(sK2,sK0))))) ) | $spl61),
  inference(resolution,[],[f2149,f129])).
fof(f2504,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK4(sK2,sK0)))) | s(bool,X12) != s(bool,X13)) ) | $spl61),
  inference(resolution,[],[f2149,f127])).
fof(f2502,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK4(sK2,sK0))))) ) | $spl61),
  inference(resolution,[],[f2149,f125])).
fof(f2332,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | s(bool,X4) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK1,sK0)),s('type.extreal.extreal',sK1)))) ) | $spl84),
  inference(cnf_transformation,[],[f2332_D])).
fof(f2332_D,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | s(bool,X4) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK1,sK0)),s('type.extreal.extreal',sK1)))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f2477,plain,(
  ( ! [X33,X34] : (p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X34) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK1,sK0)),s('type.extreal.extreal',sK1)))) ) | $spl75),
  inference(resolution,[],[f2028,f190])).
fof(f2464,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK1,sK0)),s('type.extreal.extreal',sK1))) | s(bool,X24) != s(bool,X26)) ) | $spl75),
  inference(resolution,[],[f2028,f131])).
fof(f2462,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK1,sK0)),s('type.extreal.extreal',sK1)))) ) | $spl75),
  inference(resolution,[],[f2028,f130])).
fof(f2460,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK1,sK0)),s('type.extreal.extreal',sK1)))) ) | $spl75),
  inference(resolution,[],[f2028,f129])).
fof(f2459,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK1,sK0)),s('type.extreal.extreal',sK1))) | s(bool,X9) != s(bool,X11)) ) | $spl75),
  inference(resolution,[],[f2028,f127])).
fof(f2457,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK1,sK0)),s('type.extreal.extreal',sK1)))) ) | $spl75),
  inference(resolution,[],[f2028,f125])).
fof(f2435,plain,(
  ( ! [X33,X34] : (p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X34) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK1,sK0))))) ) | $spl73),
  inference(resolution,[],[f2026,f190])).
fof(f2422,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK1,sK0)))) | s(bool,X24) != s(bool,X26)) ) | $spl73),
  inference(resolution,[],[f2026,f131])).
fof(f2420,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK1,sK0))))) ) | $spl73),
  inference(resolution,[],[f2026,f130])).
fof(f2418,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK1,sK0))))) ) | $spl73),
  inference(resolution,[],[f2026,f129])).
fof(f2417,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK1,sK0)))) | s(bool,X9) != s(bool,X11)) ) | $spl73),
  inference(resolution,[],[f2026,f127])).
fof(f2415,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK1,sK0))))) ) | $spl73),
  inference(resolution,[],[f2026,f125])).
fof(f2406,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK1,sK0))))) ) | $spl88),
  inference(cnf_transformation,[],[f2406_D])).
fof(f2406_D,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK1,sK0))))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f2397,plain,(
  ( ! [X26,X25] : (p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X26) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK1,sK0))))) ) | $spl72),
  inference(resolution,[],[f2025,f190])).
fof(f2025,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK1,sK0))))) | $spl72),
  inference(cnf_transformation,[],[f2025_D])).
fof(f2025_D,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK1,sK0))))) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f2387,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK1,sK0)))) | s(bool,X23) != s(bool,X24)) ) | $spl72),
  inference(resolution,[],[f2025,f131])).
fof(f2384,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK1,sK0))))) ) | $spl72),
  inference(resolution,[],[f2025,f130])).
fof(f2383,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK1,sK0))))) ) | $spl72),
  inference(resolution,[],[f2025,f129])).
fof(f2382,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK1,sK0)))) | s(bool,X8) != s(bool,X9)) ) | $spl72),
  inference(resolution,[],[f2025,f127])).
fof(f2380,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK1,sK0))))) ) | $spl72),
  inference(resolution,[],[f2025,f125])).
fof(f2361,plain,(
  ( ! [X33,X34] : (p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X34) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK1,sK0))))) ) | $spl73),
  inference(resolution,[],[f2026,f190])).
fof(f2348,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK1,sK0)))) | s(bool,X24) != s(bool,X26)) ) | $spl73),
  inference(resolution,[],[f2026,f131])).
fof(f2346,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK1,sK0))))) ) | $spl73),
  inference(resolution,[],[f2026,f130])).
fof(f2344,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK1,sK0))))) ) | $spl73),
  inference(resolution,[],[f2026,f129])).
fof(f2343,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK1,sK0)))) | s(bool,X9) != s(bool,X11)) ) | $spl73),
  inference(resolution,[],[f2026,f127])).
fof(f2341,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK1,sK0))))) ) | $spl73),
  inference(resolution,[],[f2026,f125])).
fof(f2315,plain,(
  ( ! [X33,X34] : (p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X34) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK1,sK0)),s('type.extreal.extreal',sK1)))) ) | $spl75),
  inference(resolution,[],[f2028,f190])).
fof(f2302,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK1,sK0)),s('type.extreal.extreal',sK1))) | s(bool,X24) != s(bool,X26)) ) | $spl75),
  inference(resolution,[],[f2028,f131])).
fof(f2300,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK1,sK0)),s('type.extreal.extreal',sK1)))) ) | $spl75),
  inference(resolution,[],[f2028,f130])).
fof(f2298,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK1,sK0)),s('type.extreal.extreal',sK1)))) ) | $spl75),
  inference(resolution,[],[f2028,f129])).
fof(f2297,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK1,sK0)),s('type.extreal.extreal',sK1))) | s(bool,X9) != s(bool,X11)) ) | $spl75),
  inference(resolution,[],[f2028,f127])).
fof(f2295,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK1,sK0)),s('type.extreal.extreal',sK1)))) ) | $spl75),
  inference(resolution,[],[f2028,f125])).
fof(f2194,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | s(bool,X4) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK2,sK0))))) ) | $spl80),
  inference(cnf_transformation,[],[f2194_D])).
fof(f2194_D,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | s(bool,X4) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK2,sK0))))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f2274,plain,(
  ( ! [X33,X34] : (p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X34) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK2,sK0))))) ) | $spl63),
  inference(resolution,[],[f2012,f190])).
fof(f2261,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK2,sK0)))) | s(bool,X24) != s(bool,X26)) ) | $spl63),
  inference(resolution,[],[f2012,f131])).
fof(f2259,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK2,sK0))))) ) | $spl63),
  inference(resolution,[],[f2012,f130])).
fof(f2257,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK2,sK0))))) ) | $spl63),
  inference(resolution,[],[f2012,f129])).
fof(f2256,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK2,sK0)))) | s(bool,X9) != s(bool,X11)) ) | $spl63),
  inference(resolution,[],[f2012,f127])).
fof(f2254,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK2,sK0))))) ) | $spl63),
  inference(resolution,[],[f2012,f125])).
fof(f2224,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK2,sK0))))) ) | $spl82),
  inference(cnf_transformation,[],[f2224_D])).
fof(f2224_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK2,sK0))))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f2217,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | p(s(bool,X25)) | s(bool,X25) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK2,sK0))))) ) | $spl62),
  inference(resolution,[],[f2011,f190])).
fof(f2011,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK2,sK0))))) | $spl62),
  inference(cnf_transformation,[],[f2011_D])).
fof(f2011_D,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK2,sK0))))) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f2207,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK2,sK0)))) | s(bool,X22) != s(bool,X23)) ) | $spl62),
  inference(resolution,[],[f2011,f131])).
fof(f2204,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK2,sK0))))) ) | $spl62),
  inference(resolution,[],[f2011,f130])).
fof(f2203,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK2,sK0))))) ) | $spl62),
  inference(resolution,[],[f2011,f129])).
fof(f2202,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK2,sK0)))) | s(bool,X7) != s(bool,X8)) ) | $spl62),
  inference(resolution,[],[f2011,f127])).
fof(f2200,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK2,sK0))))) ) | $spl62),
  inference(resolution,[],[f2011,f125])).
fof(f2183,plain,(
  ( ! [X33,X34] : (p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X34) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK2,sK0))))) ) | $spl63),
  inference(resolution,[],[f2012,f190])).
fof(f2170,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK2,sK0)))) | s(bool,X24) != s(bool,X26)) ) | $spl63),
  inference(resolution,[],[f2012,f131])).
fof(f2168,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK2,sK0))))) ) | $spl63),
  inference(resolution,[],[f2012,f130])).
fof(f2166,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK2,sK0))))) ) | $spl63),
  inference(resolution,[],[f2012,f129])).
fof(f2165,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK2,sK0)))) | s(bool,X9) != s(bool,X11)) ) | $spl63),
  inference(resolution,[],[f2012,f127])).
fof(f2163,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK2,sK0))))) ) | $spl63),
  inference(resolution,[],[f2012,f125])).
fof(f2074,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | s(bool,X4) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK2,sK0)),s('type.extreal.extreal',sK1)))) ) | $spl76),
  inference(cnf_transformation,[],[f2074_D])).
fof(f2074_D,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | s(bool,X4) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK2,sK0)),s('type.extreal.extreal',sK1)))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f2135,plain,(
  ( ! [X33,X34] : (p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X34) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK2,sK0)),s('type.extreal.extreal',sK1)))) ) | $spl61),
  inference(resolution,[],[f2010,f190])).
fof(f2122,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK2,sK0)),s('type.extreal.extreal',sK1))) | s(bool,X24) != s(bool,X26)) ) | $spl61),
  inference(resolution,[],[f2010,f131])).
fof(f2120,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK2,sK0)),s('type.extreal.extreal',sK1)))) ) | $spl61),
  inference(resolution,[],[f2010,f130])).
fof(f2118,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK2,sK0)),s('type.extreal.extreal',sK1)))) ) | $spl61),
  inference(resolution,[],[f2010,f129])).
fof(f2117,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK2,sK0)),s('type.extreal.extreal',sK1))) | s(bool,X9) != s(bool,X11)) ) | $spl61),
  inference(resolution,[],[f2010,f127])).
fof(f2115,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK2,sK0)),s('type.extreal.extreal',sK1)))) ) | $spl61),
  inference(resolution,[],[f2010,f125])).
fof(f2097,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | p(s(bool,X25)) | s(bool,X25) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK2,sK0)),s('type.extreal.extreal',sK1)))) ) | $spl60),
  inference(resolution,[],[f2009,f190])).
fof(f2087,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK2,sK0)),s('type.extreal.extreal',sK1))) | s(bool,X22) != s(bool,X23)) ) | $spl60),
  inference(resolution,[],[f2009,f131])).
fof(f2084,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK2,sK0)),s('type.extreal.extreal',sK1)))) ) | $spl60),
  inference(resolution,[],[f2009,f130])).
fof(f2083,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK2,sK0)),s('type.extreal.extreal',sK1)))) ) | $spl60),
  inference(resolution,[],[f2009,f129])).
fof(f2082,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK2,sK0)),s('type.extreal.extreal',sK1))) | s(bool,X7) != s(bool,X8)) ) | $spl60),
  inference(resolution,[],[f2009,f127])).
fof(f2080,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK2,sK0)),s('type.extreal.extreal',sK1)))) ) | $spl60),
  inference(resolution,[],[f2009,f125])).
fof(f2054,plain,(
  ( ! [X33,X34] : (p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X34) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK2,sK0)),s('type.extreal.extreal',sK1)))) ) | $spl61),
  inference(resolution,[],[f2010,f190])).
fof(f2041,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK2,sK0)),s('type.extreal.extreal',sK1))) | s(bool,X24) != s(bool,X26)) ) | $spl61),
  inference(resolution,[],[f2010,f131])).
fof(f2039,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK2,sK0)),s('type.extreal.extreal',sK1)))) ) | $spl61),
  inference(resolution,[],[f2010,f130])).
fof(f2037,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK2,sK0)),s('type.extreal.extreal',sK1)))) ) | $spl61),
  inference(resolution,[],[f2010,f129])).
fof(f2036,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK2,sK0)),s('type.extreal.extreal',sK1))) | s(bool,X9) != s(bool,X11)) ) | $spl61),
  inference(resolution,[],[f2010,f127])).
fof(f2034,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK2,sK0)),s('type.extreal.extreal',sK1)))) ) | $spl61),
  inference(resolution,[],[f2010,f125])).
fof(f2026,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK1,sK0))))) | $spl73),
  inference(cnf_transformation,[],[f2026_D])).
fof(f2026_D,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK1,sK0))))) <=> ~$spl73),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl73])])).
fof(f2028,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK1,sK0)),s('type.extreal.extreal',sK1)))) | $spl75),
  inference(cnf_transformation,[],[f2028_D])).
fof(f2028_D,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK1,sK0)),s('type.extreal.extreal',sK1)))) <=> ~$spl75),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl75])])).
fof(f2001,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK1,sK0)),s('type.extreal.extreal',sK1)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK1,sK0)))))),
  inference(equality_resolution,[],[f240])).
fof(f240,plain,(
  ( ! [X5] : (s('type.extreal.extreal',sK1) != s('type.extreal.extreal',X5) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(X5,sK0)),s('type.extreal.extreal',X5)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X5),s('type.extreal.extreal',sK4(X5,sK0)))))) )),
  inference(superposition,[],[f116,f123])).
fof(f2022,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK7('F_0',sK0,'type.extreal.extreal'),sK0)),s('type.extreal.extreal',sK1)))) | $spl69),
  inference(cnf_transformation,[],[f2022_D])).
fof(f2022_D,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK7('F_0',sK0,'type.extreal.extreal'),sK0)),s('type.extreal.extreal',sK1)))) <=> ~$spl69),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl69])])).
fof(f2024,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK7('F_0',sK0,'type.extreal.extreal'),sK0))))) | $spl71),
  inference(cnf_transformation,[],[f2024_D])).
fof(f2024_D,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK7('F_0',sK0,'type.extreal.extreal'),sK0))))) <=> ~$spl71),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl71])])).
fof(f2020,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK7('F_0',sK0,'type.extreal.extreal'),sK0))))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK7('F_0',sK0,'type.extreal.extreal'),sK0)),s('type.extreal.extreal',sK1)))) | $spl20),
  inference(forward_demodulation,[],[f2019,f594])).
fof(f594,plain,(
  s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK7('F_0',sK0,'type.extreal.extreal')) | $spl20),
  inference(resolution,[],[f435,f181])).
fof(f181,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f53])).
fof(f53,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.extreal.inf_const_alt',aHL_FALSITY)).
fof(f2019,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK7('F_0',sK0,'type.extreal.extreal'),sK0)),s('type.extreal.extreal',sK1)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK7('F_0',sK0,'type.extreal.extreal')),s('type.extreal.extreal',sK4(sK7('F_0',sK0,'type.extreal.extreal'),sK0))))) | $spl20),
  inference(forward_demodulation,[],[f2003,f594])).
fof(f2003,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK7('F_0',sK0,'type.extreal.extreal'),sK0)),s('type.extreal.extreal',sK7('F_0',sK0,'type.extreal.extreal'))))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK7('F_0',sK0,'type.extreal.extreal')),s('type.extreal.extreal',sK4(sK7('F_0',sK0,'type.extreal.extreal'),sK0))))) | $spl20),
  inference(trivial_inequality_removal,[],[f1999])).
fof(f1999,plain,(
  s('type.extreal.extreal',sK1) != s('type.extreal.extreal',sK1) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK7('F_0',sK0,'type.extreal.extreal'),sK0)),s('type.extreal.extreal',sK7('F_0',sK0,'type.extreal.extreal'))))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK7('F_0',sK0,'type.extreal.extreal')),s('type.extreal.extreal',sK4(sK7('F_0',sK0,'type.extreal.extreal'),sK0))))) | $spl20),
  inference(superposition,[],[f240,f594])).
fof(f2016,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK6('F_0',sK0,'type.extreal.extreal'),sK0)),s('type.extreal.extreal',sK1)))) | $spl65),
  inference(cnf_transformation,[],[f2016_D])).
fof(f2016_D,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK6('F_0',sK0,'type.extreal.extreal'),sK0)),s('type.extreal.extreal',sK1)))) <=> ~$spl65),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl65])])).
fof(f2018,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK6('F_0',sK0,'type.extreal.extreal'),sK0))))) | $spl67),
  inference(cnf_transformation,[],[f2018_D])).
fof(f2018_D,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK6('F_0',sK0,'type.extreal.extreal'),sK0))))) <=> ~$spl67),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl67])])).
fof(f2014,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK6('F_0',sK0,'type.extreal.extreal'),sK0))))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK6('F_0',sK0,'type.extreal.extreal'),sK0)),s('type.extreal.extreal',sK1)))) | $spl22),
  inference(forward_demodulation,[],[f2013,f691])).
fof(f691,plain,(
  s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK6('F_0',sK0,'type.extreal.extreal')) | $spl22),
  inference(resolution,[],[f436,f181])).
fof(f2013,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK6('F_0',sK0,'type.extreal.extreal'),sK0)),s('type.extreal.extreal',sK1)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK6('F_0',sK0,'type.extreal.extreal')),s('type.extreal.extreal',sK4(sK6('F_0',sK0,'type.extreal.extreal'),sK0))))) | $spl22),
  inference(forward_demodulation,[],[f2005,f691])).
fof(f2005,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK6('F_0',sK0,'type.extreal.extreal'),sK0)),s('type.extreal.extreal',sK6('F_0',sK0,'type.extreal.extreal'))))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK6('F_0',sK0,'type.extreal.extreal')),s('type.extreal.extreal',sK4(sK6('F_0',sK0,'type.extreal.extreal'),sK0))))) | $spl22),
  inference(trivial_inequality_removal,[],[f1997])).
fof(f1997,plain,(
  s('type.extreal.extreal',sK1) != s('type.extreal.extreal',sK1) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK6('F_0',sK0,'type.extreal.extreal'),sK0)),s('type.extreal.extreal',sK6('F_0',sK0,'type.extreal.extreal'))))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK6('F_0',sK0,'type.extreal.extreal')),s('type.extreal.extreal',sK4(sK6('F_0',sK0,'type.extreal.extreal'),sK0))))) | $spl22),
  inference(superposition,[],[f240,f691])).
fof(f2010,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK2,sK0)),s('type.extreal.extreal',sK1)))) | $spl61),
  inference(cnf_transformation,[],[f2010_D])).
fof(f2010_D,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK2,sK0)),s('type.extreal.extreal',sK1)))) <=> ~$spl61),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl61])])).
fof(f2012,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK2,sK0))))) | $spl63),
  inference(cnf_transformation,[],[f2012_D])).
fof(f2012_D,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK2,sK0))))) <=> ~$spl63),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl63])])).
fof(f2008,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK4(sK2,sK0))))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK2,sK0)),s('type.extreal.extreal',sK1))))),
  inference(forward_demodulation,[],[f2007,f390])).
fof(f2007,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK2,sK0)),s('type.extreal.extreal',sK1)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK2),s('type.extreal.extreal',sK4(sK2,sK0)))))),
  inference(forward_demodulation,[],[f2006,f390])).
fof(f2006,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK2,sK0)),s('type.extreal.extreal',sK2)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK2),s('type.extreal.extreal',sK4(sK2,sK0)))))),
  inference(trivial_inequality_removal,[],[f1996])).
fof(f1996,plain,(
  s('type.extreal.extreal',sK1) != s('type.extreal.extreal',sK1) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK5(sK2,sK0)),s('type.extreal.extreal',sK2)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK2),s('type.extreal.extreal',sK4(sK2,sK0)))))),
  inference(superposition,[],[f240,f390])).
fof(f1972,plain,(
  ( ! [X2,X0] : (s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK6(X0,sK0,'type.extreal.extreal')))) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) | $spl58),
  inference(cnf_transformation,[],[f1972_D])).
fof(f1972_D,plain,(
  ( ! [X2,X0] : (s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK6(X0,sK0,'type.extreal.extreal')))) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f1937,plain,(
  ( ! [X43,X41,X42] : (s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK6(X41,sK0,'type.extreal.extreal')))) | p(s(bool,X42)) | p(s(bool,X43)) | s(bool,X41) != s(bool,X43)) )),
  inference(resolution,[],[f1023,f190])).
fof(f1023,plain,(
  ( ! [X45] : (p(s(bool,X45)) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK6(X45,sK0,'type.extreal.extreal'))))) )),
  inference(resolution,[],[f195,f139])).
fof(f195,plain,(
  ( ! [X3] : (p(s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK6(X3,sK0,'type.extreal.extreal'))))) | p(s(bool,X3))) )),
  inference(resolution,[],[f114,f184])).
fof(f184,plain,(
  ( ! [X2,X0,X5,X1] : (p(s(bool,X2)) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK6(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f112])).
fof(f112,plain,(
  ! [X0,X1,X2] : (((p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK6(X2,X1,X0))))) & ~p(s(bool,X2))) | ! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | p(s(bool,X2))) & ((p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK7(X2,X1,X0))))) & ~p(s(bool,X2))) | ! [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK6,sK7])],[f111])).
fof(f111,plain,(
  ! [X0,X1,X2] : ((? [X6] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) & ~p(s(bool,X2))) | ! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | p(s(bool,X2))) & ((? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ~p(s(bool,X2))) | ! [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2)))))),
  inference(rectify,[],[f110])).
fof(f110,plain,(
  ! [X0,X1,X2] : ((? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & ~p(s(bool,X2))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,X2))) & ((? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ~p(s(bool,X2))) | ! [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2)))))),
  inference(flattening,[],[f109])).
fof(f109,plain,(
  ! [X0,X1,X2] : ((? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & ~p(s(bool,X2))) | (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,X2)))) & ((? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ~p(s(bool,X2))) | ! [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2)))))),
  inference(nnf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0,X1,X2] : (! [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2))) <=> (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,X2))))),
  inference(ennf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0,X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) => p(s(bool,X2))) <=> (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) => p(s(bool,X2))))),
  inference(flattening,[],[f54])).
fof(f54,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) => p(s(bool,X2))) <=> (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) => p(s(bool,X2))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X9] : ! [X12,X13] : (! [X8] : (p(s(bool,happ(s(fun(X9,bool),X12),s(X9,X8)))) => p(s(bool,X13))) <=> (? [X8] : p(s(bool,happ(s(fun(X9,bool),X12),s(X9,X8)))) => p(s(bool,X13))))),
  file('thm.extreal.inf_const_alt','thm.bool.LEFT_FORALL_IMP_THM')).
fof(f1927,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK6(X28,sK0,'type.extreal.extreal')))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) )),
  inference(resolution,[],[f1023,f131])).
fof(f1924,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK6(X16,sK0,'type.extreal.extreal')))) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) )),
  inference(resolution,[],[f1023,f130])).
fof(f1923,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK6(X12,sK0,'type.extreal.extreal')))) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) )),
  inference(resolution,[],[f1023,f129])).
fof(f1922,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK6(X8,sK0,'type.extreal.extreal')))) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) )),
  inference(resolution,[],[f1023,f127])).
fof(f1920,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK6(X0,sK0,'type.extreal.extreal')))) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) )),
  inference(resolution,[],[f1023,f125])).
fof(f1662,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_inf_1'(s(fun('type.extreal.extreal',bool),sK0))),s('type.extreal.extreal',sK1)))) ) | $spl56),
  inference(cnf_transformation,[],[f1662_D])).
fof(f1662_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_inf_1'(s(fun('type.extreal.extreal',bool),sK0))),s('type.extreal.extreal',sK1)))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f1645,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | p(s(bool,X25)) | s(bool,X25) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_inf_1'(s(fun('type.extreal.extreal',bool),sK0))),s('type.extreal.extreal',sK1)))) )),
  inference(resolution,[],[f1627,f190])).
fof(f1627,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_inf_1'(s(fun('type.extreal.extreal',bool),sK0))),s('type.extreal.extreal',sK1))))),
  inference(equality_resolution,[],[f486])).
fof(f486,plain,(
  ( ! [X0] : (s('type.extreal.extreal',X0) != s('type.extreal.extreal','const.extreal.extreal_inf_1'(s(fun('type.extreal.extreal',bool),sK0))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',sK1))))) )),
  inference(resolution,[],[f447,f117])).
fof(f117,plain,(
  ( ! [X0,X7,X1] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X7)))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal',X7)))) | s('type.extreal.extreal',X1) != s('type.extreal.extreal','const.extreal.extreal_inf_1'(s(fun('type.extreal.extreal',bool),X0)))) )),
  inference(cnf_transformation,[],[f77])).
fof(f1635,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_inf_1'(s(fun('type.extreal.extreal',bool),sK0))),s('type.extreal.extreal',sK1))) | s(bool,X22) != s(bool,X23)) )),
  inference(resolution,[],[f1627,f131])).
fof(f1632,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_inf_1'(s(fun('type.extreal.extreal',bool),sK0))),s('type.extreal.extreal',sK1)))) )),
  inference(resolution,[],[f1627,f130])).
fof(f1631,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_inf_1'(s(fun('type.extreal.extreal',bool),sK0))),s('type.extreal.extreal',sK1)))) )),
  inference(resolution,[],[f1627,f129])).
fof(f1630,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_inf_1'(s(fun('type.extreal.extreal',bool),sK0))),s('type.extreal.extreal',sK1))) | s(bool,X7) != s(bool,X8)) )),
  inference(resolution,[],[f1627,f127])).
fof(f1628,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_inf_1'(s(fun('type.extreal.extreal',bool),sK0))),s('type.extreal.extreal',sK1)))) )),
  inference(resolution,[],[f1627,f125])).
fof(f1540,plain,(
  ( ! [X222] : (s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK6(happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',X222)),sK0,'type.extreal.extreal')) | s('type.extreal.extreal',sK1) = s('type.extreal.extreal',X222)) ) | $spl54),
  inference(cnf_transformation,[],[f1540_D])).
fof(f1540_D,plain,(
  ( ! [X222] : (s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK6(happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',X222)),sK0,'type.extreal.extreal')) | s('type.extreal.extreal',sK1) = s('type.extreal.extreal',X222)) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f1539,plain,(
  ( ! [X222,X223] : (s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK6(happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',X222)),sK0,'type.extreal.extreal')) | p(s(bool,X223)) | s('type.extreal.extreal',sK1) = s('type.extreal.extreal',X222)) ) | $spl40),
  inference(subsumption_resolution,[],[f1482,f135])).
fof(f135,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f35])).
fof(f35,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f34])).
fof(f34,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X10] : (s(bool,'F_0') = s(bool,X10) <=> ~p(s(bool,X10)))),
  file('thm.extreal.inf_const_alt','thm.bool.EQ_CLAUSES|split|2')).
fof(f1482,plain,(
  ( ! [X222,X223] : (s(bool,'F_0') != s(bool,X223) | s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK6(happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',X222)),sK0,'type.extreal.extreal')) | p(s(bool,X223)) | s('type.extreal.extreal',sK1) = s('type.extreal.extreal',X222)) ) | $spl40),
  inference(superposition,[],[f721,f404])).
fof(f404,plain,(
  ( ! [X41] : (s(bool,'F_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',X41))) | s('type.extreal.extreal',sK1) = s('type.extreal.extreal',X41)) )),
  inference(resolution,[],[f115,f135])).
fof(f1538,plain,(
  ( ! [X217] : (s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK7(happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',X217)),sK0,'type.extreal.extreal')) | s('type.extreal.extreal',sK1) = s('type.extreal.extreal',X217)) ) | $spl52),
  inference(cnf_transformation,[],[f1538_D])).
fof(f1538_D,plain,(
  ( ! [X217] : (s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK7(happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',X217)),sK0,'type.extreal.extreal')) | s('type.extreal.extreal',sK1) = s('type.extreal.extreal',X217)) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f1537,plain,(
  ( ! [X218,X217] : (s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK7(happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',X217)),sK0,'type.extreal.extreal')) | p(s(bool,X218)) | s('type.extreal.extreal',sK1) = s('type.extreal.extreal',X217)) ) | $spl34),
  inference(subsumption_resolution,[],[f1479,f135])).
fof(f1479,plain,(
  ( ! [X218,X217] : (s(bool,'F_0') != s(bool,X218) | s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK7(happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',X217)),sK0,'type.extreal.extreal')) | p(s(bool,X218)) | s('type.extreal.extreal',sK1) = s('type.extreal.extreal',X217)) ) | $spl34),
  inference(superposition,[],[f624,f404])).
fof(f437,plain,(
  ( ! [X11,X9] : (s(bool,X11) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',X9))) | ~p(s(bool,X11)) | s('type.extreal.extreal',sK1) = s('type.extreal.extreal',X9)) ) | $spl24),
  inference(cnf_transformation,[],[f437_D])).
fof(f437_D,plain,(
  ( ! [X11,X9] : (s(bool,X11) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',X9))) | ~p(s(bool,X11)) | s('type.extreal.extreal',sK1) = s('type.extreal.extreal',X9)) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f1531,plain,(
  ( ! [X202,X203,X201] : (p(s(bool,X202)) | ~p(s(bool,X203)) | s(bool,X203) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',X201))) | s('type.extreal.extreal',sK1) = s('type.extreal.extreal',X201)) )),
  inference(subsumption_resolution,[],[f1471,f168])).
fof(f168,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f100])).
fof(f100,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | ~p(s(bool,X0)) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,'F_0')) & p(s(bool,X0)))))),
  inference(flattening,[],[f99])).
fof(f99,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) | ~p(s(bool,X0))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,'F_0')) & p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f48])).
fof(f48,plain,(
  ! [X0] : ((p(s(bool,'F_0')) & p(s(bool,X0))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X10] : ((p(s(bool,'F_0')) & p(s(bool,X10))) <=> p(s(bool,'F_0')))),
  file('thm.extreal.inf_const_alt','thm.bool.AND_CLAUSES|split|2')).
fof(f1471,plain,(
  ( ! [X202,X203,X201] : (p(s(bool,'F_0')) | p(s(bool,X202)) | ~p(s(bool,X203)) | s(bool,X203) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',X201))) | s('type.extreal.extreal',sK1) = s('type.extreal.extreal',X201)) )),
  inference(superposition,[],[f190,f404])).
fof(f1521,plain,(
  ( ! [X146,X149,X147,X148] : (~p(s(bool,X147)) | ~p(s(bool,X148)) | s(bool,X148) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',X146))) | s(bool,X147) != s(bool,X149) | s('type.extreal.extreal',sK1) = s('type.extreal.extreal',X146)) )),
  inference(subsumption_resolution,[],[f1438,f152])).
fof(f152,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f90])).
fof(f90,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f89])).
fof(f89,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f66])).
fof(f66,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f43])).
fof(f43,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f42])).
fof(f42,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X10] : ((p(s(bool,X10)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X10)))),
  file('thm.extreal.inf_const_alt','thm.bool.IMP_CLAUSES|split|4')).
fof(f1438,plain,(
  ( ! [X146,X149,X147,X148] : (p(s(bool,'F_0')) | ~p(s(bool,X147)) | ~p(s(bool,X148)) | s(bool,X148) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',X146))) | s(bool,X147) != s(bool,X149) | s('type.extreal.extreal',sK1) = s('type.extreal.extreal',X146)) )),
  inference(superposition,[],[f131,f404])).
fof(f1519,plain,(
  ( ! [X116,X114,X117,X115] : (p(s(bool,X115)) | ~p(s(bool,X116)) | ~p(s(bool,X117)) | s(bool,X116) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',X114))) | s('type.extreal.extreal',sK1) = s('type.extreal.extreal',X114)) )),
  inference(subsumption_resolution,[],[f1430,f168])).
fof(f1430,plain,(
  ( ! [X116,X114,X117,X115] : (p(s(bool,'F_0')) | p(s(bool,X115)) | ~p(s(bool,X116)) | ~p(s(bool,X117)) | s(bool,X116) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',X114))) | s('type.extreal.extreal',sK1) = s('type.extreal.extreal',X114)) )),
  inference(superposition,[],[f130,f404])).
fof(f1515,plain,(
  ( ! [X94,X97,X95,X96] : (p(s(bool,X95)) | ~p(s(bool,X96)) | s(bool,X95) != s(bool,X97) | s(bool,X96) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',X94))) | s('type.extreal.extreal',sK1) = s('type.extreal.extreal',X94)) )),
  inference(subsumption_resolution,[],[f1425,f168])).
fof(f1425,plain,(
  ( ! [X94,X97,X95,X96] : (p(s(bool,'F_0')) | p(s(bool,X95)) | ~p(s(bool,X96)) | s(bool,X95) != s(bool,X97) | s(bool,X96) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',X94))) | s('type.extreal.extreal',sK1) = s('type.extreal.extreal',X94)) )),
  inference(superposition,[],[f129,f404])).
fof(f1512,plain,(
  ( ! [X68,X66,X69,X67] : (~p(s(bool,X67)) | ~p(s(bool,X68)) | s(bool,X68) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',X66))) | s(bool,X67) != s(bool,X69) | s('type.extreal.extreal',sK1) = s('type.extreal.extreal',X66)) )),
  inference(subsumption_resolution,[],[f1418,f152])).
fof(f1418,plain,(
  ( ! [X68,X66,X69,X67] : (p(s(bool,'F_0')) | ~p(s(bool,X67)) | ~p(s(bool,X68)) | s(bool,X68) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',X66))) | s(bool,X67) != s(bool,X69) | s('type.extreal.extreal',sK1) = s('type.extreal.extreal',X66)) )),
  inference(superposition,[],[f127,f404])).
fof(f1508,plain,(
  ( ! [X39,X41,X38,X40] : (p(s(bool,X39)) | ~p(s(bool,X40)) | s(bool,X39) != s(bool,X41) | s(bool,X40) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',X38))) | s('type.extreal.extreal',sK1) = s('type.extreal.extreal',X38)) )),
  inference(subsumption_resolution,[],[f1411,f168])).
fof(f1411,plain,(
  ( ! [X39,X41,X38,X40] : (p(s(bool,'F_0')) | p(s(bool,X39)) | ~p(s(bool,X40)) | s(bool,X39) != s(bool,X41) | s(bool,X40) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',X38))) | s('type.extreal.extreal',sK1) = s('type.extreal.extreal',X38)) )),
  inference(superposition,[],[f125,f404])).
fof(f1148,plain,(
  ( ! [X13,X11] : (p(s(bool,X11)) | p(s(bool,X13)) | s(bool,X13) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK7(X11,sK0,'type.extreal.extreal'))))) ) | $spl50),
  inference(cnf_transformation,[],[f1148_D])).
fof(f1148_D,plain,(
  ( ! [X13,X11] : (p(s(bool,X11)) | p(s(bool,X13)) | s(bool,X13) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK7(X11,sK0,'type.extreal.extreal'))))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f1133,plain,(
  ( ! [X54,X52,X53] : (p(s(bool,X52)) | p(s(bool,X53)) | p(s(bool,X54)) | s(bool,X54) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK7(X52,sK0,'type.extreal.extreal'))))) )),
  inference(resolution,[],[f196,f190])).
fof(f196,plain,(
  ( ! [X4] : (p(s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK7(X4,sK0,'type.extreal.extreal'))))) | p(s(bool,X4))) )),
  inference(resolution,[],[f114,f186])).
fof(f186,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK7(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f112])).
fof(f1123,plain,(
  ( ! [X39,X41,X42,X40] : (p(s(bool,X39)) | p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X40) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK7(X39,sK0,'type.extreal.extreal')))) | s(bool,X41) != s(bool,X42)) )),
  inference(resolution,[],[f196,f131])).
fof(f1120,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,X27)) | p(s(bool,X28)) | ~p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X30) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK7(X27,sK0,'type.extreal.extreal'))))) )),
  inference(resolution,[],[f196,f130])).
fof(f1119,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X23)) | p(s(bool,X24)) | p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK7(X23,sK0,'type.extreal.extreal'))))) )),
  inference(resolution,[],[f196,f129])).
fof(f1118,plain,(
  ( ! [X21,X19,X22,X20] : (p(s(bool,X19)) | p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X20) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK7(X19,sK0,'type.extreal.extreal')))) | s(bool,X21) != s(bool,X22)) )),
  inference(resolution,[],[f196,f127])).
fof(f1116,plain,(
  ( ! [X14,X12,X13,X11] : (p(s(bool,X11)) | p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK7(X11,sK0,'type.extreal.extreal'))))) )),
  inference(resolution,[],[f196,f125])).
fof(f1045,plain,(
  ( ! [X13,X11] : (p(s(bool,X11)) | p(s(bool,X13)) | s(bool,X13) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK6(X11,sK0,'type.extreal.extreal'))))) ) | $spl48),
  inference(cnf_transformation,[],[f1045_D])).
fof(f1045_D,plain,(
  ( ! [X13,X11] : (p(s(bool,X11)) | p(s(bool,X13)) | s(bool,X13) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK6(X11,sK0,'type.extreal.extreal'))))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f1030,plain,(
  ( ! [X54,X52,X53] : (p(s(bool,X52)) | p(s(bool,X53)) | p(s(bool,X54)) | s(bool,X54) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK6(X52,sK0,'type.extreal.extreal'))))) )),
  inference(resolution,[],[f195,f190])).
fof(f1020,plain,(
  ( ! [X39,X41,X42,X40] : (p(s(bool,X39)) | p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X40) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK6(X39,sK0,'type.extreal.extreal')))) | s(bool,X41) != s(bool,X42)) )),
  inference(resolution,[],[f195,f131])).
fof(f1017,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,X27)) | p(s(bool,X28)) | ~p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X30) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK6(X27,sK0,'type.extreal.extreal'))))) )),
  inference(resolution,[],[f195,f130])).
fof(f1016,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X23)) | p(s(bool,X24)) | p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK6(X23,sK0,'type.extreal.extreal'))))) )),
  inference(resolution,[],[f195,f129])).
fof(f1015,plain,(
  ( ! [X21,X19,X22,X20] : (p(s(bool,X19)) | p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X20) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK6(X19,sK0,'type.extreal.extreal')))) | s(bool,X21) != s(bool,X22)) )),
  inference(resolution,[],[f195,f127])).
fof(f1013,plain,(
  ( ! [X14,X12,X13,X11] : (p(s(bool,X11)) | p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK6(X11,sK0,'type.extreal.extreal'))))) )),
  inference(resolution,[],[f195,f125])).
fof(f1044,plain,(
  ( ! [X10] : (p(s(bool,X10)) | p(s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK7(X10,sK0,'type.extreal.extreal')))))) ) | $spl46),
  inference(cnf_transformation,[],[f1044_D])).
fof(f1044_D,plain,(
  ( ! [X10] : (p(s(bool,X10)) | p(s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK7(X10,sK0,'type.extreal.extreal')))))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f1012,plain,(
  ( ! [X10,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | p(s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK7(X10,sK0,'type.extreal.extreal')))))) )),
  inference(resolution,[],[f195,f186])).
fof(f998,plain,(
  $false | $spl4),
  inference(trivial_inequality_removal,[],[f987])).
fof(f987,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl4),
  inference(superposition,[],[f772,f501])).
fof(f772,plain,(
  ( ! [X10] : (s(bool,X10) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK1)))) ) | $spl4),
  inference(forward_demodulation,[],[f225,f390])).
fof(f997,plain,(
  $false | $spl4),
  inference(equality_resolution,[],[f772])).
fof(f725,plain,(
  s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK6('F_0',sK0,'type.extreal.extreal')) | $spl44),
  inference(cnf_transformation,[],[f725_D])).
fof(f725_D,plain,(
  s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK6('F_0',sK0,'type.extreal.extreal')) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f719,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK6('F_0',sK0,'type.extreal.extreal'))) ) | $spl22),
  inference(duplicate_literal_removal,[],[f709])).
fof(f709,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK6('F_0',sK0,'type.extreal.extreal')) | p(s(bool,X2))) ) | $spl22),
  inference(superposition,[],[f436,f135])).
fof(f720,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK6('F_0',sK0,'type.extreal.extreal'))) ) | $spl22),
  inference(duplicate_literal_removal,[],[f708])).
fof(f708,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK6('F_0',sK0,'type.extreal.extreal')) | p(s(bool,X1))) ) | $spl22),
  inference(superposition,[],[f436,f137])).
fof(f137,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f37])).
fof(f37,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f36])).
fof(f36,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X10] : (s(bool,'F_0') = s(bool,X10) <=> ~p(s(bool,X10)))),
  file('thm.extreal.inf_const_alt','thm.bool.EQ_CLAUSES|split|3')).
fof(f727,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK6('F_0',sK0,'type.extreal.extreal'))) ) | $spl22),
  inference(subsumption_resolution,[],[f707,f138])).
fof(f138,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f83])).
fof(f707,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK6('F_0',sK0,'type.extreal.extreal')) | s(bool,'T_0') = s(bool,X0)) ) | $spl22),
  inference(superposition,[],[f436,f142])).
fof(f142,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.extreal.inf_const_alt',aHL_BOOL_CASES)).
fof(f695,plain,(
  ( ! [X48] : (s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK6('F_0',sK0,'type.extreal.extreal')) | ~p(s(bool,X48))) ) | $spl22),
  inference(resolution,[],[f436,f152])).
fof(f726,plain,(
  ( ! [X47] : (s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK6('F_0',sK0,'type.extreal.extreal')) | p(s(bool,X47))) ) | $spl22),
  inference(subsumption_resolution,[],[f694,f154])).
fof(f694,plain,(
  ( ! [X47] : (s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK6('F_0',sK0,'type.extreal.extreal')) | p(s(bool,X47)) | ~p(s(bool,'T_0'))) ) | $spl22),
  inference(resolution,[],[f436,f159])).
fof(f159,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f94])).
fof(f693,plain,(
  ( ! [X46] : (s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK6('F_0',sK0,'type.extreal.extreal')) | p(s(bool,X46))) ) | $spl22),
  inference(resolution,[],[f436,f168])).
fof(f692,plain,(
  ( ! [X45] : (s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK6('F_0',sK0,'type.extreal.extreal')) | p(s(bool,X45))) ) | $spl22),
  inference(resolution,[],[f436,f170])).
fof(f170,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0] : ((~p(s(bool,X0)) | ~p(s(bool,'F_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,X0)) & p(s(bool,'F_0')))))),
  inference(flattening,[],[f101])).
fof(f101,plain,(
  ! [X0] : (((~p(s(bool,X0)) | ~p(s(bool,'F_0'))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,X0)) & p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0] : ((p(s(bool,X0)) & p(s(bool,'F_0'))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X10] : ((p(s(bool,X10)) & p(s(bool,'F_0'))) <=> p(s(bool,'F_0')))),
  file('thm.extreal.inf_const_alt','thm.bool.AND_CLAUSES|split|3')).
fof(f721,plain,(
  ( ! [X2,X0] : (s(bool,X0) != s(bool,X2) | s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK6(X0,sK0,'type.extreal.extreal')) | p(s(bool,X2))) ) | $spl40),
  inference(cnf_transformation,[],[f721_D])).
fof(f721_D,plain,(
  ( ! [X2,X0] : (s(bool,X0) != s(bool,X2) | s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK6(X0,sK0,'type.extreal.extreal')) | p(s(bool,X2))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f689,plain,(
  ( ! [X43,X41,X42] : (s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK6(X41,sK0,'type.extreal.extreal')) | p(s(bool,X42)) | p(s(bool,X43)) | s(bool,X41) != s(bool,X43)) ) | $spl22),
  inference(resolution,[],[f436,f190])).
fof(f679,plain,(
  ( ! [X30,X28,X31,X29] : (s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK6(X28,sK0,'type.extreal.extreal')) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) ) | $spl22),
  inference(resolution,[],[f436,f131])).
fof(f226,plain,(
  ( ! [X8,X9] : (s(bool,X8) != s(bool,X9) | p(s(bool,X8)) | ~p(s(bool,X9))) ) | $spl6),
  inference(cnf_transformation,[],[f226_D])).
fof(f226_D,plain,(
  ( ! [X8,X9] : (s(bool,X8) != s(bool,X9) | p(s(bool,X8)) | ~p(s(bool,X9))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f722,plain,(
  ( ! [X4,X7] : (s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK6(X4,sK0,'type.extreal.extreal')) | s(bool,X4) != s(bool,X7)) ) | $spl42),
  inference(cnf_transformation,[],[f722_D])).
fof(f722_D,plain,(
  ( ! [X4,X7] : (s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK6(X4,sK0,'type.extreal.extreal')) | s(bool,X4) != s(bool,X7)) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f678,plain,(
  ( ! [X26,X24,X27,X25] : (s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK6(X24,sK0,'type.extreal.extreal')) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X26) | s(bool,X24) != s(bool,X27)) ) | $spl22),
  inference(resolution,[],[f436,f131])).
fof(f676,plain,(
  ( ! [X19,X17,X18,X16] : (s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK6(X16,sK0,'type.extreal.extreal')) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) ) | $spl22),
  inference(resolution,[],[f436,f130])).
fof(f675,plain,(
  ( ! [X14,X12,X15,X13] : (s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK6(X12,sK0,'type.extreal.extreal')) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl22),
  inference(resolution,[],[f436,f129])).
fof(f674,plain,(
  ( ! [X10,X8,X11,X9] : (s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK6(X8,sK0,'type.extreal.extreal')) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) ) | $spl22),
  inference(resolution,[],[f436,f127])).
fof(f673,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK6(X4,sK0,'type.extreal.extreal')) | p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X5) != s(bool,X6) | s(bool,X4) != s(bool,X7)) ) | $spl22),
  inference(resolution,[],[f436,f127])).
fof(f672,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK6(X0,sK0,'type.extreal.extreal')) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl22),
  inference(resolution,[],[f436,f125])).
fof(f628,plain,(
  s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK7('F_0',sK0,'type.extreal.extreal')) | $spl38),
  inference(cnf_transformation,[],[f628_D])).
fof(f628_D,plain,(
  s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK7('F_0',sK0,'type.extreal.extreal')) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f622,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK7('F_0',sK0,'type.extreal.extreal'))) ) | $spl20),
  inference(duplicate_literal_removal,[],[f612])).
fof(f612,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK7('F_0',sK0,'type.extreal.extreal')) | p(s(bool,X2))) ) | $spl20),
  inference(superposition,[],[f435,f135])).
fof(f623,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK7('F_0',sK0,'type.extreal.extreal'))) ) | $spl20),
  inference(duplicate_literal_removal,[],[f611])).
fof(f611,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK7('F_0',sK0,'type.extreal.extreal')) | p(s(bool,X1))) ) | $spl20),
  inference(superposition,[],[f435,f137])).
fof(f630,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK7('F_0',sK0,'type.extreal.extreal'))) ) | $spl20),
  inference(subsumption_resolution,[],[f610,f138])).
fof(f610,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK7('F_0',sK0,'type.extreal.extreal')) | s(bool,'T_0') = s(bool,X0)) ) | $spl20),
  inference(superposition,[],[f435,f142])).
fof(f598,plain,(
  ( ! [X48] : (s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK7('F_0',sK0,'type.extreal.extreal')) | ~p(s(bool,X48))) ) | $spl20),
  inference(resolution,[],[f435,f152])).
fof(f629,plain,(
  ( ! [X47] : (s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK7('F_0',sK0,'type.extreal.extreal')) | p(s(bool,X47))) ) | $spl20),
  inference(subsumption_resolution,[],[f597,f154])).
fof(f597,plain,(
  ( ! [X47] : (s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK7('F_0',sK0,'type.extreal.extreal')) | p(s(bool,X47)) | ~p(s(bool,'T_0'))) ) | $spl20),
  inference(resolution,[],[f435,f159])).
fof(f596,plain,(
  ( ! [X46] : (s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK7('F_0',sK0,'type.extreal.extreal')) | p(s(bool,X46))) ) | $spl20),
  inference(resolution,[],[f435,f168])).
fof(f595,plain,(
  ( ! [X45] : (s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK7('F_0',sK0,'type.extreal.extreal')) | p(s(bool,X45))) ) | $spl20),
  inference(resolution,[],[f435,f170])).
fof(f624,plain,(
  ( ! [X2,X0] : (s(bool,X0) != s(bool,X2) | s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK7(X0,sK0,'type.extreal.extreal')) | p(s(bool,X2))) ) | $spl34),
  inference(cnf_transformation,[],[f624_D])).
fof(f624_D,plain,(
  ( ! [X2,X0] : (s(bool,X0) != s(bool,X2) | s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK7(X0,sK0,'type.extreal.extreal')) | p(s(bool,X2))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f592,plain,(
  ( ! [X43,X41,X42] : (s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK7(X41,sK0,'type.extreal.extreal')) | p(s(bool,X42)) | p(s(bool,X43)) | s(bool,X41) != s(bool,X43)) ) | $spl20),
  inference(resolution,[],[f435,f190])).
fof(f582,plain,(
  ( ! [X30,X28,X31,X29] : (s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK7(X28,sK0,'type.extreal.extreal')) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) ) | $spl20),
  inference(resolution,[],[f435,f131])).
fof(f625,plain,(
  ( ! [X4,X7] : (s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK7(X4,sK0,'type.extreal.extreal')) | s(bool,X4) != s(bool,X7)) ) | $spl36),
  inference(cnf_transformation,[],[f625_D])).
fof(f625_D,plain,(
  ( ! [X4,X7] : (s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK7(X4,sK0,'type.extreal.extreal')) | s(bool,X4) != s(bool,X7)) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f581,plain,(
  ( ! [X26,X24,X27,X25] : (s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK7(X24,sK0,'type.extreal.extreal')) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X26) | s(bool,X24) != s(bool,X27)) ) | $spl20),
  inference(resolution,[],[f435,f131])).
fof(f579,plain,(
  ( ! [X19,X17,X18,X16] : (s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK7(X16,sK0,'type.extreal.extreal')) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) ) | $spl20),
  inference(resolution,[],[f435,f130])).
fof(f578,plain,(
  ( ! [X14,X12,X15,X13] : (s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK7(X12,sK0,'type.extreal.extreal')) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl20),
  inference(resolution,[],[f435,f129])).
fof(f577,plain,(
  ( ! [X10,X8,X11,X9] : (s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK7(X8,sK0,'type.extreal.extreal')) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) ) | $spl20),
  inference(resolution,[],[f435,f127])).
fof(f576,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK7(X4,sK0,'type.extreal.extreal')) | p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X5) != s(bool,X6) | s(bool,X4) != s(bool,X7)) ) | $spl20),
  inference(resolution,[],[f435,f127])).
fof(f575,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK7(X0,sK0,'type.extreal.extreal')) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl20),
  inference(resolution,[],[f435,f125])).
fof(f524,plain,(
  $false | $spl18),
  inference(resolution,[],[f434,f447])).
fof(f517,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,X6) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK1)))) ) | $spl30),
  inference(cnf_transformation,[],[f517_D])).
fof(f517_D,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,X6) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK1)))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f508,plain,(
  ( ! [X30,X29] : (p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X30) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK1)))) )),
  inference(resolution,[],[f447,f190])).
fof(f498,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK1))) | s(bool,X27) != s(bool,X28)) )),
  inference(resolution,[],[f447,f131])).
fof(f518,plain,(
  ( ! [X10] : (s(bool,X10) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK1)))) ) | $spl32),
  inference(cnf_transformation,[],[f518_D])).
fof(f518_D,plain,(
  ( ! [X10] : (s(bool,X10) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK1)))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f497,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X23) != s(bool,X24) | s(bool,X25) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK1)))) )),
  inference(resolution,[],[f447,f131])).
fof(f495,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X19) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK1)))) )),
  inference(resolution,[],[f447,f130])).
fof(f494,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK1)))) )),
  inference(resolution,[],[f447,f129])).
fof(f493,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK1))) | s(bool,X12) != s(bool,X13)) )),
  inference(resolution,[],[f447,f127])).
fof(f492,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK1)))) )),
  inference(resolution,[],[f447,f127])).
fof(f491,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK1)))) )),
  inference(resolution,[],[f447,f125])).
fof(f440,plain,(
  ( ! [X33] : (s('type.extreal.extreal',sK1) = s('type.extreal.extreal',X33)) ) | $spl28),
  inference(cnf_transformation,[],[f440_D])).
fof(f440_D,plain,(
  ( ! [X33] : (s('type.extreal.extreal',sK1) = s('type.extreal.extreal',X33)) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f417,plain,(
  ( ! [X64,X62,X63] : (s('type.extreal.extreal',sK1) = s('type.extreal.extreal',X62) | ~p(s(bool,X63)) | p(s(bool,X64)) | s(bool,X63) != s(bool,X64)) )),
  inference(resolution,[],[f115,f190])).
fof(f416,plain,(
  ( ! [X61,X59,X60] : (s('type.extreal.extreal',sK1) = s('type.extreal.extreal',X59) | p(s(bool,X60)) | ~p(s(bool,X61)) | s(bool,X61) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',X59)))) )),
  inference(resolution,[],[f115,f190])).
fof(f403,plain,(
  ( ! [X39,X37,X38,X40] : (s('type.extreal.extreal',sK1) = s('type.extreal.extreal',X37) | ~p(s(bool,X38)) | ~p(s(bool,X39)) | s(bool,X39) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',X37))) | s(bool,X38) != s(bool,X40)) )),
  inference(resolution,[],[f115,f131])).
fof(f439,plain,(
  ( ! [X33,X36,X34] : (s('type.extreal.extreal',sK1) = s('type.extreal.extreal',X33) | ~p(s(bool,X34)) | p(s(bool,X36)) | s(bool,X34) != s(bool,X36)) )),
  inference(subsumption_resolution,[],[f402,f190])).
fof(f402,plain,(
  ( ! [X35,X33,X36,X34] : (s('type.extreal.extreal',sK1) = s('type.extreal.extreal',X33) | ~p(s(bool,X34)) | ~p(s(bool,X35)) | p(s(bool,X36)) | s(bool,X34) != s(bool,X36)) )),
  inference(resolution,[],[f115,f130])).
fof(f401,plain,(
  ( ! [X30,X31,X29,X32] : (s('type.extreal.extreal',sK1) = s('type.extreal.extreal',X29) | p(s(bool,X30)) | ~p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',X29)))) )),
  inference(resolution,[],[f115,f130])).
fof(f438,plain,(
  ( ! [X13,X16] : (s('type.extreal.extreal',sK1) = s('type.extreal.extreal',X13) | s(bool,X16) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',X13)))) ) | $spl26),
  inference(cnf_transformation,[],[f438_D])).
fof(f438_D,plain,(
  ( ! [X13,X16] : (s('type.extreal.extreal',sK1) = s('type.extreal.extreal',X13) | s(bool,X16) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',X13)))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f400,plain,(
  ( ! [X28,X26,X27,X25] : (s('type.extreal.extreal',sK1) = s('type.extreal.extreal',X25) | ~p(s(bool,X26)) | p(s(bool,X27)) | s(bool,X28) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',X25))) | s(bool,X26) != s(bool,X27)) )),
  inference(resolution,[],[f115,f129])).
fof(f399,plain,(
  ( ! [X24,X23,X21,X22] : (s('type.extreal.extreal',sK1) = s('type.extreal.extreal',X21) | p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X24) | s(bool,X23) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',X21)))) )),
  inference(resolution,[],[f115,f129])).
fof(f398,plain,(
  ( ! [X19,X17,X20,X18] : (s('type.extreal.extreal',sK1) = s('type.extreal.extreal',X17) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',X17))) | s(bool,X18) != s(bool,X20)) )),
  inference(resolution,[],[f115,f127])).
fof(f397,plain,(
  ( ! [X14,X15,X13,X16] : (s('type.extreal.extreal',sK1) = s('type.extreal.extreal',X13) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X16) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',X13))) | s(bool,X14) != s(bool,X15)) )),
  inference(resolution,[],[f115,f125])).
fof(f396,plain,(
  ( ! [X12,X10,X11,X9] : (s('type.extreal.extreal',sK1) = s('type.extreal.extreal',X9) | p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',X9)))) )),
  inference(resolution,[],[f115,f125])).
fof(f434,plain,(
  ( ! [X6] : (~p(s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',X6))))) ) | $spl18),
  inference(cnf_transformation,[],[f434_D])).
fof(f434_D,plain,(
  ( ! [X6] : (~p(s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',X6))))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f436,plain,(
  ( ! [X7] : (p(s(bool,X7)) | s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK6(X7,sK0,'type.extreal.extreal'))) ) | $spl22),
  inference(cnf_transformation,[],[f436_D])).
fof(f436_D,plain,(
  ( ! [X7] : (p(s(bool,X7)) | s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK6(X7,sK0,'type.extreal.extreal'))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f395,plain,(
  ( ! [X8,X7] : (s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK6(X7,sK0,'type.extreal.extreal')) | p(s(bool,X7)) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',X8))))) )),
  inference(resolution,[],[f115,f184])).
fof(f435,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK7(X5,sK0,'type.extreal.extreal'))) ) | $spl20),
  inference(cnf_transformation,[],[f435_D])).
fof(f435_D,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK7(X5,sK0,'type.extreal.extreal'))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f394,plain,(
  ( ! [X6,X5] : (s('type.extreal.extreal',sK1) = s('type.extreal.extreal',sK7(X5,sK0,'type.extreal.extreal')) | p(s(bool,X5)) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',X6))))) )),
  inference(resolution,[],[f115,f186])).
fof(f328,plain,(
  ( ! [X7] : (s(bool,'T_0') != s(bool,X7)) ) | $spl14),
  inference(cnf_transformation,[],[f328_D])).
fof(f328_D,plain,(
  ( ! [X7] : (s(bool,'T_0') != s(bool,X7)) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f358,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') != s(bool,X13) | ~p(s(bool,X12))) ) | $spl8),
  inference(duplicate_literal_removal,[],[f339])).
fof(f339,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') != s(bool,X13) | ~p(s(bool,X12)) | ~p(s(bool,X12))) ) | $spl8),
  inference(superposition,[],[f227,f139])).
fof(f359,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') != s(bool,X11) | ~p(s(bool,X10))) ) | $spl8),
  inference(duplicate_literal_removal,[],[f338])).
fof(f338,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') != s(bool,X11) | ~p(s(bool,X10)) | ~p(s(bool,X10))) ) | $spl8),
  inference(superposition,[],[f227,f141])).
fof(f141,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f39])).
fof(f39,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X10] : (s(bool,'T_0') = s(bool,X10) <=> p(s(bool,X10)))),
  file('thm.extreal.inf_const_alt','thm.bool.EQ_CLAUSES|split|0')).
fof(f330,plain,(
  ( ! [X9] : (s(bool,'F_0') != s(bool,X9)) ) | $spl16),
  inference(cnf_transformation,[],[f330_D])).
fof(f330_D,plain,(
  ( ! [X9] : (s(bool,'F_0') != s(bool,X9)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f363,plain,(
  ( ! [X8,X9] : (s(bool,'F_0') != s(bool,X9) | ~p(s(bool,X8))) ) | $spl8),
  inference(subsumption_resolution,[],[f337,f227])).
fof(f337,plain,(
  ( ! [X8,X9] : (s(bool,'F_0') != s(bool,X9) | ~p(s(bool,X8)) | s(bool,'T_0') = s(bool,X8)) ) | $spl8),
  inference(superposition,[],[f227,f142])).
fof(f362,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X7) | ~p(s(bool,X6))) ) | $spl8),
  inference(subsumption_resolution,[],[f336,f227])).
fof(f336,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X7) | ~p(s(bool,X6)) | s(bool,'F_0') = s(bool,X6)) ) | $spl8),
  inference(superposition,[],[f227,f142])).
fof(f324,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,X17) | p(s(bool,X16))) ) | $spl2),
  inference(duplicate_literal_removal,[],[f307])).
fof(f307,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,X17) | p(s(bool,X16)) | p(s(bool,X16))) ) | $spl2),
  inference(superposition,[],[f224,f135])).
fof(f325,plain,(
  ( ! [X14,X15] : (s(bool,'F_0') != s(bool,X15) | p(s(bool,X14))) ) | $spl2),
  inference(duplicate_literal_removal,[],[f306])).
fof(f306,plain,(
  ( ! [X14,X15] : (s(bool,'F_0') != s(bool,X15) | p(s(bool,X14)) | p(s(bool,X14))) ) | $spl2),
  inference(superposition,[],[f224,f137])).
fof(f329,plain,(
  ( ! [X8,X9] : (s(bool,'F_0') != s(bool,X9) | p(s(bool,X8))) ) | $spl2),
  inference(subsumption_resolution,[],[f303,f224])).
fof(f303,plain,(
  ( ! [X8,X9] : (s(bool,'F_0') != s(bool,X9) | p(s(bool,X8)) | s(bool,'T_0') = s(bool,X8)) ) | $spl2),
  inference(superposition,[],[f224,f142])).
fof(f327,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X7) | p(s(bool,X6))) ) | $spl2),
  inference(subsumption_resolution,[],[f302,f224])).
fof(f302,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X7) | p(s(bool,X6)) | s(bool,'F_0') = s(bool,X6)) ) | $spl2),
  inference(superposition,[],[f224,f142])).
fof(f294,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f268,f229])).
fof(f268,plain,(
  ( ! [X42] : (~p(s(bool,X42))) ) | $spl12),
  inference(resolution,[],[f229,f152])).
fof(f264,plain,(
  $false | $spl12),
  inference(resolution,[],[f229,f181])).
fof(f223,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,X6) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK2)))) ) | $spl0),
  inference(cnf_transformation,[],[f223_D])).
fof(f223_D,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,X6) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK2)))) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f214,plain,(
  ( ! [X30,X29] : (p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X30) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK2)))) )),
  inference(resolution,[],[f114,f190])).
fof(f204,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK2))) | s(bool,X27) != s(bool,X28)) )),
  inference(resolution,[],[f114,f131])).
fof(f225,plain,(
  ( ! [X10] : (s(bool,X10) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK2)))) ) | $spl4),
  inference(cnf_transformation,[],[f225_D])).
fof(f225_D,plain,(
  ( ! [X10] : (s(bool,X10) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK2)))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f203,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X23) != s(bool,X24) | s(bool,X25) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK2)))) )),
  inference(resolution,[],[f114,f131])).
fof(f201,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X19) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK2)))) )),
  inference(resolution,[],[f114,f130])).
fof(f200,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK2)))) )),
  inference(resolution,[],[f114,f129])).
fof(f199,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK2))) | s(bool,X12) != s(bool,X13)) )),
  inference(resolution,[],[f114,f127])).
fof(f198,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK2)))) )),
  inference(resolution,[],[f114,f127])).
fof(f197,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,happ(s(fun('type.extreal.extreal',bool),sK0),s('type.extreal.extreal',sK2)))) )),
  inference(resolution,[],[f114,f125])).
% SZS output end Proof for thm.extreal
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 103
% Passive clauses: 651
% Generated clauses: 2687
% Final active clauses: 76
% Final passive clauses: 383
% Input formulas: 27
% Initial clauses: 76
% 
% Selected by SInE selection: 27
% SInE iterations: 1
% 
% Duplicate literals: 42
% Trivial inequalities: 46
% Fw subsumption resolutions: 387
% Fw demodulations: 197
% Bw demodulations: 9
% 
% Simple tautologies: 92
% Equational tautologies: 1
% Forward subsumptions: 1072
% 
% Binary resolution: 754
% Forward superposition: 729
% Backward superposition: 437
% Equality resolution: 15
% 
% Split clauses: 211
% Split components: 450
% SAT solver clauses: 221
% SAT solver unit clauses: 7
% SAT solver binary clauses: 55
% 
% Sat splits: 211
% Sat splitting refutations: 9
% 
% Memory used [KB]: 2174
% Time elapsed: 0.080 s
% ------------------------------
% ------------------------------
% Success in time 11.76 s
