Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.rat
% (79820)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 686
% Passive clauses: 10048
% Generated clauses: 18959
% Final active clauses: 678
% Final passive clauses: 8872
% Input formulas: 69
% Initial clauses: 260
% 
% Duplicate literals: 285
% Trivial inequalities: 89
% Fw subsumption resolutions: 1753
% Fw demodulations: 256
% Bw demodulations: 6
% 
% Simple tautologies: 261
% Equational tautologies: 55
% Forward subsumptions: 6047
% 
% Binary resolution: 11098
% Factoring: 1000
% Forward superposition: 1505
% Backward superposition: 2514
% Self superposition: 22
% Equality factoring: 26
% Equality resolution: 17
% 
% Split clauses: 278
% Split components: 557
% SAT solver clauses: 584
% SAT solver unit clauses: 24
% SAT solver binary clauses: 180
% 
% Sat splits: 278
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 584
% TWLsolver calls for satisfiability: 365
% 
% Memory used [KB]: 10490
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.rat
% (79850)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 500
% Passive clauses: 3428
% Generated clauses: 12995
% Final active clauses: 480
% Final passive clauses: 2827
% Input formulas: 69
% Initial clauses: 260
% 
% Duplicate literals: 200
% Trivial inequalities: 23
% Fw subsumption resolutions: 920
% Fw demodulations: 425
% Bw demodulations: 116
% Condensations: 246
% Global subsumptions: 22
% 
% Simple tautologies: 293
% Equational tautologies: 22
% Forward subsumptions: 3575
% 
% Binary resolution: 4617
% Factoring: 100
% Forward superposition: 1221
% Backward superposition: 1353
% Self superposition: 22
% Equality factoring: 26
% Equality resolution: 4
% 
% Split clauses: 20
% Split components: 40
% SAT solver clauses: 3986
% SAT solver unit clauses: 250
% SAT solver binary clauses: 105
% SAT solver learnt clauses: 44
% SAT solver learnt literals: 10
% 
% Sat splits: 20
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 3898
% TWLsolver calls for satisfiability: 2120
% 
% Memory used [KB]: 14967
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.rat
% (79879)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1381
% Passive clauses: 7345
% Generated clauses: 56948
% Final active clauses: 889
% Final passive clauses: 2403
% Input formulas: 69
% Initial clauses: 260
% 
% Duplicate literals: 2237
% Trivial inequalities: 72
% Fw subsumption resolutions: 2556
% Fw demodulations: 4171
% Bw demodulations: 3018
% Condensations: 4024
% Global subsumptions: 2368
% 
% Simple tautologies: 176
% Equational tautologies: 1388
% Forward subsumptions: 28720
% Bw demodulations to eq. taut.: 456
% 
% Unit resulting resolution: 21193
% Factoring: 12
% Forward superposition: 7595
% Backward superposition: 6893
% Self superposition: 152
% Equality factoring: 133
% Equality resolution: 4
% 
% Split clauses: 275
% Split components: 564
% SAT solver clauses: 13521
% SAT solver unit clauses: 3115
% SAT solver binary clauses: 3434
% 
% Sat splits: 275
% Sat splitting refutations: 3255
% 
% Memory used [KB]: 171681
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.rat
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.rat
% SZS output start Proof for thm.rat
fof(f8229,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f610,f740,f611,f612,f613,f747,f614,f615,f904,f936,f943,f940,f938,f944,f1099,f1105,f935,f933,f1273,f1403,f1381,f1070,f1405,f1487,f1700,f742,f1836,f748,f751,f2105,f2139,f2153,f2143,f2141,f2218,f2312,f2314,f2315,f2346,f2313,f2361,f2316,f2138,f2136,f2903,f2913,f2928,f2937,f2945,f2952,f2960,f2969,f2978,f2983,f3017,f3026,f3021,f3019,f3064,f3210,f3217,f3224,f3232,f3239,f3245,f3252,f3260,f3268,f2906,f3272,f3016,f3014,f3522,f3552,f3559,f3567,f3574,f3580,f3597,f3605,f3613,f3622,f3533,f3758,f3757,f3756,f3755,f3754,f3738,f3821,f3822,f3823,f3824,f3970,f3969,f3968,f3966,f3948,f3947,f3946,f4053,f4054,f4055,f4056,f4057,f4058,f4068,f4075,f4083,f4090,f4096,f4103,f4111,f4119,f4440,f749,f4441,f745,f4442,f741,f743,f5203,f5237,f5246,f5241,f5239,f5284,f5451,f5458,f5465,f5473,f5480,f5497,f5504,f5512,f5520,f3525,f5539,f5236,f5234,f5577,f5611,f5620,f5615,f5613,f5672,f5823,f2907,f5830,f2915,f5851,f2931,f5859,f2939,f5866,f2947,f5872,f2955,f5879,f2964,f5899,f2973,f5907,f2982,f3625,f5940,f5610,f5608,f5995,f6029,f6046,f6033,f6031,f6092,f6192,f6251,f6252,f6259,f6260,f6679,f6028,f6026,f6725,f6759,f6773,f6763,f6761,f6852,f7030,f7091,f7092,f753,f7519,f6758,f6756,f7585,f7586,f7800,f7814,f7801,f7816,f7802,f7818,f7803,f7820,f7804,f7822,f7805,f7824,f7851,f7865,f7852,f7867,f7853,f7869,f7854,f7871,f7855,f7873,f7856,f7875,f3535,f7895,f7904,f7896,f7906,f7897,f3534,f7908,f8183,f8216,f8214,f8217,f8221,f8219,f8224])).
fof(f8224,plain,(
  $false | ($spl0 | $spl5)),
  inference(subsumption_resolution,[],[f8223,f740])).
fof(f8223,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK19(sK13))))) | $spl5),
  inference(forward_demodulation,[],[f8222,f8043])).
fof(f8043,plain,(
  s('type.num.num',sK19(sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK19(sK13))))) | $spl5),
  inference(subsumption_resolution,[],[f8019,f353])).
fof(f353,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(cnf_transformation,[],[f220])).
fof(f220,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) & p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) & ! [X2] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2)))))) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))))))))) & ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13,sK14])],[f219])).
fof(f219,plain,(
  ? [X0] : (p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.|48|_0')))) & p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) & ! [X2] : (~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2)))))) | p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))))))))) & ? [X1] : ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X1)))))),
  inference(rectify,[],[f161])).
fof(f161,plain,(
  ? [X0] : (p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.|48|_0')))) & p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) & ! [X1] : (~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))) | p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))))))) & ? [X2] : ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2)))))),
  inference(flattening,[],[f160])).
fof(f160,plain,(
  ? [X0] : ((p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.|48|_0')))) & p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) & ! [X1] : (~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))) | p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))))))) & ? [X2] : ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2)))))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ~! [X0] : ((p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.|48|_0')))) & p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) & ! [X1] : (p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))) => p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))))))) => ! [X2] : p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ((p(s(bool,happ(s(fun('type.num.num',bool),X6),s('type.num.num','const.num.|48|_0')))) & p(s(bool,happ(s(fun('type.num.num',bool),X6),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) & ! [X7] : (p(s(bool,happ(s(fun('type.num.num',bool),X6),s('type.num.num','const.num.SUC_1'(s('type.num.num',X7)))))) => p(s(bool,happ(s(fun('type.num.num',bool),X6),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X7)))))))))) => ! [X8] : p(s(bool,happ(s(fun('type.num.num',bool),X6),s('type.num.num',X8)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ((p(s(bool,happ(s(fun('type.num.num',bool),X6),s('type.num.num','const.num.|48|_0')))) & p(s(bool,happ(s(fun('type.num.num',bool),X6),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) & ! [X7] : (p(s(bool,happ(s(fun('type.num.num',bool),X6),s('type.num.num','const.num.SUC_1'(s('type.num.num',X7)))))) => p(s(bool,happ(s(fun('type.num.num',bool),X6),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X7)))))))))) => ! [X8] : p(s(bool,happ(s(fun('type.num.num',bool),X6),s('type.num.num',X8)))))),
  file('thm.rat.rat_of_num_ind','thm.rat.rat_of_num_ind')).
fof(f8019,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num',sK19(sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK19(sK13))))) | $spl5),
  inference(superposition,[],[f743,f365])).
fof(f365,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(X0)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f231])).
fof(f231,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK20])],[f75])).
fof(f75,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | ? [X1] : s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X22] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | ? [X7] : s('type.num.num',X22) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))))),
  file('thm.rat.rat_of_num_ind','thm.arithmetic.num_CASES')).
fof(f8222,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK19(sK13)))))))) | $spl5),
  inference(subsumption_resolution,[],[f8185,f743])).
fof(f8185,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK19(sK13))))))) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK19(sK13)))))))) | $spl5),
  inference(superposition,[],[f354,f8043])).
fof(f354,plain,(
  ( ! [X2] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X2)))))))) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))))))) )),
  inference(cnf_transformation,[],[f220])).
fof(f8219,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK19(sK13)) | $spl144),
  inference(cnf_transformation,[],[f8219_D])).
fof(f8219_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK19(sK13)) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f8221,plain,(
  s('type.num.num',sK19(sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.num.SUC_1'(s('type.num.num',sK20(sK19(sK13)))))))) | $spl146),
  inference(cnf_transformation,[],[f8221_D])).
fof(f8221_D,plain,(
  s('type.num.num',sK19(sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.num.SUC_1'(s('type.num.num',sK20(sK19(sK13)))))))) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f8217,plain,(
  s('type.num.num',sK19(sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.num.SUC_1'(s('type.num.num',sK20(sK19(sK13)))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK19(sK13)) | $spl5),
  inference(forward_demodulation,[],[f8184,f8043])).
fof(f8184,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK19(sK13)) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK19(sK13))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.num.SUC_1'(s('type.num.num',sK20(sK19(sK13)))))))) | $spl5),
  inference(superposition,[],[f8043,f365])).
fof(f8214,plain,(
  s('type.num.num',sK20(sK19(sK13))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK20(sK19(sK13)))))) | $spl140),
  inference(cnf_transformation,[],[f8214_D])).
fof(f8214_D,plain,(
  s('type.num.num',sK20(sK19(sK13))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK20(sK19(sK13)))))) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f8216,plain,(
  s('type.num.num',sK19(sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | $spl142),
  inference(cnf_transformation,[],[f8216_D])).
fof(f8216_D,plain,(
  s('type.num.num',sK19(sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f8183,plain,(
  s('type.num.num',sK19(sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK20(sK19(sK13))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK20(sK19(sK13)))))) | $spl5),
  inference(superposition,[],[f8043,f365])).
fof(f7908,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK19(sK13)))) | $spl139),
  inference(cnf_transformation,[],[f7908_D])).
fof(f7908_D,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK19(sK13)))) <=> ~$spl139),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl139])])).
fof(f3534,plain,(
  ( ! [X1] : (p(s(bool,X1))) ) | $spl70),
  inference(cnf_transformation,[],[f3534_D])).
fof(f3534_D,plain,(
  ( ! [X1] : (p(s(bool,X1))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f7897,plain,(
  ( ! [X17] : (p(s(bool,X17)) | s(bool,'F_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK19(sK13))))) ) | $spl0),
  inference(resolution,[],[f7753,f1765])).
fof(f1765,plain,(
  ( ! [X121,X120] : (sP6(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK19(sK13))),X120,X121)) ) | $spl0),
  inference(resolution,[],[f740,f435])).
fof(f435,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f265])).
fof(f265,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(flattening,[],[f264])).
fof(f264,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(nnf_transformation,[],[f208])).
fof(f208,plain,(
  ! [X0,X1,X2] : (sP6(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f7753,plain,(
  ( ! [X4,X5] : (~sP6(X5,'F_0',X4) | p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5)) )),
  inference(subsumption_resolution,[],[f7747,f427])).
fof(f427,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f260])).
fof(f260,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f259])).
fof(f259,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f258])).
fof(f258,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f210])).
fof(f210,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f7747,plain,(
  ( ! [X4,X5] : (~sP8(X4,X5,'F_0') | ~sP6(X5,'F_0',X4) | p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5)) )),
  inference(resolution,[],[f7712,f439])).
fof(f439,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | s(bool,X0) != s(bool,X1) | ~sP9(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f266])).
fof(f266,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP9(X2,X1,X0)) & (~sP9(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f212])).
fof(f212,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP9(X2,X1,X0))),
  inference(definition_folding,[],[f83,f211,f210,f209,f208])).
fof(f209,plain,(
  ! [X1,X0,X2] : (sP7(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f211,plain,(
  ! [X2,X1,X0] : (sP9(X2,X1,X0) <=> (sP6(X0,X1,X2) & sP7(X1,X0,X2) & sP8(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f83,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f82])).
fof(f82,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X14,X15,X16] : ((p(s(bool,X16)) <=> s(bool,X14) = s(bool,X15)) <=> ((p(s(bool,X16)) | p(s(bool,X15)) | p(s(bool,X14))) & (p(s(bool,X16)) | ~p(s(bool,X14)) | ~p(s(bool,X15))) & (p(s(bool,X15)) | ~p(s(bool,X14)) | ~p(s(bool,X16))) & (p(s(bool,X14)) | ~p(s(bool,X15)) | ~p(s(bool,X16)))))),
  file('thm.rat.rat_of_num_ind','thm.sat.sat_2|49|_prove_ltb')).
fof(f7712,plain,(
  ( ! [X4,X5] : (sP9(X4,'F_0',X5) | ~sP8(X4,X5,'F_0') | ~sP6(X5,'F_0',X4)) )),
  inference(subsumption_resolution,[],[f7710,f595])).
fof(f595,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f157])).
fof(f157,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.rat.rat_of_num_ind',aHL_FALSITY)).
fof(f7710,plain,(
  ( ! [X4,X5] : (sP9(X4,'F_0',X5) | p(s(bool,'F_0')) | ~sP8(X4,X5,'F_0') | ~sP6(X5,'F_0',X4)) )),
  inference(resolution,[],[f7707,f422])).
fof(f422,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | p(s(bool,X1)) | ~sP8(X0,X2,X1) | ~sP7(X1,X2,X0) | ~sP6(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f257])).
fof(f257,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | (sP6(X2,X1,X0) & sP7(X1,X2,X0) & sP8(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP6(X2,X1,X0) | ~sP7(X1,X2,X0) | ~sP8(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f256])).
fof(f256,plain,(
  ! [X2,X1,X0] : ((~sP9(X2,X1,X0) | (sP6(X0,X1,X2) & sP7(X1,X0,X2) & sP8(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP6(X0,X1,X2) | ~sP7(X1,X0,X2) | ~sP8(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP9(X2,X1,X0)))),
  inference(flattening,[],[f255])).
fof(f255,plain,(
  ! [X2,X1,X0] : ((~sP9(X2,X1,X0) | (sP6(X0,X1,X2) & sP7(X1,X0,X2) & sP8(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP6(X0,X1,X2) | ~sP7(X1,X0,X2) | ~sP8(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP9(X2,X1,X0)))),
  inference(nnf_transformation,[],[f211])).
fof(f7707,plain,(
  ( ! [X4,X5] : (sP7('F_0',X4,X5)) )),
  inference(subsumption_resolution,[],[f7691,f429])).
fof(f429,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f263])).
fof(f263,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f262])).
fof(f262,plain,(
  ! [X1,X0,X2] : ((~sP7(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP7(X1,X0,X2)))),
  inference(flattening,[],[f261])).
fof(f261,plain,(
  ! [X1,X0,X2] : ((~sP7(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP7(X1,X0,X2)))),
  inference(nnf_transformation,[],[f209])).
fof(f7691,plain,(
  ( ! [X4,X5] : (sP7('F_0',X4,X5) | p(s(bool,X5))) )),
  inference(resolution,[],[f7688,f427])).
fof(f7688,plain,(
  ( ! [X6,X7] : (~sP8(X7,X6,'F_0') | sP7('F_0',X6,X7)) )),
  inference(subsumption_resolution,[],[f7684,f430])).
fof(f430,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f263])).
fof(f7684,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | ~sP8(X7,X6,'F_0') | sP7('F_0',X6,X7)) )),
  inference(resolution,[],[f7678,f418])).
fof(f418,plain,(
  ( ! [X2,X0,X1] : (sP7(X1,X2,X0) | ~sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f257])).
fof(f7678,plain,(
  ( ! [X2,X3] : (sP9(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP8(X3,X2,'F_0')) )),
  inference(subsumption_resolution,[],[f7677,f581])).
fof(f581,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f348])).
fof(f348,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f142])).
fof(f142,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X12] : (s(bool,'T_0') = s(bool,X12) <=> p(s(bool,X12)))),
  file('thm.rat.rat_of_num_ind','thm.bool.EQ_CLAUSES|split|1')).
fof(f7677,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | sP9(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP8(X3,X2,'F_0')) )),
  inference(subsumption_resolution,[],[f7674,f7647])).
fof(f7647,plain,(
  ( ! [X4,X5] : (sP6(X4,'F_0',X5) | s(bool,'T_0') != s(bool,X4)) )),
  inference(subsumption_resolution,[],[f7642,f433])).
fof(f433,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f265])).
fof(f7642,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | p(s(bool,X5)) | sP6(X4,'F_0',X5)) )),
  inference(resolution,[],[f4144,f417])).
fof(f417,plain,(
  ( ! [X2,X0,X1] : (sP6(X2,X1,X0) | ~sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f257])).
fof(f4144,plain,(
  ( ! [X4,X5] : (sP9(X5,'F_0',X4) | s(bool,'T_0') != s(bool,X4) | p(s(bool,X5))) )),
  inference(superposition,[],[f4139,f436])).
fof(f436,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f266])).
fof(f4139,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(subsumption_resolution,[],[f4135,f352])).
fof(f352,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))))),
  inference(cnf_transformation,[],[f220])).
fof(f4135,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))))),
  inference(superposition,[],[f715,f581])).
fof(f715,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))),
  inference(resolution,[],[f352,f576])).
fof(f576,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f346])).
fof(f346,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f139])).
fof(f139,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f138])).
fof(f138,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X12] : (s(bool,'F_0') = s(bool,X12) <=> ~p(s(bool,X12)))),
  file('thm.rat.rat_of_num_ind','thm.bool.EQ_CLAUSES|split|2')).
fof(f7674,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | sP9(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP8(X3,X2,'F_0') | ~sP6(X2,'F_0',X3)) )),
  inference(resolution,[],[f7648,f421])).
fof(f421,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~p(s(bool,X2)) | ~sP8(X0,X2,X1) | ~sP7(X1,X2,X0) | ~sP6(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f257])).
fof(f7648,plain,(
  ( ! [X6,X7] : (sP7('F_0',X6,X7) | s(bool,'T_0') != s(bool,X6)) )),
  inference(subsumption_resolution,[],[f7643,f429])).
fof(f7643,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X6) | p(s(bool,X7)) | sP7('F_0',X6,X7)) )),
  inference(resolution,[],[f4144,f418])).
fof(f7906,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16(sK13)))) | $spl137),
  inference(cnf_transformation,[],[f7906_D])).
fof(f7906_D,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16(sK13)))) <=> ~$spl137),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl137])])).
fof(f7896,plain,(
  ( ! [X16] : (p(s(bool,X16)) | s(bool,'F_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16(sK13))))) ) | $spl8),
  inference(resolution,[],[f7753,f1901])).
fof(f1901,plain,(
  ( ! [X121,X120] : (sP6(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16(sK13))),X120,X121)) ) | $spl8),
  inference(resolution,[],[f747,f435])).
fof(f7904,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl135),
  inference(cnf_transformation,[],[f7904_D])).
fof(f7904_D,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl135),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl135])])).
fof(f7895,plain,(
  ( ! [X15] : (p(s(bool,X15)) | s(bool,'F_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(resolution,[],[f7753,f1009])).
fof(f1009,plain,(
  ( ! [X121,X120] : (sP6(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X120,X121)) )),
  inference(resolution,[],[f353,f435])).
fof(f3535,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl72),
  inference(cnf_transformation,[],[f3535_D])).
fof(f3535_D,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f7875,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK35(sK13,'type.num.num')))) | $spl132),
  inference(cnf_transformation,[],[f7875_D])).
fof(f7875_D,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK35(sK13,'type.num.num')))) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f7856,plain,(
  ( ! [X14] : (s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK35(sK13,'type.num.num')))) | ~p(s(bool,X14))) )),
  inference(resolution,[],[f7752,f2771])).
fof(f2771,plain,(
  ( ! [X101,X100] : (sP8(X100,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK35(sK13,'type.num.num'))),X101)) )),
  inference(resolution,[],[f763,f426])).
fof(f426,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f260])).
fof(f763,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK35(sK13,'type.num.num')))))),
  inference(resolution,[],[f355,f526])).
fof(f526,plain,(
  ( ! [X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK35(X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f313])).
fof(f313,plain,(
  ! [X0,X1] : ((! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK34(X1,X0)))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK35(X1,X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK34,sK35])],[f312])).
fof(f312,plain,(
  ! [X0,X1] : ((! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(rectify,[],[f311])).
fof(f311,plain,(
  ! [X0,X1] : ((! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(nnf_transformation,[],[f184])).
fof(f184,plain,(
  ! [X0,X1] : (? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(ennf_transformation,[],[f117])).
fof(f117,plain,(
  ! [X0,X1] : (~! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(flattening,[],[f116])).
fof(f116,plain,(
  ! [X0] : ! [X1] : (~! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(rectify,[],[f59])).
fof(f59,axiom,(
  ! [X13] : ! [X6] : (~! [X10] : p(s(bool,happ(s(fun(X13,bool),X6),s(X13,X10)))) <=> ? [X10] : ~p(s(bool,happ(s(fun(X13,bool),X6),s(X13,X10)))))),
  file('thm.rat.rat_of_num_ind','thm.bool.NOT_FORALL_THM')).
fof(f355,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))))),
  inference(cnf_transformation,[],[f220])).
fof(f7752,plain,(
  ( ! [X2,X3] : (~sP8(X2,X3,'F_0') | s(bool,'F_0') = s(bool,X3) | ~p(s(bool,X2))) )),
  inference(subsumption_resolution,[],[f7746,f433])).
fof(f7746,plain,(
  ( ! [X2,X3] : (~sP8(X2,X3,'F_0') | ~sP6(X3,'F_0',X2) | s(bool,'F_0') = s(bool,X3) | ~p(s(bool,X2))) )),
  inference(resolution,[],[f7712,f438])).
fof(f438,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X2)) | ~sP9(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f266])).
fof(f7873,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK34(sK13,'type.num.num')))) | $spl130),
  inference(cnf_transformation,[],[f7873_D])).
fof(f7873_D,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK34(sK13,'type.num.num')))) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f7855,plain,(
  ( ! [X13] : (s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK34(sK13,'type.num.num')))) | ~p(s(bool,X13))) )),
  inference(resolution,[],[f7752,f2612])).
fof(f2612,plain,(
  ( ! [X101,X100] : (sP8(X100,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK34(sK13,'type.num.num'))),X101)) )),
  inference(resolution,[],[f762,f426])).
fof(f762,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK34(sK13,'type.num.num')))))),
  inference(resolution,[],[f355,f525])).
fof(f525,plain,(
  ( ! [X0,X5,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK34(X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))) )),
  inference(cnf_transformation,[],[f313])).
fof(f7871,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK20(sK15(sK13))))) | $spl128),
  inference(cnf_transformation,[],[f7871_D])).
fof(f7871_D,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK20(sK15(sK13))))) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f7854,plain,(
  ( ! [X12] : (s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK20(sK15(sK13))))) | ~p(s(bool,X12))) ) | ($spl13 | $spl60)),
  inference(resolution,[],[f7752,f7287])).
fof(f7287,plain,(
  ( ! [X101,X100] : (sP8(X100,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK20(sK15(sK13)))),X101)) ) | ($spl13 | $spl60)),
  inference(resolution,[],[f6766,f426])).
fof(f6766,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK20(sK15(sK13)))))) | ($spl13 | $spl60)),
  inference(forward_demodulation,[],[f6765,f3014])).
fof(f6765,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK20(sK15(sK13))))))))) | ($spl13 | $spl60)),
  inference(subsumption_resolution,[],[f6764,f751])).
fof(f6764,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15(sK13))))) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK20(sK15(sK13))))))))) | ($spl13 | $spl60)),
  inference(forward_demodulation,[],[f6727,f2102])).
fof(f2102,plain,(
  s('type.num.num',sK15(sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK15(sK13))))) | $spl13),
  inference(subsumption_resolution,[],[f2080,f352])).
fof(f2080,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | s('type.num.num',sK15(sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK15(sK13))))) | $spl13),
  inference(superposition,[],[f751,f365])).
fof(f6727,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK15(sK13)))))))) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK20(sK15(sK13))))))))) | $spl60),
  inference(superposition,[],[f354,f3014])).
fof(f7869,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK20(sK20(sK14))))) | $spl126),
  inference(cnf_transformation,[],[f7869_D])).
fof(f7869_D,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK20(sK20(sK14))))) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f7853,plain,(
  ( ! [X11] : (s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK20(sK20(sK14))))) | ~p(s(bool,X11))) ) | ($spl16 | $spl30)),
  inference(resolution,[],[f7752,f6447])).
fof(f6447,plain,(
  ( ! [X101,X100] : (sP8(X100,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK20(sK20(sK14)))),X101)) ) | ($spl16 | $spl30)),
  inference(resolution,[],[f6036,f426])).
fof(f6036,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK20(sK20(sK14)))))) | ($spl16 | $spl30)),
  inference(forward_demodulation,[],[f6035,f2136])).
fof(f6035,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK20(sK20(sK14))))))))) | ($spl16 | $spl30)),
  inference(subsumption_resolution,[],[f6034,f2146])).
fof(f2146,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK20(sK14))))) | $spl16),
  inference(forward_demodulation,[],[f2145,f933])).
fof(f2145,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK20(sK14)))))))) | $spl16),
  inference(subsumption_resolution,[],[f2144,f355])).
fof(f2144,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK20(sK14)))))))) | $spl16),
  inference(forward_demodulation,[],[f2107,f889])).
fof(f889,plain,(
  s('type.num.num',sK14) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK14))))),
  inference(subsumption_resolution,[],[f865,f352])).
fof(f865,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | s('type.num.num',sK14) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK14))))),
  inference(superposition,[],[f355,f365])).
fof(f2107,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK14))))))) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK20(sK14)))))))) | $spl16),
  inference(superposition,[],[f354,f933])).
fof(f6034,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK20(sK14))))) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK20(sK20(sK14))))))))) | ($spl16 | $spl30)),
  inference(forward_demodulation,[],[f5997,f933])).
fof(f5997,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK20(sK14)))))))) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK20(sK20(sK14))))))))) | $spl30),
  inference(superposition,[],[f354,f2136])).
fof(f7867,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK20(sK14)))) | $spl124),
  inference(cnf_transformation,[],[f7867_D])).
fof(f7867_D,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK20(sK14)))) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f7852,plain,(
  ( ! [X10] : (s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK20(sK14)))) | ~p(s(bool,X10))) ) | $spl16),
  inference(resolution,[],[f7752,f3368])).
fof(f3368,plain,(
  ( ! [X101,X100] : (sP8(X100,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK20(sK14))),X101)) ) | $spl16),
  inference(resolution,[],[f2146,f426])).
fof(f7865,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15(sK13)))) | $spl122),
  inference(cnf_transformation,[],[f7865_D])).
fof(f7865_D,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15(sK13)))) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f7851,plain,(
  ( ! [X9] : (s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15(sK13)))) | ~p(s(bool,X9))) ) | $spl13),
  inference(resolution,[],[f7752,f2030])).
fof(f2030,plain,(
  ( ! [X101,X100] : (sP8(X100,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15(sK13))),X101)) ) | $spl13),
  inference(resolution,[],[f751,f426])).
fof(f7824,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK35(sK13,'type.num.num')))) | $spl121),
  inference(cnf_transformation,[],[f7824_D])).
fof(f7824_D,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK35(sK13,'type.num.num')))) <=> ~$spl121),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl121])])).
fof(f7805,plain,(
  ( ! [X14] : (s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK35(sK13,'type.num.num')))) | ~p(s(bool,X14))) )),
  inference(resolution,[],[f7794,f2771])).
fof(f7794,plain,(
  ( ! [X14,X15] : (~sP8(X14,X15,'F_0') | s(bool,'T_0') != s(bool,X15) | ~p(s(bool,X14))) )),
  inference(subsumption_resolution,[],[f7793,f433])).
fof(f7793,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | ~sP8(X14,X15,'F_0') | ~sP6(X15,'F_0',X14)) )),
  inference(subsumption_resolution,[],[f7792,f532])).
fof(f532,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f317])).
fof(f317,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f316])).
fof(f316,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f185])).
fof(f185,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f122])).
fof(f122,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X12] : ((p(s(bool,X12)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X12)))),
  file('thm.rat.rat_of_num_ind','thm.bool.IMP_CLAUSES|split|4')).
fof(f7792,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | p(s(bool,'F_0')) | ~sP8(X14,X15,'F_0') | ~sP6(X15,'F_0',X14)) )),
  inference(subsumption_resolution,[],[f7785,f431])).
fof(f431,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f263])).
fof(f7785,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | p(s(bool,'F_0')) | ~sP8(X14,X15,'F_0') | ~sP7('F_0',X15,X14) | ~sP6(X15,'F_0',X14)) )),
  inference(resolution,[],[f4146,f422])).
fof(f4146,plain,(
  ( ! [X8,X9] : (~sP9(X9,'F_0',X8) | ~p(s(bool,X9)) | s(bool,'T_0') != s(bool,X8)) )),
  inference(superposition,[],[f4139,f438])).
fof(f7822,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK34(sK13,'type.num.num')))) | $spl119),
  inference(cnf_transformation,[],[f7822_D])).
fof(f7822_D,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK34(sK13,'type.num.num')))) <=> ~$spl119),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl119])])).
fof(f7804,plain,(
  ( ! [X13] : (s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK34(sK13,'type.num.num')))) | ~p(s(bool,X13))) )),
  inference(resolution,[],[f7794,f2612])).
fof(f7820,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK20(sK15(sK13))))) | $spl117),
  inference(cnf_transformation,[],[f7820_D])).
fof(f7820_D,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK20(sK15(sK13))))) <=> ~$spl117),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl117])])).
fof(f7803,plain,(
  ( ! [X12] : (s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK20(sK15(sK13))))) | ~p(s(bool,X12))) ) | ($spl13 | $spl60)),
  inference(resolution,[],[f7794,f7287])).
fof(f7818,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK20(sK20(sK14))))) | $spl115),
  inference(cnf_transformation,[],[f7818_D])).
fof(f7818_D,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK20(sK20(sK14))))) <=> ~$spl115),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl115])])).
fof(f7802,plain,(
  ( ! [X11] : (s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK20(sK20(sK14))))) | ~p(s(bool,X11))) ) | ($spl16 | $spl30)),
  inference(resolution,[],[f7794,f6447])).
fof(f7816,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK20(sK14)))) | $spl113),
  inference(cnf_transformation,[],[f7816_D])).
fof(f7816_D,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK20(sK14)))) <=> ~$spl113),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl113])])).
fof(f7801,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK20(sK14)))) | ~p(s(bool,X10))) ) | $spl16),
  inference(resolution,[],[f7794,f3368])).
fof(f7814,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15(sK13)))) | $spl111),
  inference(cnf_transformation,[],[f7814_D])).
fof(f7814_D,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15(sK13)))) <=> ~$spl111),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl111])])).
fof(f7800,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15(sK13)))) | ~p(s(bool,X9))) ) | $spl13),
  inference(resolution,[],[f7794,f2030])).
fof(f7586,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | s(bool,'T_0') != s(bool,X2)) ) | $spl108),
  inference(cnf_transformation,[],[f7586_D])).
fof(f7586_D,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | s(bool,'T_0') != s(bool,X2)) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f7585,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) )),
  inference(duplicate_literal_removal,[],[f7580])).
fof(f7580,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) )),
  inference(resolution,[],[f4143,f439])).
fof(f4143,plain,(
  ( ! [X2,X3] : (sP9(X3,X2,'F_0') | s(bool,'T_0') != s(bool,X2) | p(s(bool,X3))) )),
  inference(superposition,[],[f4139,f436])).
fof(f6756,plain,(
  s('type.num.num',sK20(sK20(sK15(sK13)))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK20(sK20(sK15(sK13))))))) | $spl100),
  inference(cnf_transformation,[],[f6756_D])).
fof(f6756_D,plain,(
  s('type.num.num',sK20(sK20(sK15(sK13)))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK20(sK20(sK15(sK13))))))) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f6758,plain,(
  s('type.num.num',sK20(sK15(sK13))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | $spl102),
  inference(cnf_transformation,[],[f6758_D])).
fof(f6758_D,plain,(
  s('type.num.num',sK20(sK15(sK13))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f7519,plain,(
  s('type.num.num',sK20(sK15(sK13))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK20(sK20(sK15(sK13)))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK20(sK20(sK15(sK13))))))) | $spl60),
  inference(superposition,[],[f3014,f365])).
fof(f753,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | $spl14),
  inference(cnf_transformation,[],[f753_D])).
fof(f753_D,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f7092,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(subsumption_resolution,[],[f7086,f546])).
fof(f546,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f327])).
fof(f327,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f326])).
fof(f326,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f128])).
fof(f128,plain,(
  ! [X0] : ((p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f65])).
fof(f65,axiom,(
  ! [X12] : ((p(s(bool,'F_0')) | p(s(bool,X12))) <=> p(s(bool,X12)))),
  file('thm.rat.rat_of_num_ind','thm.bool.OR_CLAUSES|split|2')).
fof(f7086,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(resolution,[],[f1029,f488])).
fof(f488,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP12(X0,X1)) )),
  inference(cnf_transformation,[],[f290])).
fof(f290,plain,(
  ! [X0,X1] : ((~sP12(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP12(X0,X1)))),
  inference(flattening,[],[f289])).
fof(f289,plain,(
  ! [X0,X1] : ((~sP12(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP12(X0,X1)))),
  inference(nnf_transformation,[],[f217])).
fof(f217,plain,(
  ! [X0,X1] : (sP12(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP12])])).
fof(f1029,plain,(
  ( ! [X152] : (sP12(X152,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(resolution,[],[f353,f489])).
fof(f489,plain,(
  ( ! [X0,X1] : (sP12(X0,X1) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f290])).
fof(f7091,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(subsumption_resolution,[],[f7085,f546])).
fof(f7085,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(resolution,[],[f1029,f492])).
fof(f492,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP12(X0,X1)) )),
  inference(cnf_transformation,[],[f292])).
fof(f292,plain,(
  ! [X0,X1] : ((~sP12(X0,X1) | (p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP12(X0,X1)))),
  inference(flattening,[],[f291])).
fof(f291,plain,(
  ! [X0,X1] : ((~sP12(X0,X1) | ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0')))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP12(X0,X1)))),
  inference(nnf_transformation,[],[f218])).
fof(f218,plain,(
  ! [X0,X1] : (sP12(X0,X1) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(definition_folding,[],[f173,f217])).
fof(f173,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f172])).
fof(f172,plain,(
  ! [X0,X1] : (((p(s(bool,X1)) | p(s(bool,X0))) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | (p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(ennf_transformation,[],[f98])).
fof(f98,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(flattening,[],[f97])).
fof(f97,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X17,X13] : ((~(p(s(bool,X13)) | p(s(bool,X17))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X13)) => p(s(bool,'F_0'))) => (~p(s(bool,X17)) => p(s(bool,'F_0')))))),
  file('thm.rat.rat_of_num_ind','thm.sat.sat_9_prove_ltb')).
fof(f7030,plain,(
  $false | ($spl13 | $spl62)),
  inference(subsumption_resolution,[],[f6996,f751])).
fof(f6996,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15(sK13))))) | $spl62),
  inference(backward_demodulation,[],[f3016,f353])).
fof(f6852,plain,(
  $false | ($spl13 | $spl104)),
  inference(subsumption_resolution,[],[f6806,f751])).
fof(f6806,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15(sK13))))) | ($spl13 | $spl104)),
  inference(backward_demodulation,[],[f6775,f353])).
fof(f6775,plain,(
  s('type.num.num',sK15(sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ($spl13 | $spl104)),
  inference(backward_demodulation,[],[f6761,f2102])).
fof(f6761,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK20(sK15(sK13))) | $spl104),
  inference(cnf_transformation,[],[f6761_D])).
fof(f6761_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK20(sK15(sK13))) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f6763,plain,(
  s('type.num.num',sK20(sK15(sK13))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.num.SUC_1'(s('type.num.num',sK20(sK20(sK15(sK13))))))))) | $spl106),
  inference(cnf_transformation,[],[f6763_D])).
fof(f6763_D,plain,(
  s('type.num.num',sK20(sK15(sK13))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.num.SUC_1'(s('type.num.num',sK20(sK20(sK15(sK13))))))))) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f6773,plain,(
  s('type.num.num',sK20(sK15(sK13))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.num.SUC_1'(s('type.num.num',sK20(sK20(sK15(sK13))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK20(sK15(sK13))) | $spl60),
  inference(forward_demodulation,[],[f6739,f3014])).
fof(f6739,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK20(sK15(sK13))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK20(sK15(sK13)))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.num.SUC_1'(s('type.num.num',sK20(sK20(sK15(sK13))))))))) | $spl60),
  inference(superposition,[],[f365,f3014])).
fof(f6759,plain,(
  s('type.num.num',sK20(sK15(sK13))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.num.SUC_1'(s('type.num.num',sK20(sK20(sK15(sK13))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK20(sK15(sK13))) | $spl60),
  inference(forward_demodulation,[],[f6726,f3014])).
fof(f6726,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK20(sK15(sK13))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK20(sK15(sK13)))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.num.SUC_1'(s('type.num.num',sK20(sK20(sK15(sK13))))))))) | $spl60),
  inference(superposition,[],[f3014,f365])).
fof(f6725,plain,(
  s('type.num.num',sK20(sK15(sK13))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK20(sK20(sK15(sK13)))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK20(sK20(sK15(sK13))))))) | $spl60),
  inference(superposition,[],[f3014,f365])).
fof(f6026,plain,(
  s('type.num.num',sK20(sK20(sK20(sK14)))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK20(sK20(sK20(sK14))))))) | $spl92),
  inference(cnf_transformation,[],[f6026_D])).
fof(f6026_D,plain,(
  s('type.num.num',sK20(sK20(sK20(sK14)))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK20(sK20(sK20(sK14))))))) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f6028,plain,(
  s('type.num.num',sK20(sK20(sK14))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | $spl94),
  inference(cnf_transformation,[],[f6028_D])).
fof(f6028_D,plain,(
  s('type.num.num',sK20(sK20(sK14))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f6679,plain,(
  s('type.num.num',sK20(sK20(sK14))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK20(sK20(sK20(sK14)))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK20(sK20(sK20(sK14))))))) | $spl30),
  inference(superposition,[],[f2136,f365])).
fof(f6260,plain,(
  ( ! [X1] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X1))) )),
  inference(subsumption_resolution,[],[f6254,f546])).
fof(f6254,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X1))) )),
  inference(resolution,[],[f1030,f488])).
fof(f1030,plain,(
  ( ! [X153] : (sP12(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X153)) )),
  inference(resolution,[],[f353,f490])).
fof(f490,plain,(
  ( ! [X0,X1] : (sP12(X0,X1) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f290])).
fof(f6259,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f6253,f546])).
fof(f6253,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X0))) )),
  inference(resolution,[],[f1030,f492])).
fof(f6252,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(subsumption_resolution,[],[f6246,f546])).
fof(f6246,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(resolution,[],[f1029,f488])).
fof(f6251,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(subsumption_resolution,[],[f6245,f546])).
fof(f6245,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(resolution,[],[f1029,f492])).
fof(f6192,plain,(
  $false | ($spl16 | $spl96)),
  inference(subsumption_resolution,[],[f6191,f2146])).
fof(f6191,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK20(sK14))))) | ($spl16 | $spl96)),
  inference(forward_demodulation,[],[f353,f6110])).
fof(f6110,plain,(
  s('type.num.num',sK20(sK14)) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ($spl16 | $spl96)),
  inference(backward_demodulation,[],[f6031,f933])).
fof(f6092,plain,(
  $false | ($spl16 | $spl32)),
  inference(subsumption_resolution,[],[f6047,f2146])).
fof(f6047,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK20(sK14))))) | $spl32),
  inference(backward_demodulation,[],[f2138,f353])).
fof(f6031,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK20(sK20(sK14))) | $spl96),
  inference(cnf_transformation,[],[f6031_D])).
fof(f6031_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK20(sK20(sK14))) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f6033,plain,(
  s('type.num.num',sK20(sK20(sK14))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.num.SUC_1'(s('type.num.num',sK20(sK20(sK20(sK14))))))))) | $spl98),
  inference(cnf_transformation,[],[f6033_D])).
fof(f6033_D,plain,(
  s('type.num.num',sK20(sK20(sK14))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.num.SUC_1'(s('type.num.num',sK20(sK20(sK20(sK14))))))))) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f6046,plain,(
  s('type.num.num',sK20(sK20(sK14))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.num.SUC_1'(s('type.num.num',sK20(sK20(sK20(sK14))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK20(sK20(sK14))) | $spl30),
  inference(forward_demodulation,[],[f6009,f2136])).
fof(f6009,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK20(sK20(sK14))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK20(sK20(sK14)))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.num.SUC_1'(s('type.num.num',sK20(sK20(sK20(sK14))))))))) | $spl30),
  inference(superposition,[],[f365,f2136])).
fof(f6029,plain,(
  s('type.num.num',sK20(sK20(sK14))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.num.SUC_1'(s('type.num.num',sK20(sK20(sK20(sK14))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK20(sK20(sK14))) | $spl30),
  inference(forward_demodulation,[],[f5996,f2136])).
fof(f5996,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK20(sK20(sK14))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK20(sK20(sK14)))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.num.SUC_1'(s('type.num.num',sK20(sK20(sK20(sK14))))))))) | $spl30),
  inference(superposition,[],[f2136,f365])).
fof(f5995,plain,(
  s('type.num.num',sK20(sK20(sK14))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK20(sK20(sK20(sK14)))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK20(sK20(sK20(sK14))))))) | $spl30),
  inference(superposition,[],[f2136,f365])).
fof(f5608,plain,(
  s('type.num.num',sK20(sK35(sK13,'type.num.num'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK20(sK35(sK13,'type.num.num')))))) | $spl84),
  inference(cnf_transformation,[],[f5608_D])).
fof(f5608_D,plain,(
  s('type.num.num',sK20(sK35(sK13,'type.num.num'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK20(sK35(sK13,'type.num.num')))))) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f5610,plain,(
  s('type.num.num',sK35(sK13,'type.num.num')) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | $spl86),
  inference(cnf_transformation,[],[f5610_D])).
fof(f5610_D,plain,(
  s('type.num.num',sK35(sK13,'type.num.num')) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f5940,plain,(
  s('type.num.num',sK35(sK13,'type.num.num')) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK20(sK35(sK13,'type.num.num'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK20(sK35(sK13,'type.num.num'))))))),
  inference(superposition,[],[f2843,f365])).
fof(f2843,plain,(
  s('type.num.num',sK35(sK13,'type.num.num')) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK35(sK13,'type.num.num')))))),
  inference(subsumption_resolution,[],[f2821,f352])).
fof(f2821,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | s('type.num.num',sK35(sK13,'type.num.num')) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK35(sK13,'type.num.num')))))),
  inference(superposition,[],[f763,f365])).
fof(f3625,plain,(
  s('type.num.num',sK35(sK13,'type.num.num')) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK35(sK13,'type.num.num'))))) | $spl74),
  inference(cnf_transformation,[],[f3625_D])).
fof(f3625_D,plain,(
  s('type.num.num',sK35(sK13,'type.num.num')) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK35(sK13,'type.num.num'))))) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f2982,plain,(
  ( ! [X0,X1] : (sP7(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X0,X1)) ) | $spl58),
  inference(cnf_transformation,[],[f2982_D])).
fof(f2982_D,plain,(
  ( ! [X0,X1] : (sP7(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X0,X1)) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f5907,plain,(
  ( ! [X0,X1] : (sP7(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X0,X1) | s('type.num.num',sK35(sK13,'type.num.num')) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK35(sK13,'type.num.num')))))) )),
  inference(superposition,[],[f2775,f365])).
fof(f2775,plain,(
  ( ! [X109,X108] : (sP7(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK35(sK13,'type.num.num'))),X108,X109)) )),
  inference(resolution,[],[f763,f431])).
fof(f2973,plain,(
  ( ! [X0,X1] : (sP7(X0,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X1)) ) | $spl56),
  inference(cnf_transformation,[],[f2973_D])).
fof(f2973_D,plain,(
  ( ! [X0,X1] : (sP7(X0,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X1)) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f5899,plain,(
  ( ! [X0,X1] : (sP7(X0,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X1) | s('type.num.num',sK35(sK13,'type.num.num')) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK35(sK13,'type.num.num')))))) )),
  inference(superposition,[],[f2774,f365])).
fof(f2774,plain,(
  ( ! [X107,X106] : (sP7(X106,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK35(sK13,'type.num.num'))),X107)) )),
  inference(resolution,[],[f763,f430])).
fof(f2964,plain,(
  ( ! [X0,X1] : (sP8(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X0,X1)) ) | $spl54),
  inference(cnf_transformation,[],[f2964_D])).
fof(f2964_D,plain,(
  ( ! [X0,X1] : (sP8(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X0,X1)) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f5879,plain,(
  ( ! [X0,X1] : (sP8(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X0,X1) | s('type.num.num',sK35(sK13,'type.num.num')) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK35(sK13,'type.num.num')))))) )),
  inference(superposition,[],[f2772,f365])).
fof(f2772,plain,(
  ( ! [X103,X102] : (sP8(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK35(sK13,'type.num.num'))),X102,X103)) )),
  inference(resolution,[],[f763,f427])).
fof(f2955,plain,(
  ( ! [X0,X1] : (sP8(X0,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X1)) ) | $spl52),
  inference(cnf_transformation,[],[f2955_D])).
fof(f2955_D,plain,(
  ( ! [X0,X1] : (sP8(X0,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X1)) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f5872,plain,(
  ( ! [X0,X1] : (sP8(X0,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X1) | s('type.num.num',sK35(sK13,'type.num.num')) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK35(sK13,'type.num.num')))))) )),
  inference(superposition,[],[f2771,f365])).
fof(f2947,plain,(
  ( ! [X0,X1] : (sP4(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X0,X1)) ) | $spl50),
  inference(cnf_transformation,[],[f2947_D])).
fof(f2947_D,plain,(
  ( ! [X0,X1] : (sP4(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X0,X1)) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f5866,plain,(
  ( ! [X0,X1] : (sP4(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X0,X1) | s('type.num.num',sK35(sK13,'type.num.num')) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK35(sK13,'type.num.num')))))) )),
  inference(superposition,[],[f2759,f365])).
fof(f2759,plain,(
  ( ! [X76,X77] : (sP4(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK35(sK13,'type.num.num'))),X76,X77)) )),
  inference(resolution,[],[f763,f410])).
fof(f410,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f252])).
fof(f252,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(flattening,[],[f251])).
fof(f251,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(nnf_transformation,[],[f205])).
fof(f205,plain,(
  ! [X0,X1,X2] : (sP4(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f2939,plain,(
  ( ! [X0,X1] : (sP4(X0,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X1)) ) | $spl48),
  inference(cnf_transformation,[],[f2939_D])).
fof(f2939_D,plain,(
  ( ! [X0,X1] : (sP4(X0,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X1)) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f5859,plain,(
  ( ! [X0,X1] : (sP4(X0,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X1) | s('type.num.num',sK35(sK13,'type.num.num')) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK35(sK13,'type.num.num')))))) )),
  inference(superposition,[],[f2758,f365])).
fof(f2758,plain,(
  ( ! [X74,X75] : (sP4(X74,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK35(sK13,'type.num.num'))),X75)) )),
  inference(resolution,[],[f763,f409])).
fof(f409,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f252])).
fof(f2931,plain,(
  ( ! [X0,X1] : (sP2(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X0,X1)) ) | $spl46),
  inference(cnf_transformation,[],[f2931_D])).
fof(f2931_D,plain,(
  ( ! [X0,X1] : (sP2(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X0,X1)) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f5851,plain,(
  ( ! [X0,X1] : (sP2(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X0,X1) | s('type.num.num',sK35(sK13,'type.num.num')) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK35(sK13,'type.num.num')))))) )),
  inference(superposition,[],[f2745,f365])).
fof(f2745,plain,(
  ( ! [X48,X49] : (sP2(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK35(sK13,'type.num.num'))),X48,X49)) )),
  inference(resolution,[],[f763,f393])).
fof(f393,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f245])).
fof(f245,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(rectify,[],[f244])).
fof(f244,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(flattening,[],[f243])).
fof(f243,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(nnf_transformation,[],[f202])).
fof(f202,plain,(
  ! [X2,X0,X1] : (sP2(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f2915,plain,(
  ( ! [X0,X1] : (sP0(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X0,X1)) ) | $spl44),
  inference(cnf_transformation,[],[f2915_D])).
fof(f2915_D,plain,(
  ( ! [X0,X1] : (sP0(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X0,X1)) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f5830,plain,(
  ( ! [X0,X1] : (sP0(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X0,X1) | s('type.num.num',sK35(sK13,'type.num.num')) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK35(sK13,'type.num.num')))))) )),
  inference(superposition,[],[f2729,f365])).
fof(f2729,plain,(
  ( ! [X17,X16] : (sP0(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK35(sK13,'type.num.num'))),X16,X17)) )),
  inference(resolution,[],[f763,f376])).
fof(f376,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f237])).
fof(f237,plain,(
  ! [X0,X1,X2] : ((~sP0(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP0(X0,X1,X2)))),
  inference(rectify,[],[f236])).
fof(f236,plain,(
  ! [X2,X0,X1] : ((~sP0(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP0(X2,X0,X1)))),
  inference(flattening,[],[f235])).
fof(f235,plain,(
  ! [X2,X0,X1] : ((~sP0(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP0(X2,X0,X1)))),
  inference(nnf_transformation,[],[f199])).
fof(f199,plain,(
  ! [X2,X0,X1] : (sP0(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f2907,plain,(
  ( ! [X0,X1] : (sP0(X0,X1,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) ) | $spl42),
  inference(cnf_transformation,[],[f2907_D])).
fof(f2907_D,plain,(
  ( ! [X0,X1] : (sP0(X0,X1,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f5823,plain,(
  ( ! [X0,X1] : (sP0(X0,X1,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK35(sK13,'type.num.num')) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK35(sK13,'type.num.num')))))) )),
  inference(superposition,[],[f2728,f365])).
fof(f2728,plain,(
  ( ! [X14,X15] : (sP0(X14,X15,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK35(sK13,'type.num.num'))))) )),
  inference(resolution,[],[f763,f374])).
fof(f374,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f237])).
fof(f5672,plain,(
  $false | $spl88),
  inference(subsumption_resolution,[],[f5671,f601])).
fof(f601,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.rat.rat_of_num_ind',aHL_TRUTH)).
fof(f5671,plain,(
  ~p(s(bool,'T_0')) | $spl88),
  inference(forward_demodulation,[],[f5635,f717])).
fof(f717,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))),
  inference(resolution,[],[f352,f581])).
fof(f5635,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | $spl88),
  inference(backward_demodulation,[],[f5613,f763])).
fof(f5613,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK35(sK13,'type.num.num')) | $spl88),
  inference(cnf_transformation,[],[f5613_D])).
fof(f5613_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK35(sK13,'type.num.num')) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f5615,plain,(
  s('type.num.num',sK35(sK13,'type.num.num')) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.num.SUC_1'(s('type.num.num',sK20(sK35(sK13,'type.num.num')))))))) | $spl90),
  inference(cnf_transformation,[],[f5615_D])).
fof(f5615_D,plain,(
  s('type.num.num',sK35(sK13,'type.num.num')) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.num.SUC_1'(s('type.num.num',sK20(sK35(sK13,'type.num.num')))))))) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f5620,plain,(
  s('type.num.num',sK35(sK13,'type.num.num')) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.num.SUC_1'(s('type.num.num',sK20(sK35(sK13,'type.num.num')))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK35(sK13,'type.num.num'))),
  inference(forward_demodulation,[],[f5591,f2843])).
fof(f5591,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK35(sK13,'type.num.num')) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK35(sK13,'type.num.num'))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.num.SUC_1'(s('type.num.num',sK20(sK35(sK13,'type.num.num'))))))))),
  inference(superposition,[],[f365,f2843])).
fof(f5611,plain,(
  s('type.num.num',sK35(sK13,'type.num.num')) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.num.SUC_1'(s('type.num.num',sK20(sK35(sK13,'type.num.num')))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK35(sK13,'type.num.num'))),
  inference(forward_demodulation,[],[f5578,f2843])).
fof(f5578,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK35(sK13,'type.num.num')) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK35(sK13,'type.num.num'))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.num.SUC_1'(s('type.num.num',sK20(sK35(sK13,'type.num.num'))))))))),
  inference(superposition,[],[f2843,f365])).
fof(f5577,plain,(
  s('type.num.num',sK35(sK13,'type.num.num')) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK20(sK35(sK13,'type.num.num'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK20(sK35(sK13,'type.num.num'))))))),
  inference(superposition,[],[f2843,f365])).
fof(f5234,plain,(
  s('type.num.num',sK20(sK34(sK13,'type.num.num'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK20(sK34(sK13,'type.num.num')))))) | $spl76),
  inference(cnf_transformation,[],[f5234_D])).
fof(f5234_D,plain,(
  s('type.num.num',sK20(sK34(sK13,'type.num.num'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK20(sK34(sK13,'type.num.num')))))) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f5236,plain,(
  s('type.num.num',sK34(sK13,'type.num.num')) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | $spl78),
  inference(cnf_transformation,[],[f5236_D])).
fof(f5236_D,plain,(
  s('type.num.num',sK34(sK13,'type.num.num')) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f5539,plain,(
  s('type.num.num',sK34(sK13,'type.num.num')) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK20(sK34(sK13,'type.num.num'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK20(sK34(sK13,'type.num.num'))))))),
  inference(superposition,[],[f2684,f365])).
fof(f2684,plain,(
  s('type.num.num',sK34(sK13,'type.num.num')) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK34(sK13,'type.num.num')))))),
  inference(subsumption_resolution,[],[f2662,f352])).
fof(f2662,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | s('type.num.num',sK34(sK13,'type.num.num')) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK34(sK13,'type.num.num')))))),
  inference(superposition,[],[f762,f365])).
fof(f3525,plain,(
  s('type.num.num',sK34(sK13,'type.num.num')) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK34(sK13,'type.num.num'))))) | $spl68),
  inference(cnf_transformation,[],[f3525_D])).
fof(f3525_D,plain,(
  s('type.num.num',sK34(sK13,'type.num.num')) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK34(sK13,'type.num.num'))))) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f5520,plain,(
  ( ! [X0,X1] : (sP7(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X0,X1) | s('type.num.num',sK34(sK13,'type.num.num')) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK34(sK13,'type.num.num')))))) )),
  inference(superposition,[],[f2616,f365])).
fof(f2616,plain,(
  ( ! [X109,X108] : (sP7(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK34(sK13,'type.num.num'))),X108,X109)) )),
  inference(resolution,[],[f762,f431])).
fof(f5512,plain,(
  ( ! [X0,X1] : (sP7(X0,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X1) | s('type.num.num',sK34(sK13,'type.num.num')) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK34(sK13,'type.num.num')))))) )),
  inference(superposition,[],[f2615,f365])).
fof(f2615,plain,(
  ( ! [X107,X106] : (sP7(X106,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK34(sK13,'type.num.num'))),X107)) )),
  inference(resolution,[],[f762,f430])).
fof(f5504,plain,(
  ( ! [X0,X1] : (sP8(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X0,X1) | s('type.num.num',sK34(sK13,'type.num.num')) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK34(sK13,'type.num.num')))))) )),
  inference(superposition,[],[f2613,f365])).
fof(f2613,plain,(
  ( ! [X103,X102] : (sP8(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK34(sK13,'type.num.num'))),X102,X103)) )),
  inference(resolution,[],[f762,f427])).
fof(f5497,plain,(
  ( ! [X0,X1] : (sP8(X0,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X1) | s('type.num.num',sK34(sK13,'type.num.num')) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK34(sK13,'type.num.num')))))) )),
  inference(superposition,[],[f2612,f365])).
fof(f5480,plain,(
  ( ! [X0,X1] : (sP4(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X0,X1) | s('type.num.num',sK34(sK13,'type.num.num')) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK34(sK13,'type.num.num')))))) )),
  inference(superposition,[],[f2600,f365])).
fof(f2600,plain,(
  ( ! [X76,X77] : (sP4(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK34(sK13,'type.num.num'))),X76,X77)) )),
  inference(resolution,[],[f762,f410])).
fof(f5473,plain,(
  ( ! [X0,X1] : (sP4(X0,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X1) | s('type.num.num',sK34(sK13,'type.num.num')) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK34(sK13,'type.num.num')))))) )),
  inference(superposition,[],[f2599,f365])).
fof(f2599,plain,(
  ( ! [X74,X75] : (sP4(X74,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK34(sK13,'type.num.num'))),X75)) )),
  inference(resolution,[],[f762,f409])).
fof(f5465,plain,(
  ( ! [X0,X1] : (sP2(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X0,X1) | s('type.num.num',sK34(sK13,'type.num.num')) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK34(sK13,'type.num.num')))))) )),
  inference(superposition,[],[f2586,f365])).
fof(f2586,plain,(
  ( ! [X48,X49] : (sP2(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK34(sK13,'type.num.num'))),X48,X49)) )),
  inference(resolution,[],[f762,f393])).
fof(f5458,plain,(
  ( ! [X0,X1] : (sP0(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X0,X1) | s('type.num.num',sK34(sK13,'type.num.num')) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK34(sK13,'type.num.num')))))) )),
  inference(superposition,[],[f2570,f365])).
fof(f2570,plain,(
  ( ! [X17,X16] : (sP0(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK34(sK13,'type.num.num'))),X16,X17)) )),
  inference(resolution,[],[f762,f376])).
fof(f5451,plain,(
  ( ! [X0,X1] : (sP0(X0,X1,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK34(sK13,'type.num.num')) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK34(sK13,'type.num.num')))))) )),
  inference(superposition,[],[f2569,f365])).
fof(f2569,plain,(
  ( ! [X14,X15] : (sP0(X14,X15,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK34(sK13,'type.num.num'))))) )),
  inference(resolution,[],[f762,f374])).
fof(f5284,plain,(
  $false | $spl80),
  inference(subsumption_resolution,[],[f5283,f601])).
fof(f5283,plain,(
  ~p(s(bool,'T_0')) | $spl80),
  inference(forward_demodulation,[],[f5247,f717])).
fof(f5247,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | $spl80),
  inference(backward_demodulation,[],[f5239,f762])).
fof(f5239,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK34(sK13,'type.num.num')) | $spl80),
  inference(cnf_transformation,[],[f5239_D])).
fof(f5239_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK34(sK13,'type.num.num')) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f5241,plain,(
  s('type.num.num',sK34(sK13,'type.num.num')) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.num.SUC_1'(s('type.num.num',sK20(sK34(sK13,'type.num.num')))))))) | $spl82),
  inference(cnf_transformation,[],[f5241_D])).
fof(f5241_D,plain,(
  s('type.num.num',sK34(sK13,'type.num.num')) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.num.SUC_1'(s('type.num.num',sK20(sK34(sK13,'type.num.num')))))))) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f5246,plain,(
  s('type.num.num',sK34(sK13,'type.num.num')) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.num.SUC_1'(s('type.num.num',sK20(sK34(sK13,'type.num.num')))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK34(sK13,'type.num.num'))),
  inference(forward_demodulation,[],[f5217,f2684])).
fof(f5217,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK34(sK13,'type.num.num')) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK34(sK13,'type.num.num'))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.num.SUC_1'(s('type.num.num',sK20(sK34(sK13,'type.num.num'))))))))),
  inference(superposition,[],[f365,f2684])).
fof(f5237,plain,(
  s('type.num.num',sK34(sK13,'type.num.num')) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.num.SUC_1'(s('type.num.num',sK20(sK34(sK13,'type.num.num')))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK34(sK13,'type.num.num'))),
  inference(forward_demodulation,[],[f5204,f2684])).
fof(f5204,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK34(sK13,'type.num.num')) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK34(sK13,'type.num.num'))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.num.SUC_1'(s('type.num.num',sK20(sK34(sK13,'type.num.num'))))))))),
  inference(superposition,[],[f2684,f365])).
fof(f5203,plain,(
  s('type.num.num',sK34(sK13,'type.num.num')) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK20(sK34(sK13,'type.num.num'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK20(sK34(sK13,'type.num.num'))))))),
  inference(superposition,[],[f2684,f365])).
fof(f743,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK19(sK13))))))) | $spl5),
  inference(cnf_transformation,[],[f743_D])).
fof(f743_D,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK19(sK13))))))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f741,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X0))))) ) | $spl2),
  inference(cnf_transformation,[],[f741_D])).
fof(f741_D,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X0))))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f4442,plain,(
  ( ! [X4] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X4)))) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK19(sK13)))))))) )),
  inference(subsumption_resolution,[],[f4177,f606])).
fof(f606,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f527])).
fof(f527,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f314])).
fof(f314,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f118])).
fof(f118,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f61])).
fof(f61,axiom,(
  ! [X13] : ! [X10] : (s(X13,X10) = s(X13,X10) <=> p(s(bool,'T_0')))),
  file('thm.rat.rat_of_num_ind','thm.bool.REFL_CLAUSE')).
fof(f4177,plain,(
  ( ! [X4] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X4)))) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK19(sK13)))))))) )),
  inference(superposition,[],[f364,f717])).
fof(f364,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X1)))) | ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK19(X0))))))) | ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.|48|_0'))))) )),
  inference(cnf_transformation,[],[f230])).
fof(f230,plain,(
  ! [X0] : (~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.|48|_0')))) | (p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',sK19(X0))))) & ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK19(X0)))))))) | ! [X1] : p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK19])],[f229])).
fof(f229,plain,(
  ! [X0] : (~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.|48|_0')))) | ? [X2] : (p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2)))) & ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))))))) | ! [X1] : p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X1)))))),
  inference(rectify,[],[f164])).
fof(f164,plain,(
  ! [X0] : (~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.|48|_0')))) | ? [X1] : (p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X1)))) & ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))))) | ! [X2] : p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2)))))),
  inference(flattening,[],[f163])).
fof(f163,plain,(
  ! [X0] : ((~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.|48|_0')))) | ? [X1] : (p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X1)))) & ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))))) | ! [X2] : p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2)))))),
  inference(ennf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0] : ((p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.|48|_0')))) & ! [X1] : (p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X1)))) => p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))))) => ! [X2] : p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2)))))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X6] : ((p(s(bool,happ(s(fun('type.num.num',bool),X6),s('type.num.num','const.num.|48|_0')))) & ! [X7] : (p(s(bool,happ(s(fun('type.num.num',bool),X6),s('type.num.num',X7)))) => p(s(bool,happ(s(fun('type.num.num',bool),X6),s('type.num.num','const.num.SUC_1'(s('type.num.num',X7)))))))) => ! [X7] : p(s(bool,happ(s(fun('type.num.num',bool),X6),s('type.num.num',X7)))))),
  file('thm.rat.rat_of_num_ind','thm.num.INDUCTION')).
fof(f745,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK18(sK13))))))) | $spl7),
  inference(cnf_transformation,[],[f745_D])).
fof(f745_D,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK18(sK13))))))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f4441,plain,(
  ( ! [X2] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X2)))) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK18(sK13)))))))) )),
  inference(subsumption_resolution,[],[f4175,f606])).
fof(f4175,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X2)))) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK18(sK13)))))))) )),
  inference(superposition,[],[f362,f717])).
fof(f362,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X1)))) | ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK18(X0))))))) | ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.|48|_0'))))) )),
  inference(cnf_transformation,[],[f228])).
fof(f228,plain,(
  ! [X0] : ((~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',sK17(X0))))) | (p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.|48|_0')))) & ! [X3] : p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))))) & (~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK18(X0))))))) | ! [X1] : p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK17,sK18])],[f227])).
fof(f227,plain,(
  ! [X0] : ((? [X4] : ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X4)))) | (p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.|48|_0')))) & ! [X3] : p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))))) & (~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.|48|_0')))) | ? [X2] : ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2)))))) | ! [X1] : p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X1))))))),
  inference(rectify,[],[f226])).
fof(f226,plain,(
  ! [X0] : ((? [X1] : ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X1)))) | (p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.|48|_0')))) & ! [X2] : p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2)))))))) & (~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.|48|_0')))) | ? [X2] : ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2)))))) | ! [X1] : p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X1))))))),
  inference(flattening,[],[f225])).
fof(f225,plain,(
  ! [X0] : ((? [X1] : ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X1)))) | (p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.|48|_0')))) & ! [X2] : p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2)))))))) & ((~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.|48|_0')))) | ? [X2] : ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))))))) | ! [X1] : p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X1))))))),
  inference(nnf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0] : (! [X1] : p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X1)))) <=> (p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.|48|_0')))) & ! [X2] : p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))))))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6] : (! [X7] : p(s(bool,happ(s(fun('type.num.num',bool),X6),s('type.num.num',X7)))) <=> (p(s(bool,happ(s(fun('type.num.num',bool),X6),s('type.num.num','const.num.|48|_0')))) & ! [X7] : p(s(bool,happ(s(fun('type.num.num',bool),X6),s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))))))))),
  file('thm.rat.rat_of_num_ind','thm.arithmetic.FORALL_NUM')).
fof(f749,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16(sK13))))))) | $spl11),
  inference(cnf_transformation,[],[f749_D])).
fof(f749_D,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16(sK13))))))) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f4440,plain,(
  ( ! [X1] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X1)))) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16(sK13)))))))) )),
  inference(subsumption_resolution,[],[f4173,f606])).
fof(f4173,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X1)))) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16(sK13)))))))) )),
  inference(superposition,[],[f359,f717])).
fof(f359,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X1)))) | ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16(X0))))))) | ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.|48|_0'))))) )),
  inference(cnf_transformation,[],[f224])).
fof(f224,plain,(
  ! [X0] : ((~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',sK15(X0))))) | (p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.|48|_0')))) & ! [X3] : (~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X3)))) | p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))))))))) & (~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.|48|_0')))) | (p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',sK16(X0))))) & ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16(X0)))))))) | ! [X1] : p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK15,sK16])],[f223])).
fof(f223,plain,(
  ! [X0] : ((? [X4] : ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X4)))) | (p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.|48|_0')))) & ! [X3] : (~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X3)))) | p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))))))))) & (~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.|48|_0')))) | ? [X2] : (p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2)))) & ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))))))) | ! [X1] : p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X1))))))),
  inference(rectify,[],[f222])).
fof(f222,plain,(
  ! [X0] : ((? [X1] : ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X1)))) | (p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.|48|_0')))) & ! [X2] : (~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2)))) | p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))))))))) & (~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.|48|_0')))) | ? [X2] : (p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2)))) & ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))))))) | ! [X1] : p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X1))))))),
  inference(flattening,[],[f221])).
fof(f221,plain,(
  ! [X0] : ((? [X1] : ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X1)))) | (p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.|48|_0')))) & ! [X2] : (~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2)))) | p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))))))))) & ((~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.|48|_0')))) | ? [X2] : (p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2)))) & ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2)))))))) | ! [X1] : p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X1))))))),
  inference(nnf_transformation,[],[f162])).
fof(f162,plain,(
  ! [X0] : (! [X1] : p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X1)))) <=> (p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.|48|_0')))) & ! [X2] : (~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2)))) | p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2)))))))))),
  inference(ennf_transformation,[],[f72])).
fof(f72,plain,(
  ! [X0] : (! [X1] : p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X1)))) <=> (p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.|48|_0')))) & ! [X2] : (p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2)))) => p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2)))))))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : (! [X7] : p(s(bool,happ(s(fun('type.num.num',bool),X6),s('type.num.num',X7)))) <=> (p(s(bool,happ(s(fun('type.num.num',bool),X6),s('type.num.num','const.num.|48|_0')))) & ! [X7] : (p(s(bool,happ(s(fun('type.num.num',bool),X6),s('type.num.num',X7)))) => p(s(bool,happ(s(fun('type.num.num',bool),X6),s('type.num.num','const.num.SUC_1'(s('type.num.num',X7)))))))))),
  file('thm.rat.rat_of_num_ind','thm.arithmetic.FORALL_NUM_THM')).
fof(f4119,plain,(
  ( ! [X0,X1] : (sP7(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X0,X1) | s('type.num.num',sK35(sK13,'type.num.num')) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK35(sK13,'type.num.num')))))) )),
  inference(superposition,[],[f2775,f365])).
fof(f4111,plain,(
  ( ! [X0,X1] : (sP7(X0,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X1) | s('type.num.num',sK35(sK13,'type.num.num')) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK35(sK13,'type.num.num')))))) )),
  inference(superposition,[],[f2774,f365])).
fof(f4103,plain,(
  ( ! [X0,X1] : (sP8(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X0,X1) | s('type.num.num',sK35(sK13,'type.num.num')) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK35(sK13,'type.num.num')))))) )),
  inference(superposition,[],[f2772,f365])).
fof(f4096,plain,(
  ( ! [X0,X1] : (sP8(X0,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X1) | s('type.num.num',sK35(sK13,'type.num.num')) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK35(sK13,'type.num.num')))))) )),
  inference(superposition,[],[f2771,f365])).
fof(f4090,plain,(
  ( ! [X0,X1] : (sP4(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X0,X1) | s('type.num.num',sK35(sK13,'type.num.num')) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK35(sK13,'type.num.num')))))) )),
  inference(superposition,[],[f2759,f365])).
fof(f4083,plain,(
  ( ! [X0,X1] : (sP4(X0,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X1) | s('type.num.num',sK35(sK13,'type.num.num')) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK35(sK13,'type.num.num')))))) )),
  inference(superposition,[],[f2758,f365])).
fof(f4075,plain,(
  ( ! [X0,X1] : (sP2(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X0,X1) | s('type.num.num',sK35(sK13,'type.num.num')) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK35(sK13,'type.num.num')))))) )),
  inference(superposition,[],[f2745,f365])).
fof(f4068,plain,(
  ( ! [X0,X1] : (sP0(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X0,X1) | s('type.num.num',sK35(sK13,'type.num.num')) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK35(sK13,'type.num.num')))))) )),
  inference(superposition,[],[f2729,f365])).
fof(f4058,plain,(
  $false | $spl72),
  inference(subsumption_resolution,[],[f3945,f3535])).
fof(f3945,plain,(
  p(s(bool,'F_0')) | $spl72),
  inference(resolution,[],[f3535,f598])).
fof(f598,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f351])).
fof(f351,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f159])).
fof(f159,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f23])).
fof(f23,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.rat.rat_of_num_ind','thm.bool.NOT_CLAUSES|split|2')).
fof(f4057,plain,(
  $false | $spl72),
  inference(subsumption_resolution,[],[f3944,f3535])).
fof(f3944,plain,(
  p(s(bool,'F_0')) | $spl72),
  inference(resolution,[],[f3535,f596])).
fof(f596,plain,(
  p(s(bool,'F_0')) | p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f350])).
fof(f350,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f158])).
fof(f158,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f11])).
fof(f11,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.rat.rat_of_num_ind','thm.bool.NOT_CLAUSES|split|1')).
fof(f4056,plain,(
  $false | $spl72),
  inference(subsumption_resolution,[],[f3938,f3535])).
fof(f3938,plain,(
  p(s(bool,'F_0')) | $spl72),
  inference(resolution,[],[f3535,f543])).
fof(f543,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f325])).
fof(f325,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f324])).
fof(f324,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f189])).
fof(f189,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f127])).
fof(f127,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X12] : ((p(s(bool,'F_0')) => p(s(bool,X12))) <=> p(s(bool,'T_0')))),
  file('thm.rat.rat_of_num_ind','thm.bool.IMP_CLAUSES|split|2')).
fof(f4055,plain,(
  $false | $spl72),
  inference(subsumption_resolution,[],[f3937,f3535])).
fof(f3937,plain,(
  ( ! [X255] : (p(s(bool,X255))) ) | $spl72),
  inference(resolution,[],[f3535,f540])).
fof(f540,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f323])).
fof(f323,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,'T_0'))))),
  inference(flattening,[],[f322])).
fof(f322,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f188])).
fof(f188,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X12] : ((p(s(bool,X12)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  file('thm.rat.rat_of_num_ind','thm.bool.IMP_CLAUSES|split|1')).
fof(f4054,plain,(
  $false | $spl72),
  inference(subsumption_resolution,[],[f3936,f3535])).
fof(f3936,plain,(
  ( ! [X254] : (p(s(bool,X254))) ) | $spl72),
  inference(resolution,[],[f3535,f537])).
fof(f537,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f321])).
fof(f321,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,'T_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f320])).
fof(f320,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (~p(s(bool,'T_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f187])).
fof(f187,plain,(
  ! [X0] : ((~p(s(bool,'T_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(ennf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0] : ((p(s(bool,'T_0')) => p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X12] : ((p(s(bool,'T_0')) => p(s(bool,X12))) <=> p(s(bool,X12)))),
  file('thm.rat.rat_of_num_ind','thm.bool.IMP_CLAUSES|split|0')).
fof(f4053,plain,(
  $false | $spl72),
  inference(subsumption_resolution,[],[f3934,f3535])).
fof(f3934,plain,(
  ( ! [X252] : (p(s(bool,X252))) ) | $spl72),
  inference(resolution,[],[f3535,f534])).
fof(f534,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f319])).
fof(f319,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))))),
  inference(flattening,[],[f318])).
fof(f318,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f186])).
fof(f186,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f124])).
fof(f124,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X12] : ((p(s(bool,X12)) => p(s(bool,X12))) <=> p(s(bool,'T_0')))),
  file('thm.rat.rat_of_num_ind','thm.bool.IMP_CLAUSES|split|3')).
fof(f3946,plain,(
  $false | $spl72),
  inference(resolution,[],[f3535,f600])).
fof(f600,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f16])).
fof(f16,axiom,(
  p(s(bool,'T_0'))),
  file('thm.rat.rat_of_num_ind','thm.bool.TRUTH')).
fof(f3947,plain,(
  $false | $spl72),
  inference(resolution,[],[f3535,f601])).
fof(f3948,plain,(
  $false | $spl72),
  inference(resolution,[],[f3535,f606])).
fof(f3966,plain,(
  $false | $spl72),
  inference(resolution,[],[f3535,f352])).
fof(f3968,plain,(
  $false | $spl72),
  inference(resolution,[],[f3535,f353])).
fof(f3969,plain,(
  $false | ($spl8 | $spl72)),
  inference(resolution,[],[f3535,f747])).
fof(f3970,plain,(
  $false | ($spl0 | $spl72)),
  inference(resolution,[],[f3535,f740])).
fof(f3824,plain,(
  $false | $spl70),
  inference(subsumption_resolution,[],[f3740,f3534])).
fof(f3740,plain,(
  ~p(s(bool,'T_0')) | $spl70),
  inference(resolution,[],[f3534,f599])).
fof(f599,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f351])).
fof(f3823,plain,(
  $false | $spl70),
  inference(subsumption_resolution,[],[f3739,f3534])).
fof(f3739,plain,(
  ~p(s(bool,'T_0')) | $spl70),
  inference(resolution,[],[f3534,f597])).
fof(f597,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f350])).
fof(f3822,plain,(
  $false | $spl70),
  inference(subsumption_resolution,[],[f3736,f3534])).
fof(f3736,plain,(
  ( ! [X264] : (~p(s(bool,X264))) ) | $spl70),
  inference(resolution,[],[f3534,f591])).
fof(f591,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0] : ((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f151])).
fof(f151,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) => ~p(s(bool,X0)))),
  inference(flattening,[],[f150])).
fof(f150,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) => ~p(s(bool,X0)))),
  inference(rectify,[],[f64])).
fof(f64,axiom,(
  ! [X12] : ((p(s(bool,X12)) => p(s(bool,'F_0'))) => ~p(s(bool,X12)))),
  file('thm.rat.rat_of_num_ind','thm.bool.IMP_F')).
fof(f3821,plain,(
  $false | $spl70),
  inference(subsumption_resolution,[],[f3730,f3534])).
fof(f3730,plain,(
  ( ! [X258] : (~p(s(bool,X258))) ) | $spl70),
  inference(resolution,[],[f3534,f532])).
fof(f3738,plain,(
  $false | $spl70),
  inference(resolution,[],[f3534,f595])).
fof(f3754,plain,(
  $false | $spl70),
  inference(resolution,[],[f3534,f355])).
fof(f3755,plain,(
  $false | ($spl13 | $spl70)),
  inference(resolution,[],[f3534,f751])).
fof(f3756,plain,(
  $false | ($spl16 | $spl70)),
  inference(resolution,[],[f3534,f2146])).
fof(f3757,plain,(
  $false | $spl70),
  inference(resolution,[],[f3534,f762])).
fof(f3758,plain,(
  $false | $spl70),
  inference(resolution,[],[f3534,f763])).
fof(f3533,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1))) ) | $spl42),
  inference(subsumption_resolution,[],[f3526,f352])).
fof(f3526,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))))) ) | $spl42),
  inference(resolution,[],[f2907,f373])).
fof(f373,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP0(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f237])).
fof(f3622,plain,(
  ( ! [X0,X1] : (sP0(X0,X1,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK35(sK13,'type.num.num')) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK35(sK13,'type.num.num')))))) )),
  inference(superposition,[],[f2728,f365])).
fof(f3613,plain,(
  ( ! [X0,X1] : (sP7(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X0,X1) | s('type.num.num',sK34(sK13,'type.num.num')) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK34(sK13,'type.num.num')))))) )),
  inference(superposition,[],[f2616,f365])).
fof(f3605,plain,(
  ( ! [X0,X1] : (sP7(X0,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X1) | s('type.num.num',sK34(sK13,'type.num.num')) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK34(sK13,'type.num.num')))))) )),
  inference(superposition,[],[f2615,f365])).
fof(f3597,plain,(
  ( ! [X0,X1] : (sP8(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X0,X1) | s('type.num.num',sK34(sK13,'type.num.num')) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK34(sK13,'type.num.num')))))) )),
  inference(superposition,[],[f2613,f365])).
fof(f3580,plain,(
  ( ! [X0,X1] : (sP8(X0,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X1) | s('type.num.num',sK34(sK13,'type.num.num')) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK34(sK13,'type.num.num')))))) )),
  inference(superposition,[],[f2612,f365])).
fof(f3574,plain,(
  ( ! [X0,X1] : (sP4(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X0,X1) | s('type.num.num',sK34(sK13,'type.num.num')) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK34(sK13,'type.num.num')))))) )),
  inference(superposition,[],[f2600,f365])).
fof(f3567,plain,(
  ( ! [X0,X1] : (sP4(X0,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X1) | s('type.num.num',sK34(sK13,'type.num.num')) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK34(sK13,'type.num.num')))))) )),
  inference(superposition,[],[f2599,f365])).
fof(f3559,plain,(
  ( ! [X0,X1] : (sP2(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X0,X1) | s('type.num.num',sK34(sK13,'type.num.num')) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK34(sK13,'type.num.num')))))) )),
  inference(superposition,[],[f2586,f365])).
fof(f3552,plain,(
  ( ! [X0,X1] : (sP0(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X0,X1) | s('type.num.num',sK34(sK13,'type.num.num')) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK34(sK13,'type.num.num')))))) )),
  inference(superposition,[],[f2570,f365])).
fof(f3522,plain,(
  ( ! [X0,X1] : (sP0(X0,X1,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK34(sK13,'type.num.num')) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK34(sK13,'type.num.num')))))) )),
  inference(superposition,[],[f2569,f365])).
fof(f3014,plain,(
  s('type.num.num',sK20(sK15(sK13))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK20(sK15(sK13)))))) | $spl60),
  inference(cnf_transformation,[],[f3014_D])).
fof(f3014_D,plain,(
  s('type.num.num',sK20(sK15(sK13))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK20(sK15(sK13)))))) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f3016,plain,(
  s('type.num.num',sK15(sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | $spl62),
  inference(cnf_transformation,[],[f3016_D])).
fof(f3016_D,plain,(
  s('type.num.num',sK15(sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f3272,plain,(
  s('type.num.num',sK15(sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK20(sK15(sK13))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK20(sK15(sK13)))))) | $spl13),
  inference(superposition,[],[f2102,f365])).
fof(f2906,plain,(
  s('type.num.num',sK15(sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK15(sK13))))) | $spl40),
  inference(cnf_transformation,[],[f2906_D])).
fof(f2906_D,plain,(
  s('type.num.num',sK15(sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK15(sK13))))) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f3268,plain,(
  ( ! [X0,X1] : (sP7(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X0,X1) | s('type.num.num',sK15(sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK15(sK13)))))) ) | $spl13),
  inference(superposition,[],[f2034,f365])).
fof(f2034,plain,(
  ( ! [X109,X108] : (sP7(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15(sK13))),X108,X109)) ) | $spl13),
  inference(resolution,[],[f751,f431])).
fof(f3260,plain,(
  ( ! [X0,X1] : (sP7(X0,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X1) | s('type.num.num',sK15(sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK15(sK13)))))) ) | $spl13),
  inference(superposition,[],[f2033,f365])).
fof(f2033,plain,(
  ( ! [X107,X106] : (sP7(X106,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15(sK13))),X107)) ) | $spl13),
  inference(resolution,[],[f751,f430])).
fof(f3252,plain,(
  ( ! [X0,X1] : (sP8(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X0,X1) | s('type.num.num',sK15(sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK15(sK13)))))) ) | $spl13),
  inference(superposition,[],[f2031,f365])).
fof(f2031,plain,(
  ( ! [X103,X102] : (sP8(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15(sK13))),X102,X103)) ) | $spl13),
  inference(resolution,[],[f751,f427])).
fof(f3245,plain,(
  ( ! [X0,X1] : (sP8(X0,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X1) | s('type.num.num',sK15(sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK15(sK13)))))) ) | $spl13),
  inference(superposition,[],[f2030,f365])).
fof(f3239,plain,(
  ( ! [X0,X1] : (sP4(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X0,X1) | s('type.num.num',sK15(sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK15(sK13)))))) ) | $spl13),
  inference(superposition,[],[f2018,f365])).
fof(f2018,plain,(
  ( ! [X76,X77] : (sP4(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15(sK13))),X76,X77)) ) | $spl13),
  inference(resolution,[],[f751,f410])).
fof(f3232,plain,(
  ( ! [X0,X1] : (sP4(X0,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X1) | s('type.num.num',sK15(sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK15(sK13)))))) ) | $spl13),
  inference(superposition,[],[f2017,f365])).
fof(f2017,plain,(
  ( ! [X74,X75] : (sP4(X74,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15(sK13))),X75)) ) | $spl13),
  inference(resolution,[],[f751,f409])).
fof(f3224,plain,(
  ( ! [X0,X1] : (sP2(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X0,X1) | s('type.num.num',sK15(sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK15(sK13)))))) ) | $spl13),
  inference(superposition,[],[f2004,f365])).
fof(f2004,plain,(
  ( ! [X48,X49] : (sP2(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15(sK13))),X48,X49)) ) | $spl13),
  inference(resolution,[],[f751,f393])).
fof(f3217,plain,(
  ( ! [X0,X1] : (sP0(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X0,X1) | s('type.num.num',sK15(sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK15(sK13)))))) ) | $spl13),
  inference(superposition,[],[f1988,f365])).
fof(f1988,plain,(
  ( ! [X17,X16] : (sP0(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15(sK13))),X16,X17)) ) | $spl13),
  inference(resolution,[],[f751,f376])).
fof(f3210,plain,(
  ( ! [X0,X1] : (sP0(X0,X1,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK15(sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK15(sK13)))))) ) | $spl13),
  inference(superposition,[],[f1987,f365])).
fof(f1987,plain,(
  ( ! [X14,X15] : (sP0(X14,X15,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15(sK13))))) ) | $spl13),
  inference(resolution,[],[f751,f374])).
fof(f3064,plain,(
  $false | ($spl13 | $spl64)),
  inference(subsumption_resolution,[],[f3027,f352])).
fof(f3027,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | ($spl13 | $spl64)),
  inference(backward_demodulation,[],[f3019,f751])).
fof(f3019,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK15(sK13)) | $spl64),
  inference(cnf_transformation,[],[f3019_D])).
fof(f3019_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK15(sK13)) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f3021,plain,(
  s('type.num.num',sK15(sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.num.SUC_1'(s('type.num.num',sK20(sK15(sK13)))))))) | $spl66),
  inference(cnf_transformation,[],[f3021_D])).
fof(f3021_D,plain,(
  s('type.num.num',sK15(sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.num.SUC_1'(s('type.num.num',sK20(sK15(sK13)))))))) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f3026,plain,(
  s('type.num.num',sK15(sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.num.SUC_1'(s('type.num.num',sK20(sK15(sK13)))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK15(sK13)) | $spl13),
  inference(forward_demodulation,[],[f2997,f2102])).
fof(f2997,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK15(sK13)) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK15(sK13))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.num.SUC_1'(s('type.num.num',sK20(sK15(sK13)))))))) | $spl13),
  inference(superposition,[],[f365,f2102])).
fof(f3017,plain,(
  s('type.num.num',sK15(sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.num.SUC_1'(s('type.num.num',sK20(sK15(sK13)))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK15(sK13)) | $spl13),
  inference(forward_demodulation,[],[f2984,f2102])).
fof(f2984,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK15(sK13)) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK15(sK13))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.num.SUC_1'(s('type.num.num',sK20(sK15(sK13)))))))) | $spl13),
  inference(superposition,[],[f2102,f365])).
fof(f2983,plain,(
  s('type.num.num',sK15(sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK20(sK15(sK13))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK20(sK15(sK13)))))) | $spl13),
  inference(superposition,[],[f2102,f365])).
fof(f2978,plain,(
  ( ! [X0,X1] : (sP7(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X0,X1) | s('type.num.num',sK15(sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK15(sK13)))))) ) | $spl13),
  inference(superposition,[],[f2034,f365])).
fof(f2969,plain,(
  ( ! [X0,X1] : (sP7(X0,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X1) | s('type.num.num',sK15(sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK15(sK13)))))) ) | $spl13),
  inference(superposition,[],[f2033,f365])).
fof(f2960,plain,(
  ( ! [X0,X1] : (sP8(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X0,X1) | s('type.num.num',sK15(sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK15(sK13)))))) ) | $spl13),
  inference(superposition,[],[f2031,f365])).
fof(f2952,plain,(
  ( ! [X0,X1] : (sP8(X0,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X1) | s('type.num.num',sK15(sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK15(sK13)))))) ) | $spl13),
  inference(superposition,[],[f2030,f365])).
fof(f2945,plain,(
  ( ! [X0,X1] : (sP4(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X0,X1) | s('type.num.num',sK15(sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK15(sK13)))))) ) | $spl13),
  inference(superposition,[],[f2018,f365])).
fof(f2937,plain,(
  ( ! [X0,X1] : (sP4(X0,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X1) | s('type.num.num',sK15(sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK15(sK13)))))) ) | $spl13),
  inference(superposition,[],[f2017,f365])).
fof(f2928,plain,(
  ( ! [X0,X1] : (sP2(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X0,X1) | s('type.num.num',sK15(sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK15(sK13)))))) ) | $spl13),
  inference(superposition,[],[f2004,f365])).
fof(f2913,plain,(
  ( ! [X0,X1] : (sP0(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X0,X1) | s('type.num.num',sK15(sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK15(sK13)))))) ) | $spl13),
  inference(superposition,[],[f1988,f365])).
fof(f2903,plain,(
  ( ! [X0,X1] : (sP0(X0,X1,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK15(sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK15(sK13)))))) ) | $spl13),
  inference(superposition,[],[f1987,f365])).
fof(f2136,plain,(
  s('type.num.num',sK20(sK20(sK14))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK20(sK20(sK14)))))) | $spl30),
  inference(cnf_transformation,[],[f2136_D])).
fof(f2136_D,plain,(
  s('type.num.num',sK20(sK20(sK14))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK20(sK20(sK14)))))) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f2138,plain,(
  s('type.num.num',sK20(sK14)) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | $spl32),
  inference(cnf_transformation,[],[f2138_D])).
fof(f2138_D,plain,(
  s('type.num.num',sK20(sK14)) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f2316,plain,(
  s('type.num.num',sK20(sK14)) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK20(sK20(sK14))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK20(sK20(sK14)))))) | $spl16),
  inference(superposition,[],[f933,f365])).
fof(f2361,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2360,f355])).
fof(f2360,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) | $spl18),
  inference(forward_demodulation,[],[f353,f935])).
fof(f2313,plain,(
  s('type.num.num',sK14) = s('type.num.num',sK20(sK14)) | $spl38),
  inference(cnf_transformation,[],[f2313_D])).
fof(f2313_D,plain,(
  s('type.num.num',sK14) = s('type.num.num',sK20(sK14)) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f2346,plain,(
  s('type.num.num',sK14) = s('type.num.num',sK20(sK14)) | s('type.num.num',sK20(sK20(sK14))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK20(sK20(sK14)))))) | ($spl16 | $spl18)),
  inference(forward_demodulation,[],[f2316,f935])).
fof(f2315,plain,(
  $false | ($spl17 | $spl35)),
  inference(subsumption_resolution,[],[f2311,f2140])).
fof(f2140,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK20(sK14)) | $spl35),
  inference(cnf_transformation,[],[f2140_D])).
fof(f2140_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK20(sK14)) <=> ~$spl35),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl35])])).
fof(f2311,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK20(sK14)) | $spl17),
  inference(trivial_inequality_removal,[],[f2309])).
fof(f2309,plain,(
  s('type.num.num',sK20(sK14)) != s('type.num.num',sK20(sK14)) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK20(sK14)) | $spl17),
  inference(superposition,[],[f932,f365])).
fof(f932,plain,(
  s('type.num.num',sK20(sK14)) != s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK20(sK14))))) | $spl17),
  inference(cnf_transformation,[],[f932_D])).
fof(f932_D,plain,(
  s('type.num.num',sK20(sK14)) != s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK20(sK14))))) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f2314,plain,(
  s('type.num.num',sK14) != s('type.num.num',sK20(sK14)) | $spl39),
  inference(cnf_transformation,[],[f2314_D])).
fof(f2314_D,plain,(
  s('type.num.num',sK14) != s('type.num.num',sK20(sK14)) <=> ~$spl39),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl39])])).
fof(f2312,plain,(
  s('type.num.num',sK14) != s('type.num.num',sK20(sK14)) | s('type.num.num',sK20(sK20(sK14))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK20(sK20(sK14)))))) | ($spl17 | $spl18)),
  inference(forward_demodulation,[],[f2308,f935])).
fof(f2308,plain,(
  s('type.num.num',sK20(sK14)) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK20(sK20(sK14))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK20(sK20(sK14)))))) | $spl17),
  inference(superposition,[],[f932,f365])).
fof(f2218,plain,(
  $false | $spl34),
  inference(subsumption_resolution,[],[f2186,f355])).
fof(f2186,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) | $spl34),
  inference(backward_demodulation,[],[f2155,f353])).
fof(f2155,plain,(
  s('type.num.num',sK14) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | $spl34),
  inference(backward_demodulation,[],[f2141,f889])).
fof(f2141,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK20(sK14)) | $spl34),
  inference(cnf_transformation,[],[f2141_D])).
fof(f2141_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK20(sK14)) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f2143,plain,(
  s('type.num.num',sK20(sK14)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.num.SUC_1'(s('type.num.num',sK20(sK20(sK14)))))))) | $spl36),
  inference(cnf_transformation,[],[f2143_D])).
fof(f2143_D,plain,(
  s('type.num.num',sK20(sK14)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.num.SUC_1'(s('type.num.num',sK20(sK20(sK14)))))))) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f2153,plain,(
  s('type.num.num',sK20(sK14)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.num.SUC_1'(s('type.num.num',sK20(sK20(sK14)))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK20(sK14)) | $spl16),
  inference(forward_demodulation,[],[f2119,f933])).
fof(f2119,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK20(sK14)) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK20(sK14))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.num.SUC_1'(s('type.num.num',sK20(sK20(sK14)))))))) | $spl16),
  inference(superposition,[],[f365,f933])).
fof(f2139,plain,(
  s('type.num.num',sK20(sK14)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.num.SUC_1'(s('type.num.num',sK20(sK20(sK14)))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK20(sK14)) | $spl16),
  inference(forward_demodulation,[],[f2106,f933])).
fof(f2106,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK20(sK14)) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK20(sK14))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.num.SUC_1'(s('type.num.num',sK20(sK20(sK14)))))))) | $spl16),
  inference(superposition,[],[f933,f365])).
fof(f2105,plain,(
  s('type.num.num',sK20(sK14)) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK20(sK20(sK14))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK20(sK20(sK14)))))) | $spl16),
  inference(superposition,[],[f933,f365])).
fof(f751,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15(sK13))))) | $spl13),
  inference(cnf_transformation,[],[f751_D])).
fof(f751_D,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15(sK13))))) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f748,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16(sK13))))))) | $spl10),
  inference(cnf_transformation,[],[f748_D])).
fof(f748_D,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16(sK13))))))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f1836,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16(sK13))))))) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15(sK13))))) | $spl8),
  inference(resolution,[],[f747,f357])).
fof(f357,plain,(
  ( ! [X0,X3] : (p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))) | ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X3)))) | ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',sK15(X0)))))) )),
  inference(cnf_transformation,[],[f224])).
fof(f742,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK19(sK13))))))) | $spl4),
  inference(cnf_transformation,[],[f742_D])).
fof(f742_D,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK19(sK13))))))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f1700,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK19(sK13))))))) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15(sK13))))) | $spl0),
  inference(resolution,[],[f740,f357])).
fof(f1487,plain,(
  $false | $spl2),
  inference(resolution,[],[f741,f355])).
fof(f1405,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) | s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(X0))))) ) | $spl28),
  inference(cnf_transformation,[],[f1405_D])).
fof(f1405_D,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) | s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(X0))))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f1070,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl24),
  inference(cnf_transformation,[],[f1070_D])).
fof(f1070_D,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f1381,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) | s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(X0))))) )),
  inference(superposition,[],[f354,f365])).
fof(f1403,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))))))) ) | $spl26),
  inference(cnf_transformation,[],[f1403_D])).
fof(f1403_D,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))))))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f1273,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))))))) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15(sK13)))))) )),
  inference(resolution,[],[f354,f357])).
fof(f933,plain,(
  s('type.num.num',sK20(sK14)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK20(sK14))))) | $spl16),
  inference(cnf_transformation,[],[f933_D])).
fof(f933_D,plain,(
  s('type.num.num',sK20(sK14)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK20(sK14))))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f935,plain,(
  s('type.num.num',sK14) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | $spl18),
  inference(cnf_transformation,[],[f935_D])).
fof(f935_D,plain,(
  s('type.num.num',sK14) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f1105,plain,(
  s('type.num.num',sK14) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK20(sK14)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK20(sK14)))))),
  inference(superposition,[],[f889,f365])).
fof(f1099,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f1075,f352])).
fof(f1075,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | $spl20),
  inference(backward_demodulation,[],[f938,f355])).
fof(f944,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15(sK13)))))),
  inference(resolution,[],[f353,f357])).
fof(f938,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK14) | $spl20),
  inference(cnf_transformation,[],[f938_D])).
fof(f938_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK14) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f940,plain,(
  s('type.num.num',sK14) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.num.SUC_1'(s('type.num.num',sK20(sK14))))))) | $spl22),
  inference(cnf_transformation,[],[f940_D])).
fof(f940_D,plain,(
  s('type.num.num',sK14) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.num.SUC_1'(s('type.num.num',sK20(sK14))))))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f943,plain,(
  s('type.num.num',sK14) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.num.SUC_1'(s('type.num.num',sK20(sK14))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK14)),
  inference(forward_demodulation,[],[f916,f889])).
fof(f916,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK14) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK14)))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.num.SUC_1'(s('type.num.num',sK20(sK14)))))))),
  inference(superposition,[],[f365,f889])).
fof(f936,plain,(
  s('type.num.num',sK14) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.num.SUC_1'(s('type.num.num',sK20(sK14))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK14)),
  inference(forward_demodulation,[],[f905,f889])).
fof(f905,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK14) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK14)))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.num.SUC_1'(s('type.num.num',sK20(sK14)))))))),
  inference(superposition,[],[f889,f365])).
fof(f904,plain,(
  s('type.num.num',sK14) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK20(sK14)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK20(sK14)))))),
  inference(superposition,[],[f889,f365])).
fof(f615,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15(sK13)))))),
  inference(resolution,[],[f352,f357])).
fof(f614,plain,(
  ( ! [X4] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X4)))) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16(sK13)))))))) )),
  inference(resolution,[],[f352,f359])).
fof(f747,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16(sK13))))) | $spl8),
  inference(cnf_transformation,[],[f747_D])).
fof(f747_D,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16(sK13))))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f613,plain,(
  ( ! [X3] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X3)))) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16(sK13)))))) )),
  inference(resolution,[],[f352,f358])).
fof(f358,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X1)))) | p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',sK16(X0))))) | ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.|48|_0'))))) )),
  inference(cnf_transformation,[],[f224])).
fof(f612,plain,(
  ( ! [X2] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X2)))) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK18(sK13)))))))) )),
  inference(resolution,[],[f352,f362])).
fof(f611,plain,(
  ( ! [X1] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X1)))) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK19(sK13)))))))) )),
  inference(resolution,[],[f352,f364])).
fof(f740,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK19(sK13))))) | $spl0),
  inference(cnf_transformation,[],[f740_D])).
fof(f740_D,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK19(sK13))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f610,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X0)))) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK19(sK13)))))) )),
  inference(resolution,[],[f352,f363])).
fof(f363,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X1)))) | p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',sK19(X0))))) | ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.|48|_0'))))) )),
  inference(cnf_transformation,[],[f230])).
% SZS output end Proof for thm.rat
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 443
% Passive clauses: 3003
% Generated clauses: 7638
% Final active clauses: 350
% Final passive clauses: 2477
% Input formulas: 69
% Initial clauses: 250
% 
% Selected by SInE selection: 62
% SInE iterations: 1
% 
% Duplicate literals: 118
% Trivial inequalities: 25
% Fw subsumption resolutions: 889
% Fw demodulations: 335
% Bw demodulations: 73
% 
% Simple tautologies: 298
% Equational tautologies: 1
% Forward subsumptions: 2788
% 
% Binary resolution: 3851
% Forward superposition: 822
% Backward superposition: 1380
% Equality resolution: 1
% 
% Split clauses: 142
% Split components: 284
% SAT solver clauses: 179
% SAT solver unit clauses: 25
% SAT solver binary clauses: 83
% 
% Sat splits: 142
% Sat splitting refutations: 36
% 
% Memory used [KB]: 3454
% Time elapsed: 0.122 s
% ------------------------------
% ------------------------------
% Success in time 11.854 s
