Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.rich_list
% (62767)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 80
% Passive clauses: 2513
% Generated clauses: 16941
% Final active clauses: 72
% Final passive clauses: 2315
% Input formulas: 9
% Initial clauses: 14
% 
% Duplicate literals: 1126
% Trivial inequalities: 230
% Fw subsumption resolutions: 1963
% Fw demodulations: 1340
% Bw demodulations: 5
% 
% Simple tautologies: 320
% Equational tautologies: 103
% Forward subsumptions: 7328
% 
% Binary resolution: 61
% Forward superposition: 6801
% Backward superposition: 4567
% Self superposition: 380
% Equality factoring: 291
% Equality resolution: 12
% 
% Split clauses: 857
% Split components: 1726
% SAT solver clauses: 2041
% SAT solver unit clauses: 250
% SAT solver binary clauses: 655
% 
% Sat splits: 857
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 2033
% TWLsolver calls for satisfiability: 76
% 
% Memory used [KB]: 7931
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.rich_list
% (62787)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 54
% Passive clauses: 278
% Generated clauses: 7679
% Final active clauses: 52
% Final passive clauses: 217
% Input formulas: 9
% Initial clauses: 14
% 
% Duplicate literals: 934
% Trivial inequalities: 99
% Fw subsumption resolutions: 467
% Fw demodulations: 274
% Bw demodulations: 8
% Condensations: 1485
% 
% Simple tautologies: 170
% Equational tautologies: 58
% Forward subsumptions: 2766
% Fw demodulations to eq. taut.: 1
% 
% Binary resolution: 23
% Forward superposition: 2398
% Backward superposition: 1512
% Self superposition: 197
% Equality factoring: 168
% Equality resolution: 8
% 
% Split clauses: 63
% Split components: 127
% SAT solver clauses: 702
% SAT solver unit clauses: 20
% SAT solver binary clauses: 68
% 
% Sat splits: 63
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 702
% TWLsolver calls for satisfiability: 79
% 
% Memory used [KB]: 6780
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.rich_list
% (62811)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 322
% Passive clauses: 6113
% Generated clauses: 102287
% Final active clauses: 130
% Final passive clauses: 656
% Input formulas: 9
% Initial clauses: 14
% Discarded non-redundant clauses: 3949
% 
% Duplicate literals: 6400
% Trivial inequalities: 1016
% Fw subsumption resolutions: 11009
% Fw demodulations: 10579
% Bw demodulations: 2598
% Condensations: 11726
% Global subsumptions: 13
% 
% Simple tautologies: 993
% Equational tautologies: 825
% Forward subsumptions: 52924
% Fw demodulations to eq. taut.: 10
% Bw demodulations to eq. taut.: 254
% 
% Unit resulting resolution: 2221
% Forward superposition: 30967
% Backward superposition: 20986
% Self superposition: 632
% Equality factoring: 356
% Equality resolution: 45
% 
% Split clauses: 306
% Split components: 613
% SAT solver clauses: 6848
% SAT solver unit clauses: 18
% SAT solver binary clauses: 181
% 
% Sat splits: 306
% Sat splitting refutations: 108
% 
% Memory used [KB]: 52834
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.rich_list
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.rich_list
% SZS output start Proof for thm.rich_list
fof(f404,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f403])).
fof(f403,plain,(
  s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,sK2))),s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,sK3))),s(sK0,sK4))))) != s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,sK2))),s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,sK3))),s(sK0,sK4)))))),
  inference(forward_demodulation,[],[f350,f49])).
fof(f49,plain,(
  ( ! [X0,X1] : (s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,X0))),s(sK0,X1))) = s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,X1))),s(sK0,X0)))) )),
  inference(resolution,[],[f34,f37])).
fof(f37,plain,(
  ( ! [X6,X2,X0,X5,X1] : (s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,X5))),s(X1,X6))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,X6))),s(X1,X5))) | ~p(s(bool,'const.operator.COMM_1'(s(fun(X1,fun(X1,X0)),X2))))) )),
  inference(cnf_transformation,[],[f28])).
fof(f28,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.operator.COMM_1'(s(fun(X1,fun(X1,X0)),X2)))) | ! [X5,X6] : s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,X5))),s(X1,X6))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,X6))),s(X1,X5)))) & (s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,sK5(X2,X1,X0)))),s(X1,sK6(X2,X1,X0)))) != s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,sK6(X2,X1,X0)))),s(X1,sK5(X2,X1,X0)))) | p(s(bool,'const.operator.COMM_1'(s(fun(X1,fun(X1,X0)),X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5,sK6])],[f27])).
fof(f27,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.operator.COMM_1'(s(fun(X1,fun(X1,X0)),X2)))) | ! [X5,X6] : s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,X5))),s(X1,X6))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,X6))),s(X1,X5)))) & (? [X3,X4] : s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,X3))),s(X1,X4))) != s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,X4))),s(X1,X3))) | p(s(bool,'const.operator.COMM_1'(s(fun(X1,fun(X1,X0)),X2))))))),
  inference(rectify,[],[f26])).
fof(f26,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.operator.COMM_1'(s(fun(X1,fun(X1,X0)),X2)))) | ! [X3,X4] : s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,X3))),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,X4))),s(X1,X3)))) & (? [X3,X4] : s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,X3))),s(X1,X4))) != s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,X4))),s(X1,X3))) | p(s(bool,'const.operator.COMM_1'(s(fun(X1,fun(X1,X0)),X2))))))),
  inference(nnf_transformation,[],[f14])).
fof(f14,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.operator.COMM_1'(s(fun(X1,fun(X1,X0)),X2)))) <=> ! [X3,X4] : s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,X3))),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,X4))),s(X1,X3))))),
  inference(flattening,[],[f13])).
fof(f13,plain,(
  ! [X0,X1] : ! [X2] : (p(s(bool,'const.operator.COMM_1'(s(fun(X1,fun(X1,X0)),X2)))) <=> ! [X3,X4] : s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,X3))),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,X4))),s(X1,X3))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X14,X6] : ! [X7] : (p(s(bool,'const.operator.COMM_1'(s(fun(X6,fun(X6,X14)),X7)))) <=> ! [X11,X12] : s(X14,happ(s(fun(X6,X14),happ(s(fun(X6,fun(X6,X14)),X7),s(X6,X11))),s(X6,X12))) = s(X14,happ(s(fun(X6,X14),happ(s(fun(X6,fun(X6,X14)),X7),s(X6,X12))),s(X6,X11))))),
  file('thm.rich_list.rich_list_2|48||49|_prove_ltb','thm.operator.COMM_DEF')).
fof(f34,plain,(
  p(s(bool,'const.operator.COMM_1'(s(fun(sK0,fun(sK0,sK0)),sK1))))),
  inference(cnf_transformation,[],[f25])).
fof(f25,plain,(
  p(s(bool,'const.operator.COMM_1'(s(fun(sK0,fun(sK0,sK0)),sK1)))) & p(s(bool,'const.operator.ASSOC_1'(s(fun(sK0,fun(sK0,sK0)),sK1)))) & s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,sK2))),s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,sK3))),s(sK0,sK4))))) != s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,sK3))),s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,sK2))),s(sK0,sK4)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4])],[f23])).
fof(f23,plain,(
  ? [X0,X1] : (p(s(bool,'const.operator.COMM_1'(s(fun(X0,fun(X0,X0)),X1)))) & p(s(bool,'const.operator.ASSOC_1'(s(fun(X0,fun(X0,X0)),X1)))) & ? [X2,X3,X4] : s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X2))),s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X3))),s(X0,X4))))) != s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X3))),s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X2))),s(X0,X4))))))),
  inference(flattening,[],[f22])).
fof(f22,plain,(
  ? [X0,X1] : (p(s(bool,'const.operator.COMM_1'(s(fun(X0,fun(X0,X0)),X1)))) & (p(s(bool,'const.operator.ASSOC_1'(s(fun(X0,fun(X0,X0)),X1)))) & ? [X2,X3,X4] : s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X2))),s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X3))),s(X0,X4))))) != s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X3))),s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X2))),s(X0,X4)))))))),
  inference(ennf_transformation,[],[f12])).
fof(f12,plain,(
  ~! [X0,X1] : (p(s(bool,'const.operator.COMM_1'(s(fun(X0,fun(X0,X0)),X1)))) => (p(s(bool,'const.operator.ASSOC_1'(s(fun(X0,fun(X0,X0)),X1)))) => ! [X2,X3,X4] : s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X2))),s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X3))),s(X0,X4))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X3))),s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X2))),s(X0,X4)))))))),
  inference(flattening,[],[f11])).
fof(f11,plain,(
  ~! [X0] : ! [X1] : (p(s(bool,'const.operator.COMM_1'(s(fun(X0,fun(X0,X0)),X1)))) => (p(s(bool,'const.operator.ASSOC_1'(s(fun(X0,fun(X0,X0)),X1)))) => ! [X2,X3,X4] : s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X2))),s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X3))),s(X0,X4))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X3))),s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X2))),s(X0,X4)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7] : (p(s(bool,'const.operator.COMM_1'(s(fun(X6,fun(X6,X6)),X7)))) => (p(s(bool,'const.operator.ASSOC_1'(s(fun(X6,fun(X6,X6)),X7)))) => ! [X8,X9,X10] : s(X6,happ(s(fun(X6,X6),happ(s(fun(X6,fun(X6,X6)),X7),s(X6,X8))),s(X6,happ(s(fun(X6,X6),happ(s(fun(X6,fun(X6,X6)),X7),s(X6,X9))),s(X6,X10))))) = s(X6,happ(s(fun(X6,X6),happ(s(fun(X6,fun(X6,X6)),X7),s(X6,X9))),s(X6,happ(s(fun(X6,X6),happ(s(fun(X6,fun(X6,X6)),X7),s(X6,X8))),s(X6,X10)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7] : (p(s(bool,'const.operator.COMM_1'(s(fun(X6,fun(X6,X6)),X7)))) => (p(s(bool,'const.operator.ASSOC_1'(s(fun(X6,fun(X6,X6)),X7)))) => ! [X8,X9,X10] : s(X6,happ(s(fun(X6,X6),happ(s(fun(X6,fun(X6,X6)),X7),s(X6,X8))),s(X6,happ(s(fun(X6,X6),happ(s(fun(X6,fun(X6,X6)),X7),s(X6,X9))),s(X6,X10))))) = s(X6,happ(s(fun(X6,X6),happ(s(fun(X6,fun(X6,X6)),X7),s(X6,X9))),s(X6,happ(s(fun(X6,X6),happ(s(fun(X6,fun(X6,X6)),X7),s(X6,X8))),s(X6,X10)))))))),
  file('thm.rich_list.rich_list_2|48||49|_prove_ltb','thm.rich_list.rich_list_2|48||49|_prove_ltb')).
fof(f350,plain,(
  s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,sK2))),s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,sK3))),s(sK0,sK4))))) != s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,sK2))),s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,sK4))),s(sK0,sK3)))))),
  inference(superposition,[],[f36,f140])).
fof(f140,plain,(
  ( ! [X10,X8,X9] : (s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,X8))),s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,X9))),s(sK0,X10))))) = s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,X10))),s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,X8))),s(sK0,X9)))))) )),
  inference(subsumption_resolution,[],[f139,f48])).
fof(f48,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f41])).
fof(f41,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f32])).
fof(f32,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f18])).
fof(f18,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f17])).
fof(f17,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : ! [X11] : (s(X6,X11) = s(X6,X11) <=> p(s(bool,'T_0')))),
  file('thm.rich_list.rich_list_2|48||49|_prove_ltb','thm.bool.REFL_CLAUSE')).
fof(f139,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,'T_0')) | s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,X8))),s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,X9))),s(sK0,X10))))) = s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,X10))),s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,X8))),s(sK0,X9)))))) )),
  inference(forward_demodulation,[],[f103,f60])).
fof(f60,plain,(
  s(bool,'T_0') = s(bool,'const.operator.ASSOC_1'(s(fun(sK0,fun(sK0,sK0)),sK1)))),
  inference(subsumption_resolution,[],[f59,f44])).
fof(f44,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f19])).
fof(f19,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.rich_list.rich_list_2|48||49|_prove_ltb',aHL_FALSITY)).
fof(f59,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.operator.ASSOC_1'(s(fun(sK0,fun(sK0,sK0)),sK1)))),
  inference(superposition,[],[f35,f43])).
fof(f43,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.rich_list.rich_list_2|48||49|_prove_ltb',aHL_BOOL_CASES)).
fof(f35,plain,(
  p(s(bool,'const.operator.ASSOC_1'(s(fun(sK0,fun(sK0,sK0)),sK1))))),
  inference(cnf_transformation,[],[f25])).
fof(f103,plain,(
  ( ! [X10,X8,X9] : (s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,X8))),s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,X9))),s(sK0,X10))))) = s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,X10))),s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,X8))),s(sK0,X9))))) | ~p(s(bool,'const.operator.ASSOC_1'(s(fun(sK0,fun(sK0,sK0)),sK1))))) )),
  inference(superposition,[],[f49,f39])).
fof(f39,plain,(
  ( ! [X6,X0,X7,X5,X1] : (s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X5))),s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X6))),s(X0,X7))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X5))),s(X0,X6))))),s(X0,X7))) | ~p(s(bool,'const.operator.ASSOC_1'(s(fun(X0,fun(X0,X0)),X1))))) )),
  inference(cnf_transformation,[],[f31])).
fof(f31,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.operator.ASSOC_1'(s(fun(X0,fun(X0,X0)),X1)))) | ! [X5,X6,X7] : s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X5))),s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X6))),s(X0,X7))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X5))),s(X0,X6))))),s(X0,X7)))) & (s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,sK7(X1,X0)))),s(X0,sK8(X1,X0)))))),s(X0,sK9(X1,X0)))) != s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,sK7(X1,X0)))),s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,sK8(X1,X0)))),s(X0,sK9(X1,X0)))))) | p(s(bool,'const.operator.ASSOC_1'(s(fun(X0,fun(X0,X0)),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK7,sK8,sK9])],[f30])).
fof(f30,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.operator.ASSOC_1'(s(fun(X0,fun(X0,X0)),X1)))) | ! [X5,X6,X7] : s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X5))),s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X6))),s(X0,X7))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X5))),s(X0,X6))))),s(X0,X7)))) & (? [X2,X3,X4] : s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X2))),s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X3))),s(X0,X4))))) != s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X2))),s(X0,X3))))),s(X0,X4))) | p(s(bool,'const.operator.ASSOC_1'(s(fun(X0,fun(X0,X0)),X1))))))),
  inference(rectify,[],[f29])).
fof(f29,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.operator.ASSOC_1'(s(fun(X0,fun(X0,X0)),X1)))) | ! [X2,X3,X4] : s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X2))),s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X3))),s(X0,X4))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X2))),s(X0,X3))))),s(X0,X4)))) & (? [X2,X3,X4] : s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X2))),s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X3))),s(X0,X4))))) != s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X2))),s(X0,X3))))),s(X0,X4))) | p(s(bool,'const.operator.ASSOC_1'(s(fun(X0,fun(X0,X0)),X1))))))),
  inference(nnf_transformation,[],[f16])).
fof(f16,plain,(
  ! [X0,X1] : (p(s(bool,'const.operator.ASSOC_1'(s(fun(X0,fun(X0,X0)),X1)))) <=> ! [X2,X3,X4] : s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X2))),s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X3))),s(X0,X4))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X2))),s(X0,X3))))),s(X0,X4))))),
  inference(flattening,[],[f15])).
fof(f15,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.operator.ASSOC_1'(s(fun(X0,fun(X0,X0)),X1)))) <=> ! [X2,X3,X4] : s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X2))),s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X3))),s(X0,X4))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X2))),s(X0,X3))))),s(X0,X4))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6] : ! [X7] : (p(s(bool,'const.operator.ASSOC_1'(s(fun(X6,fun(X6,X6)),X7)))) <=> ! [X11,X12,X13] : s(X6,happ(s(fun(X6,X6),happ(s(fun(X6,fun(X6,X6)),X7),s(X6,X11))),s(X6,happ(s(fun(X6,X6),happ(s(fun(X6,fun(X6,X6)),X7),s(X6,X12))),s(X6,X13))))) = s(X6,happ(s(fun(X6,X6),happ(s(fun(X6,fun(X6,X6)),X7),s(X6,happ(s(fun(X6,X6),happ(s(fun(X6,fun(X6,X6)),X7),s(X6,X11))),s(X6,X12))))),s(X6,X13))))),
  file('thm.rich_list.rich_list_2|48||49|_prove_ltb','thm.operator.ASSOC_DEF')).
fof(f36,plain,(
  s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,sK2))),s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,sK3))),s(sK0,sK4))))) != s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,sK3))),s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,sK2))),s(sK0,sK4)))))),
  inference(cnf_transformation,[],[f25])).
% SZS output end Proof for thm.rich_list
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 51
% Passive clauses: 137
% Generated clauses: 359
% Final active clauses: 49
% Final passive clauses: 92
% Input formulas: 9
% Initial clauses: 14
% 
% Selected by SInE selection: 9
% SInE iterations: 1
% 
% Trivial inequalities: 5
% Fw subsumption resolutions: 26
% Fw demodulations: 54
% Bw demodulations: 2
% 
% Equational tautologies: 1
% Forward subsumptions: 122
% Fw demodulations to eq. taut.: 3
% 
% Binary resolution: 2
% Forward superposition: 160
% Backward superposition: 88
% Self superposition: 20
% 
% SAT solver clauses: 1
% 
% Memory used [KB]: 1663
% Time elapsed: 0.030 s
% ------------------------------
% ------------------------------
% Success in time 11.755 s
