Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.pred_set
% (68345)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 373
% Passive clauses: 3180
% Generated clauses: 17056
% Final active clauses: 367
% Final passive clauses: 2499
% Input formulas: 58
% Initial clauses: 153
% 
% Duplicate literals: 544
% Trivial inequalities: 79
% Fw subsumption resolutions: 1537
% Fw demodulations: 1742
% Bw demodulations: 4
% 
% Simple tautologies: 463
% Equational tautologies: 55
% Forward subsumptions: 8579
% 
% Binary resolution: 4970
% Factoring: 18
% Forward superposition: 3885
% Backward superposition: 3819
% Self superposition: 99
% Equality factoring: 62
% Equality resolution: 28
% 
% Split clauses: 498
% Split components: 1018
% SAT solver clauses: 728
% SAT solver unit clauses: 27
% SAT solver binary clauses: 164
% 
% Sat splits: 498
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 720
% TWLsolver calls for satisfiability: 165
% 
% Memory used [KB]: 10106
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.pred_set
% (68358)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 239
% Passive clauses: 1357
% Generated clauses: 10542
% Final active clauses: 208
% Final passive clauses: 1061
% Input formulas: 58
% Initial clauses: 153
% 
% Duplicate literals: 253
% Trivial inequalities: 52
% Fw subsumption resolutions: 617
% Fw demodulations: 1011
% Bw demodulations: 51
% Condensations: 574
% Global subsumptions: 18
% 
% Simple tautologies: 214
% Equational tautologies: 41
% Forward subsumptions: 4389
% 
% Binary resolution: 827
% Factoring: 10
% Forward superposition: 2773
% Backward superposition: 2614
% Self superposition: 90
% Equality factoring: 50
% Equality resolution: 23
% 
% Split clauses: 77
% Split components: 165
% SAT solver clauses: 2202
% SAT solver unit clauses: 183
% SAT solver binary clauses: 83
% SAT solver learnt clauses: 36
% SAT solver learnt literals: 10
% 
% Sat splits: 77
% Sat splitting refutations: 14
% 
% TWLsolver clauses: 2130
% TWLsolver calls for satisfiability: 1125
% 
% Memory used [KB]: 13432
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f3525,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f3524])).
fof(f3524,plain,(
  s(bool,'const.bool.IN_2'(s(fun(sK7,bool),sK10('const.bool.LET_2'(s(fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool)),'_dst_x0x1x1_3'(s(fun(fun(fun(sK7,bool),bool),fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool))),'const.pred_set.UNION_0'),s(fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool)),'const.pred_set.IMAGE_1'(s(fun(fun(sK7,bool),fun(sK7,bool)),'const.pred_set.INSERT_1'(s(sK7,sK8))))),s(fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool)),'_i_0_0'))),s(fun(fun(sK7,bool),bool),'const.pred_set.POW_1'(s(fun(sK7,bool),sK9)))),'const.pred_set.POW_1'(s(fun(sK7,bool),'const.pred_set.INSERT_2'(s(sK7,sK8),s(fun(sK7,bool),sK9)))),fun(sK7,bool))),s(fun(fun(sK7,bool),bool),'const.pred_set.POW_1'(s(fun(sK7,bool),'const.pred_set.INSERT_2'(s(sK7,sK8),s(fun(sK7,bool),sK9))))))) != s(bool,'const.bool.IN_2'(s(fun(sK7,bool),sK10('const.bool.LET_2'(s(fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool)),'_dst_x0x1x1_3'(s(fun(fun(fun(sK7,bool),bool),fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool))),'const.pred_set.UNION_0'),s(fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool)),'const.pred_set.IMAGE_1'(s(fun(fun(sK7,bool),fun(sK7,bool)),'const.pred_set.INSERT_1'(s(sK7,sK8))))),s(fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool)),'_i_0_0'))),s(fun(fun(sK7,bool),bool),'const.pred_set.POW_1'(s(fun(sK7,bool),sK9)))),'const.pred_set.POW_1'(s(fun(sK7,bool),'const.pred_set.INSERT_2'(s(sK7,sK8),s(fun(sK7,bool),sK9)))),fun(sK7,bool))),s(fun(fun(sK7,bool),bool),'const.pred_set.POW_1'(s(fun(sK7,bool),'const.pred_set.INSERT_2'(s(sK7,sK8),s(fun(sK7,bool),sK9)))))))),
  inference(forward_demodulation,[],[f3523,f424])).
fof(f424,plain,(
  ( ! [X2,X0,X1] : (s(fun(fun(X0,bool),bool),'const.pred_set.POW_1'(s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2))))) = s(fun(fun(X0,bool),bool),'const.pred_set.UNION_2'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun(fun(X0,bool),fun(X0,bool)),'const.pred_set.INSERT_1'(s(X0,X1))),s(fun(fun(X0,bool),bool),'const.pred_set.POW_1'(s(fun(X0,bool),X2))))),s(fun(fun(X0,bool),bool),'const.pred_set.POW_1'(s(fun(X0,bool),X2)))))) )),
  inference(forward_demodulation,[],[f332,f329])).
fof(f329,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2))) = s(fun(X0,bool),happ(s(fun(fun(X0,bool),fun(X0,bool)),'const.pred_set.INSERT_1'(s(X0,X1))),s(fun(X0,bool),X2)))) )),
  inference(cnf_transformation,[],[f107])).
fof(f107,plain,(
  ! [X0,X1,X2] : s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2))) = s(fun(X0,bool),happ(s(fun(fun(X0,bool),fun(X0,bool)),'const.pred_set.INSERT_1'(s(X0,X1))),s(fun(X0,bool),X2)))),
  inference(flattening,[],[f106])).
fof(f106,plain,(
  ! [X0] : ! [X1,X2] : s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2))) = s(fun(X0,bool),happ(s(fun(fun(X0,bool),fun(X0,bool)),'const.pred_set.INSERT_1'(s(X0,X1))),s(fun(X0,bool),X2)))),
  inference(rectify,[],[f59])).
fof(f59,axiom,(
  ! [X6] : ! [X23,X24] : s(fun(X6,bool),'const.pred_set.INSERT_2'(s(X6,X23),s(fun(X6,bool),X24))) = s(fun(X6,bool),happ(s(fun(fun(X6,bool),fun(X6,bool)),'const.pred_set.INSERT_1'(s(X6,X23))),s(fun(X6,bool),X24)))),
  file('thm.pred_set.POW_EQNS|split|1','ARITY_1_2_const.pred_set.INSERT')).
fof(f332,plain,(
  ( ! [X2,X0,X1] : (s(fun(fun(X0,bool),bool),'const.pred_set.POW_1'(s(fun(X0,bool),happ(s(fun(fun(X0,bool),fun(X0,bool)),'const.pred_set.INSERT_1'(s(X0,X1))),s(fun(X0,bool),X2))))) = s(fun(fun(X0,bool),bool),'const.pred_set.UNION_2'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun(fun(X0,bool),fun(X0,bool)),'const.pred_set.INSERT_1'(s(X0,X1))),s(fun(fun(X0,bool),bool),'const.pred_set.POW_1'(s(fun(X0,bool),X2))))),s(fun(fun(X0,bool),bool),'const.pred_set.POW_1'(s(fun(X0,bool),X2)))))) )),
  inference(cnf_transformation,[],[f113])).
fof(f113,plain,(
  ! [X0,X1,X2] : s(fun(fun(X0,bool),bool),'const.pred_set.POW_1'(s(fun(X0,bool),happ(s(fun(fun(X0,bool),fun(X0,bool)),'const.pred_set.INSERT_1'(s(X0,X1))),s(fun(X0,bool),X2))))) = s(fun(fun(X0,bool),bool),'const.pred_set.UNION_2'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun(fun(X0,bool),fun(X0,bool)),'const.pred_set.INSERT_1'(s(X0,X1))),s(fun(fun(X0,bool),bool),'const.pred_set.POW_1'(s(fun(X0,bool),X2))))),s(fun(fun(X0,bool),bool),'const.pred_set.POW_1'(s(fun(X0,bool),X2)))))),
  inference(flattening,[],[f112])).
fof(f112,plain,(
  ! [X0] : ! [X1,X2] : s(fun(fun(X0,bool),bool),'const.pred_set.POW_1'(s(fun(X0,bool),happ(s(fun(fun(X0,bool),fun(X0,bool)),'const.pred_set.INSERT_1'(s(X0,X1))),s(fun(X0,bool),X2))))) = s(fun(fun(X0,bool),bool),'const.pred_set.UNION_2'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun(fun(X0,bool),fun(X0,bool)),'const.pred_set.INSERT_1'(s(X0,X1))),s(fun(fun(X0,bool),bool),'const.pred_set.POW_1'(s(fun(X0,bool),X2))))),s(fun(fun(X0,bool),bool),'const.pred_set.POW_1'(s(fun(X0,bool),X2)))))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X6] : ! [X7,X8] : s(fun(fun(X6,bool),bool),'const.pred_set.POW_1'(s(fun(X6,bool),happ(s(fun(fun(X6,bool),fun(X6,bool)),'const.pred_set.INSERT_1'(s(X6,X7))),s(fun(X6,bool),X8))))) = s(fun(fun(X6,bool),bool),'const.pred_set.UNION_2'(s(fun(fun(X6,bool),bool),'const.pred_set.IMAGE_2'(s(fun(fun(X6,bool),fun(X6,bool)),'const.pred_set.INSERT_1'(s(X6,X7))),s(fun(fun(X6,bool),bool),'const.pred_set.POW_1'(s(fun(X6,bool),X8))))),s(fun(fun(X6,bool),bool),'const.pred_set.POW_1'(s(fun(X6,bool),X8)))))),
  file('thm.pred_set.POW_EQNS|split|1','thm.pred_set.POW_INSERT')).
fof(f3523,plain,(
  s(bool,'const.bool.IN_2'(s(fun(sK7,bool),sK10('const.bool.LET_2'(s(fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool)),'_dst_x0x1x1_3'(s(fun(fun(fun(sK7,bool),bool),fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool))),'const.pred_set.UNION_0'),s(fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool)),'const.pred_set.IMAGE_1'(s(fun(fun(sK7,bool),fun(sK7,bool)),'const.pred_set.INSERT_1'(s(sK7,sK8))))),s(fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool)),'_i_0_0'))),s(fun(fun(sK7,bool),bool),'const.pred_set.POW_1'(s(fun(sK7,bool),sK9)))),'const.pred_set.POW_1'(s(fun(sK7,bool),'const.pred_set.INSERT_2'(s(sK7,sK8),s(fun(sK7,bool),sK9)))),fun(sK7,bool))),s(fun(fun(sK7,bool),bool),'const.pred_set.POW_1'(s(fun(sK7,bool),'const.pred_set.INSERT_2'(s(sK7,sK8),s(fun(sK7,bool),sK9))))))) != s(bool,'const.bool.IN_2'(s(fun(sK7,bool),sK10('const.bool.LET_2'(s(fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool)),'_dst_x0x1x1_3'(s(fun(fun(fun(sK7,bool),bool),fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool))),'const.pred_set.UNION_0'),s(fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool)),'const.pred_set.IMAGE_1'(s(fun(fun(sK7,bool),fun(sK7,bool)),'const.pred_set.INSERT_1'(s(sK7,sK8))))),s(fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool)),'_i_0_0'))),s(fun(fun(sK7,bool),bool),'const.pred_set.POW_1'(s(fun(sK7,bool),sK9)))),'const.pred_set.POW_1'(s(fun(sK7,bool),'const.pred_set.INSERT_2'(s(sK7,sK8),s(fun(sK7,bool),sK9)))),fun(sK7,bool))),s(fun(fun(sK7,bool),bool),'const.pred_set.UNION_2'(s(fun(fun(sK7,bool),bool),'const.pred_set.IMAGE_2'(s(fun(fun(sK7,bool),fun(sK7,bool)),'const.pred_set.INSERT_1'(s(sK7,sK8))),s(fun(fun(sK7,bool),bool),'const.pred_set.POW_1'(s(fun(sK7,bool),sK9))))),s(fun(fun(sK7,bool),bool),'const.pred_set.POW_1'(s(fun(sK7,bool),sK9)))))))),
  inference(forward_demodulation,[],[f3522,f364])).
fof(f364,plain,(
  ( ! [X2,X0,X3,X1] : (s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X1,X0),X2),s(fun(X1,bool),X3))) = s(fun(X0,bool),happ(s(fun(fun(X1,bool),fun(X0,bool)),'const.pred_set.IMAGE_1'(s(fun(X1,X0),X2))),s(fun(X1,bool),X3)))) )),
  inference(cnf_transformation,[],[f135])).
fof(f135,plain,(
  ! [X0,X1,X2,X3] : s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X1,X0),X2),s(fun(X1,bool),X3))) = s(fun(X0,bool),happ(s(fun(fun(X1,bool),fun(X0,bool)),'const.pred_set.IMAGE_1'(s(fun(X1,X0),X2))),s(fun(X1,bool),X3)))),
  inference(flattening,[],[f134])).
fof(f134,plain,(
  ! [X0,X1] : ! [X2,X3] : s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X1,X0),X2),s(fun(X1,bool),X3))) = s(fun(X0,bool),happ(s(fun(fun(X1,bool),fun(X0,bool)),'const.pred_set.IMAGE_1'(s(fun(X1,X0),X2))),s(fun(X1,bool),X3)))),
  inference(rectify,[],[f52])).
fof(f52,axiom,(
  ! [X12,X6] : ! [X23,X24] : s(fun(X12,bool),'const.pred_set.IMAGE_2'(s(fun(X6,X12),X23),s(fun(X6,bool),X24))) = s(fun(X12,bool),happ(s(fun(fun(X6,bool),fun(X12,bool)),'const.pred_set.IMAGE_1'(s(fun(X6,X12),X23))),s(fun(X6,bool),X24)))),
  file('thm.pred_set.POW_EQNS|split|1','ARITY_1_2_const.pred_set.IMAGE')).
fof(f3522,plain,(
  s(bool,'const.bool.IN_2'(s(fun(sK7,bool),sK10('const.bool.LET_2'(s(fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool)),'_dst_x0x1x1_3'(s(fun(fun(fun(sK7,bool),bool),fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool))),'const.pred_set.UNION_0'),s(fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool)),'const.pred_set.IMAGE_1'(s(fun(fun(sK7,bool),fun(sK7,bool)),'const.pred_set.INSERT_1'(s(sK7,sK8))))),s(fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool)),'_i_0_0'))),s(fun(fun(sK7,bool),bool),'const.pred_set.POW_1'(s(fun(sK7,bool),sK9)))),'const.pred_set.POW_1'(s(fun(sK7,bool),'const.pred_set.INSERT_2'(s(sK7,sK8),s(fun(sK7,bool),sK9)))),fun(sK7,bool))),s(fun(fun(sK7,bool),bool),'const.pred_set.POW_1'(s(fun(sK7,bool),'const.pred_set.INSERT_2'(s(sK7,sK8),s(fun(sK7,bool),sK9))))))) != s(bool,'const.bool.IN_2'(s(fun(sK7,bool),sK10('const.bool.LET_2'(s(fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool)),'_dst_x0x1x1_3'(s(fun(fun(fun(sK7,bool),bool),fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool))),'const.pred_set.UNION_0'),s(fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool)),'const.pred_set.IMAGE_1'(s(fun(fun(sK7,bool),fun(sK7,bool)),'const.pred_set.INSERT_1'(s(sK7,sK8))))),s(fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool)),'_i_0_0'))),s(fun(fun(sK7,bool),bool),'const.pred_set.POW_1'(s(fun(sK7,bool),sK9)))),'const.pred_set.POW_1'(s(fun(sK7,bool),'const.pred_set.INSERT_2'(s(sK7,sK8),s(fun(sK7,bool),sK9)))),fun(sK7,bool))),s(fun(fun(sK7,bool),bool),'const.pred_set.UNION_2'(s(fun(fun(sK7,bool),bool),happ(s(fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool)),'const.pred_set.IMAGE_1'(s(fun(fun(sK7,bool),fun(sK7,bool)),'const.pred_set.INSERT_1'(s(sK7,sK8))))),s(fun(fun(sK7,bool),bool),'const.pred_set.POW_1'(s(fun(sK7,bool),sK9))))),s(fun(fun(sK7,bool),bool),'const.pred_set.POW_1'(s(fun(sK7,bool),sK9)))))))),
  inference(forward_demodulation,[],[f3518,f320])).
fof(f320,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f96])).
fof(f96,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(flattening,[],[f95])).
fof(f95,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(rectify,[],[f54])).
fof(f54,axiom,(
  ! [X25] : ! [X26] : s(X25,X26) = s(X25,happ(s(fun(X25,X25),'_i_0_0'),s(X25,X26)))),
  file('thm.pred_set.POW_EQNS|split|1','_i0_ax')).
fof(f3518,plain,(
  s(bool,'const.bool.IN_2'(s(fun(sK7,bool),sK10('const.bool.LET_2'(s(fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool)),'_dst_x0x1x1_3'(s(fun(fun(fun(sK7,bool),bool),fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool))),'const.pred_set.UNION_0'),s(fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool)),'const.pred_set.IMAGE_1'(s(fun(fun(sK7,bool),fun(sK7,bool)),'const.pred_set.INSERT_1'(s(sK7,sK8))))),s(fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool)),'_i_0_0'))),s(fun(fun(sK7,bool),bool),'const.pred_set.POW_1'(s(fun(sK7,bool),sK9)))),'const.pred_set.POW_1'(s(fun(sK7,bool),'const.pred_set.INSERT_2'(s(sK7,sK8),s(fun(sK7,bool),sK9)))),fun(sK7,bool))),s(fun(fun(sK7,bool),bool),'const.pred_set.POW_1'(s(fun(sK7,bool),'const.pred_set.INSERT_2'(s(sK7,sK8),s(fun(sK7,bool),sK9))))))) != s(bool,'const.bool.IN_2'(s(fun(sK7,bool),sK10('const.bool.LET_2'(s(fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool)),'_dst_x0x1x1_3'(s(fun(fun(fun(sK7,bool),bool),fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool))),'const.pred_set.UNION_0'),s(fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool)),'const.pred_set.IMAGE_1'(s(fun(fun(sK7,bool),fun(sK7,bool)),'const.pred_set.INSERT_1'(s(sK7,sK8))))),s(fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool)),'_i_0_0'))),s(fun(fun(sK7,bool),bool),'const.pred_set.POW_1'(s(fun(sK7,bool),sK9)))),'const.pred_set.POW_1'(s(fun(sK7,bool),'const.pred_set.INSERT_2'(s(sK7,sK8),s(fun(sK7,bool),sK9)))),fun(sK7,bool))),s(fun(fun(sK7,bool),bool),'const.pred_set.UNION_2'(s(fun(fun(sK7,bool),bool),happ(s(fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool)),'const.pred_set.IMAGE_1'(s(fun(fun(sK7,bool),fun(sK7,bool)),'const.pred_set.INSERT_1'(s(sK7,sK8))))),s(fun(fun(sK7,bool),bool),'const.pred_set.POW_1'(s(fun(sK7,bool),sK9))))),s(fun(fun(sK7,bool),bool),happ(s(fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool)),'_i_0_0'),s(fun(fun(sK7,bool),bool),'const.pred_set.POW_1'(s(fun(sK7,bool),sK9)))))))))),
  inference(backward_demodulation,[],[f3332,f3159])).
fof(f3159,plain,(
  s(bool,'const.bool.IN_2'(s(fun(sK7,bool),sK10('const.bool.LET_2'(s(fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool)),'_dst_x0x1x1_3'(s(fun(fun(fun(sK7,bool),bool),fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool))),'const.pred_set.UNION_0'),s(fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool)),'const.pred_set.IMAGE_1'(s(fun(fun(sK7,bool),fun(sK7,bool)),'const.pred_set.INSERT_1'(s(sK7,sK8))))),s(fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool)),'_i_0_0'))),s(fun(fun(sK7,bool),bool),'const.pred_set.POW_1'(s(fun(sK7,bool),sK9)))),'const.pred_set.POW_1'(s(fun(sK7,bool),'const.pred_set.INSERT_2'(s(sK7,sK8),s(fun(sK7,bool),sK9)))),fun(sK7,bool))),s(fun(fun(sK7,bool),bool),'const.pred_set.POW_1'(s(fun(sK7,bool),'const.pred_set.INSERT_2'(s(sK7,sK8),s(fun(sK7,bool),sK9))))))) != s(bool,'const.bool.IN_2'(s(fun(sK7,bool),sK10('const.bool.LET_2'(s(fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool)),'_dst_x0x1x1_3'(s(fun(fun(fun(sK7,bool),bool),fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool))),'const.pred_set.UNION_0'),s(fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool)),'const.pred_set.IMAGE_1'(s(fun(fun(sK7,bool),fun(sK7,bool)),'const.pred_set.INSERT_1'(s(sK7,sK8))))),s(fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool)),'_i_0_0'))),s(fun(fun(sK7,bool),bool),'const.pred_set.POW_1'(s(fun(sK7,bool),sK9)))),'const.pred_set.POW_1'(s(fun(sK7,bool),'const.pred_set.INSERT_2'(s(sK7,sK8),s(fun(sK7,bool),sK9)))),fun(sK7,bool))),s(fun(fun(sK7,bool),bool),happ(s(fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool)),'_dst_x0x1x1_3'(s(fun(fun(fun(sK7,bool),bool),fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool))),'const.pred_set.UNION_0'),s(fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool)),'const.pred_set.IMAGE_1'(s(fun(fun(sK7,bool),fun(sK7,bool)),'const.pred_set.INSERT_1'(s(sK7,sK8))))),s(fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool)),'_i_0_0'))),s(fun(fun(sK7,bool),bool),'const.pred_set.POW_1'(s(fun(sK7,bool),sK9)))))))),
  inference(forward_demodulation,[],[f3146,f363])).
fof(f363,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,X3))) = s(X0,'const.bool.LET_2'(s(fun(X1,X0),X2),s(X1,X3)))) )),
  inference(cnf_transformation,[],[f133])).
fof(f133,plain,(
  ! [X0,X1,X2,X3] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X3))) = s(X0,'const.bool.LET_2'(s(fun(X1,X0),X2),s(X1,X3)))),
  inference(flattening,[],[f132])).
fof(f132,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X3))) = s(X0,'const.bool.LET_2'(s(fun(X1,X0),X2),s(X1,X3)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X12,X6] : ! [X13,X14] : s(X12,happ(s(fun(X6,X12),X13),s(X6,X14))) = s(X12,'const.bool.LET_2'(s(fun(X6,X12),X13),s(X6,X14)))),
  file('thm.pred_set.POW_EQNS|split|1','thm.bool.LET_THM')).
fof(f3146,plain,(
  s(bool,'const.bool.IN_2'(s(fun(sK7,bool),sK10('const.bool.LET_2'(s(fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool)),'_dst_x0x1x1_3'(s(fun(fun(fun(sK7,bool),bool),fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool))),'const.pred_set.UNION_0'),s(fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool)),'const.pred_set.IMAGE_1'(s(fun(fun(sK7,bool),fun(sK7,bool)),'const.pred_set.INSERT_1'(s(sK7,sK8))))),s(fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool)),'_i_0_0'))),s(fun(fun(sK7,bool),bool),'const.pred_set.POW_1'(s(fun(sK7,bool),sK9)))),'const.pred_set.POW_1'(s(fun(sK7,bool),'const.pred_set.INSERT_2'(s(sK7,sK8),s(fun(sK7,bool),sK9)))),fun(sK7,bool))),s(fun(fun(sK7,bool),bool),'const.pred_set.POW_1'(s(fun(sK7,bool),'const.pred_set.INSERT_2'(s(sK7,sK8),s(fun(sK7,bool),sK9))))))) != s(bool,'const.bool.IN_2'(s(fun(sK7,bool),sK10('const.bool.LET_2'(s(fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool)),'_dst_x0x1x1_3'(s(fun(fun(fun(sK7,bool),bool),fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool))),'const.pred_set.UNION_0'),s(fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool)),'const.pred_set.IMAGE_1'(s(fun(fun(sK7,bool),fun(sK7,bool)),'const.pred_set.INSERT_1'(s(sK7,sK8))))),s(fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool)),'_i_0_0'))),s(fun(fun(sK7,bool),bool),'const.pred_set.POW_1'(s(fun(sK7,bool),sK9)))),'const.pred_set.POW_1'(s(fun(sK7,bool),'const.pred_set.INSERT_2'(s(sK7,sK8),s(fun(sK7,bool),sK9)))),fun(sK7,bool))),s(fun(fun(sK7,bool),bool),'const.bool.LET_2'(s(fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool)),'_dst_x0x1x1_3'(s(fun(fun(fun(sK7,bool),bool),fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool))),'const.pred_set.UNION_0'),s(fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool)),'const.pred_set.IMAGE_1'(s(fun(fun(sK7,bool),fun(sK7,bool)),'const.pred_set.INSERT_1'(s(sK7,sK8))))),s(fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool)),'_i_0_0'))),s(fun(fun(sK7,bool),bool),'const.pred_set.POW_1'(s(fun(sK7,bool),sK9)))))))),
  inference(unit_resulting_resolution,[],[f423,f336])).
fof(f336,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.bool.IN_2'(s(X0,sK10(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK10(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) )),
  inference(cnf_transformation,[],[f226])).
fof(f226,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (s(bool,'const.bool.IN_2'(s(X0,sK10(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK10(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK10])],[f225])).
fof(f225,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(rectify,[],[f224])).
fof(f224,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(nnf_transformation,[],[f117])).
fof(f117,plain,(
  ! [X0,X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(flattening,[],[f116])).
fof(f116,plain,(
  ! [X0] : ! [X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X6] : ! [X8,X11] : (s(fun(X6,bool),X8) = s(fun(X6,bool),X11) <=> ! [X14] : s(bool,'const.bool.IN_2'(s(X6,X14),s(fun(X6,bool),X8))) = s(bool,'const.bool.IN_2'(s(X6,X14),s(fun(X6,bool),X11))))),
  file('thm.pred_set.POW_EQNS|split|1','thm.pred_set.EXTENSION')).
fof(f423,plain,(
  s(fun(fun(sK7,bool),bool),'const.pred_set.POW_1'(s(fun(sK7,bool),'const.pred_set.INSERT_2'(s(sK7,sK8),s(fun(sK7,bool),sK9))))) != s(fun(fun(sK7,bool),bool),'const.bool.LET_2'(s(fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool)),'_dst_x0x1x1_3'(s(fun(fun(fun(sK7,bool),bool),fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool))),'const.pred_set.UNION_0'),s(fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool)),'const.pred_set.IMAGE_1'(s(fun(fun(sK7,bool),fun(sK7,bool)),'const.pred_set.INSERT_1'(s(sK7,sK8))))),s(fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool)),'_i_0_0'))),s(fun(fun(sK7,bool),bool),'const.pred_set.POW_1'(s(fun(sK7,bool),sK9)))))),
  inference(backward_demodulation,[],[f329,f261])).
fof(f261,plain,(
  s(fun(fun(sK7,bool),bool),'const.pred_set.POW_1'(s(fun(sK7,bool),happ(s(fun(fun(sK7,bool),fun(sK7,bool)),'const.pred_set.INSERT_1'(s(sK7,sK8))),s(fun(sK7,bool),sK9))))) != s(fun(fun(sK7,bool),bool),'const.bool.LET_2'(s(fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool)),'_dst_x0x1x1_3'(s(fun(fun(fun(sK7,bool),bool),fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool))),'const.pred_set.UNION_0'),s(fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool)),'const.pred_set.IMAGE_1'(s(fun(fun(sK7,bool),fun(sK7,bool)),'const.pred_set.INSERT_1'(s(sK7,sK8))))),s(fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool)),'_i_0_0'))),s(fun(fun(sK7,bool),bool),'const.pred_set.POW_1'(s(fun(sK7,bool),sK9)))))),
  inference(cnf_transformation,[],[f190])).
fof(f190,plain,(
  s(fun(fun(sK7,bool),bool),'const.pred_set.POW_1'(s(fun(sK7,bool),happ(s(fun(fun(sK7,bool),fun(sK7,bool)),'const.pred_set.INSERT_1'(s(sK7,sK8))),s(fun(sK7,bool),sK9))))) != s(fun(fun(sK7,bool),bool),'const.bool.LET_2'(s(fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool)),'_dst_x0x1x1_3'(s(fun(fun(fun(sK7,bool),bool),fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool))),'const.pred_set.UNION_0'),s(fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool)),'const.pred_set.IMAGE_1'(s(fun(fun(sK7,bool),fun(sK7,bool)),'const.pred_set.INSERT_1'(s(sK7,sK8))))),s(fun(fun(fun(sK7,bool),bool),fun(fun(sK7,bool),bool)),'_i_0_0'))),s(fun(fun(sK7,bool),bool),'const.pred_set.POW_1'(s(fun(sK7,bool),sK9)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK7,sK8,sK9])],[f152])).
fof(f152,plain,(
  ? [X0,X1,X2] : s(fun(fun(X0,bool),bool),'const.pred_set.POW_1'(s(fun(X0,bool),happ(s(fun(fun(X0,bool),fun(X0,bool)),'const.pred_set.INSERT_1'(s(X0,X1))),s(fun(X0,bool),X2))))) != s(fun(fun(X0,bool),bool),'const.bool.LET_2'(s(fun(fun(fun(X0,bool),bool),fun(fun(X0,bool),bool)),'_dst_x0x1x1_3'(s(fun(fun(fun(X0,bool),bool),fun(fun(fun(X0,bool),bool),fun(fun(X0,bool),bool))),'const.pred_set.UNION_0'),s(fun(fun(fun(X0,bool),bool),fun(fun(X0,bool),bool)),'const.pred_set.IMAGE_1'(s(fun(fun(X0,bool),fun(X0,bool)),'const.pred_set.INSERT_1'(s(X0,X1))))),s(fun(fun(fun(X0,bool),bool),fun(fun(X0,bool),bool)),'_i_0_0'))),s(fun(fun(X0,bool),bool),'const.pred_set.POW_1'(s(fun(X0,bool),X2)))))),
  inference(ennf_transformation,[],[f61])).
fof(f61,plain,(
  ~! [X0,X1,X2] : s(fun(fun(X0,bool),bool),'const.pred_set.POW_1'(s(fun(X0,bool),happ(s(fun(fun(X0,bool),fun(X0,bool)),'const.pred_set.INSERT_1'(s(X0,X1))),s(fun(X0,bool),X2))))) = s(fun(fun(X0,bool),bool),'const.bool.LET_2'(s(fun(fun(fun(X0,bool),bool),fun(fun(X0,bool),bool)),'_dst_x0x1x1_3'(s(fun(fun(fun(X0,bool),bool),fun(fun(fun(X0,bool),bool),fun(fun(X0,bool),bool))),'const.pred_set.UNION_0'),s(fun(fun(fun(X0,bool),bool),fun(fun(X0,bool),bool)),'const.pred_set.IMAGE_1'(s(fun(fun(X0,bool),fun(X0,bool)),'const.pred_set.INSERT_1'(s(X0,X1))))),s(fun(fun(fun(X0,bool),bool),fun(fun(X0,bool),bool)),'_i_0_0'))),s(fun(fun(X0,bool),bool),'const.pred_set.POW_1'(s(fun(X0,bool),X2)))))),
  inference(flattening,[],[f60])).
fof(f60,plain,(
  ~! [X0] : ! [X1,X2] : s(fun(fun(X0,bool),bool),'const.pred_set.POW_1'(s(fun(X0,bool),happ(s(fun(fun(X0,bool),fun(X0,bool)),'const.pred_set.INSERT_1'(s(X0,X1))),s(fun(X0,bool),X2))))) = s(fun(fun(X0,bool),bool),'const.bool.LET_2'(s(fun(fun(fun(X0,bool),bool),fun(fun(X0,bool),bool)),'_dst_x0x1x1_3'(s(fun(fun(fun(X0,bool),bool),fun(fun(fun(X0,bool),bool),fun(fun(X0,bool),bool))),'const.pred_set.UNION_0'),s(fun(fun(fun(X0,bool),bool),fun(fun(X0,bool),bool)),'const.pred_set.IMAGE_1'(s(fun(fun(X0,bool),fun(X0,bool)),'const.pred_set.INSERT_1'(s(X0,X1))))),s(fun(fun(fun(X0,bool),bool),fun(fun(X0,bool),bool)),'_i_0_0'))),s(fun(fun(X0,bool),bool),'const.pred_set.POW_1'(s(fun(X0,bool),X2)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : s(fun(fun(X6,bool),bool),'const.pred_set.POW_1'(s(fun(X6,bool),happ(s(fun(fun(X6,bool),fun(X6,bool)),'const.pred_set.INSERT_1'(s(X6,X7))),s(fun(X6,bool),X8))))) = s(fun(fun(X6,bool),bool),'const.bool.LET_2'(s(fun(fun(fun(X6,bool),bool),fun(fun(X6,bool),bool)),'_dst_x0x1x1_3'(s(fun(fun(fun(X6,bool),bool),fun(fun(fun(X6,bool),bool),fun(fun(X6,bool),bool))),'const.pred_set.UNION_0'),s(fun(fun(fun(X6,bool),bool),fun(fun(X6,bool),bool)),'const.pred_set.IMAGE_1'(s(fun(fun(X6,bool),fun(X6,bool)),'const.pred_set.INSERT_1'(s(X6,X7))))),s(fun(fun(fun(X6,bool),bool),fun(fun(X6,bool),bool)),'_i_0_0'))),s(fun(fun(X6,bool),bool),'const.pred_set.POW_1'(s(fun(X6,bool),X8)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : s(fun(fun(X6,bool),bool),'const.pred_set.POW_1'(s(fun(X6,bool),happ(s(fun(fun(X6,bool),fun(X6,bool)),'const.pred_set.INSERT_1'(s(X6,X7))),s(fun(X6,bool),X8))))) = s(fun(fun(X6,bool),bool),'const.bool.LET_2'(s(fun(fun(fun(X6,bool),bool),fun(fun(X6,bool),bool)),'_dst_x0x1x1_3'(s(fun(fun(fun(X6,bool),bool),fun(fun(fun(X6,bool),bool),fun(fun(X6,bool),bool))),'const.pred_set.UNION_0'),s(fun(fun(fun(X6,bool),bool),fun(fun(X6,bool),bool)),'const.pred_set.IMAGE_1'(s(fun(fun(X6,bool),fun(X6,bool)),'const.pred_set.INSERT_1'(s(X6,X7))))),s(fun(fun(fun(X6,bool),bool),fun(fun(X6,bool),bool)),'_i_0_0'))),s(fun(fun(X6,bool),bool),'const.pred_set.POW_1'(s(fun(X6,bool),X8)))))),
  file('thm.pred_set.POW_EQNS|split|1','thm.pred_set.POW_EQNS|split|1')).
fof(f3332,plain,(
  ( ! [X30,X28,X31,X29,X32] : (s(fun(X28,bool),'const.pred_set.UNION_2'(s(fun(X28,bool),happ(s(fun(X29,fun(X28,bool)),X30),s(X29,X31))),s(fun(X28,bool),happ(s(fun(X29,fun(X28,bool)),X32),s(X29,X31))))) = s(fun(X28,bool),happ(s(fun(X29,fun(X28,bool)),'_dst_x0x1x1_3'(s(fun(fun(X28,bool),fun(fun(X28,bool),fun(X28,bool))),'const.pred_set.UNION_0'),s(fun(X29,fun(X28,bool)),X30),s(fun(X29,fun(X28,bool)),X32))),s(X29,X31)))) )),
  inference(superposition,[],[f413,f331])).
fof(f331,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) = s(fun(X0,bool),happ(s(fun(fun(X0,bool),fun(X0,bool)),happ(s(fun(fun(X0,bool),fun(fun(X0,bool),fun(X0,bool))),'const.pred_set.UNION_0'),s(fun(X0,bool),X1))),s(fun(X0,bool),X2)))) )),
  inference(cnf_transformation,[],[f111])).
fof(f111,plain,(
  ! [X0,X1,X2] : s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) = s(fun(X0,bool),happ(s(fun(fun(X0,bool),fun(X0,bool)),happ(s(fun(fun(X0,bool),fun(fun(X0,bool),fun(X0,bool))),'const.pred_set.UNION_0'),s(fun(X0,bool),X1))),s(fun(X0,bool),X2)))),
  inference(flattening,[],[f110])).
fof(f110,plain,(
  ! [X0] : ! [X1,X2] : s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) = s(fun(X0,bool),happ(s(fun(fun(X0,bool),fun(X0,bool)),happ(s(fun(fun(X0,bool),fun(fun(X0,bool),fun(X0,bool))),'const.pred_set.UNION_0'),s(fun(X0,bool),X1))),s(fun(X0,bool),X2)))),
  inference(rectify,[],[f53])).
fof(f53,axiom,(
  ! [X6] : ! [X23,X24] : s(fun(X6,bool),'const.pred_set.UNION_2'(s(fun(X6,bool),X23),s(fun(X6,bool),X24))) = s(fun(X6,bool),happ(s(fun(fun(X6,bool),fun(X6,bool)),happ(s(fun(fun(X6,bool),fun(fun(X6,bool),fun(X6,bool))),'const.pred_set.UNION_0'),s(fun(X6,bool),X23))),s(fun(X6,bool),X24)))),
  file('thm.pred_set.POW_EQNS|split|1','ARITY_0_2_const.pred_set.UNION')).
fof(f413,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun(X3,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X2,X0)),X4),s(fun(X3,X1),X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,happ(s(fun(X3,X1),X5),s(X3,X7))))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))) )),
  inference(cnf_transformation,[],[f151])).
fof(f151,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun(X3,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X2,X0)),X4),s(fun(X3,X1),X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,happ(s(fun(X3,X1),X5),s(X3,X7))))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))),
  inference(flattening,[],[f150])).
fof(f150,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6,X7] : s(X0,happ(s(fun(X3,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X2,X0)),X4),s(fun(X3,X1),X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,happ(s(fun(X3,X1),X5),s(X3,X7))))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))),
  inference(rectify,[],[f55])).
fof(f55,axiom,(
  ! [X27,X28,X29,X25] : ! [X30,X26,X31,X32] : s(X27,happ(s(fun(X25,X27),'_dst_x0x1x1_3'(s(fun(X28,fun(X29,X27)),X30),s(fun(X25,X28),X26),s(fun(X25,X29),X31))),s(X25,X32))) = s(X27,happ(s(fun(X29,X27),happ(s(fun(X28,fun(X29,X27)),X30),s(X28,happ(s(fun(X25,X28),X26),s(X25,X32))))),s(X29,happ(s(fun(X25,X29),X31),s(X25,X32)))))),
  file('thm.pred_set.POW_EQNS|split|1','_dst_x0x1x1_ax')).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 179
% Passive clauses: 548
% Generated clauses: 3264
% Final active clauses: 177
% Final passive clauses: 363
% Input formulas: 58
% Initial clauses: 153
% 
% Duplicate literals: 4
% Trivial inequalities: 3
% Fw subsumption resolutions: 7
% Fw demodulations: 731
% Bw demodulations: 8
% Condensations: 5
% 
% Simple tautologies: 48
% Equational tautologies: 19
% Forward subsumptions: 1438
% Fw demodulations to eq. taut.: 358
% 
% Unit resulting resolution: 1170
% Factoring: 1
% Forward superposition: 443
% Backward superposition: 697
% Self superposition: 32
% Equality factoring: 2
% Equality resolution: 8
% 
% SAT solver clauses: 550
% SAT solver unit clauses: 287
% SAT solver binary clauses: 153
% 
% Memory used [KB]: 6908
% Time elapsed: 0.173 s
% ------------------------------
% ------------------------------
% Success in time 1.514 s
