Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.extreal
% (49404)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 199
% Passive clauses: 3130
% Generated clauses: 20330
% Final active clauses: 179
% Final passive clauses: 2748
% Input formulas: 37
% Initial clauses: 92
% 
% Duplicate literals: 904
% Trivial inequalities: 395
% Fw subsumption resolutions: 2402
% Fw demodulations: 2078
% Bw demodulations: 19
% 
% Simple tautologies: 460
% Equational tautologies: 46
% Forward subsumptions: 6658
% 
% Binary resolution: 1017
% Forward superposition: 5473
% Backward superposition: 7384
% Self superposition: 160
% Equality factoring: 123
% Equality resolution: 51
% 
% Split clauses: 692
% Split components: 1399
% SAT solver clauses: 1213
% SAT solver unit clauses: 16
% SAT solver binary clauses: 132
% 
% Sat splits: 691
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 1209
% TWLsolver calls for satisfiability: 217
% 
% Memory used [KB]: 13048
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.extreal
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.extreal
% SZS output start Proof for thm.extreal
fof(f1524,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f1523,f149,f259,f153,f267,f154,f268,f209,f286,f1504,f1506,f150,f264,f297,f300,f147,f262,f157,f270,f1481,f1482,f1299,f1354,f210,f288,f1176,f1201,f672,f705,f222,f293,f1453,f1454,f148,f260,f213,f289])).
fof(f289,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.extreal.extreal',happ(s(fun(X0,'type.extreal.extreal'),X2),s(X0,X3))) != s('type.extreal.extreal',happ(s(fun(X0,'type.extreal.extreal'),X2),s(X0,sK6(X3,X2,X1,X0)))) | s('type.extreal.extreal','const.extreal.EXTREAL_SUM_IMAGE_2'(s(fun(X0,'type.extreal.extreal'),X2),s(fun(X0,bool),X1))) = s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))))),s('type.extreal.extreal',happ(s(fun(X0,'type.extreal.extreal'),X2),s(X0,X3))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1))))) ) | $spl30),
  inference(cnf_transformation,[],[f289_D])).
fof(f289_D,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.extreal.extreal',happ(s(fun(X0,'type.extreal.extreal'),X2),s(X0,X3))) != s('type.extreal.extreal',happ(s(fun(X0,'type.extreal.extreal'),X2),s(X0,sK6(X3,X2,X1,X0)))) | s('type.extreal.extreal','const.extreal.EXTREAL_SUM_IMAGE_2'(s(fun(X0,'type.extreal.extreal'),X2),s(fun(X0,bool),X1))) = s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))))),s('type.extreal.extreal',happ(s(fun(X0,'type.extreal.extreal'),X2),s(X0,X3))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1))))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f213,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.extreal.extreal','const.extreal.EXTREAL_SUM_IMAGE_2'(s(fun(X0,'type.extreal.extreal'),X2),s(fun(X0,bool),X1))) = s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))))),s('type.extreal.extreal',happ(s(fun(X0,'type.extreal.extreal'),X2),s(X0,X3))))) | s('type.extreal.extreal',happ(s(fun(X0,'type.extreal.extreal'),X2),s(X0,X3))) != s('type.extreal.extreal',happ(s(fun(X0,'type.extreal.extreal'),X2),s(X0,sK6(X3,X2,X1,X0)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f135])).
fof(f135,plain,(
  ! [X0,X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) | ! [X2,X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | (p(s(bool,'const.bool.IN_2'(s(X0,sK6(X3,X2,X1,X0)),s(fun(X0,bool),X1)))) & s('type.extreal.extreal',happ(s(fun(X0,'type.extreal.extreal'),X2),s(X0,X3))) != s('type.extreal.extreal',happ(s(fun(X0,'type.extreal.extreal'),X2),s(X0,sK6(X3,X2,X1,X0))))) | s('type.extreal.extreal','const.extreal.EXTREAL_SUM_IMAGE_2'(s(fun(X0,'type.extreal.extreal'),X2),s(fun(X0,bool),X1))) = s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))))),s('type.extreal.extreal',happ(s(fun(X0,'type.extreal.extreal'),X2),s(X0,X3)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK6])],[f92])).
fof(f92,plain,(
  ! [X0,X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) | ! [X2,X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ? [X4] : (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) & s('type.extreal.extreal',happ(s(fun(X0,'type.extreal.extreal'),X2),s(X0,X3))) != s('type.extreal.extreal',happ(s(fun(X0,'type.extreal.extreal'),X2),s(X0,X4)))) | s('type.extreal.extreal','const.extreal.EXTREAL_SUM_IMAGE_2'(s(fun(X0,'type.extreal.extreal'),X2),s(fun(X0,bool),X1))) = s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))))),s('type.extreal.extreal',happ(s(fun(X0,'type.extreal.extreal'),X2),s(X0,X3)))))))),
  inference(flattening,[],[f91])).
fof(f91,plain,(
  ! [X0,X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) | ! [X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ? [X4] : (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) & s('type.extreal.extreal',happ(s(fun(X0,'type.extreal.extreal'),X2),s(X0,X3))) != s('type.extreal.extreal',happ(s(fun(X0,'type.extreal.extreal'),X2),s(X0,X4))))) | s('type.extreal.extreal','const.extreal.EXTREAL_SUM_IMAGE_2'(s(fun(X0,'type.extreal.extreal'),X2),s(fun(X0,bool),X1))) = s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))))),s('type.extreal.extreal',happ(s(fun(X0,'type.extreal.extreal'),X2),s(X0,X3)))))))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0,X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) => ! [X2,X3] : ((p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ! [X4] : (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) => s('type.extreal.extreal',happ(s(fun(X0,'type.extreal.extreal'),X2),s(X0,X3))) = s('type.extreal.extreal',happ(s(fun(X0,'type.extreal.extreal'),X2),s(X0,X4))))) => s('type.extreal.extreal','const.extreal.EXTREAL_SUM_IMAGE_2'(s(fun(X0,'type.extreal.extreal'),X2),s(fun(X0,bool),X1))) = s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))))),s('type.extreal.extreal',happ(s(fun(X0,'type.extreal.extreal'),X2),s(X0,X3)))))))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) => ! [X2,X3] : ((p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ! [X4] : (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) => s('type.extreal.extreal',happ(s(fun(X0,'type.extreal.extreal'),X2),s(X0,X3))) = s('type.extreal.extreal',happ(s(fun(X0,'type.extreal.extreal'),X2),s(X0,X4))))) => s('type.extreal.extreal','const.extreal.EXTREAL_SUM_IMAGE_2'(s(fun(X0,'type.extreal.extreal'),X2),s(fun(X0,bool),X1))) = s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))))),s('type.extreal.extreal',happ(s(fun(X0,'type.extreal.extreal'),X2),s(X0,X3)))))))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X6] : ! [X18] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X18)))) => ! [X8,X19] : ((p(s(bool,'const.bool.IN_2'(s(X6,X19),s(fun(X6,bool),X18)))) & ! [X20] : (p(s(bool,'const.bool.IN_2'(s(X6,X20),s(fun(X6,bool),X18)))) => s('type.extreal.extreal',happ(s(fun(X6,'type.extreal.extreal'),X8),s(X6,X19))) = s('type.extreal.extreal',happ(s(fun(X6,'type.extreal.extreal'),X8),s(X6,X20))))) => s('type.extreal.extreal','const.extreal.EXTREAL_SUM_IMAGE_2'(s(fun(X6,'type.extreal.extreal'),X8),s(fun(X6,bool),X18))) = s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X6,bool),X18))))),s('type.extreal.extreal',happ(s(fun(X6,'type.extreal.extreal'),X8),s(X6,X19)))))))),
  file('thm.extreal.EXTREAL_SUM_IMAGE_FINITE_CONST','thm.extreal.EXTREAL_SUM_IMAGE_FINITE_SAME')).
fof(f260,plain,(
  ( ! [X4] : (s('type.extreal.extreal',sK3) = s('type.extreal.extreal',happ(s(fun(sK0,'type.extreal.extreal'),sK2),s(sK0,X4)))) ) | $spl2),
  inference(cnf_transformation,[],[f260_D])).
fof(f260_D,plain,(
  ( ! [X4] : (s('type.extreal.extreal',sK3) = s('type.extreal.extreal',happ(s(fun(sK0,'type.extreal.extreal'),sK2),s(sK0,X4)))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f148,plain,(
  ( ! [X4] : (s('type.extreal.extreal',sK3) = s('type.extreal.extreal',happ(s(fun(sK0,'type.extreal.extreal'),sK2),s(sK0,X4)))) )),
  inference(cnf_transformation,[],[f99])).
fof(f99,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))) & (! [X4] : s('type.extreal.extreal',sK3) = s('type.extreal.extreal',happ(s(fun(sK0,'type.extreal.extreal'),sK2),s(sK0,X4))) & s('type.extreal.extreal','const.extreal.EXTREAL_SUM_IMAGE_2'(s(fun(sK0,'type.extreal.extreal'),sK2),s(fun(sK0,bool),sK1))) != s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.extreal.extreal',sK3))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f85])).
fof(f85,plain,(
  ? [X0,X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) & ? [X2,X3] : (! [X4] : s('type.extreal.extreal',X3) = s('type.extreal.extreal',happ(s(fun(X0,'type.extreal.extreal'),X2),s(X0,X4))) & s('type.extreal.extreal','const.extreal.EXTREAL_SUM_IMAGE_2'(s(fun(X0,'type.extreal.extreal'),X2),s(fun(X0,bool),X1))) != s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))))),s('type.extreal.extreal',X3)))))),
  inference(ennf_transformation,[],[f40])).
fof(f40,plain,(
  ~! [X0,X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) => ! [X2,X3] : (! [X4] : s('type.extreal.extreal',X3) = s('type.extreal.extreal',happ(s(fun(X0,'type.extreal.extreal'),X2),s(X0,X4))) => s('type.extreal.extreal','const.extreal.EXTREAL_SUM_IMAGE_2'(s(fun(X0,'type.extreal.extreal'),X2),s(fun(X0,bool),X1))) = s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))))),s('type.extreal.extreal',X3)))))),
  inference(flattening,[],[f39])).
fof(f39,plain,(
  ~! [X0] : ! [X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) => ! [X2,X3] : (! [X4] : s('type.extreal.extreal',X3) = s('type.extreal.extreal',happ(s(fun(X0,'type.extreal.extreal'),X2),s(X0,X4))) => s('type.extreal.extreal','const.extreal.EXTREAL_SUM_IMAGE_2'(s(fun(X0,'type.extreal.extreal'),X2),s(fun(X0,bool),X1))) = s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))))),s('type.extreal.extreal',X3)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X7)))) => ! [X8,X9] : (! [X10] : s('type.extreal.extreal',X9) = s('type.extreal.extreal',happ(s(fun(X6,'type.extreal.extreal'),X8),s(X6,X10))) => s('type.extreal.extreal','const.extreal.EXTREAL_SUM_IMAGE_2'(s(fun(X6,'type.extreal.extreal'),X8),s(fun(X6,bool),X7))) = s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X6,bool),X7))))),s('type.extreal.extreal',X9)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X7)))) => ! [X8,X9] : (! [X10] : s('type.extreal.extreal',X9) = s('type.extreal.extreal',happ(s(fun(X6,'type.extreal.extreal'),X8),s(X6,X10))) => s('type.extreal.extreal','const.extreal.EXTREAL_SUM_IMAGE_2'(s(fun(X6,'type.extreal.extreal'),X8),s(fun(X6,bool),X7))) = s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X6,bool),X7))))),s('type.extreal.extreal',X9)))))),
  file('thm.extreal.EXTREAL_SUM_IMAGE_FINITE_CONST','thm.extreal.EXTREAL_SUM_IMAGE_FINITE_CONST')).
fof(f1454,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),X1)))) | s('type.extreal.extreal','const.extreal.EXTREAL_SUM_IMAGE_2'(s(fun(sK0,'type.extreal.extreal'),sK2),s(fun(sK0,bool),X1))) = s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),X1))))),s('type.extreal.extreal',sK3))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X1))))) ) | $spl310),
  inference(cnf_transformation,[],[f1454_D])).
fof(f1454_D,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),X1)))) | s('type.extreal.extreal','const.extreal.EXTREAL_SUM_IMAGE_2'(s(fun(sK0,'type.extreal.extreal'),sK2),s(fun(sK0,bool),X1))) = s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),X1))))),s('type.extreal.extreal',sK3))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X1))))) ) <=> ~$spl310),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl310])])).
fof(f1453,plain,(
  ( ! [X0,X1] : (s('type.extreal.extreal','const.extreal.EXTREAL_SUM_IMAGE_2'(s(fun(sK0,'type.extreal.extreal'),sK2),s(fun(sK0,bool),X1))) = s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),X1))))),s('type.extreal.extreal',sK3))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),X1)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X1))))) ) | ($spl2 | $spl30)),
  inference(forward_demodulation,[],[f1452,f260])).
fof(f1452,plain,(
  ( ! [X0,X1] : (s('type.extreal.extreal','const.extreal.EXTREAL_SUM_IMAGE_2'(s(fun(sK0,'type.extreal.extreal'),sK2),s(fun(sK0,bool),X1))) = s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),X1))))),s('type.extreal.extreal',happ(s(fun(sK0,'type.extreal.extreal'),sK2),s(sK0,X0))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),X1)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X1))))) ) | ($spl2 | $spl30)),
  inference(subsumption_resolution,[],[f1448,f260])).
fof(f1448,plain,(
  ( ! [X0,X1] : (s('type.extreal.extreal',sK3) != s('type.extreal.extreal',happ(s(fun(sK0,'type.extreal.extreal'),sK2),s(sK0,X0))) | s('type.extreal.extreal','const.extreal.EXTREAL_SUM_IMAGE_2'(s(fun(sK0,'type.extreal.extreal'),sK2),s(fun(sK0,bool),X1))) = s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),X1))))),s('type.extreal.extreal',happ(s(fun(sK0,'type.extreal.extreal'),sK2),s(sK0,X0))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),X1)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X1))))) ) | ($spl2 | $spl30)),
  inference(superposition,[],[f289,f260])).
fof(f293,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) | s(X0,X1) != s(X0,X2)) ) | $spl38),
  inference(cnf_transformation,[],[f293_D])).
fof(f293_D,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) | s(X0,X1) != s(X0,X2)) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f222,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) | s(X0,X1) != s(X0,X2)) )),
  inference(cnf_transformation,[],[f141])).
fof(f141,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) | s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))) & ((s(X0,X1) != s(X0,X2) & ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3))))))))),
  inference(flattening,[],[f140])).
fof(f140,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) | (s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3)))))) & ((s(X0,X1) != s(X0,X2) & ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3))))))))),
  inference(nnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) <=> (s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))))),
  inference(flattening,[],[f79])).
fof(f79,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) <=> (s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X6] : ! [X9,X10,X18] : (p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),'const.pred_set.INSERT_2'(s(X6,X10),s(fun(X6,bool),X18)))))) <=> (s(X6,X9) = s(X6,X10) | p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X18))))))),
  file('thm.extreal.EXTREAL_SUM_IMAGE_FINITE_CONST','thm.pred_set.IN_INSERT')).
fof(f705,plain,(
  ( ! [X6,X8,X7,X5] : (s(X5,X6) != s(X5,X7) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X5,X6),s(fun(X5,bool),'const.pred_set.INSERT_2'(s(X5,X7),s(fun(X5,bool),X8)))))) ) | $spl86),
  inference(cnf_transformation,[],[f705_D])).
fof(f705_D,plain,(
  ( ! [X6,X8,X7,X5] : (s(X5,X6) != s(X5,X7) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X5,X6),s(fun(X5,bool),'const.pred_set.INSERT_2'(s(X5,X7),s(fun(X5,bool),X8)))))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f672,plain,(
  ( ! [X6,X8,X7,X5] : (s(X5,X6) != s(X5,X7) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X5,X6),s(fun(X5,bool),'const.pred_set.INSERT_2'(s(X5,X7),s(fun(X5,bool),X8)))))) ) | ($spl16 | $spl38)),
  inference(resolution,[],[f293,f270])).
fof(f1201,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))))) ) | $spl240),
  inference(cnf_transformation,[],[f1201_D])).
fof(f1201_D,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))))) ) <=> ~$spl240),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl240])])).
fof(f1176,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))))) ) | $spl86),
  inference(equality_resolution,[],[f705])).
fof(f288,plain,(
  ( ! [X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,sK4(X1,X0)),s(fun(X0,bool),sK5(X1,X0)))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1)) ) | $spl28),
  inference(cnf_transformation,[],[f288_D])).
fof(f288_D,plain,(
  ( ! [X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,sK4(X1,X0)),s(fun(X0,bool),sK5(X1,X0)))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1)) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f210,plain,(
  ( ! [X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,sK4(X1,X0)),s(fun(X0,bool),sK5(X1,X0)))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1)) )),
  inference(cnf_transformation,[],[f134])).
fof(f134,plain,(
  ! [X0,X1] : (s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1) | (s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,sK4(X1,X0)),s(fun(X0,bool),sK5(X1,X0)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,sK4(X1,X0)),s(fun(X0,bool),sK5(X1,X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4,sK5])],[f69])).
fof(f69,plain,(
  ! [X0,X1] : (s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1) | ? [X2,X3] : (s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X3))))))),
  inference(flattening,[],[f68])).
fof(f68,plain,(
  ! [X0] : ! [X1] : (s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1) | ? [X2,X3] : (s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X6] : ! [X18] : (s(fun(X6,bool),'const.pred_set.EMPTY_0') = s(fun(X6,bool),X18) | ? [X9,X11] : (s(fun(X6,bool),X18) = s(fun(X6,bool),'const.pred_set.INSERT_2'(s(X6,X9),s(fun(X6,bool),X11))) & ~p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X11))))))),
  file('thm.extreal.EXTREAL_SUM_IMAGE_FINITE_CONST','thm.pred_set.SET_CASES')).
fof(f1354,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,sK4(X1,X0)),s(fun(X0,bool),X1))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1)) ) | $spl272),
  inference(cnf_transformation,[],[f1354_D])).
fof(f1354_D,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,sK4(X1,X0)),s(fun(X0,bool),X1))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1)) ) <=> ~$spl272),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl272])])).
fof(f1299,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,sK4(X1,X0)),s(fun(X0,bool),X1))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1)) ) | ($spl28 | $spl240)),
  inference(superposition,[],[f1201,f288])).
fof(f1482,plain,(
  ( ! [X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X1)))) | s('type.extreal.extreal','const.extreal.EXTREAL_SUM_IMAGE_2'(s(fun(sK0,'type.extreal.extreal'),sK2),s(fun(sK0,bool),X1))) = s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),X1))))),s('type.extreal.extreal',sK3))) | s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),X1)) ) | $spl328),
  inference(cnf_transformation,[],[f1482_D])).
fof(f1482_D,plain,(
  ( ! [X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X1)))) | s('type.extreal.extreal','const.extreal.EXTREAL_SUM_IMAGE_2'(s(fun(sK0,'type.extreal.extreal'),sK2),s(fun(sK0,bool),X1))) = s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),X1))))),s('type.extreal.extreal',sK3))) | s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),X1)) ) <=> ~$spl328),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl328])])).
fof(f1481,plain,(
  ( ! [X1] : (s('type.extreal.extreal','const.extreal.EXTREAL_SUM_IMAGE_2'(s(fun(sK0,'type.extreal.extreal'),sK2),s(fun(sK0,bool),X1))) = s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),X1))))),s('type.extreal.extreal',sK3))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X1)))) | s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),X1)) ) | ($spl6 | $spl272 | $spl310)),
  inference(subsumption_resolution,[],[f1465,f264])).
fof(f1465,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | s('type.extreal.extreal','const.extreal.EXTREAL_SUM_IMAGE_2'(s(fun(sK0,'type.extreal.extreal'),sK2),s(fun(sK0,bool),X1))) = s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),X1))))),s('type.extreal.extreal',sK3))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X1)))) | s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),X1)) ) | ($spl272 | $spl310)),
  inference(superposition,[],[f1454,f1354])).
fof(f270,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl16),
  inference(cnf_transformation,[],[f270_D])).
fof(f270_D,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f157,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f100])).
fof(f100,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f44])).
fof(f44,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X11] : (s(bool,'T_0') = s(bool,X11) <=> p(s(bool,X11)))),
  file('thm.extreal.EXTREAL_SUM_IMAGE_FINITE_CONST','thm.bool.EQ_CLAUSES|split|1')).
fof(f262,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))) | $spl4),
  inference(cnf_transformation,[],[f262_D])).
fof(f262_D,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f147,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1))))),
  inference(cnf_transformation,[],[f99])).
fof(f300,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1))) | $spl46),
  inference(cnf_transformation,[],[f300_D])).
fof(f300_D,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1))) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f297,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1))) | ($spl4 | $spl16)),
  inference(resolution,[],[f270,f262])).
fof(f264,plain,(
  p(s(bool,'T_0')) | $spl6),
  inference(cnf_transformation,[],[f264_D])).
fof(f264_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f150,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.extreal.EXTREAL_SUM_IMAGE_FINITE_CONST',aHL_TRUTH)).
fof(f1506,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),sK1) | $spl342),
  inference(cnf_transformation,[],[f1506_D])).
fof(f1506_D,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),sK1) <=> ~$spl342),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl342])])).
fof(f1504,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),sK1) | ($spl1 | $spl6 | $spl46 | $spl328)),
  inference(subsumption_resolution,[],[f1503,f259])).
fof(f1503,plain,(
  s('type.extreal.extreal','const.extreal.EXTREAL_SUM_IMAGE_2'(s(fun(sK0,'type.extreal.extreal'),sK2),s(fun(sK0,bool),sK1))) = s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.extreal.extreal',sK3))) | s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),sK1) | ($spl6 | $spl46 | $spl328)),
  inference(subsumption_resolution,[],[f1492,f264])).
fof(f1492,plain,(
  ~p(s(bool,'T_0')) | s('type.extreal.extreal','const.extreal.EXTREAL_SUM_IMAGE_2'(s(fun(sK0,'type.extreal.extreal'),sK2),s(fun(sK0,bool),sK1))) = s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.extreal.extreal',sK3))) | s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),sK1) | ($spl46 | $spl328)),
  inference(superposition,[],[f1482,f300])).
fof(f286,plain,(
  ( ! [X0,X1] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.EXTREAL_SUM_IMAGE_2'(s(fun(X0,'type.extreal.extreal'),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) ) | $spl24),
  inference(cnf_transformation,[],[f286_D])).
fof(f286_D,plain,(
  ( ! [X0,X1] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.EXTREAL_SUM_IMAGE_2'(s(fun(X0,'type.extreal.extreal'),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f209,plain,(
  ( ! [X0,X1] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.EXTREAL_SUM_IMAGE_2'(s(fun(X0,'type.extreal.extreal'),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(cnf_transformation,[],[f67])).
fof(f67,plain,(
  ! [X0,X1] : s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.EXTREAL_SUM_IMAGE_2'(s(fun(X0,'type.extreal.extreal'),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))),
  inference(flattening,[],[f66])).
fof(f66,plain,(
  ! [X0] : ! [X1] : s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.EXTREAL_SUM_IMAGE_2'(s(fun(X0,'type.extreal.extreal'),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X6] : ! [X8] : s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.EXTREAL_SUM_IMAGE_2'(s(fun(X6,'type.extreal.extreal'),X8),s(fun(X6,bool),'const.pred_set.EMPTY_0')))),
  file('thm.extreal.EXTREAL_SUM_IMAGE_FINITE_CONST','thm.extreal.EXTREAL_SUM_IMAGE_THM|split|0')).
fof(f268,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X0)))) ) | $spl12),
  inference(cnf_transformation,[],[f268_D])).
fof(f268_D,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X0)))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f154,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X0)))) )),
  inference(cnf_transformation,[],[f43])).
fof(f43,plain,(
  ! [X0] : s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X0)))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X9] : s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X9)))),
  file('thm.extreal.EXTREAL_SUM_IMAGE_FINITE_CONST','thm.extreal.mul_lzero')).
fof(f267,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.pred_set.EMPTY_0')))) ) | $spl10),
  inference(cnf_transformation,[],[f267_D])).
fof(f267_D,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.pred_set.EMPTY_0')))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f153,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(cnf_transformation,[],[f42])).
fof(f42,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.pred_set.EMPTY_0')))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X6] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(X6,bool),'const.pred_set.EMPTY_0')))),
  file('thm.extreal.EXTREAL_SUM_IMAGE_FINITE_CONST','thm.pred_set.CARD_DEF|split|0')).
fof(f259,plain,(
  s('type.extreal.extreal','const.extreal.EXTREAL_SUM_IMAGE_2'(s(fun(sK0,'type.extreal.extreal'),sK2),s(fun(sK0,bool),sK1))) != s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.extreal.extreal',sK3))) | $spl1),
  inference(cnf_transformation,[],[f259_D])).
fof(f259_D,plain,(
  s('type.extreal.extreal','const.extreal.EXTREAL_SUM_IMAGE_2'(s(fun(sK0,'type.extreal.extreal'),sK2),s(fun(sK0,bool),sK1))) != s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.extreal.extreal',sK3))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f149,plain,(
  s('type.extreal.extreal','const.extreal.EXTREAL_SUM_IMAGE_2'(s(fun(sK0,'type.extreal.extreal'),sK2),s(fun(sK0,bool),sK1))) != s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.extreal.extreal',sK3)))),
  inference(cnf_transformation,[],[f99])).
fof(f1523,plain,(
  $false | ($spl1 | $spl10 | $spl12 | $spl24 | $spl342)),
  inference(subsumption_resolution,[],[f1522,f268])).
fof(f1522,plain,(
  s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK3))) | ($spl1 | $spl10 | $spl24 | $spl342)),
  inference(forward_demodulation,[],[f1521,f286])).
fof(f1521,plain,(
  s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK3))) != s('type.extreal.extreal','const.extreal.EXTREAL_SUM_IMAGE_2'(s(fun(sK0,'type.extreal.extreal'),sK2),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ($spl1 | $spl10 | $spl342)),
  inference(forward_demodulation,[],[f1514,f267])).
fof(f1514,plain,(
  s('type.extreal.extreal','const.extreal.EXTREAL_SUM_IMAGE_2'(s(fun(sK0,'type.extreal.extreal'),sK2),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) != s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))),s('type.extreal.extreal',sK3))) | ($spl1 | $spl342)),
  inference(backward_demodulation,[],[f1506,f259])).
% SZS output end Proof for thm.extreal
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 46
% Passive clauses: 162
% Generated clauses: 1339
% Final active clauses: 39
% Final passive clauses: 116
% Input formulas: 37
% Initial clauses: 92
% 
% Duplicate literals: 87
% Trivial inequalities: 28
% Fw subsumption resolutions: 75
% Fw demodulations: 63
% Bw demodulations: 7
% Condensations: 120
% Global subsumptions: 12
% 
% Simple tautologies: 94
% Equational tautologies: 16
% Forward subsumptions: 396
% 
% Binary resolution: 53
% Forward superposition: 373
% Backward superposition: 199
% Self superposition: 24
% Equality factoring: 26
% Equality resolution: 10
% 
% Split clauses: 4
% Split components: 8
% SAT solver clauses: 392
% SAT solver unit clauses: 100
% SAT solver binary clauses: 32
% SAT solver learnt clauses: 25
% 
% Sat splits: 4
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 343
% TWLsolver calls for satisfiability: 331
% 
% Memory used [KB]: 2302
% Time elapsed: 0.060 s
% ------------------------------
% ------------------------------
% Success in time 0.681 s
