Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.integer
% (28072)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 752
% Passive clauses: 4011
% Generated clauses: 18087
% Final active clauses: 735
% Final passive clauses: 3067
% Input formulas: 261
% Initial clauses: 525
% 
% Duplicate literals: 152
% Trivial inequalities: 28
% Fw subsumption resolutions: 1126
% Fw demodulations: 1545
% Bw demodulations: 10
% 
% Simple tautologies: 291
% Equational tautologies: 21
% Forward subsumptions: 10718
% 
% Binary resolution: 4077
% Factoring: 346
% Forward superposition: 5015
% Backward superposition: 5082
% Self superposition: 33
% Equality factoring: 2
% Equality resolution: 7
% 
% Split clauses: 286
% Split components: 574
% SAT solver clauses: 379
% SAT solver unit clauses: 15
% SAT solver binary clauses: 62
% 
% Sat splits: 286
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 379
% TWLsolver calls for satisfiability: 264
% 
% Memory used [KB]: 9210
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.integer
% (28098)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 494
% Passive clauses: 1968
% Generated clauses: 18410
% Final active clauses: 332
% Final passive clauses: 770
% Input formulas: 261
% Initial clauses: 525
% 
% Duplicate literals: 90
% Trivial inequalities: 23
% Fw subsumption resolutions: 798
% Fw demodulations: 1630
% Bw demodulations: 802
% Condensations: 50
% Global subsumptions: 25
% 
% Simple tautologies: 277
% Equational tautologies: 5
% Forward subsumptions: 10304
% Fw demodulations to eq. taut.: 5
% Bw demodulations to eq. taut.: 6
% 
% Binary resolution: 1482
% Factoring: 114
% Forward superposition: 4470
% Backward superposition: 4900
% Self superposition: 74
% Equality factoring: 2
% Equality resolution: 4
% 
% Split clauses: 25
% Split components: 50
% SAT solver clauses: 2678
% SAT solver unit clauses: 647
% SAT solver binary clauses: 199
% SAT solver learnt clauses: 53
% SAT solver learnt literals: 18
% 
% Sat splits: 25
% Sat splitting refutations: 110
% 
% TWLsolver clauses: 2575
% TWLsolver calls for satisfiability: 3429
% 
% Memory used [KB]: 14839
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.integer
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.integer
% SZS output start Proof for thm.integer
fof(f23282,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f23281])).
fof(f23281,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK15))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK14))))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK15))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK14)))))))),
  inference(forward_demodulation,[],[f23275,f1462])).
fof(f1462,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))))) = s('type.integer.int','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num',X0)))))))))) )),
  inference(forward_demodulation,[],[f1126,f1085])).
fof(f1085,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X0),s('type.integer.int',X1))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))) )),
  inference(cnf_transformation,[],[f410])).
fof(f410,plain,(
  ! [X0,X1] : s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X0),s('type.integer.int',X1))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))),
  inference(rectify,[],[f49])).
fof(f49,axiom,(
  ! [X13,X14] : s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X13),s('type.integer.int',X14))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X14)))))),
  file('thm.integer.INT_ADD_CALCULATE|split|4','thm.integer.int_sub')).
fof(f1126,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))))))) = s('type.integer.int','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num',X0)))))))))) )),
  inference(cnf_transformation,[],[f451])).
fof(f451,plain,(
  ! [X0,X1] : s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))))))) = s('type.integer.int','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num',X0)))))))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6,X7] : s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))))))) = s('type.integer.int','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num',X6))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X6),s('type.num.num',X7))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X7),s('type.num.num',X6)))))))))),
  file('thm.integer.INT_ADD_CALCULATE|split|4','thm.integer.INT_ADD_CALCULATE|split|3')).
fof(f23275,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK15))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK14))))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',sK15))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK15),s('type.num.num',sK14))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num',sK15)))))))))))),
  inference(unit_resulting_resolution,[],[f1451,f1157])).
fof(f1157,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))) | s('type.integer.int',X0) = s('type.integer.int',X1)) )),
  inference(cnf_transformation,[],[f779])).
fof(f779,plain,(
  ! [X0,X1] : ((s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))) | s('type.integer.int',X0) = s('type.integer.int',X1)) & (s('type.integer.int',X0) != s('type.integer.int',X1) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))),
  inference(nnf_transformation,[],[f473])).
fof(f473,plain,(
  ! [X0,X1] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))) <=> s('type.integer.int',X0) = s('type.integer.int',X1))),
  inference(rectify,[],[f233])).
fof(f233,axiom,(
  ! [X13,X14] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X13))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X14))) <=> s('type.integer.int',X13) = s('type.integer.int',X14))),
  file('thm.integer.INT_ADD_CALCULATE|split|4','thm.integer.INT_EQ_NEG')).
fof(f1451,plain,(
  s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK15))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK14))))) != s('type.integer.int','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',sK15))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK15),s('type.num.num',sK14))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num',sK15)))))))))),
  inference(backward_demodulation,[],[f1085,f1449])).
fof(f1449,plain,(
  s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK15))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK14))))))) != s('type.integer.int','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',sK15))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK15),s('type.num.num',sK14))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num',sK15)))))))))),
  inference(backward_demodulation,[],[f1082,f875])).
fof(f875,plain,(
  s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK14))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK15))))) != s('type.integer.int','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',sK15))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK15),s('type.num.num',sK14))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num',sK15)))))))))),
  inference(cnf_transformation,[],[f692])).
fof(f692,plain,(
  s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK14))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK15))))) != s('type.integer.int','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',sK15))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK15),s('type.num.num',sK14))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num',sK15)))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK14,sK15])],[f589])).
fof(f589,plain,(
  ? [X0,X1] : s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))))) != s('type.integer.int','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num',X0))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1)))))))))),
  inference(ennf_transformation,[],[f263])).
fof(f263,plain,(
  ~! [X0,X1] : s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))))) = s('type.integer.int','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num',X0))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1)))))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))))) = s('type.integer.int','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num',X7))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X7),s('type.num.num',X6))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X6),s('type.num.num',X7)))))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))))) = s('type.integer.int','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num',X7))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X7),s('type.num.num',X6))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X6),s('type.num.num',X7)))))))))),
  file('thm.integer.INT_ADD_CALCULATE|split|4','thm.integer.INT_ADD_CALCULATE|split|4')).
fof(f1082,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) )),
  inference(cnf_transformation,[],[f407])).
fof(f407,plain,(
  ! [X0,X1] : s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',X0)))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X14,X13] : s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X13),s('type.integer.int',X14))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X14),s('type.integer.int',X13)))),
  file('thm.integer.INT_ADD_CALCULATE|split|4','thm.integer.INT_ADD_SYM')).
% SZS output end Proof for thm.integer
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 551
% Passive clauses: 2976
% Generated clauses: 22386
% Final active clauses: 536
% Final passive clauses: 2346
% Input formulas: 261
% Initial clauses: 525
% 
% Duplicate literals: 9
% Trivial inequalities: 9
% Fw subsumption resolutions: 180
% Fw demodulations: 3871
% Bw demodulations: 90
% Condensations: 26
% 
% Simple tautologies: 120
% Equational tautologies: 35
% Forward subsumptions: 14899
% Fw demodulations to eq. taut.: 243
% Bw demodulations to eq. taut.: 4
% 
% Unit resulting resolution: 4718
% Factoring: 2
% Forward superposition: 8777
% Backward superposition: 4088
% Self superposition: 66
% Equality factoring: 2
% Equality resolution: 12
% 
% Split clauses: 9
% Split components: 18
% SAT solver clauses: 2997
% SAT solver unit clauses: 1513
% SAT solver binary clauses: 1276
% 
% Sat splits: 9
% Sat splitting refutations: 1
% 
% Memory used [KB]: 15223
% Time elapsed: 1.138 s
% ------------------------------
% ------------------------------
% Success in time 2.55 s
