Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.relation
% (67553)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 201
% Passive clauses: 3406
% Generated clauses: 18176
% Final active clauses: 191
% Final passive clauses: 2972
% Input formulas: 27
% Initial clauses: 71
% 
% Duplicate literals: 856
% Trivial inequalities: 302
% Fw subsumption resolutions: 2586
% Fw demodulations: 2726
% Bw demodulations: 6
% 
% Simple tautologies: 339
% Equational tautologies: 18
% Forward subsumptions: 6395
% 
% Binary resolution: 1182
% Forward superposition: 6393
% Backward superposition: 3742
% Self superposition: 64
% Equality factoring: 32
% Equality resolution: 46
% 
% Split clauses: 531
% Split components: 1070
% SAT solver clauses: 1169
% SAT solver unit clauses: 12
% SAT solver binary clauses: 139
% 
% Sat splits: 531
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 1168
% TWLsolver calls for satisfiability: 309
% 
% Memory used [KB]: 10490
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.relation
% (67568)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 119
% Passive clauses: 1313
% Generated clauses: 12340
% Final active clauses: 103
% Final passive clauses: 711
% Input formulas: 27
% Initial clauses: 71
% 
% Duplicate literals: 402
% Trivial inequalities: 173
% Fw subsumption resolutions: 848
% Fw demodulations: 1986
% Bw demodulations: 345
% Condensations: 491
% Global subsumptions: 8
% 
% Simple tautologies: 180
% Equational tautologies: 29
% Forward subsumptions: 4739
% Bw demodulations to eq. taut.: 152
% 
% Binary resolution: 419
% Factoring: 3
% Forward superposition: 3762
% Backward superposition: 2316
% Self superposition: 77
% Equality factoring: 48
% Equality resolution: 18
% 
% Split clauses: 129
% Split components: 266
% SAT solver clauses: 2133
% SAT solver unit clauses: 76
% SAT solver binary clauses: 44
% SAT solver learnt clauses: 16
% 
% Sat splits: 129
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 2101
% TWLsolver calls for satisfiability: 347
% 
% Memory used [KB]: 12025
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.relation
% (67578)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 557
% Passive clauses: 7563
% Generated clauses: 172141
% Final active clauses: 64
% Final passive clauses: 289
% Input formulas: 27
% Initial clauses: 71
% 
% Duplicate literals: 5466
% Trivial inequalities: 830
% Fw subsumption resolutions: 13238
% Fw demodulations: 29043
% Bw demodulations: 5633
% Condensations: 12355
% Global subsumptions: 489
% 
% Simple tautologies: 998
% Equational tautologies: 1842
% Forward subsumptions: 97303
% Fw demodulations to eq. taut.: 149
% Bw demodulations to eq. taut.: 324
% 
% Unit resulting resolution: 9407
% Forward superposition: 33041
% Backward superposition: 23266
% Self superposition: 636
% Equality factoring: 325
% Equality resolution: 75
% 
% Split clauses: 508
% Split components: 1031
% SAT solver clauses: 9655
% SAT solver unit clauses: 81
% SAT solver binary clauses: 604
% 
% Sat splits: 508
% Sat splitting refutations: 585
% 
% Memory used [KB]: 76246
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.relation
% (67950)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 701
% Passive clauses: 30727
% Generated clauses: 136881
% Final active clauses: 127
% Final passive clauses: 17622
% Input formulas: 27
% Initial clauses: 71
% 
% Selected by SInE selection: 27
% SInE iterations: 1
% 
% Duplicate literals: 4110
% Trivial inequalities: 1628
% Fw subsumption resolutions: 22419
% Fw demodulations: 15804
% Bw demodulations: 335
% 
% Simple tautologies: 3009
% Equational tautologies: 104
% Forward subsumptions: 49871
% Bw demodulations to eq. taut.: 7
% 
% Binary resolution: 11021
% Factoring: 11
% Forward superposition: 38314
% Backward superposition: 35060
% Self superposition: 505
% Equality factoring: 238
% Equality resolution: 77
% 
% Split clauses: 9641
% Split components: 20582
% SAT solver clauses: 9959
% SAT solver unit clauses: 27
% SAT solver binary clauses: 2707
% 
% Sat splits: 9641
% Sat splitting refutations: 125
% 
% Memory used [KB]: 37483
% Time elapsed: 4.0000 s
% ------------------------------
% ------------------------------
% remaining time: 142 next slice time: 6
lrs+1002_2:3_br=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:sas=minisat:stl=30:sac=on:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none:sp=reverse_arity:urr=on_4 on thm.relation
% (68090)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 76
% Passive clauses: 1982
% Generated clauses: 12320
% Final active clauses: 44
% Final passive clauses: 738
% Input formulas: 27
% Initial clauses: 71
% 
% Duplicate literals: 369
% Trivial inequalities: 283
% Fw subsumption resolutions: 1900
% Fw demodulations: 749
% Bw demodulations: 1162
% Global subsumptions: 109
% 
% Simple tautologies: 441
% Equational tautologies: 73
% Forward subsumptions: 6140
% Bw demodulations to eq. taut.: 4
% 
% Unit resulting resolution: 371
% Forward superposition: 3586
% Backward superposition: 2394
% Self superposition: 82
% Equality factoring: 54
% Equality resolution: 12
% 
% Split clauses: 17
% Split components: 35
% SAT solver clauses: 2326
% SAT solver unit clauses: 159
% SAT solver binary clauses: 136
% 
% Sat splits: 17
% Sat splitting refutations: 216
% 
% Memory used [KB]: 7675
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 135 next slice time: 15
lrs+1011_12_bs=on:bsr=unit_only:cond=on:gs=on:gsaa=from_current:gsssp=full:nwc=1.1:sas=minisat:stl=60:sos=all:sac=on:sdd=off:sser=off:sfr=on:ssfp=100000:ssfq=1.2:ssnc=none:sp=reverse_arity:updr=off_14 on thm.relation
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.relation
% SZS output start Proof for thm.relation
fof(f1121,plain,(
  $false),
  inference(subsumption_resolution,[],[f1120,f187])).
fof(f187,plain,(
  p(s(bool,'T_0'))),
  inference(condensation,[],[f148])).
fof(f148,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f95])).
fof(f95,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))))),
  inference(flattening,[],[f94])).
fof(f94,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f65])).
fof(f65,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f45])).
fof(f45,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X9] : ((p(s(bool,X9)) => p(s(bool,X9))) <=> p(s(bool,'T_0')))),
  file('thm.relation.antisymmetric_RINTER|split|0','thm.bool.IMP_CLAUSES|split|3')).
fof(f1120,plain,(
  ~p(s(bool,'T_0'))),
  inference(forward_demodulation,[],[f1119,f959])).
fof(f959,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))),
  inference(resolution,[],[f899,f120])).
fof(f120,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f32])).
fof(f32,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X9] : (s(bool,'T_0') = s(bool,X9) <=> p(s(bool,X9)))),
  file('thm.relation.antisymmetric_RINTER|split|0','thm.bool.EQ_CLAUSES|split|0')).
fof(f899,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))))),
  inference(resolution,[],[f281,f167])).
fof(f167,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))) | ~p(s(bool,'const.relation.RINTER_4'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3),s(X0,X4),s(X1,X5))))) )),
  inference(cnf_transformation,[],[f108])).
fof(f108,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((~p(s(bool,'const.relation.RINTER_4'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3),s(X0,X4),s(X1,X5)))) | (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X4))),s(X1,X5)))))) & (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X4))),s(X1,X5)))) | p(s(bool,'const.relation.RINTER_4'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3),s(X0,X4),s(X1,X5))))))),
  inference(flattening,[],[f107])).
fof(f107,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((~p(s(bool,'const.relation.RINTER_4'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3),s(X0,X4),s(X1,X5)))) | (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X4))),s(X1,X5)))))) & ((~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X4))),s(X1,X5))))) | p(s(bool,'const.relation.RINTER_4'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3),s(X0,X4),s(X1,X5))))))),
  inference(nnf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (p(s(bool,'const.relation.RINTER_4'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3),s(X0,X4),s(X1,X5)))) <=> (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X4))),s(X1,X5))))))),
  inference(flattening,[],[f58])).
fof(f58,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5] : (p(s(bool,'const.relation.RINTER_4'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3),s(X0,X4),s(X1,X5)))) <=> (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X4))),s(X1,X5))))))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X6,X18] : ! [X8,X7,X10,X11] : (p(s(bool,'const.relation.RINTER_4'(s(fun(X6,fun(X18,bool)),X8),s(fun(X6,fun(X18,bool)),X7),s(X6,X10),s(X18,X11)))) <=> (p(s(bool,happ(s(fun(X18,bool),happ(s(fun(X6,fun(X18,bool)),X8),s(X6,X10))),s(X18,X11)))) & p(s(bool,happ(s(fun(X18,bool),happ(s(fun(X6,fun(X18,bool)),X7),s(X6,X10))),s(X18,X11))))))),
  file('thm.relation.antisymmetric_RINTER|split|0','thm.relation.RINTER')).
fof(f281,plain,(
  p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))))),
  inference(forward_demodulation,[],[f245,f166])).
fof(f166,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'const.relation.RINTER_4'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3),s(X0,X4),s(X1,X5))) = s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),'const.relation.RINTER_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3))),s(X0,X4))),s(X1,X5)))) )),
  inference(cnf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(bool,'const.relation.RINTER_4'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3),s(X0,X4),s(X1,X5))) = s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),'const.relation.RINTER_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3))),s(X0,X4))),s(X1,X5)))),
  inference(flattening,[],[f56])).
fof(f56,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5] : s(bool,'const.relation.RINTER_4'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3),s(X0,X4),s(X1,X5))) = s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),'const.relation.RINTER_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3))),s(X0,X4))),s(X1,X5)))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X6,X18] : ! [X19,X20,X21,X22] : s(bool,'const.relation.RINTER_4'(s(fun(X6,fun(X18,bool)),X19),s(fun(X6,fun(X18,bool)),X20),s(X6,X21),s(X18,X22))) = s(bool,happ(s(fun(X18,bool),happ(s(fun(X6,fun(X18,bool)),'const.relation.RINTER_2'(s(fun(X6,fun(X18,bool)),X19),s(fun(X6,fun(X18,bool)),X20))),s(X6,X21))),s(X18,X22)))),
  file('thm.relation.antisymmetric_RINTER|split|0','ARITY_2_4_const.relation.RINTER')).
fof(f245,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1))),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))))),
  inference(resolution,[],[f114,f158])).
fof(f158,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK3(X1,X0)))),s(X0,sK4(X1,X0)))))) )),
  inference(cnf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) | ! [X4,X5] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X5)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X4)))) | s(X0,X4) = s(X0,X5))) & ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK3(X1,X0)))),s(X0,sK4(X1,X0))))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK4(X1,X0)))),s(X0,sK3(X1,X0))))) & s(X0,sK3(X1,X0)) != s(X0,sK4(X1,X0))) | p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3,sK4])],[f101])).
fof(f101,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) | ! [X4,X5] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X5)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X4)))) | s(X0,X4) = s(X0,X5))) & (? [X2,X3] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2)))) & s(X0,X2) != s(X0,X3)) | p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(rectify,[],[f100])).
fof(f100,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) | ! [X2,X3] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2)))) | s(X0,X2) = s(X0,X3))) & (? [X2,X3] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2)))) & s(X0,X2) != s(X0,X3)) | p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(nnf_transformation,[],[f69])).
fof(f69,plain,(
  ! [X0,X1] : (p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> ! [X2,X3] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2)))) | s(X0,X2) = s(X0,X3)))),
  inference(flattening,[],[f68])).
fof(f68,plain,(
  ! [X0,X1] : (p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> ! [X2,X3] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2))))) | s(X0,X2) = s(X0,X3)))),
  inference(ennf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0,X1] : (p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> ! [X2,X3] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2))))) => s(X0,X2) = s(X0,X3)))),
  inference(flattening,[],[f49])).
fof(f49,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> ! [X2,X3] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2))))) => s(X0,X2) = s(X0,X3)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X6] : ! [X17] : (p(s(bool,'const.relation.antisymmetric_1'(s(fun(X6,fun(X6,bool)),X17)))) <=> ! [X10,X11] : ((p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X17),s(X6,X10))),s(X6,X11)))) & p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X17),s(X6,X11))),s(X6,X10))))) => s(X6,X10) = s(X6,X11)))),
  file('thm.relation.antisymmetric_RINTER|split|0','thm.relation.antisymmetric_def')).
fof(f114,plain,(
  ~p(s(bool,'const.relation.antisymmetric_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1))))))),
  inference(cnf_transformation,[],[f75])).
fof(f75,plain,(
  p(s(bool,'const.relation.antisymmetric_1'(s(fun(sK0,fun(sK0,bool)),sK2)))) & ~p(s(bool,'const.relation.antisymmetric_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f62])).
fof(f62,plain,(
  ? [X0,X1,X2] : (p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X2)))) & ~p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.RINTER_2'(s(fun(X0,fun(X0,bool)),X2),s(fun(X0,fun(X0,bool)),X1)))))))),
  inference(ennf_transformation,[],[f30])).
fof(f30,plain,(
  ~! [X0,X1,X2] : (p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X2)))) => p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.RINTER_2'(s(fun(X0,fun(X0,bool)),X2),s(fun(X0,fun(X0,bool)),X1)))))))),
  inference(flattening,[],[f29])).
fof(f29,plain,(
  ~! [X0] : ! [X1,X2] : (p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X2)))) => p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.RINTER_2'(s(fun(X0,fun(X0,bool)),X2),s(fun(X0,fun(X0,bool)),X1)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : (p(s(bool,'const.relation.antisymmetric_1'(s(fun(X6,fun(X6,bool)),X8)))) => p(s(bool,'const.relation.antisymmetric_1'(s(fun(X6,fun(X6,bool)),'const.relation.RINTER_2'(s(fun(X6,fun(X6,bool)),X8),s(fun(X6,fun(X6,bool)),X7)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : (p(s(bool,'const.relation.antisymmetric_1'(s(fun(X6,fun(X6,bool)),X8)))) => p(s(bool,'const.relation.antisymmetric_1'(s(fun(X6,fun(X6,bool)),'const.relation.RINTER_2'(s(fun(X6,fun(X6,bool)),X8),s(fun(X6,fun(X6,bool)),X7)))))))),
  file('thm.relation.antisymmetric_RINTER|split|0','thm.relation.antisymmetric_RINTER|split|0')).
fof(f1119,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))))),
  inference(subsumption_resolution,[],[f1082,f247])).
fof(f247,plain,(
  s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)) != s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),
  inference(resolution,[],[f114,f160])).
fof(f160,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) | s(X0,sK3(X1,X0)) != s(X0,sK4(X1,X0))) )),
  inference(cnf_transformation,[],[f102])).
fof(f1082,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)) = s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),
  inference(resolution,[],[f282,f699])).
fof(f699,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),X2),s(sK0,X1),s(sK0,X0)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X0))),s(sK0,X1)))) | s(sK0,X0) = s(sK0,X1)) )),
  inference(resolution,[],[f201,f167])).
fof(f201,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X1))),s(sK0,X0)))) | s(sK0,X0) = s(sK0,X1) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X0))),s(sK0,X1))))) )),
  inference(resolution,[],[f113,f157])).
fof(f157,plain,(
  ( ! [X4,X0,X5,X1] : (s(X0,X4) = s(X0,X5) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X4)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X5)))) | ~p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(cnf_transformation,[],[f102])).
fof(f113,plain,(
  p(s(bool,'const.relation.antisymmetric_1'(s(fun(sK0,fun(sK0,bool)),sK2))))),
  inference(cnf_transformation,[],[f75])).
fof(f282,plain,(
  p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))))),
  inference(forward_demodulation,[],[f246,f166])).
fof(f246,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))))),
  inference(resolution,[],[f114,f159])).
fof(f159,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK4(X1,X0)))),s(X0,sK3(X1,X0)))))) )),
  inference(cnf_transformation,[],[f102])).
% SZS output end Proof for thm.relation
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 63
% Passive clauses: 182
% Generated clauses: 922
% Final active clauses: 58
% Final passive clauses: 157
% Input formulas: 27
% Initial clauses: 71
% 
% Duplicate literals: 2
% Trivial inequalities: 12
% Fw subsumption resolutions: 60
% Fw demodulations: 196
% Bw demodulations: 6
% Condensations: 17
% 
% Simple tautologies: 40
% Forward subsumptions: 423
% 
% Binary resolution: 165
% Forward superposition: 334
% Backward superposition: 145
% 
% SAT solver clauses: 184
% SAT solver unit clauses: 16
% SAT solver binary clauses: 47
% 
% Memory used [KB]: 2174
% Time elapsed: 0.041 s
% ------------------------------
% ------------------------------
% Success in time 16.489 s
