Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.realax
% (72644)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 113
% Passive clauses: 2642
% Generated clauses: 16809
% Final active clauses: 110
% Final passive clauses: 2415
% Input formulas: 39
% Initial clauses: 48
% 
% Duplicate literals: 778
% Trivial inequalities: 197
% Fw subsumption resolutions: 2004
% Fw demodulations: 2844
% Bw demodulations: 1
% 
% Simple tautologies: 310
% Equational tautologies: 26
% Forward subsumptions: 7397
% 
% Binary resolution: 29
% Forward superposition: 6921
% Backward superposition: 3608
% Self superposition: 213
% Equality factoring: 86
% Equality resolution: 42
% 
% Split clauses: 448
% Split components: 896
% SAT solver clauses: 722
% SAT solver unit clauses: 55
% SAT solver binary clauses: 207
% 
% Sat splits: 448
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 722
% TWLsolver calls for satisfiability: 100
% 
% Memory used [KB]: 13176
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.realax
% (72652)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 85
% Passive clauses: 731
% Generated clauses: 13002
% Final active clauses: 71
% Final passive clauses: 558
% Input formulas: 39
% Initial clauses: 48
% 
% Duplicate literals: 743
% Trivial inequalities: 196
% Fw subsumption resolutions: 1027
% Fw demodulations: 1501
% Bw demodulations: 97
% Condensations: 1609
% 
% Simple tautologies: 164
% Equational tautologies: 25
% Forward subsumptions: 5645
% Bw demodulations to eq. taut.: 2
% 
% Binary resolution: 16
% Forward superposition: 3882
% Backward superposition: 2845
% Self superposition: 200
% Equality factoring: 86
% Equality resolution: 19
% 
% Split clauses: 16
% Split components: 32
% SAT solver clauses: 1306
% SAT solver unit clauses: 62
% SAT solver binary clauses: 43
% 
% Sat splits: 16
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 1306
% TWLsolver calls for satisfiability: 287
% 
% Memory used [KB]: 11897
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.realax
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.realax
% SZS output start Proof for thm.realax
fof(f36585,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f132,f180,f178,f196,f199,f197,f743,f2172,f2117,f2175,f2188,f13224,f13230,f13231,f13232,f13261,f13084,f13226,f13309,f16782,f16788,f16789,f16799,f16801,f16820,f16825,f16784,f19835,f19839,f19837,f20787,f20788,f20789,f20790,f35082,f35088,f35089,f35128,f35142,f35144,f34933,f35084,f35205,f36369,f36375,f36376,f36416,f36430,f36432,f2173,f36216,f36371,f36493,f36576,f36578,f36579,f36580])).
fof(f36580,plain,(
  $false | $spl29),
  inference(subsumption_resolution,[],[f36574,f19648])).
fof(f19648,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'|dollar|equals_0'),s(X0,X1))),s(X0,X1)))) )),
  inference(unit_resulting_resolution,[],[f19248,f190])).
fof(f190,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f133,f137])).
fof(f137,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.realax.REAL_ADD_SYM',aHL_BOOL_CASES)).
fof(f133,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f43])).
fof(f43,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.realax.REAL_ADD_SYM',aHL_FALSITY)).
fof(f19248,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'|dollar|equals_0'),s(X0,X1))),s(X0,X1))))) )),
  inference(forward_demodulation,[],[f19247,f167])).
fof(f167,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),'_dst_x1x0_2'(s(fun(X1,fun(X2,X0)),X3),s(X2,X4))),s(X1,X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X4)))) )),
  inference(cnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X1,X0),'_dst_x1x0_2'(s(fun(X1,fun(X2,X0)),X3),s(X2,X4))),s(X1,X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X4)))),
  inference(flattening,[],[f81])).
fof(f81,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,happ(s(fun(X1,X0),'_dst_x1x0_2'(s(fun(X1,fun(X2,X0)),X3),s(X2,X4))),s(X1,X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X4)))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X41,X39,X42] : ! [X43,X40,X44] : s(X41,happ(s(fun(X39,X41),'_dst_x1x0_2'(s(fun(X39,fun(X42,X41)),X43),s(X42,X40))),s(X39,X44))) = s(X41,happ(s(fun(X42,X41),happ(s(fun(X39,fun(X42,X41)),X43),s(X39,X44))),s(X42,X40)))),
  file('thm.realax.REAL_ADD_SYM','_dst_x1x0_ax')).
fof(f19247,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'|dollar|equals_0'),s(X0,X1))),s(X0,X1))))) )),
  inference(forward_demodulation,[],[f19246,f181])).
fof(f181,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))) )),
  inference(forward_demodulation,[],[f140,f138])).
fof(f138,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))) )),
  inference(cnf_transformation,[],[f48])).
fof(f48,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))),
  inference(flattening,[],[f47])).
fof(f47,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X8] : ! [X6] : s(X8,X6) = s(X8,'const.combin.I_1'(s(X8,X6)))),
  file('thm.realax.REAL_ADD_SYM','thm.combin.I_THM')).
fof(f140,plain,(
  ( ! [X0,X1] : (s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0,X1] : s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))),
  inference(flattening,[],[f51])).
fof(f51,plain,(
  ! [X0] : ! [X1] : s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X8] : ! [X35] : s(X8,'const.combin.I_1'(s(X8,X35))) = s(X8,happ(s(fun(X8,X8),'const.combin.I_0'),s(X8,X35)))),
  file('thm.realax.REAL_ADD_SYM','ARITY_0_1_const.combin.I')).
fof(f19246,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'|dollar|equals_0'),s(X0,X1))),s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1))))))) )),
  inference(forward_demodulation,[],[f18992,f167])).
fof(f18992,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'|dollar|equals_0'),s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1))))),s(X0,X1))))) )),
  inference(unit_resulting_resolution,[],[f181,f9467])).
fof(f9467,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'|dollar|equals_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) != s(X0,X2)) )),
  inference(forward_demodulation,[],[f9466,f181])).
fof(f9466,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'|dollar|equals_0'),s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1))))),s(X0,X2)))) | s(X0,X1) != s(X0,X2)) )),
  inference(forward_demodulation,[],[f9465,f181])).
fof(f9465,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'|dollar|equals_0'),s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1))))),s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X2))))))) )),
  inference(subsumption_resolution,[],[f9437,f130])).
fof(f130,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.realax.REAL_ADD_SYM',aHL_TRUTH)).
fof(f9437,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'|dollar|equals_0'),s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1))))),s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X2))))))) )),
  inference(superposition,[],[f153,f248])).
fof(f248,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X0,fun(X0,bool)),'|dollar|equals_0'),s(fun(X0,X0),'const.combin.I_0'),s(fun(X0,X0),'const.combin.I_0')))) )),
  inference(unit_resulting_resolution,[],[f134,f190])).
fof(f134,plain,(
  ( ! [X0] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X0,fun(X0,bool)),'|dollar|equals_0'),s(fun(X0,X0),'const.combin.I_0'),s(fun(X0,X0),'const.combin.I_0'))))) )),
  inference(cnf_transformation,[],[f44])).
fof(f44,plain,(
  ! [X0] : p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X0,fun(X0,bool)),'|dollar|equals_0'),s(fun(X0,X0),'const.combin.I_0'),s(fun(X0,X0),'const.combin.I_0'))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X8] : p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X8,fun(X8,bool)),'|dollar|equals_0'),s(fun(X8,X8),'const.combin.I_0'),s(fun(X8,X8),'const.combin.I_0'))))),
  file('thm.realax.REAL_ADD_SYM','thm.quotient.IDENTITY_QUOTIENT')).
fof(f153,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X0,fun(X0,bool)),X2),s(fun(X0,X1),X3),s(fun(X1,X0),X4)))) | s(X1,X5) != s(X1,X6) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))))),s(X0,happ(s(fun(X1,X0),X4),s(X1,X6))))))) )),
  inference(cnf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0,X1,X2,X3,X4] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X0,fun(X0,bool)),X2),s(fun(X0,X1),X3),s(fun(X1,X0),X4)))) | ! [X5,X6] : ((s(X1,X5) != s(X1,X6) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))))),s(X0,happ(s(fun(X1,X0),X4),s(X1,X6))))))) & (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))))),s(X0,happ(s(fun(X1,X0),X4),s(X1,X6)))))) | s(X1,X5) = s(X1,X6))))),
  inference(nnf_transformation,[],[f107])).
fof(f107,plain,(
  ! [X0,X1,X2,X3,X4] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X0,fun(X0,bool)),X2),s(fun(X0,X1),X3),s(fun(X1,X0),X4)))) | ! [X5,X6] : (s(X1,X5) = s(X1,X6) <=> p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))))),s(X0,happ(s(fun(X1,X0),X4),s(X1,X6))))))))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X0,fun(X0,bool)),X2),s(fun(X0,X1),X3),s(fun(X1,X0),X4)))) => ! [X5,X6] : (s(X1,X5) = s(X1,X6) <=> p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))))),s(X0,happ(s(fun(X1,X0),X4),s(X1,X6))))))))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X0,fun(X0,bool)),X2),s(fun(X0,X1),X3),s(fun(X1,X0),X4)))) => ! [X5,X6] : (s(X1,X5) = s(X1,X6) <=> p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))))),s(X0,happ(s(fun(X1,X0),X4),s(X1,X6))))))))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X8,X12] : ! [X23,X24,X25] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X8,fun(X8,bool)),X23),s(fun(X8,X12),X24),s(fun(X12,X8),X25)))) => ! [X6,X7] : (s(X12,X6) = s(X12,X7) <=> p(s(bool,happ(s(fun(X8,bool),happ(s(fun(X8,fun(X8,bool)),X23),s(X8,happ(s(fun(X12,X8),X25),s(X12,X6))))),s(X8,happ(s(fun(X12,X8),X25),s(X12,X7))))))))),
  file('thm.realax.REAL_ADD_SYM','thm.quotient.EQUALS_PRS')).
fof(f36574,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),'|dollar|equals_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))))) | $spl29),
  inference(backward_demodulation,[],[f36522,f36371])).
fof(f36522,plain,(
  ( ! [X2,X3] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X2),s('type.realax.real',X3))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X3),s('type.realax.real',X2)))) )),
  inference(superposition,[],[f1991,f145])).
fof(f145,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_ABS_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_add_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.real_REP_1'(s('type.realax.real',X0))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.real_REP_1'(s('type.realax.real',X1)))))))) )),
  inference(cnf_transformation,[],[f58])).
fof(f58,plain,(
  ! [X0,X1] : s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_ABS_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_add_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.real_REP_1'(s('type.realax.real',X0))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.real_REP_1'(s('type.realax.real',X1)))))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X33,X34] : s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X33),s('type.realax.real',X34))) = s('type.realax.real','const.realax.real_ABS_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_add_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.real_REP_1'(s('type.realax.real',X33))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.real_REP_1'(s('type.realax.real',X34)))))))),
  file('thm.realax.REAL_ADD_SYM','thm.realax.real_add')).
fof(f1991,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_ABS_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_add_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.real_REP_1'(s('type.realax.real',X1))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.real_REP_1'(s('type.realax.real',X0)))))))) )),
  inference(superposition,[],[f145,f144])).
fof(f144,plain,(
  ( ! [X0,X1] : (s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_add_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X0),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1))) = s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_add_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X0)))) )),
  inference(cnf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0,X1] : s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_add_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X0),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1))) = s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_add_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X0)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X6,X7] : s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_add_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X6),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X7))) = s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_add_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X7),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X6)))),
  file('thm.realax.REAL_ADD_SYM','thm.realax.TREAL_ADD_SYM')).
fof(f36579,plain,(
  $false | $spl25),
  inference(subsumption_resolution,[],[f36563,f19648])).
fof(f36563,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),'|dollar|equals_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))))) | $spl25),
  inference(backward_demodulation,[],[f36522,f35084])).
fof(f36578,plain,(
  $false | ($spl7 | $spl15)),
  inference(subsumption_resolution,[],[f36577,f199])).
fof(f36577,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl15),
  inference(forward_demodulation,[],[f36551,f19648])).
fof(f36551,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),'|dollar|equals_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))))) | $spl15),
  inference(backward_demodulation,[],[f36522,f33956])).
fof(f33956,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),'|dollar|equals_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))))) | $spl15),
  inference(forward_demodulation,[],[f33929,f181])).
fof(f33929,plain,(
  s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),'|dollar|equals_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))))) | $spl15),
  inference(unit_resulting_resolution,[],[f16828,f33331,f332])).
fof(f332,plain,(
  ( ! [X2,X1] : (p(s(bool,X1)) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f130,f187])).
fof(f187,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f137,f137])).
fof(f33331,plain,(
  ~p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),'|dollar|equals_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))))))),
  inference(unit_resulting_resolution,[],[f129,f11188])).
fof(f11188,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'|dollar|equals_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) )),
  inference(forward_demodulation,[],[f11187,f181])).
fof(f11187,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'|dollar|equals_0'),s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1))))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) )),
  inference(forward_demodulation,[],[f11186,f181])).
fof(f11186,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'|dollar|equals_0'),s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1))))),s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X2)))))) | s(X0,X1) = s(X0,X2)) )),
  inference(subsumption_resolution,[],[f11155,f130])).
fof(f11155,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'|dollar|equals_0'),s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1))))),s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X2)))))) | s(X0,X1) = s(X0,X2)) )),
  inference(superposition,[],[f154,f248])).
fof(f154,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X0,fun(X0,bool)),X2),s(fun(X0,X1),X3),s(fun(X1,X0),X4)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))))),s(X0,happ(s(fun(X1,X0),X4),s(X1,X6)))))) | s(X1,X5) = s(X1,X6)) )),
  inference(cnf_transformation,[],[f123])).
fof(f129,plain,(
  s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK0)))),
  inference(cnf_transformation,[],[f119])).
fof(f119,plain,(
  s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK0)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f100])).
fof(f100,plain,(
  ? [X0,X1] : s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X0)))),
  inference(ennf_transformation,[],[f41])).
fof(f41,plain,(
  ~! [X0,X1] : s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X0)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X6),s('type.realax.real',X7))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X7),s('type.realax.real',X6)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X6),s('type.realax.real',X7))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X7),s('type.realax.real',X6)))),
  file('thm.realax.REAL_ADD_SYM','thm.realax.REAL_ADD_SYM')).
fof(f16828,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | $spl15),
  inference(subsumption_resolution,[],[f16502,f13226])).
fof(f16502,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | $spl15),
  inference(superposition,[],[f13226,f15768])).
fof(f15768,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(condensation,[],[f15767])).
fof(f15767,plain,(
  ( ! [X0,X1] : (s(bool,X1) = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f15317])).
fof(f15317,plain,(
  ( ! [X0,X1] : (s(bool,X1) = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f290,f137])).
fof(f290,plain,(
  ( ! [X19,X20] : (s(bool,X20) = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X19))) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') = s(bool,X20)) )),
  inference(superposition,[],[f187,f181])).
fof(f36576,plain,(
  $false | ($spl7 | $spl15)),
  inference(subsumption_resolution,[],[f36575,f199])).
fof(f36575,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl15),
  inference(forward_demodulation,[],[f36546,f19648])).
fof(f36546,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),'|dollar|equals_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))))) | $spl15),
  inference(backward_demodulation,[],[f36522,f33893])).
fof(f33893,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),'|dollar|equals_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))) | $spl15),
  inference(forward_demodulation,[],[f33866,f181])).
fof(f33866,plain,(
  s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),'|dollar|equals_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))) | $spl15),
  inference(unit_resulting_resolution,[],[f16828,f33326,f332])).
fof(f33326,plain,(
  ~p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),'|dollar|equals_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))))),
  inference(unit_resulting_resolution,[],[f129,f11188])).
fof(f36493,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),'|dollar|equals_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',sK1)))))) ) | $spl30),
  inference(cnf_transformation,[],[f36493_D])).
fof(f36493_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),'|dollar|equals_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',sK1)))))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f36371,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),'|dollar|equals_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))))) | $spl29),
  inference(cnf_transformation,[],[f36371_D])).
fof(f36371_D,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),'|dollar|equals_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f36216,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),'|dollar|equals_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))))) | s(bool,X0) = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),'|dollar|equals_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',sK1)))))) )),
  inference(equality_factoring,[],[f33960])).
fof(f33960,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),'|dollar|equals_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f33931,f130])).
fof(f33931,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),'|dollar|equals_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f33331,f187])).
fof(f2173,plain,(
  ( ! [X30] : (s(bool,'T_0') = s(bool,X30)) ) | $spl8),
  inference(cnf_transformation,[],[f2173_D])).
fof(f2173_D,plain,(
  ( ! [X30] : (s(bool,'T_0') = s(bool,X30)) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f36432,plain,(
  ( ! [X269] : (s(bool,'T_0') != s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),'|dollar|equals_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))))) | s(bool,'T_0') = s(bool,X269)) )),
  inference(subsumption_resolution,[],[f36073,f15904])).
fof(f15904,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(forward_demodulation,[],[f15699,f138])).
fof(f15699,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'const.combin.I_1'(s(bool,'T_0'))) != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X0)))) )),
  inference(condensation,[],[f15698])).
fof(f15698,plain,(
  ( ! [X92,X93] : (s(bool,'const.combin.I_1'(s(bool,'T_0'))) != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X93))) | s(bool,'T_0') = s(bool,X92) | s(bool,'T_0') = s(bool,X93)) )),
  inference(duplicate_literal_removal,[],[f15447])).
fof(f15447,plain,(
  ( ! [X92,X93] : (s(bool,'const.combin.I_1'(s(bool,'T_0'))) != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X93))) | s(bool,'T_0') = s(bool,X92) | s(bool,'T_0') = s(bool,X93) | s(bool,'T_0') = s(bool,X92)) )),
  inference(superposition,[],[f1058,f290])).
fof(f1058,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(duplicate_literal_removal,[],[f1040])).
fof(f1040,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f263])).
fof(f263,plain,(
  ( ! [X23,X22] : (s(bool,X22) = s(bool,X23) | s(bool,'T_0') = s(bool,X23) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X22)))) )),
  inference(superposition,[],[f187,f138])).
fof(f36073,plain,(
  ( ! [X269] : (s(bool,'T_0') != s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),'|dollar|equals_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))))) | s(bool,'T_0') = s(bool,X269) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X269)))) )),
  inference(superposition,[],[f15904,f33960])).
fof(f36430,plain,(
  ( ! [X260] : (s(bool,'T_0') != s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),'|dollar|equals_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))))) | s(bool,'T_0') = s(bool,X260)) )),
  inference(subsumption_resolution,[],[f36067,f12676])).
fof(f12676,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(forward_demodulation,[],[f12509,f138])).
fof(f12509,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'const.combin.I_1'(s(bool,'T_0'))) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X0)))) )),
  inference(condensation,[],[f12508])).
fof(f12508,plain,(
  ( ! [X87,X86] : (s(bool,'const.combin.I_1'(s(bool,'T_0'))) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X87))) | s(bool,'T_0') = s(bool,X86) | s(bool,'T_0') = s(bool,X87)) )),
  inference(duplicate_literal_removal,[],[f12287])).
fof(f12287,plain,(
  ( ! [X87,X86] : (s(bool,'const.combin.I_1'(s(bool,'T_0'))) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X87))) | s(bool,'T_0') = s(bool,X86) | s(bool,'T_0') = s(bool,X87) | s(bool,'T_0') = s(bool,X86)) )),
  inference(superposition,[],[f1058,f289])).
fof(f289,plain,(
  ( ! [X17,X18] : (s(bool,X18) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X17))) | s(bool,'T_0') = s(bool,X17) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f187,f139])).
fof(f139,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(flattening,[],[f49])).
fof(f49,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X39] : ! [X40] : s(X39,X40) = s(X39,happ(s(fun(X39,X39),'_i_0_0'),s(X39,X40)))),
  file('thm.realax.REAL_ADD_SYM','_i0_ax')).
fof(f36067,plain,(
  ( ! [X260] : (s(bool,'T_0') != s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),'|dollar|equals_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))))) | s(bool,'T_0') = s(bool,X260) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X260)))) )),
  inference(superposition,[],[f12676,f33960])).
fof(f36416,plain,(
  ( ! [X148] : (s(bool,'T_0') != s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),'|dollar|equals_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))))) | s(bool,'T_0') = s(bool,X148)) )),
  inference(subsumption_resolution,[],[f36042,f1959])).
fof(f1959,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(forward_demodulation,[],[f1887,f138])).
fof(f1887,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'const.combin.I_1'(s(bool,'T_0'))) != s(bool,'const.combin.I_1'(s(bool,X0)))) )),
  inference(condensation,[],[f1886])).
fof(f1886,plain,(
  ( ! [X41,X42] : (s(bool,'const.combin.I_1'(s(bool,'T_0'))) != s(bool,'const.combin.I_1'(s(bool,X42))) | s(bool,'T_0') = s(bool,X41) | s(bool,'T_0') = s(bool,X42)) )),
  inference(duplicate_literal_removal,[],[f1814])).
fof(f1814,plain,(
  ( ! [X41,X42] : (s(bool,'const.combin.I_1'(s(bool,'T_0'))) != s(bool,'const.combin.I_1'(s(bool,X42))) | s(bool,'T_0') = s(bool,X41) | s(bool,'T_0') = s(bool,X42) | s(bool,'T_0') = s(bool,X41)) )),
  inference(superposition,[],[f1058,f288])).
fof(f288,plain,(
  ( ! [X15,X16] : (s(bool,X16) = s(bool,'const.combin.I_1'(s(bool,X15))) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f187,f138])).
fof(f36042,plain,(
  ( ! [X148] : (s(bool,'T_0') != s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),'|dollar|equals_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))))) | s(bool,'T_0') = s(bool,X148) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X148)))) )),
  inference(superposition,[],[f1959,f33960])).
fof(f36376,plain,(
  ( ! [X75] : (s(bool,'T_0') != s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),'|dollar|equals_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))))) | s(bool,'T_0') = s(bool,X75)) )),
  inference(subsumption_resolution,[],[f35996,f15904])).
fof(f35996,plain,(
  ( ! [X75] : (s(bool,'T_0') != s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),'|dollar|equals_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X75))) | s(bool,'T_0') = s(bool,X75)) )),
  inference(superposition,[],[f1319,f33960])).
fof(f1319,plain,(
  ( ! [X30] : (s(bool,'T_0') != s(bool,X30) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X30)))) )),
  inference(forward_demodulation,[],[f1276,f138])).
fof(f1276,plain,(
  ( ! [X30] : (s(bool,X30) != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X30)))) )),
  inference(superposition,[],[f1058,f181])).
fof(f36375,plain,(
  ( ! [X74] : (s(bool,'T_0') != s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),'|dollar|equals_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))))) | s(bool,'T_0') = s(bool,X74)) )),
  inference(subsumption_resolution,[],[f35995,f12676])).
fof(f35995,plain,(
  ( ! [X74] : (s(bool,'T_0') != s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),'|dollar|equals_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X74))) | s(bool,'T_0') = s(bool,X74)) )),
  inference(superposition,[],[f1318,f33960])).
fof(f1318,plain,(
  ( ! [X29] : (s(bool,'T_0') != s(bool,X29) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X29)))) )),
  inference(forward_demodulation,[],[f1275,f138])).
fof(f1275,plain,(
  ( ! [X29] : (s(bool,X29) != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X29)))) )),
  inference(superposition,[],[f1058,f139])).
fof(f36369,plain,(
  ( ! [X72] : (s(bool,'T_0') != s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),'|dollar|equals_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))))) | s(bool,'T_0') = s(bool,X72)) ) | $spl7),
  inference(subsumption_resolution,[],[f35993,f1959])).
fof(f35993,plain,(
  ( ! [X72] : (s(bool,'T_0') != s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),'|dollar|equals_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X72))) | s(bool,'T_0') = s(bool,X72)) ) | $spl7),
  inference(superposition,[],[f1217,f33960])).
fof(f1217,plain,(
  ( ! [X37] : (s(bool,'T_0') != s(bool,X37) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X37)))) ) | $spl7),
  inference(subsumption_resolution,[],[f967,f199])).
fof(f967,plain,(
  ( ! [X37] : (s(bool,'T_0') != s(bool,X37) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X37)))) ) | $spl7),
  inference(superposition,[],[f199,f263])).
fof(f35205,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),'|dollar|equals_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK0)))))) ) | $spl26),
  inference(cnf_transformation,[],[f35205_D])).
fof(f35205_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),'|dollar|equals_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK0)))))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f35084,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),'|dollar|equals_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))) | $spl25),
  inference(cnf_transformation,[],[f35084_D])).
fof(f35084_D,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),'|dollar|equals_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))) <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f34933,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),'|dollar|equals_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))) | s(bool,X0) = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),'|dollar|equals_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK0)))))) )),
  inference(equality_factoring,[],[f33897])).
fof(f33897,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),'|dollar|equals_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f33868,f130])).
fof(f33868,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),'|dollar|equals_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f33326,f187])).
fof(f35144,plain,(
  ( ! [X268] : (s(bool,'T_0') != s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),'|dollar|equals_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))) | s(bool,'T_0') = s(bool,X268)) )),
  inference(subsumption_resolution,[],[f34791,f15904])).
fof(f34791,plain,(
  ( ! [X268] : (s(bool,'T_0') != s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),'|dollar|equals_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))) | s(bool,'T_0') = s(bool,X268) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X268)))) )),
  inference(superposition,[],[f15904,f33897])).
fof(f35142,plain,(
  ( ! [X259] : (s(bool,'T_0') != s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),'|dollar|equals_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))) | s(bool,'T_0') = s(bool,X259)) )),
  inference(subsumption_resolution,[],[f34785,f12676])).
fof(f34785,plain,(
  ( ! [X259] : (s(bool,'T_0') != s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),'|dollar|equals_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))) | s(bool,'T_0') = s(bool,X259) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X259)))) )),
  inference(superposition,[],[f12676,f33897])).
fof(f35128,plain,(
  ( ! [X147] : (s(bool,'T_0') != s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),'|dollar|equals_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))) | s(bool,'T_0') = s(bool,X147)) )),
  inference(subsumption_resolution,[],[f34760,f1959])).
fof(f34760,plain,(
  ( ! [X147] : (s(bool,'T_0') != s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),'|dollar|equals_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))) | s(bool,'T_0') = s(bool,X147) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X147)))) )),
  inference(superposition,[],[f1959,f33897])).
fof(f35089,plain,(
  ( ! [X75] : (s(bool,'T_0') != s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),'|dollar|equals_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))) | s(bool,'T_0') = s(bool,X75)) )),
  inference(subsumption_resolution,[],[f34715,f15904])).
fof(f34715,plain,(
  ( ! [X75] : (s(bool,'T_0') != s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),'|dollar|equals_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X75))) | s(bool,'T_0') = s(bool,X75)) )),
  inference(superposition,[],[f1319,f33897])).
fof(f35088,plain,(
  ( ! [X74] : (s(bool,'T_0') != s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),'|dollar|equals_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))) | s(bool,'T_0') = s(bool,X74)) )),
  inference(subsumption_resolution,[],[f34714,f12676])).
fof(f34714,plain,(
  ( ! [X74] : (s(bool,'T_0') != s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),'|dollar|equals_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X74))) | s(bool,'T_0') = s(bool,X74)) )),
  inference(superposition,[],[f1318,f33897])).
fof(f35082,plain,(
  ( ! [X72] : (s(bool,'T_0') != s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),'|dollar|equals_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))) | s(bool,'T_0') = s(bool,X72)) ) | $spl7),
  inference(subsumption_resolution,[],[f34712,f1959])).
fof(f34712,plain,(
  ( ! [X72] : (s(bool,'T_0') != s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),'|dollar|equals_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X72))) | s(bool,'T_0') = s(bool,X72)) ) | $spl7),
  inference(superposition,[],[f1217,f33897])).
fof(f20790,plain,(
  $false | ($spl8 | $spl15)),
  inference(subsumption_resolution,[],[f20075,f2173])).
fof(f20075,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'T_0'))) | ($spl8 | $spl15)),
  inference(backward_demodulation,[],[f2173,f16828])).
fof(f20789,plain,(
  $false | ($spl8 | $spl15)),
  inference(subsumption_resolution,[],[f20052,f2173])).
fof(f20052,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'T_0'))) | ($spl8 | $spl15)),
  inference(backward_demodulation,[],[f2173,f13226])).
fof(f20788,plain,(
  $false | ($spl8 | $spl11)),
  inference(subsumption_resolution,[],[f19941,f2173])).
fof(f19941,plain,(
  s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | ($spl8 | $spl11)),
  inference(backward_demodulation,[],[f2173,f2175])).
fof(f20787,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f19924,f130])).
fof(f19924,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(backward_demodulation,[],[f2173,f133])).
fof(f19837,plain,(
  p(s(bool,happ(s(fun(bool,bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|dollar|equals_0'),s(fun(bool,bool),'const.combin.I_0'),s(fun(bool,bool),'const.combin.I_0'))),s(bool,'F_0')))) | $spl22),
  inference(cnf_transformation,[],[f19837_D])).
fof(f19837_D,plain,(
  p(s(bool,happ(s(fun(bool,bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|dollar|equals_0'),s(fun(bool,bool),'const.combin.I_0'),s(fun(bool,bool),'const.combin.I_0'))),s(bool,'F_0')))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f19839,plain,(
  ( ! [X90] : (s(bool,'T_0') = s(bool,X90) | p(s(bool,happ(s(fun(bool,bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|dollar|equals_0'),s(fun(bool,bool),'const.combin.I_0'),s(fun(bool,bool),'const.combin.I_0'))),s(bool,'F_0'))))) )),
  inference(forward_demodulation,[],[f19838,f138])).
fof(f19838,plain,(
  ( ! [X90] : (p(s(bool,happ(s(fun(bool,bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|dollar|equals_0'),s(fun(bool,bool),'const.combin.I_0'),s(fun(bool,bool),'const.combin.I_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X90)))) )),
  inference(forward_demodulation,[],[f19683,f175])).
fof(f175,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun(X3,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X2,X0)),X4),s(fun(X3,X1),X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,happ(s(fun(X3,X1),X5),s(X3,X7))))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))) )),
  inference(cnf_transformation,[],[f97])).
fof(f97,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun(X3,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X2,X0)),X4),s(fun(X3,X1),X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,happ(s(fun(X3,X1),X5),s(X3,X7))))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))),
  inference(flattening,[],[f96])).
fof(f96,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6,X7] : s(X0,happ(s(fun(X3,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X2,X0)),X4),s(fun(X3,X1),X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,happ(s(fun(X3,X1),X5),s(X3,X7))))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X41,X42,X45,X39] : ! [X43,X40,X46,X44] : s(X41,happ(s(fun(X39,X41),'_dst_x0x1x1_3'(s(fun(X42,fun(X45,X41)),X43),s(fun(X39,X42),X40),s(fun(X39,X45),X46))),s(X39,X44))) = s(X41,happ(s(fun(X45,X41),happ(s(fun(X42,fun(X45,X41)),X43),s(X42,happ(s(fun(X39,X42),X40),s(X39,X44))))),s(X45,happ(s(fun(X39,X45),X46),s(X39,X44)))))),
  file('thm.realax.REAL_ADD_SYM','_dst_x0x1x1_ax')).
fof(f19683,plain,(
  ( ! [X90] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|dollar|equals_0'),s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))))),s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0')))))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X90)))) )),
  inference(superposition,[],[f19248,f16703])).
fof(f16703,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X0))) | s(bool,X0) = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0')))) )),
  inference(condensation,[],[f16702])).
fof(f16702,plain,(
  ( ! [X12,X11] : (s(bool,X12) = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X12)))) )),
  inference(duplicate_literal_removal,[],[f16296])).
fof(f16296,plain,(
  ( ! [X12,X11] : (s(bool,X12) = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X12)))) )),
  inference(superposition,[],[f15768,f263])).
fof(f19835,plain,(
  ( ! [X89] : (p(s(bool,happ(s(fun(bool,bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|dollar|equals_0'),s(fun(bool,bool),'const.combin.I_0'),s(fun(bool,bool),'const.combin.I_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X89)) )),
  inference(forward_demodulation,[],[f19682,f175])).
fof(f19682,plain,(
  ( ! [X89] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|dollar|equals_0'),s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))))),s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0')))))) | s(bool,'T_0') = s(bool,X89)) )),
  inference(superposition,[],[f19248,f15768])).
fof(f16784,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | $spl21),
  inference(cnf_transformation,[],[f16784_D])).
fof(f16784_D,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f16825,plain,(
  ( ! [X185] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X185)) )),
  inference(subsumption_resolution,[],[f16490,f12676])).
fof(f16490,plain,(
  ( ! [X185] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X185) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X185)))) )),
  inference(superposition,[],[f12676,f15768])).
fof(f16820,plain,(
  ( ! [X137] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X137)) )),
  inference(subsumption_resolution,[],[f16477,f1959])).
fof(f16477,plain,(
  ( ! [X137] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X137) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X137)))) )),
  inference(superposition,[],[f1959,f15768])).
fof(f16801,plain,(
  ( ! [X110] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X110)) )),
  inference(forward_demodulation,[],[f16800,f139])).
fof(f16800,plain,(
  ( ! [X110] : (s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X110)) )),
  inference(subsumption_resolution,[],[f16459,f1959])).
fof(f16459,plain,(
  ( ! [X110] : (s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X110))) | s(bool,'T_0') = s(bool,X110)) )),
  inference(superposition,[],[f9315,f15768])).
fof(f9315,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X0)))) )),
  inference(condensation,[],[f9314])).
fof(f9314,plain,(
  ( ! [X12,X11] : (s(bool,X12) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X12)))) )),
  inference(duplicate_literal_removal,[],[f9223])).
fof(f9223,plain,(
  ( ! [X12,X11] : (s(bool,X12) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X12)))) )),
  inference(superposition,[],[f8617,f263])).
fof(f8617,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(duplicate_literal_removal,[],[f8599])).
fof(f8599,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f264])).
fof(f264,plain,(
  ( ! [X24,X25] : (s(bool,X24) = s(bool,X25) | s(bool,'T_0') = s(bool,X25) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X24)))) )),
  inference(superposition,[],[f187,f139])).
fof(f16799,plain,(
  ( ! [X109] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X109)) )),
  inference(forward_demodulation,[],[f16638,f139])).
fof(f16638,plain,(
  ( ! [X109] : (s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X109)) )),
  inference(duplicate_literal_removal,[],[f16458])).
fof(f16458,plain,(
  ( ! [X109] : (s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X109) | s(bool,'T_0') = s(bool,X109)) )),
  inference(superposition,[],[f8617,f15768])).
fof(f16789,plain,(
  ( ! [X75] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X75)) )),
  inference(subsumption_resolution,[],[f16439,f15904])).
fof(f16439,plain,(
  ( ! [X75] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X75))) | s(bool,'T_0') = s(bool,X75)) )),
  inference(superposition,[],[f1319,f15768])).
fof(f16788,plain,(
  ( ! [X74] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X74)) )),
  inference(subsumption_resolution,[],[f16438,f12676])).
fof(f16438,plain,(
  ( ! [X74] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X74))) | s(bool,'T_0') = s(bool,X74)) )),
  inference(superposition,[],[f1318,f15768])).
fof(f16782,plain,(
  ( ! [X72] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X72)) ) | $spl7),
  inference(subsumption_resolution,[],[f16436,f1959])).
fof(f16436,plain,(
  ( ! [X72] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X72))) | s(bool,'T_0') = s(bool,X72)) ) | $spl7),
  inference(superposition,[],[f1217,f15768])).
fof(f13309,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0')))) ) | $spl18),
  inference(cnf_transformation,[],[f13309_D])).
fof(f13309_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0')))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f13226,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | $spl15),
  inference(cnf_transformation,[],[f13226_D])).
fof(f13226_D,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f13084,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,X0) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0')))) )),
  inference(equality_factoring,[],[f12570])).
fof(f12570,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(condensation,[],[f12569])).
fof(f12569,plain,(
  ( ! [X0,X1] : (s(bool,X1) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f12170])).
fof(f12170,plain,(
  ( ! [X0,X1] : (s(bool,X1) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f289,f137])).
fof(f13261,plain,(
  ( ! [X129] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X129)) )),
  inference(subsumption_resolution,[],[f12957,f1959])).
fof(f12957,plain,(
  ( ! [X129] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X129) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X129)))) )),
  inference(superposition,[],[f1959,f12570])).
fof(f13232,plain,(
  ( ! [X71] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X71)))) ) | $spl16),
  inference(cnf_transformation,[],[f13232_D])).
fof(f13232_D,plain,(
  ( ! [X71] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X71)))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f13231,plain,(
  ( ! [X71] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X71)))) )),
  inference(subsumption_resolution,[],[f12923,f1319])).
fof(f12923,plain,(
  ( ! [X71] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X71))) | s(bool,'T_0') = s(bool,X71)) )),
  inference(superposition,[],[f1319,f12570])).
fof(f13230,plain,(
  ( ! [X70] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X70)) )),
  inference(subsumption_resolution,[],[f12922,f12676])).
fof(f12922,plain,(
  ( ! [X70] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X70))) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f1318,f12570])).
fof(f13224,plain,(
  ( ! [X68] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X68)) ) | $spl7),
  inference(subsumption_resolution,[],[f12920,f1959])).
fof(f12920,plain,(
  ( ! [X68] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X68))) | s(bool,'T_0') = s(bool,X68)) ) | $spl7),
  inference(superposition,[],[f1217,f12570])).
fof(f2188,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | $spl12),
  inference(cnf_transformation,[],[f2188_D])).
fof(f2188_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f2175,plain,(
  s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | $spl11),
  inference(cnf_transformation,[],[f2175_D])).
fof(f2175_D,plain,(
  s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'F_0'))) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f2117,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X0) = s(bool,'const.combin.I_1'(s(bool,'F_0')))) )),
  inference(equality_factoring,[],[f1923])).
fof(f1923,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(condensation,[],[f1922])).
fof(f1922,plain,(
  ( ! [X0,X1] : (s(bool,X1) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f1759])).
fof(f1759,plain,(
  ( ! [X0,X1] : (s(bool,X1) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f288,f137])).
fof(f2172,plain,(
  ( ! [X30] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X30)) ) | $spl7),
  inference(subsumption_resolution,[],[f2064,f1959])).
fof(f2064,plain,(
  ( ! [X30] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X30))) | s(bool,'T_0') = s(bool,X30)) ) | $spl7),
  inference(superposition,[],[f1217,f1923])).
fof(f743,plain,(
  $false | $spl1),
  inference(subsumption_resolution,[],[f742,f130])).
fof(f742,plain,(
  ~p(s(bool,'T_0')) | $spl1),
  inference(backward_demodulation,[],[f731,f185])).
fof(f185,plain,(
  ~p(s(bool,'const.realax.treal_eq_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_add_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK2),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK3))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_add_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK2),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK3)))))) | $spl1),
  inference(forward_demodulation,[],[f184,f147])).
fof(f147,plain,(
  ( ! [X0,X1] : (s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_add_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X0),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1))) = s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal')),happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal'))),'const.realax.treal_add_0'),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X0))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1)))) )),
  inference(cnf_transformation,[],[f60])).
fof(f60,plain,(
  ! [X0,X1] : s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_add_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X0),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1))) = s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal')),happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal'))),'const.realax.treal_add_0'),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X0))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1)))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X35,X36] : s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_add_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X35),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X36))) = s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal')),happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal'))),'const.realax.treal_add_0'),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X35))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X36)))),
  file('thm.realax.REAL_ADD_SYM','ARITY_0_2_const.realax.treal_add')).
fof(f184,plain,(
  ~p(s(bool,'const.realax.treal_eq_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal')),happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal'))),'const.realax.treal_add_0'),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK2))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK3))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_add_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK2),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK3)))))) | $spl1),
  inference(forward_demodulation,[],[f183,f144])).
fof(f183,plain,(
  ~p(s(bool,'const.realax.treal_eq_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal')),happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal'))),'const.realax.treal_add_0'),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK2))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK3))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_add_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK3),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK2)))))) | $spl1),
  inference(forward_demodulation,[],[f182,f147])).
fof(f182,plain,(
  ~p(s(bool,'const.realax.treal_eq_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal')),happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal'))),'const.realax.treal_add_0'),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK2))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK3))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal')),happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal'))),'const.realax.treal_add_0'),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK3))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK2)))))) | $spl1),
  inference(forward_demodulation,[],[f178,f146])).
fof(f146,plain,(
  ( ! [X0,X1] : (s(bool,'const.realax.treal_eq_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X0),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1))) = s(bool,happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool)),'const.realax.treal_eq_0'),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X0))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1)))) )),
  inference(cnf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0,X1] : s(bool,'const.realax.treal_eq_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X0),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1))) = s(bool,happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool)),'const.realax.treal_eq_0'),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X0))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1)))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X35,X36] : s(bool,'const.realax.treal_eq_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X35),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X36))) = s(bool,happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool)),'const.realax.treal_eq_0'),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X35))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X36)))),
  file('thm.realax.REAL_ADD_SYM','ARITY_0_2_const.realax.treal_eq')).
fof(f731,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.realax.treal_eq_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X0),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X0)))) )),
  inference(unit_resulting_resolution,[],[f711,f190])).
fof(f711,plain,(
  ( ! [X0] : (p(s(bool,'const.realax.treal_eq_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X0),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X0))))) )),
  inference(forward_demodulation,[],[f697,f181])).
fof(f697,plain,(
  ( ! [X0] : (p(s(bool,'const.realax.treal_eq_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal')),'const.combin.I_0'),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X0))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X0))))) )),
  inference(unit_resulting_resolution,[],[f181,f148])).
fof(f148,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.realax.treal_eq_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X0),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1)))) | s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X0) != s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1)) )),
  inference(cnf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0,X1] : (s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X0) != s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1) | p(s(bool,'const.realax.treal_eq_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X0),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1)))))),
  inference(ennf_transformation,[],[f61])).
fof(f61,plain,(
  ! [X0,X1] : (s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X0) = s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1) => p(s(bool,'const.realax.treal_eq_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X0),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1)))))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X31,X32] : (s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X31) = s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X32) => p(s(bool,'const.realax.treal_eq_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X31),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X32)))))),
  file('thm.realax.REAL_ADD_SYM','thm.realax.TREAL_EQ_AP')).
fof(f197,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl4),
  inference(cnf_transformation,[],[f197_D])).
fof(f197_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f199,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl7),
  inference(cnf_transformation,[],[f199_D])).
fof(f199_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f196,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f137])).
fof(f178,plain,(
  ~p(s(bool,happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool)),'const.realax.treal_eq_0'),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal')),happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal'))),'const.realax.treal_add_0'),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK2))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK3))))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal')),happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal'))),'const.realax.treal_add_0'),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK3))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK2)))))) | $spl1),
  inference(cnf_transformation,[],[f178_D])).
fof(f178_D,plain,(
  ~p(s(bool,happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool)),'const.realax.treal_eq_0'),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal')),happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal'))),'const.realax.treal_add_0'),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK2))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK3))))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal')),happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal'))),'const.realax.treal_add_0'),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK3))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK2)))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f180,plain,(
  p(s(bool,happ(s(fun(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),bool),'const.bool.RES_FORALL_1'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.quotient.respects_1'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool)),'const.realax.treal_eq_0'))))),s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'_dst_x0x1_2'(s(fun(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),bool),'const.bool.RES_FORALL_1'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.quotient.respects_1'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool)),'const.realax.treal_eq_0'))))),s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool)),'_dst_x00x11x11_3'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool)),'const.realax.treal_eq_0'),s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal'))),'const.realax.treal_add_0'),s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal'))),'_dst_x01x10_2'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal'))),'const.realax.treal_add_0'),s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal')),'_i_0_0')))))))))) | $spl2),
  inference(cnf_transformation,[],[f180_D])).
fof(f180_D,plain,(
  p(s(bool,happ(s(fun(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),bool),'const.bool.RES_FORALL_1'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.quotient.respects_1'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool)),'const.realax.treal_eq_0'))))),s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'_dst_x0x1_2'(s(fun(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),bool),'const.bool.RES_FORALL_1'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.quotient.respects_1'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool)),'const.realax.treal_eq_0'))))),s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool)),'_dst_x00x11x11_3'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool)),'const.realax.treal_eq_0'),s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal'))),'const.realax.treal_add_0'),s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal'))),'_dst_x01x10_2'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal'))),'const.realax.treal_add_0'),s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal')),'_i_0_0')))))))))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f132,plain,(
  p(s(bool,happ(s(fun(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),bool),'const.bool.RES_FORALL_1'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.quotient.respects_1'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool)),'const.realax.treal_eq_0'))))),s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'_dst_x0x1_2'(s(fun(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),bool),'const.bool.RES_FORALL_1'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.quotient.respects_1'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool)),'const.realax.treal_eq_0'))))),s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool)),'_dst_x00x11x11_3'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool)),'const.realax.treal_eq_0'),s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal'))),'const.realax.treal_add_0'),s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal'))),'_dst_x01x10_2'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal'))),'const.realax.treal_add_0'),s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal')),'_i_0_0')))))))))) | ~p(s(bool,happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool)),'const.realax.treal_eq_0'),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal')),happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal'))),'const.realax.treal_add_0'),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK2))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK3))))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal')),happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal'))),'const.realax.treal_add_0'),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK3))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK2))))))),
  inference(cnf_transformation,[],[f120])).
fof(f120,plain,(
  ~p(s(bool,happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool)),'const.realax.treal_eq_0'),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal')),happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal'))),'const.realax.treal_add_0'),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK2))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK3))))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal')),happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal'))),'const.realax.treal_add_0'),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK3))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK2)))))) | p(s(bool,happ(s(fun(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),bool),'const.bool.RES_FORALL_1'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.quotient.respects_1'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool)),'const.realax.treal_eq_0'))))),s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'_dst_x0x1_2'(s(fun(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),bool),'const.bool.RES_FORALL_1'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.quotient.respects_1'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool)),'const.realax.treal_eq_0'))))),s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool)),'_dst_x00x11x11_3'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool)),'const.realax.treal_eq_0'),s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal'))),'const.realax.treal_add_0'),s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal'))),'_dst_x01x10_2'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal'))),'const.realax.treal_add_0'),s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal')),'_i_0_0'))))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2,sK3])],[f101])).
fof(f101,plain,(
  ? [X0,X1] : ~p(s(bool,happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool)),'const.realax.treal_eq_0'),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal')),happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal'))),'const.realax.treal_add_0'),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X0))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1))))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal')),happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal'))),'const.realax.treal_add_0'),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X0)))))) | p(s(bool,happ(s(fun(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),bool),'const.bool.RES_FORALL_1'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.quotient.respects_1'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool)),'const.realax.treal_eq_0'))))),s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'_dst_x0x1_2'(s(fun(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),bool),'const.bool.RES_FORALL_1'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.quotient.respects_1'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool)),'const.realax.treal_eq_0'))))),s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool)),'_dst_x00x11x11_3'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool)),'const.realax.treal_eq_0'),s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal'))),'const.realax.treal_add_0'),s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal'))),'_dst_x01x10_2'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal'))),'const.realax.treal_add_0'),s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal')),'_i_0_0'))))))))))),
  inference(ennf_transformation,[],[f42])).
fof(f42,plain,(
  ! [X0,X1] : p(s(bool,happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool)),'const.realax.treal_eq_0'),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal')),happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal'))),'const.realax.treal_add_0'),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X0))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1))))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal')),happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal'))),'const.realax.treal_add_0'),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X0)))))) => p(s(bool,happ(s(fun(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),bool),'const.bool.RES_FORALL_1'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.quotient.respects_1'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool)),'const.realax.treal_eq_0'))))),s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'_dst_x0x1_2'(s(fun(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),bool),'const.bool.RES_FORALL_1'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.quotient.respects_1'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool)),'const.realax.treal_eq_0'))))),s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool)),'_dst_x00x11x11_3'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool)),'const.realax.treal_eq_0'),s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal'))),'const.realax.treal_add_0'),s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal'))),'_dst_x01x10_2'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal'))),'const.realax.treal_add_0'),s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal')),'_i_0_0'))))))))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X6,X7] : p(s(bool,happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool)),'const.realax.treal_eq_0'),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal')),happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal'))),'const.realax.treal_add_0'),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X6))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X7))))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal')),happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal'))),'const.realax.treal_add_0'),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X7))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X6)))))) => p(s(bool,happ(s(fun(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),bool),'const.bool.RES_FORALL_1'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.quotient.respects_1'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool)),'const.realax.treal_eq_0'))))),s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'_dst_x0x1_2'(s(fun(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),bool),'const.bool.RES_FORALL_1'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.quotient.respects_1'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool)),'const.realax.treal_eq_0'))))),s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool)),'_dst_x00x11x11_3'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool)),'const.realax.treal_eq_0'),s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal'))),'const.realax.treal_add_0'),s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal'))),'_dst_x01x10_2'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal'))),'const.realax.treal_add_0'),s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'type.pair.prod'('type.hreal.hreal','type.hreal.hreal')),'_i_0_0'))))))))))),
  file('thm.realax.REAL_ADD_SYM','thm.realax.realax_82_prove_ltb')).
% SZS output end Proof for thm.realax
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 149
% Passive clauses: 2106
% Generated clauses: 36941
% Final active clauses: 96
% Final passive clauses: 1145
% Input formulas: 39
% Initial clauses: 48
% 
% Duplicate literals: 1644
% Trivial inequalities: 289
% Fw subsumption resolutions: 2975
% Fw demodulations: 6778
% Bw demodulations: 838
% Condensations: 3596
% 
% Simple tautologies: 316
% Equational tautologies: 228
% Forward subsumptions: 18936
% Fw demodulations to eq. taut.: 39
% Bw demodulations to eq. taut.: 25
% 
% Unit resulting resolution: 1457
% Forward superposition: 10609
% Backward superposition: 7490
% Self superposition: 284
% Equality factoring: 103
% Equality resolution: 25
% 
% Split clauses: 32
% Split components: 64
% SAT solver clauses: 2181
% SAT solver unit clauses: 230
% SAT solver binary clauses: 794
% 
% Sat splits: 32
% Sat splitting refutations: 9
% 
% Memory used [KB]: 23794
% Time elapsed: 2.015 s
% ------------------------------
% ------------------------------
% Success in time 3.329 s
