Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.while
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.while
% SZS output start Proof for thm.while
fof(f27405,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f27292,f26880,f26877,f350,f602])).
fof(f602,plain,(
  ~p(s(bool,'F_0')) | $spl3),
  inference(cnf_transformation,[],[f602_D])).
fof(f602_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f350,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f74])).
fof(f74,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.while.OLEAST_EQNS|split|2',aHL_FALSITY)).
fof(f26877,plain,(
  s(bool,'F_0') = s(bool,sK18('_dst_x0_1'(s(bool,'F_0')))) | $spl538),
  inference(cnf_transformation,[],[f26877_D])).
fof(f26877_D,plain,(
  s(bool,'F_0') = s(bool,sK18('_dst_x0_1'(s(bool,'F_0')))) <=> ~$spl538),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl538])])).
fof(f26880,plain,(
  s(bool,'F_0') = s(bool,sK18('_dst_x0_1'(s(bool,'F_0')))) | $spl3),
  inference(forward_demodulation,[],[f26879,f4386])).
fof(f4386,plain,(
  s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))),
  inference(resolution,[],[f4382,f367])).
fof(f367,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f233])).
fof(f233,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f88])).
fof(f88,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f87])).
fof(f87,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X13] : (s(bool,'F_0') = s(bool,X13) <=> ~p(s(bool,X13)))),
  file('thm.while.OLEAST_EQNS|split|2','thm.bool.EQ_CLAUSES|split|3')).
fof(f4382,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))),
  inference(equality_resolution,[],[f404])).
fof(f404,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f257])).
fof(f257,plain,(
  ! [X0] : ((s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))) & (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0)))),
  inference(nnf_transformation,[],[f104])).
fof(f104,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) <=> p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))))),
  inference(flattening,[],[f103])).
fof(f103,plain,(
  ! [X0] : (~s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) <=> p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))))),
  inference(rectify,[],[f54])).
fof(f54,axiom,(
  ! [X8] : (~s('type.num.num','const.num.|48|_0') = s('type.num.num',X8) <=> p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X8)))))),
  file('thm.while.OLEAST_EQNS|split|2','thm.arithmetic.NOT_ZERO_LT_ZERO')).
fof(f26879,plain,(
  s(bool,'F_0') = s(bool,sK18('_dst_x0_1'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))))) | $spl3),
  inference(subsumption_resolution,[],[f26855,f602])).
fof(f26855,plain,(
  p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,sK18('_dst_x0_1'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))))),
  inference(superposition,[],[f26588,f4386])).
fof(f26588,plain,(
  ( ! [X9] : (p(s(bool,X9)) | s(bool,'F_0') = s(bool,sK18('_dst_x0_1'(s(bool,X9))))) )),
  inference(superposition,[],[f4743,f486])).
fof(f486,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f154])).
fof(f154,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))),
  inference(flattening,[],[f153])).
fof(f153,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X14,X15] : ! [X16,X17] : s(X15,X16) = s(X15,happ(s(fun(X14,X15),'_dst_x0_1'(s(X15,X16))),s(X14,X17)))),
  file('thm.while.OLEAST_EQNS|split|2','_dst_x0_ax')).
fof(f4743,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',sK19(X0))))) | s(bool,'F_0') = s(bool,sK18(X0))) )),
  inference(resolution,[],[f435,f367])).
fof(f435,plain,(
  ( ! [X0] : (~p(s(bool,sK18(X0))) | p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',sK19(X0)))))) )),
  inference(cnf_transformation,[],[f279])).
fof(f279,plain,(
  ! [X0] : ((~p(s(bool,sK18(X0))) | p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',sK19(X0)))))) & (! [X2] : ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2)))) | p(s(bool,sK18(X0)))) & s('type.option.option'('type.num.num'),'const.while.OLEAST_1'(s(fun('type.num.num',bool),X0))) = s('type.option.option'('type.num.num'),'const.bool.COND_3'(s(bool,sK18(X0)),s('type.option.option'('type.num.num'),'const.option.SOME_1'(s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),X0))))),s('type.option.option'('type.num.num'),'const.option.NONE_0'))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK18,sK19])],[f278])).
fof(f278,plain,(
  ! [X0] : ? [X1] : ((~p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X3))))) & (! [X2] : ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2)))) | p(s(bool,X1))) & s('type.option.option'('type.num.num'),'const.while.OLEAST_1'(s(fun('type.num.num',bool),X0))) = s('type.option.option'('type.num.num'),'const.bool.COND_3'(s(bool,X1),s('type.option.option'('type.num.num'),'const.option.SOME_1'(s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),X0))))),s('type.option.option'('type.num.num'),'const.option.NONE_0'))))),
  inference(rectify,[],[f277])).
fof(f277,plain,(
  ! [X0] : ? [X1] : ((~p(s(bool,X1)) | ? [X2] : p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2))))) & (! [X2] : ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2)))) | p(s(bool,X1))) & s('type.option.option'('type.num.num'),'const.while.OLEAST_1'(s(fun('type.num.num',bool),X0))) = s('type.option.option'('type.num.num'),'const.bool.COND_3'(s(bool,X1),s('type.option.option'('type.num.num'),'const.option.SOME_1'(s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),X0))))),s('type.option.option'('type.num.num'),'const.option.NONE_0'))))),
  inference(flattening,[],[f276])).
fof(f276,plain,(
  ! [X0] : ? [X1] : (((~p(s(bool,X1)) | ? [X2] : p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2))))) & (! [X2] : ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2)))) | p(s(bool,X1)))) & s('type.option.option'('type.num.num'),'const.while.OLEAST_1'(s(fun('type.num.num',bool),X0))) = s('type.option.option'('type.num.num'),'const.bool.COND_3'(s(bool,X1),s('type.option.option'('type.num.num'),'const.option.SOME_1'(s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),X0))))),s('type.option.option'('type.num.num'),'const.option.NONE_0'))))),
  inference(nnf_transformation,[],[f135])).
fof(f135,plain,(
  ! [X0] : ? [X1] : ((p(s(bool,X1)) <=> ? [X2] : p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2))))) & s('type.option.option'('type.num.num'),'const.while.OLEAST_1'(s(fun('type.num.num',bool),X0))) = s('type.option.option'('type.num.num'),'const.bool.COND_3'(s(bool,X1),s('type.option.option'('type.num.num'),'const.option.SOME_1'(s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),X0))))),s('type.option.option'('type.num.num'),'const.option.NONE_0'))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X7] : ? [X10] : ((p(s(bool,X10)) <=> ? [X8] : p(s(bool,happ(s(fun('type.num.num',bool),X7),s('type.num.num',X8))))) & s('type.option.option'('type.num.num'),'const.while.OLEAST_1'(s(fun('type.num.num',bool),X7))) = s('type.option.option'('type.num.num'),'const.bool.COND_3'(s(bool,X10),s('type.option.option'('type.num.num'),'const.option.SOME_1'(s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),X7))))),s('type.option.option'('type.num.num'),'const.option.NONE_0'))))),
  file('thm.while.OLEAST_EQNS|split|2','thm.while.OLEAST_def')).
fof(f27292,plain,(
  $false | $spl538),
  inference(subsumption_resolution,[],[f27291,f343])).
fof(f343,plain,(
  s('type.option.option'('type.num.num'),'const.option.NONE_0') != s('type.option.option'('type.num.num'),'const.while.OLEAST_1'(s(fun('type.num.num',bool),'_dst_x0_1'(s(bool,'F_0')))))),
  inference(cnf_transformation,[],[f71])).
fof(f71,plain,(
  s('type.option.option'('type.num.num'),'const.option.NONE_0') != s('type.option.option'('type.num.num'),'const.while.OLEAST_1'(s(fun('type.num.num',bool),'_dst_x0_1'(s(bool,'F_0')))))),
  inference(flattening,[],[f6])).
fof(f6,negated_conjecture,(
  ~s('type.option.option'('type.num.num'),'const.option.NONE_0') = s('type.option.option'('type.num.num'),'const.while.OLEAST_1'(s(fun('type.num.num',bool),'_dst_x0_1'(s(bool,'F_0')))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  s('type.option.option'('type.num.num'),'const.option.NONE_0') = s('type.option.option'('type.num.num'),'const.while.OLEAST_1'(s(fun('type.num.num',bool),'_dst_x0_1'(s(bool,'F_0')))))),
  file('thm.while.OLEAST_EQNS|split|2','thm.while.OLEAST_EQNS|split|2')).
fof(f27291,plain,(
  s('type.option.option'('type.num.num'),'const.option.NONE_0') = s('type.option.option'('type.num.num'),'const.while.OLEAST_1'(s(fun('type.num.num',bool),'_dst_x0_1'(s(bool,'F_0'))))) | $spl538),
  inference(forward_demodulation,[],[f26926,f419])).
fof(f419,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0,X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f121])).
fof(f121,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f65])).
fof(f65,axiom,(
  ! [X12] : ! [X21,X22] : s(X12,X22) = s(X12,'const.bool.COND_3'(s(bool,'F_0'),s(X12,X21),s(X12,X22)))),
  file('thm.while.OLEAST_EQNS|split|2','thm.bool.COND_CLAUSES|split|1')).
fof(f26926,plain,(
  s('type.option.option'('type.num.num'),'const.while.OLEAST_1'(s(fun('type.num.num',bool),'_dst_x0_1'(s(bool,'F_0'))))) = s('type.option.option'('type.num.num'),'const.bool.COND_3'(s(bool,'F_0'),s('type.option.option'('type.num.num'),'const.option.SOME_1'(s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),'_dst_x0_1'(s(bool,'F_0'))))))),s('type.option.option'('type.num.num'),'const.option.NONE_0'))) | $spl538),
  inference(superposition,[],[f437,f26877])).
fof(f437,plain,(
  ( ! [X0] : (s('type.option.option'('type.num.num'),'const.while.OLEAST_1'(s(fun('type.num.num',bool),X0))) = s('type.option.option'('type.num.num'),'const.bool.COND_3'(s(bool,sK18(X0)),s('type.option.option'('type.num.num'),'const.option.SOME_1'(s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),X0))))),s('type.option.option'('type.num.num'),'const.option.NONE_0')))) )),
  inference(cnf_transformation,[],[f279])).
% SZS output end Proof for thm.while
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 864
% Passive clauses: 9872
% Generated clauses: 26730
% Final active clauses: 858
% Final passive clauses: 8698
% Input formulas: 69
% Initial clauses: 248
% 
% Duplicate literals: 573
% Trivial inequalities: 60
% Fw subsumption resolutions: 1917
% Fw demodulations: 1940
% Bw demodulations: 5
% 
% Simple tautologies: 285
% Equational tautologies: 39
% Forward subsumptions: 11461
% 
% Binary resolution: 8630
% Factoring: 1100
% Forward superposition: 4682
% Backward superposition: 7270
% Self superposition: 40
% Equality factoring: 26
% Equality resolution: 15
% 
% Split clauses: 644
% Split components: 1290
% SAT solver clauses: 903
% SAT solver unit clauses: 14
% SAT solver binary clauses: 142
% SAT solver learnt clauses: 1
% 
% Sat splits: 644
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 902
% TWLsolver calls for satisfiability: 380
% 
% Memory used [KB]: 16247
% Time elapsed: 0.591 s
% ------------------------------
% ------------------------------
% Success in time 0.6 s
