Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.ringNorm
% (106002)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 127
% Passive clauses: 1287
% Generated clauses: 24121
% Final active clauses: 123
% Final passive clauses: 1078
% Input formulas: 24
% Initial clauses: 25
% 
% Duplicate literals: 1632
% Trivial inequalities: 35
% Fw subsumption resolutions: 2437
% Fw demodulations: 4083
% Bw demodulations: 2
% 
% Simple tautologies: 211
% Equational tautologies: 18
% Forward subsumptions: 12887
% 
% Binary resolution: 11
% Forward superposition: 9304
% Backward superposition: 6241
% Self superposition: 172
% Equality factoring: 34
% Equality resolution: 4
% 
% Split clauses: 657
% Split components: 1319
% SAT solver clauses: 1160
% SAT solver unit clauses: 1
% SAT solver binary clauses: 23
% 
% Sat splits: 657
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 1160
% TWLsolver calls for satisfiability: 89
% 
% Memory used [KB]: 15479
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.ringNorm
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.ringNorm
% SZS output start Proof for thm.ringNorm
fof(f6974,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f6967,f74,f100,f6819,f6859,f83,f109,f2565,f2566,f78,f106,f92,f118,f93,f119])).
fof(f119,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_scalar3_4'(s('type.ring.ring'(X0),X1),s(X0,X2),s('type.list.list'('type.quote.index'),X3),s('type.canonical.canonical_sum'(X0),X4))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),happ(s(fun(X0,fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)))),'const.ringNorm.r_canonical_sum_scalar3_1'(s('type.ring.ring'(X0),X1))),s(X0,X2))),s('type.list.list'('type.quote.index'),X3))),s('type.canonical.canonical_sum'(X0),X4)))) ) | $spl34),
  inference(cnf_transformation,[],[f119_D])).
fof(f119_D,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_scalar3_4'(s('type.ring.ring'(X0),X1),s(X0,X2),s('type.list.list'('type.quote.index'),X3),s('type.canonical.canonical_sum'(X0),X4))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),happ(s(fun(X0,fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)))),'const.ringNorm.r_canonical_sum_scalar3_1'(s('type.ring.ring'(X0),X1))),s(X0,X2))),s('type.list.list'('type.quote.index'),X3))),s('type.canonical.canonical_sum'(X0),X4)))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f93,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_scalar3_4'(s('type.ring.ring'(X0),X1),s(X0,X2),s('type.list.list'('type.quote.index'),X3),s('type.canonical.canonical_sum'(X0),X4))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),happ(s(fun(X0,fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)))),'const.ringNorm.r_canonical_sum_scalar3_1'(s('type.ring.ring'(X0),X1))),s(X0,X2))),s('type.list.list'('type.quote.index'),X3))),s('type.canonical.canonical_sum'(X0),X4)))) )),
  inference(cnf_transformation,[],[f58])).
fof(f58,plain,(
  ! [X0,X1,X2,X3,X4] : s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_scalar3_4'(s('type.ring.ring'(X0),X1),s(X0,X2),s('type.list.list'('type.quote.index'),X3),s('type.canonical.canonical_sum'(X0),X4))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),happ(s(fun(X0,fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)))),'const.ringNorm.r_canonical_sum_scalar3_1'(s('type.ring.ring'(X0),X1))),s(X0,X2))),s('type.list.list'('type.quote.index'),X3))),s('type.canonical.canonical_sum'(X0),X4)))),
  inference(flattening,[],[f57])).
fof(f57,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_scalar3_4'(s('type.ring.ring'(X0),X1),s(X0,X2),s('type.list.list'('type.quote.index'),X3),s('type.canonical.canonical_sum'(X0),X4))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),happ(s(fun(X0,fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)))),'const.ringNorm.r_canonical_sum_scalar3_1'(s('type.ring.ring'(X0),X1))),s(X0,X2))),s('type.list.list'('type.quote.index'),X3))),s('type.canonical.canonical_sum'(X0),X4)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X6] : ! [X20,X21,X22,X23] : s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_canonical_sum_scalar3_4'(s('type.ring.ring'(X6),X20),s(X6,X21),s('type.list.list'('type.quote.index'),X22),s('type.canonical.canonical_sum'(X6),X23))) = s('type.canonical.canonical_sum'(X6),happ(s(fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6)),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6))),happ(s(fun(X6,fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6)))),'const.ringNorm.r_canonical_sum_scalar3_1'(s('type.ring.ring'(X6),X20))),s(X6,X21))),s('type.list.list'('type.quote.index'),X22))),s('type.canonical.canonical_sum'(X6),X23)))),
  file('thm.ringNorm.canonical_sum_scalar3_def|split|2','ARITY_1_4_const.ringNorm.r_canonical_sum_scalar3')).
fof(f118,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar3_4'(s('type.semi_ring.semi_ring'(X0),X1),s(X0,X2),s('type.list.list'('type.quote.index'),X3),s('type.canonical.canonical_sum'(X0),X4))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),happ(s(fun(X0,fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)))),'const.canonical.canonical_sum_scalar3_1'(s('type.semi_ring.semi_ring'(X0),X1))),s(X0,X2))),s('type.list.list'('type.quote.index'),X3))),s('type.canonical.canonical_sum'(X0),X4)))) ) | $spl32),
  inference(cnf_transformation,[],[f118_D])).
fof(f118_D,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar3_4'(s('type.semi_ring.semi_ring'(X0),X1),s(X0,X2),s('type.list.list'('type.quote.index'),X3),s('type.canonical.canonical_sum'(X0),X4))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),happ(s(fun(X0,fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)))),'const.canonical.canonical_sum_scalar3_1'(s('type.semi_ring.semi_ring'(X0),X1))),s(X0,X2))),s('type.list.list'('type.quote.index'),X3))),s('type.canonical.canonical_sum'(X0),X4)))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f92,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar3_4'(s('type.semi_ring.semi_ring'(X0),X1),s(X0,X2),s('type.list.list'('type.quote.index'),X3),s('type.canonical.canonical_sum'(X0),X4))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),happ(s(fun(X0,fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)))),'const.canonical.canonical_sum_scalar3_1'(s('type.semi_ring.semi_ring'(X0),X1))),s(X0,X2))),s('type.list.list'('type.quote.index'),X3))),s('type.canonical.canonical_sum'(X0),X4)))) )),
  inference(cnf_transformation,[],[f56])).
fof(f56,plain,(
  ! [X0,X1,X2,X3,X4] : s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar3_4'(s('type.semi_ring.semi_ring'(X0),X1),s(X0,X2),s('type.list.list'('type.quote.index'),X3),s('type.canonical.canonical_sum'(X0),X4))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),happ(s(fun(X0,fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)))),'const.canonical.canonical_sum_scalar3_1'(s('type.semi_ring.semi_ring'(X0),X1))),s(X0,X2))),s('type.list.list'('type.quote.index'),X3))),s('type.canonical.canonical_sum'(X0),X4)))),
  inference(flattening,[],[f55])).
fof(f55,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar3_4'(s('type.semi_ring.semi_ring'(X0),X1),s(X0,X2),s('type.list.list'('type.quote.index'),X3),s('type.canonical.canonical_sum'(X0),X4))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),happ(s(fun(X0,fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)))),'const.canonical.canonical_sum_scalar3_1'(s('type.semi_ring.semi_ring'(X0),X1))),s(X0,X2))),s('type.list.list'('type.quote.index'),X3))),s('type.canonical.canonical_sum'(X0),X4)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X6] : ! [X20,X21,X22,X23] : s('type.canonical.canonical_sum'(X6),'const.canonical.canonical_sum_scalar3_4'(s('type.semi_ring.semi_ring'(X6),X20),s(X6,X21),s('type.list.list'('type.quote.index'),X22),s('type.canonical.canonical_sum'(X6),X23))) = s('type.canonical.canonical_sum'(X6),happ(s(fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6)),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6))),happ(s(fun(X6,fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6)))),'const.canonical.canonical_sum_scalar3_1'(s('type.semi_ring.semi_ring'(X6),X20))),s(X6,X21))),s('type.list.list'('type.quote.index'),X22))),s('type.canonical.canonical_sum'(X6),X23)))),
  file('thm.ringNorm.canonical_sum_scalar3_def|split|2','ARITY_1_4_const.canonical.canonical_sum_scalar3')).
fof(f106,plain,(
  ( ! [X0,X1] : (s(fun(X0,fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)))),'const.ringNorm.r_canonical_sum_scalar3_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)))),'const.canonical.canonical_sum_scalar3_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))) ) | $spl8),
  inference(cnf_transformation,[],[f106_D])).
fof(f106_D,plain,(
  ( ! [X0,X1] : (s(fun(X0,fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)))),'const.ringNorm.r_canonical_sum_scalar3_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)))),'const.canonical.canonical_sum_scalar3_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f78,plain,(
  ( ! [X0,X1] : (s(fun(X0,fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)))),'const.ringNorm.r_canonical_sum_scalar3_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)))),'const.canonical.canonical_sum_scalar3_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f30])).
fof(f30,plain,(
  ! [X0,X1] : s(fun(X0,fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)))),'const.ringNorm.r_canonical_sum_scalar3_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)))),'const.canonical.canonical_sum_scalar3_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f29])).
fof(f29,plain,(
  ! [X0] : ! [X1] : s(fun(X0,fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)))),'const.ringNorm.r_canonical_sum_scalar3_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)))),'const.canonical.canonical_sum_scalar3_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6] : ! [X7] : s(fun(X6,fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6)))),'const.ringNorm.r_canonical_sum_scalar3_1'(s('type.ring.ring'(X6),X7))) = s(fun(X6,fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6)))),'const.canonical.canonical_sum_scalar3_1'(s('type.semi_ring.semi_ring'(X6),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X6),X7)))))),
  file('thm.ringNorm.canonical_sum_scalar3_def|split|2','thm.ringNorm.r_canonical_sum_scalar3_def')).
fof(f2566,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_scalar3_4'(s('type.ring.ring'(X0),X1),s(X0,X2),s('type.list.list'('type.quote.index'),X3),s('type.canonical.canonical_sum'(X0),X4))) = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar3_4'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s(X0,X2),s('type.list.list'('type.quote.index'),X3),s('type.canonical.canonical_sum'(X0),X4)))) ) | $spl374),
  inference(cnf_transformation,[],[f2566_D])).
fof(f2566_D,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_scalar3_4'(s('type.ring.ring'(X0),X1),s(X0,X2),s('type.list.list'('type.quote.index'),X3),s('type.canonical.canonical_sum'(X0),X4))) = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar3_4'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s(X0,X2),s('type.list.list'('type.quote.index'),X3),s('type.canonical.canonical_sum'(X0),X4)))) ) <=> ~$spl374),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl374])])).
fof(f2565,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_scalar3_4'(s('type.ring.ring'(X0),X1),s(X0,X2),s('type.list.list'('type.quote.index'),X3),s('type.canonical.canonical_sum'(X0),X4))) = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar3_4'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s(X0,X2),s('type.list.list'('type.quote.index'),X3),s('type.canonical.canonical_sum'(X0),X4)))) ) | ($spl8 | $spl32 | $spl34)),
  inference(forward_demodulation,[],[f2490,f119])).
fof(f2490,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar3_4'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s(X0,X2),s('type.list.list'('type.quote.index'),X3),s('type.canonical.canonical_sum'(X0),X4))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),happ(s(fun(X0,fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)))),'const.ringNorm.r_canonical_sum_scalar3_1'(s('type.ring.ring'(X0),X1))),s(X0,X2))),s('type.list.list'('type.quote.index'),X3))),s('type.canonical.canonical_sum'(X0),X4)))) ) | ($spl8 | $spl32)),
  inference(superposition,[],[f118,f106])).
fof(f109,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar3_4'(s('type.semi_ring.semi_ring'(X0),X1),s(X0,X2),s('type.list.list'('type.quote.index'),X3),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))) ) | $spl14),
  inference(cnf_transformation,[],[f109_D])).
fof(f109_D,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar3_4'(s('type.semi_ring.semi_ring'(X0),X1),s(X0,X2),s('type.list.list'('type.quote.index'),X3),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f83,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar3_4'(s('type.semi_ring.semi_ring'(X0),X1),s(X0,X2),s('type.list.list'('type.quote.index'),X3),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))) )),
  inference(cnf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0,X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar3_4'(s('type.semi_ring.semi_ring'(X0),X1),s(X0,X2),s('type.list.list'('type.quote.index'),X3),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))),
  inference(flattening,[],[f37])).
fof(f37,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar3_4'(s('type.semi_ring.semi_ring'(X0),X1),s(X0,X2),s('type.list.list'('type.quote.index'),X3),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X6] : ! [X12,X8,X9] : s('type.canonical.canonical_sum'(X6),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'(X6),'const.canonical.canonical_sum_scalar3_4'(s('type.semi_ring.semi_ring'(X6),X12),s(X6,X8),s('type.list.list'('type.quote.index'),X9),s('type.canonical.canonical_sum'(X6),'const.canonical.Nil_monom_0')))),
  file('thm.ringNorm.canonical_sum_scalar3_def|split|2','thm.canonical.canonical_sum_scalar3_def|split|2')).
fof(f6859,plain,(
  ( ! [X6,X8,X7,X5] : (s('type.canonical.canonical_sum'(X5),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'(X5),'const.ringNorm.r_canonical_sum_scalar3_4'(s('type.ring.ring'(X5),X6),s(X5,X7),s('type.list.list'('type.quote.index'),X8),s('type.canonical.canonical_sum'(X5),'const.canonical.Nil_monom_0')))) ) | $spl682),
  inference(cnf_transformation,[],[f6859_D])).
fof(f6859_D,plain,(
  ( ! [X6,X8,X7,X5] : (s('type.canonical.canonical_sum'(X5),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'(X5),'const.ringNorm.r_canonical_sum_scalar3_4'(s('type.ring.ring'(X5),X6),s(X5,X7),s('type.list.list'('type.quote.index'),X8),s('type.canonical.canonical_sum'(X5),'const.canonical.Nil_monom_0')))) ) <=> ~$spl682),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl682])])).
fof(f6819,plain,(
  ( ! [X6,X8,X7,X5] : (s('type.canonical.canonical_sum'(X5),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'(X5),'const.ringNorm.r_canonical_sum_scalar3_4'(s('type.ring.ring'(X5),X6),s(X5,X7),s('type.list.list'('type.quote.index'),X8),s('type.canonical.canonical_sum'(X5),'const.canonical.Nil_monom_0')))) ) | ($spl14 | $spl374)),
  inference(superposition,[],[f2566,f109])).
fof(f100,plain,(
  s('type.canonical.canonical_sum'(sK0),'const.canonical.Nil_monom_0') != s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_canonical_sum_scalar3_4'(s('type.ring.ring'(sK0),sK1),s(sK0,sK2),s('type.list.list'('type.quote.index'),sK3),s('type.canonical.canonical_sum'(sK0),'const.canonical.Nil_monom_0'))) | $spl1),
  inference(cnf_transformation,[],[f100_D])).
fof(f100_D,plain,(
  s('type.canonical.canonical_sum'(sK0),'const.canonical.Nil_monom_0') != s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_canonical_sum_scalar3_4'(s('type.ring.ring'(sK0),sK1),s(sK0,sK2),s('type.list.list'('type.quote.index'),sK3),s('type.canonical.canonical_sum'(sK0),'const.canonical.Nil_monom_0'))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f74,plain,(
  s('type.canonical.canonical_sum'(sK0),'const.canonical.Nil_monom_0') != s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_canonical_sum_scalar3_4'(s('type.ring.ring'(sK0),sK1),s(sK0,sK2),s('type.list.list'('type.quote.index'),sK3),s('type.canonical.canonical_sum'(sK0),'const.canonical.Nil_monom_0')))),
  inference(cnf_transformation,[],[f71])).
fof(f71,plain,(
  s('type.canonical.canonical_sum'(sK0),'const.canonical.Nil_monom_0') != s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_canonical_sum_scalar3_4'(s('type.ring.ring'(sK0),sK1),s(sK0,sK2),s('type.list.list'('type.quote.index'),sK3),s('type.canonical.canonical_sum'(sK0),'const.canonical.Nil_monom_0')))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f69])).
fof(f69,plain,(
  ? [X0,X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0') != s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_scalar3_4'(s('type.ring.ring'(X0),X1),s(X0,X2),s('type.list.list'('type.quote.index'),X3),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))),
  inference(ennf_transformation,[],[f27])).
fof(f27,plain,(
  ~! [X0,X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_scalar3_4'(s('type.ring.ring'(X0),X1),s(X0,X2),s('type.list.list'('type.quote.index'),X3),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))),
  inference(flattening,[],[f26])).
fof(f26,plain,(
  ~! [X0] : ! [X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_scalar3_4'(s('type.ring.ring'(X0),X1),s(X0,X2),s('type.list.list'('type.quote.index'),X3),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : s('type.canonical.canonical_sum'(X6),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_canonical_sum_scalar3_4'(s('type.ring.ring'(X6),X7),s(X6,X8),s('type.list.list'('type.quote.index'),X9),s('type.canonical.canonical_sum'(X6),'const.canonical.Nil_monom_0')))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : s('type.canonical.canonical_sum'(X6),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_canonical_sum_scalar3_4'(s('type.ring.ring'(X6),X7),s(X6,X8),s('type.list.list'('type.quote.index'),X9),s('type.canonical.canonical_sum'(X6),'const.canonical.Nil_monom_0')))),
  file('thm.ringNorm.canonical_sum_scalar3_def|split|2','thm.ringNorm.canonical_sum_scalar3_def|split|2')).
fof(f6967,plain,(
  $false | ($spl1 | $spl682)),
  inference(trivial_inequality_removal,[],[f6934])).
fof(f6934,plain,(
  s('type.canonical.canonical_sum'(sK0),'const.canonical.Nil_monom_0') != s('type.canonical.canonical_sum'(sK0),'const.canonical.Nil_monom_0') | ($spl1 | $spl682)),
  inference(superposition,[],[f100,f6859])).
% SZS output end Proof for thm.ringNorm
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 78
% Passive clauses: 342
% Generated clauses: 6822
% Final active clauses: 74
% Final passive clauses: 257
% Input formulas: 24
% Initial clauses: 25
% 
% Duplicate literals: 402
% Trivial inequalities: 23
% Fw subsumption resolutions: 281
% Fw demodulations: 604
% Bw demodulations: 7
% Condensations: 515
% 
% Simple tautologies: 96
% Equational tautologies: 16
% Forward subsumptions: 4058
% 
% Binary resolution: 8
% Forward superposition: 2924
% Backward superposition: 1557
% Self superposition: 98
% Equality factoring: 30
% Equality resolution: 2
% 
% Split clauses: 12
% Split components: 29
% SAT solver clauses: 508
% SAT solver unit clauses: 44
% SAT solver binary clauses: 14
% SAT solver learnt clauses: 1
% 
% Sat splits: 12
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 507
% TWLsolver calls for satisfiability: 94
% 
% Memory used [KB]: 8187
% Time elapsed: 0.235 s
% ------------------------------
% ------------------------------
% Success in time 0.861 s
