Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.extreal
% (46372)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1326
% Passive clauses: 7683
% Generated clauses: 35969
% Final active clauses: 1283
% Final passive clauses: 5977
% Input formulas: 261
% Initial clauses: 527
% 
% Duplicate literals: 212
% Trivial inequalities: 88
% Fw subsumption resolutions: 2164
% Fw demodulations: 3318
% Bw demodulations: 42
% 
% Simple tautologies: 342
% Equational tautologies: 35
% Forward subsumptions: 21909
% 
% Binary resolution: 7226
% Factoring: 673
% Forward superposition: 10090
% Backward superposition: 11416
% Self superposition: 41
% Equality factoring: 2
% Equality resolution: 22
% 
% Split clauses: 389
% Split components: 780
% SAT solver clauses: 604
% SAT solver unit clauses: 31
% SAT solver binary clauses: 114
% 
% Sat splits: 389
% 
% TWLsolver clauses: 604
% TWLsolver calls for satisfiability: 508
% 
% Memory used [KB]: 18549
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 144 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.extreal
% (46384)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 776
% Passive clauses: 4181
% Generated clauses: 28582
% Final active clauses: 748
% Final passive clauses: 3365
% Input formulas: 261
% Initial clauses: 527
% 
% Duplicate literals: 116
% Trivial inequalities: 85
% Fw subsumption resolutions: 978
% Fw demodulations: 3643
% Bw demodulations: 59
% Condensations: 192
% Global subsumptions: 35
% 
% Simple tautologies: 341
% Equational tautologies: 14
% Forward subsumptions: 14253
% Fw demodulations to eq. taut.: 12
% 
% Binary resolution: 3123
% Factoring: 266
% Forward superposition: 7794
% Backward superposition: 7422
% Self superposition: 69
% Equality factoring: 2
% Equality resolution: 20
% 
% Split clauses: 36
% Split components: 72
% SAT solver clauses: 5294
% SAT solver unit clauses: 685
% SAT solver binary clauses: 167
% SAT solver learnt clauses: 70
% SAT solver learnt literals: 18
% 
% Sat splits: 36
% Sat splitting refutations: 9
% 
% TWLsolver clauses: 5154
% TWLsolver calls for satisfiability: 5810
% 
% Memory used [KB]: 33645
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.extreal
% (46403)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1046
% Passive clauses: 9954
% Generated clauses: 123923
% Final active clauses: 839
% Final passive clauses: 4823
% Input formulas: 261
% Initial clauses: 527
% 
% Duplicate literals: 11
% Trivial inequalities: 57
% Fw subsumption resolutions: 2243
% Fw demodulations: 39291
% Bw demodulations: 1206
% Condensations: 60
% 
% Simple tautologies: 117
% Equational tautologies: 68
% Forward subsumptions: 69591
% Fw demodulations to eq. taut.: 278
% Bw demodulations to eq. taut.: 17
% 
% Unit resulting resolution: 60385
% Factoring: 1
% Forward superposition: 10948
% Backward superposition: 7958
% Self superposition: 98
% Equality factoring: 2
% Equality resolution: 19
% 
% Split clauses: 25
% Split components: 50
% SAT solver clauses: 12234
% SAT solver unit clauses: 5670
% SAT solver binary clauses: 4484
% 
% Sat splits: 25
% Sat splitting refutations: 2228
% 
% Memory used [KB]: 338885
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.extreal
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.extreal
% SZS output start Proof for thm.extreal
fof(f5946,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f538,f746,f744,f791,f790,f789,f788,f787,f786,f783,f782,f781,f780,f779,f793,f963,f978,f984,f988,f991,f995,f998,f1003,f1007,f1013,f1018,f1022,f1025,f1026,f1027,f1028,f1029,f1030,f1032,f1102,f1347,f1348,f1349,f1350,f1351,f1352,f1354,f1424,f1471,f1476,f1478,f1480,f1482,f1484,f1486,f1488,f1544,f1579,f1580,f1581,f1582,f1584,f1585,f1586,f1587,f1588,f1735,f1589,f1591,f1661,f1734,f1740,f794,f1744,f795,f1748,f796,f1752,f797,f1756,f798,f1760,f799,f1763,f1764,f1765,f1766,f1770,f800,f1774,f801,f1775,f1777,f1778,f1780,f1782,f1784,f1788,f806,f1792,f808,f1923,f1921,f1916,f1976,f1977,f1978,f1979,f1980,f1981,f1983,f2053,f2356,f2386,f2188,f1779,f2254,f2255,f2428,f2435,f2439,f2441,f2436,f2444,f2445,f2446,f2447,f2448,f2449,f2450,f2451,f2452,f2453,f2454,f2455,f2456,f2438,f2457,f2458,f2440,f2459,f2460,f2461,f2443,f2462,f2463,f2464,f2680,f2691,f2692,f2693,f2694,f2696,f2697,f2622,f2699,f2700,f2620,f2617,f2701,f2702,f2613,f2704,f2612,f2706,f2611,f2708,f2745,f2746,f2747,f2748,f2749,f2751,f2900,f2752,f2822,f2899,f2920,f2931,f2932,f2933,f2934,f2935,f3065,f2936,f3066,f2938,f3008,f3064,f3009,f3067,f3041,f3068,f3042,f3071,f3043,f3070,f3073,f3074,f980,f1489,f3075,f3076,f3077,f3078,f3079,f3080,f3081,f3083,f3153,f3167,f3266,f3267,f3268,f3269,f3270,f3272,f3273,f3343,f3458,f3461,f3465,f3467,f3468,f3469,f3470,f3471,f3472,f3474,f3544,f3674,f804,f3687,f3688,f3689,f3690,f3691,f3692,f3694,f3764,f3777,f3780,f3781,f3782,f3783,f3784,f3786,f3778,f3787,f3857,f3870,f4136,f4142,f4158,f1473,f1472,f4329,f4328,f1490,f4313,f4309,f4305,f4138,f4332,f4333,f4334,f4335,f4336,f1117,f4337,f4339,f4409,f1115,f4622,f979,f4625,f985,f4628,f989,f4630,f992,f4633,f996,f4635,f999,f4639,f1004,f4642,f1008,f4647,f1016,f1015,f4649,f1020,f4651,f1024,f5002,f2388,f4866,f2429,f4867,f2430,f981,f5047,f5048,f5051,f5052,f2390,f5075,f5076,f5077,f5078,f5079,f5081,f3432,f5082,f5152,f3431,f5268,f5269,f5270,f5271,f1116,f5272,f1118,f3633,f5273,f1120,f5275,f1119,f5345,f1121,f3632,f5754,f5760,f802,f5756,f5776])).
fof(f5776,plain,(
  $false | ($spl1 | $spl26)),
  inference(subsumption_resolution,[],[f5529,f744])).
fof(f5529,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))) | $spl26),
  inference(superposition,[],[f565,f806])).
fof(f565,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.NegInf_0')))) )),
  inference(cnf_transformation,[],[f301])).
fof(f301,plain,(
  ! [X0] : s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.NegInf_0')))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X21] : s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X21))),s('type.extreal.extreal','const.extreal.NegInf_0')))),
  file('thm.extreal.add_infty|split|2','thm.extreal.extreal_add_def|split|6')).
fof(f5756,plain,(
  s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13('const.extreal.Normal_1'(s('type.realax.real',sK34(sK11))))))) | $spl176),
  inference(cnf_transformation,[],[f5756_D])).
fof(f5756_D,plain,(
  s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13('const.extreal.Normal_1'(s('type.realax.real',sK34(sK11))))))) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f802,plain,(
  ( ! [X36] : (s('type.extreal.extreal',sK11) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X36)))) ) | $spl22),
  inference(cnf_transformation,[],[f802_D])).
fof(f802_D,plain,(
  ( ! [X36] : (s('type.extreal.extreal',sK11) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X36)))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f5760,plain,(
  ( ! [X37] : (s('type.extreal.extreal',sK11) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X37))) | s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13('const.extreal.Normal_1'(s('type.realax.real',sK34(sK11)))))))) ) | ($spl5 | $spl26)),
  inference(forward_demodulation,[],[f5759,f806])).
fof(f5759,plain,(
  ( ! [X37] : (s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13('const.extreal.Normal_1'(s('type.realax.real',sK34(sK11))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X37))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK34(sK11))))) ) | ($spl5 | $spl26)),
  inference(forward_demodulation,[],[f5758,f806])).
fof(f5758,plain,(
  ( ! [X37] : (s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK34(sK11)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13('const.extreal.Normal_1'(s('type.realax.real',sK34(sK11))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X37))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK34(sK11))))) ) | ($spl5 | $spl26)),
  inference(subsumption_resolution,[],[f5757,f638])).
fof(f638,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))) )),
  inference(cnf_transformation,[],[f342])).
fof(f342,plain,(
  ! [X0] : s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))),
  inference(flattening,[],[f341])).
fof(f341,plain,(
  ! [X0] : ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))),
  inference(rectify,[],[f141])).
fof(f141,axiom,(
  ! [X6] : ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X6)))),
  file('thm.extreal.add_infty|split|2','thm.extreal.extreal_not_infty|split|1')).
fof(f5757,plain,(
  ( ! [X37] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK34(sK11)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK34(sK11)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13('const.extreal.Normal_1'(s('type.realax.real',sK34(sK11))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X37))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK34(sK11))))) ) | ($spl5 | $spl26)),
  inference(subsumption_resolution,[],[f5518,f792])).
fof(f792,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',sK11) | $spl5),
  inference(cnf_transformation,[],[f792_D])).
fof(f792_D,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',sK11) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f5518,plain,(
  ( ! [X37] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK11) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK34(sK11)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK34(sK11)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13('const.extreal.Normal_1'(s('type.realax.real',sK34(sK11))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X37))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK34(sK11))))) ) | $spl26),
  inference(superposition,[],[f806,f584])).
fof(f584,plain,(
  ( ! [X2,X0] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(X0)))) | s('type.extreal.extreal',X0) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) )),
  inference(cnf_transformation,[],[f468])).
fof(f468,plain,(
  ! [X0] : ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0) & ! [X2] : s('type.extreal.extreal',X0) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) | s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(X0)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13])],[f467])).
fof(f467,plain,(
  ! [X0] : ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0) & ! [X2] : s('type.extreal.extreal',X0) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) | ? [X1] : s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0))),
  inference(rectify,[],[f420])).
fof(f420,plain,(
  ! [X0] : ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0) & ! [X1] : s('type.extreal.extreal',X0) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) | ? [X2] : s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0))),
  inference(flattening,[],[f419])).
fof(f419,plain,(
  ! [X0] : ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0) & ! [X1] : s('type.extreal.extreal',X0) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) | (? [X2] : s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)))),
  inference(ennf_transformation,[],[f318])).
fof(f318,plain,(
  ! [X0] : ((s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | ? [X1] : s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) => (? [X2] : s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)))),
  inference(rectify,[],[f78])).
fof(f78,axiom,(
  ! [X11] : ((s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X11) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X11) | ? [X6] : s('type.extreal.extreal',X11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X6)))) => (? [X6] : s('type.extreal.extreal',X11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X6))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X11) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X11)))),
  file('thm.extreal.add_infty|split|2','thm.extreal.extreal_46_prove_ltb')).
fof(f5754,plain,(
  ( ! [X36] : (s('type.extreal.extreal',sK11) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X36))) | s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13('const.extreal.Normal_1'(s('type.realax.real',sK34(sK11)))))))) ) | $spl26),
  inference(forward_demodulation,[],[f5753,f806])).
fof(f5753,plain,(
  ( ! [X36] : (s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13('const.extreal.Normal_1'(s('type.realax.real',sK34(sK11))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X36))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK34(sK11))))) ) | $spl26),
  inference(forward_demodulation,[],[f5752,f806])).
fof(f5752,plain,(
  ( ! [X36] : (s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK34(sK11)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13('const.extreal.Normal_1'(s('type.realax.real',sK34(sK11))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X36))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK34(sK11))))) ) | $spl26),
  inference(subsumption_resolution,[],[f5751,f576])).
fof(f576,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))) )),
  inference(cnf_transformation,[],[f312])).
fof(f312,plain,(
  ! [X0] : s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))),
  inference(flattening,[],[f311])).
fof(f311,plain,(
  ! [X0] : ~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))),
  inference(rectify,[],[f160])).
fof(f160,axiom,(
  ! [X11] : ~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X11)))),
  file('thm.extreal.add_infty|split|2','thm.extreal.extreal_distinct|split|1')).
fof(f5751,plain,(
  ( ! [X36] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK34(sK11)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK34(sK11)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13('const.extreal.Normal_1'(s('type.realax.real',sK34(sK11))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X36))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK34(sK11))))) ) | $spl26),
  inference(subsumption_resolution,[],[f5517,f537])).
fof(f537,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',sK11)),
  inference(cnf_transformation,[],[f463])).
fof(f463,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',sK11) & (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK11])],[f383])).
fof(f383,plain,(
  ? [X0] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0) & (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)))))),
  inference(ennf_transformation,[],[f264])).
fof(f264,plain,(
  ~! [X0] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0) => (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.NegInf_0'))) & s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)))))),
  inference(flattening,[],[f263])).
fof(f263,plain,(
  ~! [X0] : (~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) => (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.NegInf_0'))) & s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : (~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X6) => (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal','const.extreal.NegInf_0'))) & s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X6)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : (~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X6) => (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal','const.extreal.NegInf_0'))) & s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X6)))))),
  file('thm.extreal.add_infty|split|2','thm.extreal.add_infty|split|2')).
fof(f5517,plain,(
  ( ! [X36] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK11) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK34(sK11)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK34(sK11)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13('const.extreal.Normal_1'(s('type.realax.real',sK34(sK11))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X36))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK34(sK11))))) ) | $spl26),
  inference(superposition,[],[f806,f584])).
fof(f3632,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl160),
  inference(cnf_transformation,[],[f3632_D])).
fof(f3632_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f1121,plain,(
  ( ! [X15] : (p(s(bool,X15))) ) | $spl68),
  inference(cnf_transformation,[],[f1121_D])).
fof(f1121_D,plain,(
  ( ! [X15] : (p(s(bool,X15))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f5345,plain,(
  ( ! [X152,X153] : (p(s(bool,X152)) | ~p(s(bool,X153)) | s(bool,X153) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(resolution,[],[f960,f737])).
fof(f737,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f603])).
fof(f603,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f482])).
fof(f482,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f481])).
fof(f481,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f432])).
fof(f432,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f431])).
fof(f431,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f331])).
fof(f331,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f61])).
fof(f61,axiom,(
  ! [X6,X40,X12,X41] : ((s(bool,X6) = s(bool,X40) & (p(s(bool,X40)) => s(bool,X12) = s(bool,X41))) => ((p(s(bool,X6)) => p(s(bool,X12))) <=> (p(s(bool,X40)) => p(s(bool,X41)))))),
  file('thm.extreal.add_infty|split|2','thm.bool.IMP_CONG')).
fof(f960,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl1),
  inference(subsumption_resolution,[],[f830,f570])).
fof(f570,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))),
  inference(cnf_transformation,[],[f17])).
fof(f17,axiom,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))),
  file('thm.extreal.add_infty|split|2','thm.extreal.extreal_add_def|split|4')).
fof(f830,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl1),
  inference(superposition,[],[f744,f580])).
fof(f580,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.NegInf_0'))))) )),
  inference(cnf_transformation,[],[f466])).
fof(f466,plain,(
  ! [X0] : ((~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0)) & (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.NegInf_0'))))))),
  inference(nnf_transformation,[],[f317])).
fof(f317,plain,(
  ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.NegInf_0')))) <=> s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0))),
  inference(rectify,[],[f98])).
fof(f98,axiom,(
  ! [X6] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal','const.extreal.NegInf_0')))) <=> s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X6))),
  file('thm.extreal.add_infty|split|2','thm.extreal.le_infty|split|2')).
fof(f1119,plain,(
  ( ! [X6,X8] : (~p(s(bool,X6)) | s(bool,X6) != s(bool,X8)) ) | $spl64),
  inference(cnf_transformation,[],[f1119_D])).
fof(f1119_D,plain,(
  ( ! [X6,X8] : (~p(s(bool,X6)) | s(bool,X6) != s(bool,X8)) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f5275,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s(bool,X21) != s(bool,X23)) ) | $spl1),
  inference(resolution,[],[f960,f606])).
fof(f606,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f482])).
fof(f1120,plain,(
  ( ! [X17] : (~p(s(bool,X17))) ) | $spl66),
  inference(cnf_transformation,[],[f1120_D])).
fof(f1120_D,plain,(
  ( ! [X17] : (~p(s(bool,X17))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f5273,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(resolution,[],[f960,f605])).
fof(f605,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f482])).
fof(f3633,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl162),
  inference(cnf_transformation,[],[f3633_D])).
fof(f3633_D,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f1118,plain,(
  ( ! [X4,X3] : (~p(s(bool,X3)) | s(bool,X3) != s(bool,X4) | p(s(bool,X4))) ) | $spl62),
  inference(cnf_transformation,[],[f1118_D])).
fof(f1118_D,plain,(
  ( ! [X4,X3] : (~p(s(bool,X3)) | s(bool,X3) != s(bool,X4) | p(s(bool,X4))) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f5272,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s(bool,X12) != s(bool,X13)) ) | $spl1),
  inference(resolution,[],[f960,f604])).
fof(f604,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f482])).
fof(f1116,plain,(
  ( ! [X2,X0] : (p(s(bool,X0)) | s(bool,X0) != s(bool,X2)) ) | $spl58),
  inference(cnf_transformation,[],[f1116_D])).
fof(f1116_D,plain,(
  ( ! [X2,X0] : (p(s(bool,X0)) | s(bool,X0) != s(bool,X2)) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f5271,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(resolution,[],[f960,f604])).
fof(f5270,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s(bool,X6) != s(bool,X8)) ) | $spl1),
  inference(resolution,[],[f960,f602])).
fof(f602,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f482])).
fof(f5269,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s(bool,X3) != s(bool,X4)) ) | $spl1),
  inference(resolution,[],[f960,f600])).
fof(f600,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f482])).
fof(f5268,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(resolution,[],[f960,f600])).
fof(f3431,plain,(
  ( ! [X9] : (p(s(bool,X9)) | s(bool,X9) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))) ) | $spl156),
  inference(cnf_transformation,[],[f3431_D])).
fof(f3431_D,plain,(
  ( ! [X9] : (p(s(bool,X9)) | s(bool,X9) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f5152,plain,(
  ( ! [X161,X160] : (p(s(bool,X160)) | p(s(bool,X161)) | s(bool,X161) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))) ) | $spl1),
  inference(resolution,[],[f959,f737])).
fof(f959,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))) | $spl1),
  inference(subsumption_resolution,[],[f829,f570])).
fof(f829,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))) | $spl1),
  inference(superposition,[],[f744,f578])).
fof(f578,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0)) )),
  inference(cnf_transformation,[],[f465])).
fof(f465,plain,(
  ! [X0] : ((s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0))))) & (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)))) | s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0)))),
  inference(nnf_transformation,[],[f316])).
fof(f316,plain,(
  ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) <=> p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)))))),
  inference(flattening,[],[f315])).
fof(f315,plain,(
  ! [X0] : (~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) <=> p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)))))),
  inference(rectify,[],[f116])).
fof(f116,axiom,(
  ! [X6] : (~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X6) <=> p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X6)))))),
  file('thm.extreal.add_infty|split|2','thm.extreal.lt_infty|split|6')).
fof(f5082,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) | s(bool,X30) != s(bool,X31)) ) | $spl1),
  inference(resolution,[],[f959,f606])).
fof(f3432,plain,(
  ( ! [X13] : (s(bool,X13) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))) ) | $spl158),
  inference(cnf_transformation,[],[f3432_D])).
fof(f3432_D,plain,(
  ( ! [X13] : (s(bool,X13) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f5081,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,X27) | s(bool,X28) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))) ) | $spl1),
  inference(resolution,[],[f959,f606])).
fof(f5079,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X22) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))) ) | $spl1),
  inference(resolution,[],[f959,f605])).
fof(f5078,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X18) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))) ) | $spl1),
  inference(resolution,[],[f959,f604])).
fof(f5077,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X14) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) | s(bool,X15) != s(bool,X16)) ) | $spl1),
  inference(resolution,[],[f959,f602])).
fof(f5076,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X12) | s(bool,X13) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))) ) | $spl1),
  inference(resolution,[],[f959,f602])).
fof(f5075,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X8)) | p(s(bool,X9)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))) ) | $spl1),
  inference(resolution,[],[f959,f600])).
fof(f2390,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl87),
  inference(cnf_transformation,[],[f2390_D])).
fof(f2390_D,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) <=> ~$spl87),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl87])])).
fof(f5052,plain,(
  ( ! [X211] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X211)))))))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X211))))) ) | $spl174),
  inference(cnf_transformation,[],[f5052_D])).
fof(f5052_D,plain,(
  ( ! [X211] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X211)))))))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X211))))) ) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f5051,plain,(
  ( ! [X211] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X211)))))))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X211))))) ) | $spl2),
  inference(forward_demodulation,[],[f4888,f549])).
fof(f549,plain,(
  ( ! [X2,X0,X1] : (s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))),s('type.extreal.extreal',X2)))) )),
  inference(cnf_transformation,[],[f281])).
fof(f281,plain,(
  ! [X0,X1,X2] : s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))),s('type.extreal.extreal',X2)))),
  inference(rectify,[],[f72])).
fof(f72,axiom,(
  ! [X6,X12,X38] : s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X12),s('type.extreal.extreal',X38))))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X12))),s('type.extreal.extreal',X38)))),
  file('thm.extreal.add_infty|split|2','thm.extreal.add_assoc')).
fof(f4888,plain,(
  ( ! [X211] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))),s('type.extreal.extreal',X211)))))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X211))))) ) | $spl2),
  inference(superposition,[],[f556,f745])).
fof(f745,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) | $spl2),
  inference(cnf_transformation,[],[f745_D])).
fof(f745_D,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f556,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X1)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X0))))) )),
  inference(cnf_transformation,[],[f415])).
fof(f415,plain,(
  ! [X0,X1] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X0)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X1)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))))))),
  inference(flattening,[],[f414])).
fof(f414,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X0)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X1))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))))))),
  inference(ennf_transformation,[],[f292])).
fof(f292,plain,(
  ! [X0,X1] : ((p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X0)))) & p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X1))))) => p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))))))),
  inference(rectify,[],[f85])).
fof(f85,axiom,(
  ! [X6,X12] : ((p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X6)))) & p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X12))))) => p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X12)))))))),
  file('thm.extreal.add_infty|split|2','thm.extreal.le_add')).
fof(f5048,plain,(
  ( ! [X208] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X208)))))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X208))))) ) | $spl172),
  inference(cnf_transformation,[],[f5048_D])).
fof(f5048_D,plain,(
  ( ! [X208] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X208)))))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X208))))) ) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f5047,plain,(
  ( ! [X208] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X208)))))))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X208))))) ) | $spl2),
  inference(forward_demodulation,[],[f4885,f549])).
fof(f4885,plain,(
  ( ! [X208] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))),s('type.extreal.extreal',X208)))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X208))))) ) | $spl2),
  inference(superposition,[],[f555,f745])).
fof(f555,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X1)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X0))))) )),
  inference(cnf_transformation,[],[f413])).
fof(f413,plain,(
  ! [X0,X1] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X0)))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X1)))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))))))),
  inference(flattening,[],[f412])).
fof(f412,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X0)))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X1))))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))))))),
  inference(ennf_transformation,[],[f291])).
fof(f291,plain,(
  ! [X0,X1] : ((p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X0)))) & p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X1))))) => p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))))))),
  inference(rectify,[],[f86])).
fof(f86,axiom,(
  ! [X6,X12] : ((p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X6)))) & p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X12))))) => p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X12)))))))),
  file('thm.extreal.add_infty|split|2','thm.extreal.let_add')).
fof(f981,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) | $spl32),
  inference(cnf_transformation,[],[f981_D])).
fof(f981_D,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f2430,plain,(
  ( ! [X190] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X190),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X190)) ) | $spl92),
  inference(cnf_transformation,[],[f2430_D])).
fof(f2430_D,plain,(
  ( ! [X190] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X190),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X190)) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f4867,plain,(
  ( ! [X190] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X190),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X190)) ) | $spl2),
  inference(superposition,[],[f550,f745])).
fof(f550,plain,(
  ( ! [X0,X1] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X1) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)) )),
  inference(cnf_transformation,[],[f403])).
fof(f403,plain,(
  ! [X0,X1] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X1) | s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))),
  inference(flattening,[],[f402])).
fof(f402,plain,(
  ! [X0,X1] : ((s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X1)) | s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))),
  inference(ennf_transformation,[],[f283])).
fof(f283,plain,(
  ! [X0,X1] : ((s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X1)) => s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))),
  inference(flattening,[],[f282])).
fof(f282,plain,(
  ! [X0,X1] : ((~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) & ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X1)) => ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))),
  inference(rectify,[],[f62])).
fof(f62,axiom,(
  ! [X6,X12] : ((~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X6) & ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X12)) => ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X12))))),
  file('thm.extreal.add_infty|split|2','thm.extreal.add_not_infty|split|1')).
fof(f2429,plain,(
  ( ! [X189] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X189))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X189)) ) | $spl90),
  inference(cnf_transformation,[],[f2429_D])).
fof(f2429_D,plain,(
  ( ! [X189] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X189))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X189)) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f4866,plain,(
  ( ! [X189] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X189))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X189) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))) ) | $spl2),
  inference(superposition,[],[f550,f745])).
fof(f2388,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK11)))) | $spl85),
  inference(cnf_transformation,[],[f2388_D])).
fof(f2388_D,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK11)))) <=> ~$spl85),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl85])])).
fof(f5002,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK11)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl2),
  inference(subsumption_resolution,[],[f4799,f574])).
fof(f574,plain,(
  ( ! [X0] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.NegInf_0'))))) )),
  inference(cnf_transformation,[],[f308])).
fof(f308,plain,(
  ! [X0] : ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.NegInf_0'))))),
  inference(flattening,[],[f307])).
fof(f307,plain,(
  ! [X0] : ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.NegInf_0'))))),
  inference(rectify,[],[f145])).
fof(f145,axiom,(
  ! [X6] : ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal','const.extreal.NegInf_0'))))),
  file('thm.extreal.add_infty|split|2','thm.extreal.lt_infty|split|3')).
fof(f4799,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK11)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl2),
  inference(superposition,[],[f555,f745])).
fof(f1024,plain,(
  s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK35('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))))) | $spl54),
  inference(cnf_transformation,[],[f1024_D])).
fof(f1024_D,plain,(
  s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK35('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))))) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f4651,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) | s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK35('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))))) | $spl1),
  inference(forward_demodulation,[],[f4650,f560])).
fof(f560,plain,(
  ( ! [X0,X1] : (s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))) )),
  inference(cnf_transformation,[],[f296])).
fof(f296,plain,(
  ! [X0,X1] : s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))),
  inference(rectify,[],[f69])).
fof(f69,axiom,(
  ! [X6,X12] : s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X12))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X12),s('type.extreal.extreal',X6)))),
  file('thm.extreal.add_infty|split|2','thm.extreal.add_comm')).
fof(f4650,plain,(
  s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK35('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))) | $spl1),
  inference(forward_demodulation,[],[f4555,f560])).
fof(f4555,plain,(
  s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK35('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))) | $spl1),
  inference(trivial_inequality_removal,[],[f4552])).
fof(f4552,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.NegInf_0') | s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK35('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))) | $spl1),
  inference(superposition,[],[f744,f591])).
fof(f591,plain,(
  ( ! [X0] : (s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK35(X0)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0)) )),
  inference(cnf_transformation,[],[f480])).
fof(f480,plain,(
  ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK35(X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK35])],[f325])).
fof(f325,plain,(
  ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | ? [X1] : s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1))))),
  inference(rectify,[],[f90])).
fof(f90,axiom,(
  ! [X46] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X46) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X46) | ? [X36] : s('type.extreal.extreal',X46) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X36))))),
  file('thm.extreal.add_infty|split|2','thm.extreal.extreal_nchotomy')).
fof(f1020,plain,(
  s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK34('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))))) | $spl52),
  inference(cnf_transformation,[],[f1020_D])).
fof(f1020_D,plain,(
  s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK34('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))))) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f4649,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) | s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK34('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))))) | $spl1),
  inference(forward_demodulation,[],[f4648,f560])).
fof(f4648,plain,(
  s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK34('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))) | $spl1),
  inference(forward_demodulation,[],[f4556,f560])).
fof(f4556,plain,(
  s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK34('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))) | $spl1),
  inference(trivial_inequality_removal,[],[f4550])).
fof(f4550,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.NegInf_0') | s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK34('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))) | $spl1),
  inference(superposition,[],[f744,f590])).
fof(f590,plain,(
  ( ! [X0] : (s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK34(X0)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0)) )),
  inference(cnf_transformation,[],[f479])).
fof(f479,plain,(
  ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK34(X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK34])],[f324])).
fof(f324,plain,(
  ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | ? [X1] : s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1))))),
  inference(rectify,[],[f93])).
fof(f93,axiom,(
  ! [X6] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X6) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X6) | ? [X36] : s('type.extreal.extreal',X6) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X36))))),
  file('thm.extreal.add_infty|split|2','thm.extreal.extreal_cases')).
fof(f1015,plain,(
  s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))))) | $spl48),
  inference(cnf_transformation,[],[f1015_D])).
fof(f1015_D,plain,(
  s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))))) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f1016,plain,(
  ( ! [X37] : (s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X37))) != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))) ) | $spl50),
  inference(cnf_transformation,[],[f1016_D])).
fof(f1016_D,plain,(
  ( ! [X37] : (s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X37))) != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f4647,plain,(
  ( ! [X37] : (s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X37))) != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) | s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))) ) | $spl1),
  inference(forward_demodulation,[],[f4646,f560])).
fof(f4646,plain,(
  ( ! [X37] : (s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X37))) != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(forward_demodulation,[],[f4645,f560])).
fof(f4645,plain,(
  ( ! [X37] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) | s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X37))) != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(forward_demodulation,[],[f4557,f560])).
fof(f4557,plain,(
  ( ! [X37] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X37))) != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(trivial_inequality_removal,[],[f4548])).
fof(f4548,plain,(
  ( ! [X37] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.NegInf_0') | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X37))) != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(superposition,[],[f744,f584])).
fof(f1008,plain,(
  ( ! [X35] : (s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X35))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X35)))))))) ) | $spl46),
  inference(cnf_transformation,[],[f1008_D])).
fof(f1008_D,plain,(
  ( ! [X35] : (s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X35))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X35)))))))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f4642,plain,(
  ( ! [X35] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) | s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X35))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X35)))))))) ) | $spl1),
  inference(forward_demodulation,[],[f4641,f560])).
fof(f4641,plain,(
  ( ! [X35] : (s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X35))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X35))))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(forward_demodulation,[],[f4640,f560])).
fof(f4640,plain,(
  ( ! [X35] : (s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X35))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X35))))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(forward_demodulation,[],[f4560,f549])).
fof(f4560,plain,(
  ( ! [X35] : (s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X35))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',X35))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(trivial_inequality_removal,[],[f4544])).
fof(f4544,plain,(
  ( ! [X35] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.NegInf_0') | s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X35))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',X35))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(superposition,[],[f744,f553])).
fof(f553,plain,(
  ( ! [X0,X1] : (s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X1))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0)) )),
  inference(cnf_transformation,[],[f409])).
fof(f409,plain,(
  ! [X0,X1] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X1))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))))),
  inference(flattening,[],[f408])).
fof(f408,plain,(
  ! [X0,X1] : ((s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)) | s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X1))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))))),
  inference(ennf_transformation,[],[f289])).
fof(f289,plain,(
  ! [X0,X1] : ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0)) => s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X1))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))))),
  inference(flattening,[],[f288])).
fof(f288,plain,(
  ! [X0,X1] : ((~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) & ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)) => s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X1))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))))),
  inference(rectify,[],[f56])).
fof(f56,axiom,(
  ! [X6,X12] : ((~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X6) & ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X6)) => s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X12))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X12))))))),
  file('thm.extreal.add_infty|split|2','thm.extreal.le_addr')).
fof(f1004,plain,(
  ( ! [X33] : (s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X33))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X33),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))))))) ) | $spl44),
  inference(cnf_transformation,[],[f1004_D])).
fof(f1004_D,plain,(
  ( ! [X33] : (s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X33))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X33),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))))))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f4639,plain,(
  ( ! [X33] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) | s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X33))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X33),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))))))) ) | $spl1),
  inference(forward_demodulation,[],[f4638,f560])).
fof(f4638,plain,(
  ( ! [X33] : (s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X33))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X33),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(forward_demodulation,[],[f4561,f560])).
fof(f4561,plain,(
  ( ! [X33] : (s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X33))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X33),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(trivial_inequality_removal,[],[f4542])).
fof(f4542,plain,(
  ( ! [X33] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.NegInf_0') | s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X33))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X33),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(superposition,[],[f744,f552])).
fof(f552,plain,(
  ( ! [X0,X1] : (s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X0))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X1) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X1)) )),
  inference(cnf_transformation,[],[f407])).
fof(f407,plain,(
  ! [X0,X1] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X1) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X1) | s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X0))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))))),
  inference(flattening,[],[f406])).
fof(f406,plain,(
  ! [X0,X1] : ((s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X1) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X1)) | s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X0))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))))),
  inference(ennf_transformation,[],[f287])).
fof(f287,plain,(
  ! [X0,X1] : ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X1) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X1)) => s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X0))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))))),
  inference(flattening,[],[f286])).
fof(f286,plain,(
  ! [X0,X1] : ((~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X1) & ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X1)) => s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X0))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))))),
  inference(rectify,[],[f53])).
fof(f53,axiom,(
  ! [X6,X12] : ((~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X12) & ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X12)) => s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X6))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X12))))))),
  file('thm.extreal.add_infty|split|2','thm.extreal.lt_addl')).
fof(f999,plain,(
  ( ! [X26,X27] : (s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X26),s('type.extreal.extreal',X27))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X26),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X27),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))))))) ) | $spl42),
  inference(cnf_transformation,[],[f999_D])).
fof(f999_D,plain,(
  ( ! [X26,X27] : (s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X26),s('type.extreal.extreal',X27))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X26),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X27),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))))))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f4635,plain,(
  ( ! [X26,X27] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) | s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X26),s('type.extreal.extreal',X27))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X26),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X27),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))))))) ) | $spl1),
  inference(forward_demodulation,[],[f4634,f560])).
fof(f4634,plain,(
  ( ! [X26,X27] : (s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X26),s('type.extreal.extreal',X27))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X26),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X27),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(forward_demodulation,[],[f4564,f560])).
fof(f4564,plain,(
  ( ! [X26,X27] : (s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X26),s('type.extreal.extreal',X27))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X26),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X27),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(trivial_inequality_removal,[],[f4536])).
fof(f4536,plain,(
  ( ! [X26,X27] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.NegInf_0') | s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X26),s('type.extreal.extreal',X27))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X26),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X27),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(superposition,[],[f744,f546])).
fof(f546,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X0))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0)) )),
  inference(cnf_transformation,[],[f399])).
fof(f399,plain,(
  ! [X0,X1,X2] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X0))))))),
  inference(flattening,[],[f398])).
fof(f398,plain,(
  ! [X0,X1,X2] : ((s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)) | s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X0))))))),
  inference(ennf_transformation,[],[f278])).
fof(f278,plain,(
  ! [X0,X1,X2] : ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0)) => s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X0))))))),
  inference(flattening,[],[f277])).
fof(f277,plain,(
  ! [X0,X1,X2] : ((~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) & ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)) => s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X0))))))),
  inference(rectify,[],[f54])).
fof(f54,axiom,(
  ! [X6,X12,X38] : ((~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X6) & ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X6)) => s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X12),s('type.extreal.extreal',X38))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X12),s('type.extreal.extreal',X6))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X38),s('type.extreal.extreal',X6))))))),
  file('thm.extreal.add_infty|split|2','thm.extreal.lt_radd')).
fof(f996,plain,(
  ( ! [X23,X22] : (s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X22),s('type.extreal.extreal',X23))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X22))))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X23)))))))) ) | $spl40),
  inference(cnf_transformation,[],[f996_D])).
fof(f996_D,plain,(
  ( ! [X23,X22] : (s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X22),s('type.extreal.extreal',X23))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X22))))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X23)))))))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f4633,plain,(
  ( ! [X23,X22] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) | s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X22),s('type.extreal.extreal',X23))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X22))))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X23)))))))) ) | $spl1),
  inference(forward_demodulation,[],[f4632,f560])).
fof(f4632,plain,(
  ( ! [X23,X22] : (s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X22),s('type.extreal.extreal',X23))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X22))))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X23))))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(forward_demodulation,[],[f4631,f549])).
fof(f4631,plain,(
  ( ! [X23,X22] : (s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X22),s('type.extreal.extreal',X23))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',X22))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X23))))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(forward_demodulation,[],[f4565,f549])).
fof(f4565,plain,(
  ( ! [X23,X22] : (s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X22),s('type.extreal.extreal',X23))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',X22))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',X23))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(trivial_inequality_removal,[],[f4534])).
fof(f4534,plain,(
  ( ! [X23,X22] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.NegInf_0') | s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X22),s('type.extreal.extreal',X23))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',X22))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',X23))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(superposition,[],[f744,f545])).
fof(f545,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0)) )),
  inference(cnf_transformation,[],[f397])).
fof(f397,plain,(
  ! [X0,X1,X2] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2))))))),
  inference(flattening,[],[f396])).
fof(f396,plain,(
  ! [X0,X1,X2] : ((s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)) | s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2))))))),
  inference(ennf_transformation,[],[f276])).
fof(f276,plain,(
  ! [X0,X1,X2] : ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0)) => s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2))))))),
  inference(flattening,[],[f275])).
fof(f275,plain,(
  ! [X0,X1,X2] : ((~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) & ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)) => s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2))))))),
  inference(rectify,[],[f59])).
fof(f59,axiom,(
  ! [X6,X12,X38] : ((~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X6) & ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X6)) => s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X12),s('type.extreal.extreal',X38))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X12))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X38))))))),
  file('thm.extreal.add_infty|split|2','thm.extreal.le_ladd')).
fof(f992,plain,(
  ( ! [X19,X18] : (s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X18),s('type.extreal.extreal',X19))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X18),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X19),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))))))) ) | $spl38),
  inference(cnf_transformation,[],[f992_D])).
fof(f992_D,plain,(
  ( ! [X19,X18] : (s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X18),s('type.extreal.extreal',X19))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X18),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X19),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))))))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f4630,plain,(
  ( ! [X19,X18] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) | s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X18),s('type.extreal.extreal',X19))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X18),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X19),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))))))) ) | $spl1),
  inference(forward_demodulation,[],[f4629,f560])).
fof(f4629,plain,(
  ( ! [X19,X18] : (s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X18),s('type.extreal.extreal',X19))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X18),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X19),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(forward_demodulation,[],[f4566,f560])).
fof(f4566,plain,(
  ( ! [X19,X18] : (s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X18),s('type.extreal.extreal',X19))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X18),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X19),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(trivial_inequality_removal,[],[f4532])).
fof(f4532,plain,(
  ( ! [X19,X18] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.NegInf_0') | s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X18),s('type.extreal.extreal',X19))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X18),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X19),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(superposition,[],[f744,f544])).
fof(f544,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X0))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0)) )),
  inference(cnf_transformation,[],[f395])).
fof(f395,plain,(
  ! [X0,X1,X2] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X0))))))),
  inference(flattening,[],[f394])).
fof(f394,plain,(
  ! [X0,X1,X2] : ((s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)) | s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X0))))))),
  inference(ennf_transformation,[],[f274])).
fof(f274,plain,(
  ! [X0,X1,X2] : ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0)) => s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X0))))))),
  inference(flattening,[],[f273])).
fof(f273,plain,(
  ! [X0,X1,X2] : ((~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) & ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)) => s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X0))))))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X6,X12,X38] : ((~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X6) & ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X6)) => s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X12),s('type.extreal.extreal',X38))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X12),s('type.extreal.extreal',X6))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X38),s('type.extreal.extreal',X6))))))),
  file('thm.extreal.add_infty|split|2','thm.extreal.le_radd')).
fof(f989,plain,(
  ( ! [X14,X15] : (s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X14),s('type.extreal.extreal',X15))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X14))))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X15)))))))) ) | $spl36),
  inference(cnf_transformation,[],[f989_D])).
fof(f989_D,plain,(
  ( ! [X14,X15] : (s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X14),s('type.extreal.extreal',X15))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X14))))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X15)))))))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f4628,plain,(
  ( ! [X14,X15] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) | s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X14),s('type.extreal.extreal',X15))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X14))))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X15)))))))) ) | $spl1),
  inference(forward_demodulation,[],[f4627,f560])).
fof(f4627,plain,(
  ( ! [X14,X15] : (s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X14),s('type.extreal.extreal',X15))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X14))))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X15))))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(forward_demodulation,[],[f4626,f549])).
fof(f4626,plain,(
  ( ! [X14,X15] : (s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X14),s('type.extreal.extreal',X15))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',X14))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X15))))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(forward_demodulation,[],[f4567,f549])).
fof(f4567,plain,(
  ( ! [X14,X15] : (s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X14),s('type.extreal.extreal',X15))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',X14))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',X15))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(trivial_inequality_removal,[],[f4530])).
fof(f4530,plain,(
  ( ! [X14,X15] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.NegInf_0') | s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X14),s('type.extreal.extreal',X15))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',X14))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',X15))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(superposition,[],[f744,f543])).
fof(f543,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0)) )),
  inference(cnf_transformation,[],[f393])).
fof(f393,plain,(
  ! [X0,X1,X2] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2))))))),
  inference(flattening,[],[f392])).
fof(f392,plain,(
  ! [X0,X1,X2] : ((s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)) | s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2))))))),
  inference(ennf_transformation,[],[f272])).
fof(f272,plain,(
  ! [X0,X1,X2] : ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0)) => s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2))))))),
  inference(flattening,[],[f271])).
fof(f271,plain,(
  ! [X0,X1,X2] : ((~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) & ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)) => s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2))))))),
  inference(rectify,[],[f55])).
fof(f55,axiom,(
  ! [X6,X12,X38] : ((~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X6) & ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X6)) => s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X12),s('type.extreal.extreal',X38))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X12))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X38))))))),
  file('thm.extreal.add_infty|split|2','thm.extreal.lt_ladd')).
fof(f985,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))),s('type.extreal.extreal',X10)))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X9),s('type.extreal.extreal',X11)))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X9))))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X10),s('type.extreal.extreal',X11))))))) ) | $spl34),
  inference(cnf_transformation,[],[f985_D])).
fof(f985_D,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))),s('type.extreal.extreal',X10)))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X9),s('type.extreal.extreal',X11)))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X9))))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X10),s('type.extreal.extreal',X11))))))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f4625,plain,(
  ( ! [X10,X11,X9] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))),s('type.extreal.extreal',X10)))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X9))))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X10),s('type.extreal.extreal',X11)))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X9),s('type.extreal.extreal',X11))))) ) | $spl1),
  inference(forward_demodulation,[],[f4624,f560])).
fof(f4624,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))),s('type.extreal.extreal',X10)))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X9))))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X10),s('type.extreal.extreal',X11)))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X9),s('type.extreal.extreal',X11)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(forward_demodulation,[],[f4623,f560])).
fof(f4623,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X9))))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X10),s('type.extreal.extreal',X11)))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X9),s('type.extreal.extreal',X11)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',X10)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(forward_demodulation,[],[f4568,f549])).
fof(f4568,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',X9))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X10),s('type.extreal.extreal',X11)))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X9),s('type.extreal.extreal',X11)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',X10)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(trivial_inequality_removal,[],[f4528])).
fof(f4528,plain,(
  ( ! [X10,X11,X9] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.NegInf_0') | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',X9))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X10),s('type.extreal.extreal',X11)))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X9),s('type.extreal.extreal',X11)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',X10)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(superposition,[],[f744,f540])).
fof(f540,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X3)))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X3)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0)) )),
  inference(cnf_transformation,[],[f387])).
fof(f387,plain,(
  ! [X0,X1,X2,X3] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X3)))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X3)))))))),
  inference(flattening,[],[f386])).
fof(f386,plain,(
  ! [X0,X1,X2,X3] : ((s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X3))))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X3)))))))),
  inference(ennf_transformation,[],[f268])).
fof(f268,plain,(
  ! [X0,X1,X2,X3] : ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0) & p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) & p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X3))))) => p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X3)))))))),
  inference(flattening,[],[f267])).
fof(f267,plain,(
  ! [X0,X1,X2,X3] : ((~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) & ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) & p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) & p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X3))))) => p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X3)))))))),
  inference(rectify,[],[f58])).
fof(f58,axiom,(
  ! [X37,X6,X12,X38] : ((~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X37) & ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X37) & p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X37),s('type.extreal.extreal',X6)))) & p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X12),s('type.extreal.extreal',X38))))) => p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X37),s('type.extreal.extreal',X12))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X38)))))))),
  file('thm.extreal.add_infty|split|2','thm.extreal.let_add2')).
fof(f979,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X4),s('type.extreal.extreal',X5)))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal',X4))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X5))))))))) ) | $spl30),
  inference(cnf_transformation,[],[f979_D])).
fof(f979_D,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X4),s('type.extreal.extreal',X5)))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal',X4))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X5))))))))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f4622,plain,(
  ( ! [X4,X5,X3] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal',X4))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X5)))))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X4),s('type.extreal.extreal',X5))))) ) | $spl1),
  inference(forward_demodulation,[],[f4621,f560])).
fof(f4621,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal',X4))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X5)))))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X4),s('type.extreal.extreal',X5)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(forward_demodulation,[],[f4620,f560])).
fof(f4620,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal',X4))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X5)))))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X4),s('type.extreal.extreal',X5)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(forward_demodulation,[],[f4569,f549])).
fof(f4569,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal',X4))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',X5)))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X4),s('type.extreal.extreal',X5)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(trivial_inequality_removal,[],[f4526])).
fof(f4526,plain,(
  ( ! [X4,X5,X3] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.NegInf_0') | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal',X4))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',X5)))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X4),s('type.extreal.extreal',X5)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(superposition,[],[f744,f539])).
fof(f539,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X3)))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X3)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X1) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X1)) )),
  inference(cnf_transformation,[],[f385])).
fof(f385,plain,(
  ! [X0,X1,X2,X3] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X1) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X1) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X3)))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X3)))))))),
  inference(flattening,[],[f384])).
fof(f384,plain,(
  ! [X0,X1,X2,X3] : ((s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X1) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X1) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X3))))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X3)))))))),
  inference(ennf_transformation,[],[f266])).
fof(f266,plain,(
  ! [X0,X1,X2,X3] : ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X1) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X1) & p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) & p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X3))))) => p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X3)))))))),
  inference(flattening,[],[f265])).
fof(f265,plain,(
  ! [X0,X1,X2,X3] : ((~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X1) & ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X1) & p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) & p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X3))))) => p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X3)))))))),
  inference(rectify,[],[f52])).
fof(f52,axiom,(
  ! [X37,X6,X12,X38] : ((~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X6) & ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X6) & p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X37),s('type.extreal.extreal',X6)))) & p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X12),s('type.extreal.extreal',X38))))) => p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X37),s('type.extreal.extreal',X12))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X38)))))))),
  file('thm.extreal.add_infty|split|2','thm.extreal.let_add2_alt')).
fof(f1115,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,'F_0') != s(bool,X1)) ) | $spl56),
  inference(cnf_transformation,[],[f1115_D])).
fof(f1115_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,'F_0') != s(bool,X1)) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f4409,plain,(
  ( ! [X152,X153] : (p(s(bool,X152)) | ~p(s(bool,X153)) | s(bool,'F_0') != s(bool,X153)) ) | $spl1),
  inference(resolution,[],[f975,f737])).
fof(f975,plain,(
  ~p(s(bool,'F_0')) | $spl1),
  inference(forward_demodulation,[],[f974,f597])).
fof(f597,plain,(
  s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))),
  inference(cnf_transformation,[],[f134])).
fof(f134,axiom,(
  s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))),
  file('thm.extreal.add_infty|split|2','thm.extreal.extreal_le_def|split|4')).
fof(f974,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl1),
  inference(subsumption_resolution,[],[f883,f566])).
fof(f566,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.PosInf_0')))) )),
  inference(cnf_transformation,[],[f302])).
fof(f302,plain,(
  ! [X0] : s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.PosInf_0')))),
  inference(rectify,[],[f65])).
fof(f65,axiom,(
  ! [X6] : s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal','const.extreal.PosInf_0')))),
  file('thm.extreal.add_infty|split|2','thm.extreal.add_infty|split|0')).
fof(f883,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.PosInf_0'))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl1),
  inference(superposition,[],[f744,f729])).
fof(f729,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X0))))) )),
  inference(cnf_transformation,[],[f536])).
fof(f536,plain,(
  ! [X0] : ((~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X0)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)) & (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X0))))))),
  inference(nnf_transformation,[],[f377])).
fof(f377,plain,(
  ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X0)))) <=> s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0))),
  inference(rectify,[],[f101])).
fof(f101,axiom,(
  ! [X6] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X6)))) <=> s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X6))),
  file('thm.extreal.add_infty|split|2','thm.extreal.le_infty|split|3')).
fof(f4339,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,'F_0') != s(bool,X22) | s(bool,X21) != s(bool,X23)) ) | $spl1),
  inference(resolution,[],[f975,f606])).
fof(f4337,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,'F_0') != s(bool,X16)) ) | $spl1),
  inference(resolution,[],[f975,f605])).
fof(f1117,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,X5)) ) | $spl60),
  inference(cnf_transformation,[],[f1117_D])).
fof(f1117_D,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,X5)) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f4336,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,'F_0') != s(bool,X14) | s(bool,X12) != s(bool,X13)) ) | $spl1),
  inference(resolution,[],[f975,f604])).
fof(f4335,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,'F_0') != s(bool,X10)) ) | $spl1),
  inference(resolution,[],[f975,f604])).
fof(f4334,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,'F_0') != s(bool,X7) | s(bool,X6) != s(bool,X8)) ) | $spl1),
  inference(resolution,[],[f975,f602])).
fof(f4333,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5) | s(bool,X3) != s(bool,X4)) ) | $spl1),
  inference(resolution,[],[f975,f600])).
fof(f4332,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,'F_0') != s(bool,X1)) ) | $spl1),
  inference(resolution,[],[f975,f600])).
fof(f4138,plain,(
  s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13('const.extreal.Normal_1'(s('type.realax.real',sK13(sK11))))))) | $spl170),
  inference(cnf_transformation,[],[f4138_D])).
fof(f4138_D,plain,(
  s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13('const.extreal.Normal_1'(s('type.realax.real',sK13(sK11))))))) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f4305,plain,(
  ( ! [X218] : (s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13('const.extreal.Normal_1'(s('type.realax.real',sK13(sK11))))))) | s('type.extreal.extreal',sK11) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X218)))) ) | $spl24),
  inference(forward_demodulation,[],[f4304,f804])).
fof(f4304,plain,(
  ( ! [X218] : (s('type.extreal.extreal',sK11) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X218))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13('const.extreal.Normal_1'(s('type.realax.real',sK13(sK11)))))))) ) | $spl24),
  inference(subsumption_resolution,[],[f4303,f576])).
fof(f4303,plain,(
  ( ! [X218] : (s('type.extreal.extreal',sK11) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X218))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13('const.extreal.Normal_1'(s('type.realax.real',sK13(sK11)))))))) ) | $spl24),
  inference(subsumption_resolution,[],[f4043,f638])).
fof(f4043,plain,(
  ( ! [X218] : (s('type.extreal.extreal',sK11) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X218))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13('const.extreal.Normal_1'(s('type.realax.real',sK13(sK11)))))))) ) | $spl24),
  inference(superposition,[],[f584,f804])).
fof(f4309,plain,(
  ( ! [X219] : (s('type.extreal.extreal',sK11) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X219))) | s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13('const.extreal.Normal_1'(s('type.realax.real',sK13(sK11)))))))) ) | $spl24),
  inference(forward_demodulation,[],[f4308,f804])).
fof(f4308,plain,(
  ( ! [X219] : (s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13('const.extreal.Normal_1'(s('type.realax.real',sK13(sK11))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X219))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11))))) ) | $spl24),
  inference(forward_demodulation,[],[f4307,f804])).
fof(f4307,plain,(
  ( ! [X219] : (s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13('const.extreal.Normal_1'(s('type.realax.real',sK13(sK11))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X219))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11))))) ) | $spl24),
  inference(subsumption_resolution,[],[f4306,f576])).
fof(f4306,plain,(
  ( ! [X219] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13('const.extreal.Normal_1'(s('type.realax.real',sK13(sK11))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X219))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11))))) ) | $spl24),
  inference(subsumption_resolution,[],[f4044,f537])).
fof(f4044,plain,(
  ( ! [X219] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK11) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13('const.extreal.Normal_1'(s('type.realax.real',sK13(sK11))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X219))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11))))) ) | $spl24),
  inference(superposition,[],[f584,f804])).
fof(f4313,plain,(
  ( ! [X220] : (s('type.extreal.extreal',sK11) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X220))) | s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13('const.extreal.Normal_1'(s('type.realax.real',sK13(sK11)))))))) ) | ($spl5 | $spl24)),
  inference(forward_demodulation,[],[f4312,f804])).
fof(f4312,plain,(
  ( ! [X220] : (s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13('const.extreal.Normal_1'(s('type.realax.real',sK13(sK11))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X220))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11))))) ) | ($spl5 | $spl24)),
  inference(forward_demodulation,[],[f4311,f804])).
fof(f4311,plain,(
  ( ! [X220] : (s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13('const.extreal.Normal_1'(s('type.realax.real',sK13(sK11))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X220))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11))))) ) | ($spl5 | $spl24)),
  inference(subsumption_resolution,[],[f4310,f638])).
fof(f4310,plain,(
  ( ! [X220] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13('const.extreal.Normal_1'(s('type.realax.real',sK13(sK11))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X220))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11))))) ) | ($spl5 | $spl24)),
  inference(subsumption_resolution,[],[f4045,f792])).
fof(f4045,plain,(
  ( ! [X220] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK11) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13('const.extreal.Normal_1'(s('type.realax.real',sK13(sK11))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X220))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11))))) ) | $spl24),
  inference(superposition,[],[f584,f804])).
fof(f1490,plain,(
  ~p(s(bool,'F_0')) | $spl71),
  inference(cnf_transformation,[],[f1490_D])).
fof(f1490_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl71),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl71])])).
fof(f4328,plain,(
  ( ! [X10] : (~p(s(bool,X10)) | ~p(s(bool,'F_0'))) ) | $spl1),
  inference(duplicate_literal_removal,[],[f1111])).
fof(f1111,plain,(
  ( ! [X10] : (~p(s(bool,X10)) | ~p(s(bool,'F_0')) | ~p(s(bool,X10))) ) | $spl1),
  inference(superposition,[],[f975,f633])).
fof(f633,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f434])).
fof(f434,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f433])).
fof(f433,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f334])).
fof(f334,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X19,X20] : ((p(s(bool,X19)) => p(s(bool,X20))) => ((p(s(bool,X20)) => p(s(bool,X19))) => s(bool,X19) = s(bool,X20)))),
  file('thm.extreal.add_infty|split|2','thm.bool.IMP_ANTISYM_AX')).
fof(f4329,plain,(
  ( ! [X11] : (~p(s(bool,X11)) | ~p(s(bool,'F_0'))) ) | $spl1),
  inference(duplicate_literal_removal,[],[f1112])).
fof(f1112,plain,(
  ( ! [X11] : (~p(s(bool,X11)) | ~p(s(bool,X11)) | ~p(s(bool,'F_0'))) ) | $spl1),
  inference(superposition,[],[f975,f633])).
fof(f1472,plain,(
  ( ! [X10] : (~p(s(bool,X10)) | ~p(s(bool,'F_0'))) ) | $spl1),
  inference(duplicate_literal_removal,[],[f1111])).
fof(f1473,plain,(
  ( ! [X11] : (~p(s(bool,X11)) | ~p(s(bool,'F_0'))) ) | $spl1),
  inference(duplicate_literal_removal,[],[f1112])).
fof(f4158,plain,(
  $false | ($spl3 | $spl24)),
  inference(subsumption_resolution,[],[f3910,f746])).
fof(f3910,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) | $spl24),
  inference(superposition,[],[f564,f804])).
fof(f564,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))))) )),
  inference(cnf_transformation,[],[f300])).
fof(f300,plain,(
  ! [X0] : s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X22] : s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X22)))))),
  file('thm.extreal.add_infty|split|2','thm.extreal.extreal_add_def|split|5')).
fof(f4142,plain,(
  ( ! [X37] : (s('type.extreal.extreal',sK11) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X37))) | s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13('const.extreal.Normal_1'(s('type.realax.real',sK13(sK11)))))))) ) | ($spl5 | $spl24)),
  inference(forward_demodulation,[],[f4141,f804])).
fof(f4141,plain,(
  ( ! [X37] : (s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13('const.extreal.Normal_1'(s('type.realax.real',sK13(sK11))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X37))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11))))) ) | ($spl5 | $spl24)),
  inference(forward_demodulation,[],[f4140,f804])).
fof(f4140,plain,(
  ( ! [X37] : (s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13('const.extreal.Normal_1'(s('type.realax.real',sK13(sK11))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X37))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11))))) ) | ($spl5 | $spl24)),
  inference(subsumption_resolution,[],[f4139,f638])).
fof(f4139,plain,(
  ( ! [X37] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13('const.extreal.Normal_1'(s('type.realax.real',sK13(sK11))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X37))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11))))) ) | ($spl5 | $spl24)),
  inference(subsumption_resolution,[],[f3900,f792])).
fof(f3900,plain,(
  ( ! [X37] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK11) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13('const.extreal.Normal_1'(s('type.realax.real',sK13(sK11))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X37))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11))))) ) | $spl24),
  inference(superposition,[],[f804,f584])).
fof(f4136,plain,(
  ( ! [X36] : (s('type.extreal.extreal',sK11) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X36))) | s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13('const.extreal.Normal_1'(s('type.realax.real',sK13(sK11)))))))) ) | $spl24),
  inference(forward_demodulation,[],[f4135,f804])).
fof(f4135,plain,(
  ( ! [X36] : (s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13('const.extreal.Normal_1'(s('type.realax.real',sK13(sK11))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X36))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11))))) ) | $spl24),
  inference(forward_demodulation,[],[f4134,f804])).
fof(f4134,plain,(
  ( ! [X36] : (s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13('const.extreal.Normal_1'(s('type.realax.real',sK13(sK11))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X36))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11))))) ) | $spl24),
  inference(subsumption_resolution,[],[f4133,f576])).
fof(f4133,plain,(
  ( ! [X36] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13('const.extreal.Normal_1'(s('type.realax.real',sK13(sK11))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X36))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11))))) ) | $spl24),
  inference(subsumption_resolution,[],[f3899,f537])).
fof(f3899,plain,(
  ( ! [X36] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK11) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13('const.extreal.Normal_1'(s('type.realax.real',sK13(sK11))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X36))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11))))) ) | $spl24),
  inference(superposition,[],[f804,f584])).
fof(f3870,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,'T_0') != s(bool,X1)) ) | $spl168),
  inference(cnf_transformation,[],[f3870_D])).
fof(f3870_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,'T_0') != s(bool,X1)) ) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f3857,plain,(
  ( ! [X152,X153] : (p(s(bool,X152)) | p(s(bool,X153)) | s(bool,'T_0') != s(bool,X153)) ) | $spl150),
  inference(resolution,[],[f3069,f737])).
fof(f3069,plain,(
  p(s(bool,'T_0')) | $spl150),
  inference(cnf_transformation,[],[f3069_D])).
fof(f3069_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f3787,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,'T_0') != s(bool,X21) | s(bool,X22) != s(bool,X23)) ) | $spl150),
  inference(resolution,[],[f3069,f606])).
fof(f3778,plain,(
  ( ! [X5] : (s(bool,'T_0') != s(bool,X5)) ) | $spl166),
  inference(cnf_transformation,[],[f3778_D])).
fof(f3778_D,plain,(
  ( ! [X5] : (s(bool,'T_0') != s(bool,X5)) ) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f3786,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,X19) | s(bool,'T_0') != s(bool,X20)) ) | $spl150),
  inference(resolution,[],[f3069,f606])).
fof(f3784,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,'T_0') != s(bool,X14)) ) | $spl150),
  inference(resolution,[],[f3069,f605])).
fof(f3783,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,'T_0') != s(bool,X10)) ) | $spl150),
  inference(resolution,[],[f3069,f604])).
fof(f3782,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,'T_0') != s(bool,X6) | s(bool,X7) != s(bool,X8)) ) | $spl150),
  inference(resolution,[],[f3069,f602])).
fof(f3781,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X4) | s(bool,'T_0') != s(bool,X5)) ) | $spl150),
  inference(resolution,[],[f3069,f602])).
fof(f3780,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,'T_0') != s(bool,X1)) ) | $spl150),
  inference(resolution,[],[f3069,f600])).
fof(f3777,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,'T_0') != s(bool,X1)) ) | $spl164),
  inference(cnf_transformation,[],[f3777_D])).
fof(f3777_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,'T_0') != s(bool,X1)) ) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f3764,plain,(
  ( ! [X152,X153] : (p(s(bool,X152)) | ~p(s(bool,X153)) | s(bool,'T_0') != s(bool,X153)) ) | $spl151),
  inference(resolution,[],[f3070,f737])).
fof(f3694,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,'T_0') != s(bool,X22) | s(bool,X21) != s(bool,X23)) ) | $spl151),
  inference(resolution,[],[f3070,f606])).
fof(f3692,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,'T_0') != s(bool,X16)) ) | $spl151),
  inference(resolution,[],[f3070,f605])).
fof(f3691,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,'T_0') != s(bool,X14) | s(bool,X12) != s(bool,X13)) ) | $spl151),
  inference(resolution,[],[f3070,f604])).
fof(f3690,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,'T_0') != s(bool,X10)) ) | $spl151),
  inference(resolution,[],[f3070,f604])).
fof(f3689,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,'T_0') != s(bool,X7) | s(bool,X6) != s(bool,X8)) ) | $spl151),
  inference(resolution,[],[f3070,f602])).
fof(f3688,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,'T_0') != s(bool,X5) | s(bool,X3) != s(bool,X4)) ) | $spl151),
  inference(resolution,[],[f3070,f600])).
fof(f3687,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,'T_0') != s(bool,X1)) ) | $spl151),
  inference(resolution,[],[f3070,f600])).
fof(f804,plain,(
  s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11)))) | $spl24),
  inference(cnf_transformation,[],[f804_D])).
fof(f804_D,plain,(
  s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11)))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f3674,plain,(
  ( ! [X37] : (s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11)))) | s('type.extreal.extreal',sK11) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X37)))) ) | $spl3),
  inference(subsumption_resolution,[],[f3673,f537])).
fof(f3673,plain,(
  ( ! [X37] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK11) | s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11)))) | s('type.extreal.extreal',sK11) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X37)))) ) | $spl3),
  inference(subsumption_resolution,[],[f3570,f736])).
fof(f736,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0))))) )),
  inference(cnf_transformation,[],[f382])).
fof(f382,plain,(
  ! [X0] : p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0))))),
  inference(rectify,[],[f184])).
fof(f184,axiom,(
  ! [X6] : p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X6))))),
  file('thm.extreal.add_infty|split|2','thm.extreal.le_refl')).
fof(f3570,plain,(
  ( ! [X37] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK11) | s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11)))) | s('type.extreal.extreal',sK11) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X37)))) ) | $spl3),
  inference(superposition,[],[f2676,f584])).
fof(f2676,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl3),
  inference(subsumption_resolution,[],[f2572,f570])).
fof(f2572,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl3),
  inference(superposition,[],[f746,f580])).
fof(f3544,plain,(
  ( ! [X152,X153] : (p(s(bool,X152)) | ~p(s(bool,X153)) | s(bool,X153) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl3),
  inference(resolution,[],[f2676,f737])).
fof(f3474,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s(bool,X21) != s(bool,X23)) ) | $spl3),
  inference(resolution,[],[f2676,f606])).
fof(f3472,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl3),
  inference(resolution,[],[f2676,f605])).
fof(f3471,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s(bool,X12) != s(bool,X13)) ) | $spl3),
  inference(resolution,[],[f2676,f604])).
fof(f3470,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl3),
  inference(resolution,[],[f2676,f604])).
fof(f3469,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s(bool,X6) != s(bool,X8)) ) | $spl3),
  inference(resolution,[],[f2676,f602])).
fof(f3468,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s(bool,X3) != s(bool,X4)) ) | $spl3),
  inference(resolution,[],[f2676,f600])).
fof(f3467,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl3),
  inference(resolution,[],[f2676,f600])).
fof(f3465,plain,(
  ( ! [X37] : (s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11)))) | s('type.extreal.extreal',sK11) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X37)))) ) | $spl3),
  inference(subsumption_resolution,[],[f3464,f537])).
fof(f3464,plain,(
  ( ! [X37] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK11) | s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11)))) | s('type.extreal.extreal',sK11) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X37)))) ) | $spl3),
  inference(subsumption_resolution,[],[f3414,f574])).
fof(f3414,plain,(
  ( ! [X37] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK11) | s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11)))) | s('type.extreal.extreal',sK11) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X37)))) ) | $spl3),
  inference(superposition,[],[f2675,f584])).
fof(f2675,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))) | $spl3),
  inference(subsumption_resolution,[],[f2571,f570])).
fof(f2571,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))) | $spl3),
  inference(superposition,[],[f746,f578])).
fof(f3461,plain,(
  $false | ($spl3 | $spl22)),
  inference(subsumption_resolution,[],[f3460,f537])).
fof(f3460,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK11) | ($spl3 | $spl22)),
  inference(subsumption_resolution,[],[f3459,f802])).
fof(f3459,plain,(
  s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK35(sK11)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK11) | $spl3),
  inference(subsumption_resolution,[],[f3418,f574])).
fof(f3418,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK35(sK11)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK11) | $spl3),
  inference(superposition,[],[f2675,f591])).
fof(f3458,plain,(
  $false | ($spl3 | $spl22)),
  inference(subsumption_resolution,[],[f3457,f537])).
fof(f3457,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK11) | ($spl3 | $spl22)),
  inference(subsumption_resolution,[],[f3456,f802])).
fof(f3456,plain,(
  s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK34(sK11)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK11) | $spl3),
  inference(subsumption_resolution,[],[f3416,f574])).
fof(f3416,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK34(sK11)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK11) | $spl3),
  inference(superposition,[],[f2675,f590])).
fof(f3343,plain,(
  ( ! [X161,X160] : (p(s(bool,X160)) | p(s(bool,X161)) | s(bool,X161) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))) ) | $spl3),
  inference(resolution,[],[f2675,f737])).
fof(f3273,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) | s(bool,X30) != s(bool,X31)) ) | $spl3),
  inference(resolution,[],[f2675,f606])).
fof(f3272,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,X27) | s(bool,X28) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))) ) | $spl3),
  inference(resolution,[],[f2675,f606])).
fof(f3270,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X22) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))) ) | $spl3),
  inference(resolution,[],[f2675,f605])).
fof(f3269,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X18) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))) ) | $spl3),
  inference(resolution,[],[f2675,f604])).
fof(f3268,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X14) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) | s(bool,X15) != s(bool,X16)) ) | $spl3),
  inference(resolution,[],[f2675,f602])).
fof(f3267,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X12) | s(bool,X13) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))) ) | $spl3),
  inference(resolution,[],[f2675,f602])).
fof(f3266,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X8)) | p(s(bool,X9)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))) ) | $spl3),
  inference(resolution,[],[f2675,f600])).
fof(f3167,plain,(
  $false | ($spl3 | $spl70)),
  inference(subsumption_resolution,[],[f1489,f2629])).
fof(f2629,plain,(
  ~p(s(bool,'F_0')) | $spl3),
  inference(forward_demodulation,[],[f2628,f597])).
fof(f2628,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl3),
  inference(subsumption_resolution,[],[f2550,f567])).
fof(f567,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X0)))) )),
  inference(cnf_transformation,[],[f303])).
fof(f303,plain,(
  ! [X0] : s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X0)))),
  inference(rectify,[],[f63])).
fof(f63,axiom,(
  ! [X6] : s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X6)))),
  file('thm.extreal.add_infty|split|2','thm.extreal.add_infty|split|1')).
fof(f2550,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK11))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl3),
  inference(superposition,[],[f746,f729])).
fof(f3153,plain,(
  ( ! [X152,X153] : (p(s(bool,X152)) | ~p(s(bool,X153)) | s(bool,'F_0') != s(bool,X153)) ) | $spl3),
  inference(resolution,[],[f2629,f737])).
fof(f3083,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,'F_0') != s(bool,X22) | s(bool,X21) != s(bool,X23)) ) | $spl3),
  inference(resolution,[],[f2629,f606])).
fof(f3081,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,'F_0') != s(bool,X16)) ) | $spl3),
  inference(resolution,[],[f2629,f605])).
fof(f3080,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,'F_0') != s(bool,X14) | s(bool,X12) != s(bool,X13)) ) | $spl3),
  inference(resolution,[],[f2629,f604])).
fof(f3079,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,'F_0') != s(bool,X10)) ) | $spl3),
  inference(resolution,[],[f2629,f604])).
fof(f3078,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,'F_0') != s(bool,X7) | s(bool,X6) != s(bool,X8)) ) | $spl3),
  inference(resolution,[],[f2629,f602])).
fof(f3077,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5) | s(bool,X3) != s(bool,X4)) ) | $spl3),
  inference(resolution,[],[f2629,f600])).
fof(f3076,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,'F_0') != s(bool,X1)) ) | $spl3),
  inference(resolution,[],[f2629,f600])).
fof(f3075,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) | ~p(s(bool,'F_0')) | $spl0),
  inference(forward_demodulation,[],[f2470,f560])).
fof(f2470,plain,(
  ~p(s(bool,'F_0')) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))) | $spl0),
  inference(forward_demodulation,[],[f2318,f597])).
fof(f2318,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))) | $spl0),
  inference(superposition,[],[f729,f743])).
fof(f743,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))) | $spl0),
  inference(cnf_transformation,[],[f743_D])).
fof(f743_D,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f1489,plain,(
  p(s(bool,'F_0')) | $spl70),
  inference(cnf_transformation,[],[f1489_D])).
fof(f1489_D,plain,(
  p(s(bool,'F_0')) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f980,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) | $spl33),
  inference(cnf_transformation,[],[f980_D])).
fof(f980_D,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f3074,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) | p(s(bool,'F_0')) | $spl0),
  inference(forward_demodulation,[],[f2473,f560])).
fof(f2473,plain,(
  p(s(bool,'F_0')) | s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))) | $spl0),
  inference(forward_demodulation,[],[f2320,f597])).
fof(f2320,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))) | $spl0),
  inference(superposition,[],[f730,f743])).
fof(f730,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X0)))) | s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0)) )),
  inference(cnf_transformation,[],[f536])).
fof(f3073,plain,(
  sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),'F_0') | $spl154),
  inference(cnf_transformation,[],[f3073_D])).
fof(f3073_D,plain,(
  sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),'F_0') <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f3070,plain,(
  ~p(s(bool,'T_0')) | $spl151),
  inference(cnf_transformation,[],[f3070_D])).
fof(f3070_D,plain,(
  ~p(s(bool,'T_0')) <=> ~$spl151),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl151])])).
fof(f3043,plain,(
  ~p(s(bool,'T_0')) | sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),'F_0') | $spl71),
  inference(superposition,[],[f2047,f735])).
fof(f735,plain,(
  s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))),
  inference(cnf_transformation,[],[f189])).
fof(f189,axiom,(
  s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))),
  file('thm.extreal.add_infty|split|2','thm.extreal.extreal_le_def|split|2')).
fof(f2047,plain,(
  ( ! [X146] : (~p(s(bool,X146)) | sP10(X146,'F_0')) ) | $spl71),
  inference(resolution,[],[f1490,f699])).
fof(f699,plain,(
  ( ! [X0,X1] : (sP10(X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f524])).
fof(f524,plain,(
  ! [X0,X1] : ((~sP10(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP10(X0,X1)))),
  inference(rectify,[],[f523])).
fof(f523,plain,(
  ! [X1,X0] : ((~sP10(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP10(X1,X0)))),
  inference(flattening,[],[f522])).
fof(f522,plain,(
  ! [X1,X0] : ((~sP10(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP10(X1,X0)))),
  inference(nnf_transformation,[],[f461])).
fof(f461,plain,(
  ! [X1,X0] : (sP10(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f3071,plain,(
  ( ! [X9] : (sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X9))),s('type.extreal.extreal','const.extreal.PosInf_0')),'F_0')) ) | $spl152),
  inference(cnf_transformation,[],[f3071_D])).
fof(f3071_D,plain,(
  ( ! [X9] : (sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X9))),s('type.extreal.extreal','const.extreal.PosInf_0')),'F_0')) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f3042,plain,(
  ( ! [X9] : (~p(s(bool,'T_0')) | sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X9))),s('type.extreal.extreal','const.extreal.PosInf_0')),'F_0')) ) | $spl71),
  inference(superposition,[],[f2047,f732])).
fof(f732,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0')))) )),
  inference(cnf_transformation,[],[f379])).
fof(f379,plain,(
  ! [X0] : s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0')))),
  inference(rectify,[],[f187])).
fof(f187,axiom,(
  ! [X42] : s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X42))),s('type.extreal.extreal','const.extreal.PosInf_0')))),
  file('thm.extreal.add_infty|split|2','thm.extreal.extreal_le_def|split|3')).
fof(f3068,plain,(
  ( ! [X8] : (sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X8)),'F_0')) ) | $spl148),
  inference(cnf_transformation,[],[f3068_D])).
fof(f3068_D,plain,(
  ( ! [X8] : (sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X8)),'F_0')) ) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f3041,plain,(
  ( ! [X8] : (~p(s(bool,'T_0')) | sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X8)),'F_0')) ) | $spl71),
  inference(superposition,[],[f2047,f594])).
fof(f594,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)))) )),
  inference(cnf_transformation,[],[f328])).
fof(f328,plain,(
  ! [X0] : s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)))),
  inference(rectify,[],[f174])).
fof(f174,axiom,(
  ! [X11] : s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X11)))),
  file('thm.extreal.add_infty|split|2','thm.extreal.extreal_le_def|split|1')).
fof(f3067,plain,(
  ( ! [X24] : (sP10(X24,'F_0')) ) | $spl146),
  inference(cnf_transformation,[],[f3067_D])).
fof(f3067_D,plain,(
  ( ! [X24] : (sP10(X24,'F_0')) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f3009,plain,(
  ( ! [X233,X234,X232] : (sP10(X232,'F_0') | ~p(s(bool,X233)) | p(s(bool,X234)) | s(bool,X233) != s(bool,X234)) ) | $spl71),
  inference(resolution,[],[f2047,f737])).
fof(f3064,plain,(
  ( ! [X2,X0] : (sP10(X0,'F_0') | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) | $spl142),
  inference(cnf_transformation,[],[f3064_D])).
fof(f3064_D,plain,(
  ( ! [X2,X0] : (sP10(X0,'F_0') | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f3008,plain,(
  ( ! [X231,X229,X230] : (sP10(X229,'F_0') | p(s(bool,X230)) | ~p(s(bool,X231)) | s(bool,X229) != s(bool,X231)) ) | $spl71),
  inference(resolution,[],[f2047,f737])).
fof(f2938,plain,(
  ( ! [X30,X28,X31,X29] : (sP10(X28,'F_0') | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) ) | $spl71),
  inference(resolution,[],[f2047,f606])).
fof(f3066,plain,(
  ( ! [X24,X27,X25] : (sP10(X24,'F_0') | ~p(s(bool,X25)) | p(s(bool,X27)) | s(bool,X25) != s(bool,X27)) ) | $spl71),
  inference(subsumption_resolution,[],[f2937,f737])).
fof(f2937,plain,(
  ( ! [X26,X24,X27,X25] : (sP10(X24,'F_0') | ~p(s(bool,X25)) | ~p(s(bool,X26)) | p(s(bool,X27)) | s(bool,X25) != s(bool,X27)) ) | $spl71),
  inference(resolution,[],[f2047,f605])).
fof(f2936,plain,(
  ( ! [X23,X21,X22,X20] : (sP10(X20,'F_0') | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X20) != s(bool,X22)) ) | $spl71),
  inference(resolution,[],[f2047,f605])).
fof(f3065,plain,(
  ( ! [X4,X7] : (sP10(X4,'F_0') | s(bool,X4) != s(bool,X7)) ) | $spl144),
  inference(cnf_transformation,[],[f3065_D])).
fof(f3065_D,plain,(
  ( ! [X4,X7] : (sP10(X4,'F_0') | s(bool,X4) != s(bool,X7)) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f2935,plain,(
  ( ! [X19,X17,X18,X16] : (sP10(X16,'F_0') | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X16) != s(bool,X19) | s(bool,X17) != s(bool,X18)) ) | $spl71),
  inference(resolution,[],[f2047,f604])).
fof(f2934,plain,(
  ( ! [X14,X12,X15,X13] : (sP10(X12,'F_0') | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl71),
  inference(resolution,[],[f2047,f604])).
fof(f2933,plain,(
  ( ! [X10,X8,X11,X9] : (sP10(X8,'F_0') | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) ) | $spl71),
  inference(resolution,[],[f2047,f602])).
fof(f2932,plain,(
  ( ! [X6,X4,X7,X5] : (sP10(X4,'F_0') | ~p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X4) != s(bool,X7) | s(bool,X5) != s(bool,X6)) ) | $spl71),
  inference(resolution,[],[f2047,f600])).
fof(f2931,plain,(
  ( ! [X2,X0,X3,X1] : (sP10(X0,'F_0') | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl71),
  inference(resolution,[],[f2047,f600])).
fof(f2920,plain,(
  ( ! [X36] : (s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11)))) | s('type.extreal.extreal',sK11) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X36)))) ) | $spl5),
  inference(subsumption_resolution,[],[f2919,f792])).
fof(f2919,plain,(
  ( ! [X36] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK11) | s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11)))) | s('type.extreal.extreal',sK11) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X36)))) )),
  inference(subsumption_resolution,[],[f2847,f725])).
fof(f725,plain,(
  ( ! [X0] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X0))))) )),
  inference(cnf_transformation,[],[f372])).
fof(f372,plain,(
  ! [X0] : ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X0))))),
  inference(flattening,[],[f371])).
fof(f371,plain,(
  ! [X0] : ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X0))))),
  inference(rectify,[],[f157])).
fof(f157,axiom,(
  ! [X6] : ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X6))))),
  file('thm.extreal.add_infty|split|2','thm.extreal.lt_infty|split|4')).
fof(f2847,plain,(
  ( ! [X36] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK11) | s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11)))) | s('type.extreal.extreal',sK11) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X36)))) )),
  inference(superposition,[],[f778,f584])).
fof(f778,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.PosInf_0'))))),
  inference(trivial_inequality_removal,[],[f775])).
fof(f775,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.PosInf_0') | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.PosInf_0'))))),
  inference(superposition,[],[f537,f727])).
fof(f727,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)) )),
  inference(cnf_transformation,[],[f535])).
fof(f535,plain,(
  ! [X0] : ((s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.PosInf_0'))))) & (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0)))),
  inference(nnf_transformation,[],[f376])).
fof(f376,plain,(
  ! [X0] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0) <=> p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.PosInf_0')))))),
  inference(flattening,[],[f375])).
fof(f375,plain,(
  ! [X0] : (~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) <=> p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.PosInf_0')))))),
  inference(rectify,[],[f122])).
fof(f122,axiom,(
  ! [X6] : (~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X6) <=> p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal','const.extreal.PosInf_0')))))),
  file('thm.extreal.add_infty|split|2','thm.extreal.lt_infty|split|5')).
fof(f2899,plain,(
  ( ! [X9] : (p(s(bool,X9)) | s(bool,X9) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl138),
  inference(cnf_transformation,[],[f2899_D])).
fof(f2899_D,plain,(
  ( ! [X9] : (p(s(bool,X9)) | s(bool,X9) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f2822,plain,(
  ( ! [X161,X160] : (p(s(bool,X160)) | p(s(bool,X161)) | s(bool,X161) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.PosInf_0')))) )),
  inference(resolution,[],[f778,f737])).
fof(f2752,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s(bool,X30) != s(bool,X31)) )),
  inference(resolution,[],[f778,f606])).
fof(f2900,plain,(
  ( ! [X13] : (s(bool,X13) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl140),
  inference(cnf_transformation,[],[f2900_D])).
fof(f2900_D,plain,(
  ( ! [X13] : (s(bool,X13) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f2751,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,X27) | s(bool,X28) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.PosInf_0')))) )),
  inference(resolution,[],[f778,f606])).
fof(f2749,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X22) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.PosInf_0')))) )),
  inference(resolution,[],[f778,f605])).
fof(f2748,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X18) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.PosInf_0')))) )),
  inference(resolution,[],[f778,f604])).
fof(f2747,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X14) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s(bool,X15) != s(bool,X16)) )),
  inference(resolution,[],[f778,f602])).
fof(f2746,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X12) | s(bool,X13) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.PosInf_0')))) )),
  inference(resolution,[],[f778,f602])).
fof(f2745,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X8)) | p(s(bool,X9)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.PosInf_0')))) )),
  inference(resolution,[],[f778,f600])).
fof(f2708,plain,(
  s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK35('const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))))) | $spl136),
  inference(cnf_transformation,[],[f2708_D])).
fof(f2708_D,plain,(
  s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK35('const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))))) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f2611,plain,(
  s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK35('const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) | $spl3),
  inference(trivial_inequality_removal,[],[f2608])).
fof(f2608,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.NegInf_0') | s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK35('const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) | $spl3),
  inference(superposition,[],[f746,f591])).
fof(f2706,plain,(
  s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK34('const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))))) | $spl134),
  inference(cnf_transformation,[],[f2706_D])).
fof(f2706_D,plain,(
  s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK34('const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))))) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f2612,plain,(
  s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK34('const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) | $spl3),
  inference(trivial_inequality_removal,[],[f2606])).
fof(f2606,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.NegInf_0') | s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK34('const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) | $spl3),
  inference(superposition,[],[f746,f590])).
fof(f2704,plain,(
  s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13('const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))))) | $spl132),
  inference(cnf_transformation,[],[f2704_D])).
fof(f2704_D,plain,(
  s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13('const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))))) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f2613,plain,(
  ( ! [X37] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) | s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13('const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X37))) != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))) ) | $spl3),
  inference(trivial_inequality_removal,[],[f2604])).
fof(f2604,plain,(
  ( ! [X37] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.NegInf_0') | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) | s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13('const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X37))) != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))) ) | $spl3),
  inference(superposition,[],[f746,f584])).
fof(f2702,plain,(
  ( ! [X35] : (s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X35))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X35)))))))) ) | $spl130),
  inference(cnf_transformation,[],[f2702_D])).
fof(f2702_D,plain,(
  ( ! [X35] : (s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X35))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X35)))))))) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f2701,plain,(
  ( ! [X35] : (s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X35))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X35))))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))) ) | $spl3),
  inference(forward_demodulation,[],[f2616,f549])).
fof(f2616,plain,(
  ( ! [X35] : (s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X35))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))),s('type.extreal.extreal',X35))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))) ) | $spl3),
  inference(trivial_inequality_removal,[],[f2600])).
fof(f2600,plain,(
  ( ! [X35] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.NegInf_0') | s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X35))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))),s('type.extreal.extreal',X35))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))) ) | $spl3),
  inference(superposition,[],[f746,f553])).
fof(f2617,plain,(
  ( ! [X33] : (s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X33))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X33),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))) ) | $spl3),
  inference(trivial_inequality_removal,[],[f2598])).
fof(f2598,plain,(
  ( ! [X33] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.NegInf_0') | s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X33))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X33),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))) ) | $spl3),
  inference(superposition,[],[f746,f552])).
fof(f2620,plain,(
  ( ! [X26,X27] : (s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X26),s('type.extreal.extreal',X27))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X26),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X27),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))) ) | $spl3),
  inference(trivial_inequality_removal,[],[f2592])).
fof(f2592,plain,(
  ( ! [X26,X27] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.NegInf_0') | s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X26),s('type.extreal.extreal',X27))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X26),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X27),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))) ) | $spl3),
  inference(superposition,[],[f746,f546])).
fof(f2700,plain,(
  ( ! [X23,X22] : (s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X22),s('type.extreal.extreal',X23))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X22))))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X23)))))))) ) | $spl128),
  inference(cnf_transformation,[],[f2700_D])).
fof(f2700_D,plain,(
  ( ! [X23,X22] : (s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X22),s('type.extreal.extreal',X23))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X22))))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X23)))))))) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f2699,plain,(
  ( ! [X23,X22] : (s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X22),s('type.extreal.extreal',X23))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X22))))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X23))))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))) ) | $spl3),
  inference(forward_demodulation,[],[f2698,f549])).
fof(f2698,plain,(
  ( ! [X23,X22] : (s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X22),s('type.extreal.extreal',X23))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))),s('type.extreal.extreal',X22))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X23))))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))) ) | $spl3),
  inference(forward_demodulation,[],[f2621,f549])).
fof(f2621,plain,(
  ( ! [X23,X22] : (s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X22),s('type.extreal.extreal',X23))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))),s('type.extreal.extreal',X22))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))),s('type.extreal.extreal',X23))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))) ) | $spl3),
  inference(trivial_inequality_removal,[],[f2590])).
fof(f2590,plain,(
  ( ! [X23,X22] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.NegInf_0') | s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X22),s('type.extreal.extreal',X23))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))),s('type.extreal.extreal',X22))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))),s('type.extreal.extreal',X23))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))) ) | $spl3),
  inference(superposition,[],[f746,f545])).
fof(f2622,plain,(
  ( ! [X19,X18] : (s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X18),s('type.extreal.extreal',X19))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X18),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X19),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))) ) | $spl3),
  inference(trivial_inequality_removal,[],[f2588])).
fof(f2588,plain,(
  ( ! [X19,X18] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.NegInf_0') | s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X18),s('type.extreal.extreal',X19))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X18),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X19),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))) ) | $spl3),
  inference(superposition,[],[f746,f544])).
fof(f2697,plain,(
  ( ! [X14,X15] : (s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X14),s('type.extreal.extreal',X15))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X14))))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X15)))))))) ) | $spl126),
  inference(cnf_transformation,[],[f2697_D])).
fof(f2697_D,plain,(
  ( ! [X14,X15] : (s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X14),s('type.extreal.extreal',X15))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X14))))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X15)))))))) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f2696,plain,(
  ( ! [X14,X15] : (s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X14),s('type.extreal.extreal',X15))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X14))))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X15))))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))) ) | $spl3),
  inference(forward_demodulation,[],[f2695,f549])).
fof(f2695,plain,(
  ( ! [X14,X15] : (s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X14),s('type.extreal.extreal',X15))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))),s('type.extreal.extreal',X14))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X15))))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))) ) | $spl3),
  inference(forward_demodulation,[],[f2623,f549])).
fof(f2623,plain,(
  ( ! [X14,X15] : (s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X14),s('type.extreal.extreal',X15))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))),s('type.extreal.extreal',X14))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))),s('type.extreal.extreal',X15))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))) ) | $spl3),
  inference(trivial_inequality_removal,[],[f2586])).
fof(f2586,plain,(
  ( ! [X14,X15] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.NegInf_0') | s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X14),s('type.extreal.extreal',X15))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))),s('type.extreal.extreal',X14))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))),s('type.extreal.extreal',X15))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))) ) | $spl3),
  inference(superposition,[],[f746,f543])).
fof(f2694,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X9))))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X10),s('type.extreal.extreal',X11)))))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))),s('type.extreal.extreal',X10)))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X9),s('type.extreal.extreal',X11))))) ) | $spl124),
  inference(cnf_transformation,[],[f2694_D])).
fof(f2694_D,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X9))))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X10),s('type.extreal.extreal',X11)))))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))),s('type.extreal.extreal',X10)))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X9),s('type.extreal.extreal',X11))))) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f2693,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X9))))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X10),s('type.extreal.extreal',X11)))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X9),s('type.extreal.extreal',X11)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))),s('type.extreal.extreal',X10)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))) ) | $spl3),
  inference(forward_demodulation,[],[f2624,f549])).
fof(f2624,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))),s('type.extreal.extreal',X9))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X10),s('type.extreal.extreal',X11)))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X9),s('type.extreal.extreal',X11)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))),s('type.extreal.extreal',X10)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))) ) | $spl3),
  inference(trivial_inequality_removal,[],[f2584])).
fof(f2584,plain,(
  ( ! [X10,X11,X9] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.NegInf_0') | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))),s('type.extreal.extreal',X9))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X10),s('type.extreal.extreal',X11)))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X9),s('type.extreal.extreal',X11)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))),s('type.extreal.extreal',X10)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))) ) | $spl3),
  inference(superposition,[],[f746,f540])).
fof(f2692,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal',X4))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X5)))))))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X4),s('type.extreal.extreal',X5))))) ) | $spl122),
  inference(cnf_transformation,[],[f2692_D])).
fof(f2692_D,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal',X4))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X5)))))))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X4),s('type.extreal.extreal',X5))))) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f2691,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal',X4))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X5)))))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X4),s('type.extreal.extreal',X5)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))) ) | $spl3),
  inference(forward_demodulation,[],[f2625,f549])).
fof(f2625,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal',X4))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))),s('type.extreal.extreal',X5)))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X4),s('type.extreal.extreal',X5)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))) ) | $spl3),
  inference(trivial_inequality_removal,[],[f2582])).
fof(f2582,plain,(
  ( ! [X4,X5,X3] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.NegInf_0') | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal',X4))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))),s('type.extreal.extreal',X5)))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X4),s('type.extreal.extreal',X5)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))) ) | $spl3),
  inference(superposition,[],[f746,f539])).
fof(f2680,plain,(
  ( ! [X37] : (s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11)))) | s('type.extreal.extreal',sK11) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X37)))) ) | $spl3),
  inference(subsumption_resolution,[],[f2679,f537])).
fof(f2679,plain,(
  ( ! [X37] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK11) | s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11)))) | s('type.extreal.extreal',sK11) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X37)))) ) | $spl3),
  inference(subsumption_resolution,[],[f2574,f570])).
fof(f2574,plain,(
  ( ! [X37] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK11) | s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11)))) | s('type.extreal.extreal',sK11) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X37)))) ) | $spl3),
  inference(superposition,[],[f746,f584])).
fof(f2464,plain,(
  ( ! [X219] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X219),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X219),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))))))))) ) | $spl120),
  inference(cnf_transformation,[],[f2464_D])).
fof(f2464_D,plain,(
  ( ! [X219] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X219),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X219),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))))))))) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f2463,plain,(
  ( ! [X219] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X219),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X219),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))))))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0'))))) ) | $spl0),
  inference(forward_demodulation,[],[f2284,f560])).
fof(f2284,plain,(
  ( ! [X219] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X219),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X219),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))))))))) ) | $spl0),
  inference(superposition,[],[f558,f743])).
fof(f558,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X1))))) )),
  inference(cnf_transformation,[],[f418])).
fof(f418,plain,(
  ! [X0,X1] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X1)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))))))),
  inference(ennf_transformation,[],[f294])).
fof(f294,plain,(
  ! [X0,X1] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X1)))) => p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))))))),
  inference(rectify,[],[f82])).
fof(f82,axiom,(
  ! [X6,X12] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X12)))) => p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X12)))))))),
  file('thm.extreal.add_infty|split|2','thm.extreal.le_addr_imp')).
fof(f2462,plain,(
  ( ! [X218] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X218),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X218),s('type.extreal.extreal','const.extreal.NegInf_0'))))))) ) | $spl118),
  inference(cnf_transformation,[],[f2462_D])).
fof(f2462_D,plain,(
  ( ! [X218] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X218),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X218),s('type.extreal.extreal','const.extreal.NegInf_0'))))))) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f2443,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))))) | $spl101),
  inference(cnf_transformation,[],[f2443_D])).
fof(f2443_D,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))))) <=> ~$spl101),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl101])])).
fof(f2461,plain,(
  ( ! [X218] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X218),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X218),s('type.extreal.extreal','const.extreal.NegInf_0'))))))) ) | $spl0),
  inference(forward_demodulation,[],[f2283,f560])).
fof(f2283,plain,(
  ( ! [X218] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X218),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X218),s('type.extreal.extreal','const.extreal.NegInf_0')))))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))))))) ) | $spl0),
  inference(superposition,[],[f558,f743])).
fof(f2460,plain,(
  ( ! [X216] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X216),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X216))))) ) | $spl116),
  inference(cnf_transformation,[],[f2460_D])).
fof(f2460_D,plain,(
  ( ! [X216] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X216),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X216))))) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f2459,plain,(
  ( ! [X216] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X216),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))))))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X216))))) ) | $spl0),
  inference(forward_demodulation,[],[f2281,f560])).
fof(f2281,plain,(
  ( ! [X216] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X216),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X216))))) ) | $spl0),
  inference(superposition,[],[f557,f743])).
fof(f557,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X1)))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X0))))) )),
  inference(cnf_transformation,[],[f417])).
fof(f417,plain,(
  ! [X0,X1] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X0)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X1)))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))))))),
  inference(flattening,[],[f416])).
fof(f416,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X0)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X1))))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))))))),
  inference(ennf_transformation,[],[f293])).
fof(f293,plain,(
  ! [X0,X1] : ((p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X0)))) & p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X1))))) => p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))))))),
  inference(rectify,[],[f83])).
fof(f83,axiom,(
  ! [X6,X12] : ((p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X6)))) & p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X12))))) => p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X12)))))))),
  file('thm.extreal.add_infty|split|2','thm.extreal.lte_add')).
fof(f2440,plain,(
  ( ! [X202] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X202),s('type.extreal.extreal','const.extreal.NegInf_0')))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X202))))) ) | $spl98),
  inference(cnf_transformation,[],[f2440_D])).
fof(f2440_D,plain,(
  ( ! [X202] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X202),s('type.extreal.extreal','const.extreal.NegInf_0')))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X202))))) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f2458,plain,(
  ( ! [X214] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X214),s('type.extreal.extreal','const.extreal.NegInf_0')))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X214))))) ) | $spl0),
  inference(forward_demodulation,[],[f2279,f560])).
fof(f2279,plain,(
  ( ! [X214] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X214),s('type.extreal.extreal','const.extreal.NegInf_0')))))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X214))))) ) | $spl0),
  inference(superposition,[],[f557,f743])).
fof(f2457,plain,(
  ( ! [X213] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X213)))))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X213))))) ) | $spl114),
  inference(cnf_transformation,[],[f2457_D])).
fof(f2457_D,plain,(
  ( ! [X213] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X213)))))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X213))))) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f2438,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))))) | $spl97),
  inference(cnf_transformation,[],[f2438_D])).
fof(f2438_D,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))))) <=> ~$spl97),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl97])])).
fof(f2456,plain,(
  ( ! [X213] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X213)))))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X213))))) ) | $spl0),
  inference(forward_demodulation,[],[f2278,f560])).
fof(f2278,plain,(
  ( ! [X213] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X213)))))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X213)))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))))))) ) | $spl0),
  inference(superposition,[],[f557,f743])).
fof(f2455,plain,(
  ( ! [X212] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X212),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))))))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X212))))) ) | $spl112),
  inference(cnf_transformation,[],[f2455_D])).
fof(f2455_D,plain,(
  ( ! [X212] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X212),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))))))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X212))))) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f2454,plain,(
  ( ! [X212] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X212),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))))))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X212))))) ) | $spl0),
  inference(forward_demodulation,[],[f2277,f560])).
fof(f2277,plain,(
  ( ! [X212] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X212),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))))))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X212))))) ) | $spl0),
  inference(superposition,[],[f556,f743])).
fof(f2453,plain,(
  ( ! [X211] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X211)))))))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X211))))) ) | $spl110),
  inference(cnf_transformation,[],[f2453_D])).
fof(f2453_D,plain,(
  ( ! [X211] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X211)))))))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X211))))) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f2452,plain,(
  ( ! [X211] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X211)))))))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X211))))) ) | $spl0),
  inference(forward_demodulation,[],[f2276,f549])).
fof(f2276,plain,(
  ( ! [X211] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',X211)))))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X211))))) ) | $spl0),
  inference(superposition,[],[f556,f743])).
fof(f2451,plain,(
  ( ! [X210] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X210),s('type.extreal.extreal','const.extreal.NegInf_0')))))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X210))))) ) | $spl108),
  inference(cnf_transformation,[],[f2451_D])).
fof(f2451_D,plain,(
  ( ! [X210] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X210),s('type.extreal.extreal','const.extreal.NegInf_0')))))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X210))))) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f2450,plain,(
  ( ! [X210] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X210),s('type.extreal.extreal','const.extreal.NegInf_0')))))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X210))))) ) | $spl0),
  inference(forward_demodulation,[],[f2275,f560])).
fof(f2275,plain,(
  ( ! [X210] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X210),s('type.extreal.extreal','const.extreal.NegInf_0')))))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X210))))) ) | $spl0),
  inference(superposition,[],[f556,f743])).
fof(f2449,plain,(
  ( ! [X209] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X209)))))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X209))))) ) | $spl106),
  inference(cnf_transformation,[],[f2449_D])).
fof(f2449_D,plain,(
  ( ! [X209] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X209)))))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X209))))) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f2448,plain,(
  ( ! [X209] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X209)))))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X209))))) ) | $spl0),
  inference(forward_demodulation,[],[f2274,f560])).
fof(f2274,plain,(
  ( ! [X209] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X209)))))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X209)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))))))) ) | $spl0),
  inference(superposition,[],[f556,f743])).
fof(f2447,plain,(
  ( ! [X208] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X208)))))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X208))))) ) | $spl104),
  inference(cnf_transformation,[],[f2447_D])).
fof(f2447_D,plain,(
  ( ! [X208] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X208)))))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X208))))) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f2446,plain,(
  ( ! [X208] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X208)))))))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X208))))) ) | $spl0),
  inference(forward_demodulation,[],[f2273,f549])).
fof(f2273,plain,(
  ( ! [X208] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',X208)))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X208))))) ) | $spl0),
  inference(superposition,[],[f555,f743])).
fof(f2445,plain,(
  ( ! [X206] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X206),s('type.extreal.extreal','const.extreal.NegInf_0')))))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X206))))) ) | $spl102),
  inference(cnf_transformation,[],[f2445_D])).
fof(f2445_D,plain,(
  ( ! [X206] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X206),s('type.extreal.extreal','const.extreal.NegInf_0')))))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X206))))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f2444,plain,(
  ( ! [X206] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X206),s('type.extreal.extreal','const.extreal.NegInf_0')))))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X206))))) ) | $spl0),
  inference(forward_demodulation,[],[f2271,f560])).
fof(f2271,plain,(
  ( ! [X206] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X206),s('type.extreal.extreal','const.extreal.NegInf_0')))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X206))))) ) | $spl0),
  inference(superposition,[],[f555,f743])).
fof(f2436,plain,(
  ( ! [X201] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X201)))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X201))))) ) | $spl94),
  inference(cnf_transformation,[],[f2436_D])).
fof(f2436_D,plain,(
  ( ! [X201] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X201)))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X201))))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f2441,plain,(
  ( ! [X205] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X205)))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X205))))) ) | $spl0),
  inference(forward_demodulation,[],[f2270,f560])).
fof(f2270,plain,(
  ( ! [X205] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X205)))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X205)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))))))) ) | $spl0),
  inference(superposition,[],[f555,f743])).
fof(f2439,plain,(
  ( ! [X202] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X202),s('type.extreal.extreal','const.extreal.NegInf_0')))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X202))))) ) | $spl0),
  inference(forward_demodulation,[],[f2267,f560])).
fof(f2267,plain,(
  ( ! [X202] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X202),s('type.extreal.extreal','const.extreal.NegInf_0')))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X202))))) ) | $spl0),
  inference(superposition,[],[f554,f743])).
fof(f554,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X1)))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X0))))) )),
  inference(cnf_transformation,[],[f411])).
fof(f411,plain,(
  ! [X0,X1] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X0)))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X1)))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))))))),
  inference(flattening,[],[f410])).
fof(f410,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X0)))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X1))))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))))))),
  inference(ennf_transformation,[],[f290])).
fof(f290,plain,(
  ! [X0,X1] : ((p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X0)))) & p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X1))))) => p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))))))),
  inference(rectify,[],[f88])).
fof(f88,axiom,(
  ! [X6,X12] : ((p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X6)))) & p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X12))))) => p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X12)))))))),
  file('thm.extreal.add_infty|split|2','thm.extreal.lt_add')).
fof(f2435,plain,(
  ( ! [X201] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X201)))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X201))))) ) | $spl0),
  inference(forward_demodulation,[],[f2266,f560])).
fof(f2266,plain,(
  ( ! [X201] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X201)))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X201)))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))))))) ) | $spl0),
  inference(superposition,[],[f554,f743])).
fof(f2428,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))) | $spl88),
  inference(cnf_transformation,[],[f2428_D])).
fof(f2428_D,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f2255,plain,(
  ( ! [X190] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X190),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X190)) ) | $spl0),
  inference(superposition,[],[f550,f743])).
fof(f2254,plain,(
  ( ! [X189] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X189))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X189) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl0),
  inference(superposition,[],[f550,f743])).
fof(f1779,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl82),
  inference(cnf_transformation,[],[f1779_D])).
fof(f1779_D,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f2188,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl0),
  inference(superposition,[],[f558,f743])).
fof(f2386,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK11)))) | $spl0),
  inference(subsumption_resolution,[],[f2187,f574])).
fof(f2187,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK11)))) | $spl0),
  inference(superposition,[],[f557,f743])).
fof(f2356,plain,(
  ( ! [X36] : (s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11)))) | s('type.extreal.extreal',sK11) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X36)))) ) | ($spl0 | $spl5)),
  inference(subsumption_resolution,[],[f2355,f596])).
fof(f596,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.NegInf_0')),
  inference(cnf_transformation,[],[f330])).
fof(f330,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.NegInf_0')),
  inference(flattening,[],[f112])).
fof(f112,axiom,(
  ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.NegInf_0')),
  file('thm.extreal.add_infty|split|2','thm.extreal.extreal_distinct|split|0')).
fof(f2355,plain,(
  ( ! [X36] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.NegInf_0') | s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11)))) | s('type.extreal.extreal',sK11) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X36)))) ) | ($spl0 | $spl5)),
  inference(forward_demodulation,[],[f2354,f567])).
fof(f2354,plain,(
  ( ! [X36] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11)))) | s('type.extreal.extreal',sK11) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X36)))) ) | ($spl0 | $spl5)),
  inference(subsumption_resolution,[],[f2089,f792])).
fof(f2089,plain,(
  ( ! [X36] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK11) | s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11)))) | s('type.extreal.extreal',sK11) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X36)))) ) | $spl0),
  inference(superposition,[],[f743,f584])).
fof(f2053,plain,(
  ( ! [X152,X153] : (p(s(bool,X152)) | ~p(s(bool,X153)) | s(bool,'F_0') != s(bool,X153)) ) | $spl71),
  inference(resolution,[],[f1490,f737])).
fof(f1983,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,'F_0') != s(bool,X22) | s(bool,X21) != s(bool,X23)) ) | $spl71),
  inference(resolution,[],[f1490,f606])).
fof(f1981,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,'F_0') != s(bool,X16)) ) | $spl71),
  inference(resolution,[],[f1490,f605])).
fof(f1980,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,'F_0') != s(bool,X14) | s(bool,X12) != s(bool,X13)) ) | $spl71),
  inference(resolution,[],[f1490,f604])).
fof(f1979,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,'F_0') != s(bool,X10)) ) | $spl71),
  inference(resolution,[],[f1490,f604])).
fof(f1978,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,'F_0') != s(bool,X7) | s(bool,X6) != s(bool,X8)) ) | $spl71),
  inference(resolution,[],[f1490,f602])).
fof(f1977,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5) | s(bool,X3) != s(bool,X4)) ) | $spl71),
  inference(resolution,[],[f1490,f600])).
fof(f1976,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,'F_0') != s(bool,X1)) ) | $spl71),
  inference(resolution,[],[f1490,f600])).
fof(f1916,plain,(
  $false | $spl68),
  inference(resolution,[],[f1121,f777])).
fof(f777,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK11))))),
  inference(trivial_inequality_removal,[],[f776])).
fof(f776,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.PosInf_0') | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK11))))),
  inference(superposition,[],[f537,f729])).
fof(f1921,plain,(
  $false | $spl68),
  inference(resolution,[],[f1121,f574])).
fof(f1923,plain,(
  $false | $spl68),
  inference(resolution,[],[f1121,f725])).
fof(f808,plain,(
  s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK35(sK11)))) | $spl28),
  inference(cnf_transformation,[],[f808_D])).
fof(f808_D,plain,(
  s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK35(sK11)))) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f1792,plain,(
  ~p(s(bool,'F_0')) | s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK35(sK11))))),
  inference(forward_demodulation,[],[f1791,f597])).
fof(f1791,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK35(sK11))))),
  inference(subsumption_resolution,[],[f1721,f537])).
fof(f1721,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK35(sK11)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK11)),
  inference(superposition,[],[f777,f591])).
fof(f806,plain,(
  s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK34(sK11)))) | $spl26),
  inference(cnf_transformation,[],[f806_D])).
fof(f806_D,plain,(
  s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK34(sK11)))) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f1788,plain,(
  ~p(s(bool,'F_0')) | s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK34(sK11))))),
  inference(forward_demodulation,[],[f1787,f597])).
fof(f1787,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK34(sK11))))),
  inference(subsumption_resolution,[],[f1719,f537])).
fof(f1719,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK34(sK11)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK11)),
  inference(superposition,[],[f777,f590])).
fof(f1784,plain,(
  ( ! [X37] : (~p(s(bool,'F_0')) | s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11)))) | s('type.extreal.extreal',sK11) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X37)))) )),
  inference(forward_demodulation,[],[f1783,f597])).
fof(f1783,plain,(
  ( ! [X37] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11)))) | s('type.extreal.extreal',sK11) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X37)))) )),
  inference(subsumption_resolution,[],[f1717,f537])).
fof(f1717,plain,(
  ( ! [X37] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK11) | s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11)))) | s('type.extreal.extreal',sK11) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X37)))) )),
  inference(superposition,[],[f777,f584])).
fof(f1782,plain,(
  ( ! [X36] : (s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11)))) | s('type.extreal.extreal',sK11) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X36)))) ) | $spl5),
  inference(subsumption_resolution,[],[f1781,f792])).
fof(f1781,plain,(
  ( ! [X36] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK11) | s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11)))) | s('type.extreal.extreal',sK11) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X36)))) )),
  inference(subsumption_resolution,[],[f1716,f734])).
fof(f734,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.PosInf_0'))))) )),
  inference(cnf_transformation,[],[f381])).
fof(f381,plain,(
  ! [X0] : p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.PosInf_0'))))),
  inference(rectify,[],[f123])).
fof(f123,axiom,(
  ! [X6] : p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal','const.extreal.PosInf_0'))))),
  file('thm.extreal.add_infty|split|2','thm.extreal.le_infty|split|1')).
fof(f1716,plain,(
  ( ! [X36] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK11) | s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11)))) | s('type.extreal.extreal',sK11) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X36)))) )),
  inference(superposition,[],[f777,f584])).
fof(f1780,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl83),
  inference(cnf_transformation,[],[f1780_D])).
fof(f1780_D,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) <=> ~$spl83),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl83])])).
fof(f1778,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))))),
  inference(forward_demodulation,[],[f1715,f597])).
fof(f1715,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))))),
  inference(superposition,[],[f777,f580])).
fof(f1777,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))) | $spl80),
  inference(cnf_transformation,[],[f1777_D])).
fof(f1777_D,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f1775,plain,(
  ~p(s(bool,'F_0')) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))))),
  inference(forward_demodulation,[],[f1714,f597])).
fof(f1714,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))))),
  inference(superposition,[],[f777,f578])).
fof(f801,plain,(
  ( ! [X34] : (s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X34))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X34)))))) ) | $spl20),
  inference(cnf_transformation,[],[f801_D])).
fof(f801_D,plain,(
  ( ! [X34] : (s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X34))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X34)))))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f1774,plain,(
  ( ! [X35] : (~p(s(bool,'F_0')) | s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X35))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X35)))))) )),
  inference(forward_demodulation,[],[f1773,f597])).
fof(f1773,plain,(
  ( ! [X35] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X35))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X35)))))) )),
  inference(subsumption_resolution,[],[f1713,f537])).
fof(f1713,plain,(
  ( ! [X35] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X35))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X35))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK11)) )),
  inference(superposition,[],[f777,f553])).
fof(f800,plain,(
  ( ! [X32] : (s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X32))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X32),s('type.extreal.extreal',sK11)))))) ) | $spl18),
  inference(cnf_transformation,[],[f800_D])).
fof(f800_D,plain,(
  ( ! [X32] : (s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X32))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X32),s('type.extreal.extreal',sK11)))))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f1770,plain,(
  ( ! [X33] : (~p(s(bool,'F_0')) | s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X33))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X33),s('type.extreal.extreal',sK11)))))) )),
  inference(forward_demodulation,[],[f1769,f597])).
fof(f1769,plain,(
  ( ! [X33] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X33))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X33),s('type.extreal.extreal',sK11)))))) )),
  inference(subsumption_resolution,[],[f1711,f537])).
fof(f1711,plain,(
  ( ! [X33] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X33))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X33),s('type.extreal.extreal',sK11))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK11)) )),
  inference(superposition,[],[f777,f552])).
fof(f1766,plain,(
  ( ! [X31] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X31))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X31)) ) | $spl78),
  inference(cnf_transformation,[],[f1766_D])).
fof(f1766_D,plain,(
  ( ! [X31] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X31))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X31)) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f1765,plain,(
  ( ! [X31] : (~p(s(bool,'F_0')) | s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X31))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X31)) )),
  inference(forward_demodulation,[],[f1709,f597])).
fof(f1709,plain,(
  ( ! [X31] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X31))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X31)) )),
  inference(superposition,[],[f777,f551])).
fof(f551,plain,(
  ( ! [X0,X1] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X1) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0)) )),
  inference(cnf_transformation,[],[f405])).
fof(f405,plain,(
  ! [X0,X1] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X1) | s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))),
  inference(flattening,[],[f404])).
fof(f404,plain,(
  ! [X0,X1] : ((s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X1)) | s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))),
  inference(ennf_transformation,[],[f285])).
fof(f285,plain,(
  ! [X0,X1] : ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) & s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X1)) => s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))),
  inference(flattening,[],[f284])).
fof(f284,plain,(
  ! [X0,X1] : ((~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) & ~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X1)) => ~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X6,X12] : ((~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X6) & ~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X12)) => ~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X12))))),
  file('thm.extreal.add_infty|split|2','thm.extreal.add_not_infty|split|0')).
fof(f1764,plain,(
  ( ! [X30] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X30),s('type.extreal.extreal',sK11))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X30)) ) | $spl76),
  inference(cnf_transformation,[],[f1764_D])).
fof(f1764_D,plain,(
  ( ! [X30] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X30),s('type.extreal.extreal',sK11))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X30)) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f1763,plain,(
  ( ! [X30] : (~p(s(bool,'F_0')) | s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X30),s('type.extreal.extreal',sK11))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X30)) )),
  inference(forward_demodulation,[],[f1708,f597])).
fof(f1708,plain,(
  ( ! [X30] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X30),s('type.extreal.extreal',sK11))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X30)) )),
  inference(superposition,[],[f777,f551])).
fof(f799,plain,(
  ( ! [X24,X25] : (s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X24),s('type.extreal.extreal',X25))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X24),s('type.extreal.extreal',sK11))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X25),s('type.extreal.extreal',sK11)))))) ) | $spl16),
  inference(cnf_transformation,[],[f799_D])).
fof(f799_D,plain,(
  ( ! [X24,X25] : (s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X24),s('type.extreal.extreal',X25))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X24),s('type.extreal.extreal',sK11))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X25),s('type.extreal.extreal',sK11)))))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f1760,plain,(
  ( ! [X26,X27] : (~p(s(bool,'F_0')) | s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X26),s('type.extreal.extreal',X27))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X26),s('type.extreal.extreal',sK11))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X27),s('type.extreal.extreal',sK11)))))) )),
  inference(forward_demodulation,[],[f1759,f597])).
fof(f1759,plain,(
  ( ! [X26,X27] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X26),s('type.extreal.extreal',X27))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X26),s('type.extreal.extreal',sK11))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X27),s('type.extreal.extreal',sK11)))))) )),
  inference(subsumption_resolution,[],[f1705,f537])).
fof(f1705,plain,(
  ( ! [X26,X27] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X26),s('type.extreal.extreal',X27))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X26),s('type.extreal.extreal',sK11))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X27),s('type.extreal.extreal',sK11))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK11)) )),
  inference(superposition,[],[f777,f546])).
fof(f798,plain,(
  ( ! [X21,X20] : (s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X20),s('type.extreal.extreal',X21))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X20))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X21)))))) ) | $spl14),
  inference(cnf_transformation,[],[f798_D])).
fof(f798_D,plain,(
  ( ! [X21,X20] : (s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X20),s('type.extreal.extreal',X21))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X20))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X21)))))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f1756,plain,(
  ( ! [X23,X22] : (~p(s(bool,'F_0')) | s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X22),s('type.extreal.extreal',X23))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X22))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X23)))))) )),
  inference(forward_demodulation,[],[f1755,f597])).
fof(f1755,plain,(
  ( ! [X23,X22] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X22),s('type.extreal.extreal',X23))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X22))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X23)))))) )),
  inference(subsumption_resolution,[],[f1703,f537])).
fof(f1703,plain,(
  ( ! [X23,X22] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X22),s('type.extreal.extreal',X23))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X22))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X23))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK11)) )),
  inference(superposition,[],[f777,f545])).
fof(f797,plain,(
  ( ! [X17,X16] : (s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X16),s('type.extreal.extreal',X17))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X16),s('type.extreal.extreal',sK11))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X17),s('type.extreal.extreal',sK11)))))) ) | $spl12),
  inference(cnf_transformation,[],[f797_D])).
fof(f797_D,plain,(
  ( ! [X17,X16] : (s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X16),s('type.extreal.extreal',X17))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X16),s('type.extreal.extreal',sK11))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X17),s('type.extreal.extreal',sK11)))))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f1752,plain,(
  ( ! [X19,X18] : (~p(s(bool,'F_0')) | s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X18),s('type.extreal.extreal',X19))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X18),s('type.extreal.extreal',sK11))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X19),s('type.extreal.extreal',sK11)))))) )),
  inference(forward_demodulation,[],[f1751,f597])).
fof(f1751,plain,(
  ( ! [X19,X18] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X18),s('type.extreal.extreal',X19))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X18),s('type.extreal.extreal',sK11))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X19),s('type.extreal.extreal',sK11)))))) )),
  inference(subsumption_resolution,[],[f1701,f537])).
fof(f1701,plain,(
  ( ! [X19,X18] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X18),s('type.extreal.extreal',X19))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X18),s('type.extreal.extreal',sK11))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X19),s('type.extreal.extreal',sK11))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK11)) )),
  inference(superposition,[],[f777,f544])).
fof(f796,plain,(
  ( ! [X12,X13] : (s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X12),s('type.extreal.extreal',X13))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X12))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X13)))))) ) | $spl10),
  inference(cnf_transformation,[],[f796_D])).
fof(f796_D,plain,(
  ( ! [X12,X13] : (s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X12),s('type.extreal.extreal',X13))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X12))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X13)))))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f1748,plain,(
  ( ! [X14,X15] : (~p(s(bool,'F_0')) | s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X14),s('type.extreal.extreal',X15))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X14))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X15)))))) )),
  inference(forward_demodulation,[],[f1747,f597])).
fof(f1747,plain,(
  ( ! [X14,X15] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X14),s('type.extreal.extreal',X15))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X14))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X15)))))) )),
  inference(subsumption_resolution,[],[f1699,f537])).
fof(f1699,plain,(
  ( ! [X14,X15] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X14),s('type.extreal.extreal',X15))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X14))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X15))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK11)) )),
  inference(superposition,[],[f777,f543])).
fof(f795,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X6))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X7),s('type.extreal.extreal',X8)))))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X7)))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X8))))) ) | $spl8),
  inference(cnf_transformation,[],[f795_D])).
fof(f795_D,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X6))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X7),s('type.extreal.extreal',X8)))))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X7)))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X8))))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f1744,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,'F_0')) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X9))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X10),s('type.extreal.extreal',X11)))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X9),s('type.extreal.extreal',X11)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X10))))) )),
  inference(forward_demodulation,[],[f1743,f597])).
fof(f1743,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X9))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X10),s('type.extreal.extreal',X11)))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X9),s('type.extreal.extreal',X11)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X10))))) )),
  inference(subsumption_resolution,[],[f1697,f537])).
fof(f1697,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X9))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X10),s('type.extreal.extreal',X11)))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X9),s('type.extreal.extreal',X11)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X10)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK11)) )),
  inference(superposition,[],[f777,f540])).
fof(f794,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X2)))))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',sK11)))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))))) ) | $spl6),
  inference(cnf_transformation,[],[f794_D])).
fof(f794_D,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X2)))))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',sK11)))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f1740,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'F_0')) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal',X4))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X5)))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X4),s('type.extreal.extreal',X5)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal',sK11))))) )),
  inference(forward_demodulation,[],[f1739,f597])).
fof(f1739,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal',X4))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X5)))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X4),s('type.extreal.extreal',X5)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal',sK11))))) )),
  inference(subsumption_resolution,[],[f1695,f537])).
fof(f1695,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal',X4))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X5)))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X4),s('type.extreal.extreal',X5)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal',sK11)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK11)) )),
  inference(superposition,[],[f777,f539])).
fof(f1734,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK11)))) ) | $spl72),
  inference(cnf_transformation,[],[f1734_D])).
fof(f1734_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK11)))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f1661,plain,(
  ( ! [X152,X153] : (p(s(bool,X152)) | ~p(s(bool,X153)) | s(bool,X153) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK11)))) )),
  inference(resolution,[],[f777,f737])).
fof(f1591,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK11))) | s(bool,X21) != s(bool,X23)) )),
  inference(resolution,[],[f777,f606])).
fof(f1589,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK11)))) )),
  inference(resolution,[],[f777,f605])).
fof(f1735,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK11)))) ) | $spl74),
  inference(cnf_transformation,[],[f1735_D])).
fof(f1735_D,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK11)))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f1588,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK11))) | s(bool,X12) != s(bool,X13)) )),
  inference(resolution,[],[f777,f604])).
fof(f1587,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK11)))) )),
  inference(resolution,[],[f777,f604])).
fof(f1586,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK11))) | s(bool,X6) != s(bool,X8)) )),
  inference(resolution,[],[f777,f602])).
fof(f1585,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK11))) | s(bool,X3) != s(bool,X4)) )),
  inference(resolution,[],[f777,f600])).
fof(f1584,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK11)))) )),
  inference(resolution,[],[f777,f600])).
fof(f1582,plain,(
  $false | $spl60),
  inference(trivial_inequality_removal,[],[f1561])).
fof(f1561,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl60),
  inference(superposition,[],[f1117,f592])).
fof(f592,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.NegInf_0')))) )),
  inference(cnf_transformation,[],[f326])).
fof(f326,plain,(
  ! [X0] : s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.NegInf_0')))),
  inference(rectify,[],[f175])).
fof(f175,axiom,(
  ! [X21] : s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X21))),s('type.extreal.extreal','const.extreal.NegInf_0')))),
  file('thm.extreal.add_infty|split|2','thm.extreal.extreal_le_def|split|5')).
fof(f1581,plain,(
  $false | $spl60),
  inference(trivial_inequality_removal,[],[f1562])).
fof(f1562,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl60),
  inference(superposition,[],[f1117,f597])).
fof(f1580,plain,(
  $false | $spl60),
  inference(trivial_inequality_removal,[],[f1566])).
fof(f1566,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl60),
  inference(superposition,[],[f1117,f731])).
fof(f731,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))))) )),
  inference(cnf_transformation,[],[f378])).
fof(f378,plain,(
  ! [X0] : s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))))),
  inference(rectify,[],[f186])).
fof(f186,axiom,(
  ! [X22] : s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X22)))))),
  file('thm.extreal.add_infty|split|2','thm.extreal.extreal_le_def|split|6')).
fof(f1579,plain,(
  $false | $spl60),
  inference(equality_resolution,[],[f1117])).
fof(f1544,plain,(
  ( ! [X37] : (s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11)))) | s('type.extreal.extreal',sK11) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X37)))) ) | $spl5),
  inference(subsumption_resolution,[],[f1523,f537])).
fof(f1523,plain,(
  ( ! [X37] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK11) | s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11)))) | s('type.extreal.extreal',sK11) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X37)))) ) | $spl5),
  inference(trivial_inequality_removal,[],[f1514])).
fof(f1514,plain,(
  ( ! [X37] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.NegInf_0') | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK11) | s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11)))) | s('type.extreal.extreal',sK11) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X37)))) ) | $spl5),
  inference(superposition,[],[f792,f584])).
fof(f1488,plain,(
  $false | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f1487,f570])).
fof(f1487,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | ($spl1 | $spl4)),
  inference(forward_demodulation,[],[f744,f793])).
fof(f1486,plain,(
  $false | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f1485,f574])).
fof(f1485,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | ($spl1 | $spl4)),
  inference(forward_demodulation,[],[f959,f793])).
fof(f1484,plain,(
  $false | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f1483,f736])).
fof(f1483,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | ($spl1 | $spl4)),
  inference(forward_demodulation,[],[f960,f793])).
fof(f1482,plain,(
  $false | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f1481,f570])).
fof(f1481,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | ($spl1 | $spl4)),
  inference(forward_demodulation,[],[f884,f793])).
fof(f884,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) | $spl1),
  inference(superposition,[],[f744,f560])).
fof(f1480,plain,(
  $false | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f1479,f570])).
fof(f1479,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | ($spl1 | $spl4)),
  inference(forward_demodulation,[],[f885,f793])).
fof(f885,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) | $spl1),
  inference(superposition,[],[f744,f560])).
fof(f1478,plain,(
  $false | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f1477,f577])).
fof(f577,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))) )),
  inference(cnf_transformation,[],[f314])).
fof(f314,plain,(
  ! [X0] : s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))),
  inference(flattening,[],[f313])).
fof(f313,plain,(
  ! [X0] : ~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))),
  inference(rectify,[],[f124])).
fof(f124,axiom,(
  ! [X6] : ~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X6)))),
  file('thm.extreal.add_infty|split|2','thm.extreal.extreal_not_infty|split|0')).
fof(f1477,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK34(sK11)))) | ($spl1 | $spl4)),
  inference(forward_demodulation,[],[f966,f793])).
fof(f966,plain,(
  s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK34(sK11)))) | $spl1),
  inference(subsumption_resolution,[],[f965,f537])).
fof(f965,plain,(
  s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK34(sK11)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK11) | $spl1),
  inference(subsumption_resolution,[],[f834,f570])).
fof(f834,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK34(sK11)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK11) | $spl1),
  inference(superposition,[],[f744,f590])).
fof(f1476,plain,(
  $false | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f1475,f577])).
fof(f1475,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK35(sK11)))) | ($spl1 | $spl4)),
  inference(forward_demodulation,[],[f969,f793])).
fof(f969,plain,(
  s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK35(sK11)))) | $spl1),
  inference(subsumption_resolution,[],[f968,f537])).
fof(f968,plain,(
  s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK35(sK11)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK11) | $spl1),
  inference(subsumption_resolution,[],[f836,f570])).
fof(f836,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK35(sK11)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK11) | $spl1),
  inference(superposition,[],[f744,f591])).
fof(f1471,plain,(
  $false | ($spl3 | $spl4)),
  inference(subsumption_resolution,[],[f1470,f570])).
fof(f1470,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | ($spl3 | $spl4)),
  inference(forward_demodulation,[],[f746,f793])).
fof(f1424,plain,(
  ( ! [X152,X153] : (p(s(bool,X152)) | ~p(s(bool,X153)) | s(bool,'F_0') != s(bool,X153)) ) | $spl4),
  inference(resolution,[],[f1281,f737])).
fof(f1281,plain,(
  ~p(s(bool,'F_0')) | $spl4),
  inference(forward_demodulation,[],[f1280,f597])).
fof(f1280,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl4),
  inference(forward_demodulation,[],[f1279,f793])).
fof(f1279,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK11)))) | $spl4),
  inference(subsumption_resolution,[],[f1137,f596])).
fof(f1137,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.NegInf_0') | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK11)))) | $spl4),
  inference(superposition,[],[f793,f729])).
fof(f1354,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,'F_0') != s(bool,X22) | s(bool,X21) != s(bool,X23)) ) | $spl4),
  inference(resolution,[],[f1281,f606])).
fof(f1352,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,'F_0') != s(bool,X16)) ) | $spl4),
  inference(resolution,[],[f1281,f605])).
fof(f1351,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,'F_0') != s(bool,X14) | s(bool,X12) != s(bool,X13)) ) | $spl4),
  inference(resolution,[],[f1281,f604])).
fof(f1350,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,'F_0') != s(bool,X10)) ) | $spl4),
  inference(resolution,[],[f1281,f604])).
fof(f1349,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,'F_0') != s(bool,X7) | s(bool,X6) != s(bool,X8)) ) | $spl4),
  inference(resolution,[],[f1281,f602])).
fof(f1348,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5) | s(bool,X3) != s(bool,X4)) ) | $spl4),
  inference(resolution,[],[f1281,f600])).
fof(f1347,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,'F_0') != s(bool,X1)) ) | $spl4),
  inference(resolution,[],[f1281,f600])).
fof(f1102,plain,(
  ( ! [X152,X153] : (p(s(bool,X152)) | ~p(s(bool,X153)) | s(bool,'F_0') != s(bool,X153)) ) | $spl1),
  inference(resolution,[],[f975,f737])).
fof(f1032,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,'F_0') != s(bool,X22) | s(bool,X21) != s(bool,X23)) ) | $spl1),
  inference(resolution,[],[f975,f606])).
fof(f1030,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,'F_0') != s(bool,X16)) ) | $spl1),
  inference(resolution,[],[f975,f605])).
fof(f1029,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,'F_0') != s(bool,X14) | s(bool,X12) != s(bool,X13)) ) | $spl1),
  inference(resolution,[],[f975,f604])).
fof(f1028,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,'F_0') != s(bool,X10)) ) | $spl1),
  inference(resolution,[],[f975,f604])).
fof(f1027,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,'F_0') != s(bool,X7) | s(bool,X6) != s(bool,X8)) ) | $spl1),
  inference(resolution,[],[f975,f602])).
fof(f1026,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5) | s(bool,X3) != s(bool,X4)) ) | $spl1),
  inference(resolution,[],[f975,f600])).
fof(f1025,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,'F_0') != s(bool,X1)) ) | $spl1),
  inference(resolution,[],[f975,f600])).
fof(f1022,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) | s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK35('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))))) | $spl1),
  inference(forward_demodulation,[],[f1021,f560])).
fof(f1021,plain,(
  s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK35('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))) | $spl1),
  inference(forward_demodulation,[],[f916,f560])).
fof(f916,plain,(
  s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK35('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))) | $spl1),
  inference(trivial_inequality_removal,[],[f913])).
fof(f913,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.NegInf_0') | s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK35('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))) | $spl1),
  inference(superposition,[],[f744,f591])).
fof(f1018,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) | s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK34('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))))) | $spl1),
  inference(forward_demodulation,[],[f1017,f560])).
fof(f1017,plain,(
  s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK34('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))) | $spl1),
  inference(forward_demodulation,[],[f917,f560])).
fof(f917,plain,(
  s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK34('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))) | $spl1),
  inference(trivial_inequality_removal,[],[f911])).
fof(f911,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.NegInf_0') | s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK34('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))) | $spl1),
  inference(superposition,[],[f744,f590])).
fof(f1013,plain,(
  ( ! [X37] : (s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X37))) != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) | s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))) ) | $spl1),
  inference(forward_demodulation,[],[f1012,f560])).
fof(f1012,plain,(
  ( ! [X37] : (s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X37))) != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(forward_demodulation,[],[f1011,f560])).
fof(f1011,plain,(
  ( ! [X37] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) | s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X37))) != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(forward_demodulation,[],[f918,f560])).
fof(f918,plain,(
  ( ! [X37] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X37))) != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(trivial_inequality_removal,[],[f909])).
fof(f909,plain,(
  ( ! [X37] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.NegInf_0') | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X37))) != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(superposition,[],[f744,f584])).
fof(f1007,plain,(
  ( ! [X35] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) | s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X35))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X35)))))))) ) | $spl1),
  inference(forward_demodulation,[],[f1006,f560])).
fof(f1006,plain,(
  ( ! [X35] : (s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X35))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X35))))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(forward_demodulation,[],[f1005,f560])).
fof(f1005,plain,(
  ( ! [X35] : (s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X35))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X35))))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(forward_demodulation,[],[f921,f549])).
fof(f921,plain,(
  ( ! [X35] : (s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X35))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',X35))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(trivial_inequality_removal,[],[f905])).
fof(f905,plain,(
  ( ! [X35] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.NegInf_0') | s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X35))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',X35))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(superposition,[],[f744,f553])).
fof(f1003,plain,(
  ( ! [X33] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) | s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X33))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X33),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))))))) ) | $spl1),
  inference(forward_demodulation,[],[f1002,f560])).
fof(f1002,plain,(
  ( ! [X33] : (s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X33))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X33),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(forward_demodulation,[],[f922,f560])).
fof(f922,plain,(
  ( ! [X33] : (s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X33))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X33),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(trivial_inequality_removal,[],[f903])).
fof(f903,plain,(
  ( ! [X33] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.NegInf_0') | s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X33))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X33),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(superposition,[],[f744,f552])).
fof(f998,plain,(
  ( ! [X26,X27] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) | s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X26),s('type.extreal.extreal',X27))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X26),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X27),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))))))) ) | $spl1),
  inference(forward_demodulation,[],[f997,f560])).
fof(f997,plain,(
  ( ! [X26,X27] : (s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X26),s('type.extreal.extreal',X27))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X26),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X27),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(forward_demodulation,[],[f925,f560])).
fof(f925,plain,(
  ( ! [X26,X27] : (s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X26),s('type.extreal.extreal',X27))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X26),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X27),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(trivial_inequality_removal,[],[f897])).
fof(f897,plain,(
  ( ! [X26,X27] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.NegInf_0') | s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X26),s('type.extreal.extreal',X27))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X26),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X27),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(superposition,[],[f744,f546])).
fof(f995,plain,(
  ( ! [X23,X22] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) | s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X22),s('type.extreal.extreal',X23))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X22))))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X23)))))))) ) | $spl1),
  inference(forward_demodulation,[],[f994,f560])).
fof(f994,plain,(
  ( ! [X23,X22] : (s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X22),s('type.extreal.extreal',X23))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X22))))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X23))))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(forward_demodulation,[],[f993,f549])).
fof(f993,plain,(
  ( ! [X23,X22] : (s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X22),s('type.extreal.extreal',X23))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',X22))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X23))))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(forward_demodulation,[],[f926,f549])).
fof(f926,plain,(
  ( ! [X23,X22] : (s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X22),s('type.extreal.extreal',X23))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',X22))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',X23))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(trivial_inequality_removal,[],[f895])).
fof(f895,plain,(
  ( ! [X23,X22] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.NegInf_0') | s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X22),s('type.extreal.extreal',X23))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',X22))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',X23))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(superposition,[],[f744,f545])).
fof(f991,plain,(
  ( ! [X19,X18] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) | s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X18),s('type.extreal.extreal',X19))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X18),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X19),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))))))) ) | $spl1),
  inference(forward_demodulation,[],[f990,f560])).
fof(f990,plain,(
  ( ! [X19,X18] : (s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X18),s('type.extreal.extreal',X19))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X18),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X19),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(forward_demodulation,[],[f927,f560])).
fof(f927,plain,(
  ( ! [X19,X18] : (s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X18),s('type.extreal.extreal',X19))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X18),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X19),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(trivial_inequality_removal,[],[f893])).
fof(f893,plain,(
  ( ! [X19,X18] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.NegInf_0') | s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X18),s('type.extreal.extreal',X19))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X18),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X19),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(superposition,[],[f744,f544])).
fof(f988,plain,(
  ( ! [X14,X15] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) | s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X14),s('type.extreal.extreal',X15))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X14))))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X15)))))))) ) | $spl1),
  inference(forward_demodulation,[],[f987,f560])).
fof(f987,plain,(
  ( ! [X14,X15] : (s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X14),s('type.extreal.extreal',X15))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X14))))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X15))))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(forward_demodulation,[],[f986,f549])).
fof(f986,plain,(
  ( ! [X14,X15] : (s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X14),s('type.extreal.extreal',X15))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',X14))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X15))))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(forward_demodulation,[],[f928,f549])).
fof(f928,plain,(
  ( ! [X14,X15] : (s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X14),s('type.extreal.extreal',X15))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',X14))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',X15))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(trivial_inequality_removal,[],[f891])).
fof(f891,plain,(
  ( ! [X14,X15] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.NegInf_0') | s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X14),s('type.extreal.extreal',X15))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',X14))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',X15))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(superposition,[],[f744,f543])).
fof(f984,plain,(
  ( ! [X10,X11,X9] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))),s('type.extreal.extreal',X10)))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X9))))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X10),s('type.extreal.extreal',X11)))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X9),s('type.extreal.extreal',X11))))) ) | $spl1),
  inference(forward_demodulation,[],[f983,f560])).
fof(f983,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))),s('type.extreal.extreal',X10)))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X9))))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X10),s('type.extreal.extreal',X11)))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X9),s('type.extreal.extreal',X11)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(forward_demodulation,[],[f982,f560])).
fof(f982,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X9))))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X10),s('type.extreal.extreal',X11)))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X9),s('type.extreal.extreal',X11)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',X10)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(forward_demodulation,[],[f929,f549])).
fof(f929,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',X9))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X10),s('type.extreal.extreal',X11)))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X9),s('type.extreal.extreal',X11)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',X10)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(trivial_inequality_removal,[],[f889])).
fof(f889,plain,(
  ( ! [X10,X11,X9] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.NegInf_0') | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',X9))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X10),s('type.extreal.extreal',X11)))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X9),s('type.extreal.extreal',X11)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',X10)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(superposition,[],[f744,f540])).
fof(f978,plain,(
  ( ! [X4,X5,X3] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal',X4))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X5)))))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X4),s('type.extreal.extreal',X5))))) ) | $spl1),
  inference(forward_demodulation,[],[f977,f560])).
fof(f977,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11)))))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal',X4))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X5)))))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X4),s('type.extreal.extreal',X5)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(forward_demodulation,[],[f976,f560])).
fof(f976,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal',X4))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X5)))))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X4),s('type.extreal.extreal',X5)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(forward_demodulation,[],[f930,f549])).
fof(f930,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal',X4))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',X5)))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X4),s('type.extreal.extreal',X5)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(trivial_inequality_removal,[],[f887])).
fof(f887,plain,(
  ( ! [X4,X5,X3] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.NegInf_0') | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal',X4))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',X5)))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X4),s('type.extreal.extreal',X5)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl1),
  inference(superposition,[],[f744,f539])).
fof(f963,plain,(
  ( ! [X37] : (s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11)))) | s('type.extreal.extreal',sK11) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X37)))) ) | $spl1),
  inference(subsumption_resolution,[],[f962,f537])).
fof(f962,plain,(
  ( ! [X37] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK11) | s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11)))) | s('type.extreal.extreal',sK11) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X37)))) ) | $spl1),
  inference(subsumption_resolution,[],[f832,f570])).
fof(f832,plain,(
  ( ! [X37] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK11) | s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11)))) | s('type.extreal.extreal',sK11) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X37)))) ) | $spl1),
  inference(superposition,[],[f744,f584])).
fof(f793,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK11) | $spl4),
  inference(cnf_transformation,[],[f793_D])).
fof(f793_D,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK11) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f779,plain,(
  s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK35(sK11)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK11)),
  inference(trivial_inequality_removal,[],[f773])).
fof(f773,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.PosInf_0') | s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK35(sK11)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK11)),
  inference(superposition,[],[f537,f591])).
fof(f780,plain,(
  s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK34(sK11)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK11)),
  inference(trivial_inequality_removal,[],[f771])).
fof(f771,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.PosInf_0') | s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK34(sK11)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK11)),
  inference(superposition,[],[f537,f590])).
fof(f781,plain,(
  ( ! [X36] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK11) | s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11)))) | s('type.extreal.extreal',sK11) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X36)))) )),
  inference(trivial_inequality_removal,[],[f769])).
fof(f769,plain,(
  ( ! [X36] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.PosInf_0') | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK11) | s('type.extreal.extreal',sK11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK13(sK11)))) | s('type.extreal.extreal',sK11) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X36)))) )),
  inference(superposition,[],[f537,f584])).
fof(f782,plain,(
  ( ! [X34] : (s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X34))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X34))))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK11)) )),
  inference(trivial_inequality_removal,[],[f765])).
fof(f765,plain,(
  ( ! [X34] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.PosInf_0') | s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X34))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X34))))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK11)) )),
  inference(superposition,[],[f537,f553])).
fof(f783,plain,(
  ( ! [X32] : (s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X32))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X32),s('type.extreal.extreal',sK11))))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK11)) )),
  inference(trivial_inequality_removal,[],[f763])).
fof(f763,plain,(
  ( ! [X32] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.PosInf_0') | s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X32))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X32),s('type.extreal.extreal',sK11))))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK11)) )),
  inference(superposition,[],[f537,f552])).
fof(f786,plain,(
  ( ! [X24,X25] : (s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X24),s('type.extreal.extreal',X25))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X24),s('type.extreal.extreal',sK11))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X25),s('type.extreal.extreal',sK11))))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK11)) )),
  inference(trivial_inequality_removal,[],[f757])).
fof(f757,plain,(
  ( ! [X24,X25] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.PosInf_0') | s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X24),s('type.extreal.extreal',X25))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X24),s('type.extreal.extreal',sK11))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X25),s('type.extreal.extreal',sK11))))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK11)) )),
  inference(superposition,[],[f537,f546])).
fof(f787,plain,(
  ( ! [X21,X20] : (s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X20),s('type.extreal.extreal',X21))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X20))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X21))))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK11)) )),
  inference(trivial_inequality_removal,[],[f755])).
fof(f755,plain,(
  ( ! [X21,X20] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.PosInf_0') | s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X20),s('type.extreal.extreal',X21))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X20))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X21))))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK11)) )),
  inference(superposition,[],[f537,f545])).
fof(f788,plain,(
  ( ! [X17,X16] : (s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X16),s('type.extreal.extreal',X17))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X16),s('type.extreal.extreal',sK11))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X17),s('type.extreal.extreal',sK11))))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK11)) )),
  inference(trivial_inequality_removal,[],[f753])).
fof(f753,plain,(
  ( ! [X17,X16] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.PosInf_0') | s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X16),s('type.extreal.extreal',X17))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X16),s('type.extreal.extreal',sK11))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X17),s('type.extreal.extreal',sK11))))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK11)) )),
  inference(superposition,[],[f537,f544])).
fof(f789,plain,(
  ( ! [X12,X13] : (s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X12),s('type.extreal.extreal',X13))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X12))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X13))))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK11)) )),
  inference(trivial_inequality_removal,[],[f751])).
fof(f751,plain,(
  ( ! [X12,X13] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.PosInf_0') | s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X12),s('type.extreal.extreal',X13))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X12))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X13))))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK11)) )),
  inference(superposition,[],[f537,f543])).
fof(f790,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X6))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X7),s('type.extreal.extreal',X8)))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X8)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X7)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK11)) )),
  inference(trivial_inequality_removal,[],[f749])).
fof(f749,plain,(
  ( ! [X6,X8,X7] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.PosInf_0') | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X6))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X7),s('type.extreal.extreal',X8)))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X8)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X7)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK11)) )),
  inference(superposition,[],[f537,f540])).
fof(f791,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X2)))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',sK11)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK11)) )),
  inference(trivial_inequality_removal,[],[f747])).
fof(f747,plain,(
  ( ! [X2,X0,X1] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.PosInf_0') | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal',X2)))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',sK11)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK11)) )),
  inference(superposition,[],[f537,f539])).
fof(f744,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))) | $spl1),
  inference(cnf_transformation,[],[f744_D])).
fof(f744_D,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0'))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f746,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) | $spl3),
  inference(cnf_transformation,[],[f746_D])).
fof(f746_D,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f538,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK11))) | s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK11),s('type.extreal.extreal','const.extreal.NegInf_0')))),
  inference(cnf_transformation,[],[f463])).
% SZS output end Proof for thm.extreal
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 227
% Passive clauses: 1934
% Generated clauses: 5365
% Final active clauses: 187
% Final passive clauses: 878
% Input formulas: 261
% Initial clauses: 200
% 
% Selected by SInE selection: 89
% SInE iterations: 1
% 
% Duplicate literals: 44
% Trivial inequalities: 84
% Fw subsumption resolutions: 784
% Fw demodulations: 652
% 
% Simple tautologies: 86
% Equational tautologies: 1
% Forward subsumptions: 1419
% 
% Binary resolution: 1464
% Forward superposition: 1178
% Backward superposition: 775
% Equality resolution: 1
% 
% Split clauses: 223
% Split components: 465
% SAT solver clauses: 244
% SAT solver unit clauses: 7
% SAT solver binary clauses: 52
% 
% Sat splits: 223
% Sat splitting refutations: 20
% 
% Memory used [KB]: 3709
% Time elapsed: 0.095 s
% ------------------------------
% ------------------------------
% Success in time 11.837 s
