Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.prelim
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.prelim
% SZS output start Proof for thm.prelim
fof(f23257,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f23159,f1318,f1493,f14540,f22435,f363,f370,f9380,f9377,f237,f366,f6006,f6051,f6049,f6318])).
fof(f6318,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f6311,f231])).
fof(f231,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.LESS_0')))),
  inference(cnf_transformation,[],[f71])).
fof(f71,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.LESS_0')))),
  inference(flattening,[],[f6])).
fof(f6,negated_conjecture,(
  ~s('type.num.num','const.num.|48|_0') = s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.LESS_0')))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.LESS_0')))),
  file('thm.prelim.ordering2num_thm|split|0','thm.prelim.ordering2num_thm|split|0')).
fof(f6311,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.LESS_0'))) | $spl52),
  inference(backward_demodulation,[],[f6275,f6051])).
fof(f6275,plain,(
  s('type.prelim.ordering','const.prelim.LESS_0') = s('type.prelim.ordering',sK2('const.num.|48|_0')) | $spl52),
  inference(forward_demodulation,[],[f6203,f234])).
fof(f234,plain,(
  s('type.prelim.ordering','const.prelim.LESS_0') = s('type.prelim.ordering','const.prelim.num2ordering_1'(s('type.num.num','const.num.|48|_0')))),
  inference(cnf_transformation,[],[f10])).
fof(f10,axiom,(
  s('type.prelim.ordering','const.prelim.LESS_0') = s('type.prelim.ordering','const.prelim.num2ordering_1'(s('type.num.num','const.num.|48|_0')))),
  file('thm.prelim.ordering2num_thm|split|0','thm.prelim.num2ordering_thm|split|0')).
fof(f6203,plain,(
  s('type.prelim.ordering',sK2('const.num.|48|_0')) = s('type.prelim.ordering','const.prelim.num2ordering_1'(s('type.num.num','const.num.|48|_0'))) | $spl52),
  inference(superposition,[],[f241,f6051])).
fof(f241,plain,(
  ( ! [X0] : (s('type.prelim.ordering',X0) = s('type.prelim.ordering','const.prelim.num2ordering_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X0)))))) )),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0] : s('type.prelim.ordering',X0) = s('type.prelim.ordering','const.prelim.num2ordering_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X0)))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X7] : s('type.prelim.ordering',X7) = s('type.prelim.ordering','const.prelim.num2ordering_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X7)))))),
  file('thm.prelim.ordering2num_thm|split|0','thm.prelim.ordering_BIJ|split|0')).
fof(f6049,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | $spl50),
  inference(cnf_transformation,[],[f6049_D])).
fof(f6049_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f6051,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',sK2('const.num.|48|_0')))) | $spl52),
  inference(cnf_transformation,[],[f6051_D])).
fof(f6051_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',sK2('const.num.|48|_0')))) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f6006,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',sK2('const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(resolution,[],[f268,f300])).
fof(f300,plain,(
  ( ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f205])).
fof(f205,plain,(
  ! [X0] : ((s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))) & (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0)))),
  inference(nnf_transformation,[],[f110])).
fof(f110,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) <=> p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))))),
  inference(flattening,[],[f109])).
fof(f109,plain,(
  ! [X0] : (~s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) <=> p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))))),
  inference(rectify,[],[f63])).
fof(f63,axiom,(
  ! [X26] : (~s('type.num.num','const.num.|48|_0') = s('type.num.num',X26) <=> p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X26)))))),
  file('thm.prelim.ordering2num_thm|split|0','thm.arithmetic.NOT_ZERO_LT_ZERO')).
fof(f268,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.num.num',X0) = s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',sK2(X0))))) )),
  inference(cnf_transformation,[],[f184])).
fof(f184,plain,(
  ! [X0] : ((~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.num.num',X0) = s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',sK2(X0))))) & (! [X1] : s('type.num.num',X0) != s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X1))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2])],[f183])).
fof(f183,plain,(
  ! [X0] : ((~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | ? [X2] : s('type.num.num',X0) = s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X2)))) & (! [X1] : s('type.num.num',X0) != s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X1))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))))),
  inference(rectify,[],[f182])).
fof(f182,plain,(
  ! [X0] : ((~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | ? [X1] : s('type.num.num',X0) = s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X1)))) & (! [X1] : s('type.num.num',X0) != s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X1))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))))),
  inference(nnf_transformation,[],[f97])).
fof(f97,plain,(
  ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) <=> ? [X1] : s('type.num.num',X0) = s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X1))))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X6] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) <=> ? [X7] : s('type.num.num',X6) = s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X7))))),
  file('thm.prelim.ordering2num_thm|split|0','thm.prelim.ordering2num_ONTO')).
fof(f366,plain,(
  ~p(s(bool,'F_0')) | $spl1),
  inference(cnf_transformation,[],[f366_D])).
fof(f366_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f237,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.prelim.ordering2num_thm|split|0',aHL_FALSITY)).
fof(f9377,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK7),s('type.num.num','const.num.|48|_0'))) | $spl126),
  inference(cnf_transformation,[],[f9377_D])).
fof(f9377_D,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK7),s('type.num.num','const.num.|48|_0'))) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f9380,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK7),s('type.num.num','const.num.|48|_0'))) | ($spl1 | $spl50)),
  inference(subsumption_resolution,[],[f9379,f366])).
fof(f9379,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK7),s('type.num.num','const.num.|48|_0'))) | ($spl1 | $spl50)),
  inference(forward_demodulation,[],[f9369,f540])).
fof(f540,plain,(
  s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))),
  inference(resolution,[],[f531,f261])).
fof(f261,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f178])).
fof(f178,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f92])).
fof(f92,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f91])).
fof(f91,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X25] : (s(bool,'F_0') = s(bool,X25) <=> ~p(s(bool,X25)))),
  file('thm.prelim.ordering2num_thm|split|0','thm.bool.EQ_CLAUSES|split|2')).
fof(f531,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))),
  inference(equality_resolution,[],[f301])).
fof(f301,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f205])).
fof(f9369,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK7),s('type.num.num','const.num.|48|_0'))) | ($spl1 | $spl50)),
  inference(superposition,[],[f9253,f650])).
fof(f650,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X4)) ) | $spl1),
  inference(subsumption_resolution,[],[f643,f366])).
fof(f643,plain,(
  ( ! [X4] : (p(s(bool,'F_0')) | s(bool,X4) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f462,f540])).
fof(f462,plain,(
  ( ! [X2,X1] : (p(s(bool,X1)) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(resolution,[],[f317,f257])).
fof(f257,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f176])).
fof(f176,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f89])).
fof(f89,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X25] : (s(bool,'T_0') = s(bool,X25) <=> p(s(bool,X25)))),
  file('thm.prelim.ordering2num_thm|split|0','thm.bool.EQ_CLAUSES|split|0')).
fof(f317,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f162])).
fof(f162,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f161])).
fof(f161,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f124])).
fof(f124,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X8,X9] : ((p(s(bool,X8)) => p(s(bool,X9))) => ((p(s(bool,X9)) => p(s(bool,X8))) => s(bool,X8) = s(bool,X9)))),
  file('thm.prelim.ordering2num_thm|split|0','thm.bool.IMP_ANTISYM_AX')).
fof(f9253,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK7),s('type.num.num','const.num.|48|_0')))) | $spl50),
  inference(backward_demodulation,[],[f6049,f360])).
fof(f360,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK7),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),
  inference(cnf_transformation,[],[f229])).
fof(f229,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK7),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK7])],[f147])).
fof(f147,plain,(
  ? [X0] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ? [X26] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X26),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),
  file('thm.prelim.ordering2num_thm|split|0','thm.prelim.prelim_|48|_prove_ltb')).
fof(f370,plain,(
  p(s(bool,'T_0')) | $spl6),
  inference(cnf_transformation,[],[f370_D])).
fof(f370_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f363,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f304])).
fof(f304,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f207])).
fof(f207,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f114])).
fof(f114,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f113])).
fof(f113,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f55])).
fof(f55,axiom,(
  ! [X20] : ! [X30] : (s(X20,X30) = s(X20,X30) <=> p(s(bool,'T_0')))),
  file('thm.prelim.ordering2num_thm|split|0','thm.bool.REFL_CLAUSE')).
fof(f22435,plain,(
  ( ! [X10,X11] : (s('type.num.num',sK7) != s('type.num.num',X11) | s('type.num.num',X10) = s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num',X11),s('type.num.num','const.num.|48|_0')))) ) | $spl180),
  inference(cnf_transformation,[],[f22435_D])).
fof(f22435_D,plain,(
  ( ! [X10,X11] : (s('type.num.num',sK7) != s('type.num.num',X11) | s('type.num.num',X10) = s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num',X11),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f14540,plain,(
  ( ! [X24,X25] : (s('type.num.num',sK7) != s('type.num.num',X25) | s('type.num.num',X24) = s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num',X25),s('type.num.num','const.num.|48|_0')))) ) | ($spl6 | $spl126)),
  inference(forward_demodulation,[],[f14539,f855])).
fof(f855,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(superposition,[],[f310,f243])).
fof(f243,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f49])).
fof(f49,axiom,(
  ! [X27] : s('type.num.num',X27) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num','const.num.|48|_0')))),
  file('thm.prelim.ordering2num_thm|split|0','thm.arithmetic.ADD_|48|')).
fof(f310,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))),
  inference(rectify,[],[f50])).
fof(f50,axiom,(
  ! [X27,X26] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X26),s('type.num.num',X27))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',X26)))),
  file('thm.prelim.ordering2num_thm|split|0','thm.arithmetic.ADD_SYM')).
fof(f14539,plain,(
  ( ! [X24,X25] : (s('type.num.num',X25) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK7))) | s('type.num.num',X24) = s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num',X25),s('type.num.num','const.num.|48|_0')))) ) | ($spl6 | $spl126)),
  inference(forward_demodulation,[],[f14538,f244])).
fof(f244,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X27] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X27),s('type.num.num','const.num.|48|_0')))),
  file('thm.prelim.ordering2num_thm|split|0','thm.arithmetic.MULT_|48|')).
fof(f14538,plain,(
  ( ! [X24,X25] : (s('type.num.num',X24) = s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num',X25),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',X25) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X24),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK7)))) ) | ($spl6 | $spl126)),
  inference(subsumption_resolution,[],[f14392,f370])).
fof(f14392,plain,(
  ( ! [X24,X25] : (~p(s(bool,'T_0')) | s('type.num.num',X24) = s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num',X25),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',X25) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X24),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK7)))) ) | $spl126),
  inference(superposition,[],[f334,f9377])).
fof(f334,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num',X0)))) | s('type.num.num',X2) = s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num',X1),s('type.num.num',X0))) | s('type.num.num',X1) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num',X0))),s('type.num.num',X3)))) )),
  inference(cnf_transformation,[],[f171])).
fof(f171,plain,(
  ! [X0,X1,X2] : (! [X3] : (s('type.num.num',X1) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num',X0))),s('type.num.num',X3))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num',X0))))) | s('type.num.num',X2) = s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num',X1),s('type.num.num',X0))))),
  inference(ennf_transformation,[],[f139])).
fof(f139,plain,(
  ! [X0,X1,X2] : (? [X3] : (s('type.num.num',X1) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num',X0))),s('type.num.num',X3))) & p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num',X0))))) => s('type.num.num',X2) = s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num',X1),s('type.num.num',X0))))),
  inference(rectify,[],[f48])).
fof(f48,axiom,(
  ! [X26,X32,X33] : (? [X6] : (s('type.num.num',X32) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X33),s('type.num.num',X26))),s('type.num.num',X6))) & p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num',X26))))) => s('type.num.num',X33) = s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num',X32),s('type.num.num',X26))))),
  file('thm.prelim.ordering2num_thm|split|0','thm.arithmetic.DIV_UNIQUE')).
fof(f1493,plain,(
  $false | $spl16),
  inference(trivial_inequality_removal,[],[f1479])).
fof(f1479,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.|48|_0') | $spl16),
  inference(superposition,[],[f231,f1416])).
fof(f1416,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) ) | $spl16),
  inference(resolution,[],[f1318,f300])).
fof(f1318,plain,(
  ( ! [X4] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X4))))) ) | $spl16),
  inference(cnf_transformation,[],[f1318_D])).
fof(f1318_D,plain,(
  ( ! [X4] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X4))))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f23159,plain,(
  ( ! [X7] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X7))))) ) | $spl180),
  inference(subsumption_resolution,[],[f22898,f22757])).
fof(f22757,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num',sK7),s('type.num.num','const.num.|48|_0')))) ) | $spl180),
  inference(equality_resolution,[],[f22435])).
fof(f22898,plain,(
  ( ! [X7] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num',sK7),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X7))))) ) | $spl180),
  inference(superposition,[],[f301,f22757])).
% SZS output end Proof for thm.prelim
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 317
% Passive clauses: 3048
% Generated clauses: 22876
% Final active clauses: 216
% Final passive clauses: 2316
% Input formulas: 69
% Initial clauses: 131
% 
% Duplicate literals: 201
% Trivial inequalities: 411
% Fw subsumption resolutions: 1111
% Fw demodulations: 6003
% Bw demodulations: 83
% 
% Simple tautologies: 159
% Equational tautologies: 39
% Forward subsumptions: 11889
% Fw demodulations to eq. taut.: 6
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 316
% Factoring: 1
% Forward superposition: 5524
% Backward superposition: 8787
% Self superposition: 136
% Equality factoring: 38
% Equality resolution: 44
% 
% Split clauses: 130
% Split components: 261
% SAT solver clauses: 284
% SAT solver unit clauses: 9
% SAT solver binary clauses: 115
% SAT solver learnt clauses: 1
% 
% Sat splits: 130
% Sat splitting refutations: 36
% 
% TWLsolver clauses: 243
% TWLsolver calls for satisfiability: 106
% 
% Memory used [KB]: 13560
% Time elapsed: 0.349 s
% ------------------------------
% ------------------------------
% Success in time 0.355 s
