Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.pred_set
% (46533)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 723
% Passive clauses: 6352
% Generated clauses: 19612
% Final active clauses: 721
% Final passive clauses: 5351
% Input formulas: 48
% Initial clauses: 178
% 
% Duplicate literals: 452
% Trivial inequalities: 39
% Fw subsumption resolutions: 1578
% Fw demodulations: 2146
% Bw demodulations: 2
% 
% Simple tautologies: 252
% Equational tautologies: 45
% Forward subsumptions: 8368
% Fw demodulations to eq. taut.: 3
% 
% Binary resolution: 5717
% Factoring: 1129
% Forward superposition: 3594
% Backward superposition: 4504
% Self superposition: 54
% Equality factoring: 26
% Equality resolution: 12
% 
% Split clauses: 419
% Split components: 852
% SAT solver clauses: 637
% SAT solver unit clauses: 7
% SAT solver binary clauses: 93
% 
% Sat splits: 419
% 
% TWLsolver clauses: 637
% TWLsolver calls for satisfiability: 352
% 
% Memory used [KB]: 10874
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 294 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f9127,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f9123,f250,f440,f441,f443,f1046,f1059,f248,f437,f268,f450,f4020,f4021,f249,f435,f299,f460,f1081,f1090,f247,f438,f264,f448,f6433,f6457,f5263,f5265,f298,f461])).
fof(f461,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),'const.pred_set.GSPEC_1'(s(fun(X0,'type.pair.prod'(X0,bool)),'_dst_x1x1_2'(s(fun(X0,fun(bool,'type.pair.prod'(X0,bool))),'const.pair.,_0'),s(fun(X0,bool),'const.relation.RTC_2'(s(fun(X0,fun(X0,bool)),X2),s(X0,X1)))))))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X1))),s(X0,sK16(X2,X1,X0))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),'const.pred_set.GSPEC_1'(s(fun(X0,'type.pair.prod'(X0,bool)),'_dst_x1x1_2'(s(fun(X0,fun(bool,'type.pair.prod'(X0,bool))),'const.pair.,_0'),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X1))))))))))) ) | $spl30),
  inference(cnf_transformation,[],[f461_D])).
fof(f461_D,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),'const.pred_set.GSPEC_1'(s(fun(X0,'type.pair.prod'(X0,bool)),'_dst_x1x1_2'(s(fun(X0,fun(bool,'type.pair.prod'(X0,bool))),'const.pair.,_0'),s(fun(X0,bool),'const.relation.RTC_2'(s(fun(X0,fun(X0,bool)),X2),s(X0,X1)))))))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X1))),s(X0,sK16(X2,X1,X0))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),'const.pred_set.GSPEC_1'(s(fun(X0,'type.pair.prod'(X0,bool)),'_dst_x1x1_2'(s(fun(X0,fun(bool,'type.pair.prod'(X0,bool))),'const.pair.,_0'),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X1))))))))))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f298,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),'const.pred_set.GSPEC_1'(s(fun(X0,'type.pair.prod'(X0,bool)),'_dst_x1x1_2'(s(fun(X0,fun(bool,'type.pair.prod'(X0,bool))),'const.pair.,_0'),s(fun(X0,bool),'const.relation.RTC_2'(s(fun(X0,fun(X0,bool)),X2),s(X0,X1)))))))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X1))),s(X0,sK16(X2,X1,X0))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),'const.pred_set.GSPEC_1'(s(fun(X0,'type.pair.prod'(X0,bool)),'_dst_x1x1_2'(s(fun(X0,fun(bool,'type.pair.prod'(X0,bool))),'const.pair.,_0'),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X1))))))))))) )),
  inference(cnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),'const.pred_set.GSPEC_1'(s(fun(X0,'type.pair.prod'(X0,bool)),'_dst_x1x1_2'(s(fun(X0,fun(bool,'type.pair.prod'(X0,bool))),'const.pair.,_0'),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X1)))))))))) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X1))),s(X0,sK16(X2,X1,X0))))) & ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),'const.pred_set.GSPEC_1'(s(fun(X0,'type.pair.prod'(X0,bool)),'_dst_x1x1_2'(s(fun(X0,fun(bool,'type.pair.prod'(X0,bool))),'const.pair.,_0'),s(fun(X0,bool),'const.relation.RTC_2'(s(fun(X0,fun(X0,bool)),X2),s(X0,sK16(X2,X1,X0)))))))))))) | p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),'const.pred_set.GSPEC_1'(s(fun(X0,'type.pair.prod'(X0,bool)),'_dst_x1x1_2'(s(fun(X0,fun(bool,'type.pair.prod'(X0,bool))),'const.pair.,_0'),s(fun(X0,bool),'const.relation.RTC_2'(s(fun(X0,fun(X0,bool)),X2),s(X0,X1)))))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK16])],[f139])).
fof(f139,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),'const.pred_set.GSPEC_1'(s(fun(X0,'type.pair.prod'(X0,bool)),'_dst_x1x1_2'(s(fun(X0,fun(bool,'type.pair.prod'(X0,bool))),'const.pair.,_0'),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X1)))))))))) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X1))),s(X0,X3)))) & ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),'const.pred_set.GSPEC_1'(s(fun(X0,'type.pair.prod'(X0,bool)),'_dst_x1x1_2'(s(fun(X0,fun(bool,'type.pair.prod'(X0,bool))),'const.pair.,_0'),s(fun(X0,bool),'const.relation.RTC_2'(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))))))))))) | p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),'const.pred_set.GSPEC_1'(s(fun(X0,'type.pair.prod'(X0,bool)),'_dst_x1x1_2'(s(fun(X0,fun(bool,'type.pair.prod'(X0,bool))),'const.pair.,_0'),s(fun(X0,bool),'const.relation.RTC_2'(s(fun(X0,fun(X0,bool)),X2),s(X0,X1)))))))))))),
  inference(flattening,[],[f138])).
fof(f138,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),'const.pred_set.GSPEC_1'(s(fun(X0,'type.pair.prod'(X0,bool)),'_dst_x1x1_2'(s(fun(X0,fun(bool,'type.pair.prod'(X0,bool))),'const.pair.,_0'),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X1)))))))))) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X1))),s(X0,X3)))) & ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),'const.pred_set.GSPEC_1'(s(fun(X0,'type.pair.prod'(X0,bool)),'_dst_x1x1_2'(s(fun(X0,fun(bool,'type.pair.prod'(X0,bool))),'const.pair.,_0'),s(fun(X0,bool),'const.relation.RTC_2'(s(fun(X0,fun(X0,bool)),X2),s(X0,X3)))))))))))) | p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),'const.pred_set.GSPEC_1'(s(fun(X0,'type.pair.prod'(X0,bool)),'_dst_x1x1_2'(s(fun(X0,fun(bool,'type.pair.prod'(X0,bool))),'const.pair.,_0'),s(fun(X0,bool),'const.relation.RTC_2'(s(fun(X0,fun(X0,bool)),X2),s(X0,X1)))))))))))),
  inference(ennf_transformation,[],[f87])).
fof(f87,plain,(
  ! [X0,X1,X2] : ((p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),'const.pred_set.GSPEC_1'(s(fun(X0,'type.pair.prod'(X0,bool)),'_dst_x1x1_2'(s(fun(X0,fun(bool,'type.pair.prod'(X0,bool))),'const.pair.,_0'),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X1)))))))))) & ! [X3] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X1))),s(X0,X3)))) => p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),'const.pred_set.GSPEC_1'(s(fun(X0,'type.pair.prod'(X0,bool)),'_dst_x1x1_2'(s(fun(X0,fun(bool,'type.pair.prod'(X0,bool))),'const.pair.,_0'),s(fun(X0,bool),'const.relation.RTC_2'(s(fun(X0,fun(X0,bool)),X2),s(X0,X3)))))))))))) => p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),'const.pred_set.GSPEC_1'(s(fun(X0,'type.pair.prod'(X0,bool)),'_dst_x1x1_2'(s(fun(X0,fun(bool,'type.pair.prod'(X0,bool))),'const.pair.,_0'),s(fun(X0,bool),'const.relation.RTC_2'(s(fun(X0,fun(X0,bool)),X2),s(X0,X1)))))))))))),
  inference(flattening,[],[f86])).
fof(f86,plain,(
  ! [X0] : ! [X1,X2] : ((p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),'const.pred_set.GSPEC_1'(s(fun(X0,'type.pair.prod'(X0,bool)),'_dst_x1x1_2'(s(fun(X0,fun(bool,'type.pair.prod'(X0,bool))),'const.pair.,_0'),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X1)))))))))) & ! [X3] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X1))),s(X0,X3)))) => p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),'const.pred_set.GSPEC_1'(s(fun(X0,'type.pair.prod'(X0,bool)),'_dst_x1x1_2'(s(fun(X0,fun(bool,'type.pair.prod'(X0,bool))),'const.pair.,_0'),s(fun(X0,bool),'const.relation.RTC_2'(s(fun(X0,fun(X0,bool)),X2),s(X0,X3)))))))))))) => p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),'const.pred_set.GSPEC_1'(s(fun(X0,'type.pair.prod'(X0,bool)),'_dst_x1x1_2'(s(fun(X0,fun(bool,'type.pair.prod'(X0,bool))),'const.pair.,_0'),s(fun(X0,bool),'const.relation.RTC_2'(s(fun(X0,fun(X0,bool)),X2),s(X0,X1)))))))))))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X6] : ! [X21,X7] : ((p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),'const.pred_set.GSPEC_1'(s(fun(X6,'type.pair.prod'(X6,bool)),'_dst_x1x1_2'(s(fun(X6,fun(bool,'type.pair.prod'(X6,bool))),'const.pair.,_0'),s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X21)))))))))) & ! [X9] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X21))),s(X6,X9)))) => p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),'const.pred_set.GSPEC_1'(s(fun(X6,'type.pair.prod'(X6,bool)),'_dst_x1x1_2'(s(fun(X6,fun(bool,'type.pair.prod'(X6,bool))),'const.pair.,_0'),s(fun(X6,bool),'const.relation.RTC_2'(s(fun(X6,fun(X6,bool)),X7),s(X6,X9)))))))))))) => p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),'const.pred_set.GSPEC_1'(s(fun(X6,'type.pair.prod'(X6,bool)),'_dst_x1x1_2'(s(fun(X6,fun(bool,'type.pair.prod'(X6,bool))),'const.pair.,_0'),s(fun(X6,bool),'const.relation.RTC_2'(s(fun(X6,fun(X6,bool)),X7),s(X6,X21)))))))))))),
  file('thm.pred_set.pred_set_42|48|_prove_ltb','thm.pred_set.pred_set_4|49|9_prove_ltb')).
fof(f5265,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK15))),s(sK13,sK16(sK14,sK15,sK13))))) | $spl2186),
  inference(cnf_transformation,[],[f5265_D])).
fof(f5265_D,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK15))),s(sK13,sK16(sK14,sK15,sK13))))) <=> ~$spl2186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2186])])).
fof(f5263,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK15))),s(sK13,sK16(sK14,sK15,sK13))))) | ($spl6 | $spl9 | $spl30 | $spl462 | $spl482)),
  inference(subsumption_resolution,[],[f5262,f440])).
fof(f5262,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK15))),s(sK13,sK16(sK14,sK15,sK13))))) | ($spl9 | $spl30 | $spl462 | $spl482)),
  inference(forward_demodulation,[],[f5261,f1090])).
fof(f5261,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK15))),s(sK13,sK16(sK14,sK15,sK13))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.GSPEC_1'(s(fun(sK13,'type.pair.prod'(sK13,bool)),'_dst_x1x1_2'(s(fun(sK13,fun(bool,'type.pair.prod'(sK13,bool))),'const.pair.,_0'),s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK15)))))))))) | ($spl9 | $spl30 | $spl462)),
  inference(subsumption_resolution,[],[f5186,f443])).
fof(f5186,plain,(
  p(s(bool,'F_0')) | p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK15))),s(sK13,sK16(sK14,sK15,sK13))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.GSPEC_1'(s(fun(sK13,'type.pair.prod'(sK13,bool)),'_dst_x1x1_2'(s(fun(sK13,fun(bool,'type.pair.prod'(sK13,bool))),'const.pair.,_0'),s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK15)))))))))) | ($spl30 | $spl462)),
  inference(superposition,[],[f461,f1059])).
fof(f6457,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK15))),s(sK13,sK16(sK14,sK15,sK13)))) | $spl2598),
  inference(cnf_transformation,[],[f6457_D])).
fof(f6457_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK15))),s(sK13,sK16(sK14,sK15,sK13)))) <=> ~$spl2598),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2598])])).
fof(f6433,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK15))),s(sK13,sK16(sK14,sK15,sK13)))) | ($spl12 | $spl2186)),
  inference(resolution,[],[f5265,f448])).
fof(f448,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl12),
  inference(cnf_transformation,[],[f448_D])).
fof(f448_D,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f264,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f176])).
fof(f176,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X12] : (s(bool,'T_0') = s(bool,X12) <=> p(s(bool,X12)))),
  file('thm.pred_set.pred_set_42|48|_prove_ltb','thm.bool.EQ_CLAUSES|split|1')).
fof(f438,plain,(
  ( ! [X4] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.GSPEC_1'(s(fun(sK13,'type.pair.prod'(sK13,bool)),'_dst_x1x1_2'(s(fun(sK13,fun(bool,'type.pair.prod'(sK13,bool))),'const.pair.,_0'),s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,X4))))))))))) ) | $spl4),
  inference(cnf_transformation,[],[f438_D])).
fof(f438_D,plain,(
  ( ! [X4] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.GSPEC_1'(s(fun(sK13,'type.pair.prod'(sK13,bool)),'_dst_x1x1_2'(s(fun(sK13,fun(bool,'type.pair.prod'(sK13,bool))),'const.pair.,_0'),s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,X4))))))))))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f247,plain,(
  ( ! [X4] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.GSPEC_1'(s(fun(sK13,'type.pair.prod'(sK13,bool)),'_dst_x1x1_2'(s(fun(sK13,fun(bool,'type.pair.prod'(sK13,bool))),'const.pair.,_0'),s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,X4))))))))))) )),
  inference(cnf_transformation,[],[f173])).
fof(f173,plain,(
  ! [X4] : p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.GSPEC_1'(s(fun(sK13,'type.pair.prod'(sK13,bool)),'_dst_x1x1_2'(s(fun(sK13,fun(bool,'type.pair.prod'(sK13,bool))),'const.pair.,_0'),s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,X4)))))))))) & (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.GSPEC_1'(s(fun(sK13,'type.pair.prod'(sK13,bool)),'_dst_x1x1_2'(s(fun(sK13,fun(bool,'type.pair.prod'(sK13,bool))),'const.pair.,_0'),s(fun(sK13,bool),'const.relation.RTC_2'(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK15)))))))))) & ! [X3] : (~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK15))),s(sK13,X3)))) | p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.GSPEC_1'(s(fun(sK13,'type.pair.prod'(sK13,bool)),'_dst_x1x1_2'(s(fun(sK13,fun(bool,'type.pair.prod'(sK13,bool))),'const.pair.,_0'),s(fun(sK13,bool),'const.relation.RTC_2'(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,X3))))))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13,sK14,sK15])],[f172])).
fof(f172,plain,(
  ? [X0,X1] : (! [X4] : p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),'const.pred_set.GSPEC_1'(s(fun(X0,'type.pair.prod'(X0,bool)),'_dst_x1x1_2'(s(fun(X0,fun(bool,'type.pair.prod'(X0,bool))),'const.pair.,_0'),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4)))))))))) & ? [X2] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),'const.pred_set.GSPEC_1'(s(fun(X0,'type.pair.prod'(X0,bool)),'_dst_x1x1_2'(s(fun(X0,fun(bool,'type.pair.prod'(X0,bool))),'const.pair.,_0'),s(fun(X0,bool),'const.relation.RTC_2'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2)))))))))) & ! [X3] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),'const.pred_set.GSPEC_1'(s(fun(X0,'type.pair.prod'(X0,bool)),'_dst_x1x1_2'(s(fun(X0,fun(bool,'type.pair.prod'(X0,bool))),'const.pair.,_0'),s(fun(X0,bool),'const.relation.RTC_2'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3)))))))))))))),
  inference(rectify,[],[f127])).
fof(f127,plain,(
  ? [X0,X1] : (! [X2] : p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),'const.pred_set.GSPEC_1'(s(fun(X0,'type.pair.prod'(X0,bool)),'_dst_x1x1_2'(s(fun(X0,fun(bool,'type.pair.prod'(X0,bool))),'const.pair.,_0'),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2)))))))))) & ? [X3] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),'const.pred_set.GSPEC_1'(s(fun(X0,'type.pair.prod'(X0,bool)),'_dst_x1x1_2'(s(fun(X0,fun(bool,'type.pair.prod'(X0,bool))),'const.pair.,_0'),s(fun(X0,bool),'const.relation.RTC_2'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3)))))))))) & ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X4)))) | p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),'const.pred_set.GSPEC_1'(s(fun(X0,'type.pair.prod'(X0,bool)),'_dst_x1x1_2'(s(fun(X0,fun(bool,'type.pair.prod'(X0,bool))),'const.pair.,_0'),s(fun(X0,bool),'const.relation.RTC_2'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4)))))))))))))),
  inference(ennf_transformation,[],[f51])).
fof(f51,plain,(
  ~! [X0,X1] : (! [X2] : p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),'const.pred_set.GSPEC_1'(s(fun(X0,'type.pair.prod'(X0,bool)),'_dst_x1x1_2'(s(fun(X0,fun(bool,'type.pair.prod'(X0,bool))),'const.pair.,_0'),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2)))))))))) => ! [X3] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),'const.pred_set.GSPEC_1'(s(fun(X0,'type.pair.prod'(X0,bool)),'_dst_x1x1_2'(s(fun(X0,fun(bool,'type.pair.prod'(X0,bool))),'const.pair.,_0'),s(fun(X0,bool),'const.relation.RTC_2'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3)))))))))) => ? [X4] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X4)))) & ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),'const.pred_set.GSPEC_1'(s(fun(X0,'type.pair.prod'(X0,bool)),'_dst_x1x1_2'(s(fun(X0,fun(bool,'type.pair.prod'(X0,bool))),'const.pair.,_0'),s(fun(X0,bool),'const.relation.RTC_2'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4)))))))))))))),
  inference(flattening,[],[f50])).
fof(f50,plain,(
  ~! [X0] : ! [X1] : (! [X2] : p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),'const.pred_set.GSPEC_1'(s(fun(X0,'type.pair.prod'(X0,bool)),'_dst_x1x1_2'(s(fun(X0,fun(bool,'type.pair.prod'(X0,bool))),'const.pair.,_0'),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2)))))))))) => ! [X3] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),'const.pred_set.GSPEC_1'(s(fun(X0,'type.pair.prod'(X0,bool)),'_dst_x1x1_2'(s(fun(X0,fun(bool,'type.pair.prod'(X0,bool))),'const.pair.,_0'),s(fun(X0,bool),'const.relation.RTC_2'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3)))))))))) => ? [X4] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X4)))) & ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),'const.pred_set.GSPEC_1'(s(fun(X0,'type.pair.prod'(X0,bool)),'_dst_x1x1_2'(s(fun(X0,fun(bool,'type.pair.prod'(X0,bool))),'const.pair.,_0'),s(fun(X0,bool),'const.relation.RTC_2'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4)))))))))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7] : (! [X8] : p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),'const.pred_set.GSPEC_1'(s(fun(X6,'type.pair.prod'(X6,bool)),'_dst_x1x1_2'(s(fun(X6,fun(bool,'type.pair.prod'(X6,bool))),'const.pair.,_0'),s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X8)))))))))) => ! [X9] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),'const.pred_set.GSPEC_1'(s(fun(X6,'type.pair.prod'(X6,bool)),'_dst_x1x1_2'(s(fun(X6,fun(bool,'type.pair.prod'(X6,bool))),'const.pair.,_0'),s(fun(X6,bool),'const.relation.RTC_2'(s(fun(X6,fun(X6,bool)),X7),s(X6,X9)))))))))) => ? [X10] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X9))),s(X6,X10)))) & ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),'const.pred_set.GSPEC_1'(s(fun(X6,'type.pair.prod'(X6,bool)),'_dst_x1x1_2'(s(fun(X6,fun(bool,'type.pair.prod'(X6,bool))),'const.pair.,_0'),s(fun(X6,bool),'const.relation.RTC_2'(s(fun(X6,fun(X6,bool)),X7),s(X6,X10)))))))))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7] : (! [X8] : p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),'const.pred_set.GSPEC_1'(s(fun(X6,'type.pair.prod'(X6,bool)),'_dst_x1x1_2'(s(fun(X6,fun(bool,'type.pair.prod'(X6,bool))),'const.pair.,_0'),s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X8)))))))))) => ! [X9] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),'const.pred_set.GSPEC_1'(s(fun(X6,'type.pair.prod'(X6,bool)),'_dst_x1x1_2'(s(fun(X6,fun(bool,'type.pair.prod'(X6,bool))),'const.pair.,_0'),s(fun(X6,bool),'const.relation.RTC_2'(s(fun(X6,fun(X6,bool)),X7),s(X6,X9)))))))))) => ? [X10] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X9))),s(X6,X10)))) & ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),'const.pred_set.GSPEC_1'(s(fun(X6,'type.pair.prod'(X6,bool)),'_dst_x1x1_2'(s(fun(X6,fun(bool,'type.pair.prod'(X6,bool))),'const.pair.,_0'),s(fun(X6,bool),'const.relation.RTC_2'(s(fun(X6,fun(X6,bool)),X7),s(X6,X10)))))))))))))),
  file('thm.pred_set.pred_set_42|48|_prove_ltb','thm.pred_set.pred_set_42|48|_prove_ltb')).
fof(f1090,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.GSPEC_1'(s(fun(sK13,'type.pair.prod'(sK13,bool)),'_dst_x1x1_2'(s(fun(sK13,fun(bool,'type.pair.prod'(sK13,bool))),'const.pair.,_0'),s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,X0)))))))))) ) | $spl482),
  inference(cnf_transformation,[],[f1090_D])).
fof(f1090_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.GSPEC_1'(s(fun(sK13,'type.pair.prod'(sK13,bool)),'_dst_x1x1_2'(s(fun(sK13,fun(bool,'type.pair.prod'(sK13,bool))),'const.pair.,_0'),s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,X0)))))))))) ) <=> ~$spl482),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl482])])).
fof(f1081,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.GSPEC_1'(s(fun(sK13,'type.pair.prod'(sK13,bool)),'_dst_x1x1_2'(s(fun(sK13,fun(bool,'type.pair.prod'(sK13,bool))),'const.pair.,_0'),s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,X0)))))))))) ) | ($spl4 | $spl12)),
  inference(resolution,[],[f438,f448])).
fof(f460,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),'const.pred_set.GSPEC_1'(s(fun(X0,'type.pair.prod'(X0,bool)),'_dst_x1x1_2'(s(fun(X0,fun(bool,'type.pair.prod'(X0,bool))),'const.pair.,_0'),s(fun(X0,bool),'const.relation.RTC_2'(s(fun(X0,fun(X0,bool)),X2),s(X0,sK16(X2,X1,X0))))))))))) | p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),'const.pred_set.GSPEC_1'(s(fun(X0,'type.pair.prod'(X0,bool)),'_dst_x1x1_2'(s(fun(X0,fun(bool,'type.pair.prod'(X0,bool))),'const.pair.,_0'),s(fun(X0,bool),'const.relation.RTC_2'(s(fun(X0,fun(X0,bool)),X2),s(X0,X1)))))))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),'const.pred_set.GSPEC_1'(s(fun(X0,'type.pair.prod'(X0,bool)),'_dst_x1x1_2'(s(fun(X0,fun(bool,'type.pair.prod'(X0,bool))),'const.pair.,_0'),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X1))))))))))) ) | $spl28),
  inference(cnf_transformation,[],[f460_D])).
fof(f460_D,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),'const.pred_set.GSPEC_1'(s(fun(X0,'type.pair.prod'(X0,bool)),'_dst_x1x1_2'(s(fun(X0,fun(bool,'type.pair.prod'(X0,bool))),'const.pair.,_0'),s(fun(X0,bool),'const.relation.RTC_2'(s(fun(X0,fun(X0,bool)),X2),s(X0,sK16(X2,X1,X0))))))))))) | p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),'const.pred_set.GSPEC_1'(s(fun(X0,'type.pair.prod'(X0,bool)),'_dst_x1x1_2'(s(fun(X0,fun(bool,'type.pair.prod'(X0,bool))),'const.pair.,_0'),s(fun(X0,bool),'const.relation.RTC_2'(s(fun(X0,fun(X0,bool)),X2),s(X0,X1)))))))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),'const.pred_set.GSPEC_1'(s(fun(X0,'type.pair.prod'(X0,bool)),'_dst_x1x1_2'(s(fun(X0,fun(bool,'type.pair.prod'(X0,bool))),'const.pair.,_0'),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X1))))))))))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f299,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),'const.pred_set.GSPEC_1'(s(fun(X0,'type.pair.prod'(X0,bool)),'_dst_x1x1_2'(s(fun(X0,fun(bool,'type.pair.prod'(X0,bool))),'const.pair.,_0'),s(fun(X0,bool),'const.relation.RTC_2'(s(fun(X0,fun(X0,bool)),X2),s(X0,X1)))))))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),'const.pred_set.GSPEC_1'(s(fun(X0,'type.pair.prod'(X0,bool)),'_dst_x1x1_2'(s(fun(X0,fun(bool,'type.pair.prod'(X0,bool))),'const.pair.,_0'),s(fun(X0,bool),'const.relation.RTC_2'(s(fun(X0,fun(X0,bool)),X2),s(X0,sK16(X2,X1,X0))))))))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),'const.pred_set.GSPEC_1'(s(fun(X0,'type.pair.prod'(X0,bool)),'_dst_x1x1_2'(s(fun(X0,fun(bool,'type.pair.prod'(X0,bool))),'const.pair.,_0'),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X1))))))))))) )),
  inference(cnf_transformation,[],[f195])).
fof(f435,plain,(
  ( ! [X3] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.GSPEC_1'(s(fun(sK13,'type.pair.prod'(sK13,bool)),'_dst_x1x1_2'(s(fun(sK13,fun(bool,'type.pair.prod'(sK13,bool))),'const.pair.,_0'),s(fun(sK13,bool),'const.relation.RTC_2'(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,X3)))))))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK15))),s(sK13,X3))))) ) | $spl0),
  inference(cnf_transformation,[],[f435_D])).
fof(f435_D,plain,(
  ( ! [X3] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.GSPEC_1'(s(fun(sK13,'type.pair.prod'(sK13,bool)),'_dst_x1x1_2'(s(fun(sK13,fun(bool,'type.pair.prod'(sK13,bool))),'const.pair.,_0'),s(fun(sK13,bool),'const.relation.RTC_2'(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,X3)))))))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK15))),s(sK13,X3))))) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f249,plain,(
  ( ! [X3] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.GSPEC_1'(s(fun(sK13,'type.pair.prod'(sK13,bool)),'_dst_x1x1_2'(s(fun(sK13,fun(bool,'type.pair.prod'(sK13,bool))),'const.pair.,_0'),s(fun(sK13,bool),'const.relation.RTC_2'(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,X3)))))))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK15))),s(sK13,X3))))) )),
  inference(cnf_transformation,[],[f173])).
fof(f4021,plain,(
  ( ! [X0] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.GSPEC_1'(s(fun(sK13,'type.pair.prod'(sK13,bool)),'_dst_x1x1_2'(s(fun(sK13,fun(bool,'type.pair.prod'(sK13,bool))),'const.pair.,_0'),s(fun(sK13,bool),'const.relation.RTC_2'(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,X0)))))))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK15))),s(sK13,sK16(sK14,X0,sK13)))))) ) | $spl1658),
  inference(cnf_transformation,[],[f4021_D])).
fof(f4021_D,plain,(
  ( ! [X0] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.GSPEC_1'(s(fun(sK13,'type.pair.prod'(sK13,bool)),'_dst_x1x1_2'(s(fun(sK13,fun(bool,'type.pair.prod'(sK13,bool))),'const.pair.,_0'),s(fun(sK13,bool),'const.relation.RTC_2'(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,X0)))))))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK15))),s(sK13,sK16(sK14,X0,sK13)))))) ) <=> ~$spl1658),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1658])])).
fof(f4020,plain,(
  ( ! [X0] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.GSPEC_1'(s(fun(sK13,'type.pair.prod'(sK13,bool)),'_dst_x1x1_2'(s(fun(sK13,fun(bool,'type.pair.prod'(sK13,bool))),'const.pair.,_0'),s(fun(sK13,bool),'const.relation.RTC_2'(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,X0)))))))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK15))),s(sK13,sK16(sK14,X0,sK13)))))) ) | ($spl0 | $spl6 | $spl28 | $spl482)),
  inference(subsumption_resolution,[],[f4019,f440])).
fof(f4019,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.GSPEC_1'(s(fun(sK13,'type.pair.prod'(sK13,bool)),'_dst_x1x1_2'(s(fun(sK13,fun(bool,'type.pair.prod'(sK13,bool))),'const.pair.,_0'),s(fun(sK13,bool),'const.relation.RTC_2'(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,X0)))))))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK15))),s(sK13,sK16(sK14,X0,sK13)))))) ) | ($spl0 | $spl28 | $spl482)),
  inference(forward_demodulation,[],[f3986,f1090])).
fof(f3986,plain,(
  ( ! [X0] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.GSPEC_1'(s(fun(sK13,'type.pair.prod'(sK13,bool)),'_dst_x1x1_2'(s(fun(sK13,fun(bool,'type.pair.prod'(sK13,bool))),'const.pair.,_0'),s(fun(sK13,bool),'const.relation.RTC_2'(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,X0)))))))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.GSPEC_1'(s(fun(sK13,'type.pair.prod'(sK13,bool)),'_dst_x1x1_2'(s(fun(sK13,fun(bool,'type.pair.prod'(sK13,bool))),'const.pair.,_0'),s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,X0)))))))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK15))),s(sK13,sK16(sK14,X0,sK13)))))) ) | ($spl0 | $spl28)),
  inference(resolution,[],[f460,f435])).
fof(f450,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) | $spl16),
  inference(cnf_transformation,[],[f450_D])).
fof(f450_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f268,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f178])).
fof(f178,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f62])).
fof(f62,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f61])).
fof(f61,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X12] : (s(bool,'F_0') = s(bool,X12) <=> ~p(s(bool,X12)))),
  file('thm.pred_set.pred_set_42|48|_prove_ltb','thm.bool.EQ_CLAUSES|split|2')).
fof(f437,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.GSPEC_1'(s(fun(sK13,'type.pair.prod'(sK13,bool)),'_dst_x1x1_2'(s(fun(sK13,fun(bool,'type.pair.prod'(sK13,bool))),'const.pair.,_0'),s(fun(sK13,bool),'const.relation.RTC_2'(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK15)))))))))) | $spl3),
  inference(cnf_transformation,[],[f437_D])).
fof(f437_D,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.GSPEC_1'(s(fun(sK13,'type.pair.prod'(sK13,bool)),'_dst_x1x1_2'(s(fun(sK13,fun(bool,'type.pair.prod'(sK13,bool))),'const.pair.,_0'),s(fun(sK13,bool),'const.relation.RTC_2'(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK15)))))))))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f248,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.GSPEC_1'(s(fun(sK13,'type.pair.prod'(sK13,bool)),'_dst_x1x1_2'(s(fun(sK13,fun(bool,'type.pair.prod'(sK13,bool))),'const.pair.,_0'),s(fun(sK13,bool),'const.relation.RTC_2'(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK15))))))))))),
  inference(cnf_transformation,[],[f173])).
fof(f1059,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.GSPEC_1'(s(fun(sK13,'type.pair.prod'(sK13,bool)),'_dst_x1x1_2'(s(fun(sK13,fun(bool,'type.pair.prod'(sK13,bool))),'const.pair.,_0'),s(fun(sK13,bool),'const.relation.RTC_2'(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK15))))))))) | $spl462),
  inference(cnf_transformation,[],[f1059_D])).
fof(f1059_D,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.GSPEC_1'(s(fun(sK13,'type.pair.prod'(sK13,bool)),'_dst_x1x1_2'(s(fun(sK13,fun(bool,'type.pair.prod'(sK13,bool))),'const.pair.,_0'),s(fun(sK13,bool),'const.relation.RTC_2'(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK15))))))))) <=> ~$spl462),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl462])])).
fof(f1046,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.GSPEC_1'(s(fun(sK13,'type.pair.prod'(sK13,bool)),'_dst_x1x1_2'(s(fun(sK13,fun(bool,'type.pair.prod'(sK13,bool))),'const.pair.,_0'),s(fun(sK13,bool),'const.relation.RTC_2'(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK15))))))))) | ($spl3 | $spl16)),
  inference(resolution,[],[f437,f450])).
fof(f443,plain,(
  ~p(s(bool,'F_0')) | $spl9),
  inference(cnf_transformation,[],[f443_D])).
fof(f443_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f441,plain,(
  ~p(s(bool,'F_0'))),
  inference(global_subsumption,[],[f253,f250])).
fof(f253,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f174])).
fof(f174,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f52])).
fof(f52,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f20])).
fof(f20,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.pred_set.pred_set_42|48|_prove_ltb','thm.bool.NOT_CLAUSES|split|2')).
fof(f440,plain,(
  p(s(bool,'T_0')) | $spl6),
  inference(cnf_transformation,[],[f440_D])).
fof(f440_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f250,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pred_set.pred_set_42|48|_prove_ltb',aHL_TRUTH)).
fof(f9123,plain,(
  $false | ($spl6 | $spl9 | $spl462 | $spl1658 | $spl2598)),
  inference(subsumption_resolution,[],[f9122,f440])).
fof(f9122,plain,(
  ~p(s(bool,'T_0')) | ($spl9 | $spl462 | $spl1658 | $spl2598)),
  inference(forward_demodulation,[],[f9121,f6457])).
fof(f9121,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK15))),s(sK13,sK16(sK14,sK15,sK13))))) | ($spl9 | $spl462 | $spl1658)),
  inference(subsumption_resolution,[],[f9081,f443])).
fof(f9081,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK15))),s(sK13,sK16(sK14,sK15,sK13))))) | ($spl462 | $spl1658)),
  inference(superposition,[],[f4021,f1059])).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 374
% Passive clauses: 1705
% Generated clauses: 8804
% Final active clauses: 361
% Final passive clauses: 1291
% Input formulas: 48
% Initial clauses: 178
% 
% Duplicate literals: 187
% Trivial inequalities: 34
% Fw subsumption resolutions: 647
% Fw demodulations: 582
% Bw demodulations: 45
% Condensations: 238
% Global subsumptions: 13
% 
% Simple tautologies: 235
% Equational tautologies: 19
% Forward subsumptions: 3167
% 
% Binary resolution: 2235
% Factoring: 269
% Forward superposition: 946
% Backward superposition: 1623
% Self superposition: 47
% Equality factoring: 26
% Equality resolution: 5
% 
% Split clauses: 52
% Split components: 104
% SAT solver clauses: 2161
% SAT solver unit clauses: 204
% SAT solver binary clauses: 66
% SAT solver learnt clauses: 27
% SAT solver learnt literals: 10
% 
% Sat splits: 52
% Sat splitting refutations: 6
% 
% TWLsolver clauses: 2108
% TWLsolver calls for satisfiability: 1759
% 
% Memory used [KB]: 9850
% Time elapsed: 0.407 s
% ------------------------------
% ------------------------------
% Success in time 1.019 s
