Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.integral
% (91431)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 854
% Passive clauses: 8144
% Generated clauses: 21864
% Final active clauses: 839
% Final passive clauses: 6886
% Input formulas: 40
% Initial clauses: 177
% 
% Duplicate literals: 576
% Trivial inequalities: 25
% Fw subsumption resolutions: 2533
% Fw demodulations: 873
% Bw demodulations: 14
% 
% Simple tautologies: 260
% Equational tautologies: 60
% Forward subsumptions: 8506
% 
% Binary resolution: 10147
% Factoring: 1416
% Forward superposition: 2652
% Backward superposition: 3056
% Self superposition: 30
% Equality factoring: 34
% Equality resolution: 5
% 
% Split clauses: 677
% Split components: 1361
% SAT solver clauses: 1198
% SAT solver unit clauses: 25
% SAT solver binary clauses: 144
% 
% Sat splits: 677
% 
% TWLsolver clauses: 1197
% TWLsolver calls for satisfiability: 559
% 
% Memory used [KB]: 9210
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.integral
% (91440)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 344
% Passive clauses: 2065
% Generated clauses: 9939
% Final active clauses: 334
% Final passive clauses: 1700
% Input formulas: 40
% Initial clauses: 177
% 
% Duplicate literals: 216
% Trivial inequalities: 31
% Fw subsumption resolutions: 1041
% Fw demodulations: 525
% Bw demodulations: 27
% Condensations: 310
% Global subsumptions: 13
% 
% Simple tautologies: 262
% Equational tautologies: 21
% Forward subsumptions: 2805
% 
% Binary resolution: 2749
% Factoring: 127
% Forward superposition: 1368
% Backward superposition: 1209
% Self superposition: 30
% Equality factoring: 34
% Equality resolution: 9
% 
% Split clauses: 37
% Split components: 79
% SAT solver clauses: 2794
% SAT solver unit clauses: 179
% SAT solver binary clauses: 89
% SAT solver learnt clauses: 26
% SAT solver learnt literals: 10
% 
% Sat splits: 37
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 2742
% TWLsolver calls for satisfiability: 1651
% 
% Memory used [KB]: 11641
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.integral
% (91453)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 2732
% Passive clauses: 8977
% Generated clauses: 190313
% Final active clauses: 1616
% Final passive clauses: 2096
% Input formulas: 40
% Initial clauses: 177
% 
% Duplicate literals: 8103
% Trivial inequalities: 353
% Fw subsumption resolutions: 8040
% Fw demodulations: 30070
% Bw demodulations: 2047
% Condensations: 14544
% Global subsumptions: 461
% 
% Simple tautologies: 496
% Equational tautologies: 3911
% Forward subsumptions: 113748
% Bw demodulations to eq. taut.: 108
% 
% Unit resulting resolution: 57881
% Factoring: 1
% Forward superposition: 40033
% Backward superposition: 19277
% Self superposition: 375
% Equality factoring: 310
% Equality resolution: 12
% 
% Split clauses: 742
% Split components: 1499
% SAT solver clauses: 11982
% SAT solver unit clauses: 2389
% SAT solver binary clauses: 2708
% 
% Sat splits: 742
% Sat splitting refutations: 1058
% 
% Memory used [KB]: 60894
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.integral
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.integral
% SZS output start Proof for thm.integral
fof(f43784,plain,(
  $false),
  inference(subsumption_resolution,[],[f43783,f398])).
fof(f398,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.integral.INTEGRABLE_SUBINTERVAL',aHL_TRUTH)).
fof(f43783,plain,(
  ~p(s(bool,'T_0'))),
  inference(forward_demodulation,[],[f43782,f5303])).
fof(f5303,plain,(
  s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK16)))),
  inference(resolution,[],[f5210,f381])).
fof(f381,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f221])).
fof(f221,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X13] : (s(bool,'T_0') = s(bool,X13) <=> p(s(bool,X13)))),
  file('thm.integral.INTEGRABLE_SUBINTERVAL','thm.bool.EQ_CLAUSES|split|1')).
fof(f5210,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK16))))),
  inference(subsumption_resolution,[],[f5189,f367])).
fof(f367,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f214])).
fof(f214,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,'T_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f213])).
fof(f213,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (~p(s(bool,'T_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0] : ((~p(s(bool,'T_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(ennf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0] : ((p(s(bool,'T_0')) => p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X13] : ((p(s(bool,'T_0')) => p(s(bool,X13))) <=> p(s(bool,X13)))),
  file('thm.integral.INTEGRABLE_SUBINTERVAL','thm.bool.IMP_CLAUSES|split|0')).
fof(f5189,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK16))))),
  inference(superposition,[],[f407,f600])).
fof(f600,plain,(
  s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK15),s('type.realax.real',sK16)))),
  inference(resolution,[],[f227,f381])).
fof(f227,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK15),s('type.realax.real',sK16))))),
  inference(cnf_transformation,[],[f149])).
fof(f149,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK15),s('type.realax.real',sK16)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16),s('type.realax.real',sK14)))) & p(s(bool,'const.integral.integrable_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK13),s('type.realax.real',sK14))),s(fun('type.realax.real','type.realax.real'),sK12)))) & ~p(s(bool,'const.integral.integrable_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK15),s('type.realax.real',sK16))),s(fun('type.realax.real','type.realax.real'),sK12))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13,sK14,sK15,sK16])],[f101])).
fof(f101,plain,(
  ? [X0,X1,X2,X3,X4] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X3)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',X4)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real',X2)))) & p(s(bool,'const.integral.integrable_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s(fun('type.realax.real','type.realax.real'),X0)))) & ~p(s(bool,'const.integral.integrable_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X3),s('type.realax.real',X4))),s(fun('type.realax.real','type.realax.real'),X0)))))),
  inference(flattening,[],[f100])).
fof(f100,plain,(
  ? [X0,X1,X2,X3,X4] : ((p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X3)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',X4)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real',X2)))) & p(s(bool,'const.integral.integrable_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s(fun('type.realax.real','type.realax.real'),X0))))) & ~p(s(bool,'const.integral.integrable_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X3),s('type.realax.real',X4))),s(fun('type.realax.real','type.realax.real'),X0)))))),
  inference(ennf_transformation,[],[f42])).
fof(f42,plain,(
  ~! [X0,X1,X2,X3,X4] : ((p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X3)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',X4)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real',X2)))) & p(s(bool,'const.integral.integrable_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s(fun('type.realax.real','type.realax.real'),X0))))) => p(s(bool,'const.integral.integrable_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X3),s('type.realax.real',X4))),s(fun('type.realax.real','type.realax.real'),X0)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8,X9,X10] : ((p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real',X9)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real',X10)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X10),s('type.realax.real',X8)))) & p(s(bool,'const.integral.integrable_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X7),s('type.realax.real',X8))),s(fun('type.realax.real','type.realax.real'),X6))))) => p(s(bool,'const.integral.integrable_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X9),s('type.realax.real',X10))),s(fun('type.realax.real','type.realax.real'),X6)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8,X9,X10] : ((p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real',X9)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real',X10)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X10),s('type.realax.real',X8)))) & p(s(bool,'const.integral.integrable_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X7),s('type.realax.real',X8))),s(fun('type.realax.real','type.realax.real'),X6))))) => p(s(bool,'const.integral.integrable_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X9),s('type.realax.real',X10))),s(fun('type.realax.real','type.realax.real'),X6)))))),
  file('thm.integral.INTEGRABLE_SUBINTERVAL','thm.integral.INTEGRABLE_SUBINTERVAL')).
fof(f407,plain,(
  ( ! [X0] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK15),s('type.realax.real',X0)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X0))))) )),
  inference(resolution,[],[f226,f399])).
fof(f399,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X2)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1))))) )),
  inference(cnf_transformation,[],[f130])).
fof(f130,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X2)))))),
  inference(flattening,[],[f129])).
fof(f129,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X2)))))),
  inference(ennf_transformation,[],[f97])).
fof(f97,plain,(
  ! [X0,X1,X2] : ((p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2))))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X2)))))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X15,X16,X20] : ((p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X15),s('type.realax.real',X16)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X16),s('type.realax.real',X20))))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X15),s('type.realax.real',X20)))))),
  file('thm.integral.INTEGRABLE_SUBINTERVAL','thm.real.REAL_LE_TRANS')).
fof(f226,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15))))),
  inference(cnf_transformation,[],[f149])).
fof(f43782,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK16))))),
  inference(subsumption_resolution,[],[f43781,f371])).
fof(f371,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f216])).
fof(f216,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f215])).
fof(f215,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X13] : ((p(s(bool,'F_0')) => p(s(bool,X13))) <=> p(s(bool,'T_0')))),
  file('thm.integral.INTEGRABLE_SUBINTERVAL','thm.bool.IMP_CLAUSES|split|2')).
fof(f43781,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK16))))),
  inference(forward_demodulation,[],[f43668,f715])).
fof(f715,plain,(
  s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16),s('type.realax.real',sK14)))),
  inference(resolution,[],[f228,f381])).
fof(f228,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16),s('type.realax.real',sK14))))),
  inference(cnf_transformation,[],[f149])).
fof(f43668,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16),s('type.realax.real',sK14)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK16))))),
  inference(resolution,[],[f33209,f746])).
fof(f746,plain,(
  ( ! [X0] : (p(s(bool,'const.integral.integrable_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK13),s('type.realax.real',X0))),s(fun('type.realax.real','type.realax.real'),sK12)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',sK14)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X0))))) )),
  inference(resolution,[],[f229,f232])).
fof(f232,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.integral.integrable_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X3))),s(fun('type.realax.real','type.realax.real'),X0)))) | ~p(s(bool,'const.integral.integrable_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s(fun('type.realax.real','type.realax.real'),X0)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',X2)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X3))))) )),
  inference(cnf_transformation,[],[f105])).
fof(f105,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',X2)))) | ~p(s(bool,'const.integral.integrable_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s(fun('type.realax.real','type.realax.real'),X0)))) | p(s(bool,'const.integral.integrable_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X3))),s(fun('type.realax.real','type.realax.real'),X0)))))),
  inference(flattening,[],[f104])).
fof(f104,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',X2)))) | ~p(s(bool,'const.integral.integrable_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s(fun('type.realax.real','type.realax.real'),X0))))) | p(s(bool,'const.integral.integrable_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X3))),s(fun('type.realax.real','type.realax.real'),X0)))))),
  inference(ennf_transformation,[],[f44])).
fof(f44,plain,(
  ! [X0,X1,X2,X3] : ((p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X3)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',X2)))) & p(s(bool,'const.integral.integrable_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s(fun('type.realax.real','type.realax.real'),X0))))) => p(s(bool,'const.integral.integrable_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X3))),s(fun('type.realax.real','type.realax.real'),X0)))))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X6,X7,X8,X9] : ((p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real',X9)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real',X8)))) & p(s(bool,'const.integral.integrable_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X7),s('type.realax.real',X8))),s(fun('type.realax.real','type.realax.real'),X6))))) => p(s(bool,'const.integral.integrable_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X7),s('type.realax.real',X9))),s(fun('type.realax.real','type.realax.real'),X6)))))),
  file('thm.integral.INTEGRABLE_SUBINTERVAL','thm.integral.INTEGRABLE_SUBINTERVAL_LEFT')).
fof(f229,plain,(
  p(s(bool,'const.integral.integrable_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK13),s('type.realax.real',sK14))),s(fun('type.realax.real','type.realax.real'),sK12))))),
  inference(cnf_transformation,[],[f149])).
fof(f33209,plain,(
  ~p(s(bool,'const.integral.integrable_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK13),s('type.realax.real',sK16))),s(fun('type.realax.real','type.realax.real'),sK12))))),
  inference(subsumption_resolution,[],[f33181,f371])).
fof(f33181,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.integral.integrable_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK13),s('type.realax.real',sK16))),s(fun('type.realax.real','type.realax.real'),sK12))))),
  inference(superposition,[],[f974,f491])).
fof(f491,plain,(
  s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15)))),
  inference(resolution,[],[f226,f381])).
fof(f974,plain,(
  ( ! [X1] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',sK15)))) | ~p(s(bool,'const.integral.integrable_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',sK16))),s(fun('type.realax.real','type.realax.real'),sK12))))) )),
  inference(subsumption_resolution,[],[f872,f227])).
fof(f872,plain,(
  ( ! [X1] : (~p(s(bool,'const.integral.integrable_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',sK16))),s(fun('type.realax.real','type.realax.real'),sK12)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK15),s('type.realax.real',sK16)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',sK15))))) )),
  inference(resolution,[],[f230,f231])).
fof(f231,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.integral.integrable_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X3),s('type.realax.real',X2))),s(fun('type.realax.real','type.realax.real'),X0)))) | ~p(s(bool,'const.integral.integrable_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s(fun('type.realax.real','type.realax.real'),X0)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',X2)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X3))))) )),
  inference(cnf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',X2)))) | ~p(s(bool,'const.integral.integrable_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s(fun('type.realax.real','type.realax.real'),X0)))) | p(s(bool,'const.integral.integrable_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X3),s('type.realax.real',X2))),s(fun('type.realax.real','type.realax.real'),X0)))))),
  inference(flattening,[],[f102])).
fof(f102,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',X2)))) | ~p(s(bool,'const.integral.integrable_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s(fun('type.realax.real','type.realax.real'),X0))))) | p(s(bool,'const.integral.integrable_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X3),s('type.realax.real',X2))),s(fun('type.realax.real','type.realax.real'),X0)))))),
  inference(ennf_transformation,[],[f43])).
fof(f43,plain,(
  ! [X0,X1,X2,X3] : ((p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X3)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',X2)))) & p(s(bool,'const.integral.integrable_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s(fun('type.realax.real','type.realax.real'),X0))))) => p(s(bool,'const.integral.integrable_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X3),s('type.realax.real',X2))),s(fun('type.realax.real','type.realax.real'),X0)))))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X6,X7,X8,X9] : ((p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real',X9)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real',X8)))) & p(s(bool,'const.integral.integrable_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X7),s('type.realax.real',X8))),s(fun('type.realax.real','type.realax.real'),X6))))) => p(s(bool,'const.integral.integrable_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X9),s('type.realax.real',X8))),s(fun('type.realax.real','type.realax.real'),X6)))))),
  file('thm.integral.INTEGRABLE_SUBINTERVAL','thm.integral.INTEGRABLE_SUBINTERVAL_RIGHT')).
fof(f230,plain,(
  ~p(s(bool,'const.integral.integrable_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK15),s('type.realax.real',sK16))),s(fun('type.realax.real','type.realax.real'),sK12))))),
  inference(cnf_transformation,[],[f149])).
% SZS output end Proof for thm.integral
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 1154
% Passive clauses: 7542
% Generated clauses: 43307
% Final active clauses: 1133
% Final passive clauses: 5866
% Input formulas: 40
% Initial clauses: 177
% 
% Selected by SInE selection: 40
% SInE iterations: 1
% 
% Duplicate literals: 2215
% Trivial inequalities: 95
% Fw subsumption resolutions: 6921
% Fw demodulations: 6564
% Bw demodulations: 15
% 
% Simple tautologies: 3212
% Equational tautologies: 116
% Forward subsumptions: 16568
% 
% Binary resolution: 23484
% Factoring: 52
% Forward superposition: 1909
% Backward superposition: 2009
% Self superposition: 2
% Equality factoring: 2
% Equality resolution: 9
% 
% Split clauses: 616
% Split components: 1520
% SAT solver clauses: 702
% SAT solver unit clauses: 27
% SAT solver binary clauses: 93
% 
% Sat splits: 616
% Sat splitting refutations: 29
% 
% Memory used [KB]: 7419
% Time elapsed: 0.609 s
% ------------------------------
% ------------------------------
% Success in time 12.326 s
