Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.numRing
% (87779)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 155
% Passive clauses: 3829
% Generated clauses: 19262
% Final active clauses: 147
% Final passive clauses: 3542
% Input formulas: 118
% Initial clauses: 123
% 
% Duplicate literals: 543
% Trivial inequalities: 312
% Fw subsumption resolutions: 895
% Fw demodulations: 1124
% Bw demodulations: 4
% 
% Simple tautologies: 320
% Equational tautologies: 22
% Forward subsumptions: 11887
% Fw demodulations to eq. taut.: 12
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 31
% Forward superposition: 8800
% Backward superposition: 7081
% Self superposition: 208
% Equality factoring: 54
% Equality resolution: 26
% 
% Split clauses: 219
% Split components: 438
% SAT solver clauses: 432
% SAT solver unit clauses: 5
% SAT solver binary clauses: 100
% 
% Sat splits: 219
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 432
% TWLsolver calls for satisfiability: 54
% 
% Memory used [KB]: 20084
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.numRing
% (87795)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 130
% Passive clauses: 1338
% Generated clauses: 16061
% Final active clauses: 118
% Final passive clauses: 1064
% Input formulas: 118
% Initial clauses: 123
% 
% Duplicate literals: 583
% Trivial inequalities: 204
% Fw subsumption resolutions: 566
% Fw demodulations: 1025
% Bw demodulations: 121
% Condensations: 891
% 
% Simple tautologies: 176
% Equational tautologies: 25
% Forward subsumptions: 9131
% Fw demodulations to eq. taut.: 4
% Bw demodulations to eq. taut.: 26
% 
% Binary resolution: 43
% Forward superposition: 6133
% Backward superposition: 4713
% Self superposition: 213
% Equality factoring: 74
% Equality resolution: 18
% 
% Split clauses: 111
% Split components: 222
% SAT solver clauses: 2077
% SAT solver unit clauses: 235
% SAT solver binary clauses: 50
% 
% Sat splits: 111
% Sat splitting refutations: 30
% 
% TWLsolver clauses: 2077
% TWLsolver calls for satisfiability: 264
% 
% Memory used [KB]: 22643
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.numRing
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.numRing
% SZS output start Proof for thm.numRing
fof(f130448,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f130300])).
fof(f130300,plain,(
  s('type.canonical.canonical_sum'('type.num.num'),sK0) != s('type.canonical.canonical_sum'('type.num.num'),sK0)),
  inference(superposition,[],[f307,f129902])).
fof(f129902,plain,(
  ( ! [X2] : (s('type.canonical.canonical_sum'('type.num.num'),X2) = s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_canonical_sum_merge_2'(s('type.canonical.canonical_sum'('type.num.num'),X2),s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0')))) )),
  inference(superposition,[],[f30144,f335])).
fof(f335,plain,(
  ( ! [X2,X0,X1] : (s('type.canonical.canonical_sum'(X0),X2) = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_merge_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))) )),
  inference(cnf_transformation,[],[f134])).
fof(f134,plain,(
  ! [X0,X1,X2] : s('type.canonical.canonical_sum'(X0),X2) = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_merge_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))),
  inference(flattening,[],[f133])).
fof(f133,plain,(
  ! [X0] : ! [X1,X2] : s('type.canonical.canonical_sum'(X0),X2) = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_merge_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X7] : ! [X12,X6] : s('type.canonical.canonical_sum'(X7),X6) = s('type.canonical.canonical_sum'(X7),'const.canonical.canonical_sum_merge_3'(s('type.semi_ring.semi_ring'(X7),X12),s('type.canonical.canonical_sum'(X7),X6),s('type.canonical.canonical_sum'(X7),'const.canonical.Nil_monom_0')))),
  file('thm.numRing.num_ring_thms|split|14','thm.canonical.canonical_sum_merge_def|split|4')).
fof(f30144,plain,(
  ( ! [X0,X1] : (s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_canonical_sum_merge_2'(s('type.canonical.canonical_sum'('type.num.num'),X0),s('type.canonical.canonical_sum'('type.num.num'),X1))) = s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.canonical_sum_merge_3'(s('type.semi_ring.semi_ring'('type.num.num'),'const.semi_ring.recordtype|46|semi_ring_4'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.+_0'),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.*_0'))),s('type.canonical.canonical_sum'('type.num.num'),X0),s('type.canonical.canonical_sum'('type.num.num'),X1)))) )),
  inference(forward_demodulation,[],[f30101,f421])).
fof(f421,plain,(
  ( ! [X0,X1] : (s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_canonical_sum_merge_2'(s('type.canonical.canonical_sum'('type.num.num'),X0),s('type.canonical.canonical_sum'('type.num.num'),X1))) = s('type.canonical.canonical_sum'('type.num.num'),happ(s(fun('type.canonical.canonical_sum'('type.num.num'),'type.canonical.canonical_sum'('type.num.num')),happ(s(fun('type.canonical.canonical_sum'('type.num.num'),fun('type.canonical.canonical_sum'('type.num.num'),'type.canonical.canonical_sum'('type.num.num'))),'const.numRing.num_canonical_sum_merge_0'),s('type.canonical.canonical_sum'('type.num.num'),X0))),s('type.canonical.canonical_sum'('type.num.num'),X1)))) )),
  inference(cnf_transformation,[],[f287])).
fof(f287,plain,(
  ! [X0,X1] : s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_canonical_sum_merge_2'(s('type.canonical.canonical_sum'('type.num.num'),X0),s('type.canonical.canonical_sum'('type.num.num'),X1))) = s('type.canonical.canonical_sum'('type.num.num'),happ(s(fun('type.canonical.canonical_sum'('type.num.num'),'type.canonical.canonical_sum'('type.num.num')),happ(s(fun('type.canonical.canonical_sum'('type.num.num'),fun('type.canonical.canonical_sum'('type.num.num'),'type.canonical.canonical_sum'('type.num.num'))),'const.numRing.num_canonical_sum_merge_0'),s('type.canonical.canonical_sum'('type.num.num'),X0))),s('type.canonical.canonical_sum'('type.num.num'),X1)))),
  inference(rectify,[],[f116])).
fof(f116,axiom,(
  ! [X45,X46] : s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_canonical_sum_merge_2'(s('type.canonical.canonical_sum'('type.num.num'),X45),s('type.canonical.canonical_sum'('type.num.num'),X46))) = s('type.canonical.canonical_sum'('type.num.num'),happ(s(fun('type.canonical.canonical_sum'('type.num.num'),'type.canonical.canonical_sum'('type.num.num')),happ(s(fun('type.canonical.canonical_sum'('type.num.num'),fun('type.canonical.canonical_sum'('type.num.num'),'type.canonical.canonical_sum'('type.num.num'))),'const.numRing.num_canonical_sum_merge_0'),s('type.canonical.canonical_sum'('type.num.num'),X45))),s('type.canonical.canonical_sum'('type.num.num'),X46)))),
  file('thm.numRing.num_ring_thms|split|14','ARITY_0_2_const.numRing.num_canonical_sum_merge')).
fof(f30101,plain,(
  ( ! [X0,X1] : (s('type.canonical.canonical_sum'('type.num.num'),happ(s(fun('type.canonical.canonical_sum'('type.num.num'),'type.canonical.canonical_sum'('type.num.num')),happ(s(fun('type.canonical.canonical_sum'('type.num.num'),fun('type.canonical.canonical_sum'('type.num.num'),'type.canonical.canonical_sum'('type.num.num'))),'const.numRing.num_canonical_sum_merge_0'),s('type.canonical.canonical_sum'('type.num.num'),X0))),s('type.canonical.canonical_sum'('type.num.num'),X1))) = s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.canonical_sum_merge_3'(s('type.semi_ring.semi_ring'('type.num.num'),'const.semi_ring.recordtype|46|semi_ring_4'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.+_0'),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.*_0'))),s('type.canonical.canonical_sum'('type.num.num'),X0),s('type.canonical.canonical_sum'('type.num.num'),X1)))) )),
  inference(superposition,[],[f368,f318])).
fof(f318,plain,(
  s(fun('type.canonical.canonical_sum'('type.num.num'),fun('type.canonical.canonical_sum'('type.num.num'),'type.canonical.canonical_sum'('type.num.num'))),'const.numRing.num_canonical_sum_merge_0') = s(fun('type.canonical.canonical_sum'('type.num.num'),fun('type.canonical.canonical_sum'('type.num.num'),'type.canonical.canonical_sum'('type.num.num'))),'const.canonical.canonical_sum_merge_1'(s('type.semi_ring.semi_ring'('type.num.num'),'const.semi_ring.recordtype|46|semi_ring_4'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.+_0'),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.*_0')))))),
  inference(cnf_transformation,[],[f74])).
fof(f74,axiom,(
  s(fun('type.canonical.canonical_sum'('type.num.num'),fun('type.canonical.canonical_sum'('type.num.num'),'type.canonical.canonical_sum'('type.num.num'))),'const.numRing.num_canonical_sum_merge_0') = s(fun('type.canonical.canonical_sum'('type.num.num'),fun('type.canonical.canonical_sum'('type.num.num'),'type.canonical.canonical_sum'('type.num.num'))),'const.canonical.canonical_sum_merge_1'(s('type.semi_ring.semi_ring'('type.num.num'),'const.semi_ring.recordtype|46|semi_ring_4'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.+_0'),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.*_0')))))),
  file('thm.numRing.num_ring_thms|split|14','thm.numRing.num_canonical_sum_merge_def')).
fof(f368,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_merge_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_merge_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2))),s('type.canonical.canonical_sum'(X0),X3)))) )),
  inference(cnf_transformation,[],[f198])).
fof(f198,plain,(
  ! [X0,X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_merge_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_merge_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2))),s('type.canonical.canonical_sum'(X0),X3)))),
  inference(flattening,[],[f197])).
fof(f197,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_merge_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_merge_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2))),s('type.canonical.canonical_sum'(X0),X3)))),
  inference(rectify,[],[f113])).
fof(f113,axiom,(
  ! [X7] : ! [X45,X46,X47] : s('type.canonical.canonical_sum'(X7),'const.canonical.canonical_sum_merge_3'(s('type.semi_ring.semi_ring'(X7),X45),s('type.canonical.canonical_sum'(X7),X46),s('type.canonical.canonical_sum'(X7),X47))) = s('type.canonical.canonical_sum'(X7),happ(s(fun('type.canonical.canonical_sum'(X7),'type.canonical.canonical_sum'(X7)),happ(s(fun('type.canonical.canonical_sum'(X7),fun('type.canonical.canonical_sum'(X7),'type.canonical.canonical_sum'(X7))),'const.canonical.canonical_sum_merge_1'(s('type.semi_ring.semi_ring'(X7),X45))),s('type.canonical.canonical_sum'(X7),X46))),s('type.canonical.canonical_sum'(X7),X47)))),
  file('thm.numRing.num_ring_thms|split|14','ARITY_1_3_const.canonical.canonical_sum_merge')).
fof(f307,plain,(
  s('type.canonical.canonical_sum'('type.num.num'),sK0) != s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_canonical_sum_merge_2'(s('type.canonical.canonical_sum'('type.num.num'),sK0),s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0')))),
  inference(cnf_transformation,[],[f301])).
fof(f301,plain,(
  s('type.canonical.canonical_sum'('type.num.num'),sK0) != s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_canonical_sum_merge_2'(s('type.canonical.canonical_sum'('type.num.num'),sK0),s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0')))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f298])).
fof(f298,plain,(
  ? [X0] : s('type.canonical.canonical_sum'('type.num.num'),X0) != s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_canonical_sum_merge_2'(s('type.canonical.canonical_sum'('type.num.num'),X0),s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0')))),
  inference(ennf_transformation,[],[f120])).
fof(f120,plain,(
  ~! [X0] : s('type.canonical.canonical_sum'('type.num.num'),X0) = s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_canonical_sum_merge_2'(s('type.canonical.canonical_sum'('type.num.num'),X0),s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0')))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : s('type.canonical.canonical_sum'('type.num.num'),X6) = s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_canonical_sum_merge_2'(s('type.canonical.canonical_sum'('type.num.num'),X6),s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0')))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : s('type.canonical.canonical_sum'('type.num.num'),X6) = s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_canonical_sum_merge_2'(s('type.canonical.canonical_sum'('type.num.num'),X6),s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0')))),
  file('thm.numRing.num_ring_thms|split|14','thm.numRing.num_ring_thms|split|14')).
% SZS output end Proof for thm.numRing
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 558
% Passive clauses: 10297
% Generated clauses: 133027
% Final active clauses: 235
% Final passive clauses: 1548
% Input formulas: 118
% Initial clauses: 123
% Discarded non-redundant clauses: 6711
% 
% Duplicate literals: 3733
% Trivial inequalities: 1231
% Fw subsumption resolutions: 3344
% Fw demodulations: 14731
% Bw demodulations: 2110
% Condensations: 5221
% Global subsumptions: 207
% 
% Simple tautologies: 630
% Equational tautologies: 271
% Forward subsumptions: 91925
% Fw demodulations to eq. taut.: 25
% Bw demodulations to eq. taut.: 212
% 
% Unit resulting resolution: 2797
% Forward superposition: 49814
% Backward superposition: 44388
% Self superposition: 771
% Equality factoring: 232
% Equality resolution: 50
% 
% Split clauses: 343
% Split components: 691
% SAT solver clauses: 11426
% SAT solver unit clauses: 1255
% SAT solver binary clauses: 2780
% 
% Sat splits: 343
% Sat splitting refutations: 234
% 
% Memory used [KB]: 161831
% Time elapsed: 7.479 s
% ------------------------------
% ------------------------------
% Success in time 8.837 s
