Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.real
% (54095)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1158
% Passive clauses: 6810
% Generated clauses: 34101
% Final active clauses: 898
% Final passive clauses: 5029
% Input formulas: 107
% Initial clauses: 249
% 
% Duplicate literals: 251
% Trivial inequalities: 60
% Fw subsumption resolutions: 2325
% Fw demodulations: 3290
% Bw demodulations: 230
% 
% Simple tautologies: 220
% Equational tautologies: 34
% Forward subsumptions: 21114
% Fw demodulations to eq. taut.: 9
% 
% Binary resolution: 7306
% Factoring: 1390
% Forward superposition: 8961
% Backward superposition: 9855
% Self superposition: 31
% Equality factoring: 2
% Equality resolution: 10
% 
% Split clauses: 275
% Split components: 550
% SAT solver clauses: 524
% SAT solver unit clauses: 16
% SAT solver binary clauses: 55
% 
% Sat splits: 275
% Sat splitting refutations: 6
% 
% TWLsolver clauses: 524
% TWLsolver calls for satisfiability: 485
% 
% Memory used [KB]: 17782
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.real
% (54140)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 654
% Passive clauses: 3373
% Generated clauses: 28893
% Final active clauses: 564
% Final passive clauses: 2426
% Input formulas: 107
% Initial clauses: 249
% 
% Duplicate literals: 283
% Trivial inequalities: 164
% Fw subsumption resolutions: 1799
% Fw demodulations: 4993
% Bw demodulations: 381
% Condensations: 261
% Global subsumptions: 16
% 
% Simple tautologies: 357
% Equational tautologies: 25
% Forward subsumptions: 13624
% Fw demodulations to eq. taut.: 10
% 
% Binary resolution: 2593
% Factoring: 183
% Forward superposition: 8432
% Backward superposition: 6004
% Self superposition: 61
% Equality factoring: 36
% Equality resolution: 23
% 
% Split clauses: 68
% Split components: 136
% SAT solver clauses: 4236
% SAT solver unit clauses: 293
% SAT solver binary clauses: 82
% SAT solver learnt clauses: 32
% SAT solver learnt literals: 12
% 
% Sat splits: 68
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 4172
% TWLsolver calls for satisfiability: 1734
% 
% Memory used [KB]: 20468
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.real
% (54199)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 3377
% Passive clauses: 22238
% Generated clauses: 307135
% Final active clauses: 2894
% Final passive clauses: 7748
% Input formulas: 107
% Initial clauses: 249
% Discarded non-redundant clauses: 17853
% 
% Duplicate literals: 401
% Trivial inequalities: 185
% Fw subsumption resolutions: 3063
% Fw demodulations: 65495
% Bw demodulations: 2681
% Condensations: 322
% Global subsumptions: 13
% 
% Simple tautologies: 200
% Equational tautologies: 714
% Forward subsumptions: 211733
% Fw demodulations to eq. taut.: 608
% Bw demodulations to eq. taut.: 50
% 
% Unit resulting resolution: 45690
% Factoring: 4
% Forward superposition: 126648
% Backward superposition: 61618
% Self superposition: 105
% Equality factoring: 36
% Equality resolution: 37
% 
% Split clauses: 59
% Split components: 123
% SAT solver clauses: 22415
% SAT solver unit clauses: 12346
% SAT solver binary clauses: 5908
% 
% Sat splits: 59
% Sat splitting refutations: 44
% 
% Memory used [KB]: 145711
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.real
% (54417)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1941
% Passive clauses: 36222
% Generated clauses: 201468
% Final active clauses: 1881
% Final passive clauses: 32526
% Input formulas: 107
% Initial clauses: 182
% 
% Selected by SInE selection: 52
% SInE iterations: 1
% 
% Duplicate literals: 7667
% Trivial inequalities: 174
% Fw subsumption resolutions: 24117
% Fw demodulations: 4542
% Bw demodulations: 37
% 
% Simple tautologies: 4393
% Equational tautologies: 141
% Forward subsumptions: 113322
% 
% Binary resolution: 147807
% Factoring: 492
% Forward superposition: 12233
% Backward superposition: 3023
% Self superposition: 14
% Equality resolution: 46
% 
% Split clauses: 11478
% Split components: 23904
% SAT solver clauses: 12103
% SAT solver unit clauses: 34
% SAT solver binary clauses: 7030
% 
% Sat splits: 11478
% Sat splitting refutations: 49
% 
% Memory used [KB]: 46950
% Time elapsed: 4.0000 s
% ------------------------------
% ------------------------------
% remaining time: 142 next slice time: 6
lrs+1002_2:3_br=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:sas=minisat:stl=30:sac=on:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none:sp=reverse_arity:urr=on_4 on thm.real
% (54547)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 691
% Passive clauses: 4699
% Generated clauses: 34679
% Final active clauses: 648
% Final passive clauses: 3586
% Input formulas: 107
% Initial clauses: 249
% 
% Duplicate literals: 237
% Trivial inequalities: 88
% Fw subsumption resolutions: 1077
% Fw demodulations: 4129
% Bw demodulations: 389
% 
% Simple tautologies: 300
% Equational tautologies: 32
% Forward subsumptions: 23787
% Fw demodulations to eq. taut.: 112
% Bw demodulations to eq. taut.: 14
% 
% Unit resulting resolution: 10665
% Factoring: 19
% Forward superposition: 11187
% Backward superposition: 6423
% Self superposition: 87
% Equality factoring: 6
% Equality resolution: 8
% 
% Split clauses: 6
% Split components: 12
% SAT solver clauses: 4837
% SAT solver unit clauses: 1768
% SAT solver binary clauses: 722
% 
% Sat splits: 6
% Sat splitting refutations: 130
% 
% Memory used [KB]: 17654
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 135 next slice time: 15
lrs+1011_12_bs=on:bsr=unit_only:cond=on:gs=on:gsaa=from_current:gsssp=full:nwc=1.1:sas=minisat:stl=60:sos=all:sac=on:sdd=off:sser=off:sfr=on:ssfp=100000:ssfq=1.2:ssnc=none:sp=reverse_arity:updr=off_14 on thm.real
% (54569)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 766
% Passive clauses: 8761
% Generated clauses: 46394
% Final active clauses: 691
% Final passive clauses: 6106
% Input formulas: 107
% Initial clauses: 249
% 
% Duplicate literals: 587
% Trivial inequalities: 202
% Fw subsumption resolutions: 2761
% Bw subsumption resolutions: 62
% Fw demodulations: 12246
% Bw demodulations: 257
% Condensations: 1355
% Global subsumptions: 7
% 
% Simple tautologies: 1225
% Equational tautologies: 1
% Forward subsumptions: 19317
% Backward subsumptions: 1826
% 
% Binary resolution: 6758
% Forward superposition: 14961
% Backward superposition: 5881
% Self superposition: 9
% Equality resolution: 9
% 
% Split clauses: 5
% Split components: 10
% SAT solver clauses: 8785
% SAT solver unit clauses: 58
% SAT solver binary clauses: 747
% 
% Sat splits: 5
% Sat splitting refutations: 9
% 
% Memory used [KB]: 24050
% Time elapsed: 1.600 s
% ------------------------------
% ------------------------------
% remaining time: 119 next slice time: 9
lrs+10_8:1_bd=preordered:bs=on:ccuc=first:fsr=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:nicw=on:sas=minisat:stl=120:sos=on:sscc=on:sser=off:sfr=on:ssfp=4000:ssfq=1.0:ssnc=none:sp=reverse_arity:urr=on_7 on thm.real
% (54624)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 641
% Passive clauses: 4686
% Generated clauses: 38409
% Final active clauses: 499
% Final passive clauses: 2845
% Input formulas: 107
% Initial clauses: 249
% 
% Duplicate literals: 83
% Trivial inequalities: 84
% Fw demodulations: 8920
% Bw demodulations: 1333
% Global subsumptions: 367
% 
% Simple tautologies: 304
% Equational tautologies: 21
% Forward subsumptions: 22486
% Backward subsumptions: 135
% Fw demodulations to eq. taut.: 171
% Bw demodulations to eq. taut.: 14
% 
% Binary resolution: 2076
% Unit resulting resolution: 4545
% Factoring: 2
% Forward superposition: 12404
% Backward superposition: 7575
% Self superposition: 30
% Equality resolution: 12
% 
% Split clauses: 433
% Split components: 866
% SAT solver clauses: 7160
% SAT solver unit clauses: 1412
% SAT solver binary clauses: 3017
% 
% Sat splits: 433
% Sat splitting refutations: 353
% 
% Memory used [KB]: 28144
% Time elapsed: 1.0000 s
% ------------------------------
% ------------------------------
% remaining time: 110 next slice time: 85
lrs+1_5:4_cond=on:fsr=off:fde=none:gs=on:gsem=on:lwlo=on:nm=64:nwc=1:stl=60:sos=all:spl=off_80 on thm.real
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.real
% SZS output start Proof for thm.real
fof(f96766,plain,(
  $false),
  inference(resolution,[],[f96231,f354])).
fof(f354,plain,(
  ( ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))))))) )),
  inference(cnf_transformation,[],[f133])).
fof(f133,plain,(
  ! [X0] : p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))))))),
  inference(rectify,[],[f97])).
fof(f97,axiom,(
  ! [X6] : p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X6))))))),
  file('thm.real.lt_int|split|2','thm.real.REAL_POS')).
fof(f96231,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK2))))))),
  inference(forward_demodulation,[],[f96132,f366])).
fof(f366,plain,(
  ( ! [X0] : (s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(cnf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0] : s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(rectify,[],[f96])).
fof(f96,axiom,(
  ! [X13] : s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X13))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X13))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.real.lt_int|split|2','thm.real.REAL_NEG_LE|48|')).
fof(f96132,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK2))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(resolution,[],[f94867,f483])).
fof(f483,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0))))) )),
  inference(cnf_transformation,[],[f302])).
fof(f302,plain,(
  ! [X0,X1] : ((p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0))))) & (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1))))))),
  inference(nnf_transformation,[],[f211])).
fof(f211,plain,(
  ! [X0,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) <=> p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0)))))),
  inference(flattening,[],[f210])).
fof(f210,plain,(
  ! [X0,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) <=> p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0)))))),
  inference(rectify,[],[f92])).
fof(f92,axiom,(
  ! [X13,X15] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X13),s('type.realax.real',X15)))) <=> p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X15),s('type.realax.real',X13)))))),
  file('thm.real.lt_int|split|2','thm.real.REAL_NOT_LT')).
fof(f94867,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK2))))))))),
  inference(resolution,[],[f10525,f324])).
fof(f324,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.real.lt_int|split|2',aHL_TRUTH)).
fof(f10525,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK2))))))))),
  inference(resolution,[],[f961,f354])).
fof(f961,plain,(
  ( ! [X0] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK1)))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK2)))))))) | ~p(s(bool,'T_0'))) )),
  inference(superposition,[],[f498,f857])).
fof(f857,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK1))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK2)))))))),
  inference(trivial_inequality_removal,[],[f826])).
fof(f826,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK1))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK2)))))))),
  inference(superposition,[],[f323,f368])).
fof(f368,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.real.lt_int|split|2',aHL_BOOL_CASES)).
fof(f323,plain,(
  s(bool,'F_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK1))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK2)))))))),
  inference(cnf_transformation,[],[f259])).
fof(f259,plain,(
  s(bool,'F_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK1))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK2)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK1,sK2])],[f233])).
fof(f233,plain,(
  ? [X0,X1] : s(bool,'F_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1)))))))),
  inference(ennf_transformation,[],[f109])).
fof(f109,plain,(
  ~! [X0,X1] : s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X6))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X7)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X6))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X7)))))))),
  file('thm.real.lt_int|split|2','thm.real.lt_int|split|2')).
fof(f498,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X2)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1))))) )),
  inference(cnf_transformation,[],[f254])).
fof(f254,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X2)))))),
  inference(flattening,[],[f253])).
fof(f253,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X2))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X2)))))),
  inference(ennf_transformation,[],[f223])).
fof(f223,plain,(
  ! [X0,X1,X2] : ((p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) & p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X2))))) => p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X2)))))),
  inference(rectify,[],[f93])).
fof(f93,axiom,(
  ! [X13,X15,X18] : ((p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X13),s('type.realax.real',X15)))) & p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X15),s('type.realax.real',X18))))) => p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X13),s('type.realax.real',X18)))))),
  file('thm.real.lt_int|split|2','thm.real.REAL_LET_TRANS')).
% SZS output end Proof for thm.real
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 759
% Passive clauses: 10126
% Generated clauses: 96027
% Final active clauses: 719
% Final passive clauses: 9097
% Input formulas: 107
% Initial clauses: 350
% 
% Duplicate literals: 3343
% Trivial inequalities: 1692
% Fw demodulations: 22825
% Bw demodulations: 413
% Condensations: 4855
% Global subsumptions: 73
% 
% Simple tautologies: 2843
% Equational tautologies: 1
% Forward subsumptions: 50602
% Bw demodulations to eq. taut.: 31
% 
% Binary resolution: 17669
% Factoring: 24
% Forward superposition: 44095
% Backward superposition: 1170
% Equality resolution: 77
% 
% SAT solver clauses: 10463
% SAT solver unit clauses: 362
% SAT solver binary clauses: 1127
% SAT solver learnt clauses: 146
% SAT solver learnt literals: 222
% 
% TWLsolver clauses: 10199
% TWLsolver calls for satisfiability: 284394
% 
% Memory used [KB]: 36204
% Time elapsed: 3.366 s
% ------------------------------
% ------------------------------
% Success in time 22.341 s
