Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.option
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.option
% SZS output start Proof for thm.option
fof(f1592,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f1588,f93,f148,f146,f891,f94,f149,f147,f901])).
fof(f901,plain,(
  $false | ($spl2 | $spl15)),
  inference(subsumption_resolution,[],[f895,f365])).
fof(f365,plain,(
  s(bool,'T_0') != s(bool,sK2) | $spl15),
  inference(cnf_transformation,[],[f365_D])).
fof(f365_D,plain,(
  s(bool,'T_0') != s(bool,sK2) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f895,plain,(
  s(bool,'T_0') = s(bool,sK2) | $spl2),
  inference(resolution,[],[f149,f106])).
fof(f106,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f37])).
fof(f37,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X11] : (s(bool,'T_0') = s(bool,X11) <=> p(s(bool,X11)))),
  file('thm.option.IF_EQUALS_OPTION|split|0','thm.bool.EQ_CLAUSES|split|0')).
fof(f147,plain,(
  s('type.option.option'(sK0),'const.option.NONE_0') != s('type.option.option'(sK0),'const.bool.COND_3'(s(bool,sK2),s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK1))),s('type.option.option'(sK0),'const.option.NONE_0'))) | $spl1),
  inference(cnf_transformation,[],[f147_D])).
fof(f147_D,plain,(
  s('type.option.option'(sK0),'const.option.NONE_0') != s('type.option.option'(sK0),'const.bool.COND_3'(s(bool,sK2),s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK1))),s('type.option.option'(sK0),'const.option.NONE_0'))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f149,plain,(
  p(s(bool,sK2)) | $spl2),
  inference(cnf_transformation,[],[f149_D])).
fof(f149_D,plain,(
  p(s(bool,sK2)) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f94,plain,(
  p(s(bool,sK2)) | s('type.option.option'(sK0),'const.option.NONE_0') != s('type.option.option'(sK0),'const.bool.COND_3'(s(bool,sK2),s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK1))),s('type.option.option'(sK0),'const.option.NONE_0')))),
  inference(cnf_transformation,[],[f71])).
fof(f71,plain,(
  (s('type.option.option'(sK0),'const.option.NONE_0') = s('type.option.option'(sK0),'const.bool.COND_3'(s(bool,sK2),s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK1))),s('type.option.option'(sK0),'const.option.NONE_0'))) | ~p(s(bool,sK2))) & (s('type.option.option'(sK0),'const.option.NONE_0') != s('type.option.option'(sK0),'const.bool.COND_3'(s(bool,sK2),s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK1))),s('type.option.option'(sK0),'const.option.NONE_0'))) | p(s(bool,sK2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f70])).
fof(f70,plain,(
  ? [X0,X1,X2] : ((s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),'const.bool.COND_3'(s(bool,X2),s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1))),s('type.option.option'(X0),'const.option.NONE_0'))) | ~p(s(bool,X2))) & (s('type.option.option'(X0),'const.option.NONE_0') != s('type.option.option'(X0),'const.bool.COND_3'(s(bool,X2),s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1))),s('type.option.option'(X0),'const.option.NONE_0'))) | p(s(bool,X2))))),
  inference(nnf_transformation,[],[f65])).
fof(f65,plain,(
  ? [X0,X1,X2] : (s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),'const.bool.COND_3'(s(bool,X2),s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1))),s('type.option.option'(X0),'const.option.NONE_0'))) <~> ~p(s(bool,X2)))),
  inference(ennf_transformation,[],[f30])).
fof(f30,plain,(
  ~! [X0,X1,X2] : (s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),'const.bool.COND_3'(s(bool,X2),s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1))),s('type.option.option'(X0),'const.option.NONE_0'))) <=> ~p(s(bool,X2)))),
  inference(flattening,[],[f29])).
fof(f29,plain,(
  ~! [X0] : ! [X1,X2] : (s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),'const.bool.COND_3'(s(bool,X2),s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1))),s('type.option.option'(X0),'const.option.NONE_0'))) <=> ~p(s(bool,X2)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : (s('type.option.option'(X6),'const.option.NONE_0') = s('type.option.option'(X6),'const.bool.COND_3'(s(bool,X8),s('type.option.option'(X6),'const.option.SOME_1'(s(X6,X7))),s('type.option.option'(X6),'const.option.NONE_0'))) <=> ~p(s(bool,X8)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : (s('type.option.option'(X6),'const.option.NONE_0') = s('type.option.option'(X6),'const.bool.COND_3'(s(bool,X8),s('type.option.option'(X6),'const.option.SOME_1'(s(X6,X7))),s('type.option.option'(X6),'const.option.NONE_0'))) <=> ~p(s(bool,X8)))),
  file('thm.option.IF_EQUALS_OPTION|split|0','thm.option.IF_EQUALS_OPTION|split|0')).
fof(f891,plain,(
  $false | ($spl1 | $spl3)),
  inference(subsumption_resolution,[],[f890,f133])).
fof(f133,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0,X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f54])).
fof(f54,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X6] : ! [X9,X10] : s(X6,X10) = s(X6,'const.bool.COND_3'(s(bool,'F_0'),s(X6,X9),s(X6,X10)))),
  file('thm.option.IF_EQUALS_OPTION|split|0','thm.bool.COND_CLAUSES|split|1')).
fof(f890,plain,(
  s('type.option.option'(sK0),'const.option.NONE_0') != s('type.option.option'(sK0),'const.bool.COND_3'(s(bool,'F_0'),s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK1))),s('type.option.option'(sK0),'const.option.NONE_0'))) | ($spl1 | $spl3)),
  inference(forward_demodulation,[],[f147,f159])).
fof(f159,plain,(
  s(bool,'F_0') = s(bool,sK2) | $spl3),
  inference(resolution,[],[f110,f148])).
fof(f110,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f40])).
fof(f40,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f39])).
fof(f39,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X11] : (s(bool,'F_0') = s(bool,X11) <=> ~p(s(bool,X11)))),
  file('thm.option.IF_EQUALS_OPTION|split|0','thm.bool.EQ_CLAUSES|split|2')).
fof(f146,plain,(
  s('type.option.option'(sK0),'const.option.NONE_0') = s('type.option.option'(sK0),'const.bool.COND_3'(s(bool,sK2),s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK1))),s('type.option.option'(sK0),'const.option.NONE_0'))) | $spl0),
  inference(cnf_transformation,[],[f146_D])).
fof(f146_D,plain,(
  s('type.option.option'(sK0),'const.option.NONE_0') = s('type.option.option'(sK0),'const.bool.COND_3'(s(bool,sK2),s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK1))),s('type.option.option'(sK0),'const.option.NONE_0'))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f148,plain,(
  ~p(s(bool,sK2)) | $spl3),
  inference(cnf_transformation,[],[f148_D])).
fof(f148_D,plain,(
  ~p(s(bool,sK2)) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f93,plain,(
  ~p(s(bool,sK2)) | s('type.option.option'(sK0),'const.option.NONE_0') = s('type.option.option'(sK0),'const.bool.COND_3'(s(bool,sK2),s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK1))),s('type.option.option'(sK0),'const.option.NONE_0')))),
  inference(cnf_transformation,[],[f71])).
fof(f1588,plain,(
  $false | ($spl0 | $spl14)),
  inference(subsumption_resolution,[],[f1573,f132])).
fof(f132,plain,(
  ( ! [X0,X1] : (s('type.option.option'(X0),'const.option.NONE_0') != s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1)))) )),
  inference(cnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0,X1] : s('type.option.option'(X0),'const.option.NONE_0') != s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1)))),
  inference(flattening,[],[f52])).
fof(f52,plain,(
  ! [X0] : ! [X1] : ~s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1)))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X6] : ! [X7] : ~s('type.option.option'(X6),'const.option.NONE_0') = s('type.option.option'(X6),'const.option.SOME_1'(s(X6,X7)))),
  file('thm.option.IF_EQUALS_OPTION|split|0','thm.option.NOT_NONE_SOME')).
fof(f1573,plain,(
  s('type.option.option'(sK0),'const.option.NONE_0') = s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK1))) | ($spl0 | $spl14)),
  inference(superposition,[],[f1066,f134])).
fof(f134,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f56])).
fof(f56,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X6] : ! [X9,X10] : s(X6,X9) = s(X6,'const.bool.COND_3'(s(bool,'T_0'),s(X6,X9),s(X6,X10)))),
  file('thm.option.IF_EQUALS_OPTION|split|0','thm.bool.COND_CLAUSES|split|0')).
fof(f1066,plain,(
  s('type.option.option'(sK0),'const.option.NONE_0') = s('type.option.option'(sK0),'const.bool.COND_3'(s(bool,'T_0'),s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK1))),s('type.option.option'(sK0),'const.option.NONE_0'))) | ($spl0 | $spl14)),
  inference(forward_demodulation,[],[f146,f364])).
fof(f364,plain,(
  s(bool,'T_0') = s(bool,sK2) | $spl14),
  inference(cnf_transformation,[],[f364_D])).
fof(f364_D,plain,(
  s(bool,'T_0') = s(bool,sK2) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
% SZS output end Proof for thm.option
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 34
% Passive clauses: 317
% Generated clauses: 1454
% Final active clauses: 26
% Final passive clauses: 225
% Input formulas: 27
% Initial clauses: 51
% 
% Duplicate literals: 113
% Trivial inequalities: 26
% Fw subsumption resolutions: 147
% Fw demodulations: 104
% Bw demodulations: 3
% 
% Simple tautologies: 76
% Equational tautologies: 33
% Forward subsumptions: 617
% 
% Binary resolution: 39
% Factoring: 1
% Forward superposition: 275
% Backward superposition: 566
% Self superposition: 52
% Equality factoring: 54
% Equality resolution: 6
% 
% Split clauses: 22
% Split components: 44
% SAT solver clauses: 47
% SAT solver unit clauses: 8
% SAT solver binary clauses: 25
% SAT solver learnt clauses: 2
% SAT solver learnt literals: 1
% 
% Sat splits: 22
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 44
% TWLsolver calls for satisfiability: 15
% 
% Memory used [KB]: 1023
% Time elapsed: 0.036 s
% ------------------------------
% ------------------------------
% Success in time 0.038 s
