Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.int_arith
% (64681)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 662
% Passive clauses: 8444
% Generated clauses: 18862
% Final active clauses: 652
% Final passive clauses: 7378
% Input formulas: 60
% Initial clauses: 215
% 
% Duplicate literals: 423
% Trivial inequalities: 23
% Fw subsumption resolutions: 1670
% Fw demodulations: 658
% Bw demodulations: 9
% 
% Simple tautologies: 274
% Equational tautologies: 57
% Forward subsumptions: 6367
% Fw demodulations to eq. taut.: 15
% 
% Binary resolution: 9207
% Factoring: 1067
% Forward superposition: 2731
% Backward superposition: 2574
% Self superposition: 54
% Equality factoring: 28
% Equality resolution: 5
% 
% Split clauses: 787
% Split components: 1581
% SAT solver clauses: 1247
% SAT solver unit clauses: 20
% SAT solver binary clauses: 226
% 
% Sat splits: 787
% 
% TWLsolver clauses: 1246
% TWLsolver calls for satisfiability: 426
% 
% Memory used [KB]: 10746
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.int_arith
% (64705)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 357
% Passive clauses: 1882
% Generated clauses: 10583
% Final active clauses: 340
% Final passive clauses: 1384
% Input formulas: 60
% Initial clauses: 215
% 
% Duplicate literals: 242
% Trivial inequalities: 51
% Fw subsumption resolutions: 732
% Fw demodulations: 1023
% Bw demodulations: 145
% Condensations: 210
% Global subsumptions: 22
% 
% Simple tautologies: 299
% Equational tautologies: 18
% Forward subsumptions: 3762
% Fw demodulations to eq. taut.: 18
% Bw demodulations to eq. taut.: 2
% 
% Binary resolution: 2007
% Factoring: 102
% Forward superposition: 2074
% Backward superposition: 1647
% Self superposition: 52
% Equality factoring: 28
% Equality resolution: 9
% 
% Split clauses: 37
% Split components: 74
% SAT solver clauses: 2574
% SAT solver unit clauses: 218
% SAT solver binary clauses: 76
% SAT solver learnt clauses: 44
% SAT solver learnt literals: 10
% 
% Sat splits: 37
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 2486
% TWLsolver calls for satisfiability: 1696
% 
% Memory used [KB]: 11769
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.int_arith
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.int_arith
% SZS output start Proof for thm.int_arith
fof(f8688,plain,(
  $false),
  inference(subsumption_resolution,[],[f8687,f521])).
fof(f521,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f300,f297])).
fof(f297,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f8])).
fof(f8,axiom,(
  p(s(bool,'T_0'))),
  file('thm.int_arith.le_context_rwt|49|','thm.bool.TRUTH')).
fof(f300,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f200])).
fof(f200,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f63])).
fof(f63,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f30])).
fof(f30,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.int_arith.le_context_rwt|49|','thm.bool.NOT_CLAUSES|split|1')).
fof(f8687,plain,(
  p(s(bool,'F_0'))),
  inference(forward_demodulation,[],[f8686,f635])).
fof(f635,plain,(
  s(bool,'F_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK12),s('type.integer.int',sK14)))))),
  inference(unit_resulting_resolution,[],[f572,f323])).
fof(f323,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f204])).
fof(f204,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f83])).
fof(f83,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X9] : (s(bool,'F_0') = s(bool,X9) <=> ~p(s(bool,X9)))),
  file('thm.int_arith.le_context_rwt|49|','thm.bool.EQ_CLAUSES|split|2')).
fof(f572,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK12),s('type.integer.int',sK14))))))),
  inference(unit_resulting_resolution,[],[f522,f319])).
fof(f319,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f202])).
fof(f202,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X9] : (s(bool,'T_0') = s(bool,X9) <=> p(s(bool,X9)))),
  file('thm.int_arith.le_context_rwt|49|','thm.bool.EQ_CLAUSES|split|0')).
fof(f522,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK12),s('type.integer.int',sK14)))))),
  inference(backward_demodulation,[],[f379,f296])).
fof(f296,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)))))),
  inference(cnf_transformation,[],[f199])).
fof(f199,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) & s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13,sK14])],[f146])).
fof(f146,plain,(
  ? [X0,X1,X2] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X2),s('type.integer.int',X1)))))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) & s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X2),s('type.integer.int',X0))))))),
  inference(flattening,[],[f145])).
fof(f145,plain,(
  ? [X0,X1,X2] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X2),s('type.integer.int',X1)))))) & (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) & s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X2),s('type.integer.int',X0)))))))),
  inference(ennf_transformation,[],[f62])).
fof(f62,plain,(
  ~! [X0,X1,X2] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X2),s('type.integer.int',X1)))))) => (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) => s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X2),s('type.integer.int',X0)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X8),s('type.integer.int',X7)))))) => (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X7),s('type.integer.int',X6)))) => s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X8),s('type.integer.int',X6)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X8),s('type.integer.int',X7)))))) => (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X7),s('type.integer.int',X6)))) => s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X8),s('type.integer.int',X6)))))))),
  file('thm.int_arith.le_context_rwt|49|','thm.int_arith.le_context_rwt|49|')).
fof(f379,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) )),
  inference(cnf_transformation,[],[f111])).
fof(f111,plain,(
  ! [X0,X1] : s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',X0)))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X6,X7] : s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X6),s('type.integer.int',X7))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X7),s('type.integer.int',X6)))),
  file('thm.int_arith.le_context_rwt|49|','thm.integer.INT_ADD_SYM')).
fof(f8686,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK12),s('type.integer.int',sK14))))))),
  inference(forward_demodulation,[],[f8685,f379])).
fof(f8685,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK14),s('type.integer.int',sK12))))))),
  inference(forward_demodulation,[],[f8684,f3357])).
fof(f3357,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',X19),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X17),s('type.integer.int',X18))))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X16),s('type.integer.int',X19))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X18),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X16),s('type.integer.int',X17)))))))) )),
  inference(superposition,[],[f427,f1451])).
fof(f1451,plain,(
  ( ! [X4,X5,X3] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X4),s('type.integer.int',X5))))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X5),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X3),s('type.integer.int',X4)))))) )),
  inference(superposition,[],[f426,f379])).
fof(f426,plain,(
  ( ! [X2,X0,X1] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',X0))))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X2),s('type.integer.int',X1))),s('type.integer.int',X0)))) )),
  inference(cnf_transformation,[],[f133])).
fof(f133,plain,(
  ! [X0,X1,X2] : s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',X0))))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X2),s('type.integer.int',X1))),s('type.integer.int',X0)))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X15,X6,X7] : s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X6),s('type.integer.int',X15))))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X7),s('type.integer.int',X6))),s('type.integer.int',X15)))),
  file('thm.int_arith.le_context_rwt|49|','thm.integer.INT_ADD_ASSOC')).
fof(f427,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X2))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X2)))))) )),
  inference(cnf_transformation,[],[f134])).
fof(f134,plain,(
  ! [X0,X1,X2] : s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X2))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X2)))))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X7,X6,X15] : s(bool,'const.integer.int_le_2'(s('type.integer.int',X6),s('type.integer.int',X15))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X7),s('type.integer.int',X6))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X7),s('type.integer.int',X15)))))),
  file('thm.int_arith.le_context_rwt|49|','thm.integer.INT_LE_LADD')).
fof(f8684,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK13),s('type.integer.int',sK14))))))))),
  inference(subsumption_resolution,[],[f8661,f297])).
fof(f8661,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK13),s('type.integer.int',sK14))))))))),
  inference(superposition,[],[f1593,f571])).
fof(f571,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)))))),
  inference(unit_resulting_resolution,[],[f523,f319])).
fof(f523,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK13),s('type.integer.int',sK14))))))),
  inference(backward_demodulation,[],[f379,f294])).
fof(f294,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))))),
  inference(cnf_transformation,[],[f199])).
fof(f1593,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK13),s('type.integer.int',X0))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK12),s('type.integer.int',X1))))))) )),
  inference(subsumption_resolution,[],[f1578,f297])).
fof(f1578,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK13),s('type.integer.int',X0))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK12),s('type.integer.int',X1))))))) )),
  inference(superposition,[],[f508,f570])).
fof(f570,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))),
  inference(unit_resulting_resolution,[],[f295,f319])).
fof(f295,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))))),
  inference(cnf_transformation,[],[f199])).
fof(f508,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X2),s('type.integer.int',X3)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',X3))))))) )),
  inference(cnf_transformation,[],[f180])).
fof(f180,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X2),s('type.integer.int',X3)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',X3)))))))),
  inference(flattening,[],[f179])).
fof(f179,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X2),s('type.integer.int',X3))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',X3)))))))),
  inference(ennf_transformation,[],[f144])).
fof(f144,plain,(
  ! [X0,X1,X2,X3] : ((p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X2),s('type.integer.int',X3))))) => p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',X3)))))))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X16,X7,X6,X15] : ((p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X16),s('type.integer.int',X7)))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X6),s('type.integer.int',X15))))) => p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X16),s('type.integer.int',X6))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X7),s('type.integer.int',X15)))))))),
  file('thm.int_arith.le_context_rwt|49|','thm.integer.INT_LE_ADD2')).
% SZS output end Proof for thm.int_arith
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 515
% Passive clauses: 1634
% Generated clauses: 8320
% Final active clauses: 505
% Final passive clauses: 1056
% Input formulas: 60
% Initial clauses: 215
% 
% Duplicate literals: 178
% Trivial inequalities: 18
% Fw subsumption resolutions: 368
% Fw demodulations: 1321
% Bw demodulations: 62
% Condensations: 141
% Global subsumptions: 1
% 
% Simple tautologies: 102
% Equational tautologies: 45
% Forward subsumptions: 4406
% Fw demodulations to eq. taut.: 17
% Bw demodulations to eq. taut.: 1
% 
% Unit resulting resolution: 2314
% Factoring: 1
% Forward superposition: 2486
% Backward superposition: 1047
% Self superposition: 50
% Equality factoring: 28
% Equality resolution: 1
% 
% Split clauses: 73
% Split components: 151
% SAT solver clauses: 1785
% SAT solver unit clauses: 699
% SAT solver binary clauses: 333
% 
% Sat splits: 73
% Sat splitting refutations: 2
% 
% Memory used [KB]: 8443
% Time elapsed: 0.410 s
% ------------------------------
% ------------------------------
% Success in time 1.728 s
