Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.quantHeuristics
% (83072)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 679
% Passive clauses: 10738
% Generated clauses: 19691
% Final active clauses: 674
% Final passive clauses: 9626
% Input formulas: 52
% Initial clauses: 188
% 
% Duplicate literals: 302
% Trivial inequalities: 60
% Fw subsumption resolutions: 1634
% Fw demodulations: 226
% Bw demodulations: 5
% 
% Simple tautologies: 224
% Equational tautologies: 56
% Forward subsumptions: 6091
% 
% Binary resolution: 12179
% Factoring: 1149
% Forward superposition: 1727
% Backward superposition: 2006
% Self superposition: 22
% Equality factoring: 26
% Equality resolution: 18
% 
% Split clauses: 337
% Split components: 681
% SAT solver clauses: 700
% SAT solver unit clauses: 20
% SAT solver binary clauses: 192
% 
% Sat splits: 337
% 
% TWLsolver clauses: 700
% TWLsolver calls for satisfiability: 428
% 
% Memory used [KB]: 9594
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.quantHeuristics
% (83089)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 389
% Passive clauses: 3289
% Generated clauses: 13985
% Final active clauses: 334
% Final passive clauses: 1109
% Input formulas: 52
% Initial clauses: 188
% 
% Duplicate literals: 203
% Trivial inequalities: 62
% Fw subsumption resolutions: 1431
% Fw demodulations: 600
% Bw demodulations: 1744
% Condensations: 320
% Global subsumptions: 18
% 
% Simple tautologies: 264
% Equational tautologies: 17
% Forward subsumptions: 4521
% Bw demodulations to eq. taut.: 96
% 
% Binary resolution: 2808
% Factoring: 182
% Forward superposition: 842
% Backward superposition: 1267
% Self superposition: 24
% Equality factoring: 28
% Equality resolution: 7
% 
% Split clauses: 134
% Split components: 278
% SAT solver clauses: 3170
% SAT solver unit clauses: 195
% SAT solver binary clauses: 80
% SAT solver learnt clauses: 37
% SAT solver learnt literals: 11
% 
% Sat splits: 134
% Sat splitting refutations: 10
% 
% TWLsolver clauses: 3097
% TWLsolver calls for satisfiability: 1574
% 
% Memory used [KB]: 12409
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.quantHeuristics
% (83123)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1654
% Passive clauses: 13288
% Generated clauses: 101903
% Final active clauses: 623
% Final passive clauses: 379
% Input formulas: 52
% Initial clauses: 188
% 
% Duplicate literals: 5779
% Trivial inequalities: 336
% Fw subsumption resolutions: 5206
% Fw demodulations: 7833
% Bw demodulations: 8865
% Condensations: 9552
% Global subsumptions: 317
% 
% Simple tautologies: 347
% Equational tautologies: 1363
% Forward subsumptions: 56550
% Bw demodulations to eq. taut.: 1231
% 
% Unit resulting resolution: 14555
% Factoring: 2
% Forward superposition: 24744
% Backward superposition: 14845
% Self superposition: 250
% Equality factoring: 199
% Equality resolution: 13
% 
% Split clauses: 731
% Split components: 1481
% SAT solver clauses: 15604
% SAT solver unit clauses: 451
% SAT solver binary clauses: 1565
% 
% Sat splits: 731
% Sat splitting refutations: 534
% 
% Memory used [KB]: 49764
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.quantHeuristics
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.quantHeuristics
% SZS output start Proof for thm.quantHeuristics
fof(f27515,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f432,f433,f434,f435,f436,f438,f527,f439,f506,f525,f540,f541,f542,f543,f544,f528,f637,f545,f547,f618,f636,f740,f735,f784,f785,f786,f833,f832,f837,f839,f995,f986,f985,f1059,f1060,f1061,f1062,f1063,f1064,f1069,f1070,f1071,f1072,f1249,f1251,f1252,f1253,f1256,f1323,f1342,f1479,f1480,f2026,f2027,f2093,f2101,f2103,f2096,f2113,f2119,f2571,f2573,f2574,f2575,f2578,f2645,f2664,f2780,f2786,f2989,f2991,f2992,f2994,f2996,f3067,f3088,f4184,f5878,f5984,f5982,f5882,f5884,f5885,f5887,f5889,f5960,f5985,f6001,f6003,f6004,f6006,f6008,f6079,f6101,f6106,f6107,f6374,f6375,f7032,f7033,f7567,f7569,f7572,f7573,f7568,f7574,f7575,f7576,f7579,f7580,f7577,f7584,f7585,f7586,f7587,f7564,f7563,f7589,f7588,f887,f886,f891,f890,f7700,f7699,f7776,f7777,f7778,f7779,f7780,f7781,f7786,f7787,f7788,f7789,f7829,f7830,f7878,f7879,f7880,f7881,f7996,f7997,f7998,f7578,f8000,f8001,f7999,f8004,f8005,f8006,f8007,f8093,f8098,f8172,f8283,f8474,f8475,f8667,f8668,f8669,f8671,f8672,f8670,f8662,f8661,f8675,f8673,f8875,f9129,f8174,f10253,f10258,f10268,f4193,f10386,f10387,f11217,f11218,f11219,f11220,f13813,f13815,f13814,f13961,f13963,f13962,f16854,f16856,f16857,f16858,f16861,f16928,f16934,f16935,f16936,f16999,f16938,f16939,f16940,f17008,f16980,f16979,f16998,f17041,f17042,f17070,f17071,f17091,f17093,f17094,f17095,f17098,f17165,f17171,f17172,f17173,f17235,f17175,f17176,f17177,f17244,f17217,f17216,f17234,f17397,f17398,f17426,f17427,f17558,f17559,f17587,f17588,f17620,f17621,f17652,f17653,f17655,f16994,f17658,f17659,f17684,f17667,f17668,f17670,f17686,f17713,f17714,f17750,f17753,f17790,f17823,f17683,f17827,f17857,f17858,f17885,f17886,f17917,f17918,f17948,f17949,f17951,f17231,f17954,f17955,f17980,f17963,f17964,f17966,f17981,f18144,f18145,f18180,f18183,f18220,f18253,f17979,f18257,f19115,f19114,f19161,f19160,f19335,f19334,f19204,f19206,f19207,f19208,f19211,f19278,f19353,f19284,f19285,f19286,f19357,f19288,f19289,f19290,f19366,f19331,f19330,f19356,f19402,f19404,f19401,f19375,f19403,f19405,f19406,f19409,f19412,f19447,f19449,f19450,f19452,f19454,f19525,f19530,f19531,f19532,f19533,f19535,f19536,f19537,f19545,f19547,f19603,f19585,f19584,f19594,f19604,f19648,f19649,f19677,f19678,f19679,f19698,f19699,f19700,f19712,f19713,f19714,f19725,f19726,f19588,f19736,f19756,f19738,f19757,f19745,f19746,f19747,f19791,f19792,f19793,f19794,f19795,f19796,f19797,f19815,f17048,f19818,f17404,f19857,f19858,f19880,f19881,f19929,f19928,f19927,f19985,f19984,f19983,f19997,f19999,f20000,f20002,f20004,f20075,f20080,f20081,f20082,f20083,f20085,f20086,f20087,f20095,f20097,f20147,f20150,f20135,f20134,f20143,f20151,f20324,f20337,f19600,f20354,f20355,f20495,f20497,f20498,f20499,f20502,f20569,f20637,f20638,f20649,f20650,f20661,f20662,f20673,f20674,f20690,f20691,f20704,f20705,f20707,f20626,f20710,f20711,f20737,f20719,f20720,f20722,f20738,f20749,f20750,f20751,f20770,f20771,f20774,f20794,f20810,f20736,f20814,f20878,f20880,f20881,f20882,f20884,f20890,f20892,f20893,f20895,f20897,f20968,f20974,f20975,f20976,f20977,f20979,f20980,f20981,f21037,f21035,f19598,f21055,f21032,f21031,f21046,f21056,f21092,f21094,f21095,f21096,f21099,f21166,f21235,f21236,f21246,f21247,f21253,f21254,f21257,f21278,f17565,f21279,f21280,f21331,f17596,f21332,f21348,f17627,f21349,f21362,f21363,f838,f21365,f21224,f21368,f21369,f21395,f7590,f21377,f17685,f21378,f21380,f21396,f21407,f21408,f17723,f21409,f21427,f21428,f21431,f21495,f21511,f21394,f21515,f21576,f21577,f21578,f21579,f21580,f21668,f21670,f21671,f21672,f21673,f21796,f21797,f21798,f21799,f21800,f21813,f21842,f21848,f21850,f21851,f21853,f21855,f21926,f22000,f22013,f22015,f22018,f22020,f22045,f22046,f22047,f22069,f22070,f22071,f22083,f22084,f22085,f22095,f22096,f21991,f22106,f22126,f22108,f22127,f22115,f22116,f22117,f22160,f22161,f22162,f22163,f22164,f22165,f22166,f22178,f22180,f22183,f22185,f22220,f22221,f22243,f22244,f22320,f22326,f22327,f22328,f22349,f22357,f22358,f22359,f22466,f22468,f22469,f22455,f22456,f22457,f22531,f22532,f22533,f22526,f22527,f22528,f22627,f22629,f22630,f22622,f22623,f22624,f22714,f22716,f22717,f22682,f22683,f20138,f22693,f22724,f22695,f22725,f22702,f22703,f22704,f22804,f22805,f22806,f22814,f22795,f22796,f22797,f22798,f22799,f22800,f22801,f22816,f22827,f22835,f22836,f22895,f22897,f22898,f22888,f22889,f23033,f23034,f23035,f23028,f20152,f23029,f22476,f23098,f23099,f23100,f23087,f19686,f23088,f23089,f23096,f23108,f23275,f23276,f23277,f22467,f23270,f19701,f23271,f23272,f23273,f23283,f23353,f23354,f23355,f22628,f23348,f19715,f23349,f23350,f23351,f23361,f23460,f23461,f23462,f23428,f836,f23429,f21041,f23434,f22334,f23470,f23439,f23469,f23441,f23471,f23448,f19758,f23449,f23450,f23458,f23472,f23808,f23809,f23814,f23816,f23923,f23924,f23925,f22715,f23933,f22815,f23914,f23915,f19801,f23916,f19802,f23917,f23918,f23919,f23920,f23934,f23921,f23935,f24108,f24109,f24114,f24115,f24371,f24372,f24373,f24364,f24365,f24440,f19408,f24441,f19411,f24442,f19414,f22896,f24435,f21057,f24436,f23107,f24463,f24466,f24469,f24472,f24594,f24593,f24720,f24721,f24707,f24706,f24904,f24903,f24769,f24771,f24772,f24773,f24776,f24843,f24923,f24939,f24943,f24993,f24994,f24977,f24978,f25003,f25007,f24989,f25055,f25056,f25043,f24995,f25044,f24996,f25070,f25073,f25076,f25079,f25124,f25125,f25127,f25129,f25200,f25266,f25268,f25269,f25270,f25271,f25273,f25287,f17049,f25290,f17405,f25293,f25296,f25324,f25347,f25348,f25349,f25533,f25532,f25392,f25394,f25395,f25396,f25399,f25466,f25552,f25569,f25623,f25624,f25608,f25609,f25612,f25627,f25641,f25695,f25696,f25682,f25625,f25683,f25626,f25684,f25697,f25811,f25877,f25878,f20642,f25879,f20643,f25880,f25881,f25882,f20879,f26130,f26129,f26282,f26269,f26268,f26394,f24986,f26452,f26453,f26433,f26454,f26455,f26437,f26438,f26465,f25050,f26523,f26524,f26505,f26525,f26526,f26283,f26510,f26456,f26511,f26457,f26534,f25617,f26592,f26593,f26573,f26595,f26596,f26577,f26578,f26606,f25689,f26664,f26665,f16997,f26646,f20883,f26666,f26667,f26594,f26651,f26597,f26652,f26598,f26675,f25260,f26749,f26750,f26751,f26752,f26753,f26754,f26800,f25871,f21841,f26873,f25267,f26874,f26875,f25272,f26876,f22002,f26877,f22056,f26878,f21669,f22171,f27071,f27070,f26930,f26932,f26933,f26934,f26937,f27004,f27091,f27281,f21240,f27280,f21241,f27146,f27148,f27149,f526,f27150,f27153,f529,f27220,f27301,f27480,f27484,f27487,f27488,f27489,f27490,f27493,f27494,f27495,f27496,f530,f27497,f27498,f27499,f531,f27507,f27508])).
fof(f27508,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f27464,f2643])).
fof(f2643,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))))),
  inference(resolution,[],[f1246,f402])).
fof(f402,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f242])).
fof(f242,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f101])).
fof(f101,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X13] : (s(bool,'T_0') = s(bool,X13) <=> p(s(bool,X13)))),
  file('thm.quantHeuristics.GUESSES_WEAKEN_THM|split|3','thm.bool.EQ_CLAUSES|split|0')).
fof(f1246,plain,(
  p(s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13))))))),
  inference(resolution,[],[f538,f425])).
fof(f425,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) )),
  inference(cnf_transformation,[],[f152])).
fof(f152,plain,(
  ! [X0,X1,X2] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(ennf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(flattening,[],[f117])).
fof(f117,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : ! [X9,X10] : (p(s(bool,happ(s(fun(X6,bool),X9),s(X6,X10)))) => p(s(bool,happ(s(fun(X6,bool),X9),s(X6,'const.min.@_1'(s(fun(X6,bool),X9)))))))),
  file('thm.quantHeuristics.GUESSES_WEAKEN_THM|split|3','thm.bool.SELECT_AX')).
fof(f538,plain,(
  p(s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11)))))),
  inference(resolution,[],[f251,f253])).
fof(f253,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | p(s(bool,happ(s(fun(X0,bool),X3),s(X0,sK15(X3,X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f171])).
fof(f171,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X6] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK14(X6,X3,X2,X1,X0))))))))) & ((p(s(bool,happ(s(fun(X0,bool),X3),s(X0,sK15(X3,X2,X1,X0))))) & ! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK14,sK15])],[f170])).
fof(f170,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X6] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | ? [X7] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X7)))))))) & (? [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) & ! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f169])).
fof(f169,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) | ? [X5] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))))) & (? [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) & ! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(nnf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) | ? [X5] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))))),
  inference(ennf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) => ? [X5] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))))),
  inference(flattening,[],[f56])).
fof(f56,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) => ? [X5] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X7,X6] : ! [X8,X9] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X6,X7),X8),s(fun(X7,bool),X9)))) <=> ! [X18] : (p(s(bool,happ(s(fun(X7,bool),X9),s(X7,X18)))) => ? [X19] : p(s(bool,happ(s(fun(X7,bool),X9),s(X7,happ(s(fun(X6,X7),X8),s(X6,X19))))))))),
  file('thm.quantHeuristics.GUESSES_WEAKEN_THM|split|3','thm.quantHeuristics.GUESS_EXISTS_FORALL_REWRITES|split|0')).
fof(f251,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))))),
  inference(cnf_transformation,[],[f168])).
fof(f168,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) & ~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK10,sK11,sK12,sK13])],[f119])).
fof(f119,plain,(
  ? [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X3)))) & ~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X3)))))),
  inference(ennf_transformation,[],[f55])).
fof(f55,plain,(
  ~! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X3)))) => p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X3)))))),
  inference(flattening,[],[f54])).
fof(f54,plain,(
  ~! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X3)))) => p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X3)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X6,X7),X8),s(fun(X7,bool),X9)))) => p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X6,X7),X8),s(fun(X7,bool),X9)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X6,X7),X8),s(fun(X7,bool),X9)))) => p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X6,X7),X8),s(fun(X7,bool),X9)))))),
  file('thm.quantHeuristics.GUESSES_WEAKEN_THM|split|3','thm.quantHeuristics.GUESSES_WEAKEN_THM|split|3')).
fof(f27464,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13))))) | $spl42),
  inference(superposition,[],[f3064,f5984])).
fof(f3064,plain,(
  ( ! [X197] : (s(bool,'T_0') != s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X197)))))) )),
  inference(resolution,[],[f539,f401])).
fof(f401,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f242])).
fof(f539,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X0))))))) )),
  inference(resolution,[],[f251,f254])).
fof(f254,plain,(
  ( ! [X2,X0,X5,X3,X1] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) )),
  inference(cnf_transformation,[],[f171])).
fof(f27507,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f27506,f1451])).
fof(f1451,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(subsumption_resolution,[],[f1442,f250])).
fof(f250,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))))),
  inference(cnf_transformation,[],[f168])).
fof(f1442,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))))),
  inference(superposition,[],[f502,f402])).
fof(f502,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))),
  inference(resolution,[],[f250,f397])).
fof(f397,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f240])).
fof(f240,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f98])).
fof(f98,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f97])).
fof(f97,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X13] : (s(bool,'F_0') = s(bool,X13) <=> ~p(s(bool,X13)))),
  file('thm.quantHeuristics.GUESSES_WEAKEN_THM|split|3','thm.bool.EQ_CLAUSES|split|3')).
fof(f27506,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl42),
  inference(forward_demodulation,[],[f27463,f2643])).
fof(f27463,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13))))) | $spl42),
  inference(superposition,[],[f3062,f5984])).
fof(f3062,plain,(
  ( ! [X195] : (s(bool,'F_0') = s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X195)))))) )),
  inference(resolution,[],[f539,f398])).
fof(f398,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f240])).
fof(f531,plain,(
  ( ! [X13] : (p(s(bool,X13))) ) | $spl12),
  inference(cnf_transformation,[],[f531_D])).
fof(f531_D,plain,(
  ( ! [X13] : (p(s(bool,X13))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f27499,plain,(
  ( ! [X59,X58] : (p(s(bool,X59)) | p(s(bool,X58))) ) | $spl42),
  inference(subsumption_resolution,[],[f27456,f2617])).
fof(f2617,plain,(
  ( ! [X105,X104] : (sP4(X104,X105,happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))))) )),
  inference(resolution,[],[f1246,f313])).
fof(f313,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f200])).
fof(f200,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(flattening,[],[f199])).
fof(f199,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(nnf_transformation,[],[f159])).
fof(f159,plain,(
  ! [X0,X1,X2] : (sP4(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f27456,plain,(
  ( ! [X59,X58] : (~sP4(X58,X59,happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13))))) | p(s(bool,X59)) | p(s(bool,X58))) ) | $spl42),
  inference(superposition,[],[f3039,f5984])).
fof(f3039,plain,(
  ( ! [X158,X159,X160] : (~sP4(X158,X159,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X160))))) | p(s(bool,X159)) | p(s(bool,X158))) )),
  inference(resolution,[],[f539,f312])).
fof(f312,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f200])).
fof(f27498,plain,(
  ( ! [X57,X56] : (p(s(bool,X57)) | p(s(bool,X56))) ) | $spl42),
  inference(subsumption_resolution,[],[f27455,f2619])).
fof(f2619,plain,(
  ( ! [X109,X108] : (sP4(happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),X108,X109)) )),
  inference(resolution,[],[f1246,f315])).
fof(f315,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f200])).
fof(f27455,plain,(
  ( ! [X57,X56] : (~sP4(happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),X56,X57) | p(s(bool,X57)) | p(s(bool,X56))) ) | $spl42),
  inference(superposition,[],[f3038,f5984])).
fof(f3038,plain,(
  ( ! [X156,X157,X155] : (~sP4(happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X157)))),X155,X156) | p(s(bool,X156)) | p(s(bool,X155))) )),
  inference(resolution,[],[f539,f312])).
fof(f27497,plain,(
  ( ! [X54,X55] : (p(s(bool,X55)) | p(s(bool,X54))) ) | $spl42),
  inference(subsumption_resolution,[],[f27454,f2618])).
fof(f2618,plain,(
  ( ! [X107,X106] : (sP4(X106,happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),X107)) )),
  inference(resolution,[],[f1246,f314])).
fof(f314,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f200])).
fof(f27454,plain,(
  ( ! [X54,X55] : (~sP4(X54,happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),X55) | p(s(bool,X55)) | p(s(bool,X54))) ) | $spl42),
  inference(superposition,[],[f3037,f5984])).
fof(f3037,plain,(
  ( ! [X154,X152,X153] : (~sP4(X152,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X154)))),X153) | p(s(bool,X153)) | p(s(bool,X152))) )),
  inference(resolution,[],[f539,f312])).
fof(f530,plain,(
  ( ! [X14] : (~p(s(bool,X14))) ) | $spl10),
  inference(cnf_transformation,[],[f530_D])).
fof(f530_D,plain,(
  ( ! [X14] : (~p(s(bool,X14))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f27496,plain,(
  ( ! [X48,X49] : (~p(s(bool,X49)) | ~p(s(bool,X48))) ) | $spl42),
  inference(subsumption_resolution,[],[f27451,f2616])).
fof(f2616,plain,(
  ( ! [X103,X102] : (sP5(X102,X103,happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))))) )),
  inference(resolution,[],[f1246,f309])).
fof(f309,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f198])).
fof(f198,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f197])).
fof(f197,plain,(
  ! [X1,X0,X2] : ((~sP5(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP5(X1,X0,X2)))),
  inference(flattening,[],[f196])).
fof(f196,plain,(
  ! [X1,X0,X2] : ((~sP5(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP5(X1,X0,X2)))),
  inference(nnf_transformation,[],[f160])).
fof(f160,plain,(
  ! [X1,X0,X2] : (sP5(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f27451,plain,(
  ( ! [X48,X49] : (~sP5(X48,X49,happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13))))) | ~p(s(bool,X49)) | ~p(s(bool,X48))) ) | $spl42),
  inference(superposition,[],[f3034,f5984])).
fof(f3034,plain,(
  ( ! [X144,X145,X143] : (~sP5(X143,X144,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X145))))) | ~p(s(bool,X144)) | ~p(s(bool,X143))) )),
  inference(resolution,[],[f539,f308])).
fof(f308,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f198])).
fof(f27495,plain,(
  ( ! [X43,X42] : (~p(s(bool,X43)) | ~p(s(bool,X42))) ) | $spl42),
  inference(subsumption_resolution,[],[f27448,f2613])).
fof(f2613,plain,(
  ( ! [X97,X96] : (sP6(X96,X97,happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))))) )),
  inference(resolution,[],[f1246,f305])).
fof(f305,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(rectify,[],[f194])).
fof(f194,plain,(
  ! [X2,X0,X1] : ((~sP6(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP6(X2,X0,X1)))),
  inference(flattening,[],[f193])).
fof(f193,plain,(
  ! [X2,X0,X1] : ((~sP6(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP6(X2,X0,X1)))),
  inference(nnf_transformation,[],[f161])).
fof(f161,plain,(
  ! [X2,X0,X1] : (sP6(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f27448,plain,(
  ( ! [X43,X42] : (~sP6(X42,X43,happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13))))) | ~p(s(bool,X43)) | ~p(s(bool,X42))) ) | $spl42),
  inference(superposition,[],[f3031,f5984])).
fof(f3031,plain,(
  ( ! [X134,X136,X135] : (~sP6(X134,X135,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X136))))) | ~p(s(bool,X135)) | ~p(s(bool,X134))) )),
  inference(resolution,[],[f539,f304])).
fof(f304,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f195])).
fof(f27494,plain,(
  ( ! [X39,X38] : (~p(s(bool,X38)) | p(s(bool,X39))) ) | $spl42),
  inference(subsumption_resolution,[],[f27446,f2603])).
fof(f2603,plain,(
  ( ! [X76,X77] : (sP3(X76,happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),X77) | ~p(s(bool,X76))) )),
  inference(resolution,[],[f1246,f293])).
fof(f293,plain,(
  ( ! [X2,X0,X1] : (sP3(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f189])).
fof(f189,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP3(X2,X0,X1)) & (~sP3(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f188])).
fof(f188,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP3(X2,X0,X1)) & (~sP3(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f158])).
fof(f158,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP3(X2,X0,X1))),
  inference(definition_folding,[],[f62,f157,f156])).
fof(f156,plain,(
  ! [X2,X0,X1] : (sP2(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f157,plain,(
  ! [X2,X0,X1] : (sP3(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP2(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f62,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f61])).
fof(f61,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X22,X23,X24] : ((p(s(bool,X24)) <=> (p(s(bool,X23)) | p(s(bool,X22)))) <=> ((p(s(bool,X24)) | ~p(s(bool,X23))) & (p(s(bool,X24)) | ~p(s(bool,X22))) & (p(s(bool,X23)) | p(s(bool,X22)) | ~p(s(bool,X24)))))),
  file('thm.quantHeuristics.GUESSES_WEAKEN_THM|split|3','thm.sat.sat_25_prove_ltb')).
fof(f27446,plain,(
  ( ! [X39,X38] : (~sP3(X38,happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),X39) | ~p(s(bool,X38)) | p(s(bool,X39))) ) | $spl42),
  inference(superposition,[],[f3025,f5984])).
fof(f3025,plain,(
  ( ! [X118,X116,X117] : (~sP3(X117,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X118)))),X116) | ~p(s(bool,X117)) | p(s(bool,X116))) )),
  inference(resolution,[],[f539,f294])).
fof(f294,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP3(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f189])).
fof(f27493,plain,(
  ( ! [X37,X36] : (~p(s(bool,X36)) | p(s(bool,X37))) ) | $spl42),
  inference(subsumption_resolution,[],[f27445,f2601])).
fof(f2601,plain,(
  ( ! [X72,X73] : (sP3(X72,X73,happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13))))) | ~p(s(bool,X72))) )),
  inference(resolution,[],[f1246,f292])).
fof(f292,plain,(
  ( ! [X2,X0,X1] : (sP3(X2,X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f189])).
fof(f27445,plain,(
  ( ! [X37,X36] : (~sP3(X36,X37,happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13))))) | ~p(s(bool,X36)) | p(s(bool,X37))) ) | $spl42),
  inference(superposition,[],[f3024,f5984])).
fof(f3024,plain,(
  ( ! [X114,X115,X113] : (~sP3(X114,X113,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X115))))) | ~p(s(bool,X114)) | p(s(bool,X113))) )),
  inference(resolution,[],[f539,f294])).
fof(f27490,plain,(
  ( ! [X33,X32] : (p(s(bool,X33)) | p(s(bool,X32))) ) | $spl42),
  inference(subsumption_resolution,[],[f27443,f2595])).
fof(f2595,plain,(
  ( ! [X61,X60] : (~sP3(X60,happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),X61) | p(s(bool,X60))) )),
  inference(resolution,[],[f1246,f281])).
fof(f281,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f184])).
fof(f184,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP2(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP2(X0,X1,X2) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f183])).
fof(f183,plain,(
  ! [X2,X0,X1] : ((~sP3(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP2(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP2(X2,X0,X1) | sP3(X2,X0,X1)))),
  inference(flattening,[],[f182])).
fof(f182,plain,(
  ! [X2,X0,X1] : ((~sP3(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP2(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP2(X2,X0,X1)) | sP3(X2,X0,X1)))),
  inference(nnf_transformation,[],[f157])).
fof(f27443,plain,(
  ( ! [X33,X32] : (sP3(X32,happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),X33) | p(s(bool,X33)) | p(s(bool,X32))) ) | $spl42),
  inference(superposition,[],[f3022,f5984])).
fof(f3022,plain,(
  ( ! [X109,X107,X108] : (sP3(X107,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X108)))),X109) | p(s(bool,X109)) | p(s(bool,X107))) )),
  inference(resolution,[],[f539,f291])).
fof(f291,plain,(
  ( ! [X2,X0,X1] : (sP3(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f189])).
fof(f27489,plain,(
  ( ! [X30,X31] : (p(s(bool,X31)) | p(s(bool,X30))) ) | $spl42),
  inference(subsumption_resolution,[],[f27442,f2594])).
fof(f2594,plain,(
  ( ! [X59,X58] : (~sP3(X58,X59,happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13))))) | p(s(bool,X58))) )),
  inference(resolution,[],[f1246,f280])).
fof(f280,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X0)) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f184])).
fof(f27442,plain,(
  ( ! [X30,X31] : (sP3(X30,X31,happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13))))) | p(s(bool,X31)) | p(s(bool,X30))) ) | $spl42),
  inference(superposition,[],[f3021,f5984])).
fof(f3021,plain,(
  ( ! [X105,X106,X104] : (sP3(X104,X105,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X106))))) | p(s(bool,X105)) | p(s(bool,X104))) )),
  inference(resolution,[],[f539,f291])).
fof(f27488,plain,(
  ( ! [X26,X27] : (p(s(bool,X27)) | ~p(s(bool,X26))) ) | $spl42),
  inference(subsumption_resolution,[],[f27440,f2599])).
fof(f2599,plain,(
  ( ! [X68,X69] : (sP2(X68,X69,happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))))) )),
  inference(resolution,[],[f1246,f288])).
fof(f288,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f187])).
fof(f187,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(rectify,[],[f186])).
fof(f186,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(flattening,[],[f185])).
fof(f185,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(nnf_transformation,[],[f156])).
fof(f27440,plain,(
  ( ! [X26,X27] : (~sP2(X26,X27,happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13))))) | p(s(bool,X27)) | ~p(s(bool,X26))) ) | $spl42),
  inference(superposition,[],[f3019,f5984])).
fof(f3019,plain,(
  ( ! [X99,X100,X98] : (~sP2(X98,X99,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X100))))) | p(s(bool,X99)) | ~p(s(bool,X98))) )),
  inference(resolution,[],[f539,f287])).
fof(f287,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f187])).
fof(f27487,plain,(
  ( ! [X24,X25] : (p(s(bool,X25)) | ~p(s(bool,X24))) ) | $spl42),
  inference(subsumption_resolution,[],[f27439,f2600])).
fof(f2600,plain,(
  ( ! [X70,X71] : (sP2(X70,happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),X71)) )),
  inference(resolution,[],[f1246,f289])).
fof(f289,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f187])).
fof(f27439,plain,(
  ( ! [X24,X25] : (~sP2(X24,happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),X25) | p(s(bool,X25)) | ~p(s(bool,X24))) ) | $spl42),
  inference(superposition,[],[f3018,f5984])).
fof(f3018,plain,(
  ( ! [X97,X95,X96] : (~sP2(X95,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X97)))),X96) | p(s(bool,X96)) | ~p(s(bool,X95))) )),
  inference(resolution,[],[f539,f287])).
fof(f27484,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | ~p(s(bool,X19))) ) | $spl42),
  inference(subsumption_resolution,[],[f27436,f2589])).
fof(f2589,plain,(
  ( ! [X48,X49] : (sP1(X48,happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),X49) | ~p(s(bool,X48))) )),
  inference(resolution,[],[f1246,f276])).
fof(f276,plain,(
  ( ! [X2,X0,X1] : (sP1(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f181])).
fof(f181,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP1(X2,X0,X1)) & (~sP1(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f180])).
fof(f180,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP1(X2,X0,X1)) & (~sP1(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f155])).
fof(f155,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP1(X2,X0,X1))),
  inference(definition_folding,[],[f123,f154,f153])).
fof(f153,plain,(
  ! [X2,X0,X1] : (sP0(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f154,plain,(
  ! [X2,X0,X1] : (sP1(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP0(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f123,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f60])).
fof(f60,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f59])).
fof(f59,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X22,X23,X24] : ((p(s(bool,X24)) <=> (p(s(bool,X23)) => p(s(bool,X22)))) <=> ((p(s(bool,X24)) | p(s(bool,X23))) & (p(s(bool,X24)) | ~p(s(bool,X22))) & (~p(s(bool,X23)) | p(s(bool,X22)) | ~p(s(bool,X24)))))),
  file('thm.quantHeuristics.GUESSES_WEAKEN_THM|split|3','thm.sat.sat_27_prove_ltb')).
fof(f27436,plain,(
  ( ! [X19,X18] : (~sP1(X18,happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),X19) | ~p(s(bool,X18)) | ~p(s(bool,X19))) ) | $spl42),
  inference(superposition,[],[f3008,f5984])).
fof(f3008,plain,(
  ( ! [X66,X67,X65] : (~sP1(X66,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X67)))),X65) | ~p(s(bool,X66)) | ~p(s(bool,X65))) )),
  inference(resolution,[],[f539,f277])).
fof(f277,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP1(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f181])).
fof(f27480,plain,(
  ( ! [X6,X7] : (~p(s(bool,X7)) | ~p(s(bool,X6))) ) | $spl42),
  inference(subsumption_resolution,[],[f27430,f2586])).
fof(f2586,plain,(
  ( ! [X43,X42] : (sP0(X42,happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),X43)) )),
  inference(resolution,[],[f1246,f272])).
fof(f272,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f179])).
fof(f179,plain,(
  ! [X0,X1,X2] : ((~sP0(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP0(X0,X1,X2)))),
  inference(rectify,[],[f178])).
fof(f178,plain,(
  ! [X2,X0,X1] : ((~sP0(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP0(X2,X0,X1)))),
  inference(flattening,[],[f177])).
fof(f177,plain,(
  ! [X2,X0,X1] : ((~sP0(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP0(X2,X0,X1)))),
  inference(nnf_transformation,[],[f153])).
fof(f27430,plain,(
  ( ! [X6,X7] : (~sP0(X6,happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),X7) | ~p(s(bool,X7)) | ~p(s(bool,X6))) ) | $spl42),
  inference(superposition,[],[f3002,f5984])).
fof(f3002,plain,(
  ( ! [X47,X48,X49] : (~sP0(X47,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X49)))),X48) | ~p(s(bool,X48)) | ~p(s(bool,X47))) )),
  inference(resolution,[],[f539,f270])).
fof(f270,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP0(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f179])).
fof(f27301,plain,(
  ( ! [X0,X3,X1] : (sP8(X0,'T_0') | s(bool,X1) != s(bool,X3) | p(s(bool,X3)) | sP8(X1,X0)) ) | $spl330),
  inference(cnf_transformation,[],[f27301_D])).
fof(f27301_D,plain,(
  ( ! [X0,X3,X1] : (sP8(X0,'T_0') | s(bool,X1) != s(bool,X3) | p(s(bool,X3)) | sP8(X1,X0)) ) <=> ~$spl330),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl330])])).
fof(f27220,plain,(
  ( ! [X273,X275,X272,X274] : (sP8(X272,'T_0') | sP8(X273,X272) | p(s(bool,X274)) | p(s(bool,X275)) | s(bool,X273) != s(bool,X275)) ) | $spl144),
  inference(resolution,[],[f21146,f426])).
fof(f426,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f259])).
fof(f259,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f173])).
fof(f173,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f172])).
fof(f172,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f121])).
fof(f121,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f58])).
fof(f58,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X10,X16,X14,X17] : ((s(bool,X10) = s(bool,X16) & (p(s(bool,X16)) => s(bool,X14) = s(bool,X17))) => ((p(s(bool,X10)) => p(s(bool,X14))) <=> (p(s(bool,X16)) => p(s(bool,X17)))))),
  file('thm.quantHeuristics.GUESSES_WEAKEN_THM|split|3','thm.bool.IMP_CONG')).
fof(f21146,plain,(
  ( ! [X167,X168] : (sP8(X167,'T_0') | sP8(X168,X167) | p(s(bool,X168))) ) | $spl144),
  inference(resolution,[],[f21077,f331])).
fof(f331,plain,(
  ( ! [X0,X1] : (sP8(X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f206])).
fof(f206,plain,(
  ! [X0,X1] : ((~sP8(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1)))),
  inference(rectify,[],[f205])).
fof(f205,plain,(
  ! [X1,X0] : ((~sP8(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP8(X1,X0)))),
  inference(flattening,[],[f204])).
fof(f204,plain,(
  ! [X1,X0] : ((~sP8(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP8(X1,X0)))),
  inference(nnf_transformation,[],[f164])).
fof(f164,plain,(
  ! [X1,X0] : (sP8(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f21077,plain,(
  ( ! [X0] : (p(s(bool,X0)) | sP8(X0,'T_0')) ) | $spl144),
  inference(subsumption_resolution,[],[f21061,f386])).
fof(f386,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f233])).
fof(f233,plain,(
  ! [X0] : ((~p(s(bool,X0)) | ~p(s(bool,'F_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,X0)) & p(s(bool,'F_0')))))),
  inference(flattening,[],[f232])).
fof(f232,plain,(
  ! [X0] : (((~p(s(bool,X0)) | ~p(s(bool,'F_0'))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,X0)) & p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f93])).
fof(f93,plain,(
  ! [X0] : ((p(s(bool,X0)) & p(s(bool,'F_0'))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X13] : ((p(s(bool,X13)) & p(s(bool,'F_0'))) <=> p(s(bool,'F_0')))),
  file('thm.quantHeuristics.GUESSES_WEAKEN_THM|split|3','thm.bool.AND_CLAUSES|split|3')).
fof(f21061,plain,(
  ( ! [X0] : (sP8(X0,'T_0') | p(s(bool,'F_0')) | p(s(bool,X0))) ) | $spl144),
  inference(resolution,[],[f20984,f326])).
fof(f326,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | ~sP8(X0,X1)) )),
  inference(cnf_transformation,[],[f206])).
fof(f20984,plain,(
  ( ! [X321] : (sP8(X321,'F_0') | sP8(X321,'T_0')) ) | $spl144),
  inference(resolution,[],[f19503,f417])).
fof(f417,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f9])).
fof(f9,axiom,(
  p(s(bool,'T_0'))),
  file('thm.quantHeuristics.GUESSES_WEAKEN_THM|split|3','thm.bool.TRUTH')).
fof(f19503,plain,(
  ( ! [X173,X174] : (~p(s(bool,X174)) | sP8(X173,X174) | sP8(X173,'F_0')) ) | $spl144),
  inference(resolution,[],[f19403,f331])).
fof(f529,plain,(
  ( ! [X8,X9] : (s(bool,X8) != s(bool,X9) | ~p(s(bool,X8))) ) | $spl8),
  inference(cnf_transformation,[],[f529_D])).
fof(f529_D,plain,(
  ( ! [X8,X9] : (s(bool,X8) != s(bool,X9) | ~p(s(bool,X8))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f27153,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP8(X35,'T_0') | sP8(X36,X35) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X37) | s(bool,X38) != s(bool,X39)) ) | $spl144),
  inference(resolution,[],[f21146,f262])).
fof(f262,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f173])).
fof(f27150,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP8(X20,'T_0') | sP8(X21,X20) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X21) != s(bool,X24)) ) | $spl144),
  inference(resolution,[],[f21146,f261])).
fof(f261,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f173])).
fof(f526,plain,(
  ( ! [X3,X1] : (s(bool,X1) != s(bool,X3) | p(s(bool,X1))) ) | $spl2),
  inference(cnf_transformation,[],[f526_D])).
fof(f526_D,plain,(
  ( ! [X3,X1] : (s(bool,X1) != s(bool,X3) | p(s(bool,X1))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f27149,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP8(X15,'T_0') | sP8(X16,X15) | p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) ) | $spl144),
  inference(resolution,[],[f21146,f260])).
fof(f260,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f173])).
fof(f27148,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP8(X10,'T_0') | sP8(X11,X10) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X12) | s(bool,X13) != s(bool,X14)) ) | $spl144),
  inference(resolution,[],[f21146,f258])).
fof(f258,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f173])).
fof(f27146,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP8(X0,'T_0') | sP8(X1,X0) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) ) | $spl144),
  inference(resolution,[],[f21146,f256])).
fof(f256,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f173])).
fof(f21241,plain,(
  ( ! [X5] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X5)))),'T_0')) ) | $spl214),
  inference(cnf_transformation,[],[f21241_D])).
fof(f21241_D,plain,(
  ( ! [X5] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X5)))),'T_0')) ) <=> ~$spl214),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl214])])).
fof(f27280,plain,(
  ( ! [X35,X34] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X34)))),'T_0') | p(s(bool,X35))) ) | $spl144),
  inference(duplicate_literal_removal,[],[f27138])).
fof(f27138,plain,(
  ( ! [X35,X34] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X34)))),'T_0') | p(s(bool,X35)) | p(s(bool,X35))) ) | $spl144),
  inference(resolution,[],[f21146,f3042])).
fof(f3042,plain,(
  ( ! [X167,X168] : (~sP8(X167,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X168))))) | p(s(bool,X167))) )),
  inference(resolution,[],[f539,f326])).
fof(f21240,plain,(
  sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0') | $spl212),
  inference(cnf_transformation,[],[f21240_D])).
fof(f21240_D,plain,(
  sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0') <=> ~$spl212),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl212])])).
fof(f27281,plain,(
  ( ! [X31] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0') | p(s(bool,X31))) ) | $spl144),
  inference(duplicate_literal_removal,[],[f27135])).
fof(f27135,plain,(
  ( ! [X31] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0') | p(s(bool,X31)) | p(s(bool,X31))) ) | $spl144),
  inference(resolution,[],[f21146,f593])).
fof(f593,plain,(
  ( ! [X115] : (~sP8(X115,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | p(s(bool,X115))) )),
  inference(resolution,[],[f251,f326])).
fof(f27091,plain,(
  ( ! [X0,X3,X1] : (sP8(X0,'T_0') | s(bool,X1) != s(bool,X3) | p(s(bool,X3)) | sP8(X0,X1)) ) | $spl328),
  inference(cnf_transformation,[],[f27091_D])).
fof(f27091_D,plain,(
  ( ! [X0,X3,X1] : (sP8(X0,'T_0') | s(bool,X1) != s(bool,X3) | p(s(bool,X3)) | sP8(X0,X1)) ) <=> ~$spl328),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl328])])).
fof(f27004,plain,(
  ( ! [X273,X275,X272,X274] : (sP8(X272,'T_0') | sP8(X272,X273) | p(s(bool,X274)) | p(s(bool,X275)) | s(bool,X273) != s(bool,X275)) ) | $spl144),
  inference(resolution,[],[f21145,f426])).
fof(f21145,plain,(
  ( ! [X165,X166] : (sP8(X165,'T_0') | sP8(X165,X166) | p(s(bool,X166))) ) | $spl144),
  inference(resolution,[],[f21077,f328])).
fof(f328,plain,(
  ( ! [X0,X1] : (sP8(X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f206])).
fof(f26937,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP8(X35,'T_0') | sP8(X35,X36) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X37) | s(bool,X38) != s(bool,X39)) ) | $spl144),
  inference(resolution,[],[f21145,f262])).
fof(f26934,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP8(X20,'T_0') | sP8(X20,X21) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X21) != s(bool,X24)) ) | $spl144),
  inference(resolution,[],[f21145,f261])).
fof(f26933,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP8(X15,'T_0') | sP8(X15,X16) | p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) ) | $spl144),
  inference(resolution,[],[f21145,f260])).
fof(f26932,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP8(X10,'T_0') | sP8(X10,X11) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X12) | s(bool,X13) != s(bool,X14)) ) | $spl144),
  inference(resolution,[],[f21145,f258])).
fof(f26930,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP8(X0,'T_0') | sP8(X0,X1) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) ) | $spl144),
  inference(resolution,[],[f21145,f256])).
fof(f27070,plain,(
  ( ! [X41,X40] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X40)))),'T_0') | p(s(bool,X41))) ) | $spl144),
  inference(duplicate_literal_removal,[],[f26929])).
fof(f26929,plain,(
  ( ! [X41,X40] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X40)))),'T_0') | p(s(bool,X41)) | p(s(bool,X41))) ) | $spl144),
  inference(resolution,[],[f21145,f3043])).
fof(f3043,plain,(
  ( ! [X169,X170] : (~sP8(happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X170)))),X169) | p(s(bool,X169))) )),
  inference(resolution,[],[f539,f326])).
fof(f27071,plain,(
  ( ! [X37] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0') | p(s(bool,X37))) ) | $spl144),
  inference(duplicate_literal_removal,[],[f26926])).
fof(f26926,plain,(
  ( ! [X37] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0') | p(s(bool,X37)) | p(s(bool,X37))) ) | $spl144),
  inference(resolution,[],[f21145,f594])).
fof(f594,plain,(
  ( ! [X116] : (~sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X116) | p(s(bool,X116))) )),
  inference(resolution,[],[f251,f326])).
fof(f22171,plain,(
  sP8('F_0',happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13))))) | $spl238),
  inference(cnf_transformation,[],[f22171_D])).
fof(f22171_D,plain,(
  sP8('F_0',happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13))))) <=> ~$spl238),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl238])])).
fof(f21669,plain,(
  ( ! [X22] : (sP8(X22,'T_0')) ) | $spl222),
  inference(cnf_transformation,[],[f21669_D])).
fof(f21669_D,plain,(
  ( ! [X22] : (sP8(X22,'T_0')) ) <=> ~$spl222),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl222])])).
fof(f26878,plain,(
  ( ! [X39] : (sP8(X39,'T_0') | sP8('F_0',happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))))) ) | ($spl144 | $spl161)),
  inference(subsumption_resolution,[],[f26837,f21077])).
fof(f26837,plain,(
  ( ! [X39] : (sP8(X39,'T_0') | sP8('F_0',happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13))))) | ~p(s(bool,X39))) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f26296,f2622])).
fof(f2622,plain,(
  ( ! [X114] : (~sP8(happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),X114) | ~p(s(bool,X114))) )),
  inference(resolution,[],[f1246,f327])).
fof(f327,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | ~sP8(X0,X1)) )),
  inference(cnf_transformation,[],[f206])).
fof(f26296,plain,(
  ( ! [X287,X286] : (sP8(X286,'T_0') | sP8(X287,X286) | sP8('F_0',X287)) ) | ($spl144 | $spl161)),
  inference(subsumption_resolution,[],[f26255,f417])).
fof(f26255,plain,(
  ( ! [X287,X286] : (~p(s(bool,'T_0')) | sP8(X286,'T_0') | sP8(X287,X286) | sP8('F_0',X287)) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f21144,f25472])).
fof(f25472,plain,(
  ( ! [X292,X293] : (sP8('T_0',X292) | sP8(X293,X292) | sP8('F_0',X293)) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f20549,f21837])).
fof(f21837,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | sP8('F_0',X1)) ) | ($spl144 | $spl161)),
  inference(subsumption_resolution,[],[f21807,f377])).
fof(f377,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f227])).
fof(f227,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))))),
  inference(flattening,[],[f226])).
fof(f226,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f143])).
fof(f143,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f90])).
fof(f90,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X13] : ((p(s(bool,X13)) => p(s(bool,X13))) <=> p(s(bool,'T_0')))),
  file('thm.quantHeuristics.GUESSES_WEAKEN_THM|split|3','thm.bool.IMP_CLAUSES|split|3')).
fof(f21807,plain,(
  ( ! [X1] : (sP8('F_0',X1) | ~p(s(bool,X1)) | ~p(s(bool,'T_0'))) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f21782,f327])).
fof(f21782,plain,(
  ( ! [X22] : (sP8(X22,'T_0') | sP8('F_0',X22)) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f21171,f19599])).
fof(f19599,plain,(
  ~sP8('F_0','F_0') | $spl161),
  inference(cnf_transformation,[],[f19599_D])).
fof(f19599_D,plain,(
  ~sP8('F_0','F_0') <=> ~$spl161),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl161])])).
fof(f21171,plain,(
  ( ! [X210,X211] : (sP8(X210,'T_0') | sP8(X211,X210) | sP8(X211,'F_0')) ) | $spl144),
  inference(resolution,[],[f21077,f19503])).
fof(f20549,plain,(
  ( ! [X167,X168] : (sP8('T_0',X167) | sP8(X168,X167) | p(s(bool,X168))) ) | $spl144),
  inference(resolution,[],[f20335,f331])).
fof(f20335,plain,(
  ( ! [X0] : (p(s(bool,X0)) | sP8('T_0',X0)) ) | $spl144),
  inference(subsumption_resolution,[],[f20316,f386])).
fof(f20316,plain,(
  ( ! [X0] : (sP8('T_0',X0) | p(s(bool,'F_0')) | p(s(bool,X0))) ) | $spl144),
  inference(resolution,[],[f20090,f326])).
fof(f20090,plain,(
  ( ! [X319] : (sP8('T_0',X319) | sP8(X319,'F_0')) ) | $spl144),
  inference(resolution,[],[f19502,f417])).
fof(f19502,plain,(
  ( ! [X171,X172] : (~p(s(bool,X172)) | sP8(X172,X171) | sP8(X171,'F_0')) ) | $spl144),
  inference(resolution,[],[f19403,f328])).
fof(f21144,plain,(
  ( ! [X163,X164] : (~sP8(X164,X163) | ~p(s(bool,X164)) | sP8(X163,'T_0')) ) | $spl144),
  inference(resolution,[],[f21077,f327])).
fof(f22056,plain,(
  sP8('F_0',happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11)))) | $spl232),
  inference(cnf_transformation,[],[f22056_D])).
fof(f22056_D,plain,(
  sP8('F_0',happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11)))) <=> ~$spl232),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl232])])).
fof(f26877,plain,(
  ( ! [X38] : (sP8(X38,'T_0') | sP8('F_0',happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))))) ) | ($spl144 | $spl161)),
  inference(subsumption_resolution,[],[f26836,f21077])).
fof(f26836,plain,(
  ( ! [X38] : (sP8(X38,'T_0') | sP8('F_0',happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11)))) | ~p(s(bool,X38))) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f26296,f1300])).
fof(f1300,plain,(
  ( ! [X114] : (~sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),X114) | ~p(s(bool,X114))) )),
  inference(resolution,[],[f538,f327])).
fof(f22002,plain,(
  sP8('F_0','const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | $spl230),
  inference(cnf_transformation,[],[f22002_D])).
fof(f22002_D,plain,(
  sP8('F_0','const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) <=> ~$spl230),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl230])])).
fof(f26876,plain,(
  ( ! [X36] : (sP8(X36,'T_0') | sP8('F_0','const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) ) | ($spl144 | $spl161)),
  inference(subsumption_resolution,[],[f26834,f21077])).
fof(f26834,plain,(
  ( ! [X36] : (sP8(X36,'T_0') | sP8('F_0','const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | ~p(s(bool,X36))) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f26296,f483])).
fof(f483,plain,(
  ( ! [X111] : (~sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X111) | ~p(s(bool,X111))) )),
  inference(resolution,[],[f250,f327])).
fof(f25272,plain,(
  ( ! [X33] : (sP8('F_0',X33)) ) | $spl298),
  inference(cnf_transformation,[],[f25272_D])).
fof(f25272_D,plain,(
  ( ! [X33] : (sP8('F_0',X33)) ) <=> ~$spl298),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl298])])).
fof(f26875,plain,(
  ( ! [X35,X34] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X34)))),'T_0') | sP8('F_0',X35)) ) | ($spl144 | $spl161)),
  inference(subsumption_resolution,[],[f26831,f21837])).
fof(f26831,plain,(
  ( ! [X35,X34] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X34)))),'T_0') | sP8('F_0',X35) | p(s(bool,X35))) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f26296,f3042])).
fof(f26874,plain,(
  ( ! [X31] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0') | sP8('F_0',X31)) ) | ($spl144 | $spl161)),
  inference(subsumption_resolution,[],[f26828,f21837])).
fof(f26828,plain,(
  ( ! [X31] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0') | sP8('F_0',X31) | p(s(bool,X31))) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f26296,f593])).
fof(f25267,plain,(
  ( ! [X17,X16] : (sP8('F_0',X16) | s(bool,'F_0') = s(bool,X17) | sP8(X17,X16)) ) | $spl296),
  inference(cnf_transformation,[],[f25267_D])).
fof(f25267_D,plain,(
  ( ! [X17,X16] : (sP8('F_0',X16) | s(bool,'F_0') = s(bool,X17) | sP8(X17,X16)) ) <=> ~$spl296),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl296])])).
fof(f26873,plain,(
  ( ! [X14,X12,X13] : (sP8(X12,'T_0') | sP8('F_0',X13) | s(bool,'F_0') = s(bool,X14) | sP8(X14,X13)) ) | ($spl144 | $spl161)),
  inference(subsumption_resolution,[],[f26819,f21077])).
fof(f26819,plain,(
  ( ! [X14,X12,X13] : (sP8(X12,'T_0') | sP8('F_0',X13) | s(bool,'F_0') = s(bool,X14) | ~p(s(bool,X12)) | sP8(X14,X13)) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f26296,f16905])).
fof(f16905,plain,(
  ( ! [X212,X213,X214] : (~sP8(X213,X214) | s(bool,'F_0') = s(bool,X212) | ~p(s(bool,X214)) | sP8(X212,X213)) )),
  inference(resolution,[],[f16178,f327])).
fof(f16178,plain,(
  ( ! [X434,X435] : (sP8(X434,X435) | p(s(bool,X435)) | s(bool,'F_0') = s(bool,X434)) )),
  inference(subsumption_resolution,[],[f16177,f370])).
fof(f370,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f223])).
fof(f223,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,'T_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f222])).
fof(f222,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (~p(s(bool,'T_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f141])).
fof(f141,plain,(
  ! [X0] : ((~p(s(bool,'T_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(ennf_transformation,[],[f88])).
fof(f88,plain,(
  ! [X0] : ((p(s(bool,'T_0')) => p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X13] : ((p(s(bool,'T_0')) => p(s(bool,X13))) <=> p(s(bool,X13)))),
  file('thm.quantHeuristics.GUESSES_WEAKEN_THM|split|3','thm.bool.IMP_CLAUSES|split|0')).
fof(f16177,plain,(
  ( ! [X434,X435] : (~p(s(bool,'T_0')) | sP8(X434,X435) | p(s(bool,X435)) | s(bool,'F_0') = s(bool,X434)) )),
  inference(forward_demodulation,[],[f15581,f504])).
fof(f504,plain,(
  s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))),
  inference(resolution,[],[f250,f402])).
fof(f15581,plain,(
  ( ! [X434,X435] : (~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) | sP8(X434,X435) | p(s(bool,X435)) | s(bool,'F_0') = s(bool,X434)) )),
  inference(superposition,[],[f328,f14557])).
fof(f14557,plain,(
  ( ! [X24] : (s(bool,X24) = s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | s(bool,'F_0') = s(bool,X24)) )),
  inference(trivial_inequality_removal,[],[f14537])).
fof(f14537,plain,(
  ( ! [X24] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X24) = s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | s(bool,'F_0') = s(bool,X24)) )),
  inference(superposition,[],[f14143,f424])).
fof(f424,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.quantHeuristics.GUESSES_WEAKEN_THM|split|3',aHL_BOOL_CASES)).
fof(f14143,plain,(
  ( ! [X57] : (s(bool,'T_0') != s(bool,X57) | s(bool,X57) = s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) )),
  inference(subsumption_resolution,[],[f14142,f417])).
fof(f14142,plain,(
  ( ! [X57] : (~p(s(bool,'T_0')) | s(bool,'T_0') != s(bool,X57) | s(bool,X57) = s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) )),
  inference(forward_demodulation,[],[f14092,f504])).
fof(f14092,plain,(
  ( ! [X57] : (s(bool,'T_0') != s(bool,X57) | ~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) | s(bool,X57) = s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) )),
  inference(resolution,[],[f1700,f482])).
fof(f482,plain,(
  ( ! [X109,X110] : (~sP7('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X110,X109) | s(bool,X109) = s(bool,X110)) )),
  inference(resolution,[],[f250,f318])).
fof(f318,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X2)) | ~sP7(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f201])).
fof(f201,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP7(X2,X1,X0)) & (~sP7(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f163])).
fof(f163,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP7(X2,X1,X0))),
  inference(definition_folding,[],[f64,f162,f161,f160,f159])).
fof(f162,plain,(
  ! [X2,X1,X0] : (sP7(X2,X1,X0) <=> (sP4(X0,X1,X2) & sP5(X1,X0,X2) & sP6(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f64,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f63])).
fof(f63,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X22,X23,X24] : ((p(s(bool,X24)) <=> s(bool,X22) = s(bool,X23)) <=> ((p(s(bool,X24)) | p(s(bool,X23)) | p(s(bool,X22))) & (p(s(bool,X24)) | ~p(s(bool,X22)) | ~p(s(bool,X23))) & (p(s(bool,X23)) | ~p(s(bool,X22)) | ~p(s(bool,X24))) & (p(s(bool,X22)) | ~p(s(bool,X23)) | ~p(s(bool,X24)))))),
  file('thm.quantHeuristics.GUESSES_WEAKEN_THM|split|3','thm.sat.sat_2|49|_prove_ltb')).
fof(f1700,plain,(
  ( ! [X279,X278] : (sP7(X279,X278,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | s(bool,'T_0') != s(bool,X278) | ~p(s(bool,X279))) )),
  inference(superposition,[],[f317,f504])).
fof(f317,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f201])).
fof(f21841,plain,(
  ( ! [X9] : (s(bool,'F_0') = s(bool,X9) | sP8(X9,'T_0')) ) | $spl224),
  inference(cnf_transformation,[],[f21841_D])).
fof(f21841_D,plain,(
  ( ! [X9] : (s(bool,'F_0') = s(bool,X9) | sP8(X9,'T_0')) ) <=> ~$spl224),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl224])])).
fof(f25871,plain,(
  ( ! [X14,X13] : (sP8(X13,X14) | ~p(s(bool,X14)) | sP8('F_0',X13)) ) | $spl314),
  inference(cnf_transformation,[],[f25871_D])).
fof(f25871_D,plain,(
  ( ! [X14,X13] : (sP8(X13,X14) | ~p(s(bool,X14)) | sP8('F_0',X13)) ) <=> ~$spl314),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl314])])).
fof(f26800,plain,(
  ( ! [X17,X15,X16] : (sP8(X15,X16) | sP8('F_0',X15) | s(bool,'F_0') = s(bool,X17) | ~p(s(bool,X16)) | sP8(X17,'T_0')) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f26296,f16906])).
fof(f16906,plain,(
  ( ! [X216,X215,X217] : (~sP8(X217,X216) | s(bool,'F_0') = s(bool,X215) | ~p(s(bool,X217)) | sP8(X215,X216)) )),
  inference(resolution,[],[f16178,f327])).
fof(f26754,plain,(
  ( ! [X41,X40] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X40)))),'T_0') | sP8('F_0',X41)) ) | ($spl144 | $spl161)),
  inference(subsumption_resolution,[],[f26713,f21837])).
fof(f26713,plain,(
  ( ! [X41,X40] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X40)))),'T_0') | sP8('F_0',X41) | p(s(bool,X41))) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f26293,f3043])).
fof(f26293,plain,(
  ( ! [X279,X278] : (sP8(X278,'T_0') | sP8(X278,X279) | sP8('F_0',X279)) ) | ($spl144 | $spl161)),
  inference(subsumption_resolution,[],[f26252,f417])).
fof(f26252,plain,(
  ( ! [X279,X278] : (~p(s(bool,'T_0')) | sP8(X278,'T_0') | sP8(X278,X279) | sP8('F_0',X279)) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f21144,f24849])).
fof(f24849,plain,(
  ( ! [X292,X293] : (sP8('T_0',X292) | sP8(X292,X293) | sP8('F_0',X293)) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f20548,f21837])).
fof(f20548,plain,(
  ( ! [X165,X166] : (sP8('T_0',X165) | sP8(X165,X166) | p(s(bool,X166))) ) | $spl144),
  inference(resolution,[],[f20335,f328])).
fof(f26753,plain,(
  ( ! [X37] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0') | sP8('F_0',X37)) ) | ($spl144 | $spl161)),
  inference(subsumption_resolution,[],[f26710,f21837])).
fof(f26710,plain,(
  ( ! [X37] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0') | sP8('F_0',X37) | p(s(bool,X37))) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f26293,f594])).
fof(f26752,plain,(
  ( ! [X33] : (sP8(X33,'T_0') | sP8('F_0',happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))))) ) | ($spl144 | $spl161)),
  inference(subsumption_resolution,[],[f26705,f21077])).
fof(f26705,plain,(
  ( ! [X33] : (sP8(X33,'T_0') | sP8('F_0',happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13))))) | ~p(s(bool,X33))) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f26293,f2623])).
fof(f2623,plain,(
  ( ! [X115] : (~sP8(X115,happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13))))) | ~p(s(bool,X115))) )),
  inference(resolution,[],[f1246,f327])).
fof(f26751,plain,(
  ( ! [X32] : (sP8(X32,'T_0') | sP8('F_0',happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))))) ) | ($spl144 | $spl161)),
  inference(subsumption_resolution,[],[f26704,f21077])).
fof(f26704,plain,(
  ( ! [X32] : (sP8(X32,'T_0') | sP8('F_0',happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11)))) | ~p(s(bool,X32))) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f26293,f1301])).
fof(f1301,plain,(
  ( ! [X115] : (~sP8(X115,happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11)))) | ~p(s(bool,X115))) )),
  inference(resolution,[],[f538,f327])).
fof(f26750,plain,(
  ( ! [X30] : (sP8(X30,'T_0') | sP8('F_0','const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) ) | ($spl144 | $spl161)),
  inference(subsumption_resolution,[],[f26702,f21077])).
fof(f26702,plain,(
  ( ! [X30] : (sP8(X30,'T_0') | sP8('F_0','const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | ~p(s(bool,X30))) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f26293,f484])).
fof(f484,plain,(
  ( ! [X112] : (~sP8(X112,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | ~p(s(bool,X112))) )),
  inference(resolution,[],[f250,f327])).
fof(f26749,plain,(
  ( ! [X17,X15,X16] : (sP8(X15,'T_0') | sP8('F_0',X16) | s(bool,'F_0') = s(bool,X17) | sP8(X17,X16)) ) | ($spl144 | $spl161)),
  inference(subsumption_resolution,[],[f26695,f21077])).
fof(f26695,plain,(
  ( ! [X17,X15,X16] : (sP8(X15,'T_0') | sP8('F_0',X16) | s(bool,'F_0') = s(bool,X17) | ~p(s(bool,X15)) | sP8(X17,X16)) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f26293,f16906])).
fof(f25260,plain,(
  ( ! [X14,X13] : (sP8(X13,X14) | ~p(s(bool,X13)) | sP8('F_0',X14)) ) | $spl294),
  inference(cnf_transformation,[],[f25260_D])).
fof(f25260_D,plain,(
  ( ! [X14,X13] : (sP8(X13,X14) | ~p(s(bool,X13)) | sP8('F_0',X14)) ) <=> ~$spl294),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl294])])).
fof(f26675,plain,(
  ( ! [X17,X15,X16] : (sP8(X15,X16) | sP8('F_0',X16) | s(bool,'F_0') = s(bool,X17) | ~p(s(bool,X15)) | sP8(X17,'T_0')) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f26293,f16906])).
fof(f26598,plain,(
  ( ! [X14,X15] : (sP8(X14,'T_0') | sP8(happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X15)))),X14)) ) | $spl326),
  inference(cnf_transformation,[],[f26598_D])).
fof(f26598_D,plain,(
  ( ! [X14,X15] : (sP8(X14,'T_0') | sP8(happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X15)))),X14)) ) <=> ~$spl326),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl326])])).
fof(f26652,plain,(
  ( ! [X19,X20,X18] : (sP8(X18,'T_0') | sP8(happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X19)))),X18) | p(s(bool,X20))) ) | $spl144),
  inference(resolution,[],[f26295,f3090])).
fof(f3090,plain,(
  ( ! [X185,X186] : (~sP9(happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X186)))),X185) | p(s(bool,X185))) )),
  inference(subsumption_resolution,[],[f3051,f386])).
fof(f3051,plain,(
  ( ! [X185,X186] : (p(s(bool,'F_0')) | p(s(bool,X185)) | ~sP9(happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X186)))),X185)) )),
  inference(resolution,[],[f539,f342])).
fof(f342,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP9(X0,X1)) )),
  inference(cnf_transformation,[],[f211])).
fof(f211,plain,(
  ! [X0,X1] : ((~sP9(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP9(X0,X1)))),
  inference(flattening,[],[f210])).
fof(f210,plain,(
  ! [X0,X1] : ((~sP9(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP9(X0,X1)))),
  inference(nnf_transformation,[],[f166])).
fof(f166,plain,(
  ! [X0,X1] : (sP9(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f26295,plain,(
  ( ! [X285,X283,X284] : (sP8(X283,'T_0') | sP8(X284,X283) | sP9(X284,X285)) ) | $spl144),
  inference(subsumption_resolution,[],[f26254,f417])).
fof(f26254,plain,(
  ( ! [X285,X283,X284] : (~p(s(bool,'T_0')) | sP8(X283,'T_0') | sP8(X284,X283) | sP9(X284,X285)) ) | $spl144),
  inference(resolution,[],[f21144,f25452])).
fof(f25452,plain,(
  ( ! [X241,X239,X240] : (sP8('T_0',X239) | sP8(X240,X239) | sP9(X240,X241)) ) | $spl144),
  inference(resolution,[],[f20549,f344])).
fof(f344,plain,(
  ( ! [X0,X1] : (sP9(X0,X1) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f211])).
fof(f26597,plain,(
  ( ! [X12] : (sP8(X12,'T_0') | sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X12)) ) | $spl324),
  inference(cnf_transformation,[],[f26597_D])).
fof(f26597_D,plain,(
  ( ! [X12] : (sP8(X12,'T_0') | sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X12)) ) <=> ~$spl324),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl324])])).
fof(f26651,plain,(
  ( ! [X17,X16] : (sP8(X16,'T_0') | sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X16) | p(s(bool,X17))) ) | $spl144),
  inference(resolution,[],[f26295,f640])).
fof(f640,plain,(
  ( ! [X124] : (~sP9('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X124) | p(s(bool,X124))) )),
  inference(subsumption_resolution,[],[f602,f386])).
fof(f602,plain,(
  ( ! [X124] : (p(s(bool,'F_0')) | p(s(bool,X124)) | ~sP9('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X124)) )),
  inference(resolution,[],[f251,f342])).
fof(f26594,plain,(
  ( ! [X0,X1] : (sP8(X0,'T_0') | p(s(bool,X1)) | sP8(X1,X0)) ) | $spl322),
  inference(cnf_transformation,[],[f26594_D])).
fof(f26594_D,plain,(
  ( ! [X0,X1] : (sP8(X0,'T_0') | p(s(bool,X1)) | sP8(X1,X0)) ) <=> ~$spl322),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl322])])).
fof(f26667,plain,(
  ( ! [X10,X11,X9] : (sP8(X9,'T_0') | sP8(X10,X9) | p(s(bool,X10)) | p(s(bool,X11))) ) | $spl144),
  inference(subsumption_resolution,[],[f26648,f386])).
fof(f26648,plain,(
  ( ! [X10,X11,X9] : (sP8(X9,'T_0') | sP8(X10,X9) | p(s(bool,'F_0')) | p(s(bool,X10)) | p(s(bool,X11))) ) | $spl144),
  inference(resolution,[],[f26295,f342])).
fof(f26666,plain,(
  ( ! [X6,X8,X7] : (sP8(X6,'T_0') | sP8(X7,X6) | p(s(bool,X7)) | p(s(bool,X8))) ) | $spl144),
  inference(subsumption_resolution,[],[f26647,f386])).
fof(f26647,plain,(
  ( ! [X6,X8,X7] : (sP8(X6,'T_0') | sP8(X7,X6) | p(s(bool,'F_0')) | p(s(bool,X7)) | p(s(bool,X8))) ) | $spl144),
  inference(resolution,[],[f26295,f346])).
fof(f346,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP9(X0,X1)) )),
  inference(cnf_transformation,[],[f213])).
fof(f213,plain,(
  ! [X0,X1] : ((~sP9(X0,X1) | (p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP9(X0,X1)))),
  inference(flattening,[],[f212])).
fof(f212,plain,(
  ! [X0,X1] : ((~sP9(X0,X1) | ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0')))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP9(X0,X1)))),
  inference(nnf_transformation,[],[f167])).
fof(f167,plain,(
  ! [X0,X1] : (sP9(X0,X1) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(definition_folding,[],[f129,f166])).
fof(f129,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f128])).
fof(f128,plain,(
  ! [X0,X1] : (((p(s(bool,X1)) | p(s(bool,X0))) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | (p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X7,X6] : ((~(p(s(bool,X6)) | p(s(bool,X7))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X6)) => p(s(bool,'F_0'))) => (~p(s(bool,X7)) => p(s(bool,'F_0')))))),
  file('thm.quantHeuristics.GUESSES_WEAKEN_THM|split|3','thm.sat.sat_9_prove_ltb')).
fof(f20883,plain,(
  ( ! [X29] : (sP8(X29,'F_0')) ) | $spl202),
  inference(cnf_transformation,[],[f20883_D])).
fof(f20883_D,plain,(
  ( ! [X29] : (sP8(X29,'F_0')) ) <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f26646,plain,(
  ( ! [X4,X5,X3] : (sP8(X3,'T_0') | sP8(X4,X3) | p(s(bool,X4)) | sP8(X5,'F_0')) ) | $spl144),
  inference(resolution,[],[f26295,f19589])).
fof(f19589,plain,(
  ( ! [X184,X183] : (~sP9(X184,X183) | p(s(bool,X184)) | sP8(X183,'F_0')) ) | $spl144),
  inference(subsumption_resolution,[],[f19508,f386])).
fof(f19508,plain,(
  ( ! [X184,X183] : (sP8(X183,'F_0') | p(s(bool,'F_0')) | p(s(bool,X184)) | ~sP9(X184,X183)) ) | $spl144),
  inference(resolution,[],[f19403,f342])).
fof(f16997,plain,(
  ( ! [X290,X289] : (sP9(X289,X290)) ) | $spl102),
  inference(cnf_transformation,[],[f16997_D])).
fof(f16997_D,plain,(
  ( ! [X290,X289] : (sP9(X289,X290)) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f26665,plain,(
  ( ! [X54,X52,X53] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X52)))),'T_0') | sP9(X53,X54)) ) | $spl144),
  inference(subsumption_resolution,[],[f26637,f344])).
fof(f26637,plain,(
  ( ! [X54,X52,X53] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X52)))),'T_0') | sP9(X53,X54) | p(s(bool,X53))) ) | $spl144),
  inference(resolution,[],[f26295,f3042])).
fof(f26664,plain,(
  ( ! [X47,X46] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0') | sP9(X46,X47)) ) | $spl144),
  inference(subsumption_resolution,[],[f26634,f344])).
fof(f26634,plain,(
  ( ! [X47,X46] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0') | sP9(X46,X47) | p(s(bool,X46))) ) | $spl144),
  inference(resolution,[],[f26295,f593])).
fof(f25689,plain,(
  ( ! [X21,X22,X20] : (sP8(X20,X21) | ~p(s(bool,X21)) | sP9(X20,X22)) ) | $spl310),
  inference(cnf_transformation,[],[f25689_D])).
fof(f25689_D,plain,(
  ( ! [X21,X22,X20] : (sP8(X20,X21) | ~p(s(bool,X21)) | sP9(X20,X22)) ) <=> ~$spl310),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl310])])).
fof(f26606,plain,(
  ( ! [X24,X23,X25,X22] : (sP8(X22,X23) | sP9(X22,X24) | s(bool,'F_0') = s(bool,X25) | ~p(s(bool,X23)) | sP8(X25,'T_0')) ) | $spl144),
  inference(resolution,[],[f26295,f16906])).
fof(f26578,plain,(
  ( ! [X14,X15,X16] : (sP8(X14,'T_0') | sP8(happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X15)))),X14) | p(s(bool,X16))) ) | $spl144),
  inference(resolution,[],[f26294,f3089])).
fof(f3089,plain,(
  ( ! [X184,X183] : (~sP9(X183,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X184))))) | p(s(bool,X183))) )),
  inference(subsumption_resolution,[],[f3050,f386])).
fof(f3050,plain,(
  ( ! [X184,X183] : (p(s(bool,'F_0')) | p(s(bool,X183)) | ~sP9(X183,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X184)))))) )),
  inference(resolution,[],[f539,f342])).
fof(f26294,plain,(
  ( ! [X280,X282,X281] : (sP8(X280,'T_0') | sP8(X281,X280) | sP9(X282,X281)) ) | $spl144),
  inference(subsumption_resolution,[],[f26253,f417])).
fof(f26253,plain,(
  ( ! [X280,X282,X281] : (~p(s(bool,'T_0')) | sP8(X280,'T_0') | sP8(X281,X280) | sP9(X282,X281)) ) | $spl144),
  inference(resolution,[],[f21144,f25451])).
fof(f25451,plain,(
  ( ! [X237,X238,X236] : (sP8('T_0',X236) | sP8(X237,X236) | sP9(X238,X237)) ) | $spl144),
  inference(resolution,[],[f20549,f343])).
fof(f343,plain,(
  ( ! [X0,X1] : (sP9(X0,X1) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f211])).
fof(f26577,plain,(
  ( ! [X12,X13] : (sP8(X12,'T_0') | sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X12) | p(s(bool,X13))) ) | $spl144),
  inference(resolution,[],[f26294,f639])).
fof(f639,plain,(
  ( ! [X123] : (~sP9(X123,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | p(s(bool,X123))) )),
  inference(subsumption_resolution,[],[f601,f386])).
fof(f601,plain,(
  ( ! [X123] : (p(s(bool,'F_0')) | p(s(bool,X123)) | ~sP9(X123,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) )),
  inference(resolution,[],[f251,f342])).
fof(f26596,plain,(
  ( ! [X10,X11,X9] : (sP8(X9,'T_0') | sP8(X10,X9) | p(s(bool,X11)) | p(s(bool,X10))) ) | $spl144),
  inference(subsumption_resolution,[],[f26576,f386])).
fof(f26576,plain,(
  ( ! [X10,X11,X9] : (sP8(X9,'T_0') | sP8(X10,X9) | p(s(bool,'F_0')) | p(s(bool,X11)) | p(s(bool,X10))) ) | $spl144),
  inference(resolution,[],[f26294,f342])).
fof(f26595,plain,(
  ( ! [X6,X8,X7] : (sP8(X6,'T_0') | sP8(X7,X6) | p(s(bool,X8)) | p(s(bool,X7))) ) | $spl144),
  inference(subsumption_resolution,[],[f26575,f386])).
fof(f26575,plain,(
  ( ! [X6,X8,X7] : (sP8(X6,'T_0') | sP8(X7,X6) | p(s(bool,'F_0')) | p(s(bool,X8)) | p(s(bool,X7))) ) | $spl144),
  inference(resolution,[],[f26294,f346])).
fof(f26573,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,'T_0') | sP8(X1,X0) | p(s(bool,X1)) | sP8(X2,'F_0')) ) | $spl144),
  inference(resolution,[],[f26294,f19590])).
fof(f19590,plain,(
  ( ! [X185,X186] : (~sP9(X185,X186) | p(s(bool,X186)) | sP8(X185,'F_0')) ) | $spl144),
  inference(subsumption_resolution,[],[f19509,f386])).
fof(f19509,plain,(
  ( ! [X185,X186] : (sP8(X185,'F_0') | p(s(bool,'F_0')) | p(s(bool,X186)) | ~sP9(X185,X186)) ) | $spl144),
  inference(resolution,[],[f19403,f342])).
fof(f26593,plain,(
  ( ! [X54,X52,X53] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X52)))),'T_0') | sP9(X53,X54)) ) | $spl144),
  inference(subsumption_resolution,[],[f26565,f343])).
fof(f26565,plain,(
  ( ! [X54,X52,X53] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X52)))),'T_0') | sP9(X53,X54) | p(s(bool,X54))) ) | $spl144),
  inference(resolution,[],[f26294,f3042])).
fof(f26592,plain,(
  ( ! [X47,X46] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0') | sP9(X46,X47)) ) | $spl144),
  inference(subsumption_resolution,[],[f26562,f343])).
fof(f26562,plain,(
  ( ! [X47,X46] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0') | sP9(X46,X47) | p(s(bool,X47))) ) | $spl144),
  inference(resolution,[],[f26294,f593])).
fof(f25617,plain,(
  ( ! [X21,X22,X20] : (sP8(X20,X21) | ~p(s(bool,X21)) | sP9(X22,X20)) ) | $spl302),
  inference(cnf_transformation,[],[f25617_D])).
fof(f25617_D,plain,(
  ( ! [X21,X22,X20] : (sP8(X20,X21) | ~p(s(bool,X21)) | sP9(X22,X20)) ) <=> ~$spl302),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl302])])).
fof(f26534,plain,(
  ( ! [X24,X23,X25,X22] : (sP8(X22,X23) | sP9(X24,X22) | s(bool,'F_0') = s(bool,X25) | ~p(s(bool,X23)) | sP8(X25,'T_0')) ) | $spl144),
  inference(resolution,[],[f26294,f16906])).
fof(f26457,plain,(
  ( ! [X14,X15] : (sP8(X14,'T_0') | sP8(X14,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X15)))))) ) | $spl320),
  inference(cnf_transformation,[],[f26457_D])).
fof(f26457_D,plain,(
  ( ! [X14,X15] : (sP8(X14,'T_0') | sP8(X14,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X15)))))) ) <=> ~$spl320),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl320])])).
fof(f26511,plain,(
  ( ! [X19,X20,X18] : (sP8(X18,'T_0') | sP8(X18,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X19))))) | p(s(bool,X20))) ) | $spl144),
  inference(resolution,[],[f26292,f3090])).
fof(f26292,plain,(
  ( ! [X277,X275,X276] : (sP8(X275,'T_0') | sP8(X275,X276) | sP9(X276,X277)) ) | $spl144),
  inference(subsumption_resolution,[],[f26251,f417])).
fof(f26251,plain,(
  ( ! [X277,X275,X276] : (~p(s(bool,'T_0')) | sP8(X275,'T_0') | sP8(X275,X276) | sP9(X276,X277)) ) | $spl144),
  inference(resolution,[],[f21144,f24829])).
fof(f24829,plain,(
  ( ! [X241,X239,X240] : (sP8('T_0',X239) | sP8(X239,X240) | sP9(X240,X241)) ) | $spl144),
  inference(resolution,[],[f20548,f344])).
fof(f26456,plain,(
  ( ! [X12] : (sP8(X12,'T_0') | sP8(X12,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) ) | $spl318),
  inference(cnf_transformation,[],[f26456_D])).
fof(f26456_D,plain,(
  ( ! [X12] : (sP8(X12,'T_0') | sP8(X12,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) ) <=> ~$spl318),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl318])])).
fof(f26510,plain,(
  ( ! [X17,X16] : (sP8(X16,'T_0') | sP8(X16,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | p(s(bool,X17))) ) | $spl144),
  inference(resolution,[],[f26292,f640])).
fof(f26283,plain,(
  ( ! [X45,X46] : (sP8(X45,'T_0') | p(s(bool,X46)) | sP8(X45,X46)) ) | $spl316),
  inference(cnf_transformation,[],[f26283_D])).
fof(f26283_D,plain,(
  ( ! [X45,X46] : (sP8(X45,'T_0') | p(s(bool,X46)) | sP8(X45,X46)) ) <=> ~$spl316),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl316])])).
fof(f26526,plain,(
  ( ! [X10,X11,X9] : (sP8(X9,'T_0') | sP8(X9,X10) | p(s(bool,X10)) | p(s(bool,X11))) ) | $spl144),
  inference(subsumption_resolution,[],[f26507,f386])).
fof(f26507,plain,(
  ( ! [X10,X11,X9] : (sP8(X9,'T_0') | sP8(X9,X10) | p(s(bool,'F_0')) | p(s(bool,X10)) | p(s(bool,X11))) ) | $spl144),
  inference(resolution,[],[f26292,f342])).
fof(f26525,plain,(
  ( ! [X6,X8,X7] : (sP8(X6,'T_0') | sP8(X6,X7) | p(s(bool,X7)) | p(s(bool,X8))) ) | $spl144),
  inference(subsumption_resolution,[],[f26506,f386])).
fof(f26506,plain,(
  ( ! [X6,X8,X7] : (sP8(X6,'T_0') | sP8(X6,X7) | p(s(bool,'F_0')) | p(s(bool,X7)) | p(s(bool,X8))) ) | $spl144),
  inference(resolution,[],[f26292,f346])).
fof(f26505,plain,(
  ( ! [X4,X5,X3] : (sP8(X3,'T_0') | sP8(X3,X4) | p(s(bool,X4)) | sP8(X5,'F_0')) ) | $spl144),
  inference(resolution,[],[f26292,f19589])).
fof(f26524,plain,(
  ( ! [X66,X67,X65] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X65)))),'T_0') | sP9(X66,X67)) ) | $spl144),
  inference(subsumption_resolution,[],[f26503,f344])).
fof(f26503,plain,(
  ( ! [X66,X67,X65] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X65)))),'T_0') | sP9(X66,X67) | p(s(bool,X66))) ) | $spl144),
  inference(resolution,[],[f26292,f3043])).
fof(f26523,plain,(
  ( ! [X59,X60] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0') | sP9(X59,X60)) ) | $spl144),
  inference(subsumption_resolution,[],[f26500,f344])).
fof(f26500,plain,(
  ( ! [X59,X60] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0') | sP9(X59,X60) | p(s(bool,X59))) ) | $spl144),
  inference(resolution,[],[f26292,f594])).
fof(f25050,plain,(
  ( ! [X19,X20,X18] : (sP8(X18,X19) | ~p(s(bool,X18)) | sP9(X19,X20)) ) | $spl292),
  inference(cnf_transformation,[],[f25050_D])).
fof(f25050_D,plain,(
  ( ! [X19,X20,X18] : (sP8(X18,X19) | ~p(s(bool,X18)) | sP9(X19,X20)) ) <=> ~$spl292),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl292])])).
fof(f26465,plain,(
  ( ! [X24,X23,X25,X22] : (sP8(X22,X23) | sP9(X23,X24) | s(bool,'F_0') = s(bool,X25) | ~p(s(bool,X22)) | sP8(X25,'T_0')) ) | $spl144),
  inference(resolution,[],[f26292,f16906])).
fof(f26438,plain,(
  ( ! [X14,X15,X16] : (sP8(X14,'T_0') | sP8(X14,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X15))))) | p(s(bool,X16))) ) | $spl144),
  inference(resolution,[],[f26291,f3089])).
fof(f26291,plain,(
  ( ! [X273,X272,X274] : (sP8(X272,'T_0') | sP8(X272,X273) | sP9(X274,X273)) ) | $spl144),
  inference(subsumption_resolution,[],[f26250,f417])).
fof(f26250,plain,(
  ( ! [X273,X272,X274] : (~p(s(bool,'T_0')) | sP8(X272,'T_0') | sP8(X272,X273) | sP9(X274,X273)) ) | $spl144),
  inference(resolution,[],[f21144,f24828])).
fof(f24828,plain,(
  ( ! [X237,X238,X236] : (sP8('T_0',X236) | sP8(X236,X237) | sP9(X238,X237)) ) | $spl144),
  inference(resolution,[],[f20548,f343])).
fof(f26437,plain,(
  ( ! [X12,X13] : (sP8(X12,'T_0') | sP8(X12,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | p(s(bool,X13))) ) | $spl144),
  inference(resolution,[],[f26291,f639])).
fof(f26455,plain,(
  ( ! [X10,X11,X9] : (sP8(X9,'T_0') | sP8(X9,X10) | p(s(bool,X11)) | p(s(bool,X10))) ) | $spl144),
  inference(subsumption_resolution,[],[f26436,f386])).
fof(f26436,plain,(
  ( ! [X10,X11,X9] : (sP8(X9,'T_0') | sP8(X9,X10) | p(s(bool,'F_0')) | p(s(bool,X11)) | p(s(bool,X10))) ) | $spl144),
  inference(resolution,[],[f26291,f342])).
fof(f26454,plain,(
  ( ! [X6,X8,X7] : (sP8(X6,'T_0') | sP8(X6,X7) | p(s(bool,X8)) | p(s(bool,X7))) ) | $spl144),
  inference(subsumption_resolution,[],[f26435,f386])).
fof(f26435,plain,(
  ( ! [X6,X8,X7] : (sP8(X6,'T_0') | sP8(X6,X7) | p(s(bool,'F_0')) | p(s(bool,X8)) | p(s(bool,X7))) ) | $spl144),
  inference(resolution,[],[f26291,f346])).
fof(f26433,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,'T_0') | sP8(X0,X1) | p(s(bool,X1)) | sP8(X2,'F_0')) ) | $spl144),
  inference(resolution,[],[f26291,f19590])).
fof(f26453,plain,(
  ( ! [X66,X67,X65] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X65)))),'T_0') | sP9(X66,X67)) ) | $spl144),
  inference(subsumption_resolution,[],[f26432,f343])).
fof(f26432,plain,(
  ( ! [X66,X67,X65] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X65)))),'T_0') | sP9(X66,X67) | p(s(bool,X67))) ) | $spl144),
  inference(resolution,[],[f26291,f3043])).
fof(f26452,plain,(
  ( ! [X59,X60] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0') | sP9(X59,X60)) ) | $spl144),
  inference(subsumption_resolution,[],[f26429,f343])).
fof(f26429,plain,(
  ( ! [X59,X60] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0') | sP9(X59,X60) | p(s(bool,X60))) ) | $spl144),
  inference(resolution,[],[f26291,f594])).
fof(f24986,plain,(
  ( ! [X19,X20,X18] : (sP8(X18,X19) | ~p(s(bool,X18)) | sP9(X20,X19)) ) | $spl284),
  inference(cnf_transformation,[],[f24986_D])).
fof(f24986_D,plain,(
  ( ! [X19,X20,X18] : (sP8(X18,X19) | ~p(s(bool,X18)) | sP9(X20,X19)) ) <=> ~$spl284),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl284])])).
fof(f26394,plain,(
  ( ! [X24,X23,X25,X22] : (sP8(X22,X23) | sP9(X24,X23) | s(bool,'F_0') = s(bool,X25) | ~p(s(bool,X22)) | sP8(X25,'T_0')) ) | $spl144),
  inference(resolution,[],[f26291,f16906])).
fof(f26268,plain,(
  ( ! [X259,X260] : (~p(s(bool,X259)) | sP8(happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X260)))),'T_0')) ) | $spl144),
  inference(duplicate_literal_removal,[],[f26241])).
fof(f26241,plain,(
  ( ! [X259,X260] : (~p(s(bool,X259)) | sP8(happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X260)))),'T_0') | ~p(s(bool,X259))) ) | $spl144),
  inference(resolution,[],[f21144,f3044])).
fof(f3044,plain,(
  ( ! [X171,X172] : (sP8(X171,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X172))))) | ~p(s(bool,X171))) )),
  inference(resolution,[],[f539,f328])).
fof(f26269,plain,(
  ( ! [X256] : (~p(s(bool,X256)) | sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0')) ) | $spl144),
  inference(duplicate_literal_removal,[],[f26238])).
fof(f26238,plain,(
  ( ! [X256] : (~p(s(bool,X256)) | sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0') | ~p(s(bool,X256))) ) | $spl144),
  inference(resolution,[],[f21144,f595])).
fof(f595,plain,(
  ( ! [X117] : (sP8(X117,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | ~p(s(bool,X117))) )),
  inference(resolution,[],[f251,f328])).
fof(f26282,plain,(
  ( ! [X45,X46,X44] : (~p(s(bool,X44)) | sP8(X45,'T_0') | sP8(X45,X46) | p(s(bool,X46))) ) | $spl144),
  inference(subsumption_resolution,[],[f26170,f397])).
fof(f26170,plain,(
  ( ! [X45,X46,X44] : (~p(s(bool,X44)) | sP8(X45,'T_0') | sP8(X45,X46) | p(s(bool,X46)) | s(bool,'F_0') = s(bool,X44)) ) | $spl144),
  inference(resolution,[],[f21144,f16907])).
fof(f16907,plain,(
  ( ! [X218,X220,X219] : (sP8(X218,X219) | sP8(X219,X220) | p(s(bool,X220)) | s(bool,'F_0') = s(bool,X218)) )),
  inference(resolution,[],[f16178,f328])).
fof(f26129,plain,(
  ( ! [X300,X299] : (~p(s(bool,X299)) | sP8(happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X300)))),'T_0')) ) | $spl144),
  inference(duplicate_literal_removal,[],[f26128])).
fof(f26128,plain,(
  ( ! [X300,X299] : (~p(s(bool,X299)) | sP8(happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X300)))),'T_0') | ~p(s(bool,X299))) ) | $spl144),
  inference(resolution,[],[f21143,f3045])).
fof(f3045,plain,(
  ( ! [X173,X174] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X173)))),X174) | ~p(s(bool,X174))) )),
  inference(resolution,[],[f539,f331])).
fof(f21143,plain,(
  ( ! [X161,X162] : (~sP8(X161,X162) | ~p(s(bool,X162)) | sP8(X161,'T_0')) ) | $spl144),
  inference(resolution,[],[f21077,f327])).
fof(f26130,plain,(
  ( ! [X296] : (~p(s(bool,X296)) | sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0')) ) | $spl144),
  inference(duplicate_literal_removal,[],[f26125])).
fof(f26125,plain,(
  ( ! [X296] : (~p(s(bool,X296)) | sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0') | ~p(s(bool,X296))) ) | $spl144),
  inference(resolution,[],[f21143,f596])).
fof(f596,plain,(
  ( ! [X118] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X118) | ~p(s(bool,X118))) )),
  inference(resolution,[],[f251,f331])).
fof(f20879,plain,(
  ( ! [X22] : (sP8('T_0',X22)) ) | $spl200),
  inference(cnf_transformation,[],[f20879_D])).
fof(f20879_D,plain,(
  ( ! [X22] : (sP8('T_0',X22)) ) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f25882,plain,(
  ( ! [X35] : (sP8('T_0',X35) | sP8('F_0',happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))))) ) | ($spl144 | $spl161)),
  inference(subsumption_resolution,[],[f25844,f20335])).
fof(f25844,plain,(
  ( ! [X35] : (sP8('T_0',X35) | sP8('F_0',happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13))))) | ~p(s(bool,X35))) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f25472,f2622])).
fof(f25881,plain,(
  ( ! [X34] : (sP8('T_0',X34) | sP8('F_0',happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))))) ) | ($spl144 | $spl161)),
  inference(subsumption_resolution,[],[f25843,f20335])).
fof(f25843,plain,(
  ( ! [X34] : (sP8('T_0',X34) | sP8('F_0',happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11)))) | ~p(s(bool,X34))) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f25472,f1300])).
fof(f25880,plain,(
  ( ! [X32] : (sP8('T_0',X32) | sP8('F_0','const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) ) | ($spl144 | $spl161)),
  inference(subsumption_resolution,[],[f25841,f20335])).
fof(f25841,plain,(
  ( ! [X32] : (sP8('T_0',X32) | sP8('F_0','const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | ~p(s(bool,X32))) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f25472,f483])).
fof(f20643,plain,(
  ( ! [X5] : (sP8('T_0',happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X5)))))) ) | $spl192),
  inference(cnf_transformation,[],[f20643_D])).
fof(f20643_D,plain,(
  ( ! [X5] : (sP8('T_0',happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X5)))))) ) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f25879,plain,(
  ( ! [X30,X31] : (sP8('T_0',happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X30))))) | sP8('F_0',X31)) ) | ($spl144 | $spl161)),
  inference(subsumption_resolution,[],[f25838,f21837])).
fof(f25838,plain,(
  ( ! [X30,X31] : (sP8('T_0',happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X30))))) | sP8('F_0',X31) | p(s(bool,X31))) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f25472,f3042])).
fof(f20642,plain,(
  sP8('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | $spl190),
  inference(cnf_transformation,[],[f20642_D])).
fof(f20642_D,plain,(
  sP8('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f25878,plain,(
  ( ! [X27] : (sP8('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | sP8('F_0',X27)) ) | ($spl144 | $spl161)),
  inference(subsumption_resolution,[],[f25835,f21837])).
fof(f25835,plain,(
  ( ! [X27] : (sP8('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | sP8('F_0',X27) | p(s(bool,X27))) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f25472,f593])).
fof(f25877,plain,(
  ( ! [X14,X12,X13] : (sP8('T_0',X12) | sP8('F_0',X13) | s(bool,'F_0') = s(bool,X14) | sP8(X14,X13)) ) | ($spl144 | $spl161)),
  inference(subsumption_resolution,[],[f25828,f20335])).
fof(f25828,plain,(
  ( ! [X14,X12,X13] : (sP8('T_0',X12) | sP8('F_0',X13) | s(bool,'F_0') = s(bool,X14) | ~p(s(bool,X12)) | sP8(X14,X13)) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f25472,f16905])).
fof(f25811,plain,(
  ( ! [X14,X15,X13] : (sP8(X13,X14) | sP8('F_0',X13) | s(bool,'F_0') = s(bool,X15) | ~p(s(bool,X14)) | sP8(X15,'T_0')) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f25472,f16905])).
fof(f25697,plain,(
  ( ! [X1] : (sP9('T_0',X1)) ) | $spl312),
  inference(cnf_transformation,[],[f25697_D])).
fof(f25697_D,plain,(
  ( ! [X1] : (sP9('T_0',X1)) ) <=> ~$spl312),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl312])])).
fof(f25684,plain,(
  ( ! [X0,X1] : (sP8('T_0',X0) | sP9('T_0',X1)) ) | $spl144),
  inference(factoring,[],[f25452])).
fof(f25626,plain,(
  ( ! [X14,X15] : (sP8('T_0',X14) | sP8(happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X15)))),X14)) ) | $spl306),
  inference(cnf_transformation,[],[f25626_D])).
fof(f25626_D,plain,(
  ( ! [X14,X15] : (sP8('T_0',X14) | sP8(happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X15)))),X14)) ) <=> ~$spl306),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl306])])).
fof(f25683,plain,(
  ( ! [X19,X20,X18] : (sP8('T_0',X18) | sP8(happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X19)))),X18) | p(s(bool,X20))) ) | $spl144),
  inference(resolution,[],[f25452,f3090])).
fof(f25625,plain,(
  ( ! [X12] : (sP8('T_0',X12) | sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X12)) ) | $spl304),
  inference(cnf_transformation,[],[f25625_D])).
fof(f25625_D,plain,(
  ( ! [X12] : (sP8('T_0',X12) | sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X12)) ) <=> ~$spl304),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl304])])).
fof(f25682,plain,(
  ( ! [X17,X16] : (sP8('T_0',X16) | sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X16) | p(s(bool,X17))) ) | $spl144),
  inference(resolution,[],[f25452,f640])).
fof(f25696,plain,(
  ( ! [X47,X48,X46] : (sP8('T_0',happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X46))))) | sP9(X47,X48)) ) | $spl144),
  inference(subsumption_resolution,[],[f25668,f344])).
fof(f25668,plain,(
  ( ! [X47,X48,X46] : (sP8('T_0',happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X46))))) | sP9(X47,X48) | p(s(bool,X47))) ) | $spl144),
  inference(resolution,[],[f25452,f3042])).
fof(f25695,plain,(
  ( ! [X41,X40] : (sP8('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | sP9(X40,X41)) ) | $spl144),
  inference(subsumption_resolution,[],[f25665,f344])).
fof(f25665,plain,(
  ( ! [X41,X40] : (sP8('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | sP9(X40,X41) | p(s(bool,X40))) ) | $spl144),
  inference(resolution,[],[f25452,f593])).
fof(f25641,plain,(
  ( ! [X23,X21,X22,X20] : (sP8(X20,X21) | sP9(X20,X22) | s(bool,'F_0') = s(bool,X23) | ~p(s(bool,X21)) | sP8(X23,'T_0')) ) | $spl144),
  inference(resolution,[],[f25452,f16905])).
fof(f25627,plain,(
  ( ! [X1] : (sP9(X1,'T_0')) ) | $spl308),
  inference(cnf_transformation,[],[f25627_D])).
fof(f25627_D,plain,(
  ( ! [X1] : (sP9(X1,'T_0')) ) <=> ~$spl308),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl308])])).
fof(f25612,plain,(
  ( ! [X0,X1] : (sP8('T_0',X0) | sP9(X1,'T_0')) ) | $spl144),
  inference(factoring,[],[f25451])).
fof(f25609,plain,(
  ( ! [X14,X15,X16] : (sP8('T_0',X14) | sP8(happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X15)))),X14) | p(s(bool,X16))) ) | $spl144),
  inference(resolution,[],[f25451,f3089])).
fof(f25608,plain,(
  ( ! [X12,X13] : (sP8('T_0',X12) | sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X12) | p(s(bool,X13))) ) | $spl144),
  inference(resolution,[],[f25451,f639])).
fof(f25624,plain,(
  ( ! [X47,X48,X46] : (sP8('T_0',happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X46))))) | sP9(X47,X48)) ) | $spl144),
  inference(subsumption_resolution,[],[f25596,f343])).
fof(f25596,plain,(
  ( ! [X47,X48,X46] : (sP8('T_0',happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X46))))) | sP9(X47,X48) | p(s(bool,X48))) ) | $spl144),
  inference(resolution,[],[f25451,f3042])).
fof(f25623,plain,(
  ( ! [X41,X40] : (sP8('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | sP9(X40,X41)) ) | $spl144),
  inference(subsumption_resolution,[],[f25593,f343])).
fof(f25593,plain,(
  ( ! [X41,X40] : (sP8('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | sP9(X40,X41) | p(s(bool,X41))) ) | $spl144),
  inference(resolution,[],[f25451,f593])).
fof(f25569,plain,(
  ( ! [X23,X21,X22,X20] : (sP8(X20,X21) | sP9(X22,X20) | s(bool,'F_0') = s(bool,X23) | ~p(s(bool,X21)) | sP8(X23,'T_0')) ) | $spl144),
  inference(resolution,[],[f25451,f16905])).
fof(f25552,plain,(
  ( ! [X0,X3,X1] : (sP8('T_0',X0) | s(bool,X1) != s(bool,X3) | p(s(bool,X3)) | sP8(X1,X0)) ) | $spl300),
  inference(cnf_transformation,[],[f25552_D])).
fof(f25552_D,plain,(
  ( ! [X0,X3,X1] : (sP8('T_0',X0) | s(bool,X1) != s(bool,X3) | p(s(bool,X3)) | sP8(X1,X0)) ) <=> ~$spl300),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl300])])).
fof(f25466,plain,(
  ( ! [X273,X275,X272,X274] : (sP8('T_0',X272) | sP8(X273,X272) | p(s(bool,X274)) | p(s(bool,X275)) | s(bool,X273) != s(bool,X275)) ) | $spl144),
  inference(resolution,[],[f20549,f426])).
fof(f25399,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP8('T_0',X35) | sP8(X36,X35) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X37) | s(bool,X38) != s(bool,X39)) ) | $spl144),
  inference(resolution,[],[f20549,f262])).
fof(f25396,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP8('T_0',X20) | sP8(X21,X20) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X21) != s(bool,X24)) ) | $spl144),
  inference(resolution,[],[f20549,f261])).
fof(f25395,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP8('T_0',X15) | sP8(X16,X15) | p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) ) | $spl144),
  inference(resolution,[],[f20549,f260])).
fof(f25394,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP8('T_0',X10) | sP8(X11,X10) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X12) | s(bool,X13) != s(bool,X14)) ) | $spl144),
  inference(resolution,[],[f20549,f258])).
fof(f25392,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP8('T_0',X0) | sP8(X1,X0) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) ) | $spl144),
  inference(resolution,[],[f20549,f256])).
fof(f25532,plain,(
  ( ! [X30,X31] : (sP8('T_0',happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X30))))) | p(s(bool,X31))) ) | $spl144),
  inference(duplicate_literal_removal,[],[f25384])).
fof(f25384,plain,(
  ( ! [X30,X31] : (sP8('T_0',happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X30))))) | p(s(bool,X31)) | p(s(bool,X31))) ) | $spl144),
  inference(resolution,[],[f20549,f3042])).
fof(f25533,plain,(
  ( ! [X27] : (sP8('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | p(s(bool,X27))) ) | $spl144),
  inference(duplicate_literal_removal,[],[f25381])).
fof(f25381,plain,(
  ( ! [X27] : (sP8('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | p(s(bool,X27)) | p(s(bool,X27))) ) | $spl144),
  inference(resolution,[],[f20549,f593])).
fof(f25349,plain,(
  $false | $spl200),
  inference(subsumption_resolution,[],[f25340,f418])).
fof(f418,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.quantHeuristics.GUESSES_WEAKEN_THM|split|3',aHL_TRUTH)).
fof(f25340,plain,(
  ~p(s(bool,'T_0')) | $spl200),
  inference(resolution,[],[f20879,f2623])).
fof(f25348,plain,(
  $false | $spl200),
  inference(subsumption_resolution,[],[f25339,f418])).
fof(f25339,plain,(
  ~p(s(bool,'T_0')) | $spl200),
  inference(resolution,[],[f20879,f1301])).
fof(f25347,plain,(
  $false | $spl200),
  inference(subsumption_resolution,[],[f25337,f418])).
fof(f25337,plain,(
  ~p(s(bool,'T_0')) | $spl200),
  inference(resolution,[],[f20879,f484])).
fof(f25324,plain,(
  $false | ($spl200 | $spl287)),
  inference(resolution,[],[f20879,f24988])).
fof(f24988,plain,(
  ~sP8('T_0','T_0') | $spl287),
  inference(cnf_transformation,[],[f24988_D])).
fof(f24988_D,plain,(
  ~sP8('T_0','T_0') <=> ~$spl287),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl287])])).
fof(f25296,plain,(
  ( ! [X54,X52,X53] : (p(s(bool,X52)) | p(s(bool,X53)) | sP8(happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X54)))),'T_0')) ) | $spl144),
  inference(resolution,[],[f3039,f21138])).
fof(f21138,plain,(
  ( ! [X146,X147,X148] : (sP4(X147,X148,X146) | sP8(X146,'T_0')) ) | $spl144),
  inference(resolution,[],[f21077,f313])).
fof(f25293,plain,(
  ( ! [X47,X45,X46] : (p(s(bool,X45)) | p(s(bool,X46)) | sP8('T_0',happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X47)))))) ) | $spl144),
  inference(resolution,[],[f3039,f20541])).
fof(f20541,plain,(
  ( ! [X146,X147,X148] : (sP4(X147,X148,X146) | sP8('T_0',X146)) ) | $spl144),
  inference(resolution,[],[f20335,f313])).
fof(f17405,plain,(
  ( ! [X8,X9] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X8)))),X9) | s(bool,'F_0') = s(bool,X9)) ) | $spl116),
  inference(cnf_transformation,[],[f17405_D])).
fof(f17405_D,plain,(
  ( ! [X8,X9] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X8)))),X9) | s(bool,'F_0') = s(bool,X9)) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f25290,plain,(
  ( ! [X37,X35,X38,X36] : (p(s(bool,X35)) | p(s(bool,X36)) | sP8(happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X37)))),X38) | s(bool,'F_0') = s(bool,X38)) )),
  inference(resolution,[],[f3039,f17137])).
fof(f17137,plain,(
  ( ! [X194,X192,X195,X193] : (sP8(X192,X193) | sP4(X194,X195,X192) | s(bool,'F_0') = s(bool,X193)) )),
  inference(resolution,[],[f16182,f313])).
fof(f16182,plain,(
  ( ! [X438,X439] : (sP8(X439,X438) | p(s(bool,X439)) | s(bool,'F_0') = s(bool,X438)) )),
  inference(subsumption_resolution,[],[f16181,f370])).
fof(f16181,plain,(
  ( ! [X438,X439] : (~p(s(bool,'T_0')) | sP8(X439,X438) | p(s(bool,X439)) | s(bool,'F_0') = s(bool,X438)) )),
  inference(forward_demodulation,[],[f15583,f504])).
fof(f15583,plain,(
  ( ! [X438,X439] : (~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) | sP8(X439,X438) | p(s(bool,X439)) | s(bool,'F_0') = s(bool,X438)) )),
  inference(superposition,[],[f331,f14557])).
fof(f17049,plain,(
  ( ! [X8,X9] : (sP8(X8,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X9))))) | s(bool,'F_0') = s(bool,X8)) ) | $spl108),
  inference(cnf_transformation,[],[f17049_D])).
fof(f17049_D,plain,(
  ( ! [X8,X9] : (sP8(X8,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X9))))) | s(bool,'F_0') = s(bool,X8)) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f25287,plain,(
  ( ! [X28,X26,X27,X25] : (p(s(bool,X25)) | p(s(bool,X26)) | sP8(X27,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X28))))) | s(bool,'F_0') = s(bool,X27)) )),
  inference(resolution,[],[f3039,f16900])).
fof(f16900,plain,(
  ( ! [X194,X192,X195,X193] : (sP8(X192,X193) | sP4(X194,X195,X193) | s(bool,'F_0') = s(bool,X192)) )),
  inference(resolution,[],[f16178,f313])).
fof(f25273,plain,(
  ( ! [X37,X36] : (sP8('T_0',happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X36))))) | sP8('F_0',X37)) ) | ($spl144 | $spl161)),
  inference(subsumption_resolution,[],[f25234,f21837])).
fof(f25234,plain,(
  ( ! [X37,X36] : (sP8('T_0',happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X36))))) | sP8('F_0',X37) | p(s(bool,X37))) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f24849,f3043])).
fof(f25271,plain,(
  ( ! [X33] : (sP8('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | sP8('F_0',X33)) ) | ($spl144 | $spl161)),
  inference(subsumption_resolution,[],[f25231,f21837])).
fof(f25231,plain,(
  ( ! [X33] : (sP8('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | sP8('F_0',X33) | p(s(bool,X33))) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f24849,f594])).
fof(f25270,plain,(
  ( ! [X29] : (sP8('T_0',X29) | sP8('F_0',happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))))) ) | ($spl144 | $spl161)),
  inference(subsumption_resolution,[],[f25226,f20335])).
fof(f25226,plain,(
  ( ! [X29] : (sP8('T_0',X29) | sP8('F_0',happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13))))) | ~p(s(bool,X29))) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f24849,f2623])).
fof(f25269,plain,(
  ( ! [X28] : (sP8('T_0',X28) | sP8('F_0',happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))))) ) | ($spl144 | $spl161)),
  inference(subsumption_resolution,[],[f25225,f20335])).
fof(f25225,plain,(
  ( ! [X28] : (sP8('T_0',X28) | sP8('F_0',happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11)))) | ~p(s(bool,X28))) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f24849,f1301])).
fof(f25268,plain,(
  ( ! [X26] : (sP8('T_0',X26) | sP8('F_0','const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) ) | ($spl144 | $spl161)),
  inference(subsumption_resolution,[],[f25223,f20335])).
fof(f25223,plain,(
  ( ! [X26] : (sP8('T_0',X26) | sP8('F_0','const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | ~p(s(bool,X26))) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f24849,f484])).
fof(f25266,plain,(
  ( ! [X17,X15,X16] : (sP8('T_0',X15) | sP8('F_0',X16) | s(bool,'F_0') = s(bool,X17) | sP8(X17,X16)) ) | ($spl144 | $spl161)),
  inference(subsumption_resolution,[],[f25218,f20335])).
fof(f25218,plain,(
  ( ! [X17,X15,X16] : (sP8('T_0',X15) | sP8('F_0',X16) | s(bool,'F_0') = s(bool,X17) | ~p(s(bool,X15)) | sP8(X17,X16)) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f24849,f16906])).
fof(f25200,plain,(
  ( ! [X14,X15,X13] : (sP8(X13,X14) | sP8('F_0',X14) | s(bool,'F_0') = s(bool,X15) | ~p(s(bool,X13)) | sP8(X15,'T_0')) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f24849,f16905])).
fof(f25129,plain,(
  $false | $spl286),
  inference(subsumption_resolution,[],[f25128,f418])).
fof(f25128,plain,(
  ~p(s(bool,'T_0')) | $spl286),
  inference(subsumption_resolution,[],[f25113,f399])).
fof(f399,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f241])).
fof(f241,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f100])).
fof(f100,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f99])).
fof(f99,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X13] : (s(bool,'F_0') = s(bool,X13) <=> ~p(s(bool,X13)))),
  file('thm.quantHeuristics.GUESSES_WEAKEN_THM|split|3','thm.bool.EQ_CLAUSES|split|2')).
fof(f25113,plain,(
  ~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'F_0') | $spl286),
  inference(resolution,[],[f24989,f16174])).
fof(f16174,plain,(
  ( ! [X430,X431] : (~sP8(X431,X430) | ~p(s(bool,X431)) | s(bool,'F_0') = s(bool,X430)) )),
  inference(subsumption_resolution,[],[f16173,f377])).
fof(f16173,plain,(
  ( ! [X430,X431] : (~p(s(bool,'T_0')) | ~p(s(bool,X431)) | ~sP8(X431,X430) | s(bool,'F_0') = s(bool,X430)) )),
  inference(forward_demodulation,[],[f15579,f504])).
fof(f15579,plain,(
  ( ! [X430,X431] : (~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) | ~p(s(bool,X431)) | ~sP8(X431,X430) | s(bool,'F_0') = s(bool,X430)) )),
  inference(superposition,[],[f327,f14557])).
fof(f25127,plain,(
  $false | $spl286),
  inference(subsumption_resolution,[],[f25126,f418])).
fof(f25126,plain,(
  ~p(s(bool,'T_0')) | $spl286),
  inference(subsumption_resolution,[],[f25112,f399])).
fof(f25112,plain,(
  ~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'F_0') | $spl286),
  inference(resolution,[],[f24989,f16172])).
fof(f16172,plain,(
  ( ! [X428,X429] : (~sP8(X428,X429) | ~p(s(bool,X429)) | s(bool,'F_0') = s(bool,X428)) )),
  inference(subsumption_resolution,[],[f16171,f377])).
fof(f16171,plain,(
  ( ! [X428,X429] : (~p(s(bool,'T_0')) | ~p(s(bool,X429)) | ~sP8(X428,X429) | s(bool,'F_0') = s(bool,X428)) )),
  inference(forward_demodulation,[],[f15578,f504])).
fof(f15578,plain,(
  ( ! [X428,X429] : (~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) | ~p(s(bool,X429)) | ~sP8(X428,X429) | s(bool,'F_0') = s(bool,X428)) )),
  inference(superposition,[],[f327,f14557])).
fof(f25125,plain,(
  $false | $spl286),
  inference(subsumption_resolution,[],[f25121,f418])).
fof(f25121,plain,(
  ~p(s(bool,'T_0')) | $spl286),
  inference(duplicate_literal_removal,[],[f25110])).
fof(f25110,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'T_0')) | $spl286),
  inference(resolution,[],[f24989,f334])).
fof(f334,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | ~sP8(X1,X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f207,plain,(
  ! [X0,X1] : ((((p(s(bool,X1)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)))) | sP8(X1,X0)) & (~sP8(X1,X0) | ((~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X0)) | p(s(bool,X1))))))),
  inference(nnf_transformation,[],[f165])).
fof(f165,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> sP8(X1,X0))),
  inference(definition_folding,[],[f67,f164])).
fof(f67,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  inference(flattening,[],[f66])).
fof(f66,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  inference(rectify,[],[f48])).
fof(f48,axiom,(
  ! [X23,X24] : ((p(s(bool,X24)) <=> ~p(s(bool,X23))) <=> ((p(s(bool,X24)) | p(s(bool,X23))) & (~p(s(bool,X23)) | ~p(s(bool,X24)))))),
  file('thm.quantHeuristics.GUESSES_WEAKEN_THM|split|3','thm.sat.sat_29_prove_ltb')).
fof(f25124,plain,(
  $false | $spl286),
  inference(subsumption_resolution,[],[f25122,f418])).
fof(f25122,plain,(
  ~p(s(bool,'T_0')) | $spl286),
  inference(duplicate_literal_removal,[],[f25109])).
fof(f25109,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'T_0')) | $spl286),
  inference(resolution,[],[f24989,f327])).
fof(f25079,plain,(
  ( ! [X61,X59,X60] : (p(s(bool,X59)) | p(s(bool,X60)) | sP8(happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X61)))),'T_0')) ) | $spl144),
  inference(resolution,[],[f3038,f21140])).
fof(f21140,plain,(
  ( ! [X154,X152,X153] : (sP4(X152,X153,X154) | sP8(X152,'T_0')) ) | $spl144),
  inference(resolution,[],[f21077,f315])).
fof(f25076,plain,(
  ( ! [X54,X52,X53] : (p(s(bool,X52)) | p(s(bool,X53)) | sP8('T_0',happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X54)))))) ) | $spl144),
  inference(resolution,[],[f3038,f20543])).
fof(f20543,plain,(
  ( ! [X154,X152,X153] : (sP4(X152,X153,X154) | sP8('T_0',X152)) ) | $spl144),
  inference(resolution,[],[f20335,f315])).
fof(f25073,plain,(
  ( ! [X47,X45,X46,X44] : (p(s(bool,X44)) | p(s(bool,X45)) | sP8(happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X46)))),X47) | s(bool,'F_0') = s(bool,X47)) )),
  inference(resolution,[],[f3038,f17139])).
fof(f17139,plain,(
  ( ! [X202,X200,X203,X201] : (sP8(X200,X201) | sP4(X200,X202,X203) | s(bool,'F_0') = s(bool,X201)) )),
  inference(resolution,[],[f16182,f315])).
fof(f25070,plain,(
  ( ! [X37,X35,X36,X34] : (p(s(bool,X34)) | p(s(bool,X35)) | sP8(X36,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X37))))) | s(bool,'F_0') = s(bool,X36)) )),
  inference(resolution,[],[f3038,f16902])).
fof(f16902,plain,(
  ( ! [X202,X200,X203,X201] : (sP8(X200,X201) | sP4(X201,X202,X203) | s(bool,'F_0') = s(bool,X200)) )),
  inference(resolution,[],[f16178,f315])).
fof(f24996,plain,(
  ( ! [X14,X15] : (sP8('T_0',X14) | sP8(X14,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X15)))))) ) | $spl290),
  inference(cnf_transformation,[],[f24996_D])).
fof(f24996_D,plain,(
  ( ! [X14,X15] : (sP8('T_0',X14) | sP8(X14,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X15)))))) ) <=> ~$spl290),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl290])])).
fof(f25044,plain,(
  ( ! [X19,X20,X18] : (sP8('T_0',X18) | sP8(X18,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X19))))) | p(s(bool,X20))) ) | $spl144),
  inference(resolution,[],[f24829,f3090])).
fof(f24995,plain,(
  ( ! [X12] : (sP8('T_0',X12) | sP8(X12,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) ) | $spl288),
  inference(cnf_transformation,[],[f24995_D])).
fof(f24995_D,plain,(
  ( ! [X12] : (sP8('T_0',X12) | sP8(X12,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) ) <=> ~$spl288),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl288])])).
fof(f25043,plain,(
  ( ! [X17,X16] : (sP8('T_0',X16) | sP8(X16,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | p(s(bool,X17))) ) | $spl144),
  inference(resolution,[],[f24829,f640])).
fof(f25056,plain,(
  ( ! [X59,X60,X58] : (sP8('T_0',happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X58))))) | sP9(X59,X60)) ) | $spl144),
  inference(subsumption_resolution,[],[f25036,f344])).
fof(f25036,plain,(
  ( ! [X59,X60,X58] : (sP8('T_0',happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X58))))) | sP9(X59,X60) | p(s(bool,X59))) ) | $spl144),
  inference(resolution,[],[f24829,f3043])).
fof(f25055,plain,(
  ( ! [X52,X53] : (sP8('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | sP9(X52,X53)) ) | $spl144),
  inference(subsumption_resolution,[],[f25033,f344])).
fof(f25033,plain,(
  ( ! [X52,X53] : (sP8('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | sP9(X52,X53) | p(s(bool,X52))) ) | $spl144),
  inference(resolution,[],[f24829,f594])).
fof(f24989,plain,(
  sP8('T_0','T_0') | $spl286),
  inference(cnf_transformation,[],[f24989_D])).
fof(f24989_D,plain,(
  sP8('T_0','T_0') <=> ~$spl286),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl286])])).
fof(f25007,plain,(
  ( ! [X33,X34,X32] : (sP8(X32,X33) | sP9(X33,X34) | ~p(s(bool,X32)) | sP8('T_0','T_0')) ) | $spl144),
  inference(resolution,[],[f24829,f20546])).
fof(f20546,plain,(
  ( ! [X161,X162] : (~sP8(X161,X162) | ~p(s(bool,X162)) | sP8('T_0',X161)) ) | $spl144),
  inference(resolution,[],[f20335,f327])).
fof(f25003,plain,(
  ( ! [X21,X19,X20,X18] : (sP8(X18,X19) | sP9(X19,X20) | s(bool,'F_0') = s(bool,X21) | ~p(s(bool,X18)) | sP8(X21,'T_0')) ) | $spl144),
  inference(resolution,[],[f24829,f16905])).
fof(f24978,plain,(
  ( ! [X14,X15,X16] : (sP8('T_0',X14) | sP8(X14,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X15))))) | p(s(bool,X16))) ) | $spl144),
  inference(resolution,[],[f24828,f3089])).
fof(f24977,plain,(
  ( ! [X12,X13] : (sP8('T_0',X12) | sP8(X12,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | p(s(bool,X13))) ) | $spl144),
  inference(resolution,[],[f24828,f639])).
fof(f24994,plain,(
  ( ! [X59,X60,X58] : (sP8('T_0',happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X58))))) | sP9(X59,X60)) ) | $spl144),
  inference(subsumption_resolution,[],[f24972,f343])).
fof(f24972,plain,(
  ( ! [X59,X60,X58] : (sP8('T_0',happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X58))))) | sP9(X59,X60) | p(s(bool,X60))) ) | $spl144),
  inference(resolution,[],[f24828,f3043])).
fof(f24993,plain,(
  ( ! [X52,X53] : (sP8('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | sP9(X52,X53)) ) | $spl144),
  inference(subsumption_resolution,[],[f24969,f343])).
fof(f24969,plain,(
  ( ! [X52,X53] : (sP8('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | sP9(X52,X53) | p(s(bool,X53))) ) | $spl144),
  inference(resolution,[],[f24828,f594])).
fof(f24943,plain,(
  ( ! [X33,X34,X32] : (sP8(X32,X33) | sP9(X34,X33) | ~p(s(bool,X32)) | sP8('T_0','T_0')) ) | $spl144),
  inference(resolution,[],[f24828,f20546])).
fof(f24939,plain,(
  ( ! [X21,X19,X20,X18] : (sP8(X18,X19) | sP9(X20,X19) | s(bool,'F_0') = s(bool,X21) | ~p(s(bool,X18)) | sP8(X21,'T_0')) ) | $spl144),
  inference(resolution,[],[f24828,f16905])).
fof(f24923,plain,(
  ( ! [X0,X3,X1] : (sP8('T_0',X0) | s(bool,X1) != s(bool,X3) | p(s(bool,X3)) | sP8(X0,X1)) ) | $spl282),
  inference(cnf_transformation,[],[f24923_D])).
fof(f24923_D,plain,(
  ( ! [X0,X3,X1] : (sP8('T_0',X0) | s(bool,X1) != s(bool,X3) | p(s(bool,X3)) | sP8(X0,X1)) ) <=> ~$spl282),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl282])])).
fof(f24843,plain,(
  ( ! [X273,X275,X272,X274] : (sP8('T_0',X272) | sP8(X272,X273) | p(s(bool,X274)) | p(s(bool,X275)) | s(bool,X273) != s(bool,X275)) ) | $spl144),
  inference(resolution,[],[f20548,f426])).
fof(f24776,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP8('T_0',X35) | sP8(X35,X36) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X37) | s(bool,X38) != s(bool,X39)) ) | $spl144),
  inference(resolution,[],[f20548,f262])).
fof(f24773,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP8('T_0',X20) | sP8(X20,X21) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X21) != s(bool,X24)) ) | $spl144),
  inference(resolution,[],[f20548,f261])).
fof(f24772,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP8('T_0',X15) | sP8(X15,X16) | p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) ) | $spl144),
  inference(resolution,[],[f20548,f260])).
fof(f24771,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP8('T_0',X10) | sP8(X10,X11) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X12) | s(bool,X13) != s(bool,X14)) ) | $spl144),
  inference(resolution,[],[f20548,f258])).
fof(f24769,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP8('T_0',X0) | sP8(X0,X1) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) ) | $spl144),
  inference(resolution,[],[f20548,f256])).
fof(f24903,plain,(
  ( ! [X37,X36] : (sP8('T_0',happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X36))))) | p(s(bool,X37))) ) | $spl144),
  inference(duplicate_literal_removal,[],[f24768])).
fof(f24768,plain,(
  ( ! [X37,X36] : (sP8('T_0',happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X36))))) | p(s(bool,X37)) | p(s(bool,X37))) ) | $spl144),
  inference(resolution,[],[f20548,f3043])).
fof(f24904,plain,(
  ( ! [X33] : (sP8('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | p(s(bool,X33))) ) | $spl144),
  inference(duplicate_literal_removal,[],[f24765])).
fof(f24765,plain,(
  ( ! [X33] : (sP8('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | p(s(bool,X33)) | p(s(bool,X33))) ) | $spl144),
  inference(resolution,[],[f20548,f594])).
fof(f24706,plain,(
  ( ! [X239,X240] : (~p(s(bool,X239)) | sP8('T_0',happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X240)))))) ) | $spl144),
  inference(duplicate_literal_removal,[],[f24691])).
fof(f24691,plain,(
  ( ! [X239,X240] : (~p(s(bool,X239)) | sP8('T_0',happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X240))))) | ~p(s(bool,X239))) ) | $spl144),
  inference(resolution,[],[f20547,f3044])).
fof(f20547,plain,(
  ( ! [X163,X164] : (~sP8(X164,X163) | ~p(s(bool,X164)) | sP8('T_0',X163)) ) | $spl144),
  inference(resolution,[],[f20335,f327])).
fof(f24707,plain,(
  ( ! [X236] : (~p(s(bool,X236)) | sP8('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) ) | $spl144),
  inference(duplicate_literal_removal,[],[f24688])).
fof(f24688,plain,(
  ( ! [X236] : (~p(s(bool,X236)) | sP8('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | ~p(s(bool,X236))) ) | $spl144),
  inference(resolution,[],[f20547,f595])).
fof(f24721,plain,(
  ( ! [X45,X46] : (sP8('T_0',X45) | p(s(bool,X46)) | sP8(X45,X46)) ) | $spl280),
  inference(cnf_transformation,[],[f24721_D])).
fof(f24721_D,plain,(
  ( ! [X45,X46] : (sP8('T_0',X45) | p(s(bool,X46)) | sP8(X45,X46)) ) <=> ~$spl280),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl280])])).
fof(f24720,plain,(
  ( ! [X45,X46,X44] : (~p(s(bool,X44)) | sP8('T_0',X45) | sP8(X45,X46) | p(s(bool,X46))) ) | $spl144),
  inference(subsumption_resolution,[],[f24628,f397])).
fof(f24628,plain,(
  ( ! [X45,X46,X44] : (~p(s(bool,X44)) | sP8('T_0',X45) | sP8(X45,X46) | p(s(bool,X46)) | s(bool,'F_0') = s(bool,X44)) ) | $spl144),
  inference(resolution,[],[f20547,f16907])).
fof(f24593,plain,(
  ( ! [X253,X254] : (~p(s(bool,X253)) | sP8('T_0',happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X254)))))) ) | $spl144),
  inference(duplicate_literal_removal,[],[f24592])).
fof(f24592,plain,(
  ( ! [X253,X254] : (~p(s(bool,X253)) | sP8('T_0',happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X254))))) | ~p(s(bool,X253))) ) | $spl144),
  inference(resolution,[],[f20546,f3045])).
fof(f24594,plain,(
  ( ! [X250] : (~p(s(bool,X250)) | sP8('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) ) | $spl144),
  inference(duplicate_literal_removal,[],[f24589])).
fof(f24589,plain,(
  ( ! [X250] : (~p(s(bool,X250)) | sP8('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | ~p(s(bool,X250))) ) | $spl144),
  inference(resolution,[],[f20546,f596])).
fof(f24472,plain,(
  ( ! [X59,X57,X58] : (p(s(bool,X57)) | p(s(bool,X58)) | sP8(happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X59)))),'T_0')) ) | $spl144),
  inference(resolution,[],[f3037,f21139])).
fof(f21139,plain,(
  ( ! [X151,X149,X150] : (sP4(X150,X149,X151) | sP8(X149,'T_0')) ) | $spl144),
  inference(resolution,[],[f21077,f314])).
fof(f24469,plain,(
  ( ! [X52,X50,X51] : (p(s(bool,X50)) | p(s(bool,X51)) | sP8('T_0',happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X52)))))) ) | $spl144),
  inference(resolution,[],[f3037,f20542])).
fof(f20542,plain,(
  ( ! [X151,X149,X150] : (sP4(X150,X149,X151) | sP8('T_0',X149)) ) | $spl144),
  inference(resolution,[],[f20335,f314])).
fof(f24466,plain,(
  ( ! [X43,X41,X44,X42] : (p(s(bool,X41)) | p(s(bool,X42)) | sP8(happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X43)))),X44) | s(bool,'F_0') = s(bool,X44)) )),
  inference(resolution,[],[f3037,f17138])).
fof(f17138,plain,(
  ( ! [X198,X196,X199,X197] : (sP8(X196,X197) | sP4(X198,X196,X199) | s(bool,'F_0') = s(bool,X197)) )),
  inference(resolution,[],[f16182,f314])).
fof(f24463,plain,(
  ( ! [X33,X31,X34,X32] : (p(s(bool,X31)) | p(s(bool,X32)) | sP8(X33,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X34))))) | s(bool,'F_0') = s(bool,X33)) )),
  inference(resolution,[],[f3037,f16901])).
fof(f16901,plain,(
  ( ! [X198,X196,X199,X197] : (sP8(X196,X197) | sP4(X198,X197,X199) | s(bool,'F_0') = s(bool,X196)) )),
  inference(resolution,[],[f16178,f314])).
fof(f23107,plain,(
  ( ! [X30] : (sP8(X30,happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11)))) | sP8(X30,'F_0')) ) | $spl260),
  inference(cnf_transformation,[],[f23107_D])).
fof(f23107_D,plain,(
  ( ! [X30] : (sP8(X30,happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11)))) | sP8(X30,'F_0')) ) <=> ~$spl260),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl260])])).
fof(f24436,plain,(
  ( ! [X68,X69,X67] : (sP8(X67,happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11)))) | sP8(X67,'F_0') | p(s(bool,X68)) | ~p(s(bool,X69))) ) | $spl144),
  inference(resolution,[],[f20937,f1292])).
fof(f1292,plain,(
  ( ! [X99,X98] : (~sP5(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),X98,X99) | p(s(bool,X99)) | ~p(s(bool,X98))) )),
  inference(resolution,[],[f538,f308])).
fof(f20937,plain,(
  ( ! [X198,X196,X199,X197] : (sP8(X196,X197) | sP8(X196,'F_0') | sP5(X197,X198,X199)) ) | $spl144),
  inference(resolution,[],[f19503,f311])).
fof(f311,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f198])).
fof(f21057,plain,(
  ( ! [X55] : (sP8(X55,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | sP8(X55,'F_0')) ) | $spl208),
  inference(cnf_transformation,[],[f21057_D])).
fof(f21057_D,plain,(
  ( ! [X55] : (sP8(X55,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | sP8(X55,'F_0')) ) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f24435,plain,(
  ( ! [X66,X64,X65] : (sP8(X64,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | sP8(X64,'F_0') | p(s(bool,X65)) | ~p(s(bool,X66))) ) | $spl144),
  inference(resolution,[],[f20937,f475])).
fof(f475,plain,(
  ( ! [X95,X96] : (~sP5('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X95,X96) | p(s(bool,X96)) | ~p(s(bool,X95))) )),
  inference(resolution,[],[f250,f308])).
fof(f22896,plain,(
  ( ! [X43,X44,X42] : (sP5(X42,X43,X44)) ) | $spl258),
  inference(cnf_transformation,[],[f22896_D])).
fof(f22896_D,plain,(
  ( ! [X43,X44,X42] : (sP5(X42,X43,X44)) ) <=> ~$spl258),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl258])])).
fof(f19414,plain,(
  sP8(happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),'F_0') | $spl152),
  inference(cnf_transformation,[],[f19414_D])).
fof(f19414_D,plain,(
  sP8(happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),'F_0') <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f24442,plain,(
  ( ! [X70,X71,X69] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),'F_0') | sP5(X69,X70,X71)) ) | $spl144),
  inference(subsumption_resolution,[],[f24399,f311])).
fof(f24399,plain,(
  ( ! [X70,X71,X69] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),'F_0') | sP5(X69,X70,X71) | ~p(s(bool,X69))) ) | $spl144),
  inference(resolution,[],[f20937,f2622])).
fof(f19411,plain,(
  sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),'F_0') | $spl150),
  inference(cnf_transformation,[],[f19411_D])).
fof(f19411_D,plain,(
  sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),'F_0') <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f24441,plain,(
  ( ! [X68,X66,X67] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),'F_0') | sP5(X66,X67,X68)) ) | $spl144),
  inference(subsumption_resolution,[],[f24398,f311])).
fof(f24398,plain,(
  ( ! [X68,X66,X67] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),'F_0') | sP5(X66,X67,X68) | ~p(s(bool,X66))) ) | $spl144),
  inference(resolution,[],[f20937,f1300])).
fof(f19408,plain,(
  sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | $spl148),
  inference(cnf_transformation,[],[f19408_D])).
fof(f19408_D,plain,(
  sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f24440,plain,(
  ( ! [X61,X62,X60] : (sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | sP5(X60,X61,X62)) ) | $spl144),
  inference(subsumption_resolution,[],[f24396,f311])).
fof(f24396,plain,(
  ( ! [X61,X62,X60] : (sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | sP5(X60,X61,X62) | ~p(s(bool,X60))) ) | $spl144),
  inference(resolution,[],[f20937,f483])).
fof(f24365,plain,(
  ( ! [X59,X60,X58] : (sP8(X58,happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11)))) | sP8(X58,'F_0') | p(s(bool,X59)) | ~p(s(bool,X60))) ) | $spl144),
  inference(resolution,[],[f20936,f1293])).
fof(f1293,plain,(
  ( ! [X101,X100] : (~sP5(X100,happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),X101) | p(s(bool,X101)) | ~p(s(bool,X100))) )),
  inference(resolution,[],[f538,f308])).
fof(f20936,plain,(
  ( ! [X194,X192,X195,X193] : (sP8(X192,X193) | sP8(X192,'F_0') | sP5(X194,X193,X195)) ) | $spl144),
  inference(resolution,[],[f19503,f310])).
fof(f310,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f198])).
fof(f24364,plain,(
  ( ! [X57,X56,X55] : (sP8(X55,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | sP8(X55,'F_0') | p(s(bool,X56)) | ~p(s(bool,X57))) ) | $spl144),
  inference(resolution,[],[f20936,f476])).
fof(f476,plain,(
  ( ! [X97,X98] : (~sP5(X97,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X98) | p(s(bool,X98)) | ~p(s(bool,X97))) )),
  inference(resolution,[],[f250,f308])).
fof(f24373,plain,(
  ( ! [X70,X71,X69] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),'F_0') | sP5(X69,X70,X71)) ) | $spl144),
  inference(subsumption_resolution,[],[f24331,f310])).
fof(f24331,plain,(
  ( ! [X70,X71,X69] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),'F_0') | sP5(X69,X70,X71) | ~p(s(bool,X70))) ) | $spl144),
  inference(resolution,[],[f20936,f2622])).
fof(f24372,plain,(
  ( ! [X68,X66,X67] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),'F_0') | sP5(X66,X67,X68)) ) | $spl144),
  inference(subsumption_resolution,[],[f24330,f310])).
fof(f24330,plain,(
  ( ! [X68,X66,X67] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),'F_0') | sP5(X66,X67,X68) | ~p(s(bool,X67))) ) | $spl144),
  inference(resolution,[],[f20936,f1300])).
fof(f24371,plain,(
  ( ! [X61,X62,X60] : (sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | sP5(X60,X61,X62)) ) | $spl144),
  inference(subsumption_resolution,[],[f24328,f310])).
fof(f24328,plain,(
  ( ! [X61,X62,X60] : (sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | sP5(X60,X61,X62) | ~p(s(bool,X61))) ) | $spl144),
  inference(resolution,[],[f20936,f483])).
fof(f24115,plain,(
  ( ! [X47,X48,X46] : (~p(s(bool,X46)) | ~p(s(bool,X47)) | sP8(happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X48)))),'T_0')) ) | $spl144),
  inference(resolution,[],[f3034,f21137])).
fof(f21137,plain,(
  ( ! [X144,X145,X143] : (sP5(X144,X145,X143) | sP8(X143,'T_0')) ) | $spl144),
  inference(resolution,[],[f21077,f309])).
fof(f24114,plain,(
  ( ! [X45,X43,X44] : (~p(s(bool,X43)) | ~p(s(bool,X44)) | sP8('T_0',happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X45)))))) ) | $spl144),
  inference(resolution,[],[f3034,f20540])).
fof(f20540,plain,(
  ( ! [X144,X145,X143] : (sP5(X144,X145,X143) | sP8('T_0',X143)) ) | $spl144),
  inference(resolution,[],[f20335,f309])).
fof(f24109,plain,(
  ( ! [X30,X31,X29,X32] : (~p(s(bool,X29)) | ~p(s(bool,X30)) | sP8(happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X31)))),X32) | s(bool,'F_0') = s(bool,X32)) )),
  inference(resolution,[],[f3034,f17136])).
fof(f17136,plain,(
  ( ! [X191,X189,X190,X188] : (sP8(X188,X189) | sP5(X190,X191,X188) | s(bool,'F_0') = s(bool,X189)) )),
  inference(resolution,[],[f16182,f309])).
fof(f24108,plain,(
  ( ! [X28,X26,X27,X25] : (~p(s(bool,X25)) | ~p(s(bool,X26)) | sP8(X27,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X28))))) | s(bool,'F_0') = s(bool,X27)) )),
  inference(resolution,[],[f3034,f16899])).
fof(f16899,plain,(
  ( ! [X191,X189,X190,X188] : (sP8(X188,X189) | sP5(X190,X191,X189) | s(bool,'F_0') = s(bool,X188)) )),
  inference(resolution,[],[f16178,f309])).
fof(f23935,plain,(
  ( ! [X2,X1] : (sP6('F_0',X1,X2)) ) | $spl278),
  inference(cnf_transformation,[],[f23935_D])).
fof(f23935_D,plain,(
  ( ! [X2,X1] : (sP6('F_0',X1,X2)) ) <=> ~$spl278),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl278])])).
fof(f23921,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,'F_0') | sP6('F_0',X1,X2)) ) | $spl144),
  inference(factoring,[],[f20934])).
fof(f20934,plain,(
  ( ! [X187,X185,X186,X184] : (sP8(X184,X185) | sP8(X184,'F_0') | sP6(X185,X186,X187)) ) | $spl144),
  inference(resolution,[],[f19503,f307])).
fof(f307,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f195])).
fof(f23934,plain,(
  ( ! [X138] : (sP8(X138,happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13))))) | sP8(X138,'F_0')) ) | $spl276),
  inference(cnf_transformation,[],[f23934_D])).
fof(f23934_D,plain,(
  ( ! [X138] : (sP8(X138,happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13))))) | sP8(X138,'F_0')) ) <=> ~$spl276),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl276])])).
fof(f23920,plain,(
  ( ! [X138,X139] : (sP8(X138,happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13))))) | sP8(X138,'F_0') | ~p(s(bool,X139))) ) | $spl144),
  inference(resolution,[],[f20934,f12547])).
fof(f12547,plain,(
  ( ! [X68] : (~sP6(happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),X68,'F_0') | ~p(s(bool,X68))) )),
  inference(subsumption_resolution,[],[f12453,f397])).
fof(f12453,plain,(
  ( ! [X68] : (s(bool,'F_0') = s(bool,X68) | ~p(s(bool,X68)) | ~sP6(happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),X68,'F_0')) )),
  inference(resolution,[],[f2621,f1560])).
fof(f1560,plain,(
  ( ! [X2,X3] : (sP7(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP6(X3,X2,'F_0')) )),
  inference(subsumption_resolution,[],[f1559,f402])).
fof(f1559,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | sP7(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP6(X3,X2,'F_0')) )),
  inference(subsumption_resolution,[],[f1556,f1546])).
fof(f1546,plain,(
  ( ! [X4,X5] : (sP4(X4,'F_0',X5) | s(bool,'T_0') != s(bool,X4)) )),
  inference(subsumption_resolution,[],[f1541,f313])).
fof(f1541,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | p(s(bool,X5)) | sP4(X4,'F_0',X5)) )),
  inference(resolution,[],[f1458,f297])).
fof(f297,plain,(
  ( ! [X2,X0,X1] : (sP4(X2,X1,X0) | ~sP7(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | (sP4(X2,X1,X0) & sP5(X1,X2,X0) & sP6(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP4(X2,X1,X0) | ~sP5(X1,X2,X0) | ~sP6(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f191])).
fof(f191,plain,(
  ! [X2,X1,X0] : ((~sP7(X2,X1,X0) | (sP4(X0,X1,X2) & sP5(X1,X0,X2) & sP6(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP4(X0,X1,X2) | ~sP5(X1,X0,X2) | ~sP6(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP7(X2,X1,X0)))),
  inference(flattening,[],[f190])).
fof(f190,plain,(
  ! [X2,X1,X0] : ((~sP7(X2,X1,X0) | (sP4(X0,X1,X2) & sP5(X1,X0,X2) & sP6(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP4(X0,X1,X2) | ~sP5(X1,X0,X2) | ~sP6(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP7(X2,X1,X0)))),
  inference(nnf_transformation,[],[f162])).
fof(f1458,plain,(
  ( ! [X6,X5] : (sP7(X6,'F_0',X5) | s(bool,'T_0') != s(bool,X5) | p(s(bool,X6))) )),
  inference(superposition,[],[f1451,f316])).
fof(f316,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f201])).
fof(f1556,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | sP7(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP6(X3,X2,'F_0') | ~sP4(X2,'F_0',X3)) )),
  inference(resolution,[],[f1547,f301])).
fof(f301,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~p(s(bool,X2)) | ~sP6(X0,X2,X1) | ~sP5(X1,X2,X0) | ~sP4(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f192])).
fof(f1547,plain,(
  ( ! [X6,X7] : (sP5('F_0',X6,X7) | s(bool,'T_0') != s(bool,X6)) )),
  inference(subsumption_resolution,[],[f1542,f309])).
fof(f1542,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X6) | p(s(bool,X7)) | sP5('F_0',X6,X7)) )),
  inference(resolution,[],[f1458,f298])).
fof(f298,plain,(
  ( ! [X2,X0,X1] : (sP5(X1,X2,X0) | ~sP7(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f192])).
fof(f2621,plain,(
  ( ! [X112,X113] : (~sP7(happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),X113,X112) | s(bool,X112) = s(bool,X113)) )),
  inference(resolution,[],[f1246,f318])).
fof(f23919,plain,(
  ( ! [X136,X135,X137] : (sP8(X135,happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11)))) | sP8(X135,'F_0') | p(s(bool,X136)) | ~p(s(bool,X137))) ) | $spl144),
  inference(resolution,[],[f20934,f1289])).
fof(f1289,plain,(
  ( ! [X92,X93] : (~sP6(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),X92,X93) | p(s(bool,X93)) | ~p(s(bool,X92))) )),
  inference(resolution,[],[f538,f304])).
fof(f23918,plain,(
  ( ! [X134,X133] : (sP8(X133,happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11)))) | sP8(X133,'F_0') | ~p(s(bool,X134))) ) | $spl144),
  inference(resolution,[],[f20934,f9596])).
fof(f9596,plain,(
  ( ! [X55] : (~sP6(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),X55,'F_0') | ~p(s(bool,X55))) )),
  inference(subsumption_resolution,[],[f9520,f397])).
fof(f9520,plain,(
  ( ! [X55] : (s(bool,'F_0') = s(bool,X55) | ~p(s(bool,X55)) | ~sP6(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),X55,'F_0')) )),
  inference(resolution,[],[f1299,f1560])).
fof(f1299,plain,(
  ( ! [X112,X113] : (~sP7(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),X113,X112) | s(bool,X112) = s(bool,X113)) )),
  inference(resolution,[],[f538,f318])).
fof(f23917,plain,(
  ( ! [X132,X130,X131] : (sP8(X130,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | sP8(X130,'F_0') | p(s(bool,X131)) | ~p(s(bool,X132))) ) | $spl144),
  inference(resolution,[],[f20934,f472])).
fof(f472,plain,(
  ( ! [X90,X89] : (~sP6('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X89,X90) | p(s(bool,X90)) | ~p(s(bool,X89))) )),
  inference(resolution,[],[f250,f304])).
fof(f19802,plain,(
  ( ! [X90,X91] : (~p(s(bool,X90)) | s(bool,X90) = s(bool,X91)) ) | $spl176),
  inference(cnf_transformation,[],[f19802_D])).
fof(f19802_D,plain,(
  ( ! [X90,X91] : (~p(s(bool,X90)) | s(bool,X90) = s(bool,X91)) ) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f23916,plain,(
  ( ! [X127,X128,X129] : (sP8(X127,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | sP8(X127,'F_0') | ~p(s(bool,X128)) | s(bool,X128) = s(bool,X129)) ) | $spl144),
  inference(resolution,[],[f20934,f3259])).
fof(f3259,plain,(
  ( ! [X6,X7] : (~sP6('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X6,X7) | ~p(s(bool,X6)) | s(bool,X6) = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f3258,f477])).
fof(f477,plain,(
  ( ! [X99,X100] : (sP5(X99,X100,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) )),
  inference(resolution,[],[f250,f309])).
fof(f3258,plain,(
  ( ! [X6,X7] : (s(bool,X6) = s(bool,X7) | ~p(s(bool,X6)) | ~sP6('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X6,X7) | ~sP5(X7,X6,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) )),
  inference(subsumption_resolution,[],[f3248,f315])).
fof(f3248,plain,(
  ( ! [X6,X7] : (s(bool,X6) = s(bool,X7) | ~p(s(bool,X6)) | ~sP6('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X6,X7) | ~sP5(X7,X6,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | ~sP4(X6,X7,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) )),
  inference(resolution,[],[f482,f301])).
fof(f19801,plain,(
  ( ! [X88,X89] : (p(s(bool,X88)) | s(bool,X88) = s(bool,X89)) ) | $spl174),
  inference(cnf_transformation,[],[f19801_D])).
fof(f19801_D,plain,(
  ( ! [X88,X89] : (p(s(bool,X88)) | s(bool,X88) = s(bool,X89)) ) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f23915,plain,(
  ( ! [X125,X126,X124] : (sP8(X124,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | sP8(X124,'F_0') | p(s(bool,X125)) | s(bool,X125) = s(bool,X126)) ) | $spl144),
  inference(resolution,[],[f20934,f3261])).
fof(f3261,plain,(
  ( ! [X8,X9] : (~sP6('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X8,X9) | p(s(bool,X9)) | s(bool,X8) = s(bool,X9)) )),
  inference(subsumption_resolution,[],[f3260,f478])).
fof(f478,plain,(
  ( ! [X101,X102] : (sP4(X101,X102,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) )),
  inference(resolution,[],[f250,f313])).
fof(f3260,plain,(
  ( ! [X8,X9] : (s(bool,X8) = s(bool,X9) | p(s(bool,X9)) | ~sP6('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X8,X9) | ~sP4(X8,X9,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) )),
  inference(subsumption_resolution,[],[f3249,f311])).
fof(f3249,plain,(
  ( ! [X8,X9] : (s(bool,X8) = s(bool,X9) | p(s(bool,X9)) | ~sP6('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X8,X9) | ~sP5(X9,X8,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | ~sP4(X8,X9,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) )),
  inference(resolution,[],[f482,f302])).
fof(f302,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X1)) | ~sP6(X0,X2,X1) | ~sP5(X1,X2,X0) | ~sP4(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f192])).
fof(f23914,plain,(
  ( ! [X123,X122] : (sP8(X122,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | sP8(X122,'F_0') | ~p(s(bool,X123))) ) | $spl144),
  inference(resolution,[],[f20934,f3271])).
fof(f3271,plain,(
  ( ! [X15] : (~sP6('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X15,'F_0') | ~p(s(bool,X15))) )),
  inference(subsumption_resolution,[],[f3254,f397])).
fof(f3254,plain,(
  ( ! [X15] : (s(bool,'F_0') = s(bool,X15) | ~p(s(bool,X15)) | ~sP6('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X15,'F_0')) )),
  inference(resolution,[],[f482,f1560])).
fof(f22815,plain,(
  ( ! [X10,X8,X11] : (sP7(X8,X10,X11) | ~sP4(X11,X10,X8) | ~sP5(X10,X11,X8) | p(s(bool,X8))) ) | $spl254),
  inference(cnf_transformation,[],[f22815_D])).
fof(f22815_D,plain,(
  ( ! [X10,X8,X11] : (sP7(X8,X10,X11) | ~sP4(X11,X10,X8) | ~sP5(X10,X11,X8) | p(s(bool,X8))) ) <=> ~$spl254),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl254])])).
fof(f23933,plain,(
  ( ! [X10,X8,X11,X9] : (sP8(X8,'F_0') | sP7(X9,X10,X11) | p(s(bool,X9)) | ~sP5(X10,X11,X9) | ~sP4(X11,X10,X9)) ) | $spl144),
  inference(subsumption_resolution,[],[f23883,f19501])).
fof(f19501,plain,(
  ( ! [X169,X170] : (~sP8(X169,X170) | p(s(bool,X170)) | sP8(X169,'F_0')) ) | $spl144),
  inference(resolution,[],[f19403,f326])).
fof(f23883,plain,(
  ( ! [X10,X8,X11,X9] : (sP8(X8,X9) | sP8(X8,'F_0') | sP7(X9,X10,X11) | p(s(bool,X9)) | ~sP5(X10,X11,X9) | ~sP4(X11,X10,X9)) ) | $spl144),
  inference(resolution,[],[f20934,f303])).
fof(f303,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0)) | ~sP6(X0,X2,X1) | ~sP5(X1,X2,X0) | ~sP4(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f192])).
fof(f22715,plain,(
  ( ! [X43,X44,X42] : (sP6(X42,X43,X44)) ) | $spl248),
  inference(cnf_transformation,[],[f22715_D])).
fof(f22715_D,plain,(
  ( ! [X43,X44,X42] : (sP6(X42,X43,X44)) ) <=> ~$spl248),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl248])])).
fof(f23925,plain,(
  ( ! [X70,X71,X69] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),'F_0') | sP6(X69,X70,X71)) ) | $spl144),
  inference(subsumption_resolution,[],[f23863,f307])).
fof(f23863,plain,(
  ( ! [X70,X71,X69] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),'F_0') | sP6(X69,X70,X71) | ~p(s(bool,X69))) ) | $spl144),
  inference(resolution,[],[f20934,f2622])).
fof(f23924,plain,(
  ( ! [X68,X66,X67] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),'F_0') | sP6(X66,X67,X68)) ) | $spl144),
  inference(subsumption_resolution,[],[f23862,f307])).
fof(f23862,plain,(
  ( ! [X68,X66,X67] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),'F_0') | sP6(X66,X67,X68) | ~p(s(bool,X66))) ) | $spl144),
  inference(resolution,[],[f20934,f1300])).
fof(f23923,plain,(
  ( ! [X61,X62,X60] : (sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | sP6(X60,X61,X62)) ) | $spl144),
  inference(subsumption_resolution,[],[f23860,f307])).
fof(f23860,plain,(
  ( ! [X61,X62,X60] : (sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | sP6(X60,X61,X62) | ~p(s(bool,X60))) ) | $spl144),
  inference(resolution,[],[f20934,f483])).
fof(f23816,plain,(
  ( ! [X39,X37,X38] : (~p(s(bool,X37)) | ~p(s(bool,X38)) | sP8(happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X39)))),'T_0')) ) | $spl144),
  inference(resolution,[],[f3031,f21134])).
fof(f21134,plain,(
  ( ! [X134,X136,X135] : (sP6(X135,X136,X134) | sP8(X134,'T_0')) ) | $spl144),
  inference(resolution,[],[f21077,f305])).
fof(f23814,plain,(
  ( ! [X33,X31,X32] : (~p(s(bool,X31)) | ~p(s(bool,X32)) | sP8('T_0',happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X33)))))) ) | $spl144),
  inference(resolution,[],[f3031,f20537])).
fof(f20537,plain,(
  ( ! [X134,X136,X135] : (sP6(X135,X136,X134) | sP8('T_0',X134)) ) | $spl144),
  inference(resolution,[],[f20335,f305])).
fof(f23809,plain,(
  ( ! [X19,X17,X20,X18] : (~p(s(bool,X17)) | ~p(s(bool,X18)) | sP8(happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X19)))),X20) | s(bool,'F_0') = s(bool,X20)) )),
  inference(resolution,[],[f3031,f17133])).
fof(f17133,plain,(
  ( ! [X177,X178,X176,X179] : (sP8(X176,X177) | sP6(X178,X179,X176) | s(bool,'F_0') = s(bool,X177)) )),
  inference(resolution,[],[f16182,f305])).
fof(f23808,plain,(
  ( ! [X14,X15,X13,X16] : (~p(s(bool,X13)) | ~p(s(bool,X14)) | sP8(X15,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X16))))) | s(bool,'F_0') = s(bool,X15)) )),
  inference(resolution,[],[f3031,f16896])).
fof(f16896,plain,(
  ( ! [X177,X178,X176,X179] : (sP8(X176,X177) | sP6(X178,X179,X177) | s(bool,'F_0') = s(bool,X176)) )),
  inference(resolution,[],[f16178,f305])).
fof(f23472,plain,(
  ( ! [X2,X1] : (sP6(X1,'F_0',X2)) ) | $spl274),
  inference(cnf_transformation,[],[f23472_D])).
fof(f23472_D,plain,(
  ( ! [X2,X1] : (sP6(X1,'F_0',X2)) ) <=> ~$spl274),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl274])])).
fof(f23458,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,'F_0') | sP6(X1,'F_0',X2)) ) | $spl144),
  inference(factoring,[],[f20933])).
fof(f20933,plain,(
  ( ! [X182,X180,X183,X181] : (sP8(X180,X181) | sP8(X180,'F_0') | sP6(X182,X181,X183)) ) | $spl144),
  inference(resolution,[],[f19503,f306])).
fof(f306,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f195])).
fof(f23450,plain,(
  ( ! [X118,X116,X117] : (sP8(X116,happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11)))) | sP8(X116,'F_0') | p(s(bool,X117)) | ~p(s(bool,X118))) ) | $spl144),
  inference(resolution,[],[f20933,f1290])).
fof(f1290,plain,(
  ( ! [X94,X95] : (~sP6(X94,happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),X95) | p(s(bool,X95)) | ~p(s(bool,X94))) )),
  inference(resolution,[],[f538,f304])).
fof(f23449,plain,(
  ( ! [X114,X115,X113] : (sP8(X113,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | sP8(X113,'F_0') | p(s(bool,X114)) | ~p(s(bool,X115))) ) | $spl144),
  inference(resolution,[],[f20933,f473])).
fof(f473,plain,(
  ( ! [X92,X91] : (~sP6(X91,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X92) | p(s(bool,X92)) | ~p(s(bool,X91))) )),
  inference(resolution,[],[f250,f304])).
fof(f19758,plain,(
  ( ! [X81,X82] : (sP7(X81,X82,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | ~sP5(X82,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X81)) ) | $spl172),
  inference(cnf_transformation,[],[f19758_D])).
fof(f19758_D,plain,(
  ( ! [X81,X82] : (sP7(X81,X82,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | ~sP5(X82,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X81)) ) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f23448,plain,(
  ( ! [X111,X112,X110] : (sP8(X110,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | sP8(X110,'F_0') | sP7(X111,X112,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | ~sP5(X112,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X111)) ) | $spl144),
  inference(resolution,[],[f20933,f6665])).
fof(f6665,plain,(
  ( ! [X88,X87] : (~sP6(X87,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X88) | sP7(X87,X88,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | ~sP5(X88,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X87)) )),
  inference(subsumption_resolution,[],[f471,f480])).
fof(f480,plain,(
  ( ! [X105,X106] : (sP4('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X105,X106)) )),
  inference(resolution,[],[f250,f315])).
fof(f471,plain,(
  ( ! [X88,X87] : (sP7(X87,X88,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | ~sP6(X87,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X88) | ~sP5(X88,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X87) | ~sP4('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X88,X87)) )),
  inference(resolution,[],[f250,f301])).
fof(f23471,plain,(
  ( ! [X88,X85,X86] : (sP8(X85,X86) | s(bool,'T_0') != s(bool,X88) | s(bool,X86) != s(bool,X88) | sP8(X85,'F_0')) ) | $spl272),
  inference(cnf_transformation,[],[f23471_D])).
fof(f23471_D,plain,(
  ( ! [X88,X85,X86] : (sP8(X85,X86) | s(bool,'T_0') != s(bool,X88) | s(bool,X86) != s(bool,X88) | sP8(X85,'F_0')) ) <=> ~$spl272),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl272])])).
fof(f23441,plain,(
  ( ! [X88,X87,X85,X86] : (sP8(X85,X86) | sP8(X85,'F_0') | ~p(s(bool,X87)) | s(bool,X86) != s(bool,X88) | s(bool,'T_0') != s(bool,X88)) ) | ($spl24 | $spl144)),
  inference(resolution,[],[f20933,f6588])).
fof(f6588,plain,(
  ( ! [X12,X10,X11] : (~sP6(X11,X12,'F_0') | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,'T_0') != s(bool,X10)) ) | $spl24),
  inference(subsumption_resolution,[],[f6573,f1495])).
fof(f1495,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X3) | s(bool,X2) != s(bool,X3) | p(s(bool,X2))) ) | $spl24),
  inference(superposition,[],[f1480,f400])).
fof(f400,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f241])).
fof(f6573,plain,(
  ( ! [X12,X10,X11] : (s(bool,'T_0') != s(bool,X10) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | ~p(s(bool,X12)) | ~sP6(X11,X12,'F_0')) ) | $spl24),
  inference(resolution,[],[f1500,f1560])).
fof(f1500,plain,(
  ( ! [X17,X15,X16] : (~sP7(X17,'F_0',X15) | s(bool,'T_0') != s(bool,X16) | ~p(s(bool,X17)) | s(bool,X15) != s(bool,X16)) ) | $spl24),
  inference(superposition,[],[f1480,f318])).
fof(f23469,plain,(
  ( ! [X37,X36] : (sP8(X36,X37) | s(bool,'T_0') != s(bool,X37) | sP8(X36,'F_0')) ) | $spl268),
  inference(cnf_transformation,[],[f23469_D])).
fof(f23469_D,plain,(
  ( ! [X37,X36] : (sP8(X36,X37) | s(bool,'T_0') != s(bool,X37) | sP8(X36,'F_0')) ) <=> ~$spl268),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl268])])).
fof(f23439,plain,(
  ( ! [X80,X78,X79] : (sP8(X78,X79) | sP8(X78,'F_0') | s(bool,'T_0') != s(bool,X79) | ~p(s(bool,X80))) ) | $spl144),
  inference(resolution,[],[f20933,f2088])).
fof(f2088,plain,(
  ( ! [X14,X15] : (~sP6(X14,X15,'F_0') | s(bool,'T_0') != s(bool,X15) | ~p(s(bool,X14))) )),
  inference(subsumption_resolution,[],[f2087,f313])).
fof(f2087,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | ~sP6(X14,X15,'F_0') | ~sP4(X15,'F_0',X14)) )),
  inference(subsumption_resolution,[],[f2086,f368])).
fof(f368,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f221])).
fof(f221,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f220])).
fof(f220,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f140])).
fof(f140,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f87])).
fof(f87,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f86])).
fof(f86,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X13] : ((p(s(bool,X13)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X13)))),
  file('thm.quantHeuristics.GUESSES_WEAKEN_THM|split|3','thm.bool.IMP_CLAUSES|split|4')).
fof(f2086,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | p(s(bool,'F_0')) | ~sP6(X14,X15,'F_0') | ~sP4(X15,'F_0',X14)) )),
  inference(subsumption_resolution,[],[f2079,f311])).
fof(f2079,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | p(s(bool,'F_0')) | ~sP6(X14,X15,'F_0') | ~sP5('F_0',X15,X14) | ~sP4(X15,'F_0',X14)) )),
  inference(resolution,[],[f1460,f302])).
fof(f1460,plain,(
  ( ! [X10,X9] : (~sP7(X10,'F_0',X9) | ~p(s(bool,X10)) | s(bool,'T_0') != s(bool,X9)) )),
  inference(superposition,[],[f1451,f318])).
fof(f23470,plain,(
  ( ! [X64] : (s(bool,'T_0') = s(bool,X64)) ) | $spl270),
  inference(cnf_transformation,[],[f23470_D])).
fof(f23470_D,plain,(
  ( ! [X64] : (s(bool,'T_0') = s(bool,X64)) ) <=> ~$spl270),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl270])])).
fof(f22334,plain,(
  ( ! [X24,X23] : (sP8(X23,'F_0') | s(bool,'F_0') = s(bool,X24) | sP8(X23,X24)) ) | $spl240),
  inference(cnf_transformation,[],[f22334_D])).
fof(f22334_D,plain,(
  ( ! [X24,X23] : (sP8(X23,'F_0') | s(bool,'F_0') = s(bool,X24) | sP8(X23,X24)) ) <=> ~$spl240),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl240])])).
fof(f23434,plain,(
  ( ! [X64,X62,X65,X63] : (sP8(X62,X63) | sP8(X62,'F_0') | s(bool,'T_0') = s(bool,X64) | s(bool,'F_0') = s(bool,X63) | ~p(s(bool,X65))) ) | $spl144),
  inference(resolution,[],[f20933,f14390])).
fof(f14390,plain,(
  ( ! [X4,X5,X3] : (~sP6(X3,X5,X4) | s(bool,'T_0') = s(bool,X4) | s(bool,'F_0') = s(bool,X5) | ~p(s(bool,X3))) )),
  inference(subsumption_resolution,[],[f14389,f313])).
fof(f14389,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | s(bool,'T_0') = s(bool,X4) | s(bool,'F_0') = s(bool,X5) | ~sP6(X3,X5,X4) | ~sP4(X5,X4,X3)) )),
  inference(subsumption_resolution,[],[f14388,f309])).
fof(f14388,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | s(bool,'T_0') = s(bool,X4) | s(bool,'F_0') = s(bool,X5) | ~sP6(X3,X5,X4) | ~sP5(X4,X5,X3) | ~sP4(X5,X4,X3)) )),
  inference(subsumption_resolution,[],[f14308,f402])).
fof(f14308,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | s(bool,'T_0') = s(bool,X4) | s(bool,'F_0') = s(bool,X5) | p(s(bool,X4)) | ~sP6(X3,X5,X4) | ~sP5(X4,X5,X3) | ~sP4(X5,X4,X3)) )),
  inference(resolution,[],[f13954,f302])).
fof(f13954,plain,(
  ( ! [X50,X48,X49] : (~sP7(X50,X49,X48) | ~p(s(bool,X50)) | s(bool,'T_0') = s(bool,X49) | s(bool,'F_0') = s(bool,X48)) )),
  inference(trivial_inequality_removal,[],[f13933])).
fof(f13933,plain,(
  ( ! [X50,X48,X49] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'T_0') = s(bool,X49) | ~p(s(bool,X50)) | ~sP7(X50,X49,X48) | s(bool,'F_0') = s(bool,X48)) )),
  inference(superposition,[],[f13735,f424])).
fof(f13735,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') != s(bool,X10) | s(bool,'T_0') = s(bool,X9) | ~p(s(bool,X11)) | ~sP7(X11,X9,X10)) )),
  inference(forward_demodulation,[],[f13697,f504])).
fof(f13697,plain,(
  ( ! [X10,X11,X9] : (s(bool,X10) != s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | s(bool,'T_0') = s(bool,X9) | ~p(s(bool,X11)) | ~sP7(X11,X9,X10)) )),
  inference(superposition,[],[f13015,f318])).
fof(f13015,plain,(
  ( ! [X52] : (s(bool,X52) != s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | s(bool,'T_0') = s(bool,X52)) )),
  inference(subsumption_resolution,[],[f13014,f412])).
fof(f412,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f110])).
fof(f110,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.quantHeuristics.GUESSES_WEAKEN_THM|split|3',aHL_FALSITY)).
fof(f13014,plain,(
  ( ! [X52] : (p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,X52) | s(bool,X52) != s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) )),
  inference(forward_demodulation,[],[f12991,f613])).
fof(f613,plain,(
  s(bool,'F_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))),
  inference(resolution,[],[f251,f398])).
fof(f12991,plain,(
  ( ! [X52] : (s(bool,'T_0') = s(bool,X52) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) | s(bool,X52) != s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) )),
  inference(resolution,[],[f1565,f592])).
fof(f592,plain,(
  ( ! [X114,X113] : (~sP7('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X114,X113) | s(bool,X113) != s(bool,X114)) )),
  inference(resolution,[],[f251,f319])).
fof(f319,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | s(bool,X0) != s(bool,X1) | ~sP7(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f201])).
fof(f1565,plain,(
  ( ! [X0,X1] : (sP7(X1,X0,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | s(bool,'T_0') = s(bool,X0) | p(s(bool,X1))) )),
  inference(superposition,[],[f504,f316])).
fof(f21041,plain,(
  ( ! [X0,X3,X1] : (sP8(X0,X1) | s(bool,X1) != s(bool,X3) | ~p(s(bool,X3)) | sP8(X0,'F_0')) ) | $spl204),
  inference(cnf_transformation,[],[f21041_D])).
fof(f21041_D,plain,(
  ( ! [X0,X3,X1] : (sP8(X0,X1) | s(bool,X1) != s(bool,X3) | ~p(s(bool,X3)) | sP8(X0,'F_0')) ) <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f23429,plain,(
  ( ! [X43,X41,X44,X42,X40] : (sP8(X40,X41) | sP8(X40,'F_0') | ~p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X41) != s(bool,X42) | p(s(bool,X44))) ) | ($spl6 | $spl144)),
  inference(resolution,[],[f20933,f8286])).
fof(f8286,plain,(
  ( ! [X6,X4,X7,X5] : (~sP6(X6,X7,X4) | ~p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | p(s(bool,X4))) ) | $spl6),
  inference(subsumption_resolution,[],[f8285,f313])).
fof(f8285,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | ~sP6(X6,X7,X4) | ~sP4(X7,X4,X6)) ) | $spl6),
  inference(subsumption_resolution,[],[f8284,f311])).
fof(f8284,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | ~sP6(X6,X7,X4) | ~sP5(X4,X7,X6) | ~sP4(X7,X4,X6)) ) | $spl6),
  inference(duplicate_literal_removal,[],[f8246])).
fof(f8246,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | p(s(bool,X4)) | ~sP6(X6,X7,X4) | ~sP5(X4,X7,X6) | ~sP4(X7,X4,X6)) ) | $spl6),
  inference(resolution,[],[f1181,f302])).
fof(f1181,plain,(
  ( ! [X14,X12,X15,X13] : (~sP7(X15,X12,X13) | p(s(bool,X12)) | ~p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X13) != s(bool,X14)) ) | $spl6),
  inference(superposition,[],[f528,f318])).
fof(f836,plain,(
  ( ! [X29] : (s(bool,'F_0') != s(bool,X29)) ) | $spl18),
  inference(cnf_transformation,[],[f836_D])).
fof(f836_D,plain,(
  ( ! [X29] : (s(bool,'F_0') != s(bool,X29)) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f23428,plain,(
  ( ! [X39,X37,X38,X36] : (sP8(X36,X37) | sP8(X36,'F_0') | ~p(s(bool,X38)) | s(bool,'F_0') != s(bool,X39) | s(bool,'T_0') != s(bool,X37)) ) | ($spl24 | $spl144)),
  inference(resolution,[],[f20933,f7489])).
fof(f7489,plain,(
  ( ! [X4,X5,X3] : (~sP6(X4,X3,X5) | ~p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5) | s(bool,'T_0') != s(bool,X3)) ) | $spl24),
  inference(subsumption_resolution,[],[f7488,f397])).
fof(f7488,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5) | p(s(bool,X5)) | ~sP6(X4,X3,X5)) ) | $spl24),
  inference(subsumption_resolution,[],[f7487,f6637])).
fof(f6637,plain,(
  ( ! [X6,X8,X7] : (sP4(X6,X7,X8) | s(bool,'F_0') != s(bool,X7) | s(bool,'T_0') != s(bool,X6)) ) | $spl24),
  inference(subsumption_resolution,[],[f6601,f313])).
fof(f6601,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X6) | s(bool,'F_0') != s(bool,X7) | p(s(bool,X8)) | sP4(X6,X7,X8)) ) | $spl24),
  inference(resolution,[],[f1511,f297])).
fof(f1511,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X1,X0) | s(bool,'T_0') != s(bool,X0) | s(bool,'F_0') != s(bool,X1) | p(s(bool,X2))) ) | $spl24),
  inference(superposition,[],[f1480,f316])).
fof(f7487,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5) | p(s(bool,X5)) | ~sP6(X4,X3,X5) | ~sP4(X3,X5,X4)) ) | $spl24),
  inference(subsumption_resolution,[],[f7446,f6638])).
fof(f6638,plain,(
  ( ! [X10,X11,X9] : (sP5(X10,X9,X11) | s(bool,'F_0') != s(bool,X10) | s(bool,'T_0') != s(bool,X9)) ) | $spl24),
  inference(subsumption_resolution,[],[f6602,f309])).
fof(f6602,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') != s(bool,X9) | s(bool,'F_0') != s(bool,X10) | p(s(bool,X11)) | sP5(X10,X9,X11)) ) | $spl24),
  inference(resolution,[],[f1511,f298])).
fof(f7446,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5) | p(s(bool,X5)) | ~sP6(X4,X3,X5) | ~sP5(X5,X3,X4) | ~sP4(X3,X5,X4)) ) | $spl24),
  inference(resolution,[],[f1513,f302])).
fof(f1513,plain,(
  ( ! [X6,X8,X7] : (~sP7(X8,X7,X6) | s(bool,'T_0') != s(bool,X6) | ~p(s(bool,X8)) | s(bool,'F_0') != s(bool,X7)) ) | $spl24),
  inference(superposition,[],[f1480,f318])).
fof(f23462,plain,(
  ( ! [X70,X71,X69] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),'F_0') | sP6(X69,X70,X71)) ) | $spl144),
  inference(subsumption_resolution,[],[f23401,f306])).
fof(f23401,plain,(
  ( ! [X70,X71,X69] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),'F_0') | sP6(X69,X70,X71) | ~p(s(bool,X70))) ) | $spl144),
  inference(resolution,[],[f20933,f2622])).
fof(f23461,plain,(
  ( ! [X68,X66,X67] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),'F_0') | sP6(X66,X67,X68)) ) | $spl144),
  inference(subsumption_resolution,[],[f23400,f306])).
fof(f23400,plain,(
  ( ! [X68,X66,X67] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),'F_0') | sP6(X66,X67,X68) | ~p(s(bool,X67))) ) | $spl144),
  inference(resolution,[],[f20933,f1300])).
fof(f23460,plain,(
  ( ! [X61,X62,X60] : (sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | sP6(X60,X61,X62)) ) | $spl144),
  inference(subsumption_resolution,[],[f23398,f306])).
fof(f23398,plain,(
  ( ! [X61,X62,X60] : (sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | sP6(X60,X61,X62) | ~p(s(bool,X61))) ) | $spl144),
  inference(resolution,[],[f20933,f483])).
fof(f23361,plain,(
  ( ! [X2,X1] : (sP2('F_0',X1,X2)) ) | $spl266),
  inference(cnf_transformation,[],[f23361_D])).
fof(f23361_D,plain,(
  ( ! [X2,X1] : (sP2('F_0',X1,X2)) ) <=> ~$spl266),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl266])])).
fof(f23351,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,'F_0') | sP2('F_0',X1,X2)) ) | $spl144),
  inference(factoring,[],[f20921])).
fof(f20921,plain,(
  ( ! [X134,X132,X135,X133] : (sP8(X132,X133) | sP8(X132,'F_0') | sP2(X133,X134,X135)) ) | $spl144),
  inference(resolution,[],[f19503,f290])).
fof(f290,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f187])).
fof(f23350,plain,(
  ( ! [X50,X48,X49] : (sP8(X48,happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11)))) | sP8(X48,'F_0') | p(s(bool,X49)) | p(s(bool,X50))) ) | $spl144),
  inference(resolution,[],[f20921,f1276])).
fof(f1276,plain,(
  ( ! [X66,X67] : (~sP2(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),X66,X67) | p(s(bool,X67)) | p(s(bool,X66))) )),
  inference(resolution,[],[f538,f287])).
fof(f23349,plain,(
  ( ! [X47,X45,X46] : (sP8(X45,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | sP8(X45,'F_0') | p(s(bool,X46)) | p(s(bool,X47))) ) | $spl144),
  inference(resolution,[],[f20921,f459])).
fof(f459,plain,(
  ( ! [X64,X63] : (~sP2('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X63,X64) | p(s(bool,X64)) | p(s(bool,X63))) )),
  inference(resolution,[],[f250,f287])).
fof(f19715,plain,(
  ( ! [X26,X27] : (sP3('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X26,X27)) ) | $spl166),
  inference(cnf_transformation,[],[f19715_D])).
fof(f19715_D,plain,(
  ( ! [X26,X27] : (sP3('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X26,X27)) ) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f23348,plain,(
  ( ! [X43,X44,X42] : (sP8(X42,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | sP8(X42,'F_0') | sP3('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X43,X44)) ) | $spl144),
  inference(resolution,[],[f20921,f508])).
fof(f508,plain,(
  ( ! [X130,X129] : (~sP2('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X129,X130) | sP3('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X129,X130)) )),
  inference(resolution,[],[f250,f428])).
fof(f428,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~sP2(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(duplicate_literal_removal,[],[f283])).
fof(f283,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~sP2(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f184])).
fof(f22628,plain,(
  ( ! [X43,X44,X42] : (sP2(X42,X43,X44)) ) | $spl246),
  inference(cnf_transformation,[],[f22628_D])).
fof(f22628_D,plain,(
  ( ! [X43,X44,X42] : (sP2(X42,X43,X44)) ) <=> ~$spl246),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl246])])).
fof(f23355,plain,(
  ( ! [X70,X71,X69] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),'F_0') | sP2(X69,X70,X71)) ) | $spl144),
  inference(subsumption_resolution,[],[f23318,f290])).
fof(f23318,plain,(
  ( ! [X70,X71,X69] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),'F_0') | sP2(X69,X70,X71) | ~p(s(bool,X69))) ) | $spl144),
  inference(resolution,[],[f20921,f2622])).
fof(f23354,plain,(
  ( ! [X68,X66,X67] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),'F_0') | sP2(X66,X67,X68)) ) | $spl144),
  inference(subsumption_resolution,[],[f23317,f290])).
fof(f23317,plain,(
  ( ! [X68,X66,X67] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),'F_0') | sP2(X66,X67,X68) | ~p(s(bool,X66))) ) | $spl144),
  inference(resolution,[],[f20921,f1300])).
fof(f23353,plain,(
  ( ! [X61,X62,X60] : (sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | sP2(X60,X61,X62)) ) | $spl144),
  inference(subsumption_resolution,[],[f23315,f290])).
fof(f23315,plain,(
  ( ! [X61,X62,X60] : (sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | sP2(X60,X61,X62) | ~p(s(bool,X60))) ) | $spl144),
  inference(resolution,[],[f20921,f483])).
fof(f23283,plain,(
  ( ! [X2,X1] : (sP0('F_0',X1,X2)) ) | $spl264),
  inference(cnf_transformation,[],[f23283_D])).
fof(f23283_D,plain,(
  ( ! [X2,X1] : (sP0('F_0',X1,X2)) ) <=> ~$spl264),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl264])])).
fof(f23273,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,'F_0') | sP0('F_0',X1,X2)) ) | $spl144),
  inference(factoring,[],[f20905])).
fof(f20905,plain,(
  ( ! [X70,X68,X71,X69] : (sP8(X68,X69) | sP8(X68,'F_0') | sP0(X69,X70,X71)) ) | $spl144),
  inference(resolution,[],[f19503,f273])).
fof(f273,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f179])).
fof(f23272,plain,(
  ( ! [X50,X48,X49] : (sP8(X48,happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11)))) | sP8(X48,'F_0') | ~p(s(bool,X49)) | p(s(bool,X50))) ) | $spl144),
  inference(resolution,[],[f20905,f1262])).
fof(f1262,plain,(
  ( ! [X39,X38] : (~sP0(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),X38,X39) | ~p(s(bool,X39)) | p(s(bool,X38))) )),
  inference(resolution,[],[f538,f270])).
fof(f23271,plain,(
  ( ! [X47,X45,X46] : (sP8(X45,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | sP8(X45,'F_0') | ~p(s(bool,X46)) | p(s(bool,X47))) ) | $spl144),
  inference(resolution,[],[f20905,f445])).
fof(f445,plain,(
  ( ! [X35,X36] : (~sP0('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X35,X36) | ~p(s(bool,X36)) | p(s(bool,X35))) )),
  inference(resolution,[],[f250,f270])).
fof(f19701,plain,(
  ( ! [X28,X29] : (sP1('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X28,X29)) ) | $spl164),
  inference(cnf_transformation,[],[f19701_D])).
fof(f19701_D,plain,(
  ( ! [X28,X29] : (sP1('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X28,X29)) ) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f23270,plain,(
  ( ! [X43,X44,X42] : (sP8(X42,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | sP8(X42,'F_0') | sP1('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X43,X44)) ) | $spl144),
  inference(resolution,[],[f20905,f507])).
fof(f507,plain,(
  ( ! [X127,X128] : (~sP0('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X127,X128) | sP1('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X127,X128)) )),
  inference(resolution,[],[f250,f427])).
fof(f427,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | ~sP0(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(duplicate_literal_removal,[],[f266])).
fof(f266,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | ~sP0(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f176])).
fof(f176,plain,(
  ! [X0,X1,X2] : ((~sP1(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP0(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP0(X0,X1,X2) | sP1(X0,X1,X2)))),
  inference(rectify,[],[f175])).
fof(f175,plain,(
  ! [X2,X0,X1] : ((~sP1(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP0(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP0(X2,X0,X1) | sP1(X2,X0,X1)))),
  inference(flattening,[],[f174])).
fof(f174,plain,(
  ! [X2,X0,X1] : ((~sP1(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP0(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP0(X2,X0,X1)) | sP1(X2,X0,X1)))),
  inference(nnf_transformation,[],[f154])).
fof(f22467,plain,(
  ( ! [X43,X44,X42] : (sP0(X42,X43,X44)) ) | $spl242),
  inference(cnf_transformation,[],[f22467_D])).
fof(f22467_D,plain,(
  ( ! [X43,X44,X42] : (sP0(X42,X43,X44)) ) <=> ~$spl242),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl242])])).
fof(f23277,plain,(
  ( ! [X70,X71,X69] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),'F_0') | sP0(X69,X70,X71)) ) | $spl144),
  inference(subsumption_resolution,[],[f23240,f273])).
fof(f23240,plain,(
  ( ! [X70,X71,X69] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),'F_0') | sP0(X69,X70,X71) | ~p(s(bool,X69))) ) | $spl144),
  inference(resolution,[],[f20905,f2622])).
fof(f23276,plain,(
  ( ! [X68,X66,X67] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),'F_0') | sP0(X66,X67,X68)) ) | $spl144),
  inference(subsumption_resolution,[],[f23239,f273])).
fof(f23239,plain,(
  ( ! [X68,X66,X67] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),'F_0') | sP0(X66,X67,X68) | ~p(s(bool,X66))) ) | $spl144),
  inference(resolution,[],[f20905,f1300])).
fof(f23275,plain,(
  ( ! [X61,X62,X60] : (sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | sP0(X60,X61,X62)) ) | $spl144),
  inference(subsumption_resolution,[],[f23237,f273])).
fof(f23237,plain,(
  ( ! [X61,X62,X60] : (sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | sP0(X60,X61,X62) | ~p(s(bool,X60))) ) | $spl144),
  inference(resolution,[],[f20905,f483])).
fof(f23108,plain,(
  ( ! [X2,X1] : (sP0(X1,X2,'F_0')) ) | $spl262),
  inference(cnf_transformation,[],[f23108_D])).
fof(f23108_D,plain,(
  ( ! [X2,X1] : (sP0(X1,X2,'F_0')) ) <=> ~$spl262),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl262])])).
fof(f23096,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,'F_0') | sP0(X1,X2,'F_0')) ) | $spl144),
  inference(factoring,[],[f20904])).
fof(f20904,plain,(
  ( ! [X66,X64,X67,X65] : (sP8(X64,X65) | sP8(X64,'F_0') | sP0(X66,X67,X65)) ) | $spl144),
  inference(resolution,[],[f19503,f271])).
fof(f271,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f179])).
fof(f23089,plain,(
  ( ! [X30,X31,X32] : (sP8(X30,happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11)))) | sP8(X30,'F_0') | p(s(bool,X31)) | ~p(s(bool,X32))) ) | $spl144),
  inference(resolution,[],[f20904,f1263])).
fof(f1263,plain,(
  ( ! [X41,X40] : (~sP0(X40,X41,happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11)))) | p(s(bool,X41)) | ~p(s(bool,X40))) )),
  inference(resolution,[],[f538,f270])).
fof(f23088,plain,(
  ( ! [X28,X29,X27] : (sP8(X27,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | sP8(X27,'F_0') | p(s(bool,X28)) | ~p(s(bool,X29))) ) | $spl144),
  inference(resolution,[],[f20904,f446])).
fof(f446,plain,(
  ( ! [X37,X38] : (~sP0(X37,X38,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | p(s(bool,X38)) | ~p(s(bool,X37))) )),
  inference(resolution,[],[f250,f270])).
fof(f19686,plain,(
  ( ! [X19,X18] : (sP1(X18,X19,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | p(s(bool,X19))) ) | $spl162),
  inference(cnf_transformation,[],[f19686_D])).
fof(f19686_D,plain,(
  ( ! [X19,X18] : (sP1(X18,X19,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | p(s(bool,X19))) ) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f23087,plain,(
  ( ! [X26,X24,X25] : (sP8(X24,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | sP8(X24,'F_0') | sP1(X25,X26,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | p(s(bool,X26))) ) | $spl144),
  inference(resolution,[],[f20904,f444])).
fof(f444,plain,(
  ( ! [X33,X34] : (~sP0(X33,X34,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | sP1(X33,X34,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | p(s(bool,X34))) )),
  inference(resolution,[],[f250,f269])).
fof(f269,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | ~sP0(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f176])).
fof(f23100,plain,(
  ( ! [X70,X71,X69] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),'F_0') | sP0(X69,X70,X71)) ) | $spl144),
  inference(subsumption_resolution,[],[f23063,f271])).
fof(f23063,plain,(
  ( ! [X70,X71,X69] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),'F_0') | sP0(X69,X70,X71) | ~p(s(bool,X71))) ) | $spl144),
  inference(resolution,[],[f20904,f2622])).
fof(f23099,plain,(
  ( ! [X68,X66,X67] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),'F_0') | sP0(X66,X67,X68)) ) | $spl144),
  inference(subsumption_resolution,[],[f23062,f271])).
fof(f23062,plain,(
  ( ! [X68,X66,X67] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),'F_0') | sP0(X66,X67,X68) | ~p(s(bool,X68))) ) | $spl144),
  inference(resolution,[],[f20904,f1300])).
fof(f23098,plain,(
  ( ! [X61,X62,X60] : (sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | sP0(X60,X61,X62)) ) | $spl144),
  inference(subsumption_resolution,[],[f23060,f271])).
fof(f23060,plain,(
  ( ! [X61,X62,X60] : (sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | sP0(X60,X61,X62) | ~p(s(bool,X62))) ) | $spl144),
  inference(resolution,[],[f20904,f483])).
fof(f22476,plain,(
  ( ! [X30] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),X30) | sP8(X30,'F_0')) ) | $spl244),
  inference(cnf_transformation,[],[f22476_D])).
fof(f22476_D,plain,(
  ( ! [X30] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),X30) | sP8(X30,'F_0')) ) <=> ~$spl244),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl244])])).
fof(f23029,plain,(
  ( ! [X66,X64,X65] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),X64) | sP8(X64,'F_0') | p(s(bool,X65)) | ~p(s(bool,X66))) ) | $spl144),
  inference(resolution,[],[f20044,f1292])).
fof(f20044,plain,(
  ( ! [X198,X196,X199,X197] : (sP8(X196,X197) | sP8(X197,'F_0') | sP5(X196,X198,X199)) ) | $spl144),
  inference(resolution,[],[f19502,f311])).
fof(f20152,plain,(
  ( ! [X55] : (sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X55) | sP8(X55,'F_0')) ) | $spl184),
  inference(cnf_transformation,[],[f20152_D])).
fof(f20152_D,plain,(
  ( ! [X55] : (sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X55) | sP8(X55,'F_0')) ) <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f23028,plain,(
  ( ! [X61,X62,X63] : (sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X61) | sP8(X61,'F_0') | p(s(bool,X62)) | ~p(s(bool,X63))) ) | $spl144),
  inference(resolution,[],[f20044,f475])).
fof(f23035,plain,(
  ( ! [X52,X53,X51] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),'F_0') | sP5(X51,X52,X53)) ) | $spl144),
  inference(subsumption_resolution,[],[f22987,f311])).
fof(f22987,plain,(
  ( ! [X52,X53,X51] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),'F_0') | sP5(X51,X52,X53) | ~p(s(bool,X51))) ) | $spl144),
  inference(resolution,[],[f20044,f2623])).
fof(f23034,plain,(
  ( ! [X50,X48,X49] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),'F_0') | sP5(X48,X49,X50)) ) | $spl144),
  inference(subsumption_resolution,[],[f22986,f311])).
fof(f22986,plain,(
  ( ! [X50,X48,X49] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),'F_0') | sP5(X48,X49,X50) | ~p(s(bool,X48))) ) | $spl144),
  inference(resolution,[],[f20044,f1301])).
fof(f23033,plain,(
  ( ! [X43,X44,X42] : (sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | sP5(X42,X43,X44)) ) | $spl144),
  inference(subsumption_resolution,[],[f22984,f311])).
fof(f22984,plain,(
  ( ! [X43,X44,X42] : (sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | sP5(X42,X43,X44) | ~p(s(bool,X42))) ) | $spl144),
  inference(resolution,[],[f20044,f484])).
fof(f22889,plain,(
  ( ! [X57,X56,X55] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),X55) | sP8(X55,'F_0') | p(s(bool,X56)) | ~p(s(bool,X57))) ) | $spl144),
  inference(resolution,[],[f20043,f1293])).
fof(f20043,plain,(
  ( ! [X194,X192,X195,X193] : (sP8(X192,X193) | sP8(X193,'F_0') | sP5(X194,X192,X195)) ) | $spl144),
  inference(resolution,[],[f19502,f310])).
fof(f22888,plain,(
  ( ! [X54,X52,X53] : (sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X52) | sP8(X52,'F_0') | p(s(bool,X53)) | ~p(s(bool,X54))) ) | $spl144),
  inference(resolution,[],[f20043,f476])).
fof(f22898,plain,(
  ( ! [X52,X53,X51] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),'F_0') | sP5(X51,X52,X53)) ) | $spl144),
  inference(subsumption_resolution,[],[f22850,f310])).
fof(f22850,plain,(
  ( ! [X52,X53,X51] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),'F_0') | sP5(X51,X52,X53) | ~p(s(bool,X52))) ) | $spl144),
  inference(resolution,[],[f20043,f2623])).
fof(f22897,plain,(
  ( ! [X50,X48,X49] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),'F_0') | sP5(X48,X49,X50)) ) | $spl144),
  inference(subsumption_resolution,[],[f22849,f310])).
fof(f22849,plain,(
  ( ! [X50,X48,X49] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),'F_0') | sP5(X48,X49,X50) | ~p(s(bool,X49))) ) | $spl144),
  inference(resolution,[],[f20043,f1301])).
fof(f22895,plain,(
  ( ! [X43,X44,X42] : (sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | sP5(X42,X43,X44)) ) | $spl144),
  inference(subsumption_resolution,[],[f22847,f310])).
fof(f22847,plain,(
  ( ! [X43,X44,X42] : (sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | sP5(X42,X43,X44) | ~p(s(bool,X43))) ) | $spl144),
  inference(resolution,[],[f20043,f484])).
fof(f22836,plain,(
  $false | $spl202),
  inference(subsumption_resolution,[],[f22832,f412])).
fof(f22832,plain,(
  p(s(bool,'F_0')) | $spl202),
  inference(resolution,[],[f20883,f3043])).
fof(f22835,plain,(
  $false | $spl202),
  inference(subsumption_resolution,[],[f22829,f412])).
fof(f22829,plain,(
  p(s(bool,'F_0')) | $spl202),
  inference(resolution,[],[f20883,f594])).
fof(f22827,plain,(
  $false | ($spl161 | $spl202)),
  inference(resolution,[],[f20883,f19599])).
fof(f22816,plain,(
  ( ! [X135] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),X135) | sP8(X135,'F_0')) ) | $spl256),
  inference(cnf_transformation,[],[f22816_D])).
fof(f22816_D,plain,(
  ( ! [X135] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),X135) | sP8(X135,'F_0')) ) <=> ~$spl256),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl256])])).
fof(f22801,plain,(
  ( ! [X136,X135] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),X135) | sP8(X135,'F_0') | ~p(s(bool,X136))) ) | $spl144),
  inference(resolution,[],[f20041,f12547])).
fof(f20041,plain,(
  ( ! [X187,X185,X186,X184] : (sP8(X184,X185) | sP8(X185,'F_0') | sP6(X184,X186,X187)) ) | $spl144),
  inference(resolution,[],[f19502,f307])).
fof(f22800,plain,(
  ( ! [X134,X132,X133] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),X132) | sP8(X132,'F_0') | p(s(bool,X133)) | ~p(s(bool,X134))) ) | $spl144),
  inference(resolution,[],[f20041,f1289])).
fof(f22799,plain,(
  ( ! [X130,X131] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),X130) | sP8(X130,'F_0') | ~p(s(bool,X131))) ) | $spl144),
  inference(resolution,[],[f20041,f9596])).
fof(f22798,plain,(
  ( ! [X127,X128,X129] : (sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X127) | sP8(X127,'F_0') | p(s(bool,X128)) | ~p(s(bool,X129))) ) | $spl144),
  inference(resolution,[],[f20041,f472])).
fof(f22797,plain,(
  ( ! [X125,X126,X124] : (sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X124) | sP8(X124,'F_0') | ~p(s(bool,X125)) | s(bool,X125) = s(bool,X126)) ) | $spl144),
  inference(resolution,[],[f20041,f3259])).
fof(f22796,plain,(
  ( ! [X123,X121,X122] : (sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X121) | sP8(X121,'F_0') | p(s(bool,X122)) | s(bool,X122) = s(bool,X123)) ) | $spl144),
  inference(resolution,[],[f20041,f3261])).
fof(f22795,plain,(
  ( ! [X120,X119] : (sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X119) | sP8(X119,'F_0') | ~p(s(bool,X120))) ) | $spl144),
  inference(resolution,[],[f20041,f3271])).
fof(f22814,plain,(
  ( ! [X10,X8,X11,X9] : (sP8(X9,'F_0') | sP7(X8,X10,X11) | p(s(bool,X8)) | ~sP5(X10,X11,X8) | ~sP4(X11,X10,X8)) ) | $spl144),
  inference(subsumption_resolution,[],[f22765,f19500])).
fof(f19500,plain,(
  ( ! [X167,X168] : (~sP8(X168,X167) | p(s(bool,X168)) | sP8(X167,'F_0')) ) | $spl144),
  inference(resolution,[],[f19403,f326])).
fof(f22765,plain,(
  ( ! [X10,X8,X11,X9] : (sP8(X8,X9) | sP8(X9,'F_0') | sP7(X8,X10,X11) | p(s(bool,X8)) | ~sP5(X10,X11,X8) | ~sP4(X11,X10,X8)) ) | $spl144),
  inference(resolution,[],[f20041,f303])).
fof(f22806,plain,(
  ( ! [X52,X53,X51] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),'F_0') | sP6(X51,X52,X53)) ) | $spl144),
  inference(subsumption_resolution,[],[f22739,f307])).
fof(f22739,plain,(
  ( ! [X52,X53,X51] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),'F_0') | sP6(X51,X52,X53) | ~p(s(bool,X51))) ) | $spl144),
  inference(resolution,[],[f20041,f2623])).
fof(f22805,plain,(
  ( ! [X50,X48,X49] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),'F_0') | sP6(X48,X49,X50)) ) | $spl144),
  inference(subsumption_resolution,[],[f22738,f307])).
fof(f22738,plain,(
  ( ! [X50,X48,X49] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),'F_0') | sP6(X48,X49,X50) | ~p(s(bool,X48))) ) | $spl144),
  inference(resolution,[],[f20041,f1301])).
fof(f22804,plain,(
  ( ! [X43,X44,X42] : (sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | sP6(X42,X43,X44)) ) | $spl144),
  inference(subsumption_resolution,[],[f22736,f307])).
fof(f22736,plain,(
  ( ! [X43,X44,X42] : (sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | sP6(X42,X43,X44) | ~p(s(bool,X42))) ) | $spl144),
  inference(resolution,[],[f20041,f484])).
fof(f22704,plain,(
  ( ! [X118,X116,X117] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),X116) | sP8(X116,'F_0') | p(s(bool,X117)) | ~p(s(bool,X118))) ) | $spl144),
  inference(resolution,[],[f20040,f1290])).
fof(f20040,plain,(
  ( ! [X182,X180,X183,X181] : (sP8(X180,X181) | sP8(X181,'F_0') | sP6(X182,X180,X183)) ) | $spl144),
  inference(resolution,[],[f19502,f306])).
fof(f22703,plain,(
  ( ! [X114,X115,X113] : (sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X113) | sP8(X113,'F_0') | p(s(bool,X114)) | ~p(s(bool,X115))) ) | $spl144),
  inference(resolution,[],[f20040,f473])).
fof(f22702,plain,(
  ( ! [X111,X112,X110] : (sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X110) | sP8(X110,'F_0') | sP7(X111,X112,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | ~sP5(X112,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X111)) ) | $spl144),
  inference(resolution,[],[f20040,f6665])).
fof(f22725,plain,(
  ( ! [X88,X85,X86] : (sP8(X85,X86) | s(bool,'T_0') != s(bool,X88) | s(bool,X85) != s(bool,X88) | sP8(X86,'F_0')) ) | $spl252),
  inference(cnf_transformation,[],[f22725_D])).
fof(f22725_D,plain,(
  ( ! [X88,X85,X86] : (sP8(X85,X86) | s(bool,'T_0') != s(bool,X88) | s(bool,X85) != s(bool,X88) | sP8(X86,'F_0')) ) <=> ~$spl252),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl252])])).
fof(f22695,plain,(
  ( ! [X88,X87,X85,X86] : (sP8(X85,X86) | sP8(X86,'F_0') | ~p(s(bool,X87)) | s(bool,X85) != s(bool,X88) | s(bool,'T_0') != s(bool,X88)) ) | ($spl24 | $spl144)),
  inference(resolution,[],[f20040,f6588])).
fof(f22724,plain,(
  ( ! [X37,X36] : (sP8(X36,X37) | s(bool,'T_0') != s(bool,X36) | sP8(X37,'F_0')) ) | $spl250),
  inference(cnf_transformation,[],[f22724_D])).
fof(f22724_D,plain,(
  ( ! [X37,X36] : (sP8(X36,X37) | s(bool,'T_0') != s(bool,X36) | sP8(X37,'F_0')) ) <=> ~$spl250),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl250])])).
fof(f22693,plain,(
  ( ! [X80,X78,X79] : (sP8(X78,X79) | sP8(X79,'F_0') | s(bool,'T_0') != s(bool,X78) | ~p(s(bool,X80))) ) | $spl144),
  inference(resolution,[],[f20040,f2088])).
fof(f20138,plain,(
  ( ! [X0,X3,X1] : (sP8(X0,X1) | s(bool,X0) != s(bool,X3) | ~p(s(bool,X3)) | sP8(X1,'F_0')) ) | $spl178),
  inference(cnf_transformation,[],[f20138_D])).
fof(f20138_D,plain,(
  ( ! [X0,X3,X1] : (sP8(X0,X1) | s(bool,X0) != s(bool,X3) | ~p(s(bool,X3)) | sP8(X1,'F_0')) ) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f22683,plain,(
  ( ! [X43,X41,X44,X42,X40] : (sP8(X40,X41) | sP8(X41,'F_0') | ~p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X40) != s(bool,X42) | p(s(bool,X44))) ) | ($spl6 | $spl144)),
  inference(resolution,[],[f20040,f8286])).
fof(f22682,plain,(
  ( ! [X39,X37,X38,X36] : (sP8(X36,X37) | sP8(X37,'F_0') | ~p(s(bool,X38)) | s(bool,'F_0') != s(bool,X39) | s(bool,'T_0') != s(bool,X36)) ) | ($spl24 | $spl144)),
  inference(resolution,[],[f20040,f7489])).
fof(f22717,plain,(
  ( ! [X52,X53,X51] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),'F_0') | sP6(X51,X52,X53)) ) | $spl144),
  inference(subsumption_resolution,[],[f22649,f306])).
fof(f22649,plain,(
  ( ! [X52,X53,X51] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),'F_0') | sP6(X51,X52,X53) | ~p(s(bool,X52))) ) | $spl144),
  inference(resolution,[],[f20040,f2623])).
fof(f22716,plain,(
  ( ! [X50,X48,X49] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),'F_0') | sP6(X48,X49,X50)) ) | $spl144),
  inference(subsumption_resolution,[],[f22648,f306])).
fof(f22648,plain,(
  ( ! [X50,X48,X49] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),'F_0') | sP6(X48,X49,X50) | ~p(s(bool,X49))) ) | $spl144),
  inference(resolution,[],[f20040,f1301])).
fof(f22714,plain,(
  ( ! [X43,X44,X42] : (sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | sP6(X42,X43,X44)) ) | $spl144),
  inference(subsumption_resolution,[],[f22646,f306])).
fof(f22646,plain,(
  ( ! [X43,X44,X42] : (sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | sP6(X42,X43,X44) | ~p(s(bool,X43))) ) | $spl144),
  inference(resolution,[],[f20040,f484])).
fof(f22624,plain,(
  ( ! [X50,X48,X49] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),X48) | sP8(X48,'F_0') | p(s(bool,X49)) | p(s(bool,X50))) ) | $spl144),
  inference(resolution,[],[f20028,f1276])).
fof(f20028,plain,(
  ( ! [X134,X132,X135,X133] : (sP8(X132,X133) | sP8(X133,'F_0') | sP2(X132,X134,X135)) ) | $spl144),
  inference(resolution,[],[f19502,f290])).
fof(f22623,plain,(
  ( ! [X47,X45,X46] : (sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X45) | sP8(X45,'F_0') | p(s(bool,X46)) | p(s(bool,X47))) ) | $spl144),
  inference(resolution,[],[f20028,f459])).
fof(f22622,plain,(
  ( ! [X43,X44,X42] : (sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X42) | sP8(X42,'F_0') | sP3('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X43,X44)) ) | $spl144),
  inference(resolution,[],[f20028,f508])).
fof(f22630,plain,(
  ( ! [X52,X53,X51] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),'F_0') | sP2(X51,X52,X53)) ) | $spl144),
  inference(subsumption_resolution,[],[f22586,f290])).
fof(f22586,plain,(
  ( ! [X52,X53,X51] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),'F_0') | sP2(X51,X52,X53) | ~p(s(bool,X51))) ) | $spl144),
  inference(resolution,[],[f20028,f2623])).
fof(f22629,plain,(
  ( ! [X50,X48,X49] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),'F_0') | sP2(X48,X49,X50)) ) | $spl144),
  inference(subsumption_resolution,[],[f22585,f290])).
fof(f22585,plain,(
  ( ! [X50,X48,X49] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),'F_0') | sP2(X48,X49,X50) | ~p(s(bool,X48))) ) | $spl144),
  inference(resolution,[],[f20028,f1301])).
fof(f22627,plain,(
  ( ! [X43,X44,X42] : (sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | sP2(X42,X43,X44)) ) | $spl144),
  inference(subsumption_resolution,[],[f22583,f290])).
fof(f22583,plain,(
  ( ! [X43,X44,X42] : (sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | sP2(X42,X43,X44) | ~p(s(bool,X42))) ) | $spl144),
  inference(resolution,[],[f20028,f484])).
fof(f22528,plain,(
  ( ! [X50,X48,X49] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),X48) | sP8(X48,'F_0') | ~p(s(bool,X49)) | p(s(bool,X50))) ) | $spl144),
  inference(resolution,[],[f20012,f1262])).
fof(f20012,plain,(
  ( ! [X70,X68,X71,X69] : (sP8(X68,X69) | sP8(X69,'F_0') | sP0(X68,X70,X71)) ) | $spl144),
  inference(resolution,[],[f19502,f273])).
fof(f22527,plain,(
  ( ! [X47,X45,X46] : (sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X45) | sP8(X45,'F_0') | ~p(s(bool,X46)) | p(s(bool,X47))) ) | $spl144),
  inference(resolution,[],[f20012,f445])).
fof(f22526,plain,(
  ( ! [X43,X44,X42] : (sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X42) | sP8(X42,'F_0') | sP1('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X43,X44)) ) | $spl144),
  inference(resolution,[],[f20012,f507])).
fof(f22533,plain,(
  ( ! [X52,X53,X51] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),'F_0') | sP0(X51,X52,X53)) ) | $spl144),
  inference(subsumption_resolution,[],[f22490,f273])).
fof(f22490,plain,(
  ( ! [X52,X53,X51] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),'F_0') | sP0(X51,X52,X53) | ~p(s(bool,X51))) ) | $spl144),
  inference(resolution,[],[f20012,f2623])).
fof(f22532,plain,(
  ( ! [X50,X48,X49] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),'F_0') | sP0(X48,X49,X50)) ) | $spl144),
  inference(subsumption_resolution,[],[f22489,f273])).
fof(f22489,plain,(
  ( ! [X50,X48,X49] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),'F_0') | sP0(X48,X49,X50) | ~p(s(bool,X48))) ) | $spl144),
  inference(resolution,[],[f20012,f1301])).
fof(f22531,plain,(
  ( ! [X43,X44,X42] : (sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | sP0(X42,X43,X44)) ) | $spl144),
  inference(subsumption_resolution,[],[f22487,f273])).
fof(f22487,plain,(
  ( ! [X43,X44,X42] : (sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | sP0(X42,X43,X44) | ~p(s(bool,X42))) ) | $spl144),
  inference(resolution,[],[f20012,f484])).
fof(f22457,plain,(
  ( ! [X30,X31,X32] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),X30) | sP8(X30,'F_0') | p(s(bool,X31)) | ~p(s(bool,X32))) ) | $spl144),
  inference(resolution,[],[f20011,f1263])).
fof(f20011,plain,(
  ( ! [X66,X64,X67,X65] : (sP8(X64,X65) | sP8(X65,'F_0') | sP0(X66,X67,X64)) ) | $spl144),
  inference(resolution,[],[f19502,f271])).
fof(f22456,plain,(
  ( ! [X28,X29,X27] : (sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X27) | sP8(X27,'F_0') | p(s(bool,X28)) | ~p(s(bool,X29))) ) | $spl144),
  inference(resolution,[],[f20011,f446])).
fof(f22455,plain,(
  ( ! [X26,X24,X25] : (sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X24) | sP8(X24,'F_0') | sP1(X25,X26,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | p(s(bool,X26))) ) | $spl144),
  inference(resolution,[],[f20011,f444])).
fof(f22469,plain,(
  ( ! [X52,X53,X51] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),'F_0') | sP0(X51,X52,X53)) ) | $spl144),
  inference(subsumption_resolution,[],[f22425,f271])).
fof(f22425,plain,(
  ( ! [X52,X53,X51] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),'F_0') | sP0(X51,X52,X53) | ~p(s(bool,X53))) ) | $spl144),
  inference(resolution,[],[f20011,f2623])).
fof(f22468,plain,(
  ( ! [X50,X48,X49] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),'F_0') | sP0(X48,X49,X50)) ) | $spl144),
  inference(subsumption_resolution,[],[f22424,f271])).
fof(f22424,plain,(
  ( ! [X50,X48,X49] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),'F_0') | sP0(X48,X49,X50) | ~p(s(bool,X50))) ) | $spl144),
  inference(resolution,[],[f20011,f1301])).
fof(f22466,plain,(
  ( ! [X43,X44,X42] : (sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | sP0(X42,X43,X44)) ) | $spl144),
  inference(subsumption_resolution,[],[f22422,f271])).
fof(f22422,plain,(
  ( ! [X43,X44,X42] : (sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | sP0(X42,X43,X44) | ~p(s(bool,X44))) ) | $spl144),
  inference(resolution,[],[f20011,f484])).
fof(f22359,plain,(
  ( ! [X41] : (p(s(bool,X41)) | sP8(happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),'F_0')) ) | $spl144),
  inference(resolution,[],[f19590,f2631])).
fof(f2631,plain,(
  ( ! [X123] : (sP9(happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),X123)) )),
  inference(resolution,[],[f1246,f344])).
fof(f22358,plain,(
  ( ! [X40] : (p(s(bool,X40)) | sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),'F_0')) ) | $spl144),
  inference(resolution,[],[f19590,f1309])).
fof(f1309,plain,(
  ( ! [X123] : (sP9(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),X123)) )),
  inference(resolution,[],[f538,f344])).
fof(f22357,plain,(
  ( ! [X39] : (p(s(bool,X39)) | sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0')) ) | $spl144),
  inference(resolution,[],[f19590,f492])).
fof(f492,plain,(
  ( ! [X120] : (sP9('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X120)) )),
  inference(resolution,[],[f250,f344])).
fof(f22349,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | sP8(X26,'F_0') | sP8(X26,X27) | s(bool,'F_0') = s(bool,X27)) ) | $spl144),
  inference(resolution,[],[f19590,f17151])).
fof(f17151,plain,(
  ( ! [X241,X239,X240] : (sP8(X239,X240) | sP9(X239,X241) | s(bool,'F_0') = s(bool,X240)) )),
  inference(resolution,[],[f16182,f344])).
fof(f22328,plain,(
  ( ! [X38] : (p(s(bool,X38)) | sP8(happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),'F_0')) ) | $spl144),
  inference(resolution,[],[f19589,f2630])).
fof(f2630,plain,(
  ( ! [X122] : (sP9(X122,happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))))) )),
  inference(resolution,[],[f1246,f343])).
fof(f22327,plain,(
  ( ! [X37] : (p(s(bool,X37)) | sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),'F_0')) ) | $spl144),
  inference(resolution,[],[f19589,f1308])).
fof(f1308,plain,(
  ( ! [X122] : (sP9(X122,happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))))) )),
  inference(resolution,[],[f538,f343])).
fof(f22326,plain,(
  ( ! [X36] : (p(s(bool,X36)) | sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0')) ) | $spl144),
  inference(resolution,[],[f19589,f491])).
fof(f491,plain,(
  ( ! [X119] : (sP9(X119,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) )),
  inference(resolution,[],[f250,f343])).
fof(f22320,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | sP8(X23,'F_0') | sP8(X23,X24) | s(bool,'F_0') = s(bool,X24)) ) | $spl144),
  inference(resolution,[],[f19589,f17150])).
fof(f17150,plain,(
  ( ! [X237,X238,X236] : (sP8(X236,X237) | sP9(X238,X236) | s(bool,'F_0') = s(bool,X237)) )),
  inference(resolution,[],[f16182,f343])).
fof(f22244,plain,(
  ( ! [X47,X46] : (sP8('F_0',happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11)))) | p(s(bool,X46)) | ~p(s(bool,X47))) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f21895,f1292])).
fof(f21895,plain,(
  ( ! [X151,X149,X150] : (sP5(X149,X150,X151) | sP8('F_0',X149)) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f21837,f311])).
fof(f22243,plain,(
  ( ! [X45,X44] : (sP8('F_0','const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | p(s(bool,X44)) | ~p(s(bool,X45))) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f21895,f475])).
fof(f22221,plain,(
  ( ! [X41,X40] : (sP8('F_0',happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11)))) | p(s(bool,X40)) | ~p(s(bool,X41))) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f21894,f1293])).
fof(f21894,plain,(
  ( ! [X146,X147,X148] : (sP5(X147,X146,X148) | sP8('F_0',X146)) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f21837,f310])).
fof(f22220,plain,(
  ( ! [X39,X38] : (sP8('F_0','const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | p(s(bool,X38)) | ~p(s(bool,X39))) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f21894,f476])).
fof(f22185,plain,(
  ( ! [X37,X38,X36] : (p(s(bool,X36)) | ~p(s(bool,X37)) | sP8(happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X38)))),'T_0')) ) | $spl144),
  inference(resolution,[],[f3019,f21120])).
fof(f21120,plain,(
  ( ! [X94,X92,X93] : (sP2(X93,X94,X92) | sP8(X92,'T_0')) ) | $spl144),
  inference(resolution,[],[f21077,f288])).
fof(f22183,plain,(
  ( ! [X33,X31,X32] : (p(s(bool,X31)) | ~p(s(bool,X32)) | sP8('T_0',happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X33)))))) ) | $spl144),
  inference(resolution,[],[f3019,f20523])).
fof(f20523,plain,(
  ( ! [X94,X92,X93] : (sP2(X93,X94,X92) | sP8('T_0',X92)) ) | $spl144),
  inference(resolution,[],[f20335,f288])).
fof(f22180,plain,(
  ( ! [X24,X23,X25,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | sP8(happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X24)))),X25) | s(bool,'F_0') = s(bool,X25)) )),
  inference(resolution,[],[f3019,f17119])).
fof(f17119,plain,(
  ( ! [X123,X121,X122,X120] : (sP8(X120,X121) | sP2(X122,X123,X120) | s(bool,'F_0') = s(bool,X121)) )),
  inference(resolution,[],[f16182,f288])).
fof(f22178,plain,(
  ( ! [X17,X15,X18,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | sP8(X17,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X18))))) | s(bool,'F_0') = s(bool,X17)) )),
  inference(resolution,[],[f3019,f16882])).
fof(f16882,plain,(
  ( ! [X123,X121,X122,X120] : (sP8(X120,X121) | sP2(X122,X123,X121) | s(bool,'F_0') = s(bool,X120)) )),
  inference(resolution,[],[f16178,f288])).
fof(f22166,plain,(
  ( ! [X97] : (sP8('F_0',happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13))))) | ~p(s(bool,X97))) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f21892,f12547])).
fof(f21892,plain,(
  ( ! [X142,X140,X141] : (sP6(X140,X141,X142) | sP8('F_0',X140)) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f21837,f307])).
fof(f22165,plain,(
  ( ! [X95,X96] : (sP8('F_0',happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11)))) | p(s(bool,X95)) | ~p(s(bool,X96))) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f21892,f1289])).
fof(f22164,plain,(
  ( ! [X94] : (sP8('F_0',happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11)))) | ~p(s(bool,X94))) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f21892,f9596])).
fof(f22163,plain,(
  ( ! [X92,X93] : (sP8('F_0','const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | p(s(bool,X92)) | ~p(s(bool,X93))) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f21892,f472])).
fof(f22162,plain,(
  ( ! [X90,X91] : (sP8('F_0','const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | ~p(s(bool,X90)) | s(bool,X90) = s(bool,X91)) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f21892,f3259])).
fof(f22161,plain,(
  ( ! [X88,X89] : (sP8('F_0','const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | p(s(bool,X88)) | s(bool,X88) = s(bool,X89)) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f21892,f3261])).
fof(f22160,plain,(
  ( ! [X87] : (sP8('F_0','const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | ~p(s(bool,X87))) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f21892,f3271])).
fof(f22117,plain,(
  ( ! [X85,X86] : (sP8('F_0',happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11)))) | p(s(bool,X85)) | ~p(s(bool,X86))) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f21891,f1290])).
fof(f21891,plain,(
  ( ! [X138,X139,X137] : (sP6(X138,X137,X139) | sP8('F_0',X137)) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f21837,f306])).
fof(f22116,plain,(
  ( ! [X83,X84] : (sP8('F_0','const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | p(s(bool,X83)) | ~p(s(bool,X84))) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f21891,f473])).
fof(f22115,plain,(
  ( ! [X81,X82] : (sP8('F_0','const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | sP7(X81,X82,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | ~sP5(X82,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X81)) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f21891,f6665])).
fof(f22127,plain,(
  ( ! [X65,X63] : (sP8('F_0',X63) | s(bool,'T_0') != s(bool,X65) | s(bool,X63) != s(bool,X65)) ) | $spl236),
  inference(cnf_transformation,[],[f22127_D])).
fof(f22127_D,plain,(
  ( ! [X65,X63] : (sP8('F_0',X63) | s(bool,'T_0') != s(bool,X65) | s(bool,X63) != s(bool,X65)) ) <=> ~$spl236),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl236])])).
fof(f22108,plain,(
  ( ! [X64,X65,X63] : (sP8('F_0',X63) | ~p(s(bool,X64)) | s(bool,X63) != s(bool,X65) | s(bool,'T_0') != s(bool,X65)) ) | ($spl24 | $spl144 | $spl161)),
  inference(resolution,[],[f21891,f6588])).
fof(f22126,plain,(
  ( ! [X27] : (sP8('F_0',X27) | s(bool,'T_0') != s(bool,X27)) ) | $spl234),
  inference(cnf_transformation,[],[f22126_D])).
fof(f22126_D,plain,(
  ( ! [X27] : (sP8('F_0',X27) | s(bool,'T_0') != s(bool,X27)) ) <=> ~$spl234),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl234])])).
fof(f22106,plain,(
  ( ! [X59,X58] : (sP8('F_0',X58) | s(bool,'T_0') != s(bool,X58) | ~p(s(bool,X59))) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f21891,f2088])).
fof(f21991,plain,(
  ( ! [X2,X0] : (sP8('F_0',X0) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) | $spl228),
  inference(cnf_transformation,[],[f21991_D])).
fof(f21991_D,plain,(
  ( ! [X2,X0] : (sP8('F_0',X0) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) <=> ~$spl228),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl228])])).
fof(f22096,plain,(
  ( ! [X30,X33,X31,X32] : (sP8('F_0',X30) | ~p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X30) != s(bool,X31) | p(s(bool,X33))) ) | ($spl6 | $spl144 | $spl161)),
  inference(resolution,[],[f21891,f8286])).
fof(f22095,plain,(
  ( ! [X28,X29,X27] : (sP8('F_0',X27) | ~p(s(bool,X28)) | s(bool,'F_0') != s(bool,X29) | s(bool,'T_0') != s(bool,X27)) ) | ($spl24 | $spl144 | $spl161)),
  inference(resolution,[],[f21891,f7489])).
fof(f22085,plain,(
  ( ! [X33,X32] : (sP8('F_0',happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11)))) | p(s(bool,X32)) | p(s(bool,X33))) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f21879,f1276])).
fof(f21879,plain,(
  ( ! [X103,X101,X102] : (sP2(X101,X102,X103) | sP8('F_0',X101)) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f21837,f290])).
fof(f22084,plain,(
  ( ! [X30,X31] : (sP8('F_0','const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | p(s(bool,X30)) | p(s(bool,X31))) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f21879,f459])).
fof(f22083,plain,(
  ( ! [X28,X29] : (sP8('F_0','const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | sP3('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X28,X29)) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f21879,f508])).
fof(f22071,plain,(
  ( ! [X35,X34] : (sP8('F_0',happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11)))) | ~p(s(bool,X34)) | p(s(bool,X35))) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f21863,f1262])).
fof(f21863,plain,(
  ( ! [X54,X55,X53] : (sP0(X53,X54,X55) | sP8('F_0',X53)) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f21837,f273])).
fof(f22070,plain,(
  ( ! [X33,X32] : (sP8('F_0','const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | ~p(s(bool,X32)) | p(s(bool,X33))) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f21863,f445])).
fof(f22069,plain,(
  ( ! [X30,X31] : (sP8('F_0','const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | sP1('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X30,X31)) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f21863,f507])).
fof(f22047,plain,(
  ( ! [X23,X22] : (sP8('F_0',happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11)))) | p(s(bool,X22)) | ~p(s(bool,X23))) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f21862,f1263])).
fof(f21862,plain,(
  ( ! [X52,X50,X51] : (sP0(X51,X52,X50) | sP8('F_0',X50)) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f21837,f271])).
fof(f22046,plain,(
  ( ! [X21,X20] : (sP8('F_0','const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | p(s(bool,X20)) | ~p(s(bool,X21))) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f21862,f446])).
fof(f22045,plain,(
  ( ! [X19,X18] : (sP8('F_0','const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | sP1(X18,X19,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | p(s(bool,X19))) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f21862,f444])).
fof(f22020,plain,(
  ( ! [X39,X38,X40] : (p(s(bool,X38)) | ~p(s(bool,X39)) | sP8(happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X40)))),'T_0')) ) | $spl144),
  inference(resolution,[],[f3018,f21121])).
fof(f21121,plain,(
  ( ! [X97,X95,X96] : (sP2(X96,X95,X97) | sP8(X95,'T_0')) ) | $spl144),
  inference(resolution,[],[f21077,f289])).
fof(f22018,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X33)) | ~p(s(bool,X34)) | sP8('T_0',happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X35)))))) ) | $spl144),
  inference(resolution,[],[f3018,f20524])).
fof(f20524,plain,(
  ( ! [X97,X95,X96] : (sP2(X96,X95,X97) | sP8('T_0',X95)) ) | $spl144),
  inference(resolution,[],[f20335,f289])).
fof(f22015,plain,(
  ( ! [X28,X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | sP8(happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X27)))),X28) | s(bool,'F_0') = s(bool,X28)) )),
  inference(resolution,[],[f3018,f17120])).
fof(f17120,plain,(
  ( ! [X127,X125,X126,X124] : (sP8(X124,X125) | sP2(X126,X124,X127) | s(bool,'F_0') = s(bool,X125)) )),
  inference(resolution,[],[f16182,f289])).
fof(f22013,plain,(
  ( ! [X21,X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | sP8(X20,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X21))))) | s(bool,'F_0') = s(bool,X20)) )),
  inference(resolution,[],[f3018,f16883])).
fof(f16883,plain,(
  ( ! [X127,X125,X126,X124] : (sP8(X124,X125) | sP2(X126,X125,X127) | s(bool,'F_0') = s(bool,X124)) )),
  inference(resolution,[],[f16178,f289])).
fof(f22000,plain,(
  ( ! [X33] : (~p(s(bool,X33)) | sP8('F_0','const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) ) | ($spl144 | $spl161)),
  inference(subsumption_resolution,[],[f21980,f397])).
fof(f21980,plain,(
  ( ! [X33] : (~p(s(bool,X33)) | sP8('F_0','const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | s(bool,'F_0') = s(bool,X33)) ) | ($spl144 | $spl161)),
  inference(superposition,[],[f21837,f14557])).
fof(f21926,plain,(
  ( ! [X208,X206,X207] : (sP8('F_0',X206) | p(s(bool,X207)) | ~p(s(bool,X208)) | s(bool,X206) != s(bool,X208)) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f21837,f426])).
fof(f21855,plain,(
  ( ! [X30,X28,X31,X29] : (sP8('F_0',X28) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f21837,f262])).
fof(f21853,plain,(
  ( ! [X23,X21,X22,X20] : (sP8('F_0',X20) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X20) != s(bool,X22)) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f21837,f261])).
fof(f21851,plain,(
  ( ! [X14,X12,X15,X13] : (sP8('F_0',X12) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f21837,f260])).
fof(f21850,plain,(
  ( ! [X10,X8,X11,X9] : (sP8('F_0',X8) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f21837,f258])).
fof(f21848,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('F_0',X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f21837,f256])).
fof(f21842,plain,(
  ( ! [X8] : (sP8('F_0',X8) | ~p(s(bool,X8))) ) | $spl226),
  inference(cnf_transformation,[],[f21842_D])).
fof(f21842_D,plain,(
  ( ! [X8] : (sP8('F_0',X8) | ~p(s(bool,X8))) ) <=> ~$spl226),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl226])])).
fof(f21813,plain,(
  ( ! [X8,X9] : (sP8('F_0',X8) | s(bool,'F_0') = s(bool,X9) | ~p(s(bool,X8)) | sP8(X9,'T_0')) ) | ($spl144 | $spl161)),
  inference(resolution,[],[f21782,f16906])).
fof(f21800,plain,(
  ( ! [X31] : (sP8(X31,'T_0') | sP8(happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),'F_0')) ) | $spl144),
  inference(subsumption_resolution,[],[f21770,f21077])).
fof(f21770,plain,(
  ( ! [X31] : (sP8(X31,'T_0') | sP8(happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),'F_0') | ~p(s(bool,X31))) ) | $spl144),
  inference(resolution,[],[f21171,f2622])).
fof(f21799,plain,(
  ( ! [X30] : (sP8(X30,'T_0') | sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),'F_0')) ) | $spl144),
  inference(subsumption_resolution,[],[f21769,f21077])).
fof(f21769,plain,(
  ( ! [X30] : (sP8(X30,'T_0') | sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),'F_0') | ~p(s(bool,X30))) ) | $spl144),
  inference(resolution,[],[f21171,f1300])).
fof(f21798,plain,(
  ( ! [X28] : (sP8(X28,'T_0') | sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0')) ) | $spl144),
  inference(subsumption_resolution,[],[f21767,f21077])).
fof(f21767,plain,(
  ( ! [X28] : (sP8(X28,'T_0') | sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | ~p(s(bool,X28))) ) | $spl144),
  inference(resolution,[],[f21171,f483])).
fof(f21797,plain,(
  ( ! [X26,X27] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X26)))),'T_0') | sP8(X27,'F_0')) ) | $spl144),
  inference(subsumption_resolution,[],[f21765,f19403])).
fof(f21765,plain,(
  ( ! [X26,X27] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X26)))),'T_0') | sP8(X27,'F_0') | p(s(bool,X27))) ) | $spl144),
  inference(resolution,[],[f21171,f3042])).
fof(f21796,plain,(
  ( ! [X23] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0') | sP8(X23,'F_0')) ) | $spl144),
  inference(subsumption_resolution,[],[f21762,f19403])).
fof(f21762,plain,(
  ( ! [X23] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0') | sP8(X23,'F_0') | p(s(bool,X23))) ) | $spl144),
  inference(resolution,[],[f21171,f593])).
fof(f21673,plain,(
  ( ! [X33,X32] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X32)))),'T_0') | sP8(X33,'F_0')) ) | $spl144),
  inference(subsumption_resolution,[],[f21644,f19403])).
fof(f21644,plain,(
  ( ! [X33,X32] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X32)))),'T_0') | sP8(X33,'F_0') | p(s(bool,X33))) ) | $spl144),
  inference(resolution,[],[f21170,f3043])).
fof(f21170,plain,(
  ( ! [X208,X209] : (sP8(X208,'T_0') | sP8(X208,X209) | sP8(X209,'F_0')) ) | $spl144),
  inference(resolution,[],[f21077,f19502])).
fof(f21672,plain,(
  ( ! [X29] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0') | sP8(X29,'F_0')) ) | $spl144),
  inference(subsumption_resolution,[],[f21641,f19403])).
fof(f21641,plain,(
  ( ! [X29] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0') | sP8(X29,'F_0') | p(s(bool,X29))) ) | $spl144),
  inference(resolution,[],[f21170,f594])).
fof(f21671,plain,(
  ( ! [X25] : (sP8(X25,'T_0') | sP8(happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),'F_0')) ) | $spl144),
  inference(subsumption_resolution,[],[f21637,f21077])).
fof(f21637,plain,(
  ( ! [X25] : (sP8(X25,'T_0') | sP8(happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),'F_0') | ~p(s(bool,X25))) ) | $spl144),
  inference(resolution,[],[f21170,f2623])).
fof(f21670,plain,(
  ( ! [X24] : (sP8(X24,'T_0') | sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),'F_0')) ) | $spl144),
  inference(subsumption_resolution,[],[f21636,f21077])).
fof(f21636,plain,(
  ( ! [X24] : (sP8(X24,'T_0') | sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),'F_0') | ~p(s(bool,X24))) ) | $spl144),
  inference(resolution,[],[f21170,f1301])).
fof(f21668,plain,(
  ( ! [X22] : (sP8(X22,'T_0') | sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0')) ) | $spl144),
  inference(subsumption_resolution,[],[f21634,f21077])).
fof(f21634,plain,(
  ( ! [X22] : (sP8(X22,'T_0') | sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | ~p(s(bool,X22))) ) | $spl144),
  inference(resolution,[],[f21170,f484])).
fof(f21580,plain,(
  ( ! [X31] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),'F_0') | sP8('T_0',X31)) ) | $spl144),
  inference(subsumption_resolution,[],[f21538,f20335])).
fof(f21538,plain,(
  ( ! [X31] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),'F_0') | sP8('T_0',X31) | ~p(s(bool,X31))) ) | $spl144),
  inference(resolution,[],[f20973,f2622])).
fof(f20973,plain,(
  ( ! [X302,X303] : (sP8(X302,X303) | sP8(X302,'F_0') | sP8('T_0',X303)) ) | $spl144),
  inference(resolution,[],[f19503,f20335])).
fof(f21579,plain,(
  ( ! [X30] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),'F_0') | sP8('T_0',X30)) ) | $spl144),
  inference(subsumption_resolution,[],[f21537,f20335])).
fof(f21537,plain,(
  ( ! [X30] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),'F_0') | sP8('T_0',X30) | ~p(s(bool,X30))) ) | $spl144),
  inference(resolution,[],[f20973,f1300])).
fof(f21578,plain,(
  ( ! [X28] : (sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | sP8('T_0',X28)) ) | $spl144),
  inference(subsumption_resolution,[],[f21535,f20335])).
fof(f21535,plain,(
  ( ! [X28] : (sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | sP8('T_0',X28) | ~p(s(bool,X28))) ) | $spl144),
  inference(resolution,[],[f20973,f483])).
fof(f21577,plain,(
  ( ! [X26,X27] : (sP8(X26,'F_0') | sP8('T_0',happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X27)))))) ) | $spl144),
  inference(subsumption_resolution,[],[f21533,f19403])).
fof(f21533,plain,(
  ( ! [X26,X27] : (sP8(X26,'F_0') | sP8('T_0',happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X27))))) | p(s(bool,X26))) ) | $spl144),
  inference(resolution,[],[f20973,f3042])).
fof(f21576,plain,(
  ( ! [X23] : (sP8(X23,'F_0') | sP8('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) ) | $spl144),
  inference(subsumption_resolution,[],[f21530,f19403])).
fof(f21530,plain,(
  ( ! [X23] : (sP8(X23,'F_0') | sP8('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | p(s(bool,X23))) ) | $spl144),
  inference(resolution,[],[f20973,f593])).
fof(f21515,plain,(
  ( ! [X41,X42] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0') | p(s(bool,X41)) | p(s(bool,X42))) ) | $spl144),
  inference(resolution,[],[f21140,f589])).
fof(f589,plain,(
  ( ! [X107,X108] : (~sP4('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X107,X108) | p(s(bool,X108)) | p(s(bool,X107))) )),
  inference(resolution,[],[f251,f312])).
fof(f21394,plain,(
  ( ! [X24] : (sP8(X24,'T_0') | s(bool,'F_0') != s(bool,X24)) ) | $spl216),
  inference(cnf_transformation,[],[f21394_D])).
fof(f21394_D,plain,(
  ( ! [X24] : (sP8(X24,'T_0') | s(bool,'F_0') != s(bool,X24)) ) <=> ~$spl216),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl216])])).
fof(f21511,plain,(
  ( ! [X35,X34] : (sP8(X34,'T_0') | p(s(bool,X35)) | s(bool,'F_0') != s(bool,X34)) ) | $spl144),
  inference(resolution,[],[f21140,f1971])).
fof(f1971,plain,(
  ( ! [X4,X5] : (~sP4(X5,'F_0',X4) | p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5)) )),
  inference(duplicate_literal_removal,[],[f1966])).
fof(f1966,plain,(
  ( ! [X4,X5] : (p(s(bool,X4)) | ~sP4(X5,'F_0',X4) | p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5)) )),
  inference(resolution,[],[f1939,f319])).
fof(f1939,plain,(
  ( ! [X6,X7] : (sP7(X6,'F_0',X7) | p(s(bool,X6)) | ~sP4(X7,'F_0',X6)) )),
  inference(subsumption_resolution,[],[f1937,f307])).
fof(f1937,plain,(
  ( ! [X6,X7] : (sP7(X6,'F_0',X7) | p(s(bool,X6)) | ~sP6(X6,X7,'F_0') | ~sP4(X7,'F_0',X6)) )),
  inference(resolution,[],[f1933,f303])).
fof(f1933,plain,(
  ( ! [X4,X5] : (sP5('F_0',X4,X5)) )),
  inference(subsumption_resolution,[],[f1929,f309])).
fof(f1929,plain,(
  ( ! [X4,X5] : (sP5('F_0',X4,X5) | p(s(bool,X5))) )),
  inference(resolution,[],[f1926,f307])).
fof(f1926,plain,(
  ( ! [X6,X7] : (~sP6(X7,X6,'F_0') | sP5('F_0',X6,X7)) )),
  inference(subsumption_resolution,[],[f1922,f310])).
fof(f1922,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | ~sP6(X7,X6,'F_0') | sP5('F_0',X6,X7)) )),
  inference(resolution,[],[f1560,f298])).
fof(f21495,plain,(
  ( ! [X39,X40] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0') | p(s(bool,X39)) | p(s(bool,X40))) ) | $spl144),
  inference(resolution,[],[f21139,f588])).
fof(f588,plain,(
  ( ! [X105,X106] : (~sP4(X105,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X106) | p(s(bool,X106)) | p(s(bool,X105))) )),
  inference(resolution,[],[f251,f312])).
fof(f21431,plain,(
  ( ! [X38] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0') | p(s(bool,X38))) ) | $spl144),
  inference(resolution,[],[f21138,f3375])).
fof(f3375,plain,(
  ( ! [X13] : (~sP4(X13,'F_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | p(s(bool,X13))) )),
  inference(subsumption_resolution,[],[f3362,f398])).
fof(f3362,plain,(
  ( ! [X13] : (s(bool,'F_0') != s(bool,X13) | ~sP4(X13,'F_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | p(s(bool,X13))) )),
  inference(resolution,[],[f592,f1958])).
fof(f1958,plain,(
  ( ! [X2,X3] : (sP7(X2,'F_0',X3) | ~sP4(X3,'F_0',X2) | p(s(bool,X3))) )),
  inference(resolution,[],[f1938,f306])).
fof(f1938,plain,(
  ( ! [X4,X5] : (~sP6(X4,X5,'F_0') | sP7(X4,'F_0',X5) | ~sP4(X5,'F_0',X4)) )),
  inference(subsumption_resolution,[],[f1936,f412])).
fof(f1936,plain,(
  ( ! [X4,X5] : (sP7(X4,'F_0',X5) | p(s(bool,'F_0')) | ~sP6(X4,X5,'F_0') | ~sP4(X5,'F_0',X4)) )),
  inference(resolution,[],[f1933,f302])).
fof(f21428,plain,(
  ( ! [X33,X32] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0') | p(s(bool,X32)) | p(s(bool,X33))) ) | $spl144),
  inference(resolution,[],[f21138,f590])).
fof(f590,plain,(
  ( ! [X109,X110] : (~sP4(X109,X110,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | p(s(bool,X110)) | p(s(bool,X109))) )),
  inference(resolution,[],[f251,f312])).
fof(f21427,plain,(
  ( ! [X30,X31] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0') | p(s(bool,X30)) | s(bool,X30) != s(bool,X31)) ) | $spl144),
  inference(resolution,[],[f21138,f3371])).
fof(f3371,plain,(
  ( ! [X8,X9] : (~sP4(X8,X9,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | p(s(bool,X9)) | s(bool,X8) != s(bool,X9)) )),
  inference(subsumption_resolution,[],[f3370,f584])).
fof(f584,plain,(
  ( ! [X97,X98] : (sP6('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X97,X98)) )),
  inference(resolution,[],[f251,f307])).
fof(f3370,plain,(
  ( ! [X8,X9] : (s(bool,X8) != s(bool,X9) | p(s(bool,X9)) | ~sP6('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X8,X9) | ~sP4(X8,X9,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) )),
  inference(subsumption_resolution,[],[f3359,f311])).
fof(f3359,plain,(
  ( ! [X8,X9] : (s(bool,X8) != s(bool,X9) | p(s(bool,X9)) | ~sP6('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X8,X9) | ~sP5(X9,X8,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | ~sP4(X8,X9,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) )),
  inference(resolution,[],[f592,f302])).
fof(f21409,plain,(
  ( ! [X35,X34] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0') | ~p(s(bool,X34)) | ~p(s(bool,X35))) ) | $spl144),
  inference(resolution,[],[f21137,f585])).
fof(f585,plain,(
  ( ! [X99,X100] : (~sP5(X99,X100,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | ~p(s(bool,X100)) | ~p(s(bool,X99))) )),
  inference(resolution,[],[f251,f308])).
fof(f17723,plain,(
  ( ! [X41,X42] : (sP7('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X41,X42) | ~sP4(X42,X41,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) ) | $spl132),
  inference(cnf_transformation,[],[f17723_D])).
fof(f17723_D,plain,(
  ( ! [X41,X42] : (sP7('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X41,X42) | ~sP4(X42,X41,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f21408,plain,(
  ( ! [X33,X32] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0') | sP7('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X32,X33) | ~sP4(X33,X32,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) ) | $spl144),
  inference(resolution,[],[f21137,f7261])).
fof(f7261,plain,(
  ( ! [X92,X91] : (~sP5(X91,X92,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | sP7('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X91,X92) | ~sP4(X92,X91,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) )),
  inference(subsumption_resolution,[],[f581,f584])).
fof(f581,plain,(
  ( ! [X92,X91] : (sP7('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X91,X92) | ~sP6('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X92,X91) | ~sP5(X91,X92,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | ~sP4(X92,X91,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) )),
  inference(resolution,[],[f251,f303])).
fof(f21407,plain,(
  ( ! [X30,X31] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0') | ~p(s(bool,X30)) | s(bool,X30) != s(bool,X31)) ) | $spl144),
  inference(resolution,[],[f21137,f3369])).
fof(f3369,plain,(
  ( ! [X6,X7] : (~sP5(X7,X6,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | ~p(s(bool,X6)) | s(bool,X6) != s(bool,X7)) )),
  inference(subsumption_resolution,[],[f3368,f584])).
fof(f3368,plain,(
  ( ! [X6,X7] : (s(bool,X6) != s(bool,X7) | ~p(s(bool,X6)) | ~sP6('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X6,X7) | ~sP5(X7,X6,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) )),
  inference(subsumption_resolution,[],[f3358,f315])).
fof(f3358,plain,(
  ( ! [X6,X7] : (s(bool,X6) != s(bool,X7) | ~p(s(bool,X6)) | ~sP6('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X6,X7) | ~sP5(X7,X6,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | ~sP4(X6,X7,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) )),
  inference(resolution,[],[f592,f301])).
fof(f21396,plain,(
  ( ! [X75,X73] : (sP8(X73,'T_0') | s(bool,X73) != s(bool,X75) | s(bool,'T_0') = s(bool,X75)) ) | $spl220),
  inference(cnf_transformation,[],[f21396_D])).
fof(f21396_D,plain,(
  ( ! [X75,X73] : (sP8(X73,'T_0') | s(bool,X73) != s(bool,X75) | s(bool,'T_0') = s(bool,X75)) ) <=> ~$spl220),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl220])])).
fof(f21380,plain,(
  ( ! [X74,X75,X73] : (sP8(X73,'T_0') | ~p(s(bool,X74)) | s(bool,'T_0') = s(bool,X75) | s(bool,X73) != s(bool,X75)) ) | $spl144),
  inference(resolution,[],[f21134,f14828])).
fof(f14828,plain,(
  ( ! [X12,X10,X11] : (~sP6(X12,'T_0',X10) | ~p(s(bool,X12)) | s(bool,'T_0') = s(bool,X11) | s(bool,X10) != s(bool,X11)) )),
  inference(subsumption_resolution,[],[f14827,f313])).
fof(f14827,plain,(
  ( ! [X12,X10,X11] : (s(bool,X10) != s(bool,X11) | ~p(s(bool,X12)) | s(bool,'T_0') = s(bool,X11) | ~sP6(X12,'T_0',X10) | ~sP4('T_0',X10,X12)) )),
  inference(subsumption_resolution,[],[f14826,f309])).
fof(f14826,plain,(
  ( ! [X12,X10,X11] : (s(bool,X10) != s(bool,X11) | ~p(s(bool,X12)) | s(bool,'T_0') = s(bool,X11) | ~sP6(X12,'T_0',X10) | ~sP5(X10,'T_0',X12) | ~sP4('T_0',X10,X12)) )),
  inference(subsumption_resolution,[],[f14768,f377])).
fof(f14768,plain,(
  ( ! [X12,X10,X11] : (s(bool,X10) != s(bool,X11) | ~p(s(bool,X12)) | s(bool,'T_0') = s(bool,X11) | ~p(s(bool,'T_0')) | ~sP6(X12,'T_0',X10) | ~sP5(X10,'T_0',X12) | ~sP4('T_0',X10,X12)) )),
  inference(resolution,[],[f14564,f301])).
fof(f14564,plain,(
  ( ! [X14,X12,X13] : (~sP7(X14,X12,'T_0') | s(bool,X12) != s(bool,X13) | ~p(s(bool,X14)) | s(bool,'T_0') = s(bool,X13)) )),
  inference(forward_demodulation,[],[f14516,f504])).
fof(f14516,plain,(
  ( ! [X14,X12,X13] : (s(bool,X12) != s(bool,X13) | s(bool,X13) = s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | ~p(s(bool,X14)) | ~sP7(X14,X12,'T_0')) )),
  inference(superposition,[],[f14143,f318])).
fof(f21378,plain,(
  ( ! [X68,X69] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0') | ~p(s(bool,X68)) | ~p(s(bool,X69))) ) | $spl144),
  inference(resolution,[],[f21134,f582])).
fof(f582,plain,(
  ( ! [X94,X93] : (~sP6(X93,X94,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | ~p(s(bool,X94)) | ~p(s(bool,X93))) )),
  inference(resolution,[],[f251,f304])).
fof(f17685,plain,(
  ( ! [X90,X91] : (sP7(X90,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X91) | ~sP4(X91,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X90)) ) | $spl128),
  inference(cnf_transformation,[],[f17685_D])).
fof(f17685_D,plain,(
  ( ! [X90,X91] : (sP7(X90,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X91) | ~sP4(X91,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X90)) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f21377,plain,(
  ( ! [X66,X67] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0') | sP7(X66,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X67) | ~sP4(X67,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X66)) ) | $spl144),
  inference(resolution,[],[f21134,f6889])).
fof(f6889,plain,(
  ( ! [X90,X89] : (~sP6(X89,X90,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | sP7(X89,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X90) | ~sP4(X90,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X89)) )),
  inference(subsumption_resolution,[],[f580,f587])).
fof(f587,plain,(
  ( ! [X103,X104] : (sP5('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X103,X104)) )),
  inference(resolution,[],[f251,f311])).
fof(f580,plain,(
  ( ! [X90,X89] : (sP7(X89,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X90) | ~sP6(X89,X90,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | ~sP5('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X90,X89) | ~sP4(X90,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X89)) )),
  inference(resolution,[],[f251,f302])).
fof(f7590,plain,(
  ( ! [X60] : (s(bool,'F_0') = s(bool,X60)) ) | $spl66),
  inference(cnf_transformation,[],[f7590_D])).
fof(f7590_D,plain,(
  ( ! [X60] : (s(bool,'F_0') = s(bool,X60)) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f21395,plain,(
  ( ! [X44] : (sP8(X44,'T_0') | s(bool,'T_0') = s(bool,X44)) ) | $spl218),
  inference(cnf_transformation,[],[f21395_D])).
fof(f21395_D,plain,(
  ( ! [X44] : (sP8(X44,'T_0') | s(bool,'T_0') = s(bool,X44)) ) <=> ~$spl218),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl218])])).
fof(f21369,plain,(
  ( ! [X47,X48,X49] : (sP8(X47,'T_0') | s(bool,'T_0') = s(bool,X47) | s(bool,'F_0') = s(bool,X48) | ~p(s(bool,X49))) ) | $spl144),
  inference(resolution,[],[f21134,f14390])).
fof(f21368,plain,(
  ( ! [X45,X46,X44] : (sP8(X44,'T_0') | s(bool,'T_0') = s(bool,X44) | ~p(s(bool,X45)) | ~p(s(bool,X46))) ) | $spl144),
  inference(resolution,[],[f21134,f14047])).
fof(f14047,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X2,X1) | s(bool,'T_0') = s(bool,X1) | ~p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f14046,f315])).
fof(f14046,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X1) | ~p(s(bool,X2)) | ~sP6(X0,X2,X1) | ~sP4(X2,X1,X0)) )),
  inference(subsumption_resolution,[],[f14045,f309])).
fof(f14045,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X1) | ~p(s(bool,X2)) | ~sP6(X0,X2,X1) | ~sP5(X1,X2,X0) | ~sP4(X2,X1,X0)) )),
  inference(duplicate_literal_removal,[],[f13966])).
fof(f13966,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X1) | ~p(s(bool,X2)) | ~p(s(bool,X2)) | ~sP6(X0,X2,X1) | ~sP5(X1,X2,X0) | ~sP4(X2,X1,X0)) )),
  inference(resolution,[],[f13955,f301])).
fof(f13955,plain,(
  ( ! [X47,X45,X46] : (~sP7(X47,X46,X45) | ~p(s(bool,X47)) | s(bool,'T_0') = s(bool,X46) | ~p(s(bool,X45))) )),
  inference(trivial_inequality_removal,[],[f13932])).
fof(f13932,plain,(
  ( ! [X47,X45,X46] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'T_0') = s(bool,X46) | ~p(s(bool,X47)) | ~sP7(X47,X46,X45) | ~p(s(bool,X45))) )),
  inference(superposition,[],[f13735,f404])).
fof(f404,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f243])).
fof(f243,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X13] : (s(bool,'T_0') = s(bool,X13) <=> p(s(bool,X13)))),
  file('thm.quantHeuristics.GUESSES_WEAKEN_THM|split|3','thm.bool.EQ_CLAUSES|split|1')).
fof(f21224,plain,(
  ( ! [X2,X0] : (sP8(X0,'T_0') | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) | $spl210),
  inference(cnf_transformation,[],[f21224_D])).
fof(f21224_D,plain,(
  ( ! [X2,X0] : (sP8(X0,'T_0') | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) <=> ~$spl210),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl210])])).
fof(f21365,plain,(
  ( ! [X37,X35,X36,X34] : (sP8(X34,'T_0') | ~p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X34) != s(bool,X37) | p(s(bool,X37))) ) | ($spl6 | $spl144)),
  inference(resolution,[],[f21134,f8878])).
fof(f8878,plain,(
  ( ! [X2,X0,X3,X1] : (~sP6(X2,X1,X3) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X0) != s(bool,X3) | p(s(bool,X0))) ) | $spl6),
  inference(subsumption_resolution,[],[f8877,f315])).
fof(f8877,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X0) != s(bool,X3) | ~sP6(X2,X1,X3) | ~sP4(X1,X3,X2)) ) | $spl6),
  inference(subsumption_resolution,[],[f8876,f309])).
fof(f8876,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X0) != s(bool,X3) | ~sP6(X2,X1,X3) | ~sP5(X3,X1,X2) | ~sP4(X1,X3,X2)) ) | $spl6),
  inference(duplicate_literal_removal,[],[f8827])).
fof(f8827,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X0) != s(bool,X3) | ~p(s(bool,X1)) | ~sP6(X2,X1,X3) | ~sP5(X3,X1,X2) | ~sP4(X1,X3,X2)) ) | $spl6),
  inference(resolution,[],[f1201,f301])).
fof(f1201,plain,(
  ( ! [X10,X8,X11,X9] : (~sP7(X11,X9,X8) | p(s(bool,X10)) | ~p(s(bool,X8)) | ~p(s(bool,X11)) | s(bool,X9) != s(bool,X10)) ) | $spl6),
  inference(superposition,[],[f528,f318])).
fof(f838,plain,(
  ( ! [X37] : (s(bool,'T_0') != s(bool,X37)) ) | $spl20),
  inference(cnf_transformation,[],[f838_D])).
fof(f838_D,plain,(
  ( ! [X37] : (s(bool,'T_0') != s(bool,X37)) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f21363,plain,(
  ( ! [X28,X29,X27] : (sP8(X27,'T_0') | ~p(s(bool,X28)) | s(bool,'F_0') != s(bool,X27) | s(bool,'T_0') != s(bool,X29)) ) | ($spl24 | $spl144)),
  inference(resolution,[],[f21134,f7489])).
fof(f21362,plain,(
  ( ! [X26,X24,X25] : (sP8(X24,'T_0') | s(bool,'F_0') != s(bool,X24) | ~p(s(bool,X25)) | ~p(s(bool,X26))) ) | ($spl24 | $spl144)),
  inference(resolution,[],[f21134,f7486])).
fof(f7486,plain,(
  ( ! [X2,X0,X1] : (~sP6(X1,X0,X2) | s(bool,'F_0') != s(bool,X2) | ~p(s(bool,X0)) | ~p(s(bool,X1))) ) | $spl24),
  inference(subsumption_resolution,[],[f7485,f402])).
fof(f7485,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X1)) | s(bool,'F_0') != s(bool,X2) | ~p(s(bool,X0)) | ~sP6(X1,X0,X2)) ) | $spl24),
  inference(subsumption_resolution,[],[f7484,f6637])).
fof(f7484,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X1)) | s(bool,'F_0') != s(bool,X2) | ~p(s(bool,X0)) | ~sP6(X1,X0,X2) | ~sP4(X0,X2,X1)) ) | $spl24),
  inference(subsumption_resolution,[],[f7445,f6638])).
fof(f7445,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X1)) | s(bool,'F_0') != s(bool,X2) | ~p(s(bool,X0)) | ~sP6(X1,X0,X2) | ~sP5(X2,X0,X1) | ~sP4(X0,X2,X1)) ) | $spl24),
  inference(resolution,[],[f1513,f301])).
fof(f21349,plain,(
  ( ! [X24,X25] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0') | p(s(bool,X24)) | ~p(s(bool,X25))) ) | $spl144),
  inference(resolution,[],[f21121,f569])).
fof(f569,plain,(
  ( ! [X68,X67] : (~sP2(X67,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X68) | p(s(bool,X68)) | ~p(s(bool,X67))) )),
  inference(resolution,[],[f251,f287])).
fof(f17627,plain,(
  ( ! [X28,X27] : (sP3(X27,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X28) | p(s(bool,X28))) ) | $spl122),
  inference(cnf_transformation,[],[f17627_D])).
fof(f17627_D,plain,(
  ( ! [X28,X27] : (sP3(X27,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X28) | p(s(bool,X28))) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f21348,plain,(
  ( ! [X23,X22] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0') | sP3(X22,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X23) | p(s(bool,X23))) ) | $spl144),
  inference(resolution,[],[f21121,f567])).
fof(f567,plain,(
  ( ! [X64,X63] : (~sP2(X63,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X64) | sP3(X63,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X64) | p(s(bool,X64))) )),
  inference(resolution,[],[f251,f286])).
fof(f286,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~sP2(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f184])).
fof(f21332,plain,(
  ( ! [X21,X20] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0') | p(s(bool,X20)) | ~p(s(bool,X21))) ) | $spl144),
  inference(resolution,[],[f21120,f570])).
fof(f570,plain,(
  ( ! [X70,X69] : (~sP2(X69,X70,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | p(s(bool,X70)) | ~p(s(bool,X69))) )),
  inference(resolution,[],[f251,f287])).
fof(f17596,plain,(
  ( ! [X21,X22] : (sP3(X21,X22,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | p(s(bool,X22))) ) | $spl120),
  inference(cnf_transformation,[],[f17596_D])).
fof(f17596_D,plain,(
  ( ! [X21,X22] : (sP3(X21,X22,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | p(s(bool,X22))) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f21331,plain,(
  ( ! [X19,X18] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0') | sP3(X18,X19,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | p(s(bool,X19))) ) | $spl144),
  inference(resolution,[],[f21120,f568])).
fof(f568,plain,(
  ( ! [X66,X65] : (~sP2(X65,X66,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | sP3(X65,X66,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | p(s(bool,X66))) )),
  inference(resolution,[],[f251,f286])).
fof(f21280,plain,(
  ( ! [X30,X28,X29] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X28)))),'T_0') | ~p(s(bool,X29)) | ~p(s(bool,X30))) ) | $spl144),
  inference(resolution,[],[f21107,f3002])).
fof(f21107,plain,(
  ( ! [X54,X55,X53] : (sP0(X54,X53,X55) | sP8(X53,'T_0')) ) | $spl144),
  inference(resolution,[],[f21077,f272])).
fof(f21279,plain,(
  ( ! [X26,X27] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0') | ~p(s(bool,X26)) | ~p(s(bool,X27))) ) | $spl144),
  inference(resolution,[],[f21107,f553])).
fof(f553,plain,(
  ( ! [X35,X36] : (~sP0(X35,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X36) | ~p(s(bool,X36)) | ~p(s(bool,X35))) )),
  inference(resolution,[],[f251,f270])).
fof(f17565,plain,(
  ( ! [X30,X31] : (sP1(X30,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X31) | ~p(s(bool,X31))) ) | $spl118),
  inference(cnf_transformation,[],[f17565_D])).
fof(f17565_D,plain,(
  ( ! [X30,X31] : (sP1(X30,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X31) | ~p(s(bool,X31))) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f21278,plain,(
  ( ! [X24,X25] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0') | sP1(X24,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X25) | ~p(s(bool,X25))) ) | $spl144),
  inference(resolution,[],[f21107,f552])).
fof(f552,plain,(
  ( ! [X33,X34] : (~sP0(X33,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X34) | sP1(X33,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X34) | ~p(s(bool,X34))) )),
  inference(resolution,[],[f251,f269])).
fof(f21257,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,X25)) | ~p(s(bool,X26)) | sP8('T_0',happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X27)))))) ) | $spl144),
  inference(resolution,[],[f3002,f20510])).
fof(f20510,plain,(
  ( ! [X54,X55,X53] : (sP0(X54,X53,X55) | sP8('T_0',X53)) ) | $spl144),
  inference(resolution,[],[f20335,f272])).
fof(f21254,plain,(
  ( ! [X19,X17,X20,X18] : (~p(s(bool,X17)) | ~p(s(bool,X18)) | sP8(happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X19)))),X20) | s(bool,'F_0') = s(bool,X20)) )),
  inference(resolution,[],[f3002,f17106])).
fof(f17106,plain,(
  ( ! [X70,X68,X71,X69] : (sP8(X68,X69) | sP0(X70,X68,X71) | s(bool,'F_0') = s(bool,X69)) )),
  inference(resolution,[],[f16182,f272])).
fof(f21253,plain,(
  ( ! [X14,X15,X13,X16] : (~p(s(bool,X13)) | ~p(s(bool,X14)) | sP8(X15,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X16))))) | s(bool,'F_0') = s(bool,X15)) )),
  inference(resolution,[],[f3002,f16869])).
fof(f16869,plain,(
  ( ! [X70,X68,X71,X69] : (sP8(X68,X69) | sP0(X70,X69,X71) | s(bool,'F_0') = s(bool,X68)) )),
  inference(resolution,[],[f16178,f272])).
fof(f21247,plain,(
  ( ! [X8,X7] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X7)))),'T_0') | p(s(bool,X8))) ) | $spl144),
  inference(resolution,[],[f21152,f3090])).
fof(f21152,plain,(
  ( ! [X180,X179] : (sP9(X179,X180) | sP8(X179,'T_0')) ) | $spl144),
  inference(resolution,[],[f21077,f344])).
fof(f21246,plain,(
  ( ! [X6] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0') | p(s(bool,X6))) ) | $spl144),
  inference(resolution,[],[f21152,f640])).
fof(f21236,plain,(
  ( ! [X6,X5] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X5)))),'T_0') | p(s(bool,X6))) ) | $spl144),
  inference(resolution,[],[f21151,f3089])).
fof(f21151,plain,(
  ( ! [X177,X178] : (sP9(X178,X177) | sP8(X177,'T_0')) ) | $spl144),
  inference(resolution,[],[f21077,f343])).
fof(f21235,plain,(
  ( ! [X4] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0') | p(s(bool,X4))) ) | $spl144),
  inference(resolution,[],[f21151,f639])).
fof(f21166,plain,(
  ( ! [X198,X200,X199] : (sP8(X198,'T_0') | p(s(bool,X199)) | p(s(bool,X200)) | s(bool,X198) != s(bool,X200)) ) | $spl144),
  inference(resolution,[],[f21077,f426])).
fof(f21099,plain,(
  ( ! [X30,X28,X31,X29] : (sP8(X28,'T_0') | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) ) | $spl144),
  inference(resolution,[],[f21077,f262])).
fof(f21096,plain,(
  ( ! [X19,X17,X18,X16] : (sP8(X16,'T_0') | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) ) | $spl144),
  inference(resolution,[],[f21077,f261])).
fof(f21095,plain,(
  ( ! [X14,X12,X15,X13] : (sP8(X12,'T_0') | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl144),
  inference(resolution,[],[f21077,f260])).
fof(f21094,plain,(
  ( ! [X10,X8,X11,X9] : (sP8(X8,'T_0') | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) ) | $spl144),
  inference(resolution,[],[f21077,f258])).
fof(f21092,plain,(
  ( ! [X2,X0,X3,X1] : (sP8(X0,'T_0') | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl144),
  inference(resolution,[],[f21077,f256])).
fof(f21056,plain,(
  ( ! [X54,X55] : (~p(s(bool,X54)) | sP8(X55,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | sP8(X55,'F_0')) ) | $spl144),
  inference(subsumption_resolution,[],[f21021,f397])).
fof(f21021,plain,(
  ( ! [X54,X55] : (~p(s(bool,X54)) | sP8(X55,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | sP8(X55,'F_0') | s(bool,'F_0') = s(bool,X54)) ) | $spl144),
  inference(superposition,[],[f19503,f14557])).
fof(f21046,plain,(
  ( ! [X304] : (sP8(X304,'T_0') | sP8(X304,'F_0')) ) | $spl206),
  inference(cnf_transformation,[],[f21046_D])).
fof(f21046_D,plain,(
  ( ! [X304] : (sP8(X304,'T_0') | sP8(X304,'F_0')) ) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f21031,plain,(
  ( ! [X47,X46] : (~p(s(bool,X46)) | sP8(X47,'T_0') | sP8(X47,'F_0')) ) | $spl144),
  inference(duplicate_literal_removal,[],[f21017])).
fof(f21017,plain,(
  ( ! [X47,X46] : (~p(s(bool,X46)) | sP8(X47,'T_0') | sP8(X47,'F_0') | ~p(s(bool,X46))) ) | $spl144),
  inference(superposition,[],[f19503,f402])).
fof(f21032,plain,(
  ( ! [X45,X44] : (~p(s(bool,X44)) | sP8(X45,'T_0') | sP8(X45,'F_0')) ) | $spl144),
  inference(duplicate_literal_removal,[],[f21016])).
fof(f21016,plain,(
  ( ! [X45,X44] : (~p(s(bool,X44)) | sP8(X45,'T_0') | sP8(X45,'F_0') | ~p(s(bool,X44))) ) | $spl144),
  inference(superposition,[],[f19503,f404])).
fof(f21055,plain,(
  ( ! [X43,X42] : (~p(s(bool,X42)) | sP8(X43,'T_0') | sP8(X43,'F_0')) ) | $spl144),
  inference(subsumption_resolution,[],[f21015,f397])).
fof(f21015,plain,(
  ( ! [X43,X42] : (~p(s(bool,X42)) | sP8(X43,'T_0') | sP8(X43,'F_0') | s(bool,'F_0') = s(bool,X42)) ) | $spl144),
  inference(superposition,[],[f19503,f424])).
fof(f19598,plain,(
  ( ! [X227,X226] : (p(s(bool,X226)) | ~sP9(X226,X227) | p(s(bool,X227))) ) | $spl158),
  inference(cnf_transformation,[],[f19598_D])).
fof(f19598_D,plain,(
  ( ! [X227,X226] : (p(s(bool,X226)) | ~sP9(X226,X227) | p(s(bool,X227))) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f21035,plain,(
  ( ! [X333,X332,X334] : (sP8(X332,'F_0') | p(s(bool,X333)) | p(s(bool,X334)) | ~sP9(X333,X334)) ) | $spl144),
  inference(duplicate_literal_removal,[],[f20991])).
fof(f20991,plain,(
  ( ! [X333,X332,X334] : (sP8(X332,'F_0') | sP8(X332,'F_0') | p(s(bool,X333)) | p(s(bool,X334)) | ~sP9(X333,X334)) ) | $spl144),
  inference(resolution,[],[f19503,f342])).
fof(f21037,plain,(
  ( ! [X327,X326,X328] : (sP8(X326,'F_0') | p(s(bool,X327)) | p(s(bool,X328)) | ~sP9(X327,X328)) ) | $spl144),
  inference(duplicate_literal_removal,[],[f20989])).
fof(f20989,plain,(
  ( ! [X327,X326,X328] : (sP8(X326,'F_0') | sP8(X326,'F_0') | p(s(bool,X327)) | p(s(bool,X328)) | ~sP9(X327,X328)) ) | $spl144),
  inference(resolution,[],[f19503,f346])).
fof(f20981,plain,(
  ( ! [X318,X317] : (sP8(X317,'T_0') | sP8(X317,'F_0') | ~p(s(bool,X318))) ) | $spl144),
  inference(resolution,[],[f19503,f395])).
fof(f395,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f239])).
fof(f239,plain,(
  ! [X0] : ((~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,'T_0')) & p(s(bool,X0)))))),
  inference(flattening,[],[f238])).
fof(f238,plain,(
  ! [X0] : (((~p(s(bool,'T_0')) | ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,'T_0')) & p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f96])).
fof(f96,plain,(
  ! [X0] : ((p(s(bool,'T_0')) & p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X13] : ((p(s(bool,'T_0')) & p(s(bool,X13))) <=> p(s(bool,X13)))),
  file('thm.quantHeuristics.GUESSES_WEAKEN_THM|split|3','thm.bool.AND_CLAUSES|split|0')).
fof(f20980,plain,(
  ( ! [X315,X316] : (sP8(X315,'T_0') | sP8(X315,'F_0') | ~p(s(bool,X316))) ) | $spl144),
  inference(resolution,[],[f19503,f393])).
fof(f393,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f237])).
fof(f237,plain,(
  ! [X0] : ((~p(s(bool,X0)) | ~p(s(bool,'T_0')) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,X0)) & p(s(bool,'T_0')))))),
  inference(flattening,[],[f236])).
fof(f236,plain,(
  ! [X0] : (((~p(s(bool,X0)) | ~p(s(bool,'T_0'))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,X0)) & p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f95])).
fof(f95,plain,(
  ! [X0] : ((p(s(bool,X0)) & p(s(bool,'T_0'))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X13] : ((p(s(bool,X13)) & p(s(bool,'T_0'))) <=> p(s(bool,X13)))),
  file('thm.quantHeuristics.GUESSES_WEAKEN_THM|split|3','thm.bool.AND_CLAUSES|split|1')).
fof(f20979,plain,(
  ( ! [X313,X314] : (sP8(X313,'T_0') | sP8(X313,'F_0') | ~p(s(bool,X314))) ) | $spl144),
  inference(resolution,[],[f19503,f380])).
fof(f380,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f229])).
fof(f229,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f228])).
fof(f228,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f144])).
fof(f144,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f91])).
fof(f91,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X13] : ((p(s(bool,'F_0')) => p(s(bool,X13))) <=> p(s(bool,'T_0')))),
  file('thm.quantHeuristics.GUESSES_WEAKEN_THM|split|3','thm.bool.IMP_CLAUSES|split|2')).
fof(f20977,plain,(
  ( ! [X310,X311] : (sP8(X310,'T_0') | sP8(X310,'F_0') | ~p(s(bool,X311))) ) | $spl144),
  inference(resolution,[],[f19503,f377])).
fof(f20976,plain,(
  ( ! [X308,X309] : (sP8(X308,'T_0') | sP8(X308,'F_0') | p(s(bool,X309))) ) | $spl144),
  inference(resolution,[],[f19503,f376])).
fof(f376,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f227])).
fof(f20975,plain,(
  ( ! [X306,X307] : (sP8(X306,'T_0') | sP8(X306,'F_0') | p(s(bool,X307))) ) | $spl144),
  inference(resolution,[],[f19503,f373])).
fof(f373,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f225])).
fof(f225,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,'T_0'))))),
  inference(flattening,[],[f224])).
fof(f224,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f142])).
fof(f142,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f89])).
fof(f89,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X13] : ((p(s(bool,X13)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  file('thm.quantHeuristics.GUESSES_WEAKEN_THM|split|3','thm.bool.IMP_CLAUSES|split|1')).
fof(f20974,plain,(
  ( ! [X304,X305] : (sP8(X304,'T_0') | sP8(X304,'F_0') | p(s(bool,X305))) ) | $spl144),
  inference(resolution,[],[f19503,f370])).
fof(f20968,plain,(
  ( ! [X285,X287,X284,X286] : (sP8(X284,X285) | sP8(X284,'F_0') | p(s(bool,X286)) | ~p(s(bool,X287)) | s(bool,X285) != s(bool,X287)) ) | $spl144),
  inference(resolution,[],[f19503,f426])).
fof(f20897,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP8(X35,X36) | sP8(X35,'F_0') | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X38) | s(bool,X37) != s(bool,X39)) ) | $spl144),
  inference(resolution,[],[f19503,f262])).
fof(f20895,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP8(X25,X26) | sP8(X25,'F_0') | p(s(bool,X27)) | ~p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X26) != s(bool,X28)) ) | $spl144),
  inference(resolution,[],[f19503,f261])).
fof(f20893,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP8(X15,X16) | sP8(X15,'F_0') | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) ) | $spl144),
  inference(resolution,[],[f19503,f260])).
fof(f20892,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP8(X10,X11) | sP8(X10,'F_0') | ~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,X14)) ) | $spl144),
  inference(resolution,[],[f19503,f258])).
fof(f20890,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP8(X0,X1) | sP8(X0,'F_0') | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) ) | $spl144),
  inference(resolution,[],[f19503,f256])).
fof(f20884,plain,(
  ( ! [X33,X32] : (sP8('T_0',happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X32))))) | sP8(X33,'F_0')) ) | $spl144),
  inference(subsumption_resolution,[],[f20853,f19403])).
fof(f20853,plain,(
  ( ! [X33,X32] : (sP8('T_0',happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X32))))) | sP8(X33,'F_0') | p(s(bool,X33))) ) | $spl144),
  inference(resolution,[],[f20573,f3043])).
fof(f20573,plain,(
  ( ! [X208,X209] : (sP8('T_0',X208) | sP8(X208,X209) | sP8(X209,'F_0')) ) | $spl144),
  inference(resolution,[],[f20335,f19502])).
fof(f20882,plain,(
  ( ! [X29] : (sP8('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | sP8(X29,'F_0')) ) | $spl144),
  inference(subsumption_resolution,[],[f20850,f19403])).
fof(f20850,plain,(
  ( ! [X29] : (sP8('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | sP8(X29,'F_0') | p(s(bool,X29))) ) | $spl144),
  inference(resolution,[],[f20573,f594])).
fof(f20881,plain,(
  ( ! [X25] : (sP8('T_0',X25) | sP8(happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),'F_0')) ) | $spl144),
  inference(subsumption_resolution,[],[f20846,f20335])).
fof(f20846,plain,(
  ( ! [X25] : (sP8('T_0',X25) | sP8(happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),'F_0') | ~p(s(bool,X25))) ) | $spl144),
  inference(resolution,[],[f20573,f2623])).
fof(f20880,plain,(
  ( ! [X24] : (sP8('T_0',X24) | sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),'F_0')) ) | $spl144),
  inference(subsumption_resolution,[],[f20845,f20335])).
fof(f20845,plain,(
  ( ! [X24] : (sP8('T_0',X24) | sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),'F_0') | ~p(s(bool,X24))) ) | $spl144),
  inference(resolution,[],[f20573,f1301])).
fof(f20878,plain,(
  ( ! [X22] : (sP8('T_0',X22) | sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0')) ) | $spl144),
  inference(subsumption_resolution,[],[f20843,f20335])).
fof(f20843,plain,(
  ( ! [X22] : (sP8('T_0',X22) | sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | ~p(s(bool,X22))) ) | $spl144),
  inference(resolution,[],[f20573,f484])).
fof(f20814,plain,(
  ( ! [X41,X42] : (sP8('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | p(s(bool,X41)) | p(s(bool,X42))) ) | $spl144),
  inference(resolution,[],[f20543,f589])).
fof(f20736,plain,(
  ( ! [X24] : (sP8('T_0',X24) | s(bool,'F_0') != s(bool,X24)) ) | $spl194),
  inference(cnf_transformation,[],[f20736_D])).
fof(f20736_D,plain,(
  ( ! [X24] : (sP8('T_0',X24) | s(bool,'F_0') != s(bool,X24)) ) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f20810,plain,(
  ( ! [X35,X34] : (sP8('T_0',X34) | p(s(bool,X35)) | s(bool,'F_0') != s(bool,X34)) ) | $spl144),
  inference(resolution,[],[f20543,f1971])).
fof(f20794,plain,(
  ( ! [X39,X40] : (sP8('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | p(s(bool,X39)) | p(s(bool,X40))) ) | $spl144),
  inference(resolution,[],[f20542,f588])).
fof(f20774,plain,(
  ( ! [X38] : (sP8('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | p(s(bool,X38))) ) | $spl144),
  inference(resolution,[],[f20541,f3375])).
fof(f20771,plain,(
  ( ! [X33,X32] : (sP8('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | p(s(bool,X32)) | p(s(bool,X33))) ) | $spl144),
  inference(resolution,[],[f20541,f590])).
fof(f20770,plain,(
  ( ! [X30,X31] : (sP8('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | p(s(bool,X30)) | s(bool,X30) != s(bool,X31)) ) | $spl144),
  inference(resolution,[],[f20541,f3371])).
fof(f20751,plain,(
  ( ! [X35,X34] : (sP8('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | ~p(s(bool,X34)) | ~p(s(bool,X35))) ) | $spl144),
  inference(resolution,[],[f20540,f585])).
fof(f20750,plain,(
  ( ! [X33,X32] : (sP8('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | sP7('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X32,X33) | ~sP4(X33,X32,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) ) | $spl144),
  inference(resolution,[],[f20540,f7261])).
fof(f20749,plain,(
  ( ! [X30,X31] : (sP8('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | ~p(s(bool,X30)) | s(bool,X30) != s(bool,X31)) ) | $spl144),
  inference(resolution,[],[f20540,f3369])).
fof(f20738,plain,(
  ( ! [X75,X73] : (sP8('T_0',X73) | s(bool,X73) != s(bool,X75) | s(bool,'T_0') = s(bool,X75)) ) | $spl198),
  inference(cnf_transformation,[],[f20738_D])).
fof(f20738_D,plain,(
  ( ! [X75,X73] : (sP8('T_0',X73) | s(bool,X73) != s(bool,X75) | s(bool,'T_0') = s(bool,X75)) ) <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f20722,plain,(
  ( ! [X74,X75,X73] : (sP8('T_0',X73) | ~p(s(bool,X74)) | s(bool,'T_0') = s(bool,X75) | s(bool,X73) != s(bool,X75)) ) | $spl144),
  inference(resolution,[],[f20537,f14828])).
fof(f20720,plain,(
  ( ! [X68,X69] : (sP8('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | ~p(s(bool,X68)) | ~p(s(bool,X69))) ) | $spl144),
  inference(resolution,[],[f20537,f582])).
fof(f20719,plain,(
  ( ! [X66,X67] : (sP8('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | sP7(X66,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X67) | ~sP4(X67,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X66)) ) | $spl144),
  inference(resolution,[],[f20537,f6889])).
fof(f20737,plain,(
  ( ! [X44] : (sP8('T_0',X44) | s(bool,'T_0') = s(bool,X44)) ) | $spl196),
  inference(cnf_transformation,[],[f20737_D])).
fof(f20737_D,plain,(
  ( ! [X44] : (sP8('T_0',X44) | s(bool,'T_0') = s(bool,X44)) ) <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f20711,plain,(
  ( ! [X47,X48,X49] : (sP8('T_0',X47) | s(bool,'T_0') = s(bool,X47) | s(bool,'F_0') = s(bool,X48) | ~p(s(bool,X49))) ) | $spl144),
  inference(resolution,[],[f20537,f14390])).
fof(f20710,plain,(
  ( ! [X45,X46,X44] : (sP8('T_0',X44) | s(bool,'T_0') = s(bool,X44) | ~p(s(bool,X45)) | ~p(s(bool,X46))) ) | $spl144),
  inference(resolution,[],[f20537,f14047])).
fof(f20626,plain,(
  ( ! [X2,X0] : (sP8('T_0',X0) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) | $spl188),
  inference(cnf_transformation,[],[f20626_D])).
fof(f20626_D,plain,(
  ( ! [X2,X0] : (sP8('T_0',X0) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f20707,plain,(
  ( ! [X37,X35,X36,X34] : (sP8('T_0',X34) | ~p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X34) != s(bool,X37) | p(s(bool,X37))) ) | ($spl6 | $spl144)),
  inference(resolution,[],[f20537,f8878])).
fof(f20705,plain,(
  ( ! [X28,X29,X27] : (sP8('T_0',X27) | ~p(s(bool,X28)) | s(bool,'F_0') != s(bool,X27) | s(bool,'T_0') != s(bool,X29)) ) | ($spl24 | $spl144)),
  inference(resolution,[],[f20537,f7489])).
fof(f20704,plain,(
  ( ! [X26,X24,X25] : (sP8('T_0',X24) | s(bool,'F_0') != s(bool,X24) | ~p(s(bool,X25)) | ~p(s(bool,X26))) ) | ($spl24 | $spl144)),
  inference(resolution,[],[f20537,f7486])).
fof(f20691,plain,(
  ( ! [X24,X25] : (sP8('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | p(s(bool,X24)) | ~p(s(bool,X25))) ) | $spl144),
  inference(resolution,[],[f20524,f569])).
fof(f20690,plain,(
  ( ! [X23,X22] : (sP8('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | sP3(X22,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X23) | p(s(bool,X23))) ) | $spl144),
  inference(resolution,[],[f20524,f567])).
fof(f20674,plain,(
  ( ! [X21,X20] : (sP8('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | p(s(bool,X20)) | ~p(s(bool,X21))) ) | $spl144),
  inference(resolution,[],[f20523,f570])).
fof(f20673,plain,(
  ( ! [X19,X18] : (sP8('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | sP3(X18,X19,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | p(s(bool,X19))) ) | $spl144),
  inference(resolution,[],[f20523,f568])).
fof(f20662,plain,(
  ( ! [X26,X27] : (sP8('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | ~p(s(bool,X26)) | ~p(s(bool,X27))) ) | $spl144),
  inference(resolution,[],[f20510,f553])).
fof(f20661,plain,(
  ( ! [X24,X25] : (sP8('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | sP1(X24,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X25) | ~p(s(bool,X25))) ) | $spl144),
  inference(resolution,[],[f20510,f552])).
fof(f20650,plain,(
  ( ! [X8,X7] : (sP8('T_0',happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X7))))) | p(s(bool,X8))) ) | $spl144),
  inference(resolution,[],[f20555,f3090])).
fof(f20555,plain,(
  ( ! [X180,X179] : (sP9(X179,X180) | sP8('T_0',X179)) ) | $spl144),
  inference(resolution,[],[f20335,f344])).
fof(f20649,plain,(
  ( ! [X6] : (sP8('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | p(s(bool,X6))) ) | $spl144),
  inference(resolution,[],[f20555,f640])).
fof(f20638,plain,(
  ( ! [X6,X5] : (sP8('T_0',happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X5))))) | p(s(bool,X6))) ) | $spl144),
  inference(resolution,[],[f20554,f3089])).
fof(f20554,plain,(
  ( ! [X177,X178] : (sP9(X178,X177) | sP8('T_0',X177)) ) | $spl144),
  inference(resolution,[],[f20335,f343])).
fof(f20637,plain,(
  ( ! [X4] : (sP8('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | p(s(bool,X4))) ) | $spl144),
  inference(resolution,[],[f20554,f639])).
fof(f20569,plain,(
  ( ! [X198,X200,X199] : (sP8('T_0',X198) | p(s(bool,X199)) | p(s(bool,X200)) | s(bool,X198) != s(bool,X200)) ) | $spl144),
  inference(resolution,[],[f20335,f426])).
fof(f20502,plain,(
  ( ! [X30,X28,X31,X29] : (sP8('T_0',X28) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) ) | $spl144),
  inference(resolution,[],[f20335,f262])).
fof(f20499,plain,(
  ( ! [X19,X17,X18,X16] : (sP8('T_0',X16) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) ) | $spl144),
  inference(resolution,[],[f20335,f261])).
fof(f20498,plain,(
  ( ! [X14,X12,X15,X13] : (sP8('T_0',X12) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl144),
  inference(resolution,[],[f20335,f260])).
fof(f20497,plain,(
  ( ! [X10,X8,X11,X9] : (sP8('T_0',X8) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) ) | $spl144),
  inference(resolution,[],[f20335,f258])).
fof(f20495,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('T_0',X0) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl144),
  inference(resolution,[],[f20335,f256])).
fof(f20355,plain,(
  $false | $spl160),
  inference(subsumption_resolution,[],[f20350,f412])).
fof(f20350,plain,(
  p(s(bool,'F_0')) | $spl160),
  inference(duplicate_literal_removal,[],[f20343])).
fof(f20343,plain,(
  p(s(bool,'F_0')) | p(s(bool,'F_0')) | $spl160),
  inference(resolution,[],[f19600,f335])).
fof(f335,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | ~sP8(X1,X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f20354,plain,(
  $false | $spl160),
  inference(subsumption_resolution,[],[f20353,f412])).
fof(f20353,plain,(
  p(s(bool,'F_0')) | $spl160),
  inference(duplicate_literal_removal,[],[f20340])).
fof(f20340,plain,(
  p(s(bool,'F_0')) | p(s(bool,'F_0')) | $spl160),
  inference(resolution,[],[f19600,f326])).
fof(f19600,plain,(
  sP8('F_0','F_0') | $spl160),
  inference(cnf_transformation,[],[f19600_D])).
fof(f19600_D,plain,(
  sP8('F_0','F_0') <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f20337,plain,(
  ( ! [X10] : (sP8('T_0',X10) | p(s(bool,X10))) ) | $spl186),
  inference(cnf_transformation,[],[f20337_D])).
fof(f20337_D,plain,(
  ( ! [X10] : (sP8('T_0',X10) | p(s(bool,X10))) ) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f20324,plain,(
  ( ! [X10] : (sP8('T_0',X10) | p(s(bool,X10)) | sP8('F_0','F_0')) ) | $spl144),
  inference(resolution,[],[f20090,f19500])).
fof(f20151,plain,(
  ( ! [X54,X55] : (~p(s(bool,X54)) | sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X55) | sP8(X55,'F_0')) ) | $spl144),
  inference(subsumption_resolution,[],[f20127,f397])).
fof(f20127,plain,(
  ( ! [X54,X55] : (~p(s(bool,X54)) | sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X55) | sP8(X55,'F_0') | s(bool,'F_0') = s(bool,X54)) ) | $spl144),
  inference(superposition,[],[f19502,f14557])).
fof(f20143,plain,(
  ( ! [X302] : (sP8('T_0',X302) | sP8(X302,'F_0')) ) | $spl180),
  inference(cnf_transformation,[],[f20143_D])).
fof(f20143_D,plain,(
  ( ! [X302] : (sP8('T_0',X302) | sP8(X302,'F_0')) ) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f20134,plain,(
  ( ! [X47,X46] : (~p(s(bool,X46)) | sP8('T_0',X47) | sP8(X47,'F_0')) ) | $spl144),
  inference(duplicate_literal_removal,[],[f20123])).
fof(f20123,plain,(
  ( ! [X47,X46] : (~p(s(bool,X46)) | sP8('T_0',X47) | sP8(X47,'F_0') | ~p(s(bool,X46))) ) | $spl144),
  inference(superposition,[],[f19502,f402])).
fof(f20135,plain,(
  ( ! [X45,X44] : (~p(s(bool,X44)) | sP8('T_0',X45) | sP8(X45,'F_0')) ) | $spl144),
  inference(duplicate_literal_removal,[],[f20122])).
fof(f20122,plain,(
  ( ! [X45,X44] : (~p(s(bool,X44)) | sP8('T_0',X45) | sP8(X45,'F_0') | ~p(s(bool,X44))) ) | $spl144),
  inference(superposition,[],[f19502,f404])).
fof(f20150,plain,(
  ( ! [X43,X42] : (~p(s(bool,X42)) | sP8('T_0',X43) | sP8(X43,'F_0')) ) | $spl144),
  inference(subsumption_resolution,[],[f20121,f397])).
fof(f20121,plain,(
  ( ! [X43,X42] : (~p(s(bool,X42)) | sP8('T_0',X43) | sP8(X43,'F_0') | s(bool,'F_0') = s(bool,X42)) ) | $spl144),
  inference(superposition,[],[f19502,f424])).
fof(f20147,plain,(
  ( ! [X324] : (sP8('F_0',X324) | sP8(X324,'F_0')) ) | $spl182),
  inference(cnf_transformation,[],[f20147_D])).
fof(f20147_D,plain,(
  ( ! [X324] : (sP8('F_0',X324) | sP8(X324,'F_0')) ) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f20097,plain,(
  ( ! [X331,X332,X330] : (sP8('F_0',X330) | sP8(X330,'F_0') | p(s(bool,X331)) | p(s(bool,X332)) | ~sP9(X331,X332)) ) | $spl144),
  inference(resolution,[],[f19502,f342])).
fof(f20095,plain,(
  ( ! [X325,X324,X326] : (sP8('F_0',X324) | sP8(X324,'F_0') | p(s(bool,X325)) | p(s(bool,X326)) | ~sP9(X325,X326)) ) | $spl144),
  inference(resolution,[],[f19502,f346])).
fof(f20087,plain,(
  ( ! [X315,X316] : (sP8('T_0',X315) | sP8(X315,'F_0') | ~p(s(bool,X316))) ) | $spl144),
  inference(resolution,[],[f19502,f395])).
fof(f20086,plain,(
  ( ! [X313,X314] : (sP8('T_0',X313) | sP8(X313,'F_0') | ~p(s(bool,X314))) ) | $spl144),
  inference(resolution,[],[f19502,f393])).
fof(f20085,plain,(
  ( ! [X312,X311] : (sP8('T_0',X311) | sP8(X311,'F_0') | ~p(s(bool,X312))) ) | $spl144),
  inference(resolution,[],[f19502,f380])).
fof(f20083,plain,(
  ( ! [X308,X309] : (sP8('T_0',X308) | sP8(X308,'F_0') | ~p(s(bool,X309))) ) | $spl144),
  inference(resolution,[],[f19502,f377])).
fof(f20082,plain,(
  ( ! [X306,X307] : (sP8('T_0',X306) | sP8(X306,'F_0') | p(s(bool,X307))) ) | $spl144),
  inference(resolution,[],[f19502,f376])).
fof(f20081,plain,(
  ( ! [X304,X305] : (sP8('T_0',X304) | sP8(X304,'F_0') | p(s(bool,X305))) ) | $spl144),
  inference(resolution,[],[f19502,f373])).
fof(f20080,plain,(
  ( ! [X302,X303] : (sP8('T_0',X302) | sP8(X302,'F_0') | p(s(bool,X303))) ) | $spl144),
  inference(resolution,[],[f19502,f370])).
fof(f20075,plain,(
  ( ! [X285,X287,X284,X286] : (sP8(X284,X285) | sP8(X285,'F_0') | p(s(bool,X286)) | ~p(s(bool,X287)) | s(bool,X284) != s(bool,X287)) ) | $spl144),
  inference(resolution,[],[f19502,f426])).
fof(f20004,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP8(X35,X36) | sP8(X36,'F_0') | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X35) != s(bool,X38) | s(bool,X37) != s(bool,X39)) ) | $spl144),
  inference(resolution,[],[f19502,f262])).
fof(f20002,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP8(X25,X26) | sP8(X26,'F_0') | p(s(bool,X27)) | ~p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X25) != s(bool,X28)) ) | $spl144),
  inference(resolution,[],[f19502,f261])).
fof(f20000,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP8(X15,X16) | sP8(X16,'F_0') | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X15) != s(bool,X18)) ) | $spl144),
  inference(resolution,[],[f19502,f260])).
fof(f19999,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP8(X10,X11) | sP8(X11,'F_0') | ~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X10) != s(bool,X13) | s(bool,X12) != s(bool,X14)) ) | $spl144),
  inference(resolution,[],[f19502,f258])).
fof(f19997,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP8(X0,X1) | sP8(X1,'F_0') | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X0) != s(bool,X3)) ) | $spl144),
  inference(resolution,[],[f19502,f256])).
fof(f19983,plain,(
  ( ! [X109] : (p(s(bool,X109)) | sP8(happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),'F_0')) ) | $spl144),
  inference(duplicate_literal_removal,[],[f19981])).
fof(f19981,plain,(
  ( ! [X109] : (p(s(bool,X109)) | sP8(happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),'F_0') | p(s(bool,X109))) ) | $spl144),
  inference(resolution,[],[f19501,f2624])).
fof(f2624,plain,(
  ( ! [X116] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),X116) | p(s(bool,X116))) )),
  inference(resolution,[],[f1246,f328])).
fof(f19984,plain,(
  ( ! [X108] : (p(s(bool,X108)) | sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),'F_0')) ) | $spl144),
  inference(duplicate_literal_removal,[],[f19980])).
fof(f19980,plain,(
  ( ! [X108] : (p(s(bool,X108)) | sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),'F_0') | p(s(bool,X108))) ) | $spl144),
  inference(resolution,[],[f19501,f1302])).
fof(f1302,plain,(
  ( ! [X116] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),X116) | p(s(bool,X116))) )),
  inference(resolution,[],[f538,f328])).
fof(f19985,plain,(
  ( ! [X106] : (p(s(bool,X106)) | sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0')) ) | $spl144),
  inference(duplicate_literal_removal,[],[f19978])).
fof(f19978,plain,(
  ( ! [X106] : (p(s(bool,X106)) | sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | p(s(bool,X106))) ) | $spl144),
  inference(resolution,[],[f19501,f485])).
fof(f485,plain,(
  ( ! [X113] : (sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X113) | p(s(bool,X113))) )),
  inference(resolution,[],[f250,f328])).
fof(f19927,plain,(
  ( ! [X102] : (p(s(bool,X102)) | sP8(happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),'F_0')) ) | $spl144),
  inference(duplicate_literal_removal,[],[f19918])).
fof(f19918,plain,(
  ( ! [X102] : (p(s(bool,X102)) | sP8(happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),'F_0') | p(s(bool,X102))) ) | $spl144),
  inference(resolution,[],[f19500,f2625])).
fof(f2625,plain,(
  ( ! [X117] : (sP8(X117,happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13))))) | p(s(bool,X117))) )),
  inference(resolution,[],[f1246,f331])).
fof(f19928,plain,(
  ( ! [X101] : (p(s(bool,X101)) | sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),'F_0')) ) | $spl144),
  inference(duplicate_literal_removal,[],[f19917])).
fof(f19917,plain,(
  ( ! [X101] : (p(s(bool,X101)) | sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),'F_0') | p(s(bool,X101))) ) | $spl144),
  inference(resolution,[],[f19500,f1303])).
fof(f1303,plain,(
  ( ! [X117] : (sP8(X117,happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11)))) | p(s(bool,X117))) )),
  inference(resolution,[],[f538,f331])).
fof(f19929,plain,(
  ( ! [X99] : (p(s(bool,X99)) | sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0')) ) | $spl144),
  inference(duplicate_literal_removal,[],[f19915])).
fof(f19915,plain,(
  ( ! [X99] : (p(s(bool,X99)) | sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | p(s(bool,X99))) ) | $spl144),
  inference(resolution,[],[f19500,f486])).
fof(f486,plain,(
  ( ! [X114] : (sP8(X114,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | p(s(bool,X114))) )),
  inference(resolution,[],[f250,f331])).
fof(f19881,plain,(
  ( ! [X47,X46] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),'F_0') | p(s(bool,X46)) | ~p(s(bool,X47))) ) | $spl144),
  inference(resolution,[],[f19494,f1292])).
fof(f19494,plain,(
  ( ! [X151,X149,X150] : (sP5(X149,X150,X151) | sP8(X149,'F_0')) ) | $spl144),
  inference(resolution,[],[f19403,f311])).
fof(f19880,plain,(
  ( ! [X45,X44] : (sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | p(s(bool,X44)) | ~p(s(bool,X45))) ) | $spl144),
  inference(resolution,[],[f19494,f475])).
fof(f19858,plain,(
  ( ! [X41,X40] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),'F_0') | p(s(bool,X40)) | ~p(s(bool,X41))) ) | $spl144),
  inference(resolution,[],[f19493,f1293])).
fof(f19493,plain,(
  ( ! [X146,X147,X148] : (sP5(X147,X146,X148) | sP8(X146,'F_0')) ) | $spl144),
  inference(resolution,[],[f19403,f310])).
fof(f19857,plain,(
  ( ! [X39,X38] : (sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | p(s(bool,X38)) | ~p(s(bool,X39))) ) | $spl144),
  inference(resolution,[],[f19493,f476])).
fof(f17404,plain,(
  ( ! [X6] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X6) | s(bool,'F_0') = s(bool,X6)) ) | $spl114),
  inference(cnf_transformation,[],[f17404_D])).
fof(f17404_D,plain,(
  ( ! [X6] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X6) | s(bool,'F_0') = s(bool,X6)) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f19818,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X33)) | s(bool,X33) != s(bool,X34) | sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X35) | s(bool,'F_0') = s(bool,X35)) )),
  inference(resolution,[],[f3371,f17137])).
fof(f17048,plain,(
  ( ! [X6] : (sP8(X6,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | s(bool,'F_0') = s(bool,X6)) ) | $spl106),
  inference(cnf_transformation,[],[f17048_D])).
fof(f17048_D,plain,(
  ( ! [X6] : (sP8(X6,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | s(bool,'F_0') = s(bool,X6)) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f19815,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | s(bool,X24) != s(bool,X25) | sP8(X26,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | s(bool,'F_0') = s(bool,X26)) )),
  inference(resolution,[],[f3371,f16900])).
fof(f19797,plain,(
  ( ! [X97] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),'F_0') | ~p(s(bool,X97))) ) | $spl144),
  inference(resolution,[],[f19491,f12547])).
fof(f19491,plain,(
  ( ! [X142,X140,X141] : (sP6(X140,X141,X142) | sP8(X140,'F_0')) ) | $spl144),
  inference(resolution,[],[f19403,f307])).
fof(f19796,plain,(
  ( ! [X95,X96] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),'F_0') | p(s(bool,X95)) | ~p(s(bool,X96))) ) | $spl144),
  inference(resolution,[],[f19491,f1289])).
fof(f19795,plain,(
  ( ! [X94] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),'F_0') | ~p(s(bool,X94))) ) | $spl144),
  inference(resolution,[],[f19491,f9596])).
fof(f19794,plain,(
  ( ! [X92,X93] : (sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | p(s(bool,X92)) | ~p(s(bool,X93))) ) | $spl144),
  inference(resolution,[],[f19491,f472])).
fof(f19793,plain,(
  ( ! [X90,X91] : (sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | ~p(s(bool,X90)) | s(bool,X90) = s(bool,X91)) ) | $spl144),
  inference(resolution,[],[f19491,f3259])).
fof(f19792,plain,(
  ( ! [X88,X89] : (sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | p(s(bool,X88)) | s(bool,X88) = s(bool,X89)) ) | $spl144),
  inference(resolution,[],[f19491,f3261])).
fof(f19791,plain,(
  ( ! [X87] : (sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | ~p(s(bool,X87))) ) | $spl144),
  inference(resolution,[],[f19491,f3271])).
fof(f19747,plain,(
  ( ! [X85,X86] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),'F_0') | p(s(bool,X85)) | ~p(s(bool,X86))) ) | $spl144),
  inference(resolution,[],[f19490,f1290])).
fof(f19490,plain,(
  ( ! [X138,X139,X137] : (sP6(X138,X137,X139) | sP8(X137,'F_0')) ) | $spl144),
  inference(resolution,[],[f19403,f306])).
fof(f19746,plain,(
  ( ! [X83,X84] : (sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | p(s(bool,X83)) | ~p(s(bool,X84))) ) | $spl144),
  inference(resolution,[],[f19490,f473])).
fof(f19745,plain,(
  ( ! [X81,X82] : (sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | sP7(X81,X82,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | ~sP5(X82,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X81)) ) | $spl144),
  inference(resolution,[],[f19490,f6665])).
fof(f19757,plain,(
  ( ! [X65,X63] : (sP8(X63,'F_0') | s(bool,'T_0') != s(bool,X65) | s(bool,X63) != s(bool,X65)) ) | $spl170),
  inference(cnf_transformation,[],[f19757_D])).
fof(f19757_D,plain,(
  ( ! [X65,X63] : (sP8(X63,'F_0') | s(bool,'T_0') != s(bool,X65) | s(bool,X63) != s(bool,X65)) ) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f19738,plain,(
  ( ! [X64,X65,X63] : (sP8(X63,'F_0') | ~p(s(bool,X64)) | s(bool,X63) != s(bool,X65) | s(bool,'T_0') != s(bool,X65)) ) | ($spl24 | $spl144)),
  inference(resolution,[],[f19490,f6588])).
fof(f19756,plain,(
  ( ! [X27] : (sP8(X27,'F_0') | s(bool,'T_0') != s(bool,X27)) ) | $spl168),
  inference(cnf_transformation,[],[f19756_D])).
fof(f19756_D,plain,(
  ( ! [X27] : (sP8(X27,'F_0') | s(bool,'T_0') != s(bool,X27)) ) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f19736,plain,(
  ( ! [X59,X58] : (sP8(X58,'F_0') | s(bool,'T_0') != s(bool,X58) | ~p(s(bool,X59))) ) | $spl144),
  inference(resolution,[],[f19490,f2088])).
fof(f19588,plain,(
  ( ! [X2,X0] : (sP8(X0,'F_0') | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) | $spl154),
  inference(cnf_transformation,[],[f19588_D])).
fof(f19588_D,plain,(
  ( ! [X2,X0] : (sP8(X0,'F_0') | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f19726,plain,(
  ( ! [X30,X33,X31,X32] : (sP8(X30,'F_0') | ~p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X30) != s(bool,X31) | p(s(bool,X33))) ) | ($spl6 | $spl144)),
  inference(resolution,[],[f19490,f8286])).
fof(f19725,plain,(
  ( ! [X28,X29,X27] : (sP8(X27,'F_0') | ~p(s(bool,X28)) | s(bool,'F_0') != s(bool,X29) | s(bool,'T_0') != s(bool,X27)) ) | ($spl24 | $spl144)),
  inference(resolution,[],[f19490,f7489])).
fof(f19714,plain,(
  ( ! [X30,X31] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),'F_0') | p(s(bool,X30)) | p(s(bool,X31))) ) | $spl144),
  inference(resolution,[],[f19478,f1276])).
fof(f19478,plain,(
  ( ! [X103,X101,X102] : (sP2(X101,X102,X103) | sP8(X101,'F_0')) ) | $spl144),
  inference(resolution,[],[f19403,f290])).
fof(f19713,plain,(
  ( ! [X28,X29] : (sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | p(s(bool,X28)) | p(s(bool,X29))) ) | $spl144),
  inference(resolution,[],[f19478,f459])).
fof(f19712,plain,(
  ( ! [X26,X27] : (sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | sP3('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X26,X27)) ) | $spl144),
  inference(resolution,[],[f19478,f508])).
fof(f19700,plain,(
  ( ! [X33,X32] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),'F_0') | ~p(s(bool,X32)) | p(s(bool,X33))) ) | $spl144),
  inference(resolution,[],[f19462,f1262])).
fof(f19462,plain,(
  ( ! [X54,X55,X53] : (sP0(X53,X54,X55) | sP8(X53,'F_0')) ) | $spl144),
  inference(resolution,[],[f19403,f273])).
fof(f19699,plain,(
  ( ! [X30,X31] : (sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | ~p(s(bool,X30)) | p(s(bool,X31))) ) | $spl144),
  inference(resolution,[],[f19462,f445])).
fof(f19698,plain,(
  ( ! [X28,X29] : (sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | sP1('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X28,X29)) ) | $spl144),
  inference(resolution,[],[f19462,f507])).
fof(f19679,plain,(
  ( ! [X23,X22] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),'F_0') | p(s(bool,X22)) | ~p(s(bool,X23))) ) | $spl144),
  inference(resolution,[],[f19461,f1263])).
fof(f19461,plain,(
  ( ! [X52,X50,X51] : (sP0(X51,X52,X50) | sP8(X50,'F_0')) ) | $spl144),
  inference(resolution,[],[f19403,f271])).
fof(f19678,plain,(
  ( ! [X21,X20] : (sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | p(s(bool,X20)) | ~p(s(bool,X21))) ) | $spl144),
  inference(resolution,[],[f19461,f446])).
fof(f19677,plain,(
  ( ! [X19,X18] : (sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | sP1(X18,X19,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | p(s(bool,X19))) ) | $spl144),
  inference(resolution,[],[f19461,f444])).
fof(f19649,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,X27)) | s(bool,X27) != s(bool,X28) | sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X29) | s(bool,'F_0') = s(bool,X29)) )),
  inference(resolution,[],[f3369,f17136])).
fof(f19648,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | s(bool,X24) != s(bool,X25) | sP8(X26,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | s(bool,'F_0') = s(bool,X26)) )),
  inference(resolution,[],[f3369,f16899])).
fof(f19604,plain,(
  ( ! [X33] : (~p(s(bool,X33)) | sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0')) ) | $spl144),
  inference(subsumption_resolution,[],[f19577,f397])).
fof(f19577,plain,(
  ( ! [X33] : (~p(s(bool,X33)) | sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | s(bool,'F_0') = s(bool,X33)) ) | $spl144),
  inference(superposition,[],[f19403,f14557])).
fof(f19594,plain,(
  sP8('T_0','F_0') | $spl156),
  inference(cnf_transformation,[],[f19594_D])).
fof(f19594_D,plain,(
  sP8('T_0','F_0') <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f19584,plain,(
  ( ! [X29] : (~p(s(bool,X29)) | sP8('T_0','F_0')) ) | $spl144),
  inference(duplicate_literal_removal,[],[f19573])).
fof(f19573,plain,(
  ( ! [X29] : (~p(s(bool,X29)) | sP8('T_0','F_0') | ~p(s(bool,X29))) ) | $spl144),
  inference(superposition,[],[f19403,f402])).
fof(f19585,plain,(
  ( ! [X28] : (~p(s(bool,X28)) | sP8('T_0','F_0')) ) | $spl144),
  inference(duplicate_literal_removal,[],[f19572])).
fof(f19572,plain,(
  ( ! [X28] : (~p(s(bool,X28)) | sP8('T_0','F_0') | ~p(s(bool,X28))) ) | $spl144),
  inference(superposition,[],[f19403,f404])).
fof(f19603,plain,(
  ( ! [X27] : (~p(s(bool,X27)) | sP8('T_0','F_0')) ) | $spl144),
  inference(subsumption_resolution,[],[f19571,f397])).
fof(f19571,plain,(
  ( ! [X27] : (~p(s(bool,X27)) | sP8('T_0','F_0') | s(bool,'F_0') = s(bool,X27)) ) | $spl144),
  inference(superposition,[],[f19403,f424])).
fof(f19547,plain,(
  ( ! [X231,X230] : (sP8('F_0','F_0') | p(s(bool,X230)) | p(s(bool,X231)) | ~sP9(X230,X231)) ) | $spl144),
  inference(resolution,[],[f19403,f342])).
fof(f19545,plain,(
  ( ! [X227,X226] : (sP8('F_0','F_0') | p(s(bool,X226)) | p(s(bool,X227)) | ~sP9(X226,X227)) ) | $spl144),
  inference(resolution,[],[f19403,f346])).
fof(f19537,plain,(
  ( ! [X225] : (sP8('T_0','F_0') | ~p(s(bool,X225))) ) | $spl144),
  inference(resolution,[],[f19403,f395])).
fof(f19536,plain,(
  ( ! [X224] : (sP8('T_0','F_0') | ~p(s(bool,X224))) ) | $spl144),
  inference(resolution,[],[f19403,f393])).
fof(f19535,plain,(
  ( ! [X223] : (sP8('T_0','F_0') | ~p(s(bool,X223))) ) | $spl144),
  inference(resolution,[],[f19403,f380])).
fof(f19533,plain,(
  ( ! [X222] : (sP8('T_0','F_0') | ~p(s(bool,X222))) ) | $spl144),
  inference(resolution,[],[f19403,f377])).
fof(f19532,plain,(
  ( ! [X221] : (sP8('T_0','F_0') | p(s(bool,X221))) ) | $spl144),
  inference(resolution,[],[f19403,f376])).
fof(f19531,plain,(
  ( ! [X220] : (sP8('T_0','F_0') | p(s(bool,X220))) ) | $spl144),
  inference(resolution,[],[f19403,f373])).
fof(f19530,plain,(
  ( ! [X219] : (sP8('T_0','F_0') | p(s(bool,X219))) ) | $spl144),
  inference(resolution,[],[f19403,f370])).
fof(f19525,plain,(
  ( ! [X208,X206,X207] : (sP8(X206,'F_0') | p(s(bool,X207)) | ~p(s(bool,X208)) | s(bool,X206) != s(bool,X208)) ) | $spl144),
  inference(resolution,[],[f19403,f426])).
fof(f19454,plain,(
  ( ! [X30,X28,X31,X29] : (sP8(X28,'F_0') | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) ) | $spl144),
  inference(resolution,[],[f19403,f262])).
fof(f19452,plain,(
  ( ! [X23,X21,X22,X20] : (sP8(X20,'F_0') | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X20) != s(bool,X22)) ) | $spl144),
  inference(resolution,[],[f19403,f261])).
fof(f19450,plain,(
  ( ! [X14,X12,X15,X13] : (sP8(X12,'F_0') | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl144),
  inference(resolution,[],[f19403,f260])).
fof(f19449,plain,(
  ( ! [X10,X8,X11,X9] : (sP8(X8,'F_0') | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) ) | $spl144),
  inference(resolution,[],[f19403,f258])).
fof(f19447,plain,(
  ( ! [X2,X0,X3,X1] : (sP8(X0,'F_0') | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl144),
  inference(resolution,[],[f19403,f256])).
fof(f19412,plain,(
  ( ! [X21] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),'F_0') | ~p(s(bool,X21))) )),
  inference(subsumption_resolution,[],[f19380,f397])).
fof(f19380,plain,(
  ( ! [X21] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),'F_0') | s(bool,'F_0') = s(bool,X21) | ~p(s(bool,X21))) )),
  inference(resolution,[],[f19291,f2623])).
fof(f19291,plain,(
  ( ! [X391,X392] : (sP8(X391,X392) | sP8(X392,'F_0') | s(bool,'F_0') = s(bool,X391)) )),
  inference(resolution,[],[f16907,f412])).
fof(f19409,plain,(
  ( ! [X20] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),'F_0') | ~p(s(bool,X20))) )),
  inference(subsumption_resolution,[],[f19379,f397])).
fof(f19379,plain,(
  ( ! [X20] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),'F_0') | s(bool,'F_0') = s(bool,X20) | ~p(s(bool,X20))) )),
  inference(resolution,[],[f19291,f1301])).
fof(f19406,plain,(
  ( ! [X18] : (sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | ~p(s(bool,X18))) )),
  inference(subsumption_resolution,[],[f19377,f397])).
fof(f19377,plain,(
  ( ! [X18] : (sP8('const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | s(bool,'F_0') = s(bool,X18) | ~p(s(bool,X18))) )),
  inference(resolution,[],[f19291,f484])).
fof(f19405,plain,(
  ( ! [X14,X13] : (s(bool,'F_0') = s(bool,X13) | s(bool,'F_0') = s(bool,X14) | sP8(X14,X13)) ) | $spl146),
  inference(cnf_transformation,[],[f19405_D])).
fof(f19405_D,plain,(
  ( ! [X14,X13] : (s(bool,'F_0') = s(bool,X13) | s(bool,'F_0') = s(bool,X14) | sP8(X14,X13)) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f19403,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | sP8(X2,'F_0')) ) | $spl144),
  inference(cnf_transformation,[],[f19403_D])).
fof(f19403_D,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | sP8(X2,'F_0')) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f19375,plain,(
  ( ! [X14,X12,X13] : (sP8(X12,'F_0') | s(bool,'F_0') = s(bool,X13) | s(bool,'F_0') = s(bool,X14) | ~p(s(bool,X12)) | sP8(X14,X13)) )),
  inference(resolution,[],[f19291,f16905])).
fof(f19401,plain,(
  ( ! [X8,X9] : (sP8(X8,'F_0') | s(bool,'F_0') = s(bool,X9) | ~p(s(bool,X8))) )),
  inference(duplicate_literal_removal,[],[f19373])).
fof(f19373,plain,(
  ( ! [X8,X9] : (sP8(X8,'F_0') | s(bool,'F_0') = s(bool,X9) | ~p(s(bool,X8)) | s(bool,'F_0') = s(bool,X9)) )),
  inference(resolution,[],[f19291,f16172])).
fof(f19404,plain,(
  ( ! [X4,X5] : (sP8(X4,'F_0') | ~p(s(bool,X4)) | ~p(s(bool,X5))) )),
  inference(subsumption_resolution,[],[f19371,f397])).
fof(f19371,plain,(
  ( ! [X4,X5] : (sP8(X4,'F_0') | s(bool,'F_0') = s(bool,X5) | ~p(s(bool,X4)) | ~p(s(bool,X5))) )),
  inference(resolution,[],[f19291,f334])).
fof(f19402,plain,(
  ( ! [X2,X3] : (sP8(X2,'F_0') | ~p(s(bool,X3)) | ~p(s(bool,X2))) )),
  inference(subsumption_resolution,[],[f19370,f397])).
fof(f19370,plain,(
  ( ! [X2,X3] : (sP8(X2,'F_0') | s(bool,'F_0') = s(bool,X3) | ~p(s(bool,X3)) | ~p(s(bool,X2))) )),
  inference(resolution,[],[f19291,f327])).
fof(f19356,plain,(
  ( ! [X368,X369] : (sP8(X368,X369) | s(bool,'F_0') = s(bool,X368) | sP8(X369,'F_0')) ) | $spl142),
  inference(cnf_transformation,[],[f19356_D])).
fof(f19356_D,plain,(
  ( ! [X368,X369] : (sP8(X368,X369) | s(bool,'F_0') = s(bool,X368) | sP8(X369,'F_0')) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f19330,plain,(
  ( ! [X74,X72,X73] : (p(s(bool,X72)) | sP8(X73,X74) | sP8(X74,'F_0') | s(bool,'F_0') = s(bool,X73)) )),
  inference(duplicate_literal_removal,[],[f19322])).
fof(f19322,plain,(
  ( ! [X74,X72,X73] : (p(s(bool,X72)) | sP8(X73,X74) | sP8(X74,'F_0') | s(bool,'F_0') = s(bool,X73) | p(s(bool,X72))) )),
  inference(superposition,[],[f16907,f398])).
fof(f19331,plain,(
  ( ! [X70,X71,X69] : (p(s(bool,X69)) | sP8(X70,X71) | sP8(X71,'F_0') | s(bool,'F_0') = s(bool,X70)) )),
  inference(duplicate_literal_removal,[],[f19321])).
fof(f19321,plain,(
  ( ! [X70,X71,X69] : (p(s(bool,X69)) | sP8(X70,X71) | sP8(X71,'F_0') | s(bool,'F_0') = s(bool,X70) | p(s(bool,X69))) )),
  inference(superposition,[],[f16907,f400])).
fof(f19366,plain,(
  ( ! [X68,X66,X67] : (p(s(bool,X66)) | sP8(X67,X68) | sP8(X68,'F_0') | s(bool,'F_0') = s(bool,X67)) )),
  inference(subsumption_resolution,[],[f19320,f401])).
fof(f19320,plain,(
  ( ! [X68,X66,X67] : (p(s(bool,X66)) | sP8(X67,X68) | sP8(X68,'F_0') | s(bool,'F_0') = s(bool,X67) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f16907,f424])).
fof(f19290,plain,(
  ( ! [X389,X388,X390] : (sP8(X388,X389) | sP8(X389,'F_0') | s(bool,'F_0') = s(bool,X388) | p(s(bool,X390))) )),
  inference(resolution,[],[f16907,f411])).
fof(f411,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f150])).
fof(f150,plain,(
  ! [X0] : (~p(s(bool,'F_0')) | p(s(bool,X0)))),
  inference(ennf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0] : (p(s(bool,'F_0')) => p(s(bool,X0)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X13] : (p(s(bool,'F_0')) => p(s(bool,X13)))),
  file('thm.quantHeuristics.GUESSES_WEAKEN_THM|split|3','thm.bool.FALSITY')).
fof(f19289,plain,(
  ( ! [X385,X387,X386] : (sP8(X385,X386) | sP8(X386,'F_0') | s(bool,'F_0') = s(bool,X385) | p(s(bool,X387))) )),
  inference(resolution,[],[f16907,f390])).
fof(f390,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f235])).
fof(f235,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | ~p(s(bool,X0)) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,'F_0')) & p(s(bool,X0)))))),
  inference(flattening,[],[f234])).
fof(f234,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) | ~p(s(bool,X0))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,'F_0')) & p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f94])).
fof(f94,plain,(
  ! [X0] : ((p(s(bool,'F_0')) & p(s(bool,X0))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X13] : ((p(s(bool,'F_0')) & p(s(bool,X13))) <=> p(s(bool,'F_0')))),
  file('thm.quantHeuristics.GUESSES_WEAKEN_THM|split|3','thm.bool.AND_CLAUSES|split|2')).
fof(f19288,plain,(
  ( ! [X383,X382,X384] : (sP8(X382,X383) | sP8(X383,'F_0') | s(bool,'F_0') = s(bool,X382) | p(s(bool,X384))) )),
  inference(resolution,[],[f16907,f386])).
fof(f19357,plain,(
  ( ! [X379,X380,X381] : (sP8(X379,X380) | sP8(X380,'F_0') | s(bool,'F_0') = s(bool,X379) | p(s(bool,X381))) )),
  inference(subsumption_resolution,[],[f19287,f370])).
fof(f19287,plain,(
  ( ! [X379,X380,X381] : (sP8(X379,X380) | sP8(X380,'F_0') | s(bool,'F_0') = s(bool,X379) | p(s(bool,X381)) | ~p(s(bool,'T_0'))) )),
  inference(resolution,[],[f16907,f381])).
fof(f381,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f229])).
fof(f19286,plain,(
  ( ! [X377,X376,X378] : (sP8(X376,X377) | sP8(X377,'F_0') | s(bool,'F_0') = s(bool,X376) | ~p(s(bool,X378))) )),
  inference(resolution,[],[f16907,f368])).
fof(f19285,plain,(
  ( ! [X372,X374,X373,X375] : (sP8(X372,X373) | sP8(X373,'F_0') | s(bool,'F_0') = s(bool,X372) | sP9(X374,X375)) )),
  inference(resolution,[],[f16907,f350])).
fof(f350,plain,(
  ( ! [X0,X1] : (sP9(X0,X1) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f213])).
fof(f19284,plain,(
  ( ! [X368,X370,X369,X371] : (sP8(X368,X369) | sP8(X369,'F_0') | s(bool,'F_0') = s(bool,X368) | sP9(X370,X371)) )),
  inference(resolution,[],[f16907,f345])).
fof(f345,plain,(
  ( ! [X0,X1] : (sP9(X0,X1) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f211])).
fof(f19353,plain,(
  ( ! [X4,X2,X0,X1] : (sP8(X0,X1) | s(bool,X2) != s(bool,X4) | p(s(bool,X4)) | s(bool,'F_0') = s(bool,X0) | sP8(X1,X2)) ) | $spl140),
  inference(cnf_transformation,[],[f19353_D])).
fof(f19353_D,plain,(
  ( ! [X4,X2,X0,X1] : (sP8(X0,X1) | s(bool,X2) != s(bool,X4) | p(s(bool,X4)) | s(bool,'F_0') = s(bool,X0) | sP8(X1,X2)) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f19278,plain,(
  ( ! [X350,X346,X349,X347,X348] : (sP8(X346,X347) | sP8(X347,X348) | s(bool,'F_0') = s(bool,X346) | p(s(bool,X349)) | p(s(bool,X350)) | s(bool,X348) != s(bool,X350)) )),
  inference(resolution,[],[f16907,f426])).
fof(f19211,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (sP8(X42,X43) | sP8(X43,X44) | s(bool,'F_0') = s(bool,X42) | p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X44) != s(bool,X45) | s(bool,X46) != s(bool,X47)) )),
  inference(resolution,[],[f16907,f262])).
fof(f19208,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (sP8(X24,X25) | sP8(X25,X26) | s(bool,'F_0') = s(bool,X24) | p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X26) != s(bool,X29)) )),
  inference(resolution,[],[f16907,f261])).
fof(f19207,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (sP8(X18,X19) | sP8(X19,X20) | s(bool,'F_0') = s(bool,X18) | p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X20) != s(bool,X22)) )),
  inference(resolution,[],[f16907,f260])).
fof(f19206,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (sP8(X12,X13) | sP8(X13,X14) | s(bool,'F_0') = s(bool,X12) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X14) != s(bool,X15) | s(bool,X16) != s(bool,X17)) )),
  inference(resolution,[],[f16907,f258])).
fof(f19204,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP8(X0,X1) | sP8(X1,X2) | s(bool,'F_0') = s(bool,X0) | p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X2) != s(bool,X4)) )),
  inference(resolution,[],[f16907,f256])).
fof(f19334,plain,(
  ( ! [X47,X45,X46] : (sP8(X45,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X46))))) | p(s(bool,X47)) | s(bool,'F_0') = s(bool,X45)) )),
  inference(duplicate_literal_removal,[],[f19203])).
fof(f19203,plain,(
  ( ! [X47,X45,X46] : (sP8(X45,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X46))))) | p(s(bool,X47)) | s(bool,'F_0') = s(bool,X45) | p(s(bool,X47))) )),
  inference(resolution,[],[f16907,f3043])).
fof(f19335,plain,(
  ( ! [X39,X40] : (sP8(X39,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | p(s(bool,X40)) | s(bool,'F_0') = s(bool,X39)) )),
  inference(duplicate_literal_removal,[],[f19200])).
fof(f19200,plain,(
  ( ! [X39,X40] : (sP8(X39,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | p(s(bool,X40)) | s(bool,'F_0') = s(bool,X39) | p(s(bool,X40))) )),
  inference(resolution,[],[f16907,f594])).
fof(f19160,plain,(
  ( ! [X125,X126,X124] : (s(bool,'F_0') = s(bool,X124) | ~p(s(bool,X125)) | sP8(X124,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X126)))))) )),
  inference(duplicate_literal_removal,[],[f19153])).
fof(f19153,plain,(
  ( ! [X125,X126,X124] : (s(bool,'F_0') = s(bool,X124) | ~p(s(bool,X125)) | sP8(X124,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X126))))) | ~p(s(bool,X125))) )),
  inference(resolution,[],[f16906,f3044])).
fof(f19161,plain,(
  ( ! [X118,X119] : (s(bool,'F_0') = s(bool,X118) | ~p(s(bool,X119)) | sP8(X118,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) )),
  inference(duplicate_literal_removal,[],[f19150])).
fof(f19150,plain,(
  ( ! [X118,X119] : (s(bool,'F_0') = s(bool,X118) | ~p(s(bool,X119)) | sP8(X118,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | ~p(s(bool,X119))) )),
  inference(resolution,[],[f16906,f595])).
fof(f19114,plain,(
  ( ! [X138,X139,X137] : (s(bool,'F_0') = s(bool,X137) | ~p(s(bool,X138)) | sP8(X137,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X139)))))) )),
  inference(duplicate_literal_removal,[],[f19113])).
fof(f19113,plain,(
  ( ! [X138,X139,X137] : (s(bool,'F_0') = s(bool,X137) | ~p(s(bool,X138)) | sP8(X137,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X139))))) | ~p(s(bool,X138))) )),
  inference(resolution,[],[f16905,f3045])).
fof(f19115,plain,(
  ( ! [X132,X131] : (s(bool,'F_0') = s(bool,X131) | ~p(s(bool,X132)) | sP8(X131,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) )),
  inference(duplicate_literal_removal,[],[f19110])).
fof(f19110,plain,(
  ( ! [X132,X131] : (s(bool,'F_0') = s(bool,X131) | ~p(s(bool,X132)) | sP8(X131,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | ~p(s(bool,X132))) )),
  inference(resolution,[],[f16905,f596])).
fof(f18257,plain,(
  ( ! [X54,X56,X55] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X54) | s(bool,'F_0') = s(bool,X54) | p(s(bool,X55)) | p(s(bool,X56))) )),
  inference(resolution,[],[f17139,f589])).
fof(f17979,plain,(
  ( ! [X33,X32] : (sP8(X32,X33) | s(bool,'F_0') != s(bool,X32) | s(bool,'F_0') = s(bool,X33)) ) | $spl134),
  inference(cnf_transformation,[],[f17979_D])).
fof(f17979_D,plain,(
  ( ! [X33,X32] : (sP8(X32,X33) | s(bool,'F_0') != s(bool,X32) | s(bool,'F_0') = s(bool,X33)) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f18253,plain,(
  ( ! [X45,X43,X44] : (sP8(X43,X44) | s(bool,'F_0') = s(bool,X44) | p(s(bool,X45)) | s(bool,'F_0') != s(bool,X43)) )),
  inference(resolution,[],[f17139,f1971])).
fof(f18220,plain,(
  ( ! [X52,X53,X51] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X51) | s(bool,'F_0') = s(bool,X51) | p(s(bool,X52)) | p(s(bool,X53))) )),
  inference(resolution,[],[f17138,f588])).
fof(f18183,plain,(
  ( ! [X50,X49] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X49) | s(bool,'F_0') = s(bool,X49) | p(s(bool,X50))) )),
  inference(resolution,[],[f17137,f3375])).
fof(f18180,plain,(
  ( ! [X41,X42,X40] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X40) | s(bool,'F_0') = s(bool,X40) | p(s(bool,X41)) | p(s(bool,X42))) )),
  inference(resolution,[],[f17137,f590])).
fof(f18145,plain,(
  ( ! [X45,X43,X44] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X43) | s(bool,'F_0') = s(bool,X43) | ~p(s(bool,X44)) | ~p(s(bool,X45))) )),
  inference(resolution,[],[f17136,f585])).
fof(f18144,plain,(
  ( ! [X41,X42,X40] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X40) | s(bool,'F_0') = s(bool,X40) | sP7('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X41,X42) | ~sP4(X42,X41,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) )),
  inference(resolution,[],[f17136,f7261])).
fof(f17981,plain,(
  ( ! [X99,X102,X100] : (sP8(X99,X100) | s(bool,X99) != s(bool,X102) | s(bool,'T_0') = s(bool,X102) | s(bool,'F_0') = s(bool,X100)) ) | $spl138),
  inference(cnf_transformation,[],[f17981_D])).
fof(f17981_D,plain,(
  ( ! [X99,X102,X100] : (sP8(X99,X100) | s(bool,X99) != s(bool,X102) | s(bool,'T_0') = s(bool,X102) | s(bool,'F_0') = s(bool,X100)) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f17966,plain,(
  ( ! [X101,X99,X102,X100] : (sP8(X99,X100) | s(bool,'F_0') = s(bool,X100) | ~p(s(bool,X101)) | s(bool,'T_0') = s(bool,X102) | s(bool,X99) != s(bool,X102)) )),
  inference(resolution,[],[f17133,f14828])).
fof(f17964,plain,(
  ( ! [X94,X92,X93] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X92) | s(bool,'F_0') = s(bool,X92) | ~p(s(bool,X93)) | ~p(s(bool,X94))) )),
  inference(resolution,[],[f17133,f582])).
fof(f17963,plain,(
  ( ! [X90,X91,X89] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X89) | s(bool,'F_0') = s(bool,X89) | sP7(X90,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X91) | ~sP4(X91,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X90)) )),
  inference(resolution,[],[f17133,f6889])).
fof(f17980,plain,(
  ( ! [X59,X58] : (sP8(X58,X59) | s(bool,'T_0') = s(bool,X58) | s(bool,'F_0') = s(bool,X59)) ) | $spl136),
  inference(cnf_transformation,[],[f17980_D])).
fof(f17980_D,plain,(
  ( ! [X59,X58] : (sP8(X58,X59) | s(bool,'T_0') = s(bool,X58) | s(bool,'F_0') = s(bool,X59)) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f17955,plain,(
  ( ! [X64,X62,X65,X63] : (sP8(X62,X63) | s(bool,'F_0') = s(bool,X63) | s(bool,'T_0') = s(bool,X62) | s(bool,'F_0') = s(bool,X64) | ~p(s(bool,X65))) )),
  inference(resolution,[],[f17133,f14390])).
fof(f17954,plain,(
  ( ! [X61,X59,X60,X58] : (sP8(X58,X59) | s(bool,'F_0') = s(bool,X59) | s(bool,'T_0') = s(bool,X58) | ~p(s(bool,X60)) | ~p(s(bool,X61))) )),
  inference(resolution,[],[f17133,f14047])).
fof(f17231,plain,(
  ( ! [X0,X3,X1] : (sP8(X0,X1) | s(bool,X0) != s(bool,X3) | p(s(bool,X3)) | s(bool,'F_0') = s(bool,X1)) ) | $spl110),
  inference(cnf_transformation,[],[f17231_D])).
fof(f17231_D,plain,(
  ( ! [X0,X3,X1] : (sP8(X0,X1) | s(bool,X0) != s(bool,X3) | p(s(bool,X3)) | s(bool,'F_0') = s(bool,X1)) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f17951,plain,(
  ( ! [X47,X45,X48,X46,X49] : (sP8(X45,X46) | s(bool,'F_0') = s(bool,X46) | ~p(s(bool,X47)) | ~p(s(bool,X48)) | s(bool,X45) != s(bool,X49) | p(s(bool,X49))) ) | $spl6),
  inference(resolution,[],[f17133,f8878])).
fof(f17949,plain,(
  ( ! [X39,X37,X38,X36] : (sP8(X36,X37) | s(bool,'F_0') = s(bool,X37) | ~p(s(bool,X38)) | s(bool,'F_0') != s(bool,X36) | s(bool,'T_0') != s(bool,X39)) ) | $spl24),
  inference(resolution,[],[f17133,f7489])).
fof(f17948,plain,(
  ( ! [X35,X33,X34,X32] : (sP8(X32,X33) | s(bool,'F_0') = s(bool,X33) | s(bool,'F_0') != s(bool,X32) | ~p(s(bool,X34)) | ~p(s(bool,X35))) ) | $spl24),
  inference(resolution,[],[f17133,f7486])).
fof(f17918,plain,(
  ( ! [X30,X31,X29] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X29) | s(bool,'F_0') = s(bool,X29) | p(s(bool,X30)) | ~p(s(bool,X31))) )),
  inference(resolution,[],[f17120,f569])).
fof(f17917,plain,(
  ( ! [X28,X26,X27] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X26) | s(bool,'F_0') = s(bool,X26) | sP3(X27,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X28) | p(s(bool,X28))) )),
  inference(resolution,[],[f17120,f567])).
fof(f17886,plain,(
  ( ! [X24,X23,X25] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X23) | s(bool,'F_0') = s(bool,X23) | p(s(bool,X24)) | ~p(s(bool,X25))) )),
  inference(resolution,[],[f17119,f570])).
fof(f17885,plain,(
  ( ! [X21,X22,X20] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X20) | s(bool,'F_0') = s(bool,X20) | sP3(X21,X22,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | p(s(bool,X22))) )),
  inference(resolution,[],[f17119,f568])).
fof(f17858,plain,(
  ( ! [X33,X34,X32] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X32) | s(bool,'F_0') = s(bool,X32) | ~p(s(bool,X33)) | ~p(s(bool,X34))) )),
  inference(resolution,[],[f17106,f553])).
fof(f17857,plain,(
  ( ! [X30,X31,X29] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X29) | s(bool,'F_0') = s(bool,X29) | sP1(X30,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X31) | ~p(s(bool,X31))) )),
  inference(resolution,[],[f17106,f552])).
fof(f17827,plain,(
  ( ! [X54,X56,X55] : (sP8(X54,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | s(bool,'F_0') = s(bool,X54) | p(s(bool,X55)) | p(s(bool,X56))) )),
  inference(resolution,[],[f16902,f589])).
fof(f17683,plain,(
  ( ! [X33,X32] : (sP8(X32,X33) | s(bool,'F_0') != s(bool,X33) | s(bool,'F_0') = s(bool,X32)) ) | $spl124),
  inference(cnf_transformation,[],[f17683_D])).
fof(f17683_D,plain,(
  ( ! [X33,X32] : (sP8(X32,X33) | s(bool,'F_0') != s(bool,X33) | s(bool,'F_0') = s(bool,X32)) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f17823,plain,(
  ( ! [X45,X43,X44] : (sP8(X43,X44) | s(bool,'F_0') = s(bool,X43) | p(s(bool,X45)) | s(bool,'F_0') != s(bool,X44)) )),
  inference(resolution,[],[f16902,f1971])).
fof(f17790,plain,(
  ( ! [X52,X53,X51] : (sP8(X51,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | s(bool,'F_0') = s(bool,X51) | p(s(bool,X52)) | p(s(bool,X53))) )),
  inference(resolution,[],[f16901,f588])).
fof(f17753,plain,(
  ( ! [X50,X49] : (sP8(X49,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | s(bool,'F_0') = s(bool,X49) | p(s(bool,X50))) )),
  inference(resolution,[],[f16900,f3375])).
fof(f17750,plain,(
  ( ! [X41,X42,X40] : (sP8(X40,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | s(bool,'F_0') = s(bool,X40) | p(s(bool,X41)) | p(s(bool,X42))) )),
  inference(resolution,[],[f16900,f590])).
fof(f17714,plain,(
  ( ! [X45,X43,X44] : (sP8(X43,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | s(bool,'F_0') = s(bool,X43) | ~p(s(bool,X44)) | ~p(s(bool,X45))) )),
  inference(resolution,[],[f16899,f585])).
fof(f17713,plain,(
  ( ! [X41,X42,X40] : (sP8(X40,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | s(bool,'F_0') = s(bool,X40) | sP7('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X41,X42) | ~sP4(X42,X41,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) )),
  inference(resolution,[],[f16899,f7261])).
fof(f17686,plain,(
  ( ! [X99,X102,X100] : (sP8(X99,X100) | s(bool,X100) != s(bool,X102) | s(bool,'T_0') = s(bool,X102) | s(bool,'F_0') = s(bool,X99)) ) | $spl130),
  inference(cnf_transformation,[],[f17686_D])).
fof(f17686_D,plain,(
  ( ! [X99,X102,X100] : (sP8(X99,X100) | s(bool,X100) != s(bool,X102) | s(bool,'T_0') = s(bool,X102) | s(bool,'F_0') = s(bool,X99)) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f17670,plain,(
  ( ! [X101,X99,X102,X100] : (sP8(X99,X100) | s(bool,'F_0') = s(bool,X99) | ~p(s(bool,X101)) | s(bool,'T_0') = s(bool,X102) | s(bool,X100) != s(bool,X102)) )),
  inference(resolution,[],[f16896,f14828])).
fof(f17668,plain,(
  ( ! [X94,X92,X93] : (sP8(X92,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | s(bool,'F_0') = s(bool,X92) | ~p(s(bool,X93)) | ~p(s(bool,X94))) )),
  inference(resolution,[],[f16896,f582])).
fof(f17667,plain,(
  ( ! [X90,X91,X89] : (sP8(X89,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | s(bool,'F_0') = s(bool,X89) | sP7(X90,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X91) | ~sP4(X91,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X90)) )),
  inference(resolution,[],[f16896,f6889])).
fof(f17684,plain,(
  ( ! [X59,X58] : (sP8(X58,X59) | s(bool,'T_0') = s(bool,X59) | s(bool,'F_0') = s(bool,X58)) ) | $spl126),
  inference(cnf_transformation,[],[f17684_D])).
fof(f17684_D,plain,(
  ( ! [X59,X58] : (sP8(X58,X59) | s(bool,'T_0') = s(bool,X59) | s(bool,'F_0') = s(bool,X58)) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f17659,plain,(
  ( ! [X64,X62,X65,X63] : (sP8(X62,X63) | s(bool,'F_0') = s(bool,X62) | s(bool,'T_0') = s(bool,X63) | s(bool,'F_0') = s(bool,X64) | ~p(s(bool,X65))) )),
  inference(resolution,[],[f16896,f14390])).
fof(f17658,plain,(
  ( ! [X61,X59,X60,X58] : (sP8(X58,X59) | s(bool,'F_0') = s(bool,X58) | s(bool,'T_0') = s(bool,X59) | ~p(s(bool,X60)) | ~p(s(bool,X61))) )),
  inference(resolution,[],[f16896,f14047])).
fof(f16994,plain,(
  ( ! [X0,X3,X1] : (sP8(X0,X1) | s(bool,X1) != s(bool,X3) | p(s(bool,X3)) | s(bool,'F_0') = s(bool,X0)) ) | $spl100),
  inference(cnf_transformation,[],[f16994_D])).
fof(f16994_D,plain,(
  ( ! [X0,X3,X1] : (sP8(X0,X1) | s(bool,X1) != s(bool,X3) | p(s(bool,X3)) | s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f17655,plain,(
  ( ! [X47,X45,X48,X46,X49] : (sP8(X45,X46) | s(bool,'F_0') = s(bool,X45) | ~p(s(bool,X47)) | ~p(s(bool,X48)) | s(bool,X46) != s(bool,X49) | p(s(bool,X49))) ) | $spl6),
  inference(resolution,[],[f16896,f8878])).
fof(f17653,plain,(
  ( ! [X39,X37,X38,X36] : (sP8(X36,X37) | s(bool,'F_0') = s(bool,X36) | ~p(s(bool,X38)) | s(bool,'F_0') != s(bool,X37) | s(bool,'T_0') != s(bool,X39)) ) | $spl24),
  inference(resolution,[],[f16896,f7489])).
fof(f17652,plain,(
  ( ! [X35,X33,X34,X32] : (sP8(X32,X33) | s(bool,'F_0') = s(bool,X32) | s(bool,'F_0') != s(bool,X33) | ~p(s(bool,X34)) | ~p(s(bool,X35))) ) | $spl24),
  inference(resolution,[],[f16896,f7486])).
fof(f17621,plain,(
  ( ! [X30,X31,X29] : (sP8(X29,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | s(bool,'F_0') = s(bool,X29) | p(s(bool,X30)) | ~p(s(bool,X31))) )),
  inference(resolution,[],[f16883,f569])).
fof(f17620,plain,(
  ( ! [X28,X26,X27] : (sP8(X26,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | s(bool,'F_0') = s(bool,X26) | sP3(X27,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X28) | p(s(bool,X28))) )),
  inference(resolution,[],[f16883,f567])).
fof(f17588,plain,(
  ( ! [X24,X23,X25] : (sP8(X23,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | s(bool,'F_0') = s(bool,X23) | p(s(bool,X24)) | ~p(s(bool,X25))) )),
  inference(resolution,[],[f16882,f570])).
fof(f17587,plain,(
  ( ! [X21,X22,X20] : (sP8(X20,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | s(bool,'F_0') = s(bool,X20) | sP3(X21,X22,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | p(s(bool,X22))) )),
  inference(resolution,[],[f16882,f568])).
fof(f17559,plain,(
  ( ! [X33,X34,X32] : (sP8(X32,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | s(bool,'F_0') = s(bool,X32) | ~p(s(bool,X33)) | ~p(s(bool,X34))) )),
  inference(resolution,[],[f16869,f553])).
fof(f17558,plain,(
  ( ! [X30,X31,X29] : (sP8(X29,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | s(bool,'F_0') = s(bool,X29) | sP1(X30,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X31) | ~p(s(bool,X31))) )),
  inference(resolution,[],[f16869,f552])).
fof(f17427,plain,(
  ( ! [X14,X12,X13] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X12)))),X13) | s(bool,'F_0') = s(bool,X13) | p(s(bool,X14))) )),
  inference(resolution,[],[f17151,f3090])).
fof(f17426,plain,(
  ( ! [X10,X11] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X10) | s(bool,'F_0') = s(bool,X10) | p(s(bool,X11))) )),
  inference(resolution,[],[f17151,f640])).
fof(f17398,plain,(
  ( ! [X10,X8,X9] : (sP8(happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X8)))),X9) | s(bool,'F_0') = s(bool,X9) | p(s(bool,X10))) )),
  inference(resolution,[],[f17150,f3089])).
fof(f17397,plain,(
  ( ! [X6,X7] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X6) | s(bool,'F_0') = s(bool,X6) | p(s(bool,X7))) )),
  inference(resolution,[],[f17150,f639])).
fof(f17234,plain,(
  ( ! [X288] : (sP8('F_0',X288) | s(bool,'F_0') = s(bool,X288)) ) | $spl112),
  inference(cnf_transformation,[],[f17234_D])).
fof(f17234_D,plain,(
  ( ! [X288] : (sP8('F_0',X288) | s(bool,'F_0') = s(bool,X288)) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f17216,plain,(
  ( ! [X52,X53] : (p(s(bool,X52)) | sP8('F_0',X53) | s(bool,'F_0') = s(bool,X53)) )),
  inference(duplicate_literal_removal,[],[f17208])).
fof(f17208,plain,(
  ( ! [X52,X53] : (p(s(bool,X52)) | sP8('F_0',X53) | s(bool,'F_0') = s(bool,X53) | p(s(bool,X52))) )),
  inference(superposition,[],[f16182,f398])).
fof(f17217,plain,(
  ( ! [X50,X51] : (p(s(bool,X50)) | sP8('F_0',X51) | s(bool,'F_0') = s(bool,X51)) )),
  inference(duplicate_literal_removal,[],[f17207])).
fof(f17207,plain,(
  ( ! [X50,X51] : (p(s(bool,X50)) | sP8('F_0',X51) | s(bool,'F_0') = s(bool,X51) | p(s(bool,X50))) )),
  inference(superposition,[],[f16182,f400])).
fof(f17244,plain,(
  ( ! [X48,X49] : (p(s(bool,X48)) | sP8('F_0',X49) | s(bool,'F_0') = s(bool,X49)) )),
  inference(subsumption_resolution,[],[f17206,f401])).
fof(f17206,plain,(
  ( ! [X48,X49] : (p(s(bool,X48)) | sP8('F_0',X49) | s(bool,'F_0') = s(bool,X49) | s(bool,'T_0') = s(bool,X48)) )),
  inference(superposition,[],[f16182,f424])).
fof(f17177,plain,(
  ( ! [X302,X303] : (sP8('F_0',X302) | s(bool,'F_0') = s(bool,X302) | p(s(bool,X303))) )),
  inference(resolution,[],[f16182,f411])).
fof(f17176,plain,(
  ( ! [X300,X301] : (sP8('F_0',X300) | s(bool,'F_0') = s(bool,X300) | p(s(bool,X301))) )),
  inference(resolution,[],[f16182,f390])).
fof(f17175,plain,(
  ( ! [X298,X299] : (sP8('F_0',X298) | s(bool,'F_0') = s(bool,X298) | p(s(bool,X299))) )),
  inference(resolution,[],[f16182,f386])).
fof(f17235,plain,(
  ( ! [X296,X297] : (sP8('F_0',X296) | s(bool,'F_0') = s(bool,X296) | p(s(bool,X297))) )),
  inference(subsumption_resolution,[],[f17174,f370])).
fof(f17174,plain,(
  ( ! [X296,X297] : (sP8('F_0',X296) | s(bool,'F_0') = s(bool,X296) | p(s(bool,X297)) | ~p(s(bool,'T_0'))) )),
  inference(resolution,[],[f16182,f381])).
fof(f17173,plain,(
  ( ! [X294,X295] : (sP8('F_0',X294) | s(bool,'F_0') = s(bool,X294) | ~p(s(bool,X295))) )),
  inference(resolution,[],[f16182,f368])).
fof(f17172,plain,(
  ( ! [X292,X293,X291] : (sP8('F_0',X291) | s(bool,'F_0') = s(bool,X291) | sP9(X292,X293)) )),
  inference(resolution,[],[f16182,f350])).
fof(f17171,plain,(
  ( ! [X288,X290,X289] : (sP8('F_0',X288) | s(bool,'F_0') = s(bool,X288) | sP9(X289,X290)) )),
  inference(resolution,[],[f16182,f345])).
fof(f17165,plain,(
  ( ! [X273,X275,X272,X274] : (sP8(X272,X273) | s(bool,'F_0') = s(bool,X273) | p(s(bool,X274)) | p(s(bool,X275)) | s(bool,X272) != s(bool,X275)) )),
  inference(resolution,[],[f16182,f426])).
fof(f17098,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP8(X35,X36) | s(bool,'F_0') = s(bool,X36) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X35) != s(bool,X37) | s(bool,X38) != s(bool,X39)) )),
  inference(resolution,[],[f16182,f262])).
fof(f17095,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP8(X20,X21) | s(bool,'F_0') = s(bool,X21) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X20) != s(bool,X24)) )),
  inference(resolution,[],[f16182,f261])).
fof(f17094,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP8(X15,X16) | s(bool,'F_0') = s(bool,X16) | p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X15) != s(bool,X18)) )),
  inference(resolution,[],[f16182,f260])).
fof(f17093,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP8(X10,X11) | s(bool,'F_0') = s(bool,X11) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X10) != s(bool,X12) | s(bool,X13) != s(bool,X14)) )),
  inference(resolution,[],[f16182,f258])).
fof(f17091,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP8(X0,X1) | s(bool,'F_0') = s(bool,X1) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X0) != s(bool,X3)) )),
  inference(resolution,[],[f16182,f256])).
fof(f17071,plain,(
  ( ! [X14,X12,X13] : (sP8(X12,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X13))))) | s(bool,'F_0') = s(bool,X12) | p(s(bool,X14))) )),
  inference(resolution,[],[f16914,f3090])).
fof(f16914,plain,(
  ( ! [X241,X239,X240] : (sP8(X239,X240) | sP9(X240,X241) | s(bool,'F_0') = s(bool,X239)) )),
  inference(resolution,[],[f16178,f344])).
fof(f17070,plain,(
  ( ! [X10,X11] : (sP8(X10,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | s(bool,'F_0') = s(bool,X10) | p(s(bool,X11))) )),
  inference(resolution,[],[f16914,f640])).
fof(f17042,plain,(
  ( ! [X10,X8,X9] : (sP8(X8,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X9))))) | s(bool,'F_0') = s(bool,X8) | p(s(bool,X10))) )),
  inference(resolution,[],[f16913,f3089])).
fof(f16913,plain,(
  ( ! [X237,X238,X236] : (sP8(X236,X237) | sP9(X238,X237) | s(bool,'F_0') = s(bool,X236)) )),
  inference(resolution,[],[f16178,f343])).
fof(f17041,plain,(
  ( ! [X6,X7] : (sP8(X6,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | s(bool,'F_0') = s(bool,X6) | p(s(bool,X7))) )),
  inference(resolution,[],[f16913,f639])).
fof(f16998,plain,(
  ( ! [X288] : (sP8(X288,'F_0') | s(bool,'F_0') = s(bool,X288)) ) | $spl104),
  inference(cnf_transformation,[],[f16998_D])).
fof(f16998_D,plain,(
  ( ! [X288] : (sP8(X288,'F_0') | s(bool,'F_0') = s(bool,X288)) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f16979,plain,(
  ( ! [X52,X53] : (p(s(bool,X52)) | sP8(X53,'F_0') | s(bool,'F_0') = s(bool,X53)) )),
  inference(duplicate_literal_removal,[],[f16971])).
fof(f16971,plain,(
  ( ! [X52,X53] : (p(s(bool,X52)) | sP8(X53,'F_0') | s(bool,'F_0') = s(bool,X53) | p(s(bool,X52))) )),
  inference(superposition,[],[f16178,f398])).
fof(f16980,plain,(
  ( ! [X50,X51] : (p(s(bool,X50)) | sP8(X51,'F_0') | s(bool,'F_0') = s(bool,X51)) )),
  inference(duplicate_literal_removal,[],[f16970])).
fof(f16970,plain,(
  ( ! [X50,X51] : (p(s(bool,X50)) | sP8(X51,'F_0') | s(bool,'F_0') = s(bool,X51) | p(s(bool,X50))) )),
  inference(superposition,[],[f16178,f400])).
fof(f17008,plain,(
  ( ! [X48,X49] : (p(s(bool,X48)) | sP8(X49,'F_0') | s(bool,'F_0') = s(bool,X49)) )),
  inference(subsumption_resolution,[],[f16969,f401])).
fof(f16969,plain,(
  ( ! [X48,X49] : (p(s(bool,X48)) | sP8(X49,'F_0') | s(bool,'F_0') = s(bool,X49) | s(bool,'T_0') = s(bool,X48)) )),
  inference(superposition,[],[f16178,f424])).
fof(f16940,plain,(
  ( ! [X302,X303] : (sP8(X302,'F_0') | s(bool,'F_0') = s(bool,X302) | p(s(bool,X303))) )),
  inference(resolution,[],[f16178,f411])).
fof(f16939,plain,(
  ( ! [X300,X301] : (sP8(X300,'F_0') | s(bool,'F_0') = s(bool,X300) | p(s(bool,X301))) )),
  inference(resolution,[],[f16178,f390])).
fof(f16938,plain,(
  ( ! [X298,X299] : (sP8(X298,'F_0') | s(bool,'F_0') = s(bool,X298) | p(s(bool,X299))) )),
  inference(resolution,[],[f16178,f386])).
fof(f16999,plain,(
  ( ! [X296,X297] : (sP8(X296,'F_0') | s(bool,'F_0') = s(bool,X296) | p(s(bool,X297))) )),
  inference(subsumption_resolution,[],[f16937,f370])).
fof(f16937,plain,(
  ( ! [X296,X297] : (sP8(X296,'F_0') | s(bool,'F_0') = s(bool,X296) | p(s(bool,X297)) | ~p(s(bool,'T_0'))) )),
  inference(resolution,[],[f16178,f381])).
fof(f16936,plain,(
  ( ! [X294,X295] : (sP8(X294,'F_0') | s(bool,'F_0') = s(bool,X294) | ~p(s(bool,X295))) )),
  inference(resolution,[],[f16178,f368])).
fof(f16935,plain,(
  ( ! [X292,X293,X291] : (sP8(X291,'F_0') | s(bool,'F_0') = s(bool,X291) | sP9(X292,X293)) )),
  inference(resolution,[],[f16178,f350])).
fof(f16934,plain,(
  ( ! [X288,X290,X289] : (sP8(X288,'F_0') | s(bool,'F_0') = s(bool,X288) | sP9(X289,X290)) )),
  inference(resolution,[],[f16178,f345])).
fof(f16928,plain,(
  ( ! [X273,X275,X272,X274] : (sP8(X272,X273) | s(bool,'F_0') = s(bool,X272) | p(s(bool,X274)) | p(s(bool,X275)) | s(bool,X273) != s(bool,X275)) )),
  inference(resolution,[],[f16178,f426])).
fof(f16861,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP8(X35,X36) | s(bool,'F_0') = s(bool,X35) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X37) | s(bool,X38) != s(bool,X39)) )),
  inference(resolution,[],[f16178,f262])).
fof(f16858,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP8(X20,X21) | s(bool,'F_0') = s(bool,X20) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X21) != s(bool,X24)) )),
  inference(resolution,[],[f16178,f261])).
fof(f16857,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP8(X15,X16) | s(bool,'F_0') = s(bool,X15) | p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) )),
  inference(resolution,[],[f16178,f260])).
fof(f16856,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP8(X10,X11) | s(bool,'F_0') = s(bool,X10) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X12) | s(bool,X13) != s(bool,X14)) )),
  inference(resolution,[],[f16178,f258])).
fof(f16854,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP8(X0,X1) | s(bool,'F_0') = s(bool,X0) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) )),
  inference(resolution,[],[f16178,f256])).
fof(f13962,plain,(
  ( ! [X30,X28,X31] : (s(bool,'T_0') = s(bool,X30) | ~p(s(bool,X28)) | ~sP7(X31,X30,X28) | ~p(s(bool,X31))) ) | $spl98),
  inference(cnf_transformation,[],[f13962_D])).
fof(f13962_D,plain,(
  ( ! [X30,X28,X31] : (s(bool,'T_0') = s(bool,X30) | ~p(s(bool,X28)) | ~sP7(X31,X30,X28) | ~p(s(bool,X31))) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f13963,plain,(
  ( ! [X35,X33,X34,X32] : (s(bool,'T_0') = s(bool,X34) | ~p(s(bool,X35)) | ~sP7(X35,X34,X32) | ~p(s(bool,X33)) | ~p(s(bool,X32))) )),
  inference(subsumption_resolution,[],[f13928,f402])).
fof(f13928,plain,(
  ( ! [X35,X33,X34,X32] : (s(bool,'T_0') != s(bool,X33) | s(bool,'T_0') = s(bool,X34) | ~p(s(bool,X35)) | ~sP7(X35,X34,X32) | ~p(s(bool,X33)) | ~p(s(bool,X32))) )),
  inference(superposition,[],[f13735,f358])).
fof(f358,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f139])).
fof(f139,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f138])).
fof(f138,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X11,X12] : ((p(s(bool,X11)) => p(s(bool,X12))) => ((p(s(bool,X12)) => p(s(bool,X11))) => s(bool,X11) = s(bool,X12)))),
  file('thm.quantHeuristics.GUESSES_WEAKEN_THM|split|3','thm.bool.IMP_ANTISYM_AX')).
fof(f13961,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'T_0') = s(bool,X30) | ~p(s(bool,X31)) | ~sP7(X31,X30,X28) | ~p(s(bool,X28)) | ~p(s(bool,X29))) )),
  inference(subsumption_resolution,[],[f13927,f402])).
fof(f13927,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'T_0') != s(bool,X29) | s(bool,'T_0') = s(bool,X30) | ~p(s(bool,X31)) | ~sP7(X31,X30,X28) | ~p(s(bool,X28)) | ~p(s(bool,X29))) )),
  inference(superposition,[],[f13735,f358])).
fof(f13814,plain,(
  ( ! [X30,X28,X31] : (s(bool,'T_0') = s(bool,X30) | ~p(s(bool,X28)) | ~sP7(X31,X28,X30) | ~p(s(bool,X31))) ) | $spl96),
  inference(cnf_transformation,[],[f13814_D])).
fof(f13814_D,plain,(
  ( ! [X30,X28,X31] : (s(bool,'T_0') = s(bool,X30) | ~p(s(bool,X28)) | ~sP7(X31,X28,X30) | ~p(s(bool,X31))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f13815,plain,(
  ( ! [X35,X33,X34,X32] : (s(bool,'T_0') = s(bool,X34) | ~p(s(bool,X35)) | ~sP7(X35,X32,X34) | ~p(s(bool,X33)) | ~p(s(bool,X32))) )),
  inference(subsumption_resolution,[],[f13780,f402])).
fof(f13780,plain,(
  ( ! [X35,X33,X34,X32] : (s(bool,'T_0') != s(bool,X33) | s(bool,'T_0') = s(bool,X34) | ~p(s(bool,X35)) | ~sP7(X35,X32,X34) | ~p(s(bool,X33)) | ~p(s(bool,X32))) )),
  inference(superposition,[],[f13734,f358])).
fof(f13734,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X7) | s(bool,'T_0') = s(bool,X6) | ~p(s(bool,X8)) | ~sP7(X8,X7,X6)) )),
  inference(forward_demodulation,[],[f13696,f504])).
fof(f13696,plain,(
  ( ! [X6,X8,X7] : (s(bool,X7) != s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | s(bool,'T_0') = s(bool,X6) | ~p(s(bool,X8)) | ~sP7(X8,X7,X6)) )),
  inference(superposition,[],[f13015,f318])).
fof(f13813,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'T_0') = s(bool,X30) | ~p(s(bool,X31)) | ~sP7(X31,X28,X30) | ~p(s(bool,X28)) | ~p(s(bool,X29))) )),
  inference(subsumption_resolution,[],[f13779,f402])).
fof(f13779,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'T_0') != s(bool,X29) | s(bool,'T_0') = s(bool,X30) | ~p(s(bool,X31)) | ~sP7(X31,X28,X30) | ~p(s(bool,X28)) | ~p(s(bool,X29))) )),
  inference(superposition,[],[f13734,f358])).
fof(f11220,plain,(
  ( ! [X584,X581,X583] : (~p(s(bool,X583)) | p(s(bool,X584)) | sP7(X584,X583,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X581)))))) ) | $spl94),
  inference(cnf_transformation,[],[f11220_D])).
fof(f11220_D,plain,(
  ( ! [X584,X581,X583] : (~p(s(bool,X583)) | p(s(bool,X584)) | sP7(X584,X583,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X581)))))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f11219,plain,(
  ( ! [X584,X582,X581,X583] : (p(s(bool,X582)) | ~p(s(bool,X583)) | sP7(X584,X583,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X581))))) | p(s(bool,X584))) ) | $spl6),
  inference(subsumption_resolution,[],[f11043,f398])).
fof(f11043,plain,(
  ( ! [X584,X582,X581,X583] : (s(bool,'F_0') != s(bool,X582) | p(s(bool,X582)) | ~p(s(bool,X583)) | sP7(X584,X583,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X581))))) | p(s(bool,X584))) ) | $spl6),
  inference(superposition,[],[f1200,f3062])).
fof(f1200,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X5) != s(bool,X6) | p(s(bool,X6)) | ~p(s(bool,X4)) | sP7(X7,X4,X5) | p(s(bool,X7))) ) | $spl6),
  inference(superposition,[],[f528,f316])).
fof(f11218,plain,(
  ( ! [X569,X571,X572] : (~p(s(bool,X571)) | p(s(bool,X572)) | sP7(X572,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X569)))),X571)) ) | $spl92),
  inference(cnf_transformation,[],[f11218_D])).
fof(f11218_D,plain,(
  ( ! [X569,X571,X572] : (~p(s(bool,X571)) | p(s(bool,X572)) | sP7(X572,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X569)))),X571)) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f11217,plain,(
  ( ! [X569,X571,X572,X570] : (p(s(bool,X570)) | ~p(s(bool,X571)) | sP7(X572,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X569)))),X571) | p(s(bool,X572))) ) | $spl6),
  inference(subsumption_resolution,[],[f11040,f398])).
fof(f11040,plain,(
  ( ! [X569,X571,X572,X570] : (s(bool,'F_0') != s(bool,X570) | p(s(bool,X570)) | ~p(s(bool,X571)) | sP7(X572,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X569)))),X571) | p(s(bool,X572))) ) | $spl6),
  inference(superposition,[],[f1199,f3062])).
fof(f1199,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X1) != s(bool,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | sP7(X3,X1,X0) | p(s(bool,X3))) ) | $spl6),
  inference(superposition,[],[f528,f316])).
fof(f10387,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,X17) | s(bool,X16) != s(bool,X17) | s(bool,'F_0') = s(bool,X16)) ) | $spl90),
  inference(cnf_transformation,[],[f10387_D])).
fof(f10387_D,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,X17) | s(bool,X16) != s(bool,X17) | s(bool,'F_0') = s(bool,X16)) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f10386,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | s(bool,'F_0') = s(bool,X16) | s(bool,'F_0') != s(bool,X17) | s(bool,X16) != s(bool,X17)) ) | ($spl54 | $spl86)),
  inference(duplicate_literal_removal,[],[f10348])).
fof(f10348,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | s(bool,'F_0') = s(bool,X16) | s(bool,'F_0') != s(bool,X17) | p(s(bool,X15)) | s(bool,X16) != s(bool,X17)) ) | ($spl54 | $spl86)),
  inference(resolution,[],[f9433,f319])).
fof(f9433,plain,(
  ( ! [X28,X29,X27] : (sP7(X27,X28,X29) | p(s(bool,X27)) | s(bool,'F_0') = s(bool,X28) | s(bool,'F_0') != s(bool,X29)) ) | ($spl54 | $spl86)),
  inference(resolution,[],[f8958,f9397])).
fof(f9397,plain,(
  ( ! [X4,X5,X3] : (sP5(X5,X4,X3) | s(bool,'F_0') != s(bool,X4)) ) | $spl86),
  inference(subsumption_resolution,[],[f9396,f309])).
fof(f9396,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | s(bool,'F_0') != s(bool,X4) | sP5(X5,X4,X3)) ) | $spl86),
  inference(subsumption_resolution,[],[f9353,f311])).
fof(f9353,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | s(bool,'F_0') != s(bool,X4) | ~p(s(bool,X5)) | sP5(X5,X4,X3)) ) | $spl86),
  inference(resolution,[],[f8673,f298])).
fof(f8958,plain,(
  ( ! [X6,X8,X7] : (~sP5(X6,X8,X7) | sP7(X7,X6,X8) | p(s(bool,X7)) | s(bool,'F_0') = s(bool,X6)) ) | $spl54),
  inference(subsumption_resolution,[],[f8946,f307])).
fof(f8946,plain,(
  ( ! [X6,X8,X7] : (s(bool,'F_0') = s(bool,X6) | sP7(X7,X6,X8) | p(s(bool,X7)) | ~sP6(X7,X8,X6) | ~sP5(X6,X8,X7)) ) | $spl54),
  inference(resolution,[],[f8937,f303])).
fof(f8937,plain,(
  ( ! [X2,X0,X1] : (sP4(X2,X0,X1) | s(bool,'F_0') = s(bool,X0)) ) | $spl54),
  inference(subsumption_resolution,[],[f8936,f313])).
fof(f8936,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X1)) | sP4(X2,X0,X1)) ) | $spl54),
  inference(subsumption_resolution,[],[f8892,f315])).
fof(f8892,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | sP4(X2,X0,X1)) ) | $spl54),
  inference(resolution,[],[f7575,f297])).
fof(f4193,plain,(
  ( ! [X14] : (s(bool,'T_0') != s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X14)))))) ) | $spl38),
  inference(cnf_transformation,[],[f4193_D])).
fof(f4193_D,plain,(
  ( ! [X14] : (s(bool,'T_0') != s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X14)))))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f10268,plain,(
  ( ! [X21,X22,X20] : (~p(s(bool,X20)) | s(bool,'F_0') != s(bool,X21) | s(bool,'T_0') != s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X22)))))) ) | $spl24),
  inference(resolution,[],[f7489,f3032])).
fof(f3032,plain,(
  ( ! [X138,X139,X137] : (sP6(X137,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X138)))),X139)) )),
  inference(resolution,[],[f539,f306])).
fof(f10258,plain,(
  ( ! [X15,X16] : (s(bool,X15) != s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X16))))) | s(bool,'T_0') != s(bool,X15)) ) | $spl88),
  inference(cnf_transformation,[],[f10258_D])).
fof(f10258_D,plain,(
  ( ! [X15,X16] : (s(bool,X15) != s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X16))))) | s(bool,'T_0') != s(bool,X15)) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f10253,plain,(
  ( ! [X14,X15,X16] : (~p(s(bool,X14)) | s(bool,X15) != s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X16))))) | s(bool,'T_0') != s(bool,X15)) ) | $spl24),
  inference(resolution,[],[f6588,f3032])).
fof(f8174,plain,(
  ( ! [X19,X17,X16] : (s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X19) | ~p(s(bool,X17)) | p(s(bool,X16))) ) | $spl78),
  inference(cnf_transformation,[],[f8174_D])).
fof(f8174_D,plain,(
  ( ! [X19,X17,X16] : (s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X19) | ~p(s(bool,X17)) | p(s(bool,X16))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f9129,plain,(
  ( ! [X19,X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X16) != s(bool,X19) | s(bool,X17) != s(bool,X19)) ) | $spl6),
  inference(duplicate_literal_removal,[],[f9084])).
fof(f9084,plain,(
  ( ! [X19,X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X16) != s(bool,X19) | s(bool,X17) != s(bool,X19) | ~p(s(bool,X18))) ) | $spl6),
  inference(resolution,[],[f1202,f317])).
fof(f1202,plain,(
  ( ! [X14,X12,X15,X13] : (~sP7(X15,X12,X13) | p(s(bool,X14)) | ~p(s(bool,X12)) | ~p(s(bool,X15)) | s(bool,X13) != s(bool,X14)) ) | $spl6),
  inference(superposition,[],[f528,f318])).
fof(f8875,plain,(
  ( ! [X19,X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X16) != s(bool,X19) | s(bool,X17) != s(bool,X19)) ) | $spl6),
  inference(duplicate_literal_removal,[],[f8831])).
fof(f8831,plain,(
  ( ! [X19,X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X16) != s(bool,X19) | s(bool,X17) != s(bool,X19) | ~p(s(bool,X18))) ) | $spl6),
  inference(resolution,[],[f1201,f317])).
fof(f8673,plain,(
  ( ! [X47,X45,X46] : (sP7(X47,X46,X45) | p(s(bool,X47)) | s(bool,'F_0') != s(bool,X45) | ~p(s(bool,X46))) ) | $spl86),
  inference(cnf_transformation,[],[f8673_D])).
fof(f8673_D,plain,(
  ( ! [X47,X45,X46] : (sP7(X47,X46,X45) | p(s(bool,X47)) | s(bool,'F_0') != s(bool,X45) | ~p(s(bool,X46))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f8675,plain,(
  ( ! [X66,X64,X67,X65] : (s(bool,'F_0') != s(bool,X65) | p(s(bool,X64)) | ~p(s(bool,X66)) | sP7(X67,X66,X65) | p(s(bool,X67))) ) | $spl6),
  inference(subsumption_resolution,[],[f8648,f401])).
fof(f8648,plain,(
  ( ! [X66,X64,X67,X65] : (s(bool,'F_0') != s(bool,X65) | p(s(bool,X64)) | ~p(s(bool,X66)) | sP7(X67,X66,X65) | p(s(bool,X67)) | s(bool,'T_0') = s(bool,X64)) ) | $spl6),
  inference(superposition,[],[f1200,f424])).
fof(f8661,plain,(
  ( ! [X50,X48,X51,X49] : (s(bool,'F_0') != s(bool,X49) | p(s(bool,X48)) | ~p(s(bool,X50)) | sP7(X51,X50,X49) | p(s(bool,X51))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f8644])).
fof(f8644,plain,(
  ( ! [X50,X48,X51,X49] : (s(bool,'F_0') != s(bool,X49) | p(s(bool,X48)) | ~p(s(bool,X50)) | sP7(X51,X50,X49) | p(s(bool,X51)) | p(s(bool,X48))) ) | $spl6),
  inference(superposition,[],[f1200,f400])).
fof(f8662,plain,(
  ( ! [X47,X45,X46,X44] : (s(bool,'F_0') != s(bool,X45) | p(s(bool,X44)) | ~p(s(bool,X46)) | sP7(X47,X46,X45) | p(s(bool,X47))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f8643])).
fof(f8643,plain,(
  ( ! [X47,X45,X46,X44] : (s(bool,'F_0') != s(bool,X45) | p(s(bool,X44)) | ~p(s(bool,X46)) | sP7(X47,X46,X45) | p(s(bool,X47)) | p(s(bool,X44))) ) | $spl6),
  inference(superposition,[],[f1200,f398])).
fof(f8670,plain,(
  ( ! [X83,X82] : (sP7(X83,X82,'F_0') | p(s(bool,X83)) | ~p(s(bool,X82))) ) | $spl84),
  inference(cnf_transformation,[],[f8670_D])).
fof(f8670_D,plain,(
  ( ! [X83,X82] : (sP7(X83,X82,'F_0') | p(s(bool,X83)) | ~p(s(bool,X82))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f8672,plain,(
  ( ! [X90,X88,X91,X89] : (s(bool,X88) != s(bool,X89) | p(s(bool,X89)) | ~p(s(bool,X90)) | sP7(X91,X90,'F_0') | p(s(bool,X91))) ) | $spl6),
  inference(subsumption_resolution,[],[f8629,f528])).
fof(f8629,plain,(
  ( ! [X90,X88,X91,X89] : (s(bool,X88) != s(bool,X89) | p(s(bool,X89)) | ~p(s(bool,X90)) | sP7(X91,X90,'F_0') | p(s(bool,X91)) | p(s(bool,X88))) ) | $spl6),
  inference(superposition,[],[f1200,f398])).
fof(f8671,plain,(
  ( ! [X87,X85,X86,X84] : (s(bool,X84) != s(bool,X85) | p(s(bool,X85)) | ~p(s(bool,X86)) | sP7(X87,X86,'F_0') | p(s(bool,X87))) ) | $spl6),
  inference(subsumption_resolution,[],[f8628,f528])).
fof(f8628,plain,(
  ( ! [X87,X85,X86,X84] : (s(bool,X84) != s(bool,X85) | p(s(bool,X85)) | ~p(s(bool,X86)) | sP7(X87,X86,'F_0') | p(s(bool,X87)) | p(s(bool,X84))) ) | $spl6),
  inference(superposition,[],[f1200,f400])).
fof(f8669,plain,(
  ( ! [X80,X83,X81,X82] : (s(bool,X80) != s(bool,X81) | p(s(bool,X81)) | ~p(s(bool,X82)) | sP7(X83,X82,'F_0') | p(s(bool,X83))) ) | ($spl6 | $spl24)),
  inference(subsumption_resolution,[],[f8627,f1495])).
fof(f8627,plain,(
  ( ! [X80,X83,X81,X82] : (s(bool,X80) != s(bool,X81) | p(s(bool,X81)) | ~p(s(bool,X82)) | sP7(X83,X82,'F_0') | p(s(bool,X83)) | s(bool,'T_0') = s(bool,X80)) ) | $spl6),
  inference(superposition,[],[f1200,f424])).
fof(f8668,plain,(
  ( ! [X66,X64,X67] : (sP7(X67,X66,X64) | s(bool,'T_0') = s(bool,X64) | p(s(bool,X67)) | ~p(s(bool,X66))) ) | $spl82),
  inference(cnf_transformation,[],[f8668_D])).
fof(f8668_D,plain,(
  ( ! [X66,X64,X67] : (sP7(X67,X66,X64) | s(bool,'T_0') = s(bool,X64) | p(s(bool,X67)) | ~p(s(bool,X66))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f8667,plain,(
  ( ! [X66,X64,X67,X65] : (p(s(bool,X65)) | ~p(s(bool,X66)) | sP7(X67,X66,X64) | p(s(bool,X67)) | s(bool,'T_0') = s(bool,X64)) ) | $spl6),
  inference(subsumption_resolution,[],[f8623,f398])).
fof(f8623,plain,(
  ( ! [X66,X64,X67,X65] : (s(bool,'F_0') != s(bool,X65) | p(s(bool,X65)) | ~p(s(bool,X66)) | sP7(X67,X66,X64) | p(s(bool,X67)) | s(bool,'T_0') = s(bool,X64)) ) | $spl6),
  inference(superposition,[],[f1200,f424])).
fof(f8475,plain,(
  ( ! [X66,X64,X67] : (sP7(X67,X64,X66) | s(bool,'T_0') = s(bool,X64) | p(s(bool,X67)) | ~p(s(bool,X66))) ) | $spl80),
  inference(cnf_transformation,[],[f8475_D])).
fof(f8475_D,plain,(
  ( ! [X66,X64,X67] : (sP7(X67,X64,X66) | s(bool,'T_0') = s(bool,X64) | p(s(bool,X67)) | ~p(s(bool,X66))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f8474,plain,(
  ( ! [X66,X64,X67,X65] : (p(s(bool,X65)) | ~p(s(bool,X66)) | sP7(X67,X64,X66) | p(s(bool,X67)) | s(bool,'T_0') = s(bool,X64)) ) | $spl6),
  inference(subsumption_resolution,[],[f8430,f398])).
fof(f8430,plain,(
  ( ! [X66,X64,X67,X65] : (s(bool,'F_0') != s(bool,X65) | p(s(bool,X65)) | ~p(s(bool,X66)) | sP7(X67,X64,X66) | p(s(bool,X67)) | s(bool,'T_0') = s(bool,X64)) ) | $spl6),
  inference(superposition,[],[f1199,f424])).
fof(f8283,plain,(
  ( ! [X19,X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X19)) ) | $spl6),
  inference(duplicate_literal_removal,[],[f8249])).
fof(f8249,plain,(
  ( ! [X19,X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X19) | ~p(s(bool,X18))) ) | $spl6),
  inference(resolution,[],[f1181,f317])).
fof(f8172,plain,(
  ( ! [X19,X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X19)) ) | $spl6),
  inference(duplicate_literal_removal,[],[f8140])).
fof(f8140,plain,(
  ( ! [X19,X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X19) | ~p(s(bool,X18))) ) | $spl6),
  inference(resolution,[],[f1180,f317])).
fof(f1180,plain,(
  ( ! [X10,X8,X11,X9] : (~sP7(X11,X9,X8) | p(s(bool,X8)) | ~p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X9) != s(bool,X10)) ) | $spl6),
  inference(superposition,[],[f528,f318])).
fof(f8098,plain,(
  ( ! [X18] : (p(s(bool,X18)) | sP7(X18,'T_0','F_0')) ) | $spl76),
  inference(cnf_transformation,[],[f8098_D])).
fof(f8098_D,plain,(
  ( ! [X18] : (p(s(bool,X18)) | sP7(X18,'T_0','F_0')) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f8093,plain,(
  ( ! [X19,X18] : (p(s(bool,X18)) | sP7(X18,'T_0','F_0') | s(bool,'T_0') != s(bool,X19)) ) | ($spl24 | $spl56)),
  inference(duplicate_literal_removal,[],[f8088])).
fof(f8088,plain,(
  ( ! [X19,X18] : (p(s(bool,X18)) | sP7(X18,'T_0','F_0') | s(bool,'T_0') != s(bool,X19) | s(bool,'T_0') != s(bool,X19)) ) | ($spl24 | $spl56)),
  inference(resolution,[],[f7934,f6377])).
fof(f6377,plain,(
  ( ! [X14,X12,X13] : (sP5(X13,'F_0',X14) | s(bool,X12) != s(bool,X13) | s(bool,'T_0') != s(bool,X12)) ) | $spl24),
  inference(subsumption_resolution,[],[f6353,f309])).
fof(f6353,plain,(
  ( ! [X14,X12,X13] : (s(bool,'T_0') != s(bool,X12) | s(bool,X12) != s(bool,X13) | p(s(bool,X14)) | sP5(X13,'F_0',X14)) ) | $spl24),
  inference(resolution,[],[f1497,f298])).
fof(f1497,plain,(
  ( ! [X6,X8,X7] : (sP7(X8,X6,'F_0') | s(bool,'T_0') != s(bool,X7) | s(bool,X6) != s(bool,X7) | p(s(bool,X8))) ) | $spl24),
  inference(superposition,[],[f1480,f316])).
fof(f7934,plain,(
  ( ! [X4,X5] : (~sP5('T_0',X5,X4) | p(s(bool,X4)) | sP7(X4,'T_0',X5)) ) | $spl56),
  inference(subsumption_resolution,[],[f7926,f307])).
fof(f7926,plain,(
  ( ! [X4,X5] : (sP7(X4,'T_0',X5) | p(s(bool,X4)) | ~sP6(X4,X5,'T_0') | ~sP5('T_0',X5,X4)) ) | $spl56),
  inference(resolution,[],[f7917,f303])).
fof(f7917,plain,(
  ( ! [X0,X1] : (sP4(X1,'T_0',X0)) ) | $spl56),
  inference(subsumption_resolution,[],[f7916,f313])).
fof(f7916,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | sP4(X1,'T_0',X0)) ) | $spl56),
  inference(subsumption_resolution,[],[f7890,f315])).
fof(f7890,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | sP4(X1,'T_0',X0)) ) | $spl56),
  inference(resolution,[],[f7577,f297])).
fof(f8007,plain,(
  ( ! [X103,X102] : (p(s(bool,X102)) | p(s(bool,X103)) | sP7(X103,X102,happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))))) ) | $spl74),
  inference(cnf_transformation,[],[f8007_D])).
fof(f8007_D,plain,(
  ( ! [X103,X102] : (p(s(bool,X102)) | p(s(bool,X103)) | sP7(X103,X102,happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f8006,plain,(
  ( ! [X103,X101,X102] : (s(bool,'T_0') != s(bool,X101) | p(s(bool,X102)) | sP7(X103,X102,happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13))))) | p(s(bool,X103))) ) | $spl6),
  inference(subsumption_resolution,[],[f7958,f401])).
fof(f7958,plain,(
  ( ! [X103,X101,X102] : (s(bool,'T_0') != s(bool,X101) | p(s(bool,X102)) | ~p(s(bool,X101)) | sP7(X103,X102,happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13))))) | p(s(bool,X103))) ) | $spl6),
  inference(superposition,[],[f1179,f2643])).
fof(f1179,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X5) != s(bool,X6) | p(s(bool,X4)) | ~p(s(bool,X6)) | sP7(X7,X4,X5) | p(s(bool,X7))) ) | $spl6),
  inference(superposition,[],[f528,f316])).
fof(f8005,plain,(
  ( ! [X99,X100] : (p(s(bool,X99)) | p(s(bool,X100)) | sP7(X100,X99,happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))))) ) | $spl72),
  inference(cnf_transformation,[],[f8005_D])).
fof(f8005_D,plain,(
  ( ! [X99,X100] : (p(s(bool,X99)) | p(s(bool,X100)) | sP7(X100,X99,happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f8004,plain,(
  ( ! [X99,X100,X98] : (s(bool,'T_0') != s(bool,X98) | p(s(bool,X99)) | sP7(X100,X99,happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11)))) | p(s(bool,X100))) ) | $spl6),
  inference(subsumption_resolution,[],[f7957,f401])).
fof(f7957,plain,(
  ( ! [X99,X100,X98] : (s(bool,'T_0') != s(bool,X98) | p(s(bool,X99)) | ~p(s(bool,X98)) | sP7(X100,X99,happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11)))) | p(s(bool,X100))) ) | $spl6),
  inference(superposition,[],[f1179,f1321])).
fof(f1321,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))))),
  inference(resolution,[],[f538,f402])).
fof(f7999,plain,(
  ( ! [X70,X71] : (sP7(X71,X70,'T_0') | p(s(bool,X71)) | p(s(bool,X70))) ) | $spl70),
  inference(cnf_transformation,[],[f7999_D])).
fof(f7999_D,plain,(
  ( ! [X70,X71] : (sP7(X71,X70,'T_0') | p(s(bool,X71)) | p(s(bool,X70))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f8001,plain,(
  ( ! [X78,X76,X79,X77] : (s(bool,X76) != s(bool,X77) | p(s(bool,X78)) | sP7(X79,X78,'T_0') | p(s(bool,X79)) | ~p(s(bool,X76))) ) | $spl6),
  inference(subsumption_resolution,[],[f7951,f528])).
fof(f7951,plain,(
  ( ! [X78,X76,X79,X77] : (s(bool,X76) != s(bool,X77) | p(s(bool,X78)) | ~p(s(bool,X77)) | sP7(X79,X78,'T_0') | p(s(bool,X79)) | ~p(s(bool,X76))) ) | $spl6),
  inference(superposition,[],[f1179,f402])).
fof(f8000,plain,(
  ( ! [X74,X72,X75,X73] : (s(bool,X72) != s(bool,X73) | p(s(bool,X74)) | sP7(X75,X74,'T_0') | p(s(bool,X75)) | ~p(s(bool,X72))) ) | $spl6),
  inference(subsumption_resolution,[],[f7950,f528])).
fof(f7950,plain,(
  ( ! [X74,X72,X75,X73] : (s(bool,X72) != s(bool,X73) | p(s(bool,X74)) | ~p(s(bool,X73)) | sP7(X75,X74,'T_0') | p(s(bool,X75)) | ~p(s(bool,X72))) ) | $spl6),
  inference(superposition,[],[f1179,f404])).
fof(f7578,plain,(
  ( ! [X68,X69] : (s(bool,X68) != s(bool,X69) | s(bool,'F_0') = s(bool,X68)) ) | $spl58),
  inference(cnf_transformation,[],[f7578_D])).
fof(f7578_D,plain,(
  ( ! [X68,X69] : (s(bool,X68) != s(bool,X69) | s(bool,'F_0') = s(bool,X68)) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f7998,plain,(
  ( ! [X70,X68,X71,X69] : (s(bool,X68) != s(bool,X69) | p(s(bool,X70)) | sP7(X71,X70,'T_0') | p(s(bool,X71)) | s(bool,'F_0') = s(bool,X68)) ) | $spl6),
  inference(subsumption_resolution,[],[f7949,f1243])).
fof(f1243,plain,(
  ( ! [X47,X46] : (s(bool,X46) != s(bool,X47) | p(s(bool,X47)) | s(bool,'F_0') = s(bool,X46)) ) | $spl6),
  inference(subsumption_resolution,[],[f1216,f370])).
fof(f1216,plain,(
  ( ! [X47,X46] : (s(bool,X46) != s(bool,X47) | p(s(bool,X47)) | ~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,X46)) ) | $spl6),
  inference(superposition,[],[f528,f424])).
fof(f7949,plain,(
  ( ! [X70,X68,X71,X69] : (s(bool,X68) != s(bool,X69) | p(s(bool,X70)) | ~p(s(bool,X69)) | sP7(X71,X70,'T_0') | p(s(bool,X71)) | s(bool,'F_0') = s(bool,X68)) ) | $spl6),
  inference(superposition,[],[f1179,f424])).
fof(f7997,plain,(
  ( ! [X62,X60,X63] : (sP7(X63,X62,X60) | s(bool,'F_0') = s(bool,X60) | p(s(bool,X63)) | p(s(bool,X62))) ) | $spl68),
  inference(cnf_transformation,[],[f7997_D])).
fof(f7997_D,plain,(
  ( ! [X62,X60,X63] : (sP7(X63,X62,X60) | s(bool,'F_0') = s(bool,X60) | p(s(bool,X63)) | p(s(bool,X62))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f7996,plain,(
  ( ! [X61,X62,X60,X63] : (s(bool,'T_0') != s(bool,X61) | p(s(bool,X62)) | sP7(X63,X62,X60) | p(s(bool,X63)) | s(bool,'F_0') = s(bool,X60)) ) | $spl6),
  inference(subsumption_resolution,[],[f7947,f401])).
fof(f7947,plain,(
  ( ! [X61,X62,X60,X63] : (s(bool,'T_0') != s(bool,X61) | p(s(bool,X62)) | ~p(s(bool,X61)) | sP7(X63,X62,X60) | p(s(bool,X63)) | s(bool,'F_0') = s(bool,X60)) ) | $spl6),
  inference(superposition,[],[f1179,f424])).
fof(f7881,plain,(
  $false | $spl20),
  inference(trivial_inequality_removal,[],[f7873])).
fof(f7873,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl20),
  inference(superposition,[],[f838,f504])).
fof(f7880,plain,(
  $false | $spl20),
  inference(trivial_inequality_removal,[],[f7875])).
fof(f7875,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl20),
  inference(superposition,[],[f838,f1321])).
fof(f7879,plain,(
  $false | $spl20),
  inference(trivial_inequality_removal,[],[f7876])).
fof(f7876,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl20),
  inference(superposition,[],[f838,f2643])).
fof(f7878,plain,(
  $false | $spl20),
  inference(equality_resolution,[],[f838])).
fof(f7830,plain,(
  $false | $spl18),
  inference(trivial_inequality_removal,[],[f7825])).
fof(f7825,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl18),
  inference(superposition,[],[f836,f613])).
fof(f7829,plain,(
  $false | $spl18),
  inference(equality_resolution,[],[f836])).
fof(f7789,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7714,f885])).
fof(f885,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(equality_resolution,[],[f529])).
fof(f7714,plain,(
  ( ! [X249,X246,X248,X247] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X246,X247),X248),s(fun(X247,bool),X249))))) ) | $spl8),
  inference(resolution,[],[f885,f420])).
fof(f420,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | p(s(bool,happ(s(fun(X0,bool),X3),s(X0,sK17(X3,X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f248])).
fof(f248,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X6] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | s(X0,X6) = s(X0,happ(s(fun(X1,X0),X2),s(X1,sK16(X6,X3,X2,X1,X0)))))) & ((p(s(bool,happ(s(fun(X0,bool),X3),s(X0,sK17(X3,X2,X1,X0))))) & ! [X5] : s(X0,sK17(X3,X2,X1,X0)) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK16,sK17])],[f247])).
fof(f247,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X6] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | ? [X7] : s(X0,X6) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))))) & (? [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) & ! [X5] : s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f246])).
fof(f246,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) | ? [X5] : s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))) & (? [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) & ! [X5] : s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(nnf_transformation,[],[f151])).
fof(f151,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) | ? [X5] : s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))),
  inference(ennf_transformation,[],[f114])).
fof(f114,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) => ? [X5] : s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))),
  inference(flattening,[],[f113])).
fof(f113,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) => ? [X5] : s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X7,X6] : ! [X20,X21] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X6,X7),X20),s(fun(X7,bool),X21)))) <=> ! [X18] : (p(s(bool,happ(s(fun(X7,bool),X21),s(X7,X18)))) => ? [X19] : s(X7,X18) = s(X7,happ(s(fun(X6,X7),X20),s(X6,X19)))))),
  file('thm.quantHeuristics.GUESSES_WEAKEN_THM|split|3','thm.quantHeuristics.GUESS_REWRITES|split|4')).
fof(f7788,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7713,f885])).
fof(f7713,plain,(
  ( ! [X243,X244,X242,X245] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X242,X243),X244),s(fun(X243,bool),X245))))) ) | $spl8),
  inference(resolution,[],[f885,f253])).
fof(f7787,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7709,f885])).
fof(f7709,plain,(
  ( ! [X227,X225,X228,X226] : (p(s(bool,happ(s(fun(X225,bool),X226),s(X225,sK15(X226,X227,X228,X225)))))) ) | $spl8),
  inference(resolution,[],[f885,f253])).
fof(f7786,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7707,f885])).
fof(f7707,plain,(
  ( ! [X218,X216,X219,X217] : (p(s(bool,happ(s(fun(X216,bool),X217),s(X216,sK17(X217,X218,X219,X216)))))) ) | $spl8),
  inference(resolution,[],[f885,f420])).
fof(f7781,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7698,f885])).
fof(f7698,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(resolution,[],[f885,f415])).
fof(f415,plain,(
  p(s(bool,'F_0')) | p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f245])).
fof(f245,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f112])).
fof(f112,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f25])).
fof(f25,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.quantHeuristics.GUESSES_WEAKEN_THM|split|3','thm.bool.NOT_CLAUSES|split|1')).
fof(f7780,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7697,f885])).
fof(f7697,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(resolution,[],[f885,f413])).
fof(f413,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f244])).
fof(f244,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f111])).
fof(f111,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f26])).
fof(f26,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.quantHeuristics.GUESSES_WEAKEN_THM|split|3','thm.bool.NOT_CLAUSES|split|2')).
fof(f7779,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7693,f885])).
fof(f7693,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(resolution,[],[f885,f379])).
fof(f379,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f229])).
fof(f7778,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7691,f885])).
fof(f7691,plain,(
  ( ! [X205] : (p(s(bool,X205))) ) | $spl8),
  inference(resolution,[],[f885,f376])).
fof(f7777,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7690,f885])).
fof(f7690,plain,(
  ( ! [X204] : (p(s(bool,X204))) ) | $spl8),
  inference(resolution,[],[f885,f373])).
fof(f7776,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7689,f885])).
fof(f7689,plain,(
  ( ! [X203] : (p(s(bool,X203))) ) | $spl8),
  inference(resolution,[],[f885,f370])).
fof(f7699,plain,(
  $false | $spl8),
  inference(resolution,[],[f885,f417])).
fof(f7700,plain,(
  $false | $spl8),
  inference(resolution,[],[f885,f418])).
fof(f890,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') != s(bool,X37) | ~p(s(bool,X36))) ) | $spl8),
  inference(subsumption_resolution,[],[f855,f529])).
fof(f855,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') != s(bool,X37) | ~p(s(bool,X36)) | s(bool,'F_0') = s(bool,X36)) ) | $spl8),
  inference(superposition,[],[f529,f424])).
fof(f891,plain,(
  ( ! [X39,X38] : (s(bool,'F_0') != s(bool,X39) | ~p(s(bool,X38))) ) | $spl8),
  inference(subsumption_resolution,[],[f856,f529])).
fof(f856,plain,(
  ( ! [X39,X38] : (s(bool,'F_0') != s(bool,X39) | ~p(s(bool,X38)) | s(bool,'T_0') = s(bool,X38)) ) | $spl8),
  inference(superposition,[],[f529,f424])).
fof(f886,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') != s(bool,X35) | ~p(s(bool,X34))) ) | $spl8),
  inference(duplicate_literal_removal,[],[f854])).
fof(f854,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') != s(bool,X35) | ~p(s(bool,X34)) | ~p(s(bool,X34))) ) | $spl8),
  inference(superposition,[],[f529,f404])).
fof(f887,plain,(
  ( ! [X33,X32] : (s(bool,'T_0') != s(bool,X33) | ~p(s(bool,X32))) ) | $spl8),
  inference(duplicate_literal_removal,[],[f853])).
fof(f853,plain,(
  ( ! [X33,X32] : (s(bool,'T_0') != s(bool,X33) | ~p(s(bool,X32)) | ~p(s(bool,X32))) ) | $spl8),
  inference(superposition,[],[f529,f402])).
fof(f7588,plain,(
  ( ! [X54,X55,X53] : (sP7(X55,X53,X54) | p(s(bool,X55)) | s(bool,'T_0') != s(bool,X53) | p(s(bool,X54))) ) | $spl64),
  inference(cnf_transformation,[],[f7588_D])).
fof(f7588_D,plain,(
  ( ! [X54,X55,X53] : (sP7(X55,X53,X54) | p(s(bool,X55)) | s(bool,'T_0') != s(bool,X53) | p(s(bool,X54))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f7589,plain,(
  ( ! [X61,X62,X60,X63] : (s(bool,'T_0') != s(bool,X61) | p(s(bool,X62)) | sP7(X63,X61,X62) | p(s(bool,X63)) | s(bool,'F_0') = s(bool,X60)) ) | $spl6),
  inference(subsumption_resolution,[],[f7549,f398])).
fof(f7549,plain,(
  ( ! [X61,X62,X60,X63] : (s(bool,'T_0') != s(bool,X61) | p(s(bool,X62)) | ~p(s(bool,X60)) | sP7(X63,X61,X62) | p(s(bool,X63)) | s(bool,'F_0') = s(bool,X60)) ) | $spl6),
  inference(superposition,[],[f1178,f424])).
fof(f1178,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X1) != s(bool,X2) | p(s(bool,X0)) | ~p(s(bool,X2)) | sP7(X3,X1,X0) | p(s(bool,X3))) ) | $spl6),
  inference(superposition,[],[f528,f316])).
fof(f7563,plain,(
  ( ! [X59,X57,X58,X56] : (s(bool,'T_0') != s(bool,X57) | p(s(bool,X58)) | ~p(s(bool,X56)) | sP7(X59,X57,X58) | p(s(bool,X59))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f7548])).
fof(f7548,plain,(
  ( ! [X59,X57,X58,X56] : (s(bool,'T_0') != s(bool,X57) | p(s(bool,X58)) | ~p(s(bool,X56)) | sP7(X59,X57,X58) | p(s(bool,X59)) | ~p(s(bool,X56))) ) | $spl6),
  inference(superposition,[],[f1178,f404])).
fof(f7564,plain,(
  ( ! [X54,X52,X55,X53] : (s(bool,'T_0') != s(bool,X53) | p(s(bool,X54)) | ~p(s(bool,X52)) | sP7(X55,X53,X54) | p(s(bool,X55))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f7547])).
fof(f7547,plain,(
  ( ! [X54,X52,X55,X53] : (s(bool,'T_0') != s(bool,X53) | p(s(bool,X54)) | ~p(s(bool,X52)) | sP7(X55,X53,X54) | p(s(bool,X55)) | ~p(s(bool,X52))) ) | $spl6),
  inference(superposition,[],[f1178,f402])).
fof(f7587,plain,(
  ( ! [X103,X102] : (p(s(bool,X102)) | p(s(bool,X103)) | sP7(X103,happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),X102)) ) | $spl62),
  inference(cnf_transformation,[],[f7587_D])).
fof(f7587_D,plain,(
  ( ! [X103,X102] : (p(s(bool,X102)) | p(s(bool,X103)) | sP7(X103,happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),X102)) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f7586,plain,(
  ( ! [X103,X101,X102] : (s(bool,'T_0') != s(bool,X101) | p(s(bool,X102)) | sP7(X103,happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),X102) | p(s(bool,X103))) ) | $spl6),
  inference(subsumption_resolution,[],[f7535,f401])).
fof(f7535,plain,(
  ( ! [X103,X101,X102] : (s(bool,'T_0') != s(bool,X101) | p(s(bool,X102)) | ~p(s(bool,X101)) | sP7(X103,happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))),X102) | p(s(bool,X103))) ) | $spl6),
  inference(superposition,[],[f1178,f2643])).
fof(f7585,plain,(
  ( ! [X99,X100] : (p(s(bool,X99)) | p(s(bool,X100)) | sP7(X100,happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),X99)) ) | $spl60),
  inference(cnf_transformation,[],[f7585_D])).
fof(f7585_D,plain,(
  ( ! [X99,X100] : (p(s(bool,X99)) | p(s(bool,X100)) | sP7(X100,happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),X99)) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f7584,plain,(
  ( ! [X99,X100,X98] : (s(bool,'T_0') != s(bool,X98) | p(s(bool,X99)) | sP7(X100,happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),X99) | p(s(bool,X100))) ) | $spl6),
  inference(subsumption_resolution,[],[f7534,f401])).
fof(f7534,plain,(
  ( ! [X99,X100,X98] : (s(bool,'T_0') != s(bool,X98) | p(s(bool,X99)) | ~p(s(bool,X98)) | sP7(X100,happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),X99) | p(s(bool,X100))) ) | $spl6),
  inference(superposition,[],[f1178,f1321])).
fof(f7577,plain,(
  ( ! [X70,X71] : (sP7(X71,'T_0',X70) | p(s(bool,X71)) | p(s(bool,X70))) ) | $spl56),
  inference(cnf_transformation,[],[f7577_D])).
fof(f7577_D,plain,(
  ( ! [X70,X71] : (sP7(X71,'T_0',X70) | p(s(bool,X71)) | p(s(bool,X70))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f7580,plain,(
  ( ! [X78,X76,X79,X77] : (s(bool,X76) != s(bool,X77) | p(s(bool,X78)) | sP7(X79,'T_0',X78) | p(s(bool,X79)) | ~p(s(bool,X76))) ) | $spl6),
  inference(subsumption_resolution,[],[f7528,f528])).
fof(f7528,plain,(
  ( ! [X78,X76,X79,X77] : (s(bool,X76) != s(bool,X77) | p(s(bool,X78)) | ~p(s(bool,X77)) | sP7(X79,'T_0',X78) | p(s(bool,X79)) | ~p(s(bool,X76))) ) | $spl6),
  inference(superposition,[],[f1178,f402])).
fof(f7579,plain,(
  ( ! [X74,X72,X75,X73] : (s(bool,X72) != s(bool,X73) | p(s(bool,X74)) | sP7(X75,'T_0',X74) | p(s(bool,X75)) | ~p(s(bool,X72))) ) | $spl6),
  inference(subsumption_resolution,[],[f7527,f528])).
fof(f7527,plain,(
  ( ! [X74,X72,X75,X73] : (s(bool,X72) != s(bool,X73) | p(s(bool,X74)) | ~p(s(bool,X73)) | sP7(X75,'T_0',X74) | p(s(bool,X75)) | ~p(s(bool,X72))) ) | $spl6),
  inference(superposition,[],[f1178,f404])).
fof(f7576,plain,(
  ( ! [X70,X68,X71,X69] : (s(bool,X68) != s(bool,X69) | p(s(bool,X70)) | sP7(X71,'T_0',X70) | p(s(bool,X71)) | s(bool,'F_0') = s(bool,X68)) ) | $spl6),
  inference(subsumption_resolution,[],[f7526,f1243])).
fof(f7526,plain,(
  ( ! [X70,X68,X71,X69] : (s(bool,X68) != s(bool,X69) | p(s(bool,X70)) | ~p(s(bool,X69)) | sP7(X71,'T_0',X70) | p(s(bool,X71)) | s(bool,'F_0') = s(bool,X68)) ) | $spl6),
  inference(superposition,[],[f1178,f424])).
fof(f7575,plain,(
  ( ! [X62,X60,X63] : (sP7(X63,X60,X62) | s(bool,'F_0') = s(bool,X60) | p(s(bool,X63)) | p(s(bool,X62))) ) | $spl54),
  inference(cnf_transformation,[],[f7575_D])).
fof(f7575_D,plain,(
  ( ! [X62,X60,X63] : (sP7(X63,X60,X62) | s(bool,'F_0') = s(bool,X60) | p(s(bool,X63)) | p(s(bool,X62))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f7574,plain,(
  ( ! [X61,X62,X60,X63] : (s(bool,'T_0') != s(bool,X61) | p(s(bool,X62)) | sP7(X63,X60,X62) | p(s(bool,X63)) | s(bool,'F_0') = s(bool,X60)) ) | $spl6),
  inference(subsumption_resolution,[],[f7524,f401])).
fof(f7524,plain,(
  ( ! [X61,X62,X60,X63] : (s(bool,'T_0') != s(bool,X61) | p(s(bool,X62)) | ~p(s(bool,X61)) | sP7(X63,X60,X62) | p(s(bool,X63)) | s(bool,'F_0') = s(bool,X60)) ) | $spl6),
  inference(superposition,[],[f1178,f424])).
fof(f7568,plain,(
  ( ! [X37,X38,X34] : (p(s(bool,X37)) | ~p(s(bool,X34)) | p(s(bool,X38)) | sP7(X38,X34,X37)) ) | $spl52),
  inference(cnf_transformation,[],[f7568_D])).
fof(f7568_D,plain,(
  ( ! [X37,X38,X34] : (p(s(bool,X37)) | ~p(s(bool,X34)) | p(s(bool,X38)) | sP7(X38,X34,X37)) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f7573,plain,(
  ( ! [X59,X57,X58,X56] : (s(bool,'T_0') != s(bool,X57) | p(s(bool,X58)) | sP7(X59,X56,X58) | p(s(bool,X59)) | ~p(s(bool,X56))) ) | $spl6),
  inference(subsumption_resolution,[],[f7523,f401])).
fof(f7523,plain,(
  ( ! [X59,X57,X58,X56] : (s(bool,'T_0') != s(bool,X57) | p(s(bool,X58)) | ~p(s(bool,X57)) | sP7(X59,X56,X58) | p(s(bool,X59)) | ~p(s(bool,X56))) ) | $spl6),
  inference(superposition,[],[f1178,f404])).
fof(f7572,plain,(
  ( ! [X54,X52,X55,X53] : (s(bool,'T_0') != s(bool,X53) | p(s(bool,X54)) | sP7(X55,X52,X54) | p(s(bool,X55)) | ~p(s(bool,X52))) ) | $spl6),
  inference(subsumption_resolution,[],[f7522,f401])).
fof(f7522,plain,(
  ( ! [X54,X52,X55,X53] : (s(bool,'T_0') != s(bool,X53) | p(s(bool,X54)) | ~p(s(bool,X53)) | sP7(X55,X52,X54) | p(s(bool,X55)) | ~p(s(bool,X52))) ) | $spl6),
  inference(superposition,[],[f1178,f402])).
fof(f7569,plain,(
  ( ! [X39,X43,X41,X42,X40] : (s(bool,X40) != s(bool,X41) | p(s(bool,X42)) | sP7(X43,X39,X42) | p(s(bool,X43)) | ~p(s(bool,X40)) | ~p(s(bool,X39))) ) | $spl6),
  inference(subsumption_resolution,[],[f7519,f528])).
fof(f7519,plain,(
  ( ! [X39,X43,X41,X42,X40] : (s(bool,X40) != s(bool,X41) | p(s(bool,X42)) | ~p(s(bool,X41)) | sP7(X43,X39,X42) | p(s(bool,X43)) | ~p(s(bool,X40)) | ~p(s(bool,X39))) ) | $spl6),
  inference(superposition,[],[f1178,f358])).
fof(f7567,plain,(
  ( ! [X37,X35,X38,X36,X34] : (s(bool,X35) != s(bool,X36) | p(s(bool,X37)) | sP7(X38,X34,X37) | p(s(bool,X38)) | ~p(s(bool,X34)) | ~p(s(bool,X35))) ) | $spl6),
  inference(subsumption_resolution,[],[f7518,f528])).
fof(f7518,plain,(
  ( ! [X37,X35,X38,X36,X34] : (s(bool,X35) != s(bool,X36) | p(s(bool,X37)) | ~p(s(bool,X36)) | sP7(X38,X34,X37) | p(s(bool,X38)) | ~p(s(bool,X34)) | ~p(s(bool,X35))) ) | $spl6),
  inference(superposition,[],[f1178,f358])).
fof(f7033,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') != s(bool,X4) | s(bool,X3) != s(bool,X4) | ~p(s(bool,X3))) ) | $spl50),
  inference(cnf_transformation,[],[f7033_D])).
fof(f7033_D,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') != s(bool,X4) | s(bool,X3) != s(bool,X4) | ~p(s(bool,X3))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f7032,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | s(bool,'F_0') != s(bool,X4) | p(s(bool,X5)) | s(bool,X3) != s(bool,X4)) ) | $spl24),
  inference(duplicate_literal_removal,[],[f6994])).
fof(f6994,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | s(bool,'F_0') != s(bool,X4) | p(s(bool,X5)) | p(s(bool,X5)) | s(bool,X3) != s(bool,X4)) ) | $spl24),
  inference(resolution,[],[f6696,f319])).
fof(f6696,plain,(
  ( ! [X6,X8,X7] : (sP7(X6,X7,X8) | ~p(s(bool,X8)) | s(bool,'F_0') != s(bool,X7) | p(s(bool,X6))) ) | $spl24),
  inference(resolution,[],[f6674,f307])).
fof(f6674,plain,(
  ( ! [X4,X5,X3] : (~sP6(X5,X4,X3) | sP7(X5,X3,X4) | ~p(s(bool,X4)) | s(bool,'F_0') != s(bool,X3)) ) | $spl24),
  inference(subsumption_resolution,[],[f6673,f402])).
fof(f6673,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X3) | s(bool,'T_0') != s(bool,X4) | sP7(X5,X3,X4) | ~p(s(bool,X4)) | ~sP6(X5,X4,X3)) ) | $spl24),
  inference(subsumption_resolution,[],[f6667,f6637])).
fof(f6667,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X3) | s(bool,'T_0') != s(bool,X4) | sP7(X5,X3,X4) | ~p(s(bool,X4)) | ~sP6(X5,X4,X3) | ~sP4(X4,X3,X5)) ) | $spl24),
  inference(resolution,[],[f6638,f301])).
fof(f6375,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X7) | s(bool,'T_0') != s(bool,X6) | s(bool,X6) != s(bool,X7)) ) | $spl48),
  inference(cnf_transformation,[],[f6375_D])).
fof(f6375_D,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X7) | s(bool,'T_0') != s(bool,X6) | s(bool,X6) != s(bool,X7)) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f6374,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X6) | s(bool,X6) != s(bool,X7) | p(s(bool,X8)) | s(bool,'F_0') != s(bool,X7)) ) | $spl24),
  inference(duplicate_literal_removal,[],[f6351])).
fof(f6351,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X6) | s(bool,X6) != s(bool,X7) | p(s(bool,X8)) | p(s(bool,X8)) | s(bool,'F_0') != s(bool,X7)) ) | $spl24),
  inference(resolution,[],[f1497,f319])).
fof(f6107,plain,(
  $false | $spl40),
  inference(subsumption_resolution,[],[f6083,f418])).
fof(f6083,plain,(
  ~p(s(bool,'T_0')) | $spl40),
  inference(superposition,[],[f5982,f2643])).
fof(f6106,plain,(
  $false | $spl40),
  inference(subsumption_resolution,[],[f6082,f418])).
fof(f6082,plain,(
  ~p(s(bool,'T_0')) | $spl40),
  inference(superposition,[],[f5982,f1321])).
fof(f6101,plain,(
  ( ! [X11,X9] : (~p(s(bool,X9)) | s(bool,X9) != s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,X11)))) ) | $spl46),
  inference(cnf_transformation,[],[f6101_D])).
fof(f6101_D,plain,(
  ( ! [X11,X9] : (~p(s(bool,X9)) | s(bool,X9) != s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,X11)))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f6079,plain,(
  ( ! [X208,X206,X207] : (p(s(bool,X206)) | ~p(s(bool,X207)) | s(bool,X207) != s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,X208)))) ) | $spl40),
  inference(resolution,[],[f5982,f426])).
fof(f6008,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X36) != s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,X37))) | s(bool,X35) != s(bool,X38)) ) | $spl40),
  inference(resolution,[],[f5982,f262])).
fof(f6006,plain,(
  ( ! [X30,X28,X31,X29] : (p(s(bool,X28)) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,X31)))) ) | $spl40),
  inference(resolution,[],[f5982,f261])).
fof(f6004,plain,(
  ( ! [X23,X21,X22,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X20) != s(bool,X22) | s(bool,X21) != s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,X23)))) ) | $spl40),
  inference(resolution,[],[f5982,f260])).
fof(f6003,plain,(
  ( ! [X19,X17,X18,X16] : (~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X17) != s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,X18))) | s(bool,X16) != s(bool,X19)) ) | $spl40),
  inference(resolution,[],[f5982,f258])).
fof(f6001,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,X11)))) ) | $spl40),
  inference(resolution,[],[f5982,f256])).
fof(f5985,plain,(
  ( ! [X10,X8] : (s(sK11,X8) = s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,sK16(X8,sK13,sK12,sK10,sK11)))) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,X8)))) ) | $spl44),
  inference(cnf_transformation,[],[f5985_D])).
fof(f5985_D,plain,(
  ( ! [X10,X8] : (s(sK11,X8) = s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,sK16(X8,sK13,sK12,sK10,sK11)))) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,X8)))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f5960,plain,(
  ( ! [X216,X215,X214] : (s(sK11,X214) = s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,sK16(X214,sK13,sK12,sK10,sK11)))) | p(s(bool,X215)) | ~p(s(bool,X216)) | s(bool,X216) != s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,X214)))) )),
  inference(resolution,[],[f431,f426])).
fof(f431,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,X0)))) | s(sK11,X0) = s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,sK16(X0,sK13,sK12,sK10,sK11))))) )),
  inference(resolution,[],[f250,f419])).
fof(f419,plain,(
  ( ! [X6,X2,X0,X3,X1] : (s(X0,X6) = s(X0,happ(s(fun(X1,X0),X2),s(X1,sK16(X6,X3,X2,X1,X0)))) | ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | ~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f248])).
fof(f5889,plain,(
  ( ! [X39,X37,X38,X36] : (s(sK11,X36) = s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,sK16(X36,sK13,sK12,sK10,sK11)))) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X38) != s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,X36))) | s(bool,X37) != s(bool,X39)) )),
  inference(resolution,[],[f431,f262])).
fof(f5887,plain,(
  ( ! [X30,X28,X31,X29] : (s(sK11,X28) = s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,sK16(X28,sK13,sK12,sK10,sK11)))) | p(s(bool,X29)) | ~p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X30) != s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,X28)))) )),
  inference(resolution,[],[f431,f261])).
fof(f5885,plain,(
  ( ! [X23,X21,X22,X20] : (s(sK11,X20) = s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,sK16(X20,sK13,sK12,sK10,sK11)))) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,X20)))) )),
  inference(resolution,[],[f431,f260])).
fof(f5884,plain,(
  ( ! [X19,X17,X18,X16] : (s(sK11,X16) = s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,sK16(X16,sK13,sK12,sK10,sK11)))) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X18) != s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,X16))) | s(bool,X17) != s(bool,X19)) )),
  inference(resolution,[],[f431,f258])).
fof(f5882,plain,(
  ( ! [X10,X8,X11,X9] : (s(sK11,X8) = s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,sK16(X8,sK13,sK12,sK10,sK11)))) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,X8)))) )),
  inference(resolution,[],[f431,f256])).
fof(f5982,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,X0))))) ) | $spl40),
  inference(cnf_transformation,[],[f5982_D])).
fof(f5982_D,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,X0))))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f5984,plain,(
  s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13))) = s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,sK16('const.min.@_1'(s(fun(sK11,bool),sK13)),sK13,sK12,sK10,sK11)))) | $spl42),
  inference(cnf_transformation,[],[f5984_D])).
fof(f5984_D,plain,(
  s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13))) = s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,sK16('const.min.@_1'(s(fun(sK11,bool),sK13)),sK13,sK12,sK10,sK11)))) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f5878,plain,(
  ( ! [X0] : (s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13))) = s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,sK16('const.min.@_1'(s(fun(sK11,bool),sK13)),sK13,sK12,sK10,sK11)))) | ~p(s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,X0))))) )),
  inference(resolution,[],[f431,f425])).
fof(f4184,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') != s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X14))))) | ~p(s(bool,X15))) )),
  inference(resolution,[],[f3032,f2088])).
fof(f3088,plain,(
  ( ! [X4,X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X4)))))) ) | $spl36),
  inference(cnf_transformation,[],[f3088_D])).
fof(f3088_D,plain,(
  ( ! [X4,X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X4)))))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f3067,plain,(
  ( ! [X200,X199,X201] : (p(s(bool,X199)) | ~p(s(bool,X200)) | s(bool,X200) != s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X201)))))) )),
  inference(resolution,[],[f539,f426])).
fof(f2996,plain,(
  ( ! [X30,X28,X31,X29] : (~p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X29) != s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X30))))) | s(bool,X28) != s(bool,X31)) )),
  inference(resolution,[],[f539,f262])).
fof(f2994,plain,(
  ( ! [X24,X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X24)))))) )),
  inference(resolution,[],[f539,f261])).
fof(f2992,plain,(
  ( ! [X14,X15,X13,X16] : (p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X16)))))) )),
  inference(resolution,[],[f539,f260])).
fof(f2991,plain,(
  ( ! [X12,X10,X11,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X11))))) | s(bool,X9) != s(bool,X12)) )),
  inference(resolution,[],[f539,f258])).
fof(f2989,plain,(
  ( ! [X4,X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X4)))))) )),
  inference(resolution,[],[f539,f256])).
fof(f2786,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13))))) | $spl35),
  inference(cnf_transformation,[],[f2786_D])).
fof(f2786_D,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13))))) <=> ~$spl35),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl35])])).
fof(f2780,plain,(
  ( ! [X8] : (p(s(bool,X8)) | s(bool,'F_0') != s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))))) )),
  inference(resolution,[],[f2619,f1971])).
fof(f2664,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))))) ) | $spl32),
  inference(cnf_transformation,[],[f2664_D])).
fof(f2664_D,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f2645,plain,(
  ( ! [X128,X129] : (p(s(bool,X128)) | p(s(bool,X129)) | s(bool,X129) != s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))))) )),
  inference(resolution,[],[f1246,f426])).
fof(f2578,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13))))) | s(bool,X26) != s(bool,X27)) )),
  inference(resolution,[],[f1246,f262])).
fof(f2575,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))))) )),
  inference(resolution,[],[f1246,f261])).
fof(f2574,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))))) )),
  inference(resolution,[],[f1246,f260])).
fof(f2573,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13))))) | s(bool,X11) != s(bool,X12)) )),
  inference(resolution,[],[f1246,f258])).
fof(f2571,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))))) )),
  inference(resolution,[],[f1246,f256])).
fof(f2119,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11)))) | $spl31),
  inference(cnf_transformation,[],[f2119_D])).
fof(f2119_D,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11)))) <=> ~$spl31),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl31])])).
fof(f2113,plain,(
  ( ! [X14] : (p(s(bool,X14)) | s(bool,'F_0') != s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))))) )),
  inference(resolution,[],[f1971,f1297])).
fof(f1297,plain,(
  ( ! [X109,X108] : (sP4(happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))),X108,X109)) )),
  inference(resolution,[],[f538,f315])).
fof(f2096,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | $spl27),
  inference(cnf_transformation,[],[f2096_D])).
fof(f2096_D,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) <=> ~$spl27),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl27])])).
fof(f2103,plain,(
  ( ! [X8] : (sP7(X8,'F_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) ) | $spl28),
  inference(cnf_transformation,[],[f2103_D])).
fof(f2103_D,plain,(
  ( ! [X8] : (sP7(X8,'F_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f2101,plain,(
  ( ! [X8] : (sP7(X8,'F_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) )),
  inference(resolution,[],[f1554,f583])).
fof(f583,plain,(
  ( ! [X95,X96] : (sP6(X95,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X96)) )),
  inference(resolution,[],[f251,f306])).
fof(f1554,plain,(
  ( ! [X4,X5] : (~sP6(X5,X4,'F_0') | sP7(X5,'F_0',X4) | s(bool,'T_0') != s(bool,X4)) )),
  inference(subsumption_resolution,[],[f1553,f412])).
fof(f1553,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | sP7(X5,'F_0',X4) | p(s(bool,'F_0')) | ~sP6(X5,X4,'F_0')) )),
  inference(subsumption_resolution,[],[f1550,f311])).
fof(f1550,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | sP7(X5,'F_0',X4) | p(s(bool,'F_0')) | ~sP6(X5,X4,'F_0') | ~sP5('F_0',X4,X5)) )),
  inference(resolution,[],[f1546,f302])).
fof(f2093,plain,(
  ( ! [X8] : (s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | ~p(s(bool,X8))) )),
  inference(resolution,[],[f2088,f583])).
fof(f2027,plain,(
  $false | $spl4),
  inference(trivial_inequality_removal,[],[f2009])).
fof(f2009,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl4),
  inference(superposition,[],[f527,f504])).
fof(f2026,plain,(
  $false | $spl4),
  inference(equality_resolution,[],[f527])).
fof(f1480,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | s(bool,'T_0') != s(bool,X2)) ) | $spl24),
  inference(cnf_transformation,[],[f1480_D])).
fof(f1480_D,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | s(bool,'T_0') != s(bool,X2)) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f1479,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) )),
  inference(duplicate_literal_removal,[],[f1474])).
fof(f1474,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) )),
  inference(resolution,[],[f1457,f319])).
fof(f1457,plain,(
  ( ! [X4,X3] : (sP7(X4,X3,'F_0') | s(bool,'T_0') != s(bool,X3) | p(s(bool,X4))) )),
  inference(superposition,[],[f1451,f316])).
fof(f1342,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))))) ) | $spl22),
  inference(cnf_transformation,[],[f1342_D])).
fof(f1342_D,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f1323,plain,(
  ( ! [X128,X129] : (p(s(bool,X128)) | p(s(bool,X129)) | s(bool,X129) != s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))))) )),
  inference(resolution,[],[f538,f426])).
fof(f1256,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11)))) | s(bool,X26) != s(bool,X27)) )),
  inference(resolution,[],[f538,f262])).
fof(f1253,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))))) )),
  inference(resolution,[],[f538,f261])).
fof(f1252,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))))) )),
  inference(resolution,[],[f538,f260])).
fof(f1251,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11)))) | s(bool,X11) != s(bool,X12)) )),
  inference(resolution,[],[f538,f258])).
fof(f1249,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,sK15(sK13,sK12,sK10,sK11))))) )),
  inference(resolution,[],[f538,f256])).
fof(f1072,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f1001,f530])).
fof(f1001,plain,(
  ( ! [X249,X246,X248,X247] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X246,X247),X248),s(fun(X247,bool),X249))))) ) | $spl10),
  inference(resolution,[],[f530,f420])).
fof(f1071,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f1000,f530])).
fof(f1000,plain,(
  ( ! [X243,X244,X242,X245] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X242,X243),X244),s(fun(X243,bool),X245))))) ) | $spl10),
  inference(resolution,[],[f530,f253])).
fof(f1070,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f996,f530])).
fof(f996,plain,(
  ( ! [X227,X225,X228,X226] : (p(s(bool,happ(s(fun(X225,bool),X226),s(X225,sK15(X226,X227,X228,X225)))))) ) | $spl10),
  inference(resolution,[],[f530,f253])).
fof(f1069,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f993,f530])).
fof(f993,plain,(
  ( ! [X218,X216,X219,X217] : (p(s(bool,happ(s(fun(X216,bool),X217),s(X216,sK17(X217,X218,X219,X216)))))) ) | $spl10),
  inference(resolution,[],[f530,f420])).
fof(f1064,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f984,f530])).
fof(f984,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(resolution,[],[f530,f415])).
fof(f1063,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f983,f530])).
fof(f983,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(resolution,[],[f530,f413])).
fof(f1062,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f979,f530])).
fof(f979,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(resolution,[],[f530,f379])).
fof(f1061,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f977,f530])).
fof(f977,plain,(
  ( ! [X205] : (p(s(bool,X205))) ) | $spl10),
  inference(resolution,[],[f530,f376])).
fof(f1060,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f976,f530])).
fof(f976,plain,(
  ( ! [X204] : (p(s(bool,X204))) ) | $spl10),
  inference(resolution,[],[f530,f373])).
fof(f1059,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f975,f530])).
fof(f975,plain,(
  ( ! [X203] : (p(s(bool,X203))) ) | $spl10),
  inference(resolution,[],[f530,f370])).
fof(f985,plain,(
  $false | $spl10),
  inference(resolution,[],[f530,f417])).
fof(f986,plain,(
  $false | $spl10),
  inference(resolution,[],[f530,f418])).
fof(f995,plain,(
  $false | $spl10),
  inference(resolution,[],[f530,f250])).
fof(f839,plain,(
  ( ! [X39,X38] : (s(bool,'F_0') != s(bool,X39) | p(s(bool,X38))) ) | $spl2),
  inference(subsumption_resolution,[],[f802,f526])).
fof(f802,plain,(
  ( ! [X39,X38] : (s(bool,'F_0') != s(bool,X39) | p(s(bool,X38)) | s(bool,'T_0') = s(bool,X38)) ) | $spl2),
  inference(superposition,[],[f526,f424])).
fof(f837,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') != s(bool,X37) | p(s(bool,X36))) ) | $spl2),
  inference(subsumption_resolution,[],[f801,f526])).
fof(f801,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') != s(bool,X37) | p(s(bool,X36)) | s(bool,'F_0') = s(bool,X36)) ) | $spl2),
  inference(superposition,[],[f526,f424])).
fof(f832,plain,(
  ( ! [X30,X31] : (s(bool,'F_0') != s(bool,X31) | p(s(bool,X30))) ) | $spl2),
  inference(duplicate_literal_removal,[],[f798])).
fof(f798,plain,(
  ( ! [X30,X31] : (s(bool,'F_0') != s(bool,X31) | p(s(bool,X30)) | p(s(bool,X30))) ) | $spl2),
  inference(superposition,[],[f526,f400])).
fof(f833,plain,(
  ( ! [X28,X29] : (s(bool,'F_0') != s(bool,X29) | p(s(bool,X28))) ) | $spl2),
  inference(duplicate_literal_removal,[],[f797])).
fof(f797,plain,(
  ( ! [X28,X29] : (s(bool,'F_0') != s(bool,X29) | p(s(bool,X28)) | p(s(bool,X28))) ) | $spl2),
  inference(superposition,[],[f526,f398])).
fof(f786,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f737,f531])).
fof(f737,plain,(
  ~p(s(bool,'T_0')) | $spl12),
  inference(resolution,[],[f531,f416])).
fof(f416,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f245])).
fof(f785,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f736,f531])).
fof(f736,plain,(
  ~p(s(bool,'T_0')) | $spl12),
  inference(resolution,[],[f531,f414])).
fof(f414,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f244])).
fof(f784,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f730,f531])).
fof(f730,plain,(
  ( ! [X206] : (~p(s(bool,X206))) ) | $spl12),
  inference(resolution,[],[f531,f368])).
fof(f735,plain,(
  $false | $spl12),
  inference(resolution,[],[f531,f412])).
fof(f740,plain,(
  $false | $spl12),
  inference(resolution,[],[f531,f251])).
fof(f636,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) ) | $spl14),
  inference(cnf_transformation,[],[f636_D])).
fof(f636_D,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f618,plain,(
  ( ! [X130,X129] : (p(s(bool,X129)) | ~p(s(bool,X130)) | s(bool,X130) != s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) )),
  inference(resolution,[],[f251,f426])).
fof(f547,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | s(bool,X22) != s(bool,X24)) )),
  inference(resolution,[],[f251,f262])).
fof(f545,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) )),
  inference(resolution,[],[f251,f261])).
fof(f637,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) ) | $spl16),
  inference(cnf_transformation,[],[f637_D])).
fof(f637_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f528,plain,(
  ( ! [X4,X5] : (s(bool,X4) != s(bool,X5) | p(s(bool,X4)) | ~p(s(bool,X5))) ) | $spl6),
  inference(cnf_transformation,[],[f528_D])).
fof(f528_D,plain,(
  ( ! [X4,X5] : (s(bool,X4) != s(bool,X5) | p(s(bool,X4)) | ~p(s(bool,X5))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f544,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X15) != s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | s(bool,X13) != s(bool,X14)) )),
  inference(resolution,[],[f251,f260])).
fof(f543,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) )),
  inference(resolution,[],[f251,f260])).
fof(f542,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | s(bool,X7) != s(bool,X9)) )),
  inference(resolution,[],[f251,f258])).
fof(f541,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X6) != s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | s(bool,X4) != s(bool,X5)) )),
  inference(resolution,[],[f251,f256])).
fof(f540,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) )),
  inference(resolution,[],[f251,f256])).
fof(f525,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) ) | $spl0),
  inference(cnf_transformation,[],[f525_D])).
fof(f525_D,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f506,plain,(
  ( ! [X125,X126] : (p(s(bool,X125)) | p(s(bool,X126)) | s(bool,X126) != s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) )),
  inference(resolution,[],[f250,f426])).
fof(f439,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | s(bool,X23) != s(bool,X24)) )),
  inference(resolution,[],[f250,f262])).
fof(f527,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) ) | $spl4),
  inference(cnf_transformation,[],[f527_D])).
fof(f527_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f438,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) )),
  inference(resolution,[],[f250,f262])).
fof(f436,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) )),
  inference(resolution,[],[f250,f261])).
fof(f435,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) )),
  inference(resolution,[],[f250,f260])).
fof(f434,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | s(bool,X8) != s(bool,X9)) )),
  inference(resolution,[],[f250,f258])).
fof(f433,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X5) | s(bool,X6) != s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) )),
  inference(resolution,[],[f250,f258])).
fof(f432,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) )),
  inference(resolution,[],[f250,f256])).
% SZS output end Proof for thm.quantHeuristics
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 729
% Passive clauses: 3583
% Generated clauses: 27028
% Final active clauses: 708
% Final passive clauses: 2460
% Input formulas: 52
% Initial clauses: 176
% 
% Selected by SInE selection: 46
% SInE iterations: 1
% 
% Duplicate literals: 2013
% Trivial inequalities: 250
% Fw subsumption resolutions: 4134
% Fw demodulations: 2508
% Bw demodulations: 8
% 
% Simple tautologies: 3144
% Equational tautologies: 238
% Forward subsumptions: 10792
% 
% Binary resolution: 14530
% Factoring: 48
% Forward superposition: 1829
% Backward superposition: 1633
% Self superposition: 2
% Equality factoring: 2
% Equality resolution: 25
% 
% Split clauses: 718
% Split components: 1606
% SAT solver clauses: 851
% SAT solver unit clauses: 53
% SAT solver binary clauses: 168
% 
% Sat splits: 718
% Sat splitting refutations: 55
% 
% Memory used [KB]: 6652
% Time elapsed: 0.449 s
% ------------------------------
% ------------------------------
% Success in time 12.178 s
