Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.util_prob
% (52084)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1094
% Passive clauses: 7459
% Generated clauses: 41810
% Final active clauses: 1000
% Final passive clauses: 5902
% Input formulas: 261
% Initial clauses: 599
% 
% Duplicate literals: 271
% Trivial inequalities: 61
% Fw subsumption resolutions: 2196
% Fw demodulations: 3155
% Bw demodulations: 74
% 
% Simple tautologies: 343
% Equational tautologies: 30
% Forward subsumptions: 27766
% Fw demodulations to eq. taut.: 1
% 
% Binary resolution: 7628
% Factoring: 1367
% Forward superposition: 13503
% Backward superposition: 12663
% Self superposition: 49
% Equality factoring: 2
% Equality resolution: 12
% 
% Split clauses: 432
% Split components: 866
% SAT solver clauses: 668
% SAT solver unit clauses: 16
% SAT solver binary clauses: 83
% 
% Sat splits: 432
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 665
% TWLsolver calls for satisfiability: 428
% 
% Memory used [KB]: 19829
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.util_prob
% (52103)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 653
% Passive clauses: 3378
% Generated clauses: 28615
% Final active clauses: 589
% Final passive clauses: 2256
% Input formulas: 261
% Initial clauses: 599
% 
% Duplicate literals: 116
% Trivial inequalities: 32
% Fw subsumption resolutions: 1021
% Fw demodulations: 3423
% Bw demodulations: 492
% Condensations: 91
% Global subsumptions: 27
% 
% Simple tautologies: 336
% Equational tautologies: 12
% Forward subsumptions: 16178
% Fw demodulations to eq. taut.: 21
% Bw demodulations to eq. taut.: 3
% 
% Binary resolution: 2838
% Factoring: 209
% Forward superposition: 9169
% Backward superposition: 6893
% Self superposition: 71
% Equality factoring: 2
% Equality resolution: 9
% 
% Split clauses: 19
% Split components: 38
% SAT solver clauses: 4315
% SAT solver unit clauses: 690
% SAT solver binary clauses: 203
% SAT solver learnt clauses: 59
% SAT solver learnt literals: 27
% 
% Sat splits: 19
% Sat splitting refutations: 66
% 
% TWLsolver clauses: 4204
% TWLsolver calls for satisfiability: 4507
% 
% Memory used [KB]: 23283
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.util_prob
% (52125)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 953
% Passive clauses: 10896
% Generated clauses: 129797
% Final active clauses: 918
% Final passive clauses: 7539
% Input formulas: 261
% Initial clauses: 599
% Discarded non-redundant clauses: 5504
% 
% Duplicate literals: 12
% Trivial inequalities: 24
% Fw subsumption resolutions: 1053
% Fw demodulations: 46462
% Bw demodulations: 442
% Condensations: 78
% Global subsumptions: 1
% 
% Simple tautologies: 150
% Equational tautologies: 37
% Forward subsumptions: 66230
% Fw demodulations to eq. taut.: 263
% Bw demodulations to eq. taut.: 2
% 
% Unit resulting resolution: 48680
% Factoring: 2
% Forward superposition: 19764
% Backward superposition: 12541
% Self superposition: 121
% Equality factoring: 2
% Equality resolution: 12
% 
% Split clauses: 4
% Split components: 8
% SAT solver clauses: 11922
% SAT solver unit clauses: 8649
% SAT solver binary clauses: 2087
% 
% Sat splits: 4
% Sat splitting refutations: 1015
% 
% Memory used [KB]: 320208
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.util_prob
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.util_prob
% SZS output start Proof for thm.util_prob
fof(f34133,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f966,f967,f976,f1141,f978,f979,f980,f981,f982,f984,f1144,f985,f1092,f1109,f1110,f1111,f1178,f1179,f1180,f1184,f1185,f1186,f1187,f1188,f1341,f1189,f1191,f1293,f1355,f1359,f1361,f1308,f1310,f1311,f1525,f1675,f1676,f1677,f1678,f1679,f1680,f1682,f1784,f1787,f1788,f1789,f1834,f1835,f1836,f1837,f1838,f1840,f1841,f1948,f1954,f1955,f1956,f2001,f2002,f2003,f2004,f2005,f1809,f2006,f2008,f2110,f2113,f2114,f2115,f2148,f2149,f2150,f2151,f2152,f2153,f2155,f2257,f2260,f2261,f2262,f2307,f2308,f2309,f2310,f2311,f2313,f2314,f2421,f2448,f2427,f2428,f2429,f2296,f2474,f2475,f2476,f2477,f2478,f2282,f2479,f2481,f2583,f2281,f2586,f2587,f2588,f2297,f2625,f2626,f2627,f2639,f2640,f2641,f2653,f2654,f2655,f2953,f2952,f2875,f3028,f3029,f3030,f3042,f3043,f3044,f3055,f3056,f3057,f3069,f3070,f3071,f3281,f3277,f3261,f3260,f3254,f3440,f3441,f3442,f3455,f3456,f3457,f3467,f3468,f3469,f3470,f3471,f3472,f3474,f3576,f3627,f3628,f3629,f3630,f3631,f3633,f1145,f3604,f3634,f3741,f3872,f3874,f3875,f3877,f3879,f3981,f4025,f4026,f4027,f4033,f4039,f4040,f4041,f4047,f4053,f4054,f4055,f4062,f4068,f4069,f4070,f4076,f4082,f4083,f4084,f4090,f4104,f4105,f4106,f4117,f4118,f4119,f4126,f4131,f4132,f4133,f4141,f4146,f4147,f4148,f4156,f4161,f4162,f4163,f4171,f4173,f4175,f4176,f4178,f4180,f4282,f4331,f4332,f4333,f4338,f4339,f4340,f4357,f4359,f4360,f4361,f4364,f4471,f4505,f4531,f4533,f4534,f4536,f4538,f4640,f4309,f4688,f4689,f4690,f4695,f4696,f4697,f4714,f4715,f4716,f4721,f4722,f4723,f4805,f4806,f4807,f4812,f4813,f4814,f4912,f4913,f4914,f4919,f4920,f4921,f4937,f4938,f4939,f4944,f4945,f4946,f4978,f4979,f4980,f4985,f4986,f4987,f5085,f5086,f5087,f5092,f5093,f5094,f5162,f5163,f5164,f5169,f5170,f5171,f5189,f5190,f5191,f5196,f5197,f5198,f5209,f5210,f5211,f5218,f5223,f5224,f5225,f5232,f5239,f5240,f5241,f5248,f5257,f5258,f5259,f5264,f5265,f5266,f5279,f5280,f5281,f5288,f5297,f5298,f5299,f5304,f5305,f5306,f5318,f5319,f5320,f5327,f5333,f5334,f5335,f5342,f5350,f5351,f5352,f5359,f5367,f5368,f5369,f5376,f5535,f5537,f5538,f5540,f5542,f5644,f5651,f5652,f5653,f5658,f5659,f5660,f5703,f5704,f5705,f5714,f5716,f5717,f5718,f5721,f5828,f5838,f5839,f5840,f5845,f5846,f5847,f5899,f5900,f5901,f5915,f5916,f5917,f5935,f5936,f5937,f5946,f5948,f5949,f5951,f5953,f6055,f6062,f6063,f6064,f6069,f6070,f6071,f6109,f6110,f6111,f6127,f6128,f6129,f6300,f6332,f6333,f6334,f6349,f6350,f6351,f6367,f6368,f6369,f6384,f6385,f6386,f6402,f6403,f6404,f6412,f6562,f6413,f6564,f6566,f6417,f6419,f6420,f6422,f6424,f6526,f6533,f6534,f6535,f6540,f6541,f6542,f6595,f6596,f6597,f6611,f6612,f6613,f6621,f6623,f6624,f6625,f6628,f6735,f5867,f6745,f6746,f6747,f5692,f6752,f6753,f6754,f5694,f6804,f6805,f6806,f6818,f6819,f6820,f6832,f6833,f6834,f6850,f6852,f6853,f6854,f6857,f6964,f7002,f6974,f6975,f6976,f6580,f6981,f6982,f6983,f7042,f7043,f7044,f7057,f7058,f7059,f7067,f7069,f7070,f7072,f7074,f7176,f7183,f7184,f7185,f7190,f7191,f7192,f7230,f7231,f7232,f7244,f7245,f7246,f7258,f7260,f7261,f7263,f7265,f7367,f7374,f7375,f7376,f7381,f1173,f7382,f7383,f7420,f7421,f7422,f7435,f7436,f7437,f7449,f7450,f7451,f7464,f7465,f7466,f7478,f7480,f7481,f7483,f7485,f7587,f7642,f7677,f1365,f7689,f7691,f7692,f7693,f7696,f7803,f7845,f7879,f7880,f7886,f7895,f7896,f7902,f7912,f7913,f7921,f7934,f7935,f7941,f7947,f7949,f7950,f7952,f7954,f8056,f8858,f8863,f8859,f8866,f8870,f8867,f8873,f8877,f8874,f8880,f8884,f8881,f8900,f8905,f8901,f8908,f8912,f8909,f8915,f8919,f8916,f8922,f8926,f8923,f8930,f8938,f8931,f8945,f8950,f8946,f8954,f8962,f8955,f8969,f8974,f8970,f8977,f8981,f8978,f8984,f8988,f8985,f8991,f8995,f8992,f8998,f9002,f8999,f9005,f9009,f9006,f9012,f9016,f9013,f9019,f9023,f9020,f9026,f9030,f9027,f9033,f9037,f9034,f9040,f9044,f9041,f9047,f9052,f9049,f9055,f9060,f9057,f8860,f9063,f9068,f9065,f9071,f9076,f9073,f8902,f9482,f9483,f9484,f9485,f9486,f9487,f9488,f9489,f9509,f1133,f1131,f9510,f1137,f9521,f9523,f9524,f9525,f9528,f9635,f1142,f9653,f9654,f1175,f9698,f9700,f9701,f9703,f9705,f9807,f10054,f10055,f10069,f10070,f9494,f10071,f10072,f9498,f10073,f10074,f9500,f10075,f9492,f10076,f9502,f9496,f10132,f1333,f1331,f10133,f1335,f10134,f1337,f1339,f10138,f10140,f10141,f10143,f10145,f10247,f1340,f10296,f10297,f10298,f1358,f10262,f10264,f10265,f10327,f10328,f10329,f4112,f10343,f10344,f10345,f4346,f10351,f10352,f4347,f10445,f10446,f10447,f4703,f10453,f10454,f4704,f10469,f10470,f10471,f4729,f10477,f10478,f4730,f10516,f10518,f10519,f10521,f10523,f10625,f10673,f10712,f10714,f10715,f10716,f10719,f10826,f10871,f10905,f10906,f10907,f4819,f10913,f10914,f4820,f10927,f10928,f10929,f5711,f10942,f10943,f10944,f5907,f10958,f10959,f10960,f5925,f10979,f10980,f10981,f5943,f10994,f10995,f10996,f6119,f11012,f11013,f11014,f4927,f11020,f11021,f4928,f11034,f11035,f11036,f6137,f11051,f11052,f11053,f6341,f11067,f11068,f11069,f6358,f11084,f11085,f11086,f6375,f11100,f11101,f11102,f6393,f11117,f11118,f11119,f4952,f11125,f11126,f4953,f11140,f11141,f11142,f6410,f11156,f11157,f11158,f6604,f11171,f11172,f11173,f6618,f11184,f11185,f11186,f6811,f11197,f11198,f11199,f6825,f11212,f11213,f11214,f4994,f11220,f11221,f4995,f11234,f11235,f11236,f6839,f11247,f11248,f11249,f7049,f11260,f11261,f11262,f7064,f11273,f11274,f11275,f7237,f11286,f11287,f11288,f7251,f11301,f11302,f11303,f5101,f11309,f11310,f5102,f11323,f11324,f11325,f7427,f11336,f11337,f11338,f7442,f11349,f11350,f11351,f7456,f11362,f11363,f11364,f7471,f11377,f11378,f11379,f5179,f11385,f11386,f5180,f11402,f11403,f11411,f11420,f11421,f11429,f11438,f11439,f11440,f5203,f11446,f11447,f5204,f11463,f11464,f11471,f11483,f11484,f11491,f11502,f11503,f11509,f11521,f11522,f11529,f11540,f11541,f11542,f5271,f11548,f11549,f5272,f11564,f11565,f11571,f11583,f11584,f11591,f11601,f11602,f11607,f11618,f11619,f11620,f5311,f11626,f11627,f5312,f11641,f11642,f11647,f11657,f11658,f11663,f11673,f11674,f11679,f11682,f11684,f11685,f11687,f11689,f11791,f5679,f11798,f11799,f11800,f5693,f11806,f11807,f5695,f11845,f11846,f11851,f11860,f11861,f11866,f11875,f11876,f11881,f11888,f11890,f11891,f11893,f11895,f11997,f6567,f12004,f12005,f12006,f6581,f12012,f12013,f1366,f12051,f12052,f12057,f12063,f12065,f12066,f12068,f12070,f12172,f8094,f12232,f12233,f12238,f12247,f12248,f12253,f12261,f12262,f12267,f12276,f1176,f12277,f12282,f12292,f12294,f12295,f12297,f12299,f12401,f12448,f12427,f12461,f1177,f12471,f12473,f12474,f12476,f12478,f12580,f12635,f12654,f12656,f12657,f12659,f12661,f12763,f12842,f12844,f12845,f12846,f12849,f12956,f12962,f12963,f12964,f13016,f13018,f13019,f13020,f13023,f13130,f3775,f13193,f13194,f13195,f13201,f13203,f13205,f13206,f13208,f13210,f13312,f13315,f13316,f13317,f13353,f13354,f13355,f2633,f13519,f13520,f13548,f13549,f13550,f2647,f13561,f13562,f13563,f2662,f13575,f13576,f13577,f3036,f13588,f1356,f13589,f13590,f3050,f13602,f13604,f13605,f13607,f13609,f13711,f12818,f13785,f13786,f3064,f13798,f13799,f3079,f13812,f13813,f3450,f13826,f1360,f13827,f3465,f13836,f13838,f13839,f13841,f13843,f13945,f13996,f13998,f13999,f14000,f14003,f14110,f14154,f14193,f14195,f14196,f14198,f14200,f14302,f14349,f14328,f14362,f1362,f14372,f14374,f14375,f14377,f14379,f14481,f14536,f14639,f14641,f14642,f14644,f14646,f14748,f14799,f15403,f15405,f15406,f15407,f15410,f15517,f15673,f15675,f15676,f15678,f15680,f15782,f15825,f15878,f15880,f15881,f15882,f15885,f15992,f15563,f16834,f17239,f17241,f17242,f17244,f17246,f17348,f17382,f17421,f17423,f17424,f17426,f17428,f17530,f17564,f18210,f18213,f18214,f1135,f18217,f18219,f18220,f18222,f18224,f18326,f1808,f19318,f19320,f19321,f19322,f19325,f19432,f19703,f19704,f19705,f19706,f1139,f20041,f20660,f20662,f20663,f20664,f20667,f20774,f20931,f21079,f21081,f21082,f21083,f21086,f21193,f19589,f21564,f20426,f20404,f20298,f20273,f20215,f20108,f19976,f16746,f16843,f16911,f16934,f16576,f19132,f19058,f18950,f18928,f18823,f18798,f18710,f16938,f16748,f16845,f16578,f16915,f17974,f17681,f17162,f17071,f17067,f17053,f17052,f17050,f17038,f17034,f17021,f17019,f17017,f17006,f17005,f17002,f16986,f16984,f16969,f16967,f16956,f16954,f16952,f16940,f16936,f16931,f16929,f16917,f16913,f16910,f16908,f16890,f16889,f16886,f16868,f16866,f16864,f16846,f16844,f16861,f16841,f16862,f16837,f16783,f16749,f16747,f16741,f16760,f16737,f16761,f16735,f16712,f16713,f16684,f16685,f16657,f16655,f16642,f16656,f16637,f16623,f16635,f16636,f16605,f16617,f16577,f16591,f16589,f16573,f16590,f16568,f21583,f21585,f16535,f16550,f21587,f15021,f15022,f15000,f15001,f14979,f14980,f14960,f14961,f14864,f21592,f14862,f21593,f14845,f14831,f21595,f1823,f14828,f21596,f21594,f14812,f14615,f21598,f14614,f14599,f21600,f14598,f14597,f14596,f14582,f21601,f14581,f14580,f21603,f14566,f21602,f14565,f14553,f21604,f14549,f1824,f21607,f21609,f21610,f21612,f21614,f21716,f3603,f21949,f21951,f21952,f21953,f21956,f22063,f22220,f22415,f22416,f22424,f22417,f22418,f22426,f22419,f22420,f22428,f22438,f22440,f22441,f22442,f22445,f22552,f23283,f23285,f23286,f23288,f23290,f23392,f23524,f23767,f23769,f23770,f23772,f23774,f23876,f24181,f24182,f24185,f24186,f24201,f24202,f24205,f24206,f24214,f24215,f24218,f24219,f24263,f24265,f24266,f24267,f24270,f24377,f24535,f24702,f24728,f24992,f24993,f24994,f24995,f25010,f25011,f25012,f25013,f25023,f25024,f25025,f25026,f25038,f25187,f25238,f25239,f25240,f25241,f25254,f25255,f25256,f25257,f25266,f25301,f25302,f25303,f25304,f25387,f25390,f25547,f25550,f25571,f25574,f25778,f25781,f25905,f25908,f25933,f25936,f25962,f25965,f25990,f25993,f26066,f26069,f24190,f26506,f26708,f26866,f26981,f26985,f27012,f27160,f27165,f27187,f27192,f27472,f27476,f27565,f27569,f27596,f27600,f27666,f27672,f27697,f27703,f27726,f24188,f27732,f27858,f27860,f27861,f27862,f27865,f27972,f28005,f28098,f28100,f28101,f28217,f28113,f28115,f28233,f28558,f28522,f28488,f28485,f28483,f28467,f28466,f28460,f28654,f28656,f28657,f28658,f28661,f28768,f1975,f28879,f28881,f28882,f28215,f28996,f28894,f28230,f28896,f28231,f29007,f28220,f29880,f29882,f29883,f29885,f29887,f29989,f24008,f30231,f30244,f30254,f30262,f30274,f30282,f30292,f30302,f30313,f24704,f31332,f31334,f28219,f31175,f31177,f31178,f1143,f31179,f31182,f1146,f31289,f31335,f31460,f31462,f31463,f31469,f32091,f32093,f32094,f32099,f32104,f32259,f32303,f32305,f32306,f32311,f32316,f32512,f32514,f32515,f32521,f32572,f32674,f32723,f32725,f32726,f32732,f32738,f1147,f32769,f32797,f32823,f1148,f24997,f32865,f32867,f32868,f32874,f32997,f32999,f33000,f33007,f33065,f33067,f33068,f33074,f33231,f1364,f33233,f1367,f33234,f33241,f1368,f33278,f34120])).
fof(f34120,plain,(
  $false | $spl37),
  inference(trivial_inequality_removal,[],[f33970])).
fof(f33970,plain,(
  s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))) != s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))) | $spl37),
  inference(resolution,[],[f1333,f691])).
fof(f691,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | s(fun(X1,bool),X4) != s(fun(X1,bool),'const.pred_set.IMAGE_2'(s(fun(X0,X1),X2),s(fun(X0,bool),X3)))) )),
  inference(cnf_transformation,[],[f504])).
fof(f504,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | s(fun(X1,bool),X4) = s(fun(X1,bool),'const.pred_set.IMAGE_2'(s(fun(X0,X1),X2),s(fun(X0,bool),X3)))) & (s(fun(X1,bool),X4) != s(fun(X1,bool),'const.pred_set.IMAGE_2'(s(fun(X0,X1),X2),s(fun(X0,bool),X3))) | p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))))),
  inference(nnf_transformation,[],[f294])).
fof(f294,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) <=> s(fun(X1,bool),X4) = s(fun(X1,bool),'const.pred_set.IMAGE_2'(s(fun(X0,X1),X2),s(fun(X0,bool),X3))))),
  inference(flattening,[],[f293])).
fof(f293,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) <=> s(fun(X1,bool),X4) = s(fun(X1,bool),'const.pred_set.IMAGE_2'(s(fun(X0,X1),X2),s(fun(X0,bool),X3))))),
  inference(rectify,[],[f106])).
fof(f106,axiom,(
  ! [X7,X6] : ! [X9,X8,X10] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(X7,X6),X9),s(fun(X7,bool),X8),s(fun(X6,bool),X10)))) <=> s(fun(X6,bool),X10) = s(fun(X6,bool),'const.pred_set.IMAGE_2'(s(fun(X7,X6),X9),s(fun(X7,bool),X8))))),
  file('thm.util_prob.INJ_IMAGE_BIJ','thm.pred_set.IMAGE_SURJ')).
fof(f33278,plain,(
  $false | $spl35),
  inference(subsumption_resolution,[],[f1331,f968])).
fof(f968,plain,(
  p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))))))),
  inference(resolution,[],[f647,f693])).
fof(f693,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X4),s(fun(X1,bool),X3),s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X1,X0),X4),s(fun(X1,bool),X3)))))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X4),s(fun(X1,bool),X3),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f420])).
fof(f420,plain,(
  ! [X0,X1,X2,X3,X4] : (~p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X4),s(fun(X1,bool),X3),s(fun(X0,bool),X2)))) | p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X4),s(fun(X1,bool),X3),s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X1,X0),X4),s(fun(X1,bool),X3)))))))),
  inference(ennf_transformation,[],[f298])).
fof(f298,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X4),s(fun(X1,bool),X3),s(fun(X0,bool),X2)))) => p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X4),s(fun(X1,bool),X3),s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X1,X0),X4),s(fun(X1,bool),X3)))))))),
  inference(flattening,[],[f297])).
fof(f297,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X4),s(fun(X1,bool),X3),s(fun(X0,bool),X2)))) => p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X4),s(fun(X1,bool),X3),s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X1,X0),X4),s(fun(X1,bool),X3)))))))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X6,X7] : ! [X10,X8,X9] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X7,X6),X9),s(fun(X7,bool),X8),s(fun(X6,bool),X10)))) => p(s(bool,'const.pred_set.INJ_3'(s(fun(X7,X6),X9),s(fun(X7,bool),X8),s(fun(X6,bool),'const.pred_set.IMAGE_2'(s(fun(X7,X6),X9),s(fun(X7,bool),X8)))))))),
  file('thm.util_prob.INJ_IMAGE_BIJ','thm.pred_set.pred_set_|49|76_prove_ltb')).
fof(f647,plain,(
  p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20))))),
  inference(cnf_transformation,[],[f484])).
fof(f484,plain,(
  p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) & ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK16,sK17,sK18,sK19,sK20])],[f405])).
fof(f405,plain,(
  ? [X0,X1,X2,X3] : (? [X4] : p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X3),s(fun(X1,bool),X2),s(fun(X0,bool),X4)))) & ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X1,X0),X3),s(fun(X1,bool),X2),s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X1,X0),X3),s(fun(X1,bool),X2)))))))),
  inference(ennf_transformation,[],[f264])).
fof(f264,plain,(
  ~! [X0,X1,X2,X3] : (? [X4] : p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X3),s(fun(X1,bool),X2),s(fun(X0,bool),X4)))) => p(s(bool,'const.pred_set.BIJ_3'(s(fun(X1,X0),X3),s(fun(X1,bool),X2),s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X1,X0),X3),s(fun(X1,bool),X2)))))))),
  inference(flattening,[],[f263])).
fof(f263,plain,(
  ~! [X0,X1] : ! [X2,X3] : (? [X4] : p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X3),s(fun(X1,bool),X2),s(fun(X0,bool),X4)))) => p(s(bool,'const.pred_set.BIJ_3'(s(fun(X1,X0),X3),s(fun(X1,bool),X2),s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X1,X0),X3),s(fun(X1,bool),X2)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9] : (? [X10] : p(s(bool,'const.pred_set.INJ_3'(s(fun(X7,X6),X9),s(fun(X7,bool),X8),s(fun(X6,bool),X10)))) => p(s(bool,'const.pred_set.BIJ_3'(s(fun(X7,X6),X9),s(fun(X7,bool),X8),s(fun(X6,bool),'const.pred_set.IMAGE_2'(s(fun(X7,X6),X9),s(fun(X7,bool),X8)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9] : (? [X10] : p(s(bool,'const.pred_set.INJ_3'(s(fun(X7,X6),X9),s(fun(X7,bool),X8),s(fun(X6,bool),X10)))) => p(s(bool,'const.pred_set.BIJ_3'(s(fun(X7,X6),X9),s(fun(X7,bool),X8),s(fun(X6,bool),'const.pred_set.IMAGE_2'(s(fun(X7,X6),X9),s(fun(X7,bool),X8)))))))),
  file('thm.util_prob.INJ_IMAGE_BIJ','thm.util_prob.INJ_IMAGE_BIJ')).
fof(f1368,plain,(
  ( ! [X8,X9] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X8),X9),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))),s(fun(X8,bool),'const.pred_set.EMPTY_0'))))) ) | $spl62),
  inference(cnf_transformation,[],[f1368_D])).
fof(f1368_D,plain,(
  ( ! [X8,X9] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X8),X9),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))),s(fun(X8,bool),'const.pred_set.EMPTY_0'))))) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f33241,plain,(
  ( ! [X4,X3] : (sP8('const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X3,X4)) ) | $spl492),
  inference(cnf_transformation,[],[f33241_D])).
fof(f33241_D,plain,(
  ( ! [X4,X3] : (sP8('const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X3,X4)) ) <=> ~$spl492),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl492])])).
fof(f33234,plain,(
  ( ! [X19,X17,X18,X16] : (sP8('const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X16,X17) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X18),X19),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))),s(fun(X18,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1245,f678])).
fof(f678,plain,(
  ( ! [X2,X0,X3,X1] : (s(fun(X1,bool),'const.pred_set.EMPTY_0') = s(fun(X1,bool),X3) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(cnf_transformation,[],[f496])).
fof(f496,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) | s(fun(X1,bool),'const.pred_set.EMPTY_0') = s(fun(X1,bool),X3)) & (s(fun(X1,bool),'const.pred_set.EMPTY_0') != s(fun(X1,bool),X3) | p(s(bool,'const.pred_set.BIJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))))),
  inference(nnf_transformation,[],[f278])).
fof(f278,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) <=> s(fun(X1,bool),'const.pred_set.EMPTY_0') = s(fun(X1,bool),X3))),
  inference(flattening,[],[f277])).
fof(f277,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) <=> s(fun(X1,bool),'const.pred_set.EMPTY_0') = s(fun(X1,bool),X3))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X6,X7] : ! [X9,X8] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(X7,X6),X9),s(fun(X7,bool),X8),s(fun(X6,bool),'const.pred_set.EMPTY_0')))) <=> s(fun(X7,bool),'const.pred_set.EMPTY_0') = s(fun(X7,bool),X8))),
  file('thm.util_prob.INJ_IMAGE_BIJ','thm.pred_set.BIJ_EMPTY|split|1')).
fof(f1245,plain,(
  ( ! [X134,X133] : (sP8('const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18)))),X133,X134)) )),
  inference(resolution,[],[f648,f772])).
fof(f772,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f538])).
fof(f538,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f537])).
fof(f537,plain,(
  ! [X1,X0,X2] : ((~sP8(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP8(X1,X0,X2)))),
  inference(flattening,[],[f536])).
fof(f536,plain,(
  ! [X1,X0,X2] : ((~sP8(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP8(X1,X0,X2)))),
  inference(nnf_transformation,[],[f470])).
fof(f470,plain,(
  ! [X1,X0,X2] : (sP8(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f648,plain,(
  ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))))))),
  inference(cnf_transformation,[],[f484])).
fof(f1367,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X6,sK16),X7),s(fun(X6,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))))))) ) | $spl60),
  inference(cnf_transformation,[],[f1367_D])).
fof(f1367_D,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X6,sK16),X7),s(fun(X6,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))))))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f33233,plain,(
  ( ! [X14,X12,X15,X13] : (sP8('const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X12,X13) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X14,sK16),X15),s(fun(X14,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))))))) )),
  inference(superposition,[],[f1245,f680])).
fof(f680,plain,(
  ( ! [X2,X0,X3,X1] : (s(fun(X1,bool),'const.pred_set.EMPTY_0') = s(fun(X1,bool),X3) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(X1,bool),X3))))) )),
  inference(cnf_transformation,[],[f497])).
fof(f497,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(X1,bool),X3)))) | s(fun(X1,bool),'const.pred_set.EMPTY_0') = s(fun(X1,bool),X3)) & (s(fun(X1,bool),'const.pred_set.EMPTY_0') != s(fun(X1,bool),X3) | p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(X1,bool),X3))))))),
  inference(nnf_transformation,[],[f280])).
fof(f280,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(X1,bool),X3)))) <=> s(fun(X1,bool),'const.pred_set.EMPTY_0') = s(fun(X1,bool),X3))),
  inference(flattening,[],[f279])).
fof(f279,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(X1,bool),X3)))) <=> s(fun(X1,bool),'const.pred_set.EMPTY_0') = s(fun(X1,bool),X3))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X7,X6] : ! [X9,X8] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(X7,X6),X9),s(fun(X7,bool),'const.pred_set.EMPTY_0'),s(fun(X6,bool),X8)))) <=> s(fun(X6,bool),'const.pred_set.EMPTY_0') = s(fun(X6,bool),X8))),
  file('thm.util_prob.INJ_IMAGE_BIJ','thm.pred_set.BIJ_EMPTY|split|0')).
fof(f1364,plain,(
  ( ! [X2,X1] : (~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK16,X1),X2),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))),s(fun(X1,bool),'const.pred_set.EMPTY_0'))))) ) | $spl56),
  inference(cnf_transformation,[],[f1364_D])).
fof(f1364_D,plain,(
  ( ! [X2,X1] : (~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK16,X1),X2),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))),s(fun(X1,bool),'const.pred_set.EMPTY_0'))))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f33231,plain,(
  ( ! [X6,X4,X5,X3] : (sP8('const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X3,X4) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK16,X5),X6),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))),s(fun(X5,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1245,f948])).
fof(f948,plain,(
  ( ! [X2,X0,X3,X1] : (s(fun(X1,bool),'const.pred_set.EMPTY_0') = s(fun(X1,bool),X3) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(cnf_transformation,[],[f645])).
fof(f645,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) | s(fun(X1,bool),'const.pred_set.EMPTY_0') = s(fun(X1,bool),X3)) & (s(fun(X1,bool),'const.pred_set.EMPTY_0') != s(fun(X1,bool),X3) | p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))))),
  inference(nnf_transformation,[],[f392])).
fof(f392,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) <=> s(fun(X1,bool),'const.pred_set.EMPTY_0') = s(fun(X1,bool),X3))),
  inference(flattening,[],[f391])).
fof(f391,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) <=> s(fun(X1,bool),'const.pred_set.EMPTY_0') = s(fun(X1,bool),X3))),
  inference(rectify,[],[f66])).
fof(f66,axiom,(
  ! [X6,X7] : ! [X9,X8] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X7,X6),X9),s(fun(X7,bool),X8),s(fun(X6,bool),'const.pred_set.EMPTY_0')))) <=> s(fun(X7,bool),'const.pred_set.EMPTY_0') = s(fun(X7,bool),X8))),
  file('thm.util_prob.INJ_IMAGE_BIJ','thm.pred_set.INJ_EMPTY|split|1')).
fof(f33074,plain,(
  ( ! [X4,X3] : (sP8(X3,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X4)) ) | $spl490),
  inference(cnf_transformation,[],[f33074_D])).
fof(f33074_D,plain,(
  ( ! [X4,X3] : (sP8(X3,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X4)) ) <=> ~$spl490),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl490])])).
fof(f33068,plain,(
  ( ! [X19,X17,X18,X16] : (sP8(X16,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X17) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X18),X19),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))),s(fun(X18,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1244,f678])).
fof(f1244,plain,(
  ( ! [X132,X131] : (sP8(X131,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18)))),X132)) )),
  inference(resolution,[],[f648,f771])).
fof(f771,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f538])).
fof(f33067,plain,(
  ( ! [X14,X12,X15,X13] : (sP8(X12,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X13) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X14,sK16),X15),s(fun(X14,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))))))) )),
  inference(superposition,[],[f1244,f680])).
fof(f33065,plain,(
  ( ! [X6,X4,X5,X3] : (sP8(X3,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X4) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK16,X5),X6),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))),s(fun(X5,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1244,f948])).
fof(f33007,plain,(
  ( ! [X4,X3] : (sP9('const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X3,X4)) ) | $spl488),
  inference(cnf_transformation,[],[f33007_D])).
fof(f33007_D,plain,(
  ( ! [X4,X3] : (sP9('const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X3,X4)) ) <=> ~$spl488),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl488])])).
fof(f33000,plain,(
  ( ! [X19,X17,X18,X16] : (sP9('const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X16,X17) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X18),X19),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))),s(fun(X18,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1242,f678])).
fof(f1242,plain,(
  ( ! [X127,X128] : (sP9('const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18)))),X127,X128)) )),
  inference(resolution,[],[f648,f768])).
fof(f768,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f535])).
fof(f535,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f534])).
fof(f534,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP9(X2,X0,X1)))),
  inference(flattening,[],[f533])).
fof(f533,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP9(X2,X0,X1)))),
  inference(nnf_transformation,[],[f471])).
fof(f471,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f32999,plain,(
  ( ! [X14,X12,X15,X13] : (sP9('const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X12,X13) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X14,sK16),X15),s(fun(X14,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))))))) )),
  inference(superposition,[],[f1242,f680])).
fof(f32997,plain,(
  ( ! [X6,X4,X5,X3] : (sP9('const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X3,X4) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK16,X5),X6),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))),s(fun(X5,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1242,f948])).
fof(f32874,plain,(
  ( ! [X4,X3] : (sP9(X3,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X4)) ) | $spl486),
  inference(cnf_transformation,[],[f32874_D])).
fof(f32874_D,plain,(
  ( ! [X4,X3] : (sP9(X3,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X4)) ) <=> ~$spl486),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl486])])).
fof(f32868,plain,(
  ( ! [X19,X17,X18,X16] : (sP9(X16,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X17) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X18),X19),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))),s(fun(X18,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1241,f678])).
fof(f1241,plain,(
  ( ! [X125,X126] : (sP9(X125,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18)))),X126)) )),
  inference(resolution,[],[f648,f767])).
fof(f767,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f535])).
fof(f32867,plain,(
  ( ! [X14,X12,X15,X13] : (sP9(X12,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X13) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X14,sK16),X15),s(fun(X14,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))))))) )),
  inference(superposition,[],[f1241,f680])).
fof(f32865,plain,(
  ( ! [X6,X4,X5,X3] : (sP9(X3,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X4) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK16,X5),X6),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))),s(fun(X5,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1241,f948])).
fof(f24997,plain,(
  s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | $spl458),
  inference(cnf_transformation,[],[f24997_D])).
fof(f24997_D,plain,(
  s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) <=> ~$spl458),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl458])])).
fof(f1148,plain,(
  ( ! [X24] : (p(s(bool,X24))) ) | $spl24),
  inference(cnf_transformation,[],[f1148_D])).
fof(f1148_D,plain,(
  ( ! [X24] : (p(s(bool,X24))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f32823,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | ($spl69 | $spl77)),
  inference(resolution,[],[f21671,f22011])).
fof(f22011,plain,(
  ( ! [X196,X194,X195] : (sP7(X195,X196,X194) | s(bool,X194) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) | $spl77),
  inference(resolution,[],[f4632,f774])).
fof(f774,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f540])).
fof(f540,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(flattening,[],[f539])).
fof(f539,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(nnf_transformation,[],[f469])).
fof(f469,plain,(
  ! [X0,X1,X2] : (sP7(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f4632,plain,(
  ( ! [X205] : (p(s(bool,X205)) | s(bool,X205) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) | $spl77),
  inference(resolution,[],[f2297,f861])).
fof(f861,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f441])).
fof(f441,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f440])).
fof(f440,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f342])).
fof(f342,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f58])).
fof(f58,axiom,(
  ! [X25,X26] : ((p(s(bool,X25)) => p(s(bool,X26))) => ((p(s(bool,X26)) => p(s(bool,X25))) => s(bool,X25) = s(bool,X26)))),
  file('thm.util_prob.INJ_IMAGE_BIJ','thm.bool.IMP_ANTISYM_AX')).
fof(f21671,plain,(
  ( ! [X140,X139] : (~sP7(X139,X140,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X140)) | p(s(bool,X139))) ) | $spl69),
  inference(resolution,[],[f1824,f773])).
fof(f773,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f540])).
fof(f32797,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | ($spl69 | $spl77)),
  inference(resolution,[],[f21670,f22013])).
fof(f22013,plain,(
  ( ! [X202,X200,X201] : (sP7(X200,X201,X202) | s(bool,X200) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) | $spl77),
  inference(resolution,[],[f4632,f776])).
fof(f776,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f540])).
fof(f21670,plain,(
  ( ! [X138,X137] : (~sP7('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X137,X138) | p(s(bool,X138)) | p(s(bool,X137))) ) | $spl69),
  inference(resolution,[],[f1824,f773])).
fof(f32769,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | ($spl69 | $spl77)),
  inference(resolution,[],[f21669,f22012])).
fof(f22012,plain,(
  ( ! [X198,X199,X197] : (sP7(X198,X197,X199) | s(bool,X197) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) | $spl77),
  inference(resolution,[],[f4632,f775])).
fof(f775,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f540])).
fof(f21669,plain,(
  ( ! [X136,X135] : (~sP7(X135,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X136) | p(s(bool,X136)) | p(s(bool,X135))) ) | $spl69),
  inference(resolution,[],[f1824,f773])).
fof(f1147,plain,(
  ( ! [X25] : (~p(s(bool,X25))) ) | $spl22),
  inference(cnf_transformation,[],[f1147_D])).
fof(f1147_D,plain,(
  ( ! [X25] : (~p(s(bool,X25))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f32738,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | ($spl69 | $spl77)),
  inference(resolution,[],[f21666,f22010])).
fof(f22010,plain,(
  ( ! [X191,X192,X193] : (sP8(X192,X193,X191) | s(bool,X191) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) | $spl77),
  inference(resolution,[],[f4632,f770])).
fof(f770,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f538])).
fof(f21666,plain,(
  ( ! [X130,X129] : (~sP8(X129,X130,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X130)) | ~p(s(bool,X129))) ) | $spl69),
  inference(resolution,[],[f1824,f769])).
fof(f769,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f538])).
fof(f32732,plain,(
  ( ! [X4,X3] : (sP5('const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X3,X4)) ) | $spl484),
  inference(cnf_transformation,[],[f32732_D])).
fof(f32732_D,plain,(
  ( ! [X4,X3] : (sP5('const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X3,X4)) ) <=> ~$spl484),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl484])])).
fof(f32726,plain,(
  ( ! [X19,X17,X18,X16] : (sP5('const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X16,X17) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X18),X19),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))),s(fun(X18,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1229,f678])).
fof(f1229,plain,(
  ( ! [X101,X102] : (sP5('const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18)))),X101,X102)) )),
  inference(resolution,[],[f648,f751])).
fof(f751,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f527])).
fof(f527,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(flattening,[],[f526])).
fof(f526,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(nnf_transformation,[],[f466])).
fof(f466,plain,(
  ! [X0,X1,X2] : (sP5(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f32725,plain,(
  ( ! [X14,X12,X15,X13] : (sP5('const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X12,X13) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X14,sK16),X15),s(fun(X14,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))))))) )),
  inference(superposition,[],[f1229,f680])).
fof(f32723,plain,(
  ( ! [X6,X4,X5,X3] : (sP5('const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X3,X4) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK16,X5),X6),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))),s(fun(X5,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1229,f948])).
fof(f32674,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | ($spl69 | $spl77)),
  inference(resolution,[],[f21663,f22007])).
fof(f22007,plain,(
  ( ! [X182,X184,X183] : (sP9(X183,X184,X182) | s(bool,X182) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) | $spl77),
  inference(resolution,[],[f4632,f766])).
fof(f766,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f535])).
fof(f21663,plain,(
  ( ! [X123,X124] : (~sP9(X123,X124,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X124)) | ~p(s(bool,X123))) ) | $spl69),
  inference(resolution,[],[f1824,f765])).
fof(f765,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f535])).
fof(f32572,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | ($spl69 | $spl77)),
  inference(resolution,[],[f21650,f21994])).
fof(f21994,plain,(
  ( ! [X144,X145,X143] : (sP5(X144,X145,X143) | s(bool,X143) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) | $spl77),
  inference(resolution,[],[f4632,f749])).
fof(f749,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f527])).
fof(f21650,plain,(
  ( ! [X97,X98] : (~sP5(X97,X98,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X98)) | ~p(s(bool,X97))) ) | $spl69),
  inference(resolution,[],[f1824,f748])).
fof(f748,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f527])).
fof(f32521,plain,(
  ( ! [X4,X3] : (sP5(X3,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X4)) ) | $spl482),
  inference(cnf_transformation,[],[f32521_D])).
fof(f32521_D,plain,(
  ( ! [X4,X3] : (sP5(X3,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X4)) ) <=> ~$spl482),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl482])])).
fof(f32515,plain,(
  ( ! [X19,X17,X18,X16] : (sP5(X16,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X17) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X18),X19),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))),s(fun(X18,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1228,f678])).
fof(f1228,plain,(
  ( ! [X99,X100] : (sP5(X99,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18)))),X100)) )),
  inference(resolution,[],[f648,f750])).
fof(f750,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f527])).
fof(f32514,plain,(
  ( ! [X14,X12,X15,X13] : (sP5(X12,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X13) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X14,sK16),X15),s(fun(X14,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))))))) )),
  inference(superposition,[],[f1228,f680])).
fof(f32512,plain,(
  ( ! [X6,X4,X5,X3] : (sP5(X3,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X4) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK16,X5),X6),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))),s(fun(X5,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1228,f948])).
fof(f32316,plain,(
  ( ! [X8,X9] : (p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | ($spl69 | $spl77)),
  inference(resolution,[],[f21637,f21977])).
fof(f21977,plain,(
  ( ! [X94,X92,X93] : (sP3(X93,X94,X92) | s(bool,X92) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) | $spl77),
  inference(resolution,[],[f4632,f732])).
fof(f732,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f520])).
fof(f520,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f519])).
fof(f519,plain,(
  ! [X2,X0,X1] : ((~sP3(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP3(X2,X0,X1)))),
  inference(flattening,[],[f518])).
fof(f518,plain,(
  ! [X2,X0,X1] : ((~sP3(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP3(X2,X0,X1)))),
  inference(nnf_transformation,[],[f463])).
fof(f463,plain,(
  ! [X2,X0,X1] : (sP3(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f21637,plain,(
  ( ! [X72,X71] : (~sP3(X71,X72,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X72)) | ~p(s(bool,X71))) ) | $spl69),
  inference(resolution,[],[f1824,f731])).
fof(f731,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f520])).
fof(f32311,plain,(
  ( ! [X4,X3] : (sP3('const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X3,X4)) ) | $spl480),
  inference(cnf_transformation,[],[f32311_D])).
fof(f32311_D,plain,(
  ( ! [X4,X3] : (sP3('const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X3,X4)) ) <=> ~$spl480),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl480])])).
fof(f32306,plain,(
  ( ! [X19,X17,X18,X16] : (sP3('const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X16,X17) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X18),X19),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))),s(fun(X18,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1215,f678])).
fof(f1215,plain,(
  ( ! [X74,X73] : (sP3('const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18)))),X73,X74)) )),
  inference(resolution,[],[f648,f734])).
fof(f734,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f520])).
fof(f32305,plain,(
  ( ! [X14,X12,X15,X13] : (sP3('const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X12,X13) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X14,sK16),X15),s(fun(X14,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))))))) )),
  inference(superposition,[],[f1215,f680])).
fof(f32303,plain,(
  ( ! [X6,X4,X5,X3] : (sP3('const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X3,X4) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK16,X5),X6),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))),s(fun(X5,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1215,f948])).
fof(f32259,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | ($spl69 | $spl77)),
  inference(resolution,[],[f21636,f21978])).
fof(f21978,plain,(
  ( ! [X97,X95,X96] : (sP3(X96,X95,X97) | s(bool,X95) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) | $spl77),
  inference(resolution,[],[f4632,f733])).
fof(f733,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f520])).
fof(f21636,plain,(
  ( ! [X70,X69] : (~sP3(X69,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X70) | p(s(bool,X70)) | ~p(s(bool,X69))) ) | $spl69),
  inference(resolution,[],[f1824,f731])).
fof(f32104,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | ($spl69 | $spl77)),
  inference(resolution,[],[f21620,f21964])).
fof(f21964,plain,(
  ( ! [X54,X55,X53] : (sP1(X54,X53,X55) | s(bool,X53) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) | $spl77),
  inference(resolution,[],[f4632,f716])).
fof(f716,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f512])).
fof(f512,plain,(
  ! [X0,X1,X2] : ((~sP1(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1,X2)))),
  inference(rectify,[],[f511])).
fof(f511,plain,(
  ! [X2,X0,X1] : ((~sP1(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP1(X2,X0,X1)))),
  inference(flattening,[],[f510])).
fof(f510,plain,(
  ! [X2,X0,X1] : ((~sP1(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP1(X2,X0,X1)))),
  inference(nnf_transformation,[],[f460])).
fof(f460,plain,(
  ! [X2,X0,X1] : (sP1(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f21620,plain,(
  ( ! [X37,X38] : (~sP1(X37,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X38) | ~p(s(bool,X38)) | ~p(s(bool,X37))) ) | $spl69),
  inference(resolution,[],[f1824,f714])).
fof(f714,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP1(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f512])).
fof(f32099,plain,(
  ( ! [X4,X3] : (sP1('const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X3,X4)) ) | $spl478),
  inference(cnf_transformation,[],[f32099_D])).
fof(f32099_D,plain,(
  ( ! [X4,X3] : (sP1('const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X3,X4)) ) <=> ~$spl478),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl478])])).
fof(f32094,plain,(
  ( ! [X19,X17,X18,X16] : (sP1('const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X16,X17) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X18),X19),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))),s(fun(X18,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1199,f678])).
fof(f1199,plain,(
  ( ! [X41,X42] : (sP1('const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18)))),X41,X42)) )),
  inference(resolution,[],[f648,f717])).
fof(f717,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f512])).
fof(f32093,plain,(
  ( ! [X14,X12,X15,X13] : (sP1('const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X12,X13) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X14,sK16),X15),s(fun(X14,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))))))) )),
  inference(superposition,[],[f1199,f680])).
fof(f32091,plain,(
  ( ! [X6,X4,X5,X3] : (sP1('const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X3,X4) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK16,X5),X6),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))),s(fun(X5,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1199,f948])).
fof(f31469,plain,(
  ( ! [X4,X3] : (sP1(X3,X4,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) | $spl476),
  inference(cnf_transformation,[],[f31469_D])).
fof(f31469_D,plain,(
  ( ! [X4,X3] : (sP1(X3,X4,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) <=> ~$spl476),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl476])])).
fof(f31463,plain,(
  ( ! [X19,X17,X18,X16] : (sP1(X16,X17,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X18),X19),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))),s(fun(X18,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1198,f678])).
fof(f1198,plain,(
  ( ! [X39,X40] : (sP1(X39,X40,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18)))))) )),
  inference(resolution,[],[f648,f715])).
fof(f715,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f512])).
fof(f31462,plain,(
  ( ! [X14,X12,X15,X13] : (sP1(X12,X13,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X14,sK16),X15),s(fun(X14,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))))))) )),
  inference(superposition,[],[f1198,f680])).
fof(f31460,plain,(
  ( ! [X6,X4,X5,X3] : (sP1(X3,X4,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK16,X5),X6),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))),s(fun(X5,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1198,f948])).
fof(f31335,plain,(
  ( ! [X13] : (p(s(bool,X13)) | s(bool,X13) != s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18)))))) ) | $spl474),
  inference(cnf_transformation,[],[f31335_D])).
fof(f31335_D,plain,(
  ( ! [X13] : (p(s(bool,X13)) | s(bool,X13) != s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18)))))) ) <=> ~$spl474),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl474])])).
fof(f31289,plain,(
  ( ! [X251,X250] : (p(s(bool,X250)) | p(s(bool,X251)) | s(bool,X251) != s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18)))))) )),
  inference(resolution,[],[f968,f957])).
fof(f957,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f703])).
fof(f703,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f506])).
fof(f506,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f505])).
fof(f505,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f424])).
fof(f424,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f423])).
fof(f423,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f309])).
fof(f309,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f63])).
fof(f63,axiom,(
  ! [X13,X28,X24,X29] : ((s(bool,X13) = s(bool,X28) & (p(s(bool,X28)) => s(bool,X24) = s(bool,X29))) => ((p(s(bool,X13)) => p(s(bool,X24))) <=> (p(s(bool,X28)) => p(s(bool,X29)))))),
  file('thm.util_prob.INJ_IMAGE_BIJ','thm.bool.IMP_CONG')).
fof(f1146,plain,(
  ( ! [X19,X20] : (s(bool,X19) != s(bool,X20) | ~p(s(bool,X19))) ) | $spl20),
  inference(cnf_transformation,[],[f1146_D])).
fof(f1146_D,plain,(
  ( ! [X19,X20] : (s(bool,X19) != s(bool,X20) | ~p(s(bool,X19))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f31182,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X33) != s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))))) | s(bool,X34) != s(bool,X35)) )),
  inference(resolution,[],[f968,f706])).
fof(f706,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f506])).
fof(f31179,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X26) != s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18)))))) )),
  inference(resolution,[],[f968,f705])).
fof(f705,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f506])).
fof(f1143,plain,(
  ( ! [X14,X12] : (s(bool,X12) != s(bool,X14) | p(s(bool,X12))) ) | $spl14),
  inference(cnf_transformation,[],[f1143_D])).
fof(f1143_D,plain,(
  ( ! [X14,X12] : (s(bool,X12) != s(bool,X14) | p(s(bool,X12))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f31178,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18)))))) )),
  inference(resolution,[],[f968,f704])).
fof(f704,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f506])).
fof(f31177,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))))) | s(bool,X19) != s(bool,X20)) )),
  inference(resolution,[],[f968,f702])).
fof(f702,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f506])).
fof(f31175,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18)))))) )),
  inference(resolution,[],[f968,f700])).
fof(f700,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f506])).
fof(f28219,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18)))))) | $spl465),
  inference(cnf_transformation,[],[f28219_D])).
fof(f28219_D,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18)))))) <=> ~$spl465),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl465])])).
fof(f31334,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18)))))) | $spl472),
  inference(cnf_transformation,[],[f31334_D])).
fof(f31334_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18)))))) <=> ~$spl472),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl472])])).
fof(f31332,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18)))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18)))))) | $spl6),
  inference(forward_demodulation,[],[f31173,f1137])).
fof(f31173,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18)))))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))))))),
  inference(resolution,[],[f968,f946])).
fof(f946,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X3))),s('type.num.num','const.pred_set.CARD_1'(s(fun(X1,bool),X4)))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X1,bool),X4)))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))) )),
  inference(cnf_transformation,[],[f456])).
fof(f456,plain,(
  ! [X0,X1,X2,X3,X4] : (~p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X1,bool),X4)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X3))),s('type.num.num','const.pred_set.CARD_1'(s(fun(X1,bool),X4)))))))),
  inference(flattening,[],[f455])).
fof(f455,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X1,bool),X4))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X3))),s('type.num.num','const.pred_set.CARD_1'(s(fun(X1,bool),X4)))))))),
  inference(ennf_transformation,[],[f388])).
fof(f388,plain,(
  ! [X0,X1,X2,X3,X4] : ((p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) & p(s(bool,'const.pred_set.FINITE_1'(s(fun(X1,bool),X4))))) => p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X3))),s('type.num.num','const.pred_set.CARD_1'(s(fun(X1,bool),X4)))))))),
  inference(flattening,[],[f387])).
fof(f387,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : ((p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) & p(s(bool,'const.pred_set.FINITE_1'(s(fun(X1,bool),X4))))) => p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X3))),s('type.num.num','const.pred_set.CARD_1'(s(fun(X1,bool),X4)))))))),
  inference(rectify,[],[f223])).
fof(f223,axiom,(
  ! [X7,X6] : ! [X9,X8,X10] : ((p(s(bool,'const.pred_set.INJ_3'(s(fun(X7,X6),X9),s(fun(X7,bool),X8),s(fun(X6,bool),X10)))) & p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X10))))) => p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X7,bool),X8))),s('type.num.num','const.pred_set.CARD_1'(s(fun(X6,bool),X10)))))))),
  file('thm.util_prob.INJ_IMAGE_BIJ','thm.pred_set.INJ_CARD')).
fof(f24704,plain,(
  s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) = s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20))) | $spl456),
  inference(cnf_transformation,[],[f24704_D])).
fof(f24704_D,plain,(
  s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) = s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20))) <=> ~$spl456),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl456])])).
fof(f30313,plain,(
  ( ! [X19,X18] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) = s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20))) | p(s(bool,X18)) | ~p(s(bool,X19))) ) | $spl4),
  inference(resolution,[],[f29941,f1037])).
fof(f1037,plain,(
  ( ! [X138,X139] : (~sP8('const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)),X138,X139) | p(s(bool,X139)) | ~p(s(bool,X138))) )),
  inference(resolution,[],[f647,f769])).
fof(f29941,plain,(
  ( ! [X191,X192,X193] : (sP8(X191,X192,X193) | s(bool,X191) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) ) | $spl4),
  inference(resolution,[],[f28754,f772])).
fof(f28754,plain,(
  ( ! [X221] : (~p(s(bool,X221)) | s(bool,X221) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) ) | $spl4),
  inference(resolution,[],[f1134,f864])).
fof(f864,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f441])).
fof(f1134,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) | $spl4),
  inference(cnf_transformation,[],[f1134_D])).
fof(f1134_D,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f30302,plain,(
  ( ! [X17,X16] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) = s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20))) | p(s(bool,X16)) | ~p(s(bool,X17))) ) | $spl4),
  inference(resolution,[],[f29940,f1038])).
fof(f1038,plain,(
  ( ! [X140,X141] : (~sP8(X140,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)),X141) | p(s(bool,X141)) | ~p(s(bool,X140))) )),
  inference(resolution,[],[f647,f769])).
fof(f29940,plain,(
  ( ! [X189,X190,X188] : (sP8(X189,X188,X190) | s(bool,X188) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) ) | $spl4),
  inference(resolution,[],[f28754,f771])).
fof(f30292,plain,(
  ( ! [X19,X18] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) = s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20))) | p(s(bool,X18)) | ~p(s(bool,X19))) ) | $spl4),
  inference(resolution,[],[f29938,f1034])).
fof(f1034,plain,(
  ( ! [X132,X133] : (~sP9('const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)),X132,X133) | p(s(bool,X133)) | ~p(s(bool,X132))) )),
  inference(resolution,[],[f647,f765])).
fof(f29938,plain,(
  ( ! [X182,X184,X183] : (sP9(X182,X183,X184) | s(bool,X182) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) ) | $spl4),
  inference(resolution,[],[f28754,f768])).
fof(f30282,plain,(
  ( ! [X17,X16] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) = s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20))) | p(s(bool,X16)) | ~p(s(bool,X17))) ) | $spl4),
  inference(resolution,[],[f29937,f1035])).
fof(f1035,plain,(
  ( ! [X134,X135] : (~sP9(X134,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)),X135) | p(s(bool,X135)) | ~p(s(bool,X134))) )),
  inference(resolution,[],[f647,f765])).
fof(f29937,plain,(
  ( ! [X180,X181,X179] : (sP9(X180,X179,X181) | s(bool,X179) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) ) | $spl4),
  inference(resolution,[],[f28754,f767])).
fof(f30274,plain,(
  ( ! [X21,X22] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) = s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20))) | p(s(bool,X21)) | ~p(s(bool,X22))) ) | $spl4),
  inference(resolution,[],[f29925,f1021])).
fof(f1021,plain,(
  ( ! [X107,X106] : (~sP5('const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)),X106,X107) | p(s(bool,X107)) | ~p(s(bool,X106))) )),
  inference(resolution,[],[f647,f748])).
fof(f29925,plain,(
  ( ! [X144,X145,X143] : (sP5(X143,X144,X145) | s(bool,X143) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) ) | $spl4),
  inference(resolution,[],[f28754,f751])).
fof(f30262,plain,(
  ( ! [X19,X20] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) = s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20))) | p(s(bool,X19)) | ~p(s(bool,X20))) ) | $spl4),
  inference(resolution,[],[f29924,f1022])).
fof(f1022,plain,(
  ( ! [X109,X108] : (~sP5(X108,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)),X109) | p(s(bool,X109)) | ~p(s(bool,X108))) )),
  inference(resolution,[],[f647,f748])).
fof(f29924,plain,(
  ( ! [X142,X140,X141] : (sP5(X141,X140,X142) | s(bool,X140) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) ) | $spl4),
  inference(resolution,[],[f28754,f750])).
fof(f30254,plain,(
  ( ! [X24,X23] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) = s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20))) | p(s(bool,X23)) | p(s(bool,X24))) ) | $spl4),
  inference(resolution,[],[f29911,f1005])).
fof(f1005,plain,(
  ( ! [X74,X75] : (~sP3('const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)),X74,X75) | p(s(bool,X75)) | p(s(bool,X74))) )),
  inference(resolution,[],[f647,f731])).
fof(f29911,plain,(
  ( ! [X103,X101,X102] : (sP3(X101,X102,X103) | s(bool,X101) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) ) | $spl4),
  inference(resolution,[],[f28754,f734])).
fof(f30244,plain,(
  ( ! [X21,X22] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) = s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20))) | ~p(s(bool,X21)) | p(s(bool,X22))) ) | $spl4),
  inference(resolution,[],[f29895,f991])).
fof(f991,plain,(
  ( ! [X47,X46] : (~sP1('const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)),X46,X47) | ~p(s(bool,X47)) | p(s(bool,X46))) )),
  inference(resolution,[],[f647,f714])).
fof(f29895,plain,(
  ( ! [X54,X55,X53] : (sP1(X53,X54,X55) | s(bool,X53) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) ) | $spl4),
  inference(resolution,[],[f28754,f717])).
fof(f30231,plain,(
  ( ! [X15,X16] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) = s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20))) | p(s(bool,X15)) | ~p(s(bool,X16))) ) | $spl4),
  inference(resolution,[],[f29894,f992])).
fof(f992,plain,(
  ( ! [X48,X49] : (~sP1(X48,X49,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20))) | p(s(bool,X49)) | ~p(s(bool,X48))) )),
  inference(resolution,[],[f647,f714])).
fof(f29894,plain,(
  ( ! [X52,X50,X51] : (sP1(X51,X52,X50) | s(bool,X50) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) ) | $spl4),
  inference(resolution,[],[f28754,f715])).
fof(f24008,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) | $spl448),
  inference(cnf_transformation,[],[f24008_D])).
fof(f24008_D,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) <=> ~$spl448),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl448])])).
fof(f29989,plain,(
  ( ! [X333,X331,X332] : (s(bool,X331) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | p(s(bool,X332)) | ~p(s(bool,X333)) | s(bool,X331) != s(bool,X333)) ) | $spl4),
  inference(resolution,[],[f28754,f957])).
fof(f29887,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) ) | $spl4),
  inference(resolution,[],[f28754,f706])).
fof(f29885,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,X20) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X20) != s(bool,X22)) ) | $spl4),
  inference(resolution,[],[f28754,f705])).
fof(f29883,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl4),
  inference(resolution,[],[f28754,f704])).
fof(f29882,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) ) | $spl4),
  inference(resolution,[],[f28754,f702])).
fof(f29880,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl4),
  inference(resolution,[],[f28754,f700])).
fof(f28220,plain,(
  ( ! [X4,X5,X3] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(X3,bool),X4))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X3),X5),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))),s(fun(X3,bool),X4))))) ) | $spl466),
  inference(cnf_transformation,[],[f28220_D])).
fof(f28220_D,plain,(
  ( ! [X4,X5,X3] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(X3,bool),X4))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X3),X5),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))),s(fun(X3,bool),X4))))) ) <=> ~$spl466),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl466])])).
fof(f29007,plain,(
  ( ! [X12,X13,X11] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(X11,bool),X12))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X11),X13),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))),s(fun(X11,bool),X12)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))))))) ) | $spl6),
  inference(subsumption_resolution,[],[f28898,f676])).
fof(f676,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X1,bool),X4)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X3)))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))) )),
  inference(cnf_transformation,[],[f410])).
fof(f410,plain,(
  ! [X0,X1,X2,X3,X4] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X3)))) | p(s(bool,'const.pred_set.FINITE_1'(s(fun(X1,bool),X4)))))),
  inference(flattening,[],[f409])).
fof(f409,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X3))))) | p(s(bool,'const.pred_set.FINITE_1'(s(fun(X1,bool),X4)))))),
  inference(ennf_transformation,[],[f274])).
fof(f274,plain,(
  ! [X0,X1,X2,X3,X4] : ((p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) & p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X3))))) => p(s(bool,'const.pred_set.FINITE_1'(s(fun(X1,bool),X4)))))),
  inference(flattening,[],[f273])).
fof(f273,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : ((p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) & p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X3))))) => p(s(bool,'const.pred_set.FINITE_1'(s(fun(X1,bool),X4)))))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X7,X6] : ! [X9,X8,X10] : ((p(s(bool,'const.pred_set.BIJ_3'(s(fun(X7,X6),X9),s(fun(X7,bool),X8),s(fun(X6,bool),X10)))) & p(s(bool,'const.pred_set.FINITE_1'(s(fun(X7,bool),X8))))) => p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X10)))))),
  file('thm.util_prob.INJ_IMAGE_BIJ','thm.pred_set.BIJ_FINITE')).
fof(f28898,plain,(
  ( ! [X12,X13,X11] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(X11,bool),X12))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X11,bool),X12)))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X11),X13),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))),s(fun(X11,bool),X12)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))))))) ) | $spl6),
  inference(superposition,[],[f682,f1137])).
fof(f682,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X2))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(X1,bool),X3))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X1,bool),X3)))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X4),s(fun(X0,bool),X2),s(fun(X1,bool),X3)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f413])).
fof(f413,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))) | ! [X3,X4] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X4),s(fun(X0,bool),X2),s(fun(X1,bool),X3)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X1,bool),X3)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X2))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(X1,bool),X3)))))),
  inference(flattening,[],[f412])).
fof(f412,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))) | ! [X3,X4] : ((~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X4),s(fun(X0,bool),X2),s(fun(X1,bool),X3)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X1,bool),X3))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X2))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(X1,bool),X3)))))),
  inference(ennf_transformation,[],[f282])).
fof(f282,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))) => ! [X3,X4] : ((p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X4),s(fun(X0,bool),X2),s(fun(X1,bool),X3)))) & p(s(bool,'const.pred_set.FINITE_1'(s(fun(X1,bool),X3))))) => s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X2))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(X1,bool),X3)))))),
  inference(flattening,[],[f281])).
fof(f281,plain,(
  ! [X0,X1] : ! [X2] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))) => ! [X3,X4] : ((p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X4),s(fun(X0,bool),X2),s(fun(X1,bool),X3)))) & p(s(bool,'const.pred_set.FINITE_1'(s(fun(X1,bool),X3))))) => s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X2))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(X1,bool),X3)))))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X7,X6] : ! [X14] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X7,bool),X14)))) => ! [X10,X9] : ((p(s(bool,'const.pred_set.BIJ_3'(s(fun(X7,X6),X9),s(fun(X7,bool),X14),s(fun(X6,bool),X10)))) & p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X10))))) => s('type.num.num','const.pred_set.CARD_1'(s(fun(X6,bool),X10))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(X7,bool),X14)))))),
  file('thm.util_prob.INJ_IMAGE_BIJ','thm.pred_set.FINITE_BIJ_CARD_EQ')).
fof(f28231,plain,(
  ( ! [X6,X7] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X6,bool),'const.pred_set.IMAGE_2'(s(fun(sK16,X6),X7),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18))))))) ) | $spl470),
  inference(cnf_transformation,[],[f28231_D])).
fof(f28231_D,plain,(
  ( ! [X6,X7] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X6,bool),'const.pred_set.IMAGE_2'(s(fun(sK16,X6),X7),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18))))))) ) <=> ~$spl470),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl470])])).
fof(f28896,plain,(
  ( ! [X6,X7] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X6,bool),'const.pred_set.IMAGE_2'(s(fun(sK16,X6),X7),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18)))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))))))) ) | $spl6),
  inference(superposition,[],[f695,f1137])).
fof(f695,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X1,X0),X2),s(fun(X1,bool),X3))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(X1,bool),X3)))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X1,bool),X3))))) )),
  inference(cnf_transformation,[],[f421])).
fof(f421,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X1,bool),X3)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X1,X0),X2),s(fun(X1,bool),X3))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(X1,bool),X3)))))))),
  inference(ennf_transformation,[],[f302])).
fof(f302,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X1,bool),X3)))) => p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X1,X0),X2),s(fun(X1,bool),X3))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(X1,bool),X3)))))))),
  inference(flattening,[],[f301])).
fof(f301,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X1,bool),X3)))) => p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X1,X0),X2),s(fun(X1,bool),X3))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(X1,bool),X3)))))))),
  inference(rectify,[],[f222])).
fof(f222,axiom,(
  ! [X6,X7] : ! [X9,X8] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X7,bool),X8)))) => p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X6,bool),'const.pred_set.IMAGE_2'(s(fun(X7,X6),X9),s(fun(X7,bool),X8))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(X7,bool),X8)))))))),
  file('thm.util_prob.INJ_IMAGE_BIJ','thm.pred_set.CARD_IMAGE')).
fof(f28230,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18)))))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,sK16),X2),s(fun(X0,bool),X1),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))))))) ) | $spl468),
  inference(cnf_transformation,[],[f28230_D])).
fof(f28230_D,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18)))))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,sK16),X2),s(fun(X0,bool),X1),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))))))) ) <=> ~$spl468),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl468])])).
fof(f28894,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18)))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18)))))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,sK16),X2),s(fun(X0,bool),X1),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))))))) ) | $spl6),
  inference(superposition,[],[f946,f1137])).
fof(f28996,plain,(
  ( ! [X4,X5,X3] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(X3,bool),X4))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X3),X5),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))),s(fun(X3,bool),X4)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))))))) ) | $spl6),
  inference(subsumption_resolution,[],[f28888,f676])).
fof(f28888,plain,(
  ( ! [X4,X5,X3] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(X3,bool),X4))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X3,bool),X4)))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X3),X5),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))),s(fun(X3,bool),X4)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))))))) ) | $spl6),
  inference(superposition,[],[f1137,f682])).
fof(f28215,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18))) | $spl462),
  inference(cnf_transformation,[],[f28215_D])).
fof(f28215_D,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18))) <=> ~$spl462),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl462])])).
fof(f28882,plain,(
  ( ! [X8,X9] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X8),X9),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))),s(fun(X8,bool),'const.pred_set.EMPTY_0'))))) ) | $spl6),
  inference(superposition,[],[f1137,f678])).
fof(f28881,plain,(
  ( ! [X6,X7] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X6,sK16),X7),s(fun(X6,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))))))) ) | $spl6),
  inference(superposition,[],[f1137,f680])).
fof(f28879,plain,(
  ( ! [X2,X1] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK16,X1),X2),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))),s(fun(X1,bool),'const.pred_set.EMPTY_0'))))) ) | $spl6),
  inference(superposition,[],[f1137,f948])).
fof(f1975,plain,(
  ( ! [X22] : (s(bool,X22) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | p(s(bool,X22))) ) | $spl70),
  inference(cnf_transformation,[],[f1975_D])).
fof(f1975_D,plain,(
  ( ! [X22] : (s(bool,X22) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | p(s(bool,X22))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f28768,plain,(
  ( ! [X259,X260] : (p(s(bool,X259)) | p(s(bool,X260)) | s(bool,X260) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) ) | $spl4),
  inference(resolution,[],[f1134,f957])).
fof(f28661,plain,(
  ( ! [X43,X44,X42] : (p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X42) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | s(bool,X43) != s(bool,X44)) ) | $spl4),
  inference(resolution,[],[f1134,f706])).
fof(f28658,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X33)) | ~p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X35) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) ) | $spl4),
  inference(resolution,[],[f1134,f705])).
fof(f28657,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | p(s(bool,X31)) | s(bool,X30) != s(bool,X32) | s(bool,X31) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) ) | $spl4),
  inference(resolution,[],[f1134,f704])).
fof(f28656,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | s(bool,X28) != s(bool,X29)) ) | $spl4),
  inference(resolution,[],[f1134,f702])).
fof(f28654,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) ) | $spl4),
  inference(resolution,[],[f1134,f700])).
fof(f28460,plain,(
  $false | $spl22),
  inference(resolution,[],[f1147,f950])).
fof(f950,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(X1,bool),X3))))) )),
  inference(cnf_transformation,[],[f394])).
fof(f394,plain,(
  ! [X0,X1,X2,X3] : p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(X1,bool),X3))))),
  inference(flattening,[],[f393])).
fof(f393,plain,(
  ! [X0,X1] : ! [X2,X3] : p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(X1,bool),X3))))),
  inference(rectify,[],[f67])).
fof(f67,axiom,(
  ! [X7,X6] : ! [X9,X8] : p(s(bool,'const.pred_set.INJ_3'(s(fun(X7,X6),X9),s(fun(X7,bool),'const.pred_set.EMPTY_0'),s(fun(X6,bool),X8))))),
  file('thm.util_prob.INJ_IMAGE_BIJ','thm.pred_set.INJ_EMPTY|split|0')).
fof(f28466,plain,(
  $false | $spl22),
  inference(resolution,[],[f1147,f951])).
fof(f951,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f396])).
fof(f396,plain,(
  ! [X0,X1] : p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X1))))),
  inference(flattening,[],[f395])).
fof(f395,plain,(
  ! [X0] : ! [X1] : p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X1))))),
  inference(rectify,[],[f88])).
fof(f88,axiom,(
  ! [X7] : ! [X8] : p(s(bool,'const.pred_set.INJ_3'(s(fun(X7,X7),'_i_0_0'),s(fun(X7,bool),X8),s(fun(X7,bool),X8))))),
  file('thm.util_prob.INJ_IMAGE_BIJ','thm.pred_set.INJ_ID')).
fof(f28467,plain,(
  $false | $spl22),
  inference(resolution,[],[f1147,f647])).
fof(f28483,plain,(
  $false | $spl22),
  inference(resolution,[],[f1147,f686])).
fof(f686,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f290])).
fof(f290,plain,(
  ! [X0,X1] : p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X1))))),
  inference(flattening,[],[f289])).
fof(f289,plain,(
  ! [X0] : ! [X1] : p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X1))))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X7] : ! [X8] : p(s(bool,'const.pred_set.BIJ_3'(s(fun(X7,X7),'_i_0_0'),s(fun(X7,bool),X8),s(fun(X7,bool),X8))))),
  file('thm.util_prob.INJ_IMAGE_BIJ','thm.pred_set.BIJ_ID')).
fof(f28485,plain,(
  $false | ($spl22 | $spl50)),
  inference(resolution,[],[f1147,f1357])).
fof(f1357,plain,(
  p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | $spl50),
  inference(cnf_transformation,[],[f1357_D])).
fof(f1357_D,plain,(
  p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f28488,plain,(
  $false | $spl22),
  inference(resolution,[],[f1147,f698])).
fof(f698,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X0,X0),X1),s(fun(X0,bool),'const.util_prob.schroeder_close_2'(s(fun(X0,X0),X1),s(fun(X0,bool),X2))))),s(fun(X0,bool),'const.util_prob.schroeder_close_2'(s(fun(X0,X0),X1),s(fun(X0,bool),X2))))))) )),
  inference(cnf_transformation,[],[f308])).
fof(f308,plain,(
  ! [X0,X1,X2] : p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X0,X0),X1),s(fun(X0,bool),'const.util_prob.schroeder_close_2'(s(fun(X0,X0),X1),s(fun(X0,bool),X2))))),s(fun(X0,bool),'const.util_prob.schroeder_close_2'(s(fun(X0,X0),X1),s(fun(X0,bool),X2))))))),
  inference(flattening,[],[f307])).
fof(f307,plain,(
  ! [X0] : ! [X1,X2] : p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X0,X0),X1),s(fun(X0,bool),'const.util_prob.schroeder_close_2'(s(fun(X0,X0),X1),s(fun(X0,bool),X2))))),s(fun(X0,bool),'const.util_prob.schroeder_close_2'(s(fun(X0,X0),X1),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f142])).
fof(f142,axiom,(
  ! [X7] : ! [X9,X8] : p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X7,bool),'const.pred_set.IMAGE_2'(s(fun(X7,X7),X9),s(fun(X7,bool),'const.util_prob.schroeder_close_2'(s(fun(X7,X7),X9),s(fun(X7,bool),X8))))),s(fun(X7,bool),'const.util_prob.schroeder_close_2'(s(fun(X7,X7),X9),s(fun(X7,bool),X8))))))),
  file('thm.util_prob.INJ_IMAGE_BIJ','thm.util_prob.SCHROEDER_CLOSED')).
fof(f28522,plain,(
  $false | ($spl22 | $spl330)),
  inference(resolution,[],[f1147,f9492])).
fof(f28558,plain,(
  $false | ($spl10 | $spl22)),
  inference(resolution,[],[f1147,f1141])).
fof(f28233,plain,(
  ( ! [X12,X13,X11] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(X11,bool),X12))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X11),X13),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))),s(fun(X11,bool),X12)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))))))) ) | $spl6),
  inference(subsumption_resolution,[],[f28117,f676])).
fof(f28117,plain,(
  ( ! [X12,X13,X11] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(X11,bool),X12))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X11,bool),X12)))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X11),X13),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))),s(fun(X11,bool),X12)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))))))) ) | $spl6),
  inference(superposition,[],[f682,f1137])).
fof(f28115,plain,(
  ( ! [X6,X7] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X6,bool),'const.pred_set.IMAGE_2'(s(fun(sK16,X6),X7),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18)))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))))))) ) | $spl6),
  inference(superposition,[],[f695,f1137])).
fof(f28113,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18)))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18)))))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,sK16),X2),s(fun(X0,bool),X1),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))))))) ) | $spl6),
  inference(superposition,[],[f946,f1137])).
fof(f28217,plain,(
  ( ! [X4,X5,X3] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(X3,bool),X4))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X3),X5),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))),s(fun(X3,bool),X4)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))))))) ) | $spl6),
  inference(subsumption_resolution,[],[f28107,f676])).
fof(f28107,plain,(
  ( ! [X4,X5,X3] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(X3,bool),X4))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X3,bool),X4)))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X3),X5),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))),s(fun(X3,bool),X4)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))))))) ) | $spl6),
  inference(superposition,[],[f1137,f682])).
fof(f28101,plain,(
  ( ! [X8,X9] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X8),X9),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))),s(fun(X8,bool),'const.pred_set.EMPTY_0'))))) ) | $spl6),
  inference(superposition,[],[f1137,f678])).
fof(f28100,plain,(
  ( ! [X6,X7] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X6,sK16),X7),s(fun(X6,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))))))) ) | $spl6),
  inference(superposition,[],[f1137,f680])).
fof(f28098,plain,(
  ( ! [X2,X1] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK16,X1),X2),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))),s(fun(X1,bool),'const.pred_set.EMPTY_0'))))) ) | $spl6),
  inference(superposition,[],[f1137,f948])).
fof(f28005,plain,(
  ( ! [X43] : (p(s(bool,X43)) | s(bool,X43) != s(bool,'const.bool.IN_2'(s(sK17,sK61('const.pred_set.EMPTY_0',sK18,sK19,sK16,sK17)),s(fun(sK17,bool),sK18)))) ) | $spl460),
  inference(cnf_transformation,[],[f28005_D])).
fof(f28005_D,plain,(
  ( ! [X43] : (p(s(bool,X43)) | s(bool,X43) != s(bool,'const.bool.IN_2'(s(sK17,sK61('const.pred_set.EMPTY_0',sK18,sK19,sK16,sK17)),s(fun(sK17,bool),sK18)))) ) <=> ~$spl460),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl460])])).
fof(f27972,plain,(
  ( ! [X280,X281] : (p(s(bool,X280)) | p(s(bool,X281)) | s(bool,X281) != s(bool,'const.bool.IN_2'(s(sK17,sK61('const.pred_set.EMPTY_0',sK18,sK19,sK16,sK17)),s(fun(sK17,bool),sK18)))) ) | $spl330),
  inference(resolution,[],[f9492,f957])).
fof(f27865,plain,(
  ( ! [X64,X65,X63] : (p(s(bool,X63)) | ~p(s(bool,X64)) | s(bool,X63) != s(bool,'const.bool.IN_2'(s(sK17,sK61('const.pred_set.EMPTY_0',sK18,sK19,sK16,sK17)),s(fun(sK17,bool),sK18))) | s(bool,X64) != s(bool,X65)) ) | $spl330),
  inference(resolution,[],[f9492,f706])).
fof(f27862,plain,(
  ( ! [X54,X56,X55] : (p(s(bool,X54)) | ~p(s(bool,X55)) | p(s(bool,X56)) | s(bool,X56) != s(bool,'const.bool.IN_2'(s(sK17,sK61('const.pred_set.EMPTY_0',sK18,sK19,sK16,sK17)),s(fun(sK17,bool),sK18)))) ) | $spl330),
  inference(resolution,[],[f9492,f705])).
fof(f27861,plain,(
  ( ! [X52,X53,X51] : (p(s(bool,X51)) | p(s(bool,X52)) | s(bool,X51) != s(bool,X53) | s(bool,X52) != s(bool,'const.bool.IN_2'(s(sK17,sK61('const.pred_set.EMPTY_0',sK18,sK19,sK16,sK17)),s(fun(sK17,bool),sK18)))) ) | $spl330),
  inference(resolution,[],[f9492,f704])).
fof(f27860,plain,(
  ( ! [X50,X48,X49] : (p(s(bool,X48)) | ~p(s(bool,X49)) | s(bool,X48) != s(bool,'const.bool.IN_2'(s(sK17,sK61('const.pred_set.EMPTY_0',sK18,sK19,sK16,sK17)),s(fun(sK17,bool),sK18))) | s(bool,X49) != s(bool,X50)) ) | $spl330),
  inference(resolution,[],[f9492,f702])).
fof(f27858,plain,(
  ( ! [X43,X44,X42] : (p(s(bool,X42)) | p(s(bool,X43)) | s(bool,X42) != s(bool,X44) | s(bool,X43) != s(bool,'const.bool.IN_2'(s(sK17,sK61('const.pred_set.EMPTY_0',sK18,sK19,sK16,sK17)),s(fun(sK17,bool),sK18)))) ) | $spl330),
  inference(resolution,[],[f9492,f700])).
fof(f27732,plain,(
  ( ! [X23,X22] : (p(s(bool,X22)) | p(s(bool,X23)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | ($spl69 | $spl77)),
  inference(resolution,[],[f4595,f24325])).
fof(f24325,plain,(
  ( ! [X196,X194,X195] : (sP7(X195,X196,X194) | s(bool,X194) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | $spl69),
  inference(resolution,[],[f21708,f774])).
fof(f21708,plain,(
  ( ! [X205] : (p(s(bool,X205)) | s(bool,X205) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | $spl69),
  inference(resolution,[],[f1824,f861])).
fof(f4595,plain,(
  ( ! [X140,X139] : (~sP7(X139,X140,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X140)) | p(s(bool,X139))) ) | $spl77),
  inference(resolution,[],[f2297,f773])).
fof(f24188,plain,(
  s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) | $spl450),
  inference(cnf_transformation,[],[f24188_D])).
fof(f24188_D,plain,(
  s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) <=> ~$spl450),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl450])])).
fof(f27726,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)))) ) | ($spl9 | $spl77)),
  inference(resolution,[],[f4595,f21141])).
fof(f21141,plain,(
  ( ! [X196,X194,X195] : (sP7(X195,X196,X194) | s(bool,X194) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)))) ) | $spl9),
  inference(resolution,[],[f19811,f774])).
fof(f19811,plain,(
  ( ! [X205] : (p(s(bool,X205)) | s(bool,X205) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)))) ) | $spl9),
  inference(resolution,[],[f1139,f861])).
fof(f27703,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | p(s(bool,X25)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | ($spl69 | $spl77)),
  inference(resolution,[],[f4594,f24327])).
fof(f24327,plain,(
  ( ! [X202,X200,X201] : (sP7(X200,X201,X202) | s(bool,X200) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | $spl69),
  inference(resolution,[],[f21708,f776])).
fof(f4594,plain,(
  ( ! [X138,X137] : (~sP7('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X137,X138) | p(s(bool,X138)) | p(s(bool,X137))) ) | $spl77),
  inference(resolution,[],[f2297,f773])).
fof(f27697,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)))) ) | ($spl9 | $spl77)),
  inference(resolution,[],[f4594,f21143])).
fof(f21143,plain,(
  ( ! [X202,X200,X201] : (sP7(X200,X201,X202) | s(bool,X200) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)))) ) | $spl9),
  inference(resolution,[],[f19811,f776])).
fof(f27672,plain,(
  ( ! [X23,X22] : (p(s(bool,X22)) | p(s(bool,X23)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | ($spl69 | $spl77)),
  inference(resolution,[],[f4593,f24326])).
fof(f24326,plain,(
  ( ! [X198,X199,X197] : (sP7(X198,X197,X199) | s(bool,X197) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | $spl69),
  inference(resolution,[],[f21708,f775])).
fof(f4593,plain,(
  ( ! [X136,X135] : (~sP7(X135,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X136) | p(s(bool,X136)) | p(s(bool,X135))) ) | $spl77),
  inference(resolution,[],[f2297,f773])).
fof(f27666,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)))) ) | ($spl9 | $spl77)),
  inference(resolution,[],[f4593,f21142])).
fof(f21142,plain,(
  ( ! [X198,X199,X197] : (sP7(X198,X197,X199) | s(bool,X197) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)))) ) | $spl9),
  inference(resolution,[],[f19811,f775])).
fof(f27600,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | ($spl69 | $spl77)),
  inference(resolution,[],[f4590,f24324])).
fof(f24324,plain,(
  ( ! [X191,X192,X193] : (sP8(X192,X193,X191) | s(bool,X191) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | $spl69),
  inference(resolution,[],[f21708,f770])).
fof(f4590,plain,(
  ( ! [X130,X129] : (~sP8(X129,X130,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X130)) | ~p(s(bool,X129))) ) | $spl77),
  inference(resolution,[],[f2297,f769])).
fof(f27596,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)))) ) | ($spl9 | $spl77)),
  inference(resolution,[],[f4590,f21140])).
fof(f21140,plain,(
  ( ! [X191,X192,X193] : (sP8(X192,X193,X191) | s(bool,X191) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)))) ) | $spl9),
  inference(resolution,[],[f19811,f770])).
fof(f27569,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | ($spl69 | $spl77)),
  inference(resolution,[],[f4587,f24321])).
fof(f24321,plain,(
  ( ! [X182,X184,X183] : (sP9(X183,X184,X182) | s(bool,X182) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | $spl69),
  inference(resolution,[],[f21708,f766])).
fof(f4587,plain,(
  ( ! [X123,X124] : (~sP9(X123,X124,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X124)) | ~p(s(bool,X123))) ) | $spl77),
  inference(resolution,[],[f2297,f765])).
fof(f27565,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)))) ) | ($spl9 | $spl77)),
  inference(resolution,[],[f4587,f21137])).
fof(f21137,plain,(
  ( ! [X182,X184,X183] : (sP9(X183,X184,X182) | s(bool,X182) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)))) ) | $spl9),
  inference(resolution,[],[f19811,f766])).
fof(f27476,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | ($spl69 | $spl77)),
  inference(resolution,[],[f4574,f24308])).
fof(f24308,plain,(
  ( ! [X144,X145,X143] : (sP5(X144,X145,X143) | s(bool,X143) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | $spl69),
  inference(resolution,[],[f21708,f749])).
fof(f4574,plain,(
  ( ! [X97,X98] : (~sP5(X97,X98,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X98)) | ~p(s(bool,X97))) ) | $spl77),
  inference(resolution,[],[f2297,f748])).
fof(f27472,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)))) ) | ($spl9 | $spl77)),
  inference(resolution,[],[f4574,f21124])).
fof(f21124,plain,(
  ( ! [X144,X145,X143] : (sP5(X144,X145,X143) | s(bool,X143) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)))) ) | $spl9),
  inference(resolution,[],[f19811,f749])).
fof(f27192,plain,(
  ( ! [X19,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | ($spl69 | $spl77)),
  inference(resolution,[],[f4561,f24291])).
fof(f24291,plain,(
  ( ! [X94,X92,X93] : (sP3(X93,X94,X92) | s(bool,X92) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | $spl69),
  inference(resolution,[],[f21708,f732])).
fof(f4561,plain,(
  ( ! [X72,X71] : (~sP3(X71,X72,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X72)) | ~p(s(bool,X71))) ) | $spl77),
  inference(resolution,[],[f2297,f731])).
fof(f27187,plain,(
  ( ! [X8,X9] : (p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)))) ) | ($spl9 | $spl77)),
  inference(resolution,[],[f4561,f21107])).
fof(f21107,plain,(
  ( ! [X94,X92,X93] : (sP3(X93,X94,X92) | s(bool,X92) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)))) ) | $spl9),
  inference(resolution,[],[f19811,f732])).
fof(f27165,plain,(
  ( ! [X21,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | ($spl69 | $spl77)),
  inference(resolution,[],[f4560,f24292])).
fof(f24292,plain,(
  ( ! [X97,X95,X96] : (sP3(X96,X95,X97) | s(bool,X95) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | $spl69),
  inference(resolution,[],[f21708,f733])).
fof(f4560,plain,(
  ( ! [X70,X69] : (~sP3(X69,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X70) | p(s(bool,X70)) | ~p(s(bool,X69))) ) | $spl77),
  inference(resolution,[],[f2297,f731])).
fof(f27160,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)))) ) | ($spl9 | $spl77)),
  inference(resolution,[],[f4560,f21108])).
fof(f21108,plain,(
  ( ! [X97,X95,X96] : (sP3(X96,X95,X97) | s(bool,X95) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)))) ) | $spl9),
  inference(resolution,[],[f19811,f733])).
fof(f27012,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) = s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) ) | $spl4),
  inference(resolution,[],[f1038,f23827])).
fof(f23827,plain,(
  ( ! [X189,X190,X188] : (sP8(X189,X188,X190) | s(bool,X188) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) ) | $spl4),
  inference(resolution,[],[f22538,f771])).
fof(f22538,plain,(
  ( ! [X221] : (~p(s(bool,X221)) | s(bool,X221) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) ) | $spl4),
  inference(resolution,[],[f1134,f864])).
fof(f26985,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | ($spl69 | $spl77)),
  inference(resolution,[],[f4544,f24278])).
fof(f24278,plain,(
  ( ! [X54,X55,X53] : (sP1(X54,X53,X55) | s(bool,X53) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | $spl69),
  inference(resolution,[],[f21708,f716])).
fof(f4544,plain,(
  ( ! [X37,X38] : (~sP1(X37,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X38) | ~p(s(bool,X38)) | ~p(s(bool,X37))) ) | $spl77),
  inference(resolution,[],[f2297,f714])).
fof(f26981,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)))) ) | ($spl9 | $spl77)),
  inference(resolution,[],[f4544,f21094])).
fof(f21094,plain,(
  ( ! [X54,X55,X53] : (sP1(X54,X53,X55) | s(bool,X53) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)))) ) | $spl9),
  inference(resolution,[],[f19811,f716])).
fof(f26866,plain,(
  ( ! [X14,X15] : (p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) = s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) ) | $spl4),
  inference(resolution,[],[f1037,f23828])).
fof(f23828,plain,(
  ( ! [X191,X192,X193] : (sP8(X191,X192,X193) | s(bool,X191) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) ) | $spl4),
  inference(resolution,[],[f22538,f772])).
fof(f26708,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) = s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) ) | $spl4),
  inference(resolution,[],[f1035,f23824])).
fof(f23824,plain,(
  ( ! [X180,X181,X179] : (sP9(X180,X179,X181) | s(bool,X179) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) ) | $spl4),
  inference(resolution,[],[f22538,f767])).
fof(f26506,plain,(
  ( ! [X14,X15] : (p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) = s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) ) | $spl4),
  inference(resolution,[],[f1034,f23825])).
fof(f23825,plain,(
  ( ! [X182,X184,X183] : (sP9(X182,X183,X184) | s(bool,X182) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) ) | $spl4),
  inference(resolution,[],[f22538,f768])).
fof(f24190,plain,(
  s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | $spl452),
  inference(cnf_transformation,[],[f24190_D])).
fof(f24190_D,plain,(
  s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) <=> ~$spl452),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl452])])).
fof(f26069,plain,(
  ( ! [X23,X22] : (p(s(bool,X22)) | p(s(bool,X23)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | ($spl9 | $spl69)),
  inference(resolution,[],[f19774,f24325])).
fof(f19774,plain,(
  ( ! [X140,X139] : (~sP7(X139,X140,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) | p(s(bool,X140)) | p(s(bool,X139))) ) | $spl9),
  inference(resolution,[],[f1139,f773])).
fof(f26066,plain,(
  ( ! [X17,X16] : (p(s(bool,X16)) | p(s(bool,X17)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)))) ) | ($spl9 | $spl77)),
  inference(resolution,[],[f19774,f22011])).
fof(f25993,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | p(s(bool,X25)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | ($spl9 | $spl69)),
  inference(resolution,[],[f19773,f24327])).
fof(f19773,plain,(
  ( ! [X138,X137] : (~sP7('const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)),X137,X138) | p(s(bool,X138)) | p(s(bool,X137))) ) | $spl9),
  inference(resolution,[],[f1139,f773])).
fof(f25990,plain,(
  ( ! [X19,X18] : (p(s(bool,X18)) | p(s(bool,X19)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)))) ) | ($spl9 | $spl77)),
  inference(resolution,[],[f19773,f22013])).
fof(f25965,plain,(
  ( ! [X23,X22] : (p(s(bool,X22)) | p(s(bool,X23)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | ($spl9 | $spl69)),
  inference(resolution,[],[f19772,f24326])).
fof(f19772,plain,(
  ( ! [X136,X135] : (~sP7(X135,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)),X136) | p(s(bool,X136)) | p(s(bool,X135))) ) | $spl9),
  inference(resolution,[],[f1139,f773])).
fof(f25962,plain,(
  ( ! [X17,X16] : (p(s(bool,X16)) | p(s(bool,X17)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)))) ) | ($spl9 | $spl77)),
  inference(resolution,[],[f19772,f22012])).
fof(f25936,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | ($spl9 | $spl69)),
  inference(resolution,[],[f19769,f24324])).
fof(f19769,plain,(
  ( ! [X130,X129] : (~sP8(X129,X130,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) | ~p(s(bool,X130)) | ~p(s(bool,X129))) ) | $spl9),
  inference(resolution,[],[f1139,f769])).
fof(f25933,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)))) ) | ($spl9 | $spl77)),
  inference(resolution,[],[f19769,f22010])).
fof(f25908,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | ($spl9 | $spl69)),
  inference(resolution,[],[f19766,f24321])).
fof(f19766,plain,(
  ( ! [X123,X124] : (~sP9(X123,X124,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) | ~p(s(bool,X124)) | ~p(s(bool,X123))) ) | $spl9),
  inference(resolution,[],[f1139,f765])).
fof(f25905,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)))) ) | ($spl9 | $spl77)),
  inference(resolution,[],[f19766,f22007])).
fof(f25781,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | ($spl9 | $spl69)),
  inference(resolution,[],[f19753,f24308])).
fof(f19753,plain,(
  ( ! [X97,X98] : (~sP5(X97,X98,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) | ~p(s(bool,X98)) | ~p(s(bool,X97))) ) | $spl9),
  inference(resolution,[],[f1139,f748])).
fof(f25778,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)))) ) | ($spl9 | $spl77)),
  inference(resolution,[],[f19753,f21994])).
fof(f25574,plain,(
  ( ! [X19,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | ($spl9 | $spl69)),
  inference(resolution,[],[f19740,f24291])).
fof(f19740,plain,(
  ( ! [X72,X71] : (~sP3(X71,X72,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) | p(s(bool,X72)) | ~p(s(bool,X71))) ) | $spl9),
  inference(resolution,[],[f1139,f731])).
fof(f25571,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)))) ) | ($spl9 | $spl77)),
  inference(resolution,[],[f19740,f21977])).
fof(f25550,plain,(
  ( ! [X21,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | ($spl9 | $spl69)),
  inference(resolution,[],[f19739,f24292])).
fof(f19739,plain,(
  ( ! [X70,X69] : (~sP3(X69,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)),X70) | p(s(bool,X70)) | ~p(s(bool,X69))) ) | $spl9),
  inference(resolution,[],[f1139,f731])).
fof(f25547,plain,(
  ( ! [X14,X15] : (p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)))) ) | ($spl9 | $spl77)),
  inference(resolution,[],[f19739,f21978])).
fof(f25390,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | ($spl9 | $spl69)),
  inference(resolution,[],[f19723,f24278])).
fof(f19723,plain,(
  ( ! [X37,X38] : (~sP1(X37,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)),X38) | ~p(s(bool,X38)) | ~p(s(bool,X37))) ) | $spl9),
  inference(resolution,[],[f1139,f714])).
fof(f25387,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)))) ) | ($spl9 | $spl77)),
  inference(resolution,[],[f19723,f21964])).
fof(f25304,plain,(
  ( ! [X14] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X14))) ) | ($spl9 | $spl69)),
  inference(resolution,[],[f24352,f19795])).
fof(f19795,plain,(
  ( ! [X182,X181] : (~sP13(X182,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)),X181) | p(s(bool,X181))) ) | $spl9),
  inference(resolution,[],[f1139,f807])).
fof(f807,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | p(s(bool,X1)) | ~sP13(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f556])).
fof(f556,plain,(
  ! [X0,X1,X2] : ((~sP13(X0,X1,X2) | ((p(s(bool,X1)) | p(s(bool,X2))) & (p(s(bool,X1)) | p(s(bool,X0))))) & ((~p(s(bool,X1)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) & ~p(s(bool,X0))) | sP13(X0,X1,X2)))),
  inference(rectify,[],[f555])).
fof(f555,plain,(
  ! [X2,X0,X1] : ((~sP13(X2,X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X2))) | sP13(X2,X0,X1)))),
  inference(flattening,[],[f554])).
fof(f554,plain,(
  ! [X2,X0,X1] : ((~sP13(X2,X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2))))) & (((~p(s(bool,X0)) & ~p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X2)))) | sP13(X2,X0,X1)))),
  inference(nnf_transformation,[],[f478])).
fof(f478,plain,(
  ! [X2,X0,X1] : (sP13(X2,X0,X1) <=> ((p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP13])])).
fof(f24352,plain,(
  ( ! [X277,X275,X276] : (sP13(X276,X275,X277) | s(bool,X275) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | $spl69),
  inference(resolution,[],[f21708,f813])).
fof(f813,plain,(
  ( ! [X2,X0,X1] : (sP13(X2,X0,X1) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f558])).
fof(f558,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X0)) & (~p(s(bool,X1)) | ~p(s(bool,X2)))) | sP13(X2,X0,X1)) & (~sP13(X2,X0,X1) | p(s(bool,X0)) | (p(s(bool,X1)) & p(s(bool,X2)))))),
  inference(flattening,[],[f557])).
fof(f557,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X0)) & (~p(s(bool,X1)) | ~p(s(bool,X2)))) | sP13(X2,X0,X1)) & (~sP13(X2,X0,X1) | (p(s(bool,X0)) | (p(s(bool,X1)) & p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f479])).
fof(f479,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X0)) | (p(s(bool,X1)) & p(s(bool,X2)))) <=> sP13(X2,X0,X1))),
  inference(definition_folding,[],[f321,f478])).
fof(f321,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X0)) | (p(s(bool,X1)) & p(s(bool,X2)))) <=> ((p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(rectify,[],[f76])).
fof(f76,axiom,(
  ! [X7,X6,X16] : ((p(s(bool,X7)) | (p(s(bool,X6)) & p(s(bool,X16)))) <=> ((p(s(bool,X7)) | p(s(bool,X6))) & (p(s(bool,X7)) | p(s(bool,X16)))))),
  file('thm.util_prob.INJ_IMAGE_BIJ','thm.bool.LEFT_OR_OVER_AND')).
fof(f25303,plain,(
  ( ! [X13] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X13))) ) | ($spl9 | $spl69)),
  inference(resolution,[],[f24352,f19797])).
fof(f19797,plain,(
  ( ! [X185,X186] : (~sP13(X185,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)),X186) | p(s(bool,X185))) ) | $spl9),
  inference(resolution,[],[f1139,f808])).
fof(f808,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | ~sP13(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f556])).
fof(f25302,plain,(
  ( ! [X12] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X12))) ) | ($spl69 | $spl77)),
  inference(resolution,[],[f24352,f4616])).
fof(f4616,plain,(
  ( ! [X182,X181] : (~sP13(X182,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X181) | p(s(bool,X181))) ) | $spl77),
  inference(resolution,[],[f2297,f807])).
fof(f25301,plain,(
  ( ! [X11] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X11))) ) | ($spl69 | $spl77)),
  inference(resolution,[],[f24352,f4618])).
fof(f4618,plain,(
  ( ! [X185,X186] : (~sP13(X185,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X186) | p(s(bool,X185))) ) | $spl77),
  inference(resolution,[],[f2297,f808])).
fof(f25266,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) = s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) ) | $spl4),
  inference(resolution,[],[f1022,f23811])).
fof(f23811,plain,(
  ( ! [X142,X140,X141] : (sP5(X141,X140,X142) | s(bool,X140) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) ) | $spl4),
  inference(resolution,[],[f22538,f750])).
fof(f25257,plain,(
  ( ! [X17] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X17))) ) | ($spl9 | $spl69)),
  inference(resolution,[],[f24345,f19787])).
fof(f19787,plain,(
  ( ! [X165,X166] : (~sP12('const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)),X166,X165) | p(s(bool,X165))) ) | $spl9),
  inference(resolution,[],[f1139,f797])).
fof(f797,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X2)) | ~sP12(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f551])).
fof(f551,plain,(
  ! [X0,X1,X2] : ((~sP12(X0,X1,X2) | ((p(s(bool,X2)) | p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0))))) & ((~p(s(bool,X2)) & ~p(s(bool,X0))) | (~p(s(bool,X1)) & ~p(s(bool,X0))) | sP12(X0,X1,X2)))),
  inference(rectify,[],[f550])).
fof(f550,plain,(
  ! [X0,X2,X1] : ((~sP12(X0,X2,X1) | ((p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | p(s(bool,X0))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (~p(s(bool,X2)) & ~p(s(bool,X0))) | sP12(X0,X2,X1)))),
  inference(flattening,[],[f549])).
fof(f549,plain,(
  ! [X0,X2,X1] : ((~sP12(X0,X2,X1) | ((p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | p(s(bool,X0))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (~p(s(bool,X2)) & ~p(s(bool,X0)))) | sP12(X0,X2,X1)))),
  inference(nnf_transformation,[],[f476])).
fof(f476,plain,(
  ! [X0,X2,X1] : (sP12(X0,X2,X1) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | p(s(bool,X0)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP12])])).
fof(f24345,plain,(
  ( ! [X255,X254,X256] : (sP12(X254,X255,X256) | s(bool,X254) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | $spl69),
  inference(resolution,[],[f21708,f804])).
fof(f804,plain,(
  ( ! [X2,X0,X1] : (sP12(X0,X2,X1) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f553])).
fof(f553,plain,(
  ! [X0,X1,X2] : ((((~p(s(bool,X1)) | ~p(s(bool,X2))) & ~p(s(bool,X0))) | sP12(X0,X2,X1)) & (~sP12(X0,X2,X1) | (p(s(bool,X1)) & p(s(bool,X2))) | p(s(bool,X0))))),
  inference(flattening,[],[f552])).
fof(f552,plain,(
  ! [X0,X1,X2] : ((((~p(s(bool,X1)) | ~p(s(bool,X2))) & ~p(s(bool,X0))) | sP12(X0,X2,X1)) & (~sP12(X0,X2,X1) | ((p(s(bool,X1)) & p(s(bool,X2))) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f477])).
fof(f477,plain,(
  ! [X0,X1,X2] : (((p(s(bool,X1)) & p(s(bool,X2))) | p(s(bool,X0))) <=> sP12(X0,X2,X1))),
  inference(definition_folding,[],[f320,f476])).
fof(f320,plain,(
  ! [X0,X1,X2] : (((p(s(bool,X1)) & p(s(bool,X2))) | p(s(bool,X0))) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | p(s(bool,X0)))))),
  inference(rectify,[],[f85])).
fof(f85,axiom,(
  ! [X7,X6,X16] : (((p(s(bool,X6)) & p(s(bool,X16))) | p(s(bool,X7))) <=> ((p(s(bool,X6)) | p(s(bool,X7))) & (p(s(bool,X16)) | p(s(bool,X7)))))),
  file('thm.util_prob.INJ_IMAGE_BIJ','thm.bool.RIGHT_OR_OVER_AND')).
fof(f25256,plain,(
  ( ! [X16] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X16))) ) | ($spl9 | $spl69)),
  inference(resolution,[],[f24345,f19790])).
fof(f19790,plain,(
  ( ! [X171,X172] : (~sP12('const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)),X171,X172) | p(s(bool,X171))) ) | $spl9),
  inference(resolution,[],[f1139,f798])).
fof(f798,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | ~sP12(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f551])).
fof(f25255,plain,(
  ( ! [X15] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X15))) ) | ($spl69 | $spl77)),
  inference(resolution,[],[f24345,f4608])).
fof(f4608,plain,(
  ( ! [X165,X166] : (~sP12('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X166,X165) | p(s(bool,X165))) ) | $spl77),
  inference(resolution,[],[f2297,f797])).
fof(f25254,plain,(
  ( ! [X14] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X14))) ) | ($spl69 | $spl77)),
  inference(resolution,[],[f24345,f4611])).
fof(f4611,plain,(
  ( ! [X171,X172] : (~sP12('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X171,X172) | p(s(bool,X171))) ) | $spl77),
  inference(resolution,[],[f2297,f798])).
fof(f25241,plain,(
  ( ! [X15] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X15))) ) | ($spl9 | $spl69)),
  inference(resolution,[],[f24334,f19777])).
fof(f19777,plain,(
  ( ! [X146,X145] : (~sP11('const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)),X146,X145) | ~p(s(bool,X145))) ) | $spl9),
  inference(resolution,[],[f1139,f781])).
fof(f781,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X2)) | ~sP11(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f544])).
fof(f544,plain,(
  ! [X0,X1,X2] : ((~sP11(X0,X1,X2) | ((~p(s(bool,X2)) | p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0))))) & ((p(s(bool,X2)) & ~p(s(bool,X0))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | sP11(X0,X1,X2)))),
  inference(rectify,[],[f543])).
fof(f543,plain,(
  ! [X2,X1,X0] : ((~sP11(X2,X1,X0) | ((~p(s(bool,X0)) | p(s(bool,X2))) & (~p(s(bool,X1)) | p(s(bool,X2))))) & ((p(s(bool,X0)) & ~p(s(bool,X2))) | (p(s(bool,X1)) & ~p(s(bool,X2))) | sP11(X2,X1,X0)))),
  inference(flattening,[],[f542])).
fof(f542,plain,(
  ! [X2,X1,X0] : ((~sP11(X2,X1,X0) | ((~p(s(bool,X0)) | p(s(bool,X2))) & (~p(s(bool,X1)) | p(s(bool,X2))))) & (((p(s(bool,X0)) & ~p(s(bool,X2))) | (p(s(bool,X1)) & ~p(s(bool,X2)))) | sP11(X2,X1,X0)))),
  inference(nnf_transformation,[],[f474])).
fof(f474,plain,(
  ! [X2,X1,X0] : (sP11(X2,X1,X0) <=> ((~p(s(bool,X0)) | p(s(bool,X2))) & (~p(s(bool,X1)) | p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f24334,plain,(
  ( ! [X222,X223,X221] : (sP11(X221,X222,X223) | s(bool,X221) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | $spl69),
  inference(resolution,[],[f21708,f788])).
fof(f788,plain,(
  ( ! [X2,X0,X1] : (sP11(X2,X1,X0) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f546])).
fof(f546,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X0)) | p(s(bool,X1))) & ~p(s(bool,X2))) | sP11(X2,X1,X0)) & (~sP11(X2,X1,X0) | (~p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,X2))))),
  inference(flattening,[],[f545])).
fof(f545,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X0)) | p(s(bool,X1))) & ~p(s(bool,X2))) | sP11(X2,X1,X0)) & (~sP11(X2,X1,X0) | ((~p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,X2)))))),
  inference(nnf_transformation,[],[f475])).
fof(f475,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,X2))) <=> sP11(X2,X1,X0))),
  inference(definition_folding,[],[f426,f474])).
fof(f426,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,X2))) <=> ((~p(s(bool,X0)) | p(s(bool,X2))) & (~p(s(bool,X1)) | p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f318])).
fof(f318,plain,(
  ! [X0,X1,X2] : (((p(s(bool,X0)) | p(s(bool,X1))) => p(s(bool,X2))) <=> ((p(s(bool,X0)) => p(s(bool,X2))) & (p(s(bool,X1)) => p(s(bool,X2)))))),
  inference(rectify,[],[f151])).
fof(f151,axiom,(
  ! [X30,X33,X38] : (((p(s(bool,X30)) | p(s(bool,X33))) => p(s(bool,X38))) <=> ((p(s(bool,X30)) => p(s(bool,X38))) & (p(s(bool,X33)) => p(s(bool,X38)))))),
  file('thm.util_prob.INJ_IMAGE_BIJ','thm.bool.DISJ_IMP_THM')).
fof(f25240,plain,(
  ( ! [X14] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X14))) ) | ($spl9 | $spl69)),
  inference(resolution,[],[f24334,f19778])).
fof(f19778,plain,(
  ( ! [X147,X148] : (~sP11('const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)),X147,X148) | ~p(s(bool,X147))) ) | $spl9),
  inference(resolution,[],[f1139,f782])).
fof(f782,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | ~sP11(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f544])).
fof(f25239,plain,(
  ( ! [X13] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X13))) ) | ($spl69 | $spl77)),
  inference(resolution,[],[f24334,f4598])).
fof(f4598,plain,(
  ( ! [X146,X145] : (~sP11('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X146,X145) | ~p(s(bool,X145))) ) | $spl77),
  inference(resolution,[],[f2297,f781])).
fof(f25238,plain,(
  ( ! [X12] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X12))) ) | ($spl69 | $spl77)),
  inference(resolution,[],[f24334,f4599])).
fof(f4599,plain,(
  ( ! [X147,X148] : (~sP11('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X147,X148) | ~p(s(bool,X147))) ) | $spl77),
  inference(resolution,[],[f2297,f782])).
fof(f25187,plain,(
  ( ! [X14,X15] : (p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) = s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) ) | $spl4),
  inference(resolution,[],[f1021,f23812])).
fof(f23812,plain,(
  ( ! [X144,X145,X143] : (sP5(X143,X144,X145) | s(bool,X143) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) ) | $spl4),
  inference(resolution,[],[f22538,f751])).
fof(f25038,plain,(
  ( ! [X17,X16] : (p(s(bool,X16)) | p(s(bool,X17)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) = s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) ) | $spl4),
  inference(resolution,[],[f1005,f23798])).
fof(f23798,plain,(
  ( ! [X103,X101,X102] : (sP3(X101,X102,X103) | s(bool,X101) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) ) | $spl4),
  inference(resolution,[],[f22538,f734])).
fof(f25026,plain,(
  ( ! [X16] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X16))) ) | ($spl77 | $spl252)),
  inference(resolution,[],[f22038,f7427])).
fof(f22038,plain,(
  ( ! [X277,X275,X276] : (sP13(X276,X275,X277) | s(bool,X275) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) | $spl77),
  inference(resolution,[],[f4632,f813])).
fof(f25025,plain,(
  ( ! [X15] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X15))) ) | ($spl77 | $spl256)),
  inference(resolution,[],[f22038,f7456])).
fof(f25024,plain,(
  ( ! [X14] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) | p(s(bool,X14))) ) | ($spl9 | $spl77)),
  inference(resolution,[],[f22038,f19795])).
fof(f25023,plain,(
  ( ! [X13] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) | p(s(bool,X13))) ) | ($spl9 | $spl77)),
  inference(resolution,[],[f22038,f19797])).
fof(f25013,plain,(
  ( ! [X19] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X19))) ) | ($spl77 | $spl244)),
  inference(resolution,[],[f22031,f7049])).
fof(f22031,plain,(
  ( ! [X255,X254,X256] : (sP12(X254,X255,X256) | s(bool,X254) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) | $spl77),
  inference(resolution,[],[f4632,f804])).
fof(f25012,plain,(
  ( ! [X18] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X18))) ) | ($spl77 | $spl250)),
  inference(resolution,[],[f22031,f7251])).
fof(f25011,plain,(
  ( ! [X17] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) | p(s(bool,X17))) ) | ($spl9 | $spl77)),
  inference(resolution,[],[f22031,f19787])).
fof(f25010,plain,(
  ( ! [X16] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) | p(s(bool,X16))) ) | ($spl9 | $spl77)),
  inference(resolution,[],[f22031,f19790])).
fof(f24995,plain,(
  ( ! [X17] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X17))) ) | ($spl77 | $spl234)),
  inference(resolution,[],[f22020,f6618])).
fof(f22020,plain,(
  ( ! [X222,X223,X221] : (sP11(X221,X222,X223) | s(bool,X221) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) | $spl77),
  inference(resolution,[],[f4632,f788])).
fof(f24994,plain,(
  ( ! [X16] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X16))) ) | ($spl77 | $spl236)),
  inference(resolution,[],[f22020,f6811])).
fof(f24993,plain,(
  ( ! [X15] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) | ~p(s(bool,X15))) ) | ($spl9 | $spl77)),
  inference(resolution,[],[f22020,f19777])).
fof(f24992,plain,(
  ( ! [X14] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) | ~p(s(bool,X14))) ) | ($spl9 | $spl77)),
  inference(resolution,[],[f22020,f19778])).
fof(f24728,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) = s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) ) | $spl4),
  inference(resolution,[],[f992,f23781])).
fof(f23781,plain,(
  ( ! [X52,X50,X51] : (sP1(X51,X52,X50) | s(bool,X50) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) ) | $spl4),
  inference(resolution,[],[f22538,f715])).
fof(f24702,plain,(
  ( ! [X15,X16] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) = s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20))) | ~p(s(bool,X15)) | p(s(bool,X16))) ) | $spl4),
  inference(resolution,[],[f23782,f991])).
fof(f23782,plain,(
  ( ! [X54,X55,X53] : (sP1(X53,X54,X55) | s(bool,X53) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) ) | $spl4),
  inference(resolution,[],[f22538,f717])).
fof(f24535,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) | $spl454),
  inference(cnf_transformation,[],[f24535_D])).
fof(f24535_D,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) <=> ~$spl454),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl454])])).
fof(f24377,plain,(
  ( ! [X352,X354,X353] : (s(bool,X352) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X353)) | p(s(bool,X354)) | s(bool,X352) != s(bool,X354)) ) | $spl69),
  inference(resolution,[],[f21708,f957])).
fof(f24270,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) ) | $spl69),
  inference(resolution,[],[f21708,f706])).
fof(f24267,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,X16) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) ) | $spl69),
  inference(resolution,[],[f21708,f705])).
fof(f24266,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl69),
  inference(resolution,[],[f21708,f704])).
fof(f24265,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) ) | $spl69),
  inference(resolution,[],[f21708,f702])).
fof(f24263,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl69),
  inference(resolution,[],[f21708,f700])).
fof(f24219,plain,(
  ( ! [X16] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X16))) ) | ($spl9 | $spl252)),
  inference(resolution,[],[f21168,f7427])).
fof(f21168,plain,(
  ( ! [X277,X275,X276] : (sP13(X276,X275,X277) | s(bool,X275) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)))) ) | $spl9),
  inference(resolution,[],[f19811,f813])).
fof(f24218,plain,(
  ( ! [X15] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X15))) ) | ($spl9 | $spl256)),
  inference(resolution,[],[f21168,f7456])).
fof(f24215,plain,(
  ( ! [X12] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) | p(s(bool,X12))) ) | ($spl9 | $spl77)),
  inference(resolution,[],[f21168,f4616])).
fof(f24214,plain,(
  ( ! [X11] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) | p(s(bool,X11))) ) | ($spl9 | $spl77)),
  inference(resolution,[],[f21168,f4618])).
fof(f24206,plain,(
  ( ! [X19] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X19))) ) | ($spl9 | $spl244)),
  inference(resolution,[],[f21161,f7049])).
fof(f21161,plain,(
  ( ! [X255,X254,X256] : (sP12(X254,X255,X256) | s(bool,X254) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)))) ) | $spl9),
  inference(resolution,[],[f19811,f804])).
fof(f24205,plain,(
  ( ! [X18] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X18))) ) | ($spl9 | $spl250)),
  inference(resolution,[],[f21161,f7251])).
fof(f24202,plain,(
  ( ! [X15] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) | p(s(bool,X15))) ) | ($spl9 | $spl77)),
  inference(resolution,[],[f21161,f4608])).
fof(f24201,plain,(
  ( ! [X14] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) | p(s(bool,X14))) ) | ($spl9 | $spl77)),
  inference(resolution,[],[f21161,f4611])).
fof(f24186,plain,(
  ( ! [X17] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X17))) ) | ($spl9 | $spl234)),
  inference(resolution,[],[f21150,f6618])).
fof(f21150,plain,(
  ( ! [X222,X223,X221] : (sP11(X221,X222,X223) | s(bool,X221) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)))) ) | $spl9),
  inference(resolution,[],[f19811,f788])).
fof(f24185,plain,(
  ( ! [X16] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X16))) ) | ($spl9 | $spl236)),
  inference(resolution,[],[f21150,f6811])).
fof(f24182,plain,(
  ( ! [X13] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) | ~p(s(bool,X13))) ) | ($spl9 | $spl77)),
  inference(resolution,[],[f21150,f4598])).
fof(f24181,plain,(
  ( ! [X12] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) | ~p(s(bool,X12))) ) | ($spl9 | $spl77)),
  inference(resolution,[],[f21150,f4599])).
fof(f23876,plain,(
  ( ! [X333,X331,X332] : (s(bool,X331) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | p(s(bool,X332)) | ~p(s(bool,X333)) | s(bool,X331) != s(bool,X333)) ) | $spl4),
  inference(resolution,[],[f22538,f957])).
fof(f23774,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) ) | $spl4),
  inference(resolution,[],[f22538,f706])).
fof(f23772,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,X20) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X20) != s(bool,X22)) ) | $spl4),
  inference(resolution,[],[f22538,f705])).
fof(f23770,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl4),
  inference(resolution,[],[f22538,f704])).
fof(f23769,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) ) | $spl4),
  inference(resolution,[],[f22538,f702])).
fof(f23767,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl4),
  inference(resolution,[],[f22538,f700])).
fof(f23524,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20))) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) | $spl446),
  inference(cnf_transformation,[],[f23524_D])).
fof(f23524_D,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20))) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) <=> ~$spl446),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl446])])).
fof(f23392,plain,(
  ( ! [X333,X331,X332] : (s(bool,X331) = s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20))) | p(s(bool,X332)) | ~p(s(bool,X333)) | s(bool,X331) != s(bool,X333)) )),
  inference(resolution,[],[f1078,f957])).
fof(f1078,plain,(
  ( ! [X212] : (~p(s(bool,X212)) | s(bool,X212) = s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) )),
  inference(resolution,[],[f647,f864])).
fof(f23290,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20))) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) )),
  inference(resolution,[],[f1078,f706])).
fof(f23288,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,X20) = s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20))) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X20) != s(bool,X22)) )),
  inference(resolution,[],[f1078,f705])).
fof(f23286,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20))) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) )),
  inference(resolution,[],[f1078,f704])).
fof(f23285,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20))) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) )),
  inference(resolution,[],[f1078,f702])).
fof(f23283,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20))) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) )),
  inference(resolution,[],[f1078,f700])).
fof(f22552,plain,(
  ( ! [X259,X260] : (p(s(bool,X259)) | p(s(bool,X260)) | s(bool,X260) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) ) | $spl4),
  inference(resolution,[],[f1134,f957])).
fof(f22445,plain,(
  ( ! [X43,X44,X42] : (p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X42) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | s(bool,X43) != s(bool,X44)) ) | $spl4),
  inference(resolution,[],[f1134,f706])).
fof(f22442,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X33)) | ~p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X35) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) ) | $spl4),
  inference(resolution,[],[f1134,f705])).
fof(f22441,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | p(s(bool,X31)) | s(bool,X30) != s(bool,X32) | s(bool,X31) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) ) | $spl4),
  inference(resolution,[],[f1134,f704])).
fof(f22440,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | s(bool,X28) != s(bool,X29)) ) | $spl4),
  inference(resolution,[],[f1134,f702])).
fof(f22438,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) ) | $spl4),
  inference(resolution,[],[f1134,f700])).
fof(f22428,plain,(
  s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | $spl444),
  inference(cnf_transformation,[],[f22428_D])).
fof(f22428_D,plain,(
  s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) <=> ~$spl444),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl444])])).
fof(f22420,plain,(
  ( ! [X15] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | ~p(s(bool,X15))) ) | ($spl5 | $spl234)),
  inference(resolution,[],[f20731,f6618])).
fof(f20731,plain,(
  ( ! [X222,X223,X221] : (sP11(X221,X222,X223) | s(bool,X221) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) ) | $spl5),
  inference(resolution,[],[f18318,f788])).
fof(f18318,plain,(
  ( ! [X205] : (p(s(bool,X205)) | s(bool,X205) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) ) | $spl5),
  inference(resolution,[],[f1135,f861])).
fof(f22419,plain,(
  ( ! [X14] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | ~p(s(bool,X14))) ) | ($spl5 | $spl236)),
  inference(resolution,[],[f20731,f6811])).
fof(f22426,plain,(
  s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | $spl442),
  inference(cnf_transformation,[],[f22426_D])).
fof(f22426_D,plain,(
  s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) <=> ~$spl442),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl442])])).
fof(f22418,plain,(
  ( ! [X13] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | ~p(s(bool,X13))) ) | ($spl5 | $spl9)),
  inference(resolution,[],[f20731,f19777])).
fof(f22417,plain,(
  ( ! [X12] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | ~p(s(bool,X12))) ) | ($spl5 | $spl9)),
  inference(resolution,[],[f20731,f19778])).
fof(f22424,plain,(
  s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | $spl440),
  inference(cnf_transformation,[],[f22424_D])).
fof(f22424_D,plain,(
  s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) <=> ~$spl440),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl440])])).
fof(f22416,plain,(
  ( ! [X11] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | ~p(s(bool,X11))) ) | ($spl5 | $spl77)),
  inference(resolution,[],[f20731,f4598])).
fof(f22415,plain,(
  ( ! [X10] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | ~p(s(bool,X10))) ) | ($spl5 | $spl77)),
  inference(resolution,[],[f20731,f4599])).
fof(f22220,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) | $spl438),
  inference(cnf_transformation,[],[f22220_D])).
fof(f22220_D,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) <=> ~$spl438),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl438])])).
fof(f22063,plain,(
  ( ! [X352,X354,X353] : (s(bool,X352) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X353)) | p(s(bool,X354)) | s(bool,X352) != s(bool,X354)) ) | $spl77),
  inference(resolution,[],[f4632,f957])).
fof(f21956,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) ) | $spl77),
  inference(resolution,[],[f4632,f706])).
fof(f21953,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,X16) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) ) | $spl77),
  inference(resolution,[],[f4632,f705])).
fof(f21952,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl77),
  inference(resolution,[],[f4632,f704])).
fof(f21951,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) ) | $spl77),
  inference(resolution,[],[f4632,f702])).
fof(f21949,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl77),
  inference(resolution,[],[f4632,f700])).
fof(f3603,plain,(
  ( ! [X4] : (s(bool,X4) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X4))) ) | $spl98),
  inference(cnf_transformation,[],[f3603_D])).
fof(f3603_D,plain,(
  ( ! [X4] : (s(bool,X4) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X4))) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f21716,plain,(
  ( ! [X225,X226] : (p(s(bool,X225)) | ~p(s(bool,X226)) | s(bool,X226) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | $spl69),
  inference(resolution,[],[f1824,f957])).
fof(f21614,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | s(bool,X24) != s(bool,X26)) ) | $spl69),
  inference(resolution,[],[f1824,f706])).
fof(f21612,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | $spl69),
  inference(resolution,[],[f1824,f705])).
fof(f21610,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | $spl69),
  inference(resolution,[],[f1824,f704])).
fof(f21609,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | s(bool,X9) != s(bool,X11)) ) | $spl69),
  inference(resolution,[],[f1824,f702])).
fof(f21607,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | $spl69),
  inference(resolution,[],[f1824,f700])).
fof(f1824,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | $spl69),
  inference(cnf_transformation,[],[f1824_D])).
fof(f1824_D,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) <=> ~$spl69),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl69])])).
fof(f14549,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl80),
  inference(resolution,[],[f2633,f714])).
fof(f21604,plain,(
  ( ! [X8,X9] : (sP2(X8,X9,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X9))) ) | $spl436),
  inference(cnf_transformation,[],[f21604_D])).
fof(f21604_D,plain,(
  ( ! [X8,X9] : (sP2(X8,X9,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X9))) ) <=> ~$spl436),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl436])])).
fof(f14553,plain,(
  ( ! [X8,X9] : (sP2(X8,X9,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X9)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl80),
  inference(resolution,[],[f2633,f713])).
fof(f713,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~sP1(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f509])).
fof(f509,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP1(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP1(X0,X1,X2) | sP2(X0,X1,X2)))),
  inference(rectify,[],[f508])).
fof(f508,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP1(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP1(X2,X0,X1) | sP2(X2,X0,X1)))),
  inference(flattening,[],[f507])).
fof(f507,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP1(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP1(X2,X0,X1)) | sP2(X2,X0,X1)))),
  inference(nnf_transformation,[],[f461])).
fof(f461,plain,(
  ! [X2,X0,X1] : (sP2(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP1(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f14565,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X0)) | ~p(s(bool,X1))) ) | $spl82),
  inference(resolution,[],[f2647,f714])).
fof(f21602,plain,(
  ( ! [X6,X7] : (sP2('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X6,X7)) ) | $spl434),
  inference(cnf_transformation,[],[f21602_D])).
fof(f21602_D,plain,(
  ( ! [X6,X7] : (sP2('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X6,X7)) ) <=> ~$spl434),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl434])])).
fof(f14566,plain,(
  ( ! [X2,X3] : (sP2('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X2,X3) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl82),
  inference(resolution,[],[f2647,f958])).
fof(f958,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~sP1(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(duplicate_literal_removal,[],[f710])).
fof(f710,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~sP1(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f509])).
fof(f21603,plain,(
  ( ! [X4,X5] : (sP2('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X4,X5) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl82),
  inference(subsumption_resolution,[],[f14567,f719])).
fof(f719,plain,(
  ( ! [X2,X0,X1] : (sP2(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f514])).
fof(f514,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP2(X2,X0,X1)) & (~sP2(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f513])).
fof(f513,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP2(X2,X0,X1)) & (~sP2(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f462])).
fof(f462,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP2(X2,X0,X1))),
  inference(definition_folding,[],[f425,f461,f460])).
fof(f425,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f311])).
fof(f311,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f310])).
fof(f310,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f55])).
fof(f55,axiom,(
  ! [X23,X21,X22] : ((p(s(bool,X22)) <=> (p(s(bool,X21)) => p(s(bool,X23)))) <=> ((p(s(bool,X22)) | p(s(bool,X21))) & (p(s(bool,X22)) | ~p(s(bool,X23))) & (~p(s(bool,X21)) | p(s(bool,X23)) | ~p(s(bool,X22)))))),
  file('thm.util_prob.INJ_IMAGE_BIJ','thm.sat.sat_27_prove_ltb')).
fof(f14567,plain,(
  ( ! [X4,X5] : (sP2('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X4,X5) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X5))) ) | $spl82),
  inference(resolution,[],[f2647,f711])).
fof(f711,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~sP1(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f509])).
fof(f14580,plain,(
  ( ! [X6,X7] : (sP2('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X6,X7) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl82),
  inference(subsumption_resolution,[],[f14568,f720])).
fof(f720,plain,(
  ( ! [X2,X0,X1] : (sP2(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f514])).
fof(f14568,plain,(
  ( ! [X6,X7] : (sP2('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X6,X7) | p(s(bool,X6)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl82),
  inference(resolution,[],[f2647,f712])).
fof(f712,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~sP1(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f509])).
fof(f14581,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X0)) | p(s(bool,X1))) ) | $spl84),
  inference(resolution,[],[f2662,f731])).
fof(f21601,plain,(
  ( ! [X6,X7] : (sP4('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X6,X7)) ) | $spl432),
  inference(cnf_transformation,[],[f21601_D])).
fof(f21601_D,plain,(
  ( ! [X6,X7] : (sP4('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X6,X7)) ) <=> ~$spl432),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl432])])).
fof(f14582,plain,(
  ( ! [X2,X3] : (sP4('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X2,X3) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl84),
  inference(resolution,[],[f2662,f959])).
fof(f959,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~sP3(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(duplicate_literal_removal,[],[f727])).
fof(f727,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~sP3(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f517])).
fof(f517,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP3(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP3(X0,X1,X2) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f516])).
fof(f516,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP3(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP3(X2,X0,X1) | sP4(X2,X0,X1)))),
  inference(flattening,[],[f515])).
fof(f515,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP3(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP3(X2,X0,X1)) | sP4(X2,X0,X1)))),
  inference(nnf_transformation,[],[f464])).
fof(f464,plain,(
  ! [X2,X0,X1] : (sP4(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP3(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f14596,plain,(
  ( ! [X4,X5] : (sP4('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X4,X5) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl84),
  inference(subsumption_resolution,[],[f14583,f736])).
fof(f736,plain,(
  ( ! [X2,X0,X1] : (sP4(X2,X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f522])).
fof(f522,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP4(X2,X0,X1)) & (~sP4(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f521])).
fof(f521,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP4(X2,X0,X1)) & (~sP4(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f465])).
fof(f465,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP4(X2,X0,X1))),
  inference(definition_folding,[],[f313,f464,f463])).
fof(f313,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f312])).
fof(f312,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f51])).
fof(f51,axiom,(
  ! [X23,X21,X22] : ((p(s(bool,X22)) <=> (p(s(bool,X21)) | p(s(bool,X23)))) <=> ((p(s(bool,X22)) | ~p(s(bool,X21))) & (p(s(bool,X22)) | ~p(s(bool,X23))) & (p(s(bool,X21)) | p(s(bool,X23)) | ~p(s(bool,X22)))))),
  file('thm.util_prob.INJ_IMAGE_BIJ','thm.sat.sat_25_prove_ltb')).
fof(f14583,plain,(
  ( ! [X4,X5] : (sP4('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X4,X5) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X5))) ) | $spl84),
  inference(resolution,[],[f2662,f728])).
fof(f728,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~sP3(X0,X1,X2) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f517])).
fof(f14597,plain,(
  ( ! [X6,X7] : (sP4('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X6,X7) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl84),
  inference(subsumption_resolution,[],[f14584,f737])).
fof(f737,plain,(
  ( ! [X2,X0,X1] : (sP4(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f522])).
fof(f14584,plain,(
  ( ! [X6,X7] : (sP4('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X6,X7) | p(s(bool,X6)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl84),
  inference(resolution,[],[f2662,f729])).
fof(f729,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~sP3(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f517])).
fof(f14598,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X1))) ) | $spl86),
  inference(resolution,[],[f3036,f748])).
fof(f21600,plain,(
  ( ! [X2,X3] : (sP6(X2,X3,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X3))) ) | $spl430),
  inference(cnf_transformation,[],[f21600_D])).
fof(f21600_D,plain,(
  ( ! [X2,X3] : (sP6(X2,X3,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X3))) ) <=> ~$spl430),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl430])])).
fof(f14599,plain,(
  ( ! [X2,X3] : (sP6(X2,X3,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X3)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl86),
  inference(resolution,[],[f3036,f744])).
fof(f744,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP5(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f525])).
fof(f525,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (sP5(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP5(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(rectify,[],[f524])).
fof(f524,plain,(
  ! [X2,X0,X1] : ((~sP6(X2,X0,X1) | (sP5(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP5(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP6(X2,X0,X1)))),
  inference(flattening,[],[f523])).
fof(f523,plain,(
  ! [X2,X0,X1] : ((~sP6(X2,X0,X1) | (sP5(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP5(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP6(X2,X0,X1)))),
  inference(nnf_transformation,[],[f467])).
fof(f467,plain,(
  ! [X2,X0,X1] : (sP6(X2,X0,X1) <=> (sP5(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f14614,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X0)) | p(s(bool,X1))) ) | $spl88),
  inference(resolution,[],[f3050,f748])).
fof(f21598,plain,(
  ( ! [X2,X3] : (sP6(X2,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X3) | ~p(s(bool,X3))) ) | $spl428),
  inference(cnf_transformation,[],[f21598_D])).
fof(f21598_D,plain,(
  ( ! [X2,X3] : (sP6(X2,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X3) | ~p(s(bool,X3))) ) <=> ~$spl428),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl428])])).
fof(f14615,plain,(
  ( ! [X2,X3] : (sP6(X2,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X3) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X3))) ) | $spl88),
  inference(resolution,[],[f3050,f744])).
fof(f14812,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X1))) ) | $spl90),
  inference(resolution,[],[f3064,f765])).
fof(f21594,plain,(
  ( ! [X2,X3] : (sP10(X2,X3,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | $spl424),
  inference(cnf_transformation,[],[f21594_D])).
fof(f21594_D,plain,(
  ( ! [X2,X3] : (sP10(X2,X3,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) <=> ~$spl424),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl424])])).
fof(f21596,plain,(
  ( ! [X2,X3] : (sP10(X2,X3,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | ($spl90 | $spl94)),
  inference(subsumption_resolution,[],[f14826,f3450])).
fof(f14826,plain,(
  ( ! [X2,X3] : (sP10(X2,X3,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~sP8(X3,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X2)) ) | $spl90),
  inference(subsumption_resolution,[],[f14813,f776])).
fof(f14813,plain,(
  ( ! [X2,X3] : (sP10(X2,X3,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~sP8(X3,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X2) | ~sP7('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X3,X2)) ) | $spl90),
  inference(resolution,[],[f3064,f762])).
fof(f762,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~p(s(bool,X2)) | ~sP9(X0,X2,X1) | ~sP8(X1,X2,X0) | ~sP7(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f532])).
fof(f532,plain,(
  ! [X0,X1,X2] : ((~sP10(X0,X1,X2) | (sP7(X2,X1,X0) & sP8(X1,X2,X0) & sP9(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP7(X2,X1,X0) | ~sP8(X1,X2,X0) | ~sP9(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP10(X0,X1,X2)))),
  inference(rectify,[],[f531])).
fof(f531,plain,(
  ! [X2,X1,X0] : ((~sP10(X2,X1,X0) | (sP7(X0,X1,X2) & sP8(X1,X0,X2) & sP9(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP7(X0,X1,X2) | ~sP8(X1,X0,X2) | ~sP9(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP10(X2,X1,X0)))),
  inference(flattening,[],[f530])).
fof(f530,plain,(
  ! [X2,X1,X0] : ((~sP10(X2,X1,X0) | (sP7(X0,X1,X2) & sP8(X1,X0,X2) & sP9(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP7(X0,X1,X2) | ~sP8(X1,X0,X2) | ~sP9(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP10(X2,X1,X0)))),
  inference(nnf_transformation,[],[f472])).
fof(f472,plain,(
  ! [X2,X1,X0] : (sP10(X2,X1,X0) <=> (sP7(X0,X1,X2) & sP8(X1,X0,X2) & sP9(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f14828,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X0)) | p(s(bool,X1))) ) | $spl92),
  inference(resolution,[],[f3079,f765])).
fof(f1823,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | $spl68),
  inference(cnf_transformation,[],[f1823_D])).
fof(f1823_D,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f21595,plain,(
  ( ! [X6,X7] : (sP10('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X6,X7) | ~sP7(X7,X6,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | ~sP8(X6,X7,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | $spl426),
  inference(cnf_transformation,[],[f21595_D])).
fof(f21595_D,plain,(
  ( ! [X6,X7] : (sP10('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X6,X7) | ~sP7(X7,X6,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | ~sP8(X6,X7,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) <=> ~$spl426),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl426])])).
fof(f14831,plain,(
  ( ! [X6,X7] : (sP10('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X6,X7) | p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~sP8(X6,X7,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | ~sP7(X7,X6,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | $spl92),
  inference(resolution,[],[f3079,f764])).
fof(f764,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | p(s(bool,X0)) | ~sP9(X0,X2,X1) | ~sP8(X1,X2,X0) | ~sP7(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f532])).
fof(f14845,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X1))) ) | $spl94),
  inference(resolution,[],[f3450,f769])).
fof(f21593,plain,(
  ( ! [X2,X3] : (sP10(X2,X3,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | ($spl90 | $spl94)),
  inference(subsumption_resolution,[],[f14859,f3064])).
fof(f14859,plain,(
  ( ! [X2,X3] : (sP10(X2,X3,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~sP9(X2,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X3)) ) | $spl94),
  inference(subsumption_resolution,[],[f14846,f776])).
fof(f14846,plain,(
  ( ! [X2,X3] : (sP10(X2,X3,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~sP9(X2,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X3) | ~sP7('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X3,X2)) ) | $spl94),
  inference(resolution,[],[f3450,f762])).
fof(f14862,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X0)) | p(s(bool,X1))) ) | $spl96),
  inference(resolution,[],[f3465,f769])).
fof(f21592,plain,(
  ( ! [X4,X5] : (sP10(X4,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X5) | ~sP7(X5,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X4) | ~sP9(X4,X5,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | $spl422),
  inference(cnf_transformation,[],[f21592_D])).
fof(f21592_D,plain,(
  ( ! [X4,X5] : (sP10(X4,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X5) | ~sP7(X5,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X4) | ~sP9(X4,X5,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) <=> ~$spl422),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl422])])).
fof(f14864,plain,(
  ( ! [X4,X5] : (sP10(X4,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X5) | p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~sP9(X4,X5,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | ~sP7(X5,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X4)) ) | $spl96),
  inference(resolution,[],[f3465,f763])).
fof(f763,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | p(s(bool,X1)) | ~sP9(X0,X2,X1) | ~sP8(X1,X2,X0) | ~sP7(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f532])).
fof(f14961,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl168),
  inference(duplicate_literal_removal,[],[f14943])).
fof(f14943,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X0))) ) | $spl168),
  inference(resolution,[],[f5218,f845])).
fof(f845,plain,(
  ( ! [X0,X1] : (sP14(X1,X0) | ~p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f573])).
fof(f573,plain,(
  ! [X0,X1] : ((((p(s(bool,X1)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)))) | sP14(X1,X0)) & (~sP14(X1,X0) | ((~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X0)) | p(s(bool,X1))))))),
  inference(nnf_transformation,[],[f481])).
fof(f481,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> sP14(X1,X0))),
  inference(definition_folding,[],[f331,f480])).
fof(f480,plain,(
  ! [X1,X0] : (sP14(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP14])])).
fof(f331,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  inference(flattening,[],[f330])).
fof(f330,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  inference(rectify,[],[f48])).
fof(f48,axiom,(
  ! [X21,X22] : ((p(s(bool,X22)) <=> ~p(s(bool,X21))) <=> ((p(s(bool,X22)) | p(s(bool,X21))) & (~p(s(bool,X21)) | ~p(s(bool,X22)))))),
  file('thm.util_prob.INJ_IMAGE_BIJ','thm.sat.sat_29_prove_ltb')).
fof(f14960,plain,(
  ( ! [X3] : (p(s(bool,X3)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl168),
  inference(duplicate_literal_removal,[],[f14946])).
fof(f14946,plain,(
  ( ! [X3] : (p(s(bool,X3)) | p(s(bool,X3)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl168),
  inference(resolution,[],[f5218,f844])).
fof(f844,plain,(
  ( ! [X0,X1] : (sP14(X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f572])).
fof(f572,plain,(
  ! [X0,X1] : ((~sP14(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP14(X0,X1)))),
  inference(rectify,[],[f571])).
fof(f571,plain,(
  ! [X1,X0] : ((~sP14(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP14(X1,X0)))),
  inference(flattening,[],[f570])).
fof(f570,plain,(
  ! [X1,X0] : ((~sP14(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP14(X1,X0)))),
  inference(nnf_transformation,[],[f480])).
fof(f14980,plain,(
  ( ! [X1] : (p(s(bool,X1)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl170),
  inference(duplicate_literal_removal,[],[f14963])).
fof(f14963,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,X1)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl170),
  inference(resolution,[],[f5232,f846])).
fof(f846,plain,(
  ( ! [X0,X1] : (sP14(X1,X0) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f573])).
fof(f14979,plain,(
  ( ! [X2] : (p(s(bool,X2)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl170),
  inference(duplicate_literal_removal,[],[f14964])).
fof(f14964,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,X2)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl170),
  inference(resolution,[],[f5232,f841])).
fof(f841,plain,(
  ( ! [X0,X1] : (sP14(X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f572])).
fof(f15001,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl172),
  inference(duplicate_literal_removal,[],[f14982])).
fof(f14982,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X1))) ) | $spl172),
  inference(resolution,[],[f5248,f847])).
fof(f847,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | ~sP14(X1,X0)) )),
  inference(cnf_transformation,[],[f573])).
fof(f15000,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl172),
  inference(duplicate_literal_removal,[],[f14985])).
fof(f14985,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | ~p(s(bool,X4)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl172),
  inference(resolution,[],[f5248,f840])).
fof(f840,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | ~sP14(X0,X1)) )),
  inference(cnf_transformation,[],[f572])).
fof(f15022,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl178),
  inference(duplicate_literal_removal,[],[f15003])).
fof(f15003,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | ~p(s(bool,X1)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl178),
  inference(resolution,[],[f5288,f847])).
fof(f15021,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl178),
  inference(duplicate_literal_removal,[],[f15006])).
fof(f15006,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X4))) ) | $spl178),
  inference(resolution,[],[f5288,f840])).
fof(f21587,plain,(
  ( ! [X3] : (p(s(bool,X3)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | ($spl82 | $spl198)),
  inference(subsumption_resolution,[],[f21586,f2647])).
fof(f21586,plain,(
  ( ! [X2,X3] : (~sP1('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X3,X2) | p(s(bool,X3)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl198),
  inference(subsumption_resolution,[],[f16531,f714])).
fof(f16531,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | ~sP1('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X3,X2) | p(s(bool,X3)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl198),
  inference(resolution,[],[f5711,f712])).
fof(f16550,plain,(
  ( ! [X8] : (p(s(bool,X8)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl198),
  inference(duplicate_literal_removal,[],[f16534])).
fof(f16534,plain,(
  ( ! [X8] : (p(s(bool,X8)) | p(s(bool,X8)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl198),
  inference(resolution,[],[f5711,f719])).
fof(f16535,plain,(
  ( ! [X10,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl198),
  inference(resolution,[],[f5711,f720])).
fof(f21585,plain,(
  ( ! [X11] : (p(s(bool,X11)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl198),
  inference(subsumption_resolution,[],[f16536,f715])).
fof(f16536,plain,(
  ( ! [X12,X11] : (p(s(bool,X11)) | ~sP1('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X12,X11) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl198),
  inference(resolution,[],[f5711,f958])).
fof(f21583,plain,(
  ( ! [X5] : (p(s(bool,X5)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | ($spl80 | $spl202)),
  inference(subsumption_resolution,[],[f21582,f2633])).
fof(f21582,plain,(
  ( ! [X4,X5] : (~sP1(X4,X5,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X5)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl202),
  inference(subsumption_resolution,[],[f16553,f714])).
fof(f16553,plain,(
  ( ! [X4,X5] : (p(s(bool,X4)) | ~sP1(X4,X5,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X5)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl202),
  inference(resolution,[],[f5907,f713])).
fof(f16568,plain,(
  ( ! [X6,X7] : (p(s(bool,X6)) | p(s(bool,X7)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl202),
  inference(duplicate_literal_removal,[],[f16554])).
fof(f16554,plain,(
  ( ! [X6,X7] : (p(s(bool,X6)) | p(s(bool,X7)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X6))) ) | $spl202),
  inference(resolution,[],[f5907,f718])).
fof(f718,plain,(
  ( ! [X2,X0,X1] : (sP2(X2,X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f514])).
fof(f16590,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X1))) ) | $spl204),
  inference(subsumption_resolution,[],[f16569,f716])).
fof(f16569,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~sP1('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0,X1) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X1))) ) | $spl204),
  inference(resolution,[],[f5925,f711])).
fof(f16573,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | p(s(bool,X9)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl204),
  inference(resolution,[],[f5925,f719])).
fof(f16589,plain,(
  ( ! [X10] : (~p(s(bool,X10)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl204),
  inference(duplicate_literal_removal,[],[f16574])).
fof(f16574,plain,(
  ( ! [X10] : (~p(s(bool,X10)) | ~p(s(bool,X10)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl204),
  inference(resolution,[],[f5925,f720])).
fof(f16591,plain,(
  ( ! [X11] : (~p(s(bool,X11)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl204),
  inference(subsumption_resolution,[],[f16575,f716])).
fof(f16575,plain,(
  ( ! [X12,X11] : (~p(s(bool,X11)) | ~sP1('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X11,X12) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl204),
  inference(resolution,[],[f5925,f958])).
fof(f16577,plain,(
  ( ! [X14] : (~p(s(bool,X14)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | ($spl77 | $spl204)),
  inference(resolution,[],[f5925,f4549])).
fof(f4549,plain,(
  ( ! [X47,X48] : (sP2(X47,X48,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X47))) ) | $spl77),
  inference(resolution,[],[f2297,f719])).
fof(f16617,plain,(
  ( ! [X2,X1] : (~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl206),
  inference(duplicate_literal_removal,[],[f16593])).
fof(f16593,plain,(
  ( ! [X2,X1] : (~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X1))) ) | $spl206),
  inference(resolution,[],[f5943,f721])).
fof(f721,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP2(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f514])).
fof(f16605,plain,(
  ( ! [X14] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X14))) ) | ($spl204 | $spl206)),
  inference(resolution,[],[f5943,f5925])).
fof(f16636,plain,(
  ( ! [X2,X3] : (~p(s(bool,X2)) | p(s(bool,X3)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl208),
  inference(subsumption_resolution,[],[f16619,f732])).
fof(f16619,plain,(
  ( ! [X2,X3] : (~p(s(bool,X2)) | ~sP3('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X3,X2) | p(s(bool,X3)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl208),
  inference(resolution,[],[f6119,f729])).
fof(f16635,plain,(
  ( ! [X8] : (~p(s(bool,X8)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl208),
  inference(duplicate_literal_removal,[],[f16622])).
fof(f16622,plain,(
  ( ! [X8] : (~p(s(bool,X8)) | ~p(s(bool,X8)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl208),
  inference(resolution,[],[f6119,f736])).
fof(f16623,plain,(
  ( ! [X10,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl208),
  inference(resolution,[],[f6119,f737])).
fof(f16637,plain,(
  ( ! [X11] : (~p(s(bool,X11)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl208),
  inference(subsumption_resolution,[],[f16624,f732])).
fof(f16624,plain,(
  ( ! [X12,X11] : (~p(s(bool,X11)) | ~sP3('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X12,X11) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl208),
  inference(resolution,[],[f6119,f959])).
fof(f16656,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X1))) ) | $spl210),
  inference(subsumption_resolution,[],[f16638,f733])).
fof(f16638,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~sP3('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0,X1) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X1))) ) | $spl210),
  inference(resolution,[],[f6137,f728])).
fof(f16642,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | ~p(s(bool,X9)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl210),
  inference(resolution,[],[f6137,f736])).
fof(f16655,plain,(
  ( ! [X10] : (~p(s(bool,X10)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl210),
  inference(duplicate_literal_removal,[],[f16643])).
fof(f16643,plain,(
  ( ! [X10] : (~p(s(bool,X10)) | ~p(s(bool,X10)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl210),
  inference(resolution,[],[f6137,f737])).
fof(f16657,plain,(
  ( ! [X11] : (~p(s(bool,X11)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl210),
  inference(subsumption_resolution,[],[f16644,f733])).
fof(f16644,plain,(
  ( ! [X12,X11] : (~p(s(bool,X11)) | ~sP3('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X11,X12) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl210),
  inference(resolution,[],[f6137,f959])).
fof(f16685,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl212),
  inference(subsumption_resolution,[],[f16658,f749])).
fof(f16658,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~sP5(X1,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0)) ) | $spl212),
  inference(resolution,[],[f6341,f744])).
fof(f16684,plain,(
  ( ! [X10,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl212),
  inference(duplicate_literal_removal,[],[f16663])).
fof(f16663,plain,(
  ( ! [X10,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X9))) ) | $spl212),
  inference(resolution,[],[f6341,f754])).
fof(f754,plain,(
  ( ! [X2,X0,X1] : (sP6(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f529])).
fof(f529,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP6(X2,X0,X1)) & (~sP6(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f528])).
fof(f528,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP6(X2,X0,X1)) & (~sP6(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f468])).
fof(f468,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP6(X2,X0,X1))),
  inference(definition_folding,[],[f315,f467,f466])).
fof(f315,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f314])).
fof(f314,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f60])).
fof(f60,axiom,(
  ! [X23,X21,X22] : ((p(s(bool,X22)) <=> (p(s(bool,X21)) & p(s(bool,X23)))) <=> ((p(s(bool,X22)) | ~p(s(bool,X21)) | ~p(s(bool,X23))) & (p(s(bool,X21)) | ~p(s(bool,X22))) & (p(s(bool,X23)) | ~p(s(bool,X22)))))),
  file('thm.util_prob.INJ_IMAGE_BIJ','thm.sat.sat_23_prove_ltb')).
fof(f16713,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X1))) ) | $spl214),
  inference(subsumption_resolution,[],[f16686,f749])).
fof(f16686,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X1)) | ~sP5('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X1,X0)) ) | $spl214),
  inference(resolution,[],[f6358,f744])).
fof(f16712,plain,(
  ( ! [X10,X9] : (~p(s(bool,X9)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X10))) ) | $spl214),
  inference(duplicate_literal_removal,[],[f16691])).
fof(f16691,plain,(
  ( ! [X10,X9] : (~p(s(bool,X9)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X10)) | ~p(s(bool,X9))) ) | $spl214),
  inference(resolution,[],[f6358,f754])).
fof(f16735,plain,(
  ( ! [X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl216),
  inference(duplicate_literal_removal,[],[f16717])).
fof(f16717,plain,(
  ( ! [X4,X5] : (p(s(bool,X4)) | p(s(bool,X4)) | ~p(s(bool,X5)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl216),
  inference(resolution,[],[f6375,f757])).
fof(f757,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | ~sP6(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f529])).
fof(f16761,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl218),
  inference(duplicate_literal_removal,[],[f16736])).
fof(f16736,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,X0)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl218),
  inference(resolution,[],[f6393,f755])).
fof(f755,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP6(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f529])).
fof(f16737,plain,(
  ( ! [X2,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl218),
  inference(resolution,[],[f6393,f756])).
fof(f756,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X2)) | ~sP6(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f529])).
fof(f16760,plain,(
  ( ! [X7] : (p(s(bool,X7)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl218),
  inference(duplicate_literal_removal,[],[f16740])).
fof(f16740,plain,(
  ( ! [X7] : (p(s(bool,X7)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X7))) ) | $spl218),
  inference(resolution,[],[f6393,f742])).
fof(f742,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X2)) | ~sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f525])).
fof(f16741,plain,(
  ( ! [X8,X9] : (p(s(bool,X8)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X9))) ) | $spl218),
  inference(resolution,[],[f6393,f743])).
fof(f743,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f525])).
fof(f16747,plain,(
  ( ! [X11] : (p(s(bool,X11)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | ($spl77 | $spl218)),
  inference(resolution,[],[f6393,f4571])).
fof(f4571,plain,(
  ( ! [X92,X91] : (~sP6(X91,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X92) | ~p(s(bool,X91))) ) | $spl77),
  inference(resolution,[],[f2297,f743])).
fof(f16749,plain,(
  ( ! [X13] : (p(s(bool,X13)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | ($spl214 | $spl218)),
  inference(resolution,[],[f6393,f6358])).
fof(f16783,plain,(
  ( ! [X4,X5] : (p(s(bool,X4)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X5))) ) | $spl220),
  inference(duplicate_literal_removal,[],[f16765])).
fof(f16765,plain,(
  ( ! [X4,X5] : (p(s(bool,X4)) | p(s(bool,X4)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X5))) ) | $spl220),
  inference(resolution,[],[f6410,f757])).
fof(f16837,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl232),
  inference(resolution,[],[f6604,f755])).
fof(f16862,plain,(
  ( ! [X2] : (p(s(bool,X2)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl232),
  inference(duplicate_literal_removal,[],[f16838])).
fof(f16838,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,X2)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl232),
  inference(resolution,[],[f6604,f756])).
fof(f16841,plain,(
  ( ! [X8,X7] : (p(s(bool,X7)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X8))) ) | $spl232),
  inference(resolution,[],[f6604,f742])).
fof(f16861,plain,(
  ( ! [X9] : (p(s(bool,X9)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl232),
  inference(duplicate_literal_removal,[],[f16842])).
fof(f16842,plain,(
  ( ! [X9] : (p(s(bool,X9)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X9))) ) | $spl232),
  inference(resolution,[],[f6604,f743])).
fof(f16844,plain,(
  ( ! [X11] : (p(s(bool,X11)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | ($spl77 | $spl232)),
  inference(resolution,[],[f6604,f4570])).
fof(f4570,plain,(
  ( ! [X90,X89] : (~sP6(X89,X90,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X89))) ) | $spl77),
  inference(resolution,[],[f2297,f742])).
fof(f16846,plain,(
  ( ! [X13] : (p(s(bool,X13)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | ($spl212 | $spl232)),
  inference(resolution,[],[f6604,f6341])).
fof(f16864,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl234),
  inference(resolution,[],[f6618,f788])).
fof(f16866,plain,(
  ( ! [X6,X5] : (~p(s(bool,X5)) | p(s(bool,X6)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl234),
  inference(resolution,[],[f6618,f784])).
fof(f784,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f544])).
fof(f16868,plain,(
  ( ! [X8] : (~p(s(bool,X8)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl234),
  inference(resolution,[],[f6618,f961])).
fof(f961,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(duplicate_literal_removal,[],[f786])).
fof(f786,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f544])).
fof(f16886,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl236),
  inference(resolution,[],[f6811,f788])).
fof(f16889,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X7))) ) | $spl236),
  inference(resolution,[],[f6811,f785])).
fof(f785,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f544])).
fof(f16890,plain,(
  ( ! [X8] : (~p(s(bool,X8)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl236),
  inference(resolution,[],[f6811,f961])).
fof(f16908,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | p(s(bool,X3)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl238),
  inference(resolution,[],[f6825,f790])).
fof(f790,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X1)) | ~sP11(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f546])).
fof(f16910,plain,(
  ( ! [X6,X7] : (p(s(bool,X6)) | p(s(bool,X7)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl238),
  inference(resolution,[],[f6825,f782])).
fof(f16913,plain,(
  ( ! [X10] : (p(s(bool,X10)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | ($spl77 | $spl238)),
  inference(resolution,[],[f6825,f4599])).
fof(f16917,plain,(
  ( ! [X14] : (p(s(bool,X14)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | ($spl236 | $spl238)),
  inference(resolution,[],[f6825,f6811])).
fof(f16929,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl240),
  inference(resolution,[],[f6839,f789])).
fof(f789,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~sP11(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f546])).
fof(f16931,plain,(
  ( ! [X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl240),
  inference(resolution,[],[f6839,f781])).
fof(f16936,plain,(
  ( ! [X11] : (p(s(bool,X11)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | ($spl77 | $spl240)),
  inference(resolution,[],[f6839,f4598])).
fof(f16940,plain,(
  ( ! [X15] : (p(s(bool,X15)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | ($spl234 | $spl240)),
  inference(resolution,[],[f6839,f6618])).
fof(f16952,plain,(
  ( ! [X2] : (p(s(bool,X2)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl244),
  inference(resolution,[],[f7049,f804])).
fof(f16954,plain,(
  ( ! [X6,X5] : (p(s(bool,X5)) | ~p(s(bool,X6)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl244),
  inference(resolution,[],[f7049,f800])).
fof(f800,plain,(
  ( ! [X2,X0,X1] : (sP12(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f551])).
fof(f16956,plain,(
  ( ! [X8] : (p(s(bool,X8)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl244),
  inference(resolution,[],[f7049,f962])).
fof(f962,plain,(
  ( ! [X2,X0,X1] : (sP12(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(duplicate_literal_removal,[],[f802])).
fof(f802,plain,(
  ( ! [X2,X0,X1] : (sP12(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f551])).
fof(f16967,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl246),
  inference(resolution,[],[f7064,f803])).
fof(f803,plain,(
  ( ! [X2,X0,X1] : (sP12(X0,X2,X1) | ~p(s(bool,X2)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f553])).
fof(f16969,plain,(
  ( ! [X4,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl246),
  inference(resolution,[],[f7064,f799])).
fof(f799,plain,(
  ( ! [X2,X0,X1] : (sP12(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f551])).
fof(f16984,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X1))) ) | $spl248),
  inference(resolution,[],[f7237,f803])).
fof(f16986,plain,(
  ( ! [X4,X3] : (p(s(bool,X3)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X4))) ) | $spl248),
  inference(resolution,[],[f7237,f799])).
fof(f17002,plain,(
  ( ! [X2] : (p(s(bool,X2)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl250),
  inference(resolution,[],[f7251,f804])).
fof(f17005,plain,(
  ( ! [X6,X7] : (p(s(bool,X6)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X7))) ) | $spl250),
  inference(resolution,[],[f7251,f801])).
fof(f801,plain,(
  ( ! [X2,X0,X1] : (sP12(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f551])).
fof(f17006,plain,(
  ( ! [X8] : (p(s(bool,X8)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl250),
  inference(resolution,[],[f7251,f962])).
fof(f17017,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl252),
  inference(resolution,[],[f7427,f813])).
fof(f17019,plain,(
  ( ! [X3] : (p(s(bool,X3)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl252),
  inference(resolution,[],[f7427,f963])).
fof(f963,plain,(
  ( ! [X2,X0,X1] : (sP13(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(duplicate_literal_removal,[],[f809])).
fof(f809,plain,(
  ( ! [X2,X0,X1] : (sP13(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f556])).
fof(f17021,plain,(
  ( ! [X6,X5] : (p(s(bool,X5)) | ~p(s(bool,X6)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl252),
  inference(resolution,[],[f7427,f811])).
fof(f811,plain,(
  ( ! [X2,X0,X1] : (sP13(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f556])).
fof(f17034,plain,(
  ( ! [X2,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl254),
  inference(resolution,[],[f7442,f814])).
fof(f814,plain,(
  ( ! [X2,X0,X1] : (sP13(X2,X0,X1) | ~p(s(bool,X2)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f558])).
fof(f17038,plain,(
  ( ! [X8,X7] : (p(s(bool,X7)) | ~p(s(bool,X8)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl254),
  inference(resolution,[],[f7442,f812])).
fof(f812,plain,(
  ( ! [X2,X0,X1] : (sP13(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f556])).
fof(f17050,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl256),
  inference(resolution,[],[f7456,f813])).
fof(f17052,plain,(
  ( ! [X3] : (p(s(bool,X3)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl256),
  inference(resolution,[],[f7456,f963])).
fof(f17053,plain,(
  ( ! [X4,X5] : (p(s(bool,X4)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X5))) ) | $spl256),
  inference(resolution,[],[f7456,f810])).
fof(f810,plain,(
  ( ! [X2,X0,X1] : (sP13(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f556])).
fof(f17067,plain,(
  ( ! [X2,X1] : (p(s(bool,X1)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X2))) ) | $spl258),
  inference(resolution,[],[f7471,f814])).
fof(f17071,plain,(
  ( ! [X8,X7] : (p(s(bool,X7)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X8))) ) | $spl258),
  inference(resolution,[],[f7471,f812])).
fof(f17162,plain,(
  ( ! [X18] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X18))) ) | $spl198),
  inference(resolution,[],[f996,f5711])).
fof(f996,plain,(
  ( ! [X57,X56] : (sP2(X56,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)),X57) | ~p(s(bool,X56))) )),
  inference(resolution,[],[f647,f720])).
fof(f17681,plain,(
  ( ! [X15] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X15))) ) | $spl210),
  inference(resolution,[],[f1008,f6137])).
fof(f1008,plain,(
  ( ! [X80,X81] : (sP4(X80,X81,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20))) | ~p(s(bool,X80))) )),
  inference(resolution,[],[f647,f736])).
fof(f17974,plain,(
  ( ! [X15] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X15))) ) | $spl208),
  inference(resolution,[],[f1010,f6119])).
fof(f1010,plain,(
  ( ! [X85,X84] : (sP4(X84,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)),X85) | ~p(s(bool,X84))) )),
  inference(resolution,[],[f647,f737])).
fof(f16915,plain,(
  ( ! [X12] : (p(s(bool,X12)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | ($spl5 | $spl238)),
  inference(resolution,[],[f6825,f13271])).
fof(f13271,plain,(
  ( ! [X147,X148] : (~sP11('const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)),X147,X148) | ~p(s(bool,X147))) ) | $spl5),
  inference(resolution,[],[f1135,f782])).
fof(f16578,plain,(
  ( ! [X15] : (~p(s(bool,X15)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | ($spl5 | $spl204)),
  inference(resolution,[],[f5925,f13221])).
fof(f13221,plain,(
  ( ! [X47,X48] : (sP2(X47,X48,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | ~p(s(bool,X47))) ) | $spl5),
  inference(resolution,[],[f1135,f719])).
fof(f16845,plain,(
  ( ! [X12] : (p(s(bool,X12)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | ($spl5 | $spl232)),
  inference(resolution,[],[f6604,f13242])).
fof(f13242,plain,(
  ( ! [X90,X89] : (~sP6(X89,X90,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | ~p(s(bool,X89))) ) | $spl5),
  inference(resolution,[],[f1135,f742])).
fof(f16748,plain,(
  ( ! [X12] : (p(s(bool,X12)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | ($spl5 | $spl218)),
  inference(resolution,[],[f6393,f13243])).
fof(f13243,plain,(
  ( ! [X92,X91] : (~sP6(X91,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)),X92) | ~p(s(bool,X91))) ) | $spl5),
  inference(resolution,[],[f1135,f743])).
fof(f16938,plain,(
  ( ! [X13] : (p(s(bool,X13)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | ($spl5 | $spl240)),
  inference(resolution,[],[f6839,f13270])).
fof(f13270,plain,(
  ( ! [X146,X145] : (~sP11('const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)),X146,X145) | ~p(s(bool,X145))) ) | $spl5),
  inference(resolution,[],[f1135,f781])).
fof(f18710,plain,(
  ( ! [X14] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X14))) ) | ($spl5 | $spl204)),
  inference(resolution,[],[f18235,f5925])).
fof(f18235,plain,(
  ( ! [X47,X48] : (sP2(X47,X48,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | ~p(s(bool,X47))) ) | $spl5),
  inference(resolution,[],[f1135,f719])).
fof(f18798,plain,(
  ( ! [X18] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X18))) ) | ($spl5 | $spl232)),
  inference(resolution,[],[f18256,f6604])).
fof(f18256,plain,(
  ( ! [X90,X89] : (~sP6(X89,X90,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | ~p(s(bool,X89))) ) | $spl5),
  inference(resolution,[],[f1135,f742])).
fof(f18823,plain,(
  ( ! [X18] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X18))) ) | ($spl5 | $spl218)),
  inference(resolution,[],[f18257,f6393])).
fof(f18257,plain,(
  ( ! [X92,X91] : (~sP6(X91,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)),X92) | ~p(s(bool,X91))) ) | $spl5),
  inference(resolution,[],[f1135,f743])).
fof(f18928,plain,(
  ( ! [X11] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X11))) ) | ($spl5 | $spl240)),
  inference(resolution,[],[f18284,f6839])).
fof(f18284,plain,(
  ( ! [X146,X145] : (~sP11('const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)),X146,X145) | ~p(s(bool,X145))) ) | $spl5),
  inference(resolution,[],[f1135,f781])).
fof(f18950,plain,(
  ( ! [X14] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X14))) ) | ($spl5 | $spl238)),
  inference(resolution,[],[f18285,f6825])).
fof(f18285,plain,(
  ( ! [X147,X148] : (~sP11('const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)),X147,X148) | ~p(s(bool,X147))) ) | $spl5),
  inference(resolution,[],[f1135,f782])).
fof(f19058,plain,(
  ( ! [X10] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X10))) ) | $spl246),
  inference(resolution,[],[f1052,f7064])).
fof(f1052,plain,(
  ( ! [X169,X168] : (sP12(X168,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)),X169) | ~p(s(bool,X169))) )),
  inference(resolution,[],[f647,f799])).
fof(f19132,plain,(
  ( ! [X10] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X10))) ) | $spl248),
  inference(resolution,[],[f1053,f7237])).
fof(f1053,plain,(
  ( ! [X171,X170] : (sP12(X170,X171,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20))) | ~p(s(bool,X171))) )),
  inference(resolution,[],[f647,f799])).
fof(f16576,plain,(
  ( ! [X13] : (~p(s(bool,X13)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | ($spl9 | $spl204)),
  inference(resolution,[],[f5925,f2492])).
fof(f2492,plain,(
  ( ! [X47,X48] : (sP2(X47,X48,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) | ~p(s(bool,X47))) ) | $spl9),
  inference(resolution,[],[f1139,f719])).
fof(f16934,plain,(
  ( ! [X9] : (p(s(bool,X9)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | ($spl9 | $spl240)),
  inference(resolution,[],[f6839,f2541])).
fof(f2541,plain,(
  ( ! [X146,X145] : (~sP11('const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)),X146,X145) | ~p(s(bool,X145))) ) | $spl9),
  inference(resolution,[],[f1139,f781])).
fof(f16911,plain,(
  ( ! [X8] : (p(s(bool,X8)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | ($spl9 | $spl238)),
  inference(resolution,[],[f6825,f2542])).
fof(f2542,plain,(
  ( ! [X147,X148] : (~sP11('const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)),X147,X148) | ~p(s(bool,X147))) ) | $spl9),
  inference(resolution,[],[f1139,f782])).
fof(f16843,plain,(
  ( ! [X10] : (p(s(bool,X10)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | ($spl9 | $spl232)),
  inference(resolution,[],[f6604,f2513])).
fof(f2513,plain,(
  ( ! [X90,X89] : (~sP6(X89,X90,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) | ~p(s(bool,X89))) ) | $spl9),
  inference(resolution,[],[f1139,f742])).
fof(f16746,plain,(
  ( ! [X10] : (p(s(bool,X10)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | ($spl9 | $spl218)),
  inference(resolution,[],[f6393,f2514])).
fof(f2514,plain,(
  ( ! [X92,X91] : (~sP6(X91,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)),X92) | ~p(s(bool,X91))) ) | $spl9),
  inference(resolution,[],[f1139,f743])).
fof(f19976,plain,(
  ( ! [X9] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X9))) ) | $spl254),
  inference(resolution,[],[f1065,f7442])).
fof(f1065,plain,(
  ( ! [X194,X195] : (sP13('const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)),X194,X195) | ~p(s(bool,X195))) )),
  inference(resolution,[],[f647,f812])).
fof(f20108,plain,(
  ( ! [X15] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X15))) ) | $spl258),
  inference(resolution,[],[f1066,f7471])).
fof(f1066,plain,(
  ( ! [X196,X197] : (sP13(X196,X197,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20))) | ~p(s(bool,X196))) )),
  inference(resolution,[],[f647,f812])).
fof(f20215,plain,(
  ( ! [X14] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X14))) ) | ($spl9 | $spl204)),
  inference(resolution,[],[f19728,f5925])).
fof(f19728,plain,(
  ( ! [X47,X48] : (sP2(X47,X48,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) | ~p(s(bool,X47))) ) | $spl9),
  inference(resolution,[],[f1139,f719])).
fof(f20273,plain,(
  ( ! [X17] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X17))) ) | ($spl9 | $spl232)),
  inference(resolution,[],[f19749,f6604])).
fof(f19749,plain,(
  ( ! [X90,X89] : (~sP6(X89,X90,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) | ~p(s(bool,X89))) ) | $spl9),
  inference(resolution,[],[f1139,f742])).
fof(f20298,plain,(
  ( ! [X17] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X17))) ) | ($spl9 | $spl218)),
  inference(resolution,[],[f19750,f6393])).
fof(f19750,plain,(
  ( ! [X92,X91] : (~sP6(X91,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)),X92) | ~p(s(bool,X91))) ) | $spl9),
  inference(resolution,[],[f1139,f743])).
fof(f20404,plain,(
  ( ! [X10] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X10))) ) | ($spl9 | $spl240)),
  inference(resolution,[],[f19777,f6839])).
fof(f20426,plain,(
  ( ! [X13] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X13))) ) | ($spl9 | $spl238)),
  inference(resolution,[],[f19778,f6825])).
fof(f21564,plain,(
  $false | ($spl98 | $spl102)),
  inference(equality_resolution,[],[f21521])).
fof(f21521,plain,(
  ( ! [X22] : (s(bool,X22) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | ($spl98 | $spl102)),
  inference(subsumption_resolution,[],[f3775,f3603])).
fof(f19589,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) | $spl416),
  inference(cnf_transformation,[],[f19589_D])).
fof(f19589_D,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) <=> ~$spl416),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl416])])).
fof(f21193,plain,(
  ( ! [X352,X354,X353] : (s(bool,X352) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) | p(s(bool,X353)) | p(s(bool,X354)) | s(bool,X352) != s(bool,X354)) ) | $spl9),
  inference(resolution,[],[f19811,f957])).
fof(f21086,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) ) | $spl9),
  inference(resolution,[],[f19811,f706])).
fof(f21083,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,X16) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) ) | $spl9),
  inference(resolution,[],[f19811,f705])).
fof(f21082,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl9),
  inference(resolution,[],[f19811,f704])).
fof(f21081,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) ) | $spl9),
  inference(resolution,[],[f19811,f702])).
fof(f21079,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl9),
  inference(resolution,[],[f19811,f700])).
fof(f20931,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) | $spl420),
  inference(cnf_transformation,[],[f20931_D])).
fof(f20931_D,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) <=> ~$spl420),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl420])])).
fof(f20774,plain,(
  ( ! [X352,X354,X353] : (s(bool,X352) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | p(s(bool,X353)) | p(s(bool,X354)) | s(bool,X352) != s(bool,X354)) ) | $spl5),
  inference(resolution,[],[f18318,f957])).
fof(f20667,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) ) | $spl5),
  inference(resolution,[],[f18318,f706])).
fof(f20664,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,X16) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) ) | $spl5),
  inference(resolution,[],[f18318,f705])).
fof(f20663,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl5),
  inference(resolution,[],[f18318,f704])).
fof(f20662,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) ) | $spl5),
  inference(resolution,[],[f18318,f702])).
fof(f20660,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl5),
  inference(resolution,[],[f18318,f700])).
fof(f20041,plain,(
  $false | ($spl72 | $spl78)),
  inference(equality_resolution,[],[f19708])).
fof(f19708,plain,(
  ( ! [X22] : (s(bool,X22) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)))) ) | ($spl72 | $spl78)),
  inference(subsumption_resolution,[],[f2448,f2281])).
fof(f1139,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)))) | $spl9),
  inference(cnf_transformation,[],[f1139_D])).
fof(f1139_D,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f19706,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18))),s('type.num.num','const.pred_set.CARD_1'(s(fun(X3,bool),X4)))))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X3),X5),s(fun(sK16,bool),sK20),s(fun(X3,bool),X4))))) ) | $spl418),
  inference(cnf_transformation,[],[f19706_D])).
fof(f19706_D,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18))),s('type.num.num','const.pred_set.CARD_1'(s(fun(X3,bool),X4)))))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X3),X5),s(fun(sK16,bool),sK20),s(fun(X3,bool),X4))))) ) <=> ~$spl418),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl418])])).
fof(f19705,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18))),s('type.num.num','const.pred_set.CARD_1'(s(fun(X3,bool),X4)))))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X3),X5),s(fun(sK16,bool),sK20),s(fun(X3,bool),X4)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))))) ) | $spl10),
  inference(subsumption_resolution,[],[f16019,f676])).
fof(f16019,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18))),s('type.num.num','const.pred_set.CARD_1'(s(fun(X3,bool),X4)))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X3,bool),X4)))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X3),X5),s(fun(sK16,bool),sK20),s(fun(X3,bool),X4)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))))) ) | $spl10),
  inference(superposition,[],[f1141,f682])).
fof(f19704,plain,(
  ( ! [X17,X16] : (~p(s(bool,X17)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)))) | ~p(s(bool,X16))) ) | $spl72),
  inference(subsumption_resolution,[],[f19262,f864])).
fof(f19262,plain,(
  ( ! [X17,X16] : (s(bool,X16) != s(bool,X17) | ~p(s(bool,X17)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)))) | ~p(s(bool,X16))) ) | $spl72),
  inference(superposition,[],[f2281,f864])).
fof(f19703,plain,(
  ( ! [X19,X18] : (~p(s(bool,X19)) | ~p(s(bool,X18)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))))) ) | $spl72),
  inference(subsumption_resolution,[],[f19263,f864])).
fof(f19263,plain,(
  ( ! [X19,X18] : (s(bool,X18) != s(bool,X19) | ~p(s(bool,X19)) | ~p(s(bool,X18)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))))) ) | $spl72),
  inference(superposition,[],[f2281,f864])).
fof(f19432,plain,(
  ( ! [X352,X354,X353] : (s(bool,X352) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) | p(s(bool,X353)) | p(s(bool,X354)) | s(bool,X352) != s(bool,X354)) ) | $spl9),
  inference(resolution,[],[f2575,f957])).
fof(f2575,plain,(
  ( ! [X205] : (p(s(bool,X205)) | s(bool,X205) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)))) ) | $spl9),
  inference(resolution,[],[f1139,f861])).
fof(f19325,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) ) | $spl9),
  inference(resolution,[],[f2575,f706])).
fof(f19322,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,X16) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) ) | $spl9),
  inference(resolution,[],[f2575,f705])).
fof(f19321,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl9),
  inference(resolution,[],[f2575,f704])).
fof(f19320,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) ) | $spl9),
  inference(resolution,[],[f2575,f702])).
fof(f19318,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl9),
  inference(resolution,[],[f2575,f700])).
fof(f1808,plain,(
  ( ! [X4] : (s(bool,X4) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | ~p(s(bool,X4))) ) | $spl64),
  inference(cnf_transformation,[],[f1808_D])).
fof(f1808_D,plain,(
  ( ! [X4] : (s(bool,X4) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | ~p(s(bool,X4))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f18326,plain,(
  ( ! [X225,X226] : (p(s(bool,X225)) | ~p(s(bool,X226)) | s(bool,X226) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) ) | $spl5),
  inference(resolution,[],[f1135,f957])).
fof(f18224,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | s(bool,X24) != s(bool,X26)) ) | $spl5),
  inference(resolution,[],[f1135,f706])).
fof(f18222,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) ) | $spl5),
  inference(resolution,[],[f1135,f705])).
fof(f18220,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) ) | $spl5),
  inference(resolution,[],[f1135,f704])).
fof(f18219,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | s(bool,X9) != s(bool,X11)) ) | $spl5),
  inference(resolution,[],[f1135,f702])).
fof(f18217,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) ) | $spl5),
  inference(resolution,[],[f1135,f700])).
fof(f1135,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) | $spl5),
  inference(cnf_transformation,[],[f1135_D])).
fof(f1135_D,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f18214,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X3,bool),X4))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK16,bool),sK20)))))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X3),X5),s(fun(sK17,bool),sK18),s(fun(X3,bool),X4))))) ) | $spl414),
  inference(cnf_transformation,[],[f18214_D])).
fof(f18214_D,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X3,bool),X4))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK16,bool),sK20)))))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X3),X5),s(fun(sK17,bool),sK18),s(fun(X3,bool),X4))))) ) <=> ~$spl414),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl414])])).
fof(f18213,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X3,bool),X4))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK16,bool),sK20)))))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X3),X5),s(fun(sK17,bool),sK18),s(fun(X3,bool),X4)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))))) ) | $spl10),
  inference(subsumption_resolution,[],[f16006,f676])).
fof(f16006,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X3,bool),X4))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK16,bool),sK20)))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X3,bool),X4)))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X3),X5),s(fun(sK17,bool),sK18),s(fun(X3,bool),X4)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))))) ) | $spl10),
  inference(superposition,[],[f1141,f682])).
fof(f18210,plain,(
  $false | ($spl64 | $spl70)),
  inference(equality_resolution,[],[f18170])).
fof(f18170,plain,(
  ( ! [X22] : (s(bool,X22) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) ) | ($spl64 | $spl70)),
  inference(subsumption_resolution,[],[f1975,f1808])).
fof(f17564,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK17,bool),sK18),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | $spl412),
  inference(cnf_transformation,[],[f17564_D])).
fof(f17564_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK17,bool),sK18),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) <=> ~$spl412),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl412])])).
fof(f17530,plain,(
  ( ! [X222,X223] : (p(s(bool,X222)) | ~p(s(bool,X223)) | s(bool,X223) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK17,bool),sK18),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | $spl52),
  inference(resolution,[],[f14348,f957])).
fof(f14348,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK17,bool),sK18),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | $spl52),
  inference(subsumption_resolution,[],[f14192,f950])).
fof(f14192,plain,(
  ( ! [X20] : (~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK17),X20),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK17,bool),sK18),s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl52),
  inference(resolution,[],[f1360,f683])).
fof(f683,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X0),sK26(X2,X1,X0)),s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X0),X4),s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f499])).
fof(f499,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))) | ! [X4] : ~p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X0),X4),s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X0),sK26(X2,X1,X0)),s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK26])],[f498])).
fof(f498,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))) | ! [X4] : ~p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X0),X4),s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ? [X3] : p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X0),X3),s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))),
  inference(rectify,[],[f415])).
fof(f415,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))) | ! [X3] : ~p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X0),X3),s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ? [X4] : p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X0),X4),s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))),
  inference(flattening,[],[f414])).
fof(f414,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))) | ! [X3] : ~p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X0),X3),s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) | ? [X4] : p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X0),X4),s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))),
  inference(ennf_transformation,[],[f284])).
fof(f284,plain,(
  ! [X0,X1,X2] : ((p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))) & ? [X3] : p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X0),X3),s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) => ? [X4] : p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X0),X4),s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))),
  inference(flattening,[],[f283])).
fof(f283,plain,(
  ! [X0] : ! [X1,X2] : ((p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))) & ? [X3] : p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X0),X3),s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) => ? [X4] : p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X0),X4),s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X7] : ! [X8,X10] : ((p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X7,bool),X10),s(fun(X7,bool),X8)))) & ? [X9] : p(s(bool,'const.pred_set.INJ_3'(s(fun(X7,X7),X9),s(fun(X7,bool),X8),s(fun(X7,bool),X10))))) => ? [X11] : p(s(bool,'const.pred_set.BIJ_3'(s(fun(X7,X7),X11),s(fun(X7,bool),X8),s(fun(X7,bool),X10)))))),
  file('thm.util_prob.INJ_IMAGE_BIJ','thm.util_prob.SCHROEDER_BERNSTEIN_AUTO')).
fof(f17428,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK17,bool),sK18),s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | s(bool,X21) != s(bool,X23)) ) | $spl52),
  inference(resolution,[],[f14348,f706])).
fof(f17426,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK17,bool),sK18),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | $spl52),
  inference(resolution,[],[f14348,f705])).
fof(f17424,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK17,bool),sK18),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | $spl52),
  inference(resolution,[],[f14348,f704])).
fof(f17423,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK17,bool),sK18),s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | s(bool,X6) != s(bool,X8)) ) | $spl52),
  inference(resolution,[],[f14348,f702])).
fof(f17421,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK17,bool),sK18),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | $spl52),
  inference(resolution,[],[f14348,f700])).
fof(f17382,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK16,bool),sK20),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) | $spl410),
  inference(cnf_transformation,[],[f17382_D])).
fof(f17382_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK16,bool),sK20),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) <=> ~$spl410),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl410])])).
fof(f17348,plain,(
  ( ! [X222,X223] : (p(s(bool,X222)) | ~p(s(bool,X223)) | s(bool,X223) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK16,bool),sK20),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) | $spl30),
  inference(resolution,[],[f12447,f957])).
fof(f12447,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK16,bool),sK20),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | $spl30),
  inference(subsumption_resolution,[],[f12291,f950])).
fof(f12291,plain,(
  ( ! [X20] : (~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK16,sK16),X20),s(fun(sK16,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK16,bool),sK20),s(fun(sK16,bool),'const.pred_set.EMPTY_0'))))) ) | $spl30),
  inference(resolution,[],[f1176,f683])).
fof(f17246,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK16,bool),sK20),s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | s(bool,X21) != s(bool,X23)) ) | $spl30),
  inference(resolution,[],[f12447,f706])).
fof(f17244,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK16,bool),sK20),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) | $spl30),
  inference(resolution,[],[f12447,f705])).
fof(f17242,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK16,bool),sK20),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) | $spl30),
  inference(resolution,[],[f12447,f704])).
fof(f17241,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK16,bool),sK20),s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | s(bool,X6) != s(bool,X8)) ) | $spl30),
  inference(resolution,[],[f12447,f702])).
fof(f17239,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK16,bool),sK20),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) | $spl30),
  inference(resolution,[],[f12447,f700])).
fof(f16834,plain,(
  $false | $spl16),
  inference(equality_resolution,[],[f1144])).
fof(f15563,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK16,bool),sK20)))))) ) | $spl406),
  inference(cnf_transformation,[],[f15563_D])).
fof(f15563_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK16,bool),sK20)))))) ) <=> ~$spl406),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl406])])).
fof(f15992,plain,(
  ( ! [X239,X238] : (p(s(bool,X238)) | p(s(bool,X239)) | s(bool,X239) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK16,bool),sK20)))))) ) | $spl10),
  inference(resolution,[],[f1141,f957])).
fof(f15885,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK16,bool),sK20))))) | s(bool,X22) != s(bool,X23)) ) | $spl10),
  inference(resolution,[],[f1141,f706])).
fof(f15882,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK16,bool),sK20)))))) ) | $spl10),
  inference(resolution,[],[f1141,f705])).
fof(f15881,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK16,bool),sK20)))))) ) | $spl10),
  inference(resolution,[],[f1141,f704])).
fof(f15880,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK16,bool),sK20))))) | s(bool,X7) != s(bool,X8)) ) | $spl10),
  inference(resolution,[],[f1141,f702])).
fof(f15878,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK16,bool),sK20)))))) ) | $spl10),
  inference(resolution,[],[f1141,f700])).
fof(f15825,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK16,bool),sK20)))))) ) | $spl408),
  inference(cnf_transformation,[],[f15825_D])).
fof(f15825_D,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK16,bool),sK20)))))) ) <=> ~$spl408),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl408])])).
fof(f15782,plain,(
  ( ! [X223,X224] : (p(s(bool,X223)) | ~p(s(bool,X224)) | s(bool,X224) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK16,bool),sK20)))))) ) | $spl11),
  inference(resolution,[],[f1140,f957])).
fof(f1140,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK16,bool),sK20)))))) | $spl11),
  inference(cnf_transformation,[],[f1140_D])).
fof(f1140_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK16,bool),sK20)))))) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f15680,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK16,bool),sK20))))) | s(bool,X22) != s(bool,X24)) ) | $spl11),
  inference(resolution,[],[f1140,f706])).
fof(f15678,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK16,bool),sK20)))))) ) | $spl11),
  inference(resolution,[],[f1140,f705])).
fof(f15676,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK16,bool),sK20)))))) ) | $spl11),
  inference(resolution,[],[f1140,f704])).
fof(f15675,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK16,bool),sK20))))) | s(bool,X7) != s(bool,X9)) ) | $spl11),
  inference(resolution,[],[f1140,f702])).
fof(f15673,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK16,bool),sK20)))))) ) | $spl11),
  inference(resolution,[],[f1140,f700])).
fof(f15517,plain,(
  ( ! [X239,X238] : (p(s(bool,X238)) | p(s(bool,X239)) | s(bool,X239) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK16,bool),sK20)))))) ) | $spl10),
  inference(resolution,[],[f1141,f957])).
fof(f15410,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK16,bool),sK20))))) | s(bool,X22) != s(bool,X23)) ) | $spl10),
  inference(resolution,[],[f1141,f706])).
fof(f15407,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK16,bool),sK20)))))) ) | $spl10),
  inference(resolution,[],[f1141,f705])).
fof(f15406,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK16,bool),sK20)))))) ) | $spl10),
  inference(resolution,[],[f1141,f704])).
fof(f15405,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK16,bool),sK20))))) | s(bool,X7) != s(bool,X8)) ) | $spl10),
  inference(resolution,[],[f1141,f702])).
fof(f15403,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK16,bool),sK20)))))) ) | $spl10),
  inference(resolution,[],[f1141,f700])).
fof(f14799,plain,(
  ( ! [X14,X16] : (~p(s(bool,X14)) | s(bool,X14) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),X16)))) ) | $spl404),
  inference(cnf_transformation,[],[f14799_D])).
fof(f14799_D,plain,(
  ( ! [X14,X16] : (~p(s(bool,X14)) | s(bool,X14) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),X16)))) ) <=> ~$spl404),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl404])])).
fof(f14748,plain,(
  ( ! [X344,X343,X345] : (p(s(bool,X343)) | ~p(s(bool,X344)) | s(bool,X344) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),X345)))) )),
  inference(resolution,[],[f1363,f957])).
fof(f1363,plain,(
  ( ! [X0] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),X0))))) )),
  inference(subsumption_resolution,[],[f1307,f674])).
fof(f674,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))) )),
  inference(cnf_transformation,[],[f495])).
fof(f495,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | (p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) & p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))))) & (~p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | ~p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))))),
  inference(flattening,[],[f494])).
fof(f494,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | (p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) & p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))))) & ((~p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | ~p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))) | p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))))),
  inference(nnf_transformation,[],[f272])).
fof(f272,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) <=> (p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) & p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))))),
  inference(flattening,[],[f271])).
fof(f271,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) <=> (p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) & p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X7,X6] : ! [X9,X8,X10] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(X7,X6),X9),s(fun(X7,bool),X8),s(fun(X6,bool),X10)))) <=> (p(s(bool,'const.pred_set.INJ_3'(s(fun(X7,X6),X9),s(fun(X7,bool),X8),s(fun(X6,bool),X10)))) & p(s(bool,'const.pred_set.SURJ_3'(s(fun(X7,X6),X9),s(fun(X7,bool),X8),s(fun(X6,bool),X10))))))),
  file('thm.util_prob.INJ_IMAGE_BIJ','thm.pred_set.BIJ_DEF')).
fof(f1307,plain,(
  ( ! [X0] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),X0)))) | ~p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),X0))))) )),
  inference(superposition,[],[f648,f690])).
fof(f690,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun(X1,bool),X4) = s(fun(X1,bool),'const.pred_set.IMAGE_2'(s(fun(X0,X1),X2),s(fun(X0,bool),X3))) | ~p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))) )),
  inference(cnf_transformation,[],[f504])).
fof(f14646,plain,(
  ( ! [X43,X41,X42,X40] : (~p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X41) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),X42))) | s(bool,X40) != s(bool,X43)) )),
  inference(resolution,[],[f1363,f706])).
fof(f14644,plain,(
  ( ! [X35,X33,X36,X34] : (p(s(bool,X33)) | ~p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X34) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),X36)))) )),
  inference(resolution,[],[f1363,f705])).
fof(f14642,plain,(
  ( ! [X28,X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),X28)))) )),
  inference(resolution,[],[f1363,f704])).
fof(f14641,plain,(
  ( ! [X24,X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),X23))) | s(bool,X21) != s(bool,X24)) )),
  inference(resolution,[],[f1363,f702])).
fof(f14639,plain,(
  ( ! [X14,X15,X13,X16] : (p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),X16)))) )),
  inference(resolution,[],[f1363,f700])).
fof(f14536,plain,(
  ( ! [X24,X25,X22] : (~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X24),X25),s(fun(sK17,bool),sK18),s(fun(X24,bool),'const.pred_set.EMPTY_0')))) ) | $spl402),
  inference(cnf_transformation,[],[f14536_D])).
fof(f14536_D,plain,(
  ( ! [X24,X25,X22] : (~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X24),X25),s(fun(sK17,bool),sK18),s(fun(X24,bool),'const.pred_set.EMPTY_0')))) ) <=> ~$spl402),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl402])])).
fof(f14481,plain,(
  ( ! [X461,X459,X460,X462] : (p(s(bool,X459)) | ~p(s(bool,X460)) | s(bool,X460) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X461),X462),s(fun(sK17,bool),sK18),s(fun(X461,bool),'const.pred_set.EMPTY_0')))) ) | $spl54),
  inference(resolution,[],[f1362,f957])).
fof(f14379,plain,(
  ( ! [X57,X54,X58,X56,X55] : (~p(s(bool,X54)) | ~p(s(bool,X55)) | s(bool,X55) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X56),X57),s(fun(sK17,bool),sK18),s(fun(X56,bool),'const.pred_set.EMPTY_0'))) | s(bool,X54) != s(bool,X58)) ) | $spl54),
  inference(resolution,[],[f1362,f706])).
fof(f14377,plain,(
  ( ! [X47,X50,X48,X46,X49] : (p(s(bool,X46)) | ~p(s(bool,X47)) | ~p(s(bool,X48)) | s(bool,X47) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X49),X50),s(fun(sK17,bool),sK18),s(fun(X49,bool),'const.pred_set.EMPTY_0')))) ) | $spl54),
  inference(resolution,[],[f1362,f705])).
fof(f14375,plain,(
  ( ! [X39,X37,X38,X36,X40] : (p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,X38) | s(bool,X37) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X39),X40),s(fun(sK17,bool),sK18),s(fun(X39,bool),'const.pred_set.EMPTY_0')))) ) | $spl54),
  inference(resolution,[],[f1362,f704])).
fof(f14374,plain,(
  ( ! [X35,X33,X31,X34,X32] : (~p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X32) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X33),X34),s(fun(sK17,bool),sK18),s(fun(X33,bool),'const.pred_set.EMPTY_0'))) | s(bool,X31) != s(bool,X35)) ) | $spl54),
  inference(resolution,[],[f1362,f702])).
fof(f14372,plain,(
  ( ! [X24,X23,X21,X25,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X24),X25),s(fun(sK17,bool),sK18),s(fun(X24,bool),'const.pred_set.EMPTY_0')))) ) | $spl54),
  inference(resolution,[],[f1362,f700])).
fof(f1362,plain,(
  ( ! [X8,X7] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X7),X8),s(fun(sK17,bool),sK18),s(fun(X7,bool),'const.pred_set.EMPTY_0'))))) ) | $spl54),
  inference(cnf_transformation,[],[f1362_D])).
fof(f1362_D,plain,(
  ( ! [X8,X7] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X7),X8),s(fun(sK17,bool),sK18),s(fun(X7,bool),'const.pred_set.EMPTY_0'))))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f14362,plain,(
  ( ! [X12,X11] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X11,sK17),X12),s(fun(X11,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl400),
  inference(cnf_transformation,[],[f14362_D])).
fof(f14362_D,plain,(
  ( ! [X12,X11] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X11,sK17),X12),s(fun(X11,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) <=> ~$spl400),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl400])])).
fof(f14328,plain,(
  ( ! [X14,X12,X13,X11] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X11,sK17),X12),s(fun(X11,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl52),
  inference(superposition,[],[f1360,f678])).
fof(f14349,plain,(
  ( ! [X24,X25,X22] : (~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.pred_set.BIJ_3'(s(fun(X24,sK17),X25),s(fun(X24,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18)))) ) | $spl398),
  inference(cnf_transformation,[],[f14349_D])).
fof(f14349_D,plain,(
  ( ! [X24,X25,X22] : (~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.pred_set.BIJ_3'(s(fun(X24,sK17),X25),s(fun(X24,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18)))) ) <=> ~$spl398),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl398])])).
fof(f14302,plain,(
  ( ! [X461,X459,X460,X462] : (p(s(bool,X459)) | ~p(s(bool,X460)) | s(bool,X460) != s(bool,'const.pred_set.BIJ_3'(s(fun(X461,sK17),X462),s(fun(X461,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18)))) ) | $spl52),
  inference(resolution,[],[f1360,f957])).
fof(f14200,plain,(
  ( ! [X57,X54,X58,X56,X55] : (~p(s(bool,X54)) | ~p(s(bool,X55)) | s(bool,X55) != s(bool,'const.pred_set.BIJ_3'(s(fun(X56,sK17),X57),s(fun(X56,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))) | s(bool,X54) != s(bool,X58)) ) | $spl52),
  inference(resolution,[],[f1360,f706])).
fof(f14198,plain,(
  ( ! [X47,X50,X48,X46,X49] : (p(s(bool,X46)) | ~p(s(bool,X47)) | ~p(s(bool,X48)) | s(bool,X47) != s(bool,'const.pred_set.BIJ_3'(s(fun(X49,sK17),X50),s(fun(X49,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18)))) ) | $spl52),
  inference(resolution,[],[f1360,f705])).
fof(f14196,plain,(
  ( ! [X39,X37,X38,X36,X40] : (p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,X38) | s(bool,X37) != s(bool,'const.pred_set.BIJ_3'(s(fun(X39,sK17),X40),s(fun(X39,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18)))) ) | $spl52),
  inference(resolution,[],[f1360,f704])).
fof(f14195,plain,(
  ( ! [X35,X33,X31,X34,X32] : (~p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X32) != s(bool,'const.pred_set.BIJ_3'(s(fun(X33,sK17),X34),s(fun(X33,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))) | s(bool,X31) != s(bool,X35)) ) | $spl52),
  inference(resolution,[],[f1360,f702])).
fof(f14193,plain,(
  ( ! [X24,X23,X21,X25,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,'const.pred_set.BIJ_3'(s(fun(X24,sK17),X25),s(fun(X24,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18)))) ) | $spl52),
  inference(resolution,[],[f1360,f700])).
fof(f14154,plain,(
  ( ! [X10] : (p(s(bool,X10)) | s(bool,X10) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) | $spl396),
  inference(cnf_transformation,[],[f14154_D])).
fof(f14154_D,plain,(
  ( ! [X10] : (p(s(bool,X10)) | s(bool,X10) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) <=> ~$spl396),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl396])])).
fof(f14110,plain,(
  ( ! [X248,X247] : (p(s(bool,X247)) | p(s(bool,X248)) | s(bool,X248) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) | $spl50),
  inference(resolution,[],[f1357,f957])).
fof(f14003,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X30) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | s(bool,X31) != s(bool,X32)) ) | $spl50),
  inference(resolution,[],[f1357,f706])).
fof(f14000,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | p(s(bool,X23)) | s(bool,X23) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) | $spl50),
  inference(resolution,[],[f1357,f705])).
fof(f13999,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X18) != s(bool,X20) | s(bool,X19) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) | $spl50),
  inference(resolution,[],[f1357,f704])).
fof(f13998,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | s(bool,X16) != s(bool,X17)) ) | $spl50),
  inference(resolution,[],[f1357,f702])).
fof(f13996,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) | $spl50),
  inference(resolution,[],[f1357,f700])).
fof(f13945,plain,(
  ( ! [X225,X226] : (p(s(bool,X225)) | ~p(s(bool,X226)) | s(bool,X226) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | $spl69),
  inference(resolution,[],[f1824,f957])).
fof(f13843,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | s(bool,X24) != s(bool,X26)) ) | $spl69),
  inference(resolution,[],[f1824,f706])).
fof(f13841,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | $spl69),
  inference(resolution,[],[f1824,f705])).
fof(f13839,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | $spl69),
  inference(resolution,[],[f1824,f704])).
fof(f13838,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | s(bool,X9) != s(bool,X11)) ) | $spl69),
  inference(resolution,[],[f1824,f702])).
fof(f13836,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | $spl69),
  inference(resolution,[],[f1824,f700])).
fof(f3465,plain,(
  ( ! [X0,X1] : (sP8('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0,X1)) ) | $spl96),
  inference(cnf_transformation,[],[f3465_D])).
fof(f3465_D,plain,(
  ( ! [X0,X1] : (sP8('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0,X1)) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f13827,plain,(
  ( ! [X14,X12,X13,X11] : (sP8('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X11,X12) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f1736,f678])).
fof(f1736,plain,(
  ( ! [X134,X133] : (sP8('const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)),X133,X134)) ) | $spl5),
  inference(resolution,[],[f1135,f772])).
fof(f1360,plain,(
  ( ! [X6,X5] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X5,sK17),X6),s(fun(X5,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl52),
  inference(cnf_transformation,[],[f1360_D])).
fof(f1360_D,plain,(
  ( ! [X6,X5] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X5,sK17),X6),s(fun(X5,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f13826,plain,(
  ( ! [X10,X8,X7,X9] : (sP8('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X7,X8) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f1736,f680])).
fof(f3450,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X1)) ) | $spl94),
  inference(cnf_transformation,[],[f3450_D])).
fof(f3450_D,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X1)) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f13813,plain,(
  ( ! [X14,X12,X13,X11] : (sP8(X11,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X12) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f1735,f678])).
fof(f1735,plain,(
  ( ! [X132,X131] : (sP8(X131,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)),X132)) ) | $spl5),
  inference(resolution,[],[f1135,f771])).
fof(f13812,plain,(
  ( ! [X10,X8,X7,X9] : (sP8(X7,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X8) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f1735,f680])).
fof(f3079,plain,(
  ( ! [X0,X1] : (sP9('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0,X1)) ) | $spl92),
  inference(cnf_transformation,[],[f3079_D])).
fof(f3079_D,plain,(
  ( ! [X0,X1] : (sP9('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0,X1)) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f13799,plain,(
  ( ! [X14,X12,X13,X11] : (sP9('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X11,X12) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f1733,f678])).
fof(f1733,plain,(
  ( ! [X127,X128] : (sP9('const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)),X127,X128)) ) | $spl5),
  inference(resolution,[],[f1135,f768])).
fof(f13798,plain,(
  ( ! [X10,X8,X7,X9] : (sP9('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X7,X8) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f1733,f680])).
fof(f3064,plain,(
  ( ! [X0,X1] : (sP9(X0,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X1)) ) | $spl90),
  inference(cnf_transformation,[],[f3064_D])).
fof(f3064_D,plain,(
  ( ! [X0,X1] : (sP9(X0,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X1)) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f13786,plain,(
  ( ! [X14,X12,X13,X11] : (sP9(X11,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X12) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f1732,f678])).
fof(f1732,plain,(
  ( ! [X125,X126] : (sP9(X125,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)),X126)) ) | $spl5),
  inference(resolution,[],[f1135,f767])).
fof(f13785,plain,(
  ( ! [X10,X8,X7,X9] : (sP9(X7,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X8) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f1732,f680])).
fof(f12818,plain,(
  ( ! [X19,X17,X20] : (~p(s(bool,X17)) | s(bool,X17) != s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X19),X20),s(fun(sK17,bool),sK18),s(fun(X19,bool),'const.pred_set.EMPTY_0')))) ) | $spl392),
  inference(cnf_transformation,[],[f12818_D])).
fof(f12818_D,plain,(
  ( ! [X19,X17,X20] : (~p(s(bool,X17)) | s(bool,X17) != s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X19),X20),s(fun(sK17,bool),sK18),s(fun(X19,bool),'const.pred_set.EMPTY_0')))) ) <=> ~$spl392),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl392])])).
fof(f13711,plain,(
  ( ! [X455,X457,X454,X456] : (p(s(bool,X454)) | ~p(s(bool,X455)) | s(bool,X455) != s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X456),X457),s(fun(sK17,bool),sK18),s(fun(X456,bool),'const.pred_set.EMPTY_0')))) ) | $spl48),
  inference(resolution,[],[f1356,f957])).
fof(f13609,plain,(
  ( ! [X52,X50,X53,X51,X49] : (~p(s(bool,X49)) | ~p(s(bool,X50)) | s(bool,X50) != s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X51),X52),s(fun(sK17,bool),sK18),s(fun(X51,bool),'const.pred_set.EMPTY_0'))) | s(bool,X49) != s(bool,X53)) ) | $spl48),
  inference(resolution,[],[f1356,f706])).
fof(f13607,plain,(
  ( ! [X45,X43,X41,X44,X42] : (p(s(bool,X41)) | ~p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X42) != s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X44),X45),s(fun(sK17,bool),sK18),s(fun(X44,bool),'const.pred_set.EMPTY_0')))) ) | $spl48),
  inference(resolution,[],[f1356,f705])).
fof(f13605,plain,(
  ( ! [X35,X33,X31,X34,X32] : (p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,X33) | s(bool,X32) != s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X34),X35),s(fun(sK17,bool),sK18),s(fun(X34,bool),'const.pred_set.EMPTY_0')))) ) | $spl48),
  inference(resolution,[],[f1356,f704])).
fof(f13604,plain,(
  ( ! [X30,X28,X26,X29,X27] : (~p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X28),X29),s(fun(sK17,bool),sK18),s(fun(X28,bool),'const.pred_set.EMPTY_0'))) | s(bool,X26) != s(bool,X30)) ) | $spl48),
  inference(resolution,[],[f1356,f702])).
fof(f13602,plain,(
  ( ! [X19,X17,X20,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,X18) | s(bool,X17) != s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X19),X20),s(fun(sK17,bool),sK18),s(fun(X19,bool),'const.pred_set.EMPTY_0')))) ) | $spl48),
  inference(resolution,[],[f1356,f700])).
fof(f3050,plain,(
  ( ! [X0,X1] : (sP5('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0,X1)) ) | $spl88),
  inference(cnf_transformation,[],[f3050_D])).
fof(f3050_D,plain,(
  ( ! [X0,X1] : (sP5('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0,X1)) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f13590,plain,(
  ( ! [X14,X12,X13,X11] : (sP5('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X11,X12) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f1720,f678])).
fof(f1720,plain,(
  ( ! [X101,X102] : (sP5('const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)),X101,X102)) ) | $spl5),
  inference(resolution,[],[f1135,f751])).
fof(f13589,plain,(
  ( ! [X10,X8,X7,X9] : (sP5('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X7,X8) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f1720,f680])).
fof(f1356,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X0),X1),s(fun(sK17,bool),sK18),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl48),
  inference(cnf_transformation,[],[f1356_D])).
fof(f1356_D,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X0),X1),s(fun(sK17,bool),sK18),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f13588,plain,(
  ( ! [X2,X0,X3,X1] : (sP5('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0,X1) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f1720,f948])).
fof(f3036,plain,(
  ( ! [X0,X1] : (sP5(X0,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X1)) ) | $spl86),
  inference(cnf_transformation,[],[f3036_D])).
fof(f3036_D,plain,(
  ( ! [X0,X1] : (sP5(X0,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X1)) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f13577,plain,(
  ( ! [X14,X12,X13,X11] : (sP5(X11,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X12) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f1719,f678])).
fof(f1719,plain,(
  ( ! [X99,X100] : (sP5(X99,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)),X100)) ) | $spl5),
  inference(resolution,[],[f1135,f750])).
fof(f13576,plain,(
  ( ! [X10,X8,X7,X9] : (sP5(X7,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X8) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f1719,f680])).
fof(f13575,plain,(
  ( ! [X2,X0,X3,X1] : (sP5(X0,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X1) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f1719,f948])).
fof(f2662,plain,(
  ( ! [X0,X1] : (sP3('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0,X1)) ) | $spl84),
  inference(cnf_transformation,[],[f2662_D])).
fof(f2662_D,plain,(
  ( ! [X0,X1] : (sP3('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0,X1)) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f13563,plain,(
  ( ! [X14,X12,X13,X11] : (sP3('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X11,X12) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f1706,f678])).
fof(f1706,plain,(
  ( ! [X74,X73] : (sP3('const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)),X73,X74)) ) | $spl5),
  inference(resolution,[],[f1135,f734])).
fof(f13562,plain,(
  ( ! [X10,X8,X7,X9] : (sP3('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X7,X8) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f1706,f680])).
fof(f13561,plain,(
  ( ! [X2,X0,X3,X1] : (sP3('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0,X1) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f1706,f948])).
fof(f2647,plain,(
  ( ! [X0,X1] : (sP1('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0,X1)) ) | $spl82),
  inference(cnf_transformation,[],[f2647_D])).
fof(f2647_D,plain,(
  ( ! [X0,X1] : (sP1('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0,X1)) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f13550,plain,(
  ( ! [X14,X12,X13,X11] : (sP1('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X11,X12) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f1690,f678])).
fof(f1690,plain,(
  ( ! [X41,X42] : (sP1('const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)),X41,X42)) ) | $spl5),
  inference(resolution,[],[f1135,f717])).
fof(f13549,plain,(
  ( ! [X10,X8,X7,X9] : (sP1('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X7,X8) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f1690,f680])).
fof(f13548,plain,(
  ( ! [X2,X0,X3,X1] : (sP1('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0,X1) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f1690,f948])).
fof(f13520,plain,(
  $false | $spl51),
  inference(trivial_inequality_removal,[],[f13361])).
fof(f13361,plain,(
  s(fun(sK16,bool),'const.pred_set.EMPTY_0') != s(fun(sK16,bool),'const.pred_set.EMPTY_0') | $spl51),
  inference(resolution,[],[f1358,f681])).
fof(f681,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(X1,bool),X3)))) | s(fun(X1,bool),'const.pred_set.EMPTY_0') != s(fun(X1,bool),X3)) )),
  inference(cnf_transformation,[],[f497])).
fof(f13519,plain,(
  $false | $spl51),
  inference(trivial_inequality_removal,[],[f13362])).
fof(f13362,plain,(
  s(fun(sK17,bool),'const.pred_set.EMPTY_0') != s(fun(sK17,bool),'const.pred_set.EMPTY_0') | $spl51),
  inference(resolution,[],[f1358,f679])).
fof(f679,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) | s(fun(X1,bool),'const.pred_set.EMPTY_0') != s(fun(X1,bool),X3)) )),
  inference(cnf_transformation,[],[f496])).
fof(f2633,plain,(
  ( ! [X0,X1] : (sP1(X0,X1,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | $spl80),
  inference(cnf_transformation,[],[f2633_D])).
fof(f2633_D,plain,(
  ( ! [X0,X1] : (sP1(X0,X1,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f13355,plain,(
  ( ! [X14,X12,X13,X11] : (sP1(X11,X12,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f1689,f678])).
fof(f1689,plain,(
  ( ! [X39,X40] : (sP1(X39,X40,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) ) | $spl5),
  inference(resolution,[],[f1135,f715])).
fof(f13354,plain,(
  ( ! [X10,X8,X7,X9] : (sP1(X7,X8,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f1689,f680])).
fof(f13353,plain,(
  ( ! [X2,X0,X3,X1] : (sP1(X0,X1,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f1689,f948])).
fof(f13317,plain,(
  ( ! [X8,X7] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X7),X8),s(fun(sK17,bool),sK18),s(fun(X7,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f1135,f678])).
fof(f13316,plain,(
  ( ! [X6,X5] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X5,sK17),X6),s(fun(X5,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f1135,f680])).
fof(f13315,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X0),X1),s(fun(sK17,bool),sK18),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f1135,f948])).
fof(f13312,plain,(
  ( ! [X225,X226] : (p(s(bool,X225)) | ~p(s(bool,X226)) | s(bool,X226) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) ) | $spl5),
  inference(resolution,[],[f1135,f957])).
fof(f13210,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | s(bool,X24) != s(bool,X26)) ) | $spl5),
  inference(resolution,[],[f1135,f706])).
fof(f13208,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) ) | $spl5),
  inference(resolution,[],[f1135,f705])).
fof(f13206,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) ) | $spl5),
  inference(resolution,[],[f1135,f704])).
fof(f13205,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | s(bool,X9) != s(bool,X11)) ) | $spl5),
  inference(resolution,[],[f1135,f702])).
fof(f13203,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) ) | $spl5),
  inference(resolution,[],[f1135,f700])).
fof(f13201,plain,(
  ( ! [X0,X1] : (sP1(X0,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X1)) ) | $spl394),
  inference(cnf_transformation,[],[f13201_D])).
fof(f13201_D,plain,(
  ( ! [X0,X1] : (sP1(X0,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X1)) ) <=> ~$spl394),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl394])])).
fof(f13195,plain,(
  ( ! [X14,X12,X13,X11] : (sP1(X11,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X12) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl4),
  inference(superposition,[],[f1849,f678])).
fof(f1849,plain,(
  ( ! [X59,X60] : (sP1(X59,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)),X60)) ) | $spl4),
  inference(resolution,[],[f1134,f716])).
fof(f13194,plain,(
  ( ! [X10,X8,X7,X9] : (sP1(X7,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X8) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl4),
  inference(superposition,[],[f1849,f680])).
fof(f13193,plain,(
  ( ! [X2,X0,X3,X1] : (sP1(X0,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X1) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl4),
  inference(superposition,[],[f1849,f948])).
fof(f3775,plain,(
  ( ! [X22] : (p(s(bool,X22)) | s(bool,X22) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | $spl102),
  inference(cnf_transformation,[],[f3775_D])).
fof(f3775_D,plain,(
  ( ! [X22] : (p(s(bool,X22)) | s(bool,X22) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f13130,plain,(
  ( ! [X259,X260] : (p(s(bool,X259)) | p(s(bool,X260)) | s(bool,X260) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | $spl68),
  inference(resolution,[],[f1823,f957])).
fof(f13023,plain,(
  ( ! [X43,X44,X42] : (p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X42) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | s(bool,X43) != s(bool,X44)) ) | $spl68),
  inference(resolution,[],[f1823,f706])).
fof(f13020,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X33)) | ~p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X35) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | $spl68),
  inference(resolution,[],[f1823,f705])).
fof(f13019,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | p(s(bool,X31)) | s(bool,X30) != s(bool,X32) | s(bool,X31) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | $spl68),
  inference(resolution,[],[f1823,f704])).
fof(f13018,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | s(bool,X28) != s(bool,X29)) ) | $spl68),
  inference(resolution,[],[f1823,f702])).
fof(f13016,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | $spl68),
  inference(resolution,[],[f1823,f700])).
fof(f12964,plain,(
  ( ! [X8,X7] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X7),X8),s(fun(sK17,bool),sK18),s(fun(X7,bool),'const.pred_set.EMPTY_0'))))) ) | $spl4),
  inference(superposition,[],[f1134,f678])).
fof(f12963,plain,(
  ( ! [X6,X5] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X5,sK17),X6),s(fun(X5,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl4),
  inference(superposition,[],[f1134,f680])).
fof(f12962,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X0),X1),s(fun(sK17,bool),sK18),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl4),
  inference(superposition,[],[f1134,f948])).
fof(f12956,plain,(
  ( ! [X259,X260] : (p(s(bool,X259)) | p(s(bool,X260)) | s(bool,X260) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) ) | $spl4),
  inference(resolution,[],[f1134,f957])).
fof(f12849,plain,(
  ( ! [X43,X44,X42] : (p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X42) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | s(bool,X43) != s(bool,X44)) ) | $spl4),
  inference(resolution,[],[f1134,f706])).
fof(f12846,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X33)) | ~p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X35) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) ) | $spl4),
  inference(resolution,[],[f1134,f705])).
fof(f12845,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | p(s(bool,X31)) | s(bool,X30) != s(bool,X32) | s(bool,X31) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) ) | $spl4),
  inference(resolution,[],[f1134,f704])).
fof(f12844,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | s(bool,X28) != s(bool,X29)) ) | $spl4),
  inference(resolution,[],[f1134,f702])).
fof(f12842,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) ) | $spl4),
  inference(resolution,[],[f1134,f700])).
fof(f12763,plain,(
  ( ! [X455,X457,X454,X456] : (p(s(bool,X454)) | ~p(s(bool,X455)) | s(bool,X455) != s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X456),X457),s(fun(sK17,bool),sK18),s(fun(X456,bool),'const.pred_set.EMPTY_0')))) ) | $spl48),
  inference(resolution,[],[f1356,f957])).
fof(f12661,plain,(
  ( ! [X52,X50,X53,X51,X49] : (~p(s(bool,X49)) | ~p(s(bool,X50)) | s(bool,X50) != s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X51),X52),s(fun(sK17,bool),sK18),s(fun(X51,bool),'const.pred_set.EMPTY_0'))) | s(bool,X49) != s(bool,X53)) ) | $spl48),
  inference(resolution,[],[f1356,f706])).
fof(f12659,plain,(
  ( ! [X45,X43,X41,X44,X42] : (p(s(bool,X41)) | ~p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X42) != s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X44),X45),s(fun(sK17,bool),sK18),s(fun(X44,bool),'const.pred_set.EMPTY_0')))) ) | $spl48),
  inference(resolution,[],[f1356,f705])).
fof(f12657,plain,(
  ( ! [X35,X33,X31,X34,X32] : (p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,X33) | s(bool,X32) != s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X34),X35),s(fun(sK17,bool),sK18),s(fun(X34,bool),'const.pred_set.EMPTY_0')))) ) | $spl48),
  inference(resolution,[],[f1356,f704])).
fof(f12656,plain,(
  ( ! [X30,X28,X26,X29,X27] : (~p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X28),X29),s(fun(sK17,bool),sK18),s(fun(X28,bool),'const.pred_set.EMPTY_0'))) | s(bool,X26) != s(bool,X30)) ) | $spl48),
  inference(resolution,[],[f1356,f702])).
fof(f12654,plain,(
  ( ! [X19,X17,X20,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,X18) | s(bool,X17) != s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X19),X20),s(fun(sK17,bool),sK18),s(fun(X19,bool),'const.pred_set.EMPTY_0')))) ) | $spl48),
  inference(resolution,[],[f1356,f700])).
fof(f12635,plain,(
  ( ! [X24,X25,X22] : (~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X24),X25),s(fun(sK16,bool),sK20),s(fun(X24,bool),'const.pred_set.EMPTY_0')))) ) | $spl390),
  inference(cnf_transformation,[],[f12635_D])).
fof(f12635_D,plain,(
  ( ! [X24,X25,X22] : (~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X24),X25),s(fun(sK16,bool),sK20),s(fun(X24,bool),'const.pred_set.EMPTY_0')))) ) <=> ~$spl390),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl390])])).
fof(f12580,plain,(
  ( ! [X461,X459,X460,X462] : (p(s(bool,X459)) | ~p(s(bool,X460)) | s(bool,X460) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X461),X462),s(fun(sK16,bool),sK20),s(fun(X461,bool),'const.pred_set.EMPTY_0')))) ) | $spl32),
  inference(resolution,[],[f1177,f957])).
fof(f12478,plain,(
  ( ! [X57,X54,X58,X56,X55] : (~p(s(bool,X54)) | ~p(s(bool,X55)) | s(bool,X55) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X56),X57),s(fun(sK16,bool),sK20),s(fun(X56,bool),'const.pred_set.EMPTY_0'))) | s(bool,X54) != s(bool,X58)) ) | $spl32),
  inference(resolution,[],[f1177,f706])).
fof(f12476,plain,(
  ( ! [X47,X50,X48,X46,X49] : (p(s(bool,X46)) | ~p(s(bool,X47)) | ~p(s(bool,X48)) | s(bool,X47) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X49),X50),s(fun(sK16,bool),sK20),s(fun(X49,bool),'const.pred_set.EMPTY_0')))) ) | $spl32),
  inference(resolution,[],[f1177,f705])).
fof(f12474,plain,(
  ( ! [X39,X37,X38,X36,X40] : (p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,X38) | s(bool,X37) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X39),X40),s(fun(sK16,bool),sK20),s(fun(X39,bool),'const.pred_set.EMPTY_0')))) ) | $spl32),
  inference(resolution,[],[f1177,f704])).
fof(f12473,plain,(
  ( ! [X35,X33,X31,X34,X32] : (~p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X32) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X33),X34),s(fun(sK16,bool),sK20),s(fun(X33,bool),'const.pred_set.EMPTY_0'))) | s(bool,X31) != s(bool,X35)) ) | $spl32),
  inference(resolution,[],[f1177,f702])).
fof(f12471,plain,(
  ( ! [X24,X23,X21,X25,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X24),X25),s(fun(sK16,bool),sK20),s(fun(X24,bool),'const.pred_set.EMPTY_0')))) ) | $spl32),
  inference(resolution,[],[f1177,f700])).
fof(f1177,plain,(
  ( ! [X8,X7] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X7),X8),s(fun(sK16,bool),sK20),s(fun(X7,bool),'const.pred_set.EMPTY_0'))))) ) | $spl32),
  inference(cnf_transformation,[],[f1177_D])).
fof(f1177_D,plain,(
  ( ! [X8,X7] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X7),X8),s(fun(sK16,bool),sK20),s(fun(X7,bool),'const.pred_set.EMPTY_0'))))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f12461,plain,(
  ( ! [X12,X11] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X11,sK16),X12),s(fun(X11,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),'const.pred_set.EMPTY_0'))))) ) | $spl388),
  inference(cnf_transformation,[],[f12461_D])).
fof(f12461_D,plain,(
  ( ! [X12,X11] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X11,sK16),X12),s(fun(X11,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),'const.pred_set.EMPTY_0'))))) ) <=> ~$spl388),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl388])])).
fof(f12427,plain,(
  ( ! [X14,X12,X13,X11] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X11,sK16),X12),s(fun(X11,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X13),X14),s(fun(sK16,bool),sK20),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl30),
  inference(superposition,[],[f1176,f678])).
fof(f12448,plain,(
  ( ! [X24,X25,X22] : (~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.pred_set.BIJ_3'(s(fun(X24,sK16),X25),s(fun(X24,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)))) ) | $spl386),
  inference(cnf_transformation,[],[f12448_D])).
fof(f12448_D,plain,(
  ( ! [X24,X25,X22] : (~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.pred_set.BIJ_3'(s(fun(X24,sK16),X25),s(fun(X24,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)))) ) <=> ~$spl386),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl386])])).
fof(f12401,plain,(
  ( ! [X461,X459,X460,X462] : (p(s(bool,X459)) | ~p(s(bool,X460)) | s(bool,X460) != s(bool,'const.pred_set.BIJ_3'(s(fun(X461,sK16),X462),s(fun(X461,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)))) ) | $spl30),
  inference(resolution,[],[f1176,f957])).
fof(f12299,plain,(
  ( ! [X57,X54,X58,X56,X55] : (~p(s(bool,X54)) | ~p(s(bool,X55)) | s(bool,X55) != s(bool,'const.pred_set.BIJ_3'(s(fun(X56,sK16),X57),s(fun(X56,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))) | s(bool,X54) != s(bool,X58)) ) | $spl30),
  inference(resolution,[],[f1176,f706])).
fof(f12297,plain,(
  ( ! [X47,X50,X48,X46,X49] : (p(s(bool,X46)) | ~p(s(bool,X47)) | ~p(s(bool,X48)) | s(bool,X47) != s(bool,'const.pred_set.BIJ_3'(s(fun(X49,sK16),X50),s(fun(X49,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)))) ) | $spl30),
  inference(resolution,[],[f1176,f705])).
fof(f12295,plain,(
  ( ! [X39,X37,X38,X36,X40] : (p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,X38) | s(bool,X37) != s(bool,'const.pred_set.BIJ_3'(s(fun(X39,sK16),X40),s(fun(X39,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)))) ) | $spl30),
  inference(resolution,[],[f1176,f704])).
fof(f12294,plain,(
  ( ! [X35,X33,X31,X34,X32] : (~p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X32) != s(bool,'const.pred_set.BIJ_3'(s(fun(X33,sK16),X34),s(fun(X33,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))) | s(bool,X31) != s(bool,X35)) ) | $spl30),
  inference(resolution,[],[f1176,f702])).
fof(f12292,plain,(
  ( ! [X24,X23,X21,X25,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,'const.pred_set.BIJ_3'(s(fun(X24,sK16),X25),s(fun(X24,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)))) ) | $spl30),
  inference(resolution,[],[f1176,f700])).
fof(f12282,plain,(
  ( ! [X8,X7] : (~sP13('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X7,X8) | p(s(bool,X7))) ) | $spl384),
  inference(cnf_transformation,[],[f12282_D])).
fof(f12282_D,plain,(
  ( ! [X8,X7] : (~sP13('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X7,X8) | p(s(bool,X7))) ) <=> ~$spl384),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl384])])).
fof(f12277,plain,(
  ( ! [X14,X12,X13,X11] : (~sP13('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X11,X12) | p(s(bool,X11)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X13),X14),s(fun(sK16,bool),sK20),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f2562,f678])).
fof(f2562,plain,(
  ( ! [X187,X188] : (~sP13('const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)),X187,X188) | p(s(bool,X187))) ) | $spl9),
  inference(resolution,[],[f1139,f808])).
fof(f1176,plain,(
  ( ! [X6,X5] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X5,sK16),X6),s(fun(X5,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) ) | $spl30),
  inference(cnf_transformation,[],[f1176_D])).
fof(f1176_D,plain,(
  ( ! [X6,X5] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X5,sK16),X6),s(fun(X5,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f12276,plain,(
  ( ! [X10,X8,X7,X9] : (~sP13('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X7,X8) | p(s(bool,X7)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK16),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) ) | $spl9),
  inference(superposition,[],[f2562,f680])).
fof(f12267,plain,(
  ( ! [X8,X7] : (~sP13(X7,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X8) | p(s(bool,X7))) ) | $spl382),
  inference(cnf_transformation,[],[f12267_D])).
fof(f12267_D,plain,(
  ( ! [X8,X7] : (~sP13(X7,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X8) | p(s(bool,X7))) ) <=> ~$spl382),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl382])])).
fof(f12262,plain,(
  ( ! [X14,X12,X13,X11] : (~sP13(X11,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X12) | p(s(bool,X11)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X13),X14),s(fun(sK16,bool),sK20),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f2561,f678])).
fof(f2561,plain,(
  ( ! [X185,X186] : (~sP13(X185,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)),X186) | p(s(bool,X185))) ) | $spl9),
  inference(resolution,[],[f1139,f808])).
fof(f12261,plain,(
  ( ! [X10,X8,X7,X9] : (~sP13(X7,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X8) | p(s(bool,X7)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK16),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) ) | $spl9),
  inference(superposition,[],[f2561,f680])).
fof(f12253,plain,(
  ( ! [X8,X7] : (~sP13(X7,X8,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X8))) ) | $spl380),
  inference(cnf_transformation,[],[f12253_D])).
fof(f12253_D,plain,(
  ( ! [X8,X7] : (~sP13(X7,X8,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X8))) ) <=> ~$spl380),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl380])])).
fof(f12248,plain,(
  ( ! [X14,X12,X13,X11] : (~sP13(X11,X12,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X12)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X13),X14),s(fun(sK16,bool),sK20),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f2560,f678])).
fof(f2560,plain,(
  ( ! [X184,X183] : (~sP13(X184,X183,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) | p(s(bool,X183))) ) | $spl9),
  inference(resolution,[],[f1139,f807])).
fof(f12247,plain,(
  ( ! [X10,X8,X7,X9] : (~sP13(X7,X8,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X8)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK16),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) ) | $spl9),
  inference(superposition,[],[f2560,f680])).
fof(f12238,plain,(
  ( ! [X8,X7] : (~sP13(X7,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X8) | p(s(bool,X8))) ) | $spl378),
  inference(cnf_transformation,[],[f12238_D])).
fof(f12238_D,plain,(
  ( ! [X8,X7] : (~sP13(X7,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X8) | p(s(bool,X8))) ) <=> ~$spl378),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl378])])).
fof(f12233,plain,(
  ( ! [X14,X12,X13,X11] : (~sP13(X11,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X12) | p(s(bool,X12)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X13),X14),s(fun(sK16,bool),sK20),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f2559,f678])).
fof(f2559,plain,(
  ( ! [X182,X181] : (~sP13(X182,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)),X181) | p(s(bool,X181))) ) | $spl9),
  inference(resolution,[],[f1139,f807])).
fof(f12232,plain,(
  ( ! [X10,X8,X7,X9] : (~sP13(X7,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X8) | p(s(bool,X8)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK16),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) ) | $spl9),
  inference(superposition,[],[f2559,f680])).
fof(f8094,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | s(bool,X4) != s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) | $spl272),
  inference(cnf_transformation,[],[f8094_D])).
fof(f8094_D,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | s(bool,X4) != s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) <=> ~$spl272),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl272])])).
fof(f12172,plain,(
  ( ! [X225,X226] : (p(s(bool,X225)) | ~p(s(bool,X226)) | s(bool,X226) != s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) | $spl29),
  inference(resolution,[],[f1174,f957])).
fof(f1174,plain,(
  ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | $spl29),
  inference(cnf_transformation,[],[f1174_D])).
fof(f1174_D,plain,(
  ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f12070,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | s(bool,X24) != s(bool,X26)) ) | $spl29),
  inference(resolution,[],[f1174,f706])).
fof(f12068,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) | $spl29),
  inference(resolution,[],[f1174,f705])).
fof(f12066,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) | $spl29),
  inference(resolution,[],[f1174,f704])).
fof(f12065,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | s(bool,X9) != s(bool,X11)) ) | $spl29),
  inference(resolution,[],[f1174,f702])).
fof(f12063,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) | $spl29),
  inference(resolution,[],[f1174,f700])).
fof(f12057,plain,(
  ( ! [X8,X7] : (~sP12('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X7,X8) | p(s(bool,X7))) ) | $spl376),
  inference(cnf_transformation,[],[f12057_D])).
fof(f12057_D,plain,(
  ( ! [X8,X7] : (~sP12('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X7,X8) | p(s(bool,X7))) ) <=> ~$spl376),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl376])])).
fof(f12052,plain,(
  ( ! [X14,X12,X13,X11] : (~sP12('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X11,X12) | p(s(bool,X11)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X13),X14),s(fun(sK16,bool),sK20),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f2554,f678])).
fof(f2554,plain,(
  ( ! [X171,X172] : (~sP12('const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)),X171,X172) | p(s(bool,X171))) ) | $spl9),
  inference(resolution,[],[f1139,f798])).
fof(f12051,plain,(
  ( ! [X10,X8,X7,X9] : (~sP12('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X7,X8) | p(s(bool,X7)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK16),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) ) | $spl9),
  inference(superposition,[],[f2554,f680])).
fof(f1366,plain,(
  ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | $spl59),
  inference(cnf_transformation,[],[f1366_D])).
fof(f1366_D,plain,(
  ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) <=> ~$spl59),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl59])])).
fof(f12013,plain,(
  ( ! [X8,X7] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X7),X8),s(fun(sK16,bool),sK20),s(fun(X7,bool),'const.pred_set.EMPTY_0'))))) ) | $spl3),
  inference(superposition,[],[f1132,f678])).
fof(f1132,plain,(
  ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) | $spl3),
  inference(cnf_transformation,[],[f1132_D])).
fof(f1132_D,plain,(
  ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f12012,plain,(
  ( ! [X6,X5] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X5,sK16),X6),s(fun(X5,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) ) | $spl3),
  inference(superposition,[],[f1132,f680])).
fof(f6581,plain,(
  ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)))) | $spl231),
  inference(cnf_transformation,[],[f6581_D])).
fof(f6581_D,plain,(
  ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)))) <=> ~$spl231),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl231])])).
fof(f12006,plain,(
  ( ! [X8,X7] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X7),X8),s(fun(sK17,bool),sK18),s(fun(X7,bool),'const.pred_set.EMPTY_0'))))) ) | $spl3),
  inference(superposition,[],[f1132,f678])).
fof(f12005,plain,(
  ( ! [X6,X5] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X5,sK17),X6),s(fun(X5,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl3),
  inference(superposition,[],[f1132,f680])).
fof(f12004,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X0),X1),s(fun(sK17,bool),sK18),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl3),
  inference(superposition,[],[f1132,f948])).
fof(f6567,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | s(bool,X4) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) ) | $spl228),
  inference(cnf_transformation,[],[f6567_D])).
fof(f6567_D,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | s(bool,X4) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) ) <=> ~$spl228),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl228])])).
fof(f11997,plain,(
  ( ! [X225,X226] : (p(s(bool,X225)) | ~p(s(bool,X226)) | s(bool,X226) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) ) | $spl3),
  inference(resolution,[],[f1132,f957])).
fof(f11895,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20))) | s(bool,X24) != s(bool,X26)) ) | $spl3),
  inference(resolution,[],[f1132,f706])).
fof(f11893,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) ) | $spl3),
  inference(resolution,[],[f1132,f705])).
fof(f11891,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) ) | $spl3),
  inference(resolution,[],[f1132,f704])).
fof(f11890,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20))) | s(bool,X9) != s(bool,X11)) ) | $spl3),
  inference(resolution,[],[f1132,f702])).
fof(f11888,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) ) | $spl3),
  inference(resolution,[],[f1132,f700])).
fof(f11881,plain,(
  ( ! [X8,X7] : (~sP12(X7,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X8) | p(s(bool,X7))) ) | $spl374),
  inference(cnf_transformation,[],[f11881_D])).
fof(f11881_D,plain,(
  ( ! [X8,X7] : (~sP12(X7,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X8) | p(s(bool,X7))) ) <=> ~$spl374),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl374])])).
fof(f11876,plain,(
  ( ! [X14,X12,X13,X11] : (~sP12(X11,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X12) | p(s(bool,X11)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X13),X14),s(fun(sK16,bool),sK20),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f2553,f678])).
fof(f2553,plain,(
  ( ! [X169,X170] : (~sP12(X169,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)),X170) | p(s(bool,X169))) ) | $spl9),
  inference(resolution,[],[f1139,f798])).
fof(f11875,plain,(
  ( ! [X10,X8,X7,X9] : (~sP12(X7,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X8) | p(s(bool,X7)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK16),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) ) | $spl9),
  inference(superposition,[],[f2553,f680])).
fof(f11866,plain,(
  ( ! [X8,X7] : (~sP12(X7,X8,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X7))) ) | $spl372),
  inference(cnf_transformation,[],[f11866_D])).
fof(f11866_D,plain,(
  ( ! [X8,X7] : (~sP12(X7,X8,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X7))) ) <=> ~$spl372),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl372])])).
fof(f11861,plain,(
  ( ! [X14,X12,X13,X11] : (~sP12(X11,X12,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X11)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X13),X14),s(fun(sK16,bool),sK20),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f2552,f678])).
fof(f2552,plain,(
  ( ! [X167,X168] : (~sP12(X167,X168,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) | p(s(bool,X167))) ) | $spl9),
  inference(resolution,[],[f1139,f797])).
fof(f11860,plain,(
  ( ! [X10,X8,X7,X9] : (~sP12(X7,X8,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X7)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK16),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) ) | $spl9),
  inference(superposition,[],[f2552,f680])).
fof(f11851,plain,(
  ( ! [X8,X7] : (~sP12('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X7,X8) | p(s(bool,X8))) ) | $spl370),
  inference(cnf_transformation,[],[f11851_D])).
fof(f11851_D,plain,(
  ( ! [X8,X7] : (~sP12('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X7,X8) | p(s(bool,X8))) ) <=> ~$spl370),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl370])])).
fof(f11846,plain,(
  ( ! [X14,X12,X13,X11] : (~sP12('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X11,X12) | p(s(bool,X12)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X13),X14),s(fun(sK16,bool),sK20),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f2551,f678])).
fof(f2551,plain,(
  ( ! [X165,X166] : (~sP12('const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)),X166,X165) | p(s(bool,X165))) ) | $spl9),
  inference(resolution,[],[f1139,f797])).
fof(f11845,plain,(
  ( ! [X10,X8,X7,X9] : (~sP12('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X7,X8) | p(s(bool,X8)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK16),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) ) | $spl9),
  inference(superposition,[],[f2551,f680])).
fof(f5695,plain,(
  ~p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | $spl197),
  inference(cnf_transformation,[],[f5695_D])).
fof(f5695_D,plain,(
  ~p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) <=> ~$spl197),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl197])])).
fof(f11807,plain,(
  ( ! [X8,X7] : (~p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X7),X8),s(fun(sK16,bool),sK20),s(fun(X7,bool),'const.pred_set.EMPTY_0'))))) ) | $spl1),
  inference(superposition,[],[f1131,f678])).
fof(f11806,plain,(
  ( ! [X6,X5] : (~p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X5,sK16),X6),s(fun(X5,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) ) | $spl1),
  inference(superposition,[],[f1131,f680])).
fof(f5693,plain,(
  ~p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)))) | $spl195),
  inference(cnf_transformation,[],[f5693_D])).
fof(f5693_D,plain,(
  ~p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)))) <=> ~$spl195),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl195])])).
fof(f11800,plain,(
  ( ! [X8,X7] : (~p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X7),X8),s(fun(sK17,bool),sK18),s(fun(X7,bool),'const.pred_set.EMPTY_0'))))) ) | $spl1),
  inference(superposition,[],[f1131,f678])).
fof(f11799,plain,(
  ( ! [X6,X5] : (~p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X5,sK17),X6),s(fun(X5,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl1),
  inference(superposition,[],[f1131,f680])).
fof(f11798,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X0),X1),s(fun(sK17,bool),sK18),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl1),
  inference(superposition,[],[f1131,f948])).
fof(f5679,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) ) | $spl192),
  inference(cnf_transformation,[],[f5679_D])).
fof(f5679_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) ) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f11791,plain,(
  ( ! [X222,X223] : (p(s(bool,X222)) | ~p(s(bool,X223)) | s(bool,X223) != s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) ) | $spl1),
  inference(resolution,[],[f1131,f957])).
fof(f11689,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20))) | s(bool,X21) != s(bool,X23)) ) | $spl1),
  inference(resolution,[],[f1131,f706])).
fof(f11687,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) ) | $spl1),
  inference(resolution,[],[f1131,f705])).
fof(f11685,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) ) | $spl1),
  inference(resolution,[],[f1131,f704])).
fof(f11684,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20))) | s(bool,X6) != s(bool,X8)) ) | $spl1),
  inference(resolution,[],[f1131,f702])).
fof(f11682,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) ) | $spl1),
  inference(resolution,[],[f1131,f700])).
fof(f11679,plain,(
  ( ! [X8,X7] : (sP11(X7,X8,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X8))) ) | $spl368),
  inference(cnf_transformation,[],[f11679_D])).
fof(f11679_D,plain,(
  ( ! [X8,X7] : (sP11(X7,X8,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X8))) ) <=> ~$spl368),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl368])])).
fof(f11674,plain,(
  ( ! [X14,X12,X13,X11] : (sP11(X11,X12,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X12)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X13),X14),s(fun(sK16,bool),sK20),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f2544,f678])).
fof(f2544,plain,(
  ( ! [X152,X151] : (sP11(X151,X152,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) | p(s(bool,X152))) ) | $spl9),
  inference(resolution,[],[f1139,f783])).
fof(f783,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f544])).
fof(f11673,plain,(
  ( ! [X10,X8,X7,X9] : (sP11(X7,X8,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X8)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK16),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) ) | $spl9),
  inference(superposition,[],[f2544,f680])).
fof(f11663,plain,(
  ( ! [X8,X7] : (sP11(X7,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X8) | p(s(bool,X8))) ) | $spl366),
  inference(cnf_transformation,[],[f11663_D])).
fof(f11663_D,plain,(
  ( ! [X8,X7] : (sP11(X7,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X8) | p(s(bool,X8))) ) <=> ~$spl366),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl366])])).
fof(f11658,plain,(
  ( ! [X14,X12,X13,X11] : (sP11(X11,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X12) | p(s(bool,X12)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X13),X14),s(fun(sK16,bool),sK20),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f2543,f678])).
fof(f2543,plain,(
  ( ! [X149,X150] : (sP11(X149,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)),X150) | p(s(bool,X150))) ) | $spl9),
  inference(resolution,[],[f1139,f783])).
fof(f11657,plain,(
  ( ! [X10,X8,X7,X9] : (sP11(X7,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X8) | p(s(bool,X8)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK16),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) ) | $spl9),
  inference(superposition,[],[f2543,f680])).
fof(f11647,plain,(
  ( ! [X8,X7] : (~sP11('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X7,X8) | ~p(s(bool,X7))) ) | $spl364),
  inference(cnf_transformation,[],[f11647_D])).
fof(f11647_D,plain,(
  ( ! [X8,X7] : (~sP11('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X7,X8) | ~p(s(bool,X7))) ) <=> ~$spl364),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl364])])).
fof(f11642,plain,(
  ( ! [X14,X12,X13,X11] : (~sP11('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X11,X12) | ~p(s(bool,X11)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X13),X14),s(fun(sK16,bool),sK20),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f2542,f678])).
fof(f11641,plain,(
  ( ! [X10,X8,X7,X9] : (~sP11('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X7,X8) | ~p(s(bool,X7)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK16),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) ) | $spl9),
  inference(superposition,[],[f2542,f680])).
fof(f5312,plain,(
  ( ! [X0,X1] : (sP13(X0,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X1)) ) | $spl182),
  inference(cnf_transformation,[],[f5312_D])).
fof(f5312_D,plain,(
  ( ! [X0,X1] : (sP13(X0,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X1)) ) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f11627,plain,(
  ( ! [X14,X12,X13,X11] : (sP13(X11,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X12) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X13),X14),s(fun(sK16,bool),sK20),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1067,f678])).
fof(f1067,plain,(
  ( ! [X198,X199] : (sP13(X198,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)),X199)) )),
  inference(resolution,[],[f647,f813])).
fof(f11626,plain,(
  ( ! [X10,X8,X7,X9] : (sP13(X7,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X8) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK16),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) )),
  inference(superposition,[],[f1067,f680])).
fof(f5311,plain,(
  ( ! [X0,X1] : (sP13(X0,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)),X1)) ) | $spl180),
  inference(cnf_transformation,[],[f5311_D])).
fof(f5311_D,plain,(
  ( ! [X0,X1] : (sP13(X0,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)),X1)) ) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f11620,plain,(
  ( ! [X14,X12,X13,X11] : (sP13(X11,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)),X12) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1067,f678])).
fof(f11619,plain,(
  ( ! [X10,X8,X7,X9] : (sP13(X7,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)),X8) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) )),
  inference(superposition,[],[f1067,f680])).
fof(f11618,plain,(
  ( ! [X2,X0,X3,X1] : (sP13(X0,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)),X1) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1067,f948])).
fof(f11607,plain,(
  ( ! [X8,X7] : (~sP11('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X7,X8) | ~p(s(bool,X8))) ) | $spl362),
  inference(cnf_transformation,[],[f11607_D])).
fof(f11607_D,plain,(
  ( ! [X8,X7] : (~sP11('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X7,X8) | ~p(s(bool,X8))) ) <=> ~$spl362),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl362])])).
fof(f11602,plain,(
  ( ! [X14,X12,X13,X11] : (~sP11('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X11,X12) | ~p(s(bool,X12)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X13),X14),s(fun(sK16,bool),sK20),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f2541,f678])).
fof(f11601,plain,(
  ( ! [X10,X8,X7,X9] : (~sP11('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X7,X8) | ~p(s(bool,X8)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK16),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) ) | $spl9),
  inference(superposition,[],[f2541,f680])).
fof(f11591,plain,(
  ( ! [X8,X7] : (sP6('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X7,X8) | p(s(bool,X7))) ) | $spl360),
  inference(cnf_transformation,[],[f11591_D])).
fof(f11591_D,plain,(
  ( ! [X8,X7] : (sP6('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X7,X8) | p(s(bool,X7))) ) <=> ~$spl360),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl360])])).
fof(f11584,plain,(
  ( ! [X14,X12,X13,X11] : (sP6('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X11,X12) | p(s(bool,X11)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X13),X14),s(fun(sK16,bool),sK20),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f2523,f678])).
fof(f2523,plain,(
  ( ! [X109,X110] : (sP6('const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)),X109,X110) | p(s(bool,X109))) ) | $spl9),
  inference(resolution,[],[f1139,f753])).
fof(f753,plain,(
  ( ! [X2,X0,X1] : (sP6(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f529])).
fof(f11583,plain,(
  ( ! [X10,X8,X7,X9] : (sP6('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X7,X8) | p(s(bool,X7)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK16),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) ) | $spl9),
  inference(superposition,[],[f2523,f680])).
fof(f11571,plain,(
  ( ! [X8,X7] : (sP6(X7,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X8) | p(s(bool,X7))) ) | $spl358),
  inference(cnf_transformation,[],[f11571_D])).
fof(f11571_D,plain,(
  ( ! [X8,X7] : (sP6(X7,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X8) | p(s(bool,X7))) ) <=> ~$spl358),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl358])])).
fof(f11565,plain,(
  ( ! [X14,X12,X13,X11] : (sP6(X11,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X12) | p(s(bool,X11)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X13),X14),s(fun(sK16,bool),sK20),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f2522,f678])).
fof(f2522,plain,(
  ( ! [X107,X108] : (sP6(X107,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)),X108) | p(s(bool,X107))) ) | $spl9),
  inference(resolution,[],[f1139,f753])).
fof(f11564,plain,(
  ( ! [X10,X8,X7,X9] : (sP6(X7,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X8) | p(s(bool,X7)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK16),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) ) | $spl9),
  inference(superposition,[],[f2522,f680])).
fof(f5272,plain,(
  ( ! [X0,X1] : (sP12('const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X0,X1)) ) | $spl176),
  inference(cnf_transformation,[],[f5272_D])).
fof(f5272_D,plain,(
  ( ! [X0,X1] : (sP12('const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X0,X1)) ) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f11549,plain,(
  ( ! [X14,X12,X13,X11] : (sP12('const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X11,X12) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X13),X14),s(fun(sK16,bool),sK20),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1060,f678])).
fof(f1060,plain,(
  ( ! [X185,X184] : (sP12('const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)),X184,X185)) )),
  inference(resolution,[],[f647,f804])).
fof(f11548,plain,(
  ( ! [X10,X8,X7,X9] : (sP12('const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X7,X8) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK16),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) )),
  inference(superposition,[],[f1060,f680])).
fof(f5271,plain,(
  ( ! [X0,X1] : (sP12('const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)),X0,X1)) ) | $spl174),
  inference(cnf_transformation,[],[f5271_D])).
fof(f5271_D,plain,(
  ( ! [X0,X1] : (sP12('const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)),X0,X1)) ) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f11542,plain,(
  ( ! [X14,X12,X13,X11] : (sP12('const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)),X11,X12) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1060,f678])).
fof(f11541,plain,(
  ( ! [X10,X8,X7,X9] : (sP12('const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)),X7,X8) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) )),
  inference(superposition,[],[f1060,f680])).
fof(f11540,plain,(
  ( ! [X2,X0,X3,X1] : (sP12('const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)),X0,X1) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1060,f948])).
fof(f11529,plain,(
  ( ! [X8,X7] : (sP6('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X7,X8) | p(s(bool,X8))) ) | $spl356),
  inference(cnf_transformation,[],[f11529_D])).
fof(f11529_D,plain,(
  ( ! [X8,X7] : (sP6('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X7,X8) | p(s(bool,X8))) ) <=> ~$spl356),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl356])])).
fof(f11522,plain,(
  ( ! [X14,X12,X13,X11] : (sP6('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X11,X12) | p(s(bool,X12)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X13),X14),s(fun(sK16,bool),sK20),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f2521,f678])).
fof(f2521,plain,(
  ( ! [X105,X106] : (sP6('const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)),X105,X106) | p(s(bool,X106))) ) | $spl9),
  inference(resolution,[],[f1139,f752])).
fof(f752,plain,(
  ( ! [X2,X0,X1] : (sP6(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f529])).
fof(f11521,plain,(
  ( ! [X10,X8,X7,X9] : (sP6('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X7,X8) | p(s(bool,X8)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK16),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) ) | $spl9),
  inference(superposition,[],[f2521,f680])).
fof(f11509,plain,(
  ( ! [X8,X7] : (sP6(X7,X8,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X7))) ) | $spl354),
  inference(cnf_transformation,[],[f11509_D])).
fof(f11509_D,plain,(
  ( ! [X8,X7] : (sP6(X7,X8,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X7))) ) <=> ~$spl354),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl354])])).
fof(f11503,plain,(
  ( ! [X14,X12,X13,X11] : (sP6(X11,X12,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X11)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X13),X14),s(fun(sK16,bool),sK20),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f2520,f678])).
fof(f2520,plain,(
  ( ! [X103,X104] : (sP6(X103,X104,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) | p(s(bool,X103))) ) | $spl9),
  inference(resolution,[],[f1139,f752])).
fof(f11502,plain,(
  ( ! [X10,X8,X7,X9] : (sP6(X7,X8,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X7)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK16),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) ) | $spl9),
  inference(superposition,[],[f2520,f680])).
fof(f11491,plain,(
  ( ! [X8,X7] : (~sP6(X7,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X8) | ~p(s(bool,X7))) ) | $spl352),
  inference(cnf_transformation,[],[f11491_D])).
fof(f11491_D,plain,(
  ( ! [X8,X7] : (~sP6(X7,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X8) | ~p(s(bool,X7))) ) <=> ~$spl352),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl352])])).
fof(f11484,plain,(
  ( ! [X14,X12,X13,X11] : (~sP6(X11,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X12) | ~p(s(bool,X11)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X13),X14),s(fun(sK16,bool),sK20),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f2514,f678])).
fof(f11483,plain,(
  ( ! [X10,X8,X7,X9] : (~sP6(X7,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X8) | ~p(s(bool,X7)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK16),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) ) | $spl9),
  inference(superposition,[],[f2514,f680])).
fof(f11471,plain,(
  ( ! [X8,X7] : (~sP6(X7,X8,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X7))) ) | $spl350),
  inference(cnf_transformation,[],[f11471_D])).
fof(f11471_D,plain,(
  ( ! [X8,X7] : (~sP6(X7,X8,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X7))) ) <=> ~$spl350),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl350])])).
fof(f11464,plain,(
  ( ! [X14,X12,X13,X11] : (~sP6(X11,X12,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X11)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X13),X14),s(fun(sK16,bool),sK20),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f2513,f678])).
fof(f11463,plain,(
  ( ! [X10,X8,X7,X9] : (~sP6(X7,X8,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X7)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK16),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) ) | $spl9),
  inference(superposition,[],[f2513,f680])).
fof(f5204,plain,(
  ( ! [X0,X1] : (sP11('const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X0,X1)) ) | $spl166),
  inference(cnf_transformation,[],[f5204_D])).
fof(f5204_D,plain,(
  ( ! [X0,X1] : (sP11('const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X0,X1)) ) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f11447,plain,(
  ( ! [X14,X12,X13,X11] : (sP11('const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X11,X12) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X13),X14),s(fun(sK16,bool),sK20),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1049,f678])).
fof(f1049,plain,(
  ( ! [X163,X162] : (sP11('const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)),X162,X163)) )),
  inference(resolution,[],[f647,f788])).
fof(f11446,plain,(
  ( ! [X10,X8,X7,X9] : (sP11('const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X7,X8) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK16),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) )),
  inference(superposition,[],[f1049,f680])).
fof(f5203,plain,(
  ( ! [X0,X1] : (sP11('const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)),X0,X1)) ) | $spl164),
  inference(cnf_transformation,[],[f5203_D])).
fof(f5203_D,plain,(
  ( ! [X0,X1] : (sP11('const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)),X0,X1)) ) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f11440,plain,(
  ( ! [X14,X12,X13,X11] : (sP11('const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)),X11,X12) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1049,f678])).
fof(f11439,plain,(
  ( ! [X10,X8,X7,X9] : (sP11('const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)),X7,X8) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) )),
  inference(superposition,[],[f1049,f680])).
fof(f11438,plain,(
  ( ! [X2,X0,X3,X1] : (sP11('const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)),X0,X1) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1049,f948])).
fof(f11429,plain,(
  ( ! [X8,X7] : (~sP4('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X7,X8) | ~p(s(bool,X7))) ) | $spl348),
  inference(cnf_transformation,[],[f11429_D])).
fof(f11429_D,plain,(
  ( ! [X8,X7] : (~sP4('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X7,X8) | ~p(s(bool,X7))) ) <=> ~$spl348),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl348])])).
fof(f11421,plain,(
  ( ! [X14,X12,X13,X11] : (~sP4('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X11,X12) | ~p(s(bool,X11)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X13),X14),s(fun(sK16,bool),sK20),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f2498,f678])).
fof(f2498,plain,(
  ( ! [X59,X60] : (~sP4('const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)),X59,X60) | ~p(s(bool,X59))) ) | $spl9),
  inference(resolution,[],[f1139,f725])).
fof(f725,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | ~sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f517])).
fof(f11420,plain,(
  ( ! [X10,X8,X7,X9] : (~sP4('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X7,X8) | ~p(s(bool,X7)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK16),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) ) | $spl9),
  inference(superposition,[],[f2498,f680])).
fof(f11411,plain,(
  ( ! [X8,X7] : (~sP4('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X7,X8) | ~p(s(bool,X8))) ) | $spl346),
  inference(cnf_transformation,[],[f11411_D])).
fof(f11411_D,plain,(
  ( ! [X8,X7] : (~sP4('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X7,X8) | ~p(s(bool,X8))) ) <=> ~$spl346),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl346])])).
fof(f11403,plain,(
  ( ! [X14,X12,X13,X11] : (~sP4('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X11,X12) | ~p(s(bool,X12)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X13),X14),s(fun(sK16,bool),sK20),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f2497,f678])).
fof(f2497,plain,(
  ( ! [X57,X58] : (~sP4('const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)),X58,X57) | ~p(s(bool,X57))) ) | $spl9),
  inference(resolution,[],[f1139,f724])).
fof(f724,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X0)) | ~sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f517])).
fof(f11402,plain,(
  ( ! [X10,X8,X7,X9] : (~sP4('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X7,X8) | ~p(s(bool,X8)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK16),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) ) | $spl9),
  inference(superposition,[],[f2497,f680])).
fof(f5180,plain,(
  ( ! [X0,X1] : (sP7('const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X0,X1)) ) | $spl162),
  inference(cnf_transformation,[],[f5180_D])).
fof(f5180_D,plain,(
  ( ! [X0,X1] : (sP7('const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X0,X1)) ) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f11386,plain,(
  ( ! [X14,X12,X13,X11] : (sP7('const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X11,X12) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X13),X14),s(fun(sK16,bool),sK20),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1042,f678])).
fof(f1042,plain,(
  ( ! [X149,X148] : (sP7('const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)),X148,X149)) )),
  inference(resolution,[],[f647,f776])).
fof(f11385,plain,(
  ( ! [X10,X8,X7,X9] : (sP7('const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X7,X8) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK16),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) )),
  inference(superposition,[],[f1042,f680])).
fof(f5179,plain,(
  ( ! [X0,X1] : (sP7('const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)),X0,X1)) ) | $spl160),
  inference(cnf_transformation,[],[f5179_D])).
fof(f5179_D,plain,(
  ( ! [X0,X1] : (sP7('const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)),X0,X1)) ) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f11379,plain,(
  ( ! [X14,X12,X13,X11] : (sP7('const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)),X11,X12) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1042,f678])).
fof(f11378,plain,(
  ( ! [X10,X8,X7,X9] : (sP7('const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)),X7,X8) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) )),
  inference(superposition,[],[f1042,f680])).
fof(f11377,plain,(
  ( ! [X2,X0,X3,X1] : (sP7('const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)),X0,X1) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1042,f948])).
fof(f7471,plain,(
  ( ! [X0,X1] : (~sP13('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0,X1) | p(s(bool,X0))) ) | $spl258),
  inference(cnf_transformation,[],[f7471_D])).
fof(f7471_D,plain,(
  ( ! [X0,X1] : (~sP13('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0,X1) | p(s(bool,X0))) ) <=> ~$spl258),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl258])])).
fof(f11364,plain,(
  ( ! [X14,X12,X13,X11] : (~sP13('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X11,X12) | p(s(bool,X11)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2089,f678])).
fof(f2089,plain,(
  ( ! [X187,X188] : (~sP13('const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)),X187,X188) | p(s(bool,X187))) ) | $spl5),
  inference(resolution,[],[f1135,f808])).
fof(f11363,plain,(
  ( ! [X10,X8,X7,X9] : (~sP13('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X7,X8) | p(s(bool,X7)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f2089,f680])).
fof(f11362,plain,(
  ( ! [X2,X0,X3,X1] : (~sP13('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0,X1) | p(s(bool,X0)) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2089,f948])).
fof(f7456,plain,(
  ( ! [X0,X1] : (~sP13(X0,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X1) | p(s(bool,X0))) ) | $spl256),
  inference(cnf_transformation,[],[f7456_D])).
fof(f7456_D,plain,(
  ( ! [X0,X1] : (~sP13(X0,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X1) | p(s(bool,X0))) ) <=> ~$spl256),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl256])])).
fof(f11351,plain,(
  ( ! [X14,X12,X13,X11] : (~sP13(X11,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X12) | p(s(bool,X11)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2088,f678])).
fof(f2088,plain,(
  ( ! [X185,X186] : (~sP13(X185,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)),X186) | p(s(bool,X185))) ) | $spl5),
  inference(resolution,[],[f1135,f808])).
fof(f11350,plain,(
  ( ! [X10,X8,X7,X9] : (~sP13(X7,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X8) | p(s(bool,X7)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f2088,f680])).
fof(f11349,plain,(
  ( ! [X2,X0,X3,X1] : (~sP13(X0,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X1) | p(s(bool,X0)) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2088,f948])).
fof(f7442,plain,(
  ( ! [X0,X1] : (~sP13(X0,X1,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X1))) ) | $spl254),
  inference(cnf_transformation,[],[f7442_D])).
fof(f7442_D,plain,(
  ( ! [X0,X1] : (~sP13(X0,X1,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X1))) ) <=> ~$spl254),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl254])])).
fof(f11338,plain,(
  ( ! [X14,X12,X13,X11] : (~sP13(X11,X12,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X12)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2087,f678])).
fof(f2087,plain,(
  ( ! [X184,X183] : (~sP13(X184,X183,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | p(s(bool,X183))) ) | $spl5),
  inference(resolution,[],[f1135,f807])).
fof(f11337,plain,(
  ( ! [X10,X8,X7,X9] : (~sP13(X7,X8,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X8)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f2087,f680])).
fof(f11336,plain,(
  ( ! [X2,X0,X3,X1] : (~sP13(X0,X1,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X1)) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2087,f948])).
fof(f7427,plain,(
  ( ! [X0,X1] : (~sP13(X0,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X1) | p(s(bool,X1))) ) | $spl252),
  inference(cnf_transformation,[],[f7427_D])).
fof(f7427_D,plain,(
  ( ! [X0,X1] : (~sP13(X0,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X1) | p(s(bool,X1))) ) <=> ~$spl252),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl252])])).
fof(f11325,plain,(
  ( ! [X14,X12,X13,X11] : (~sP13(X11,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X12) | p(s(bool,X12)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2086,f678])).
fof(f2086,plain,(
  ( ! [X182,X181] : (~sP13(X182,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)),X181) | p(s(bool,X181))) ) | $spl5),
  inference(resolution,[],[f1135,f807])).
fof(f11324,plain,(
  ( ! [X10,X8,X7,X9] : (~sP13(X7,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X8) | p(s(bool,X8)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f2086,f680])).
fof(f11323,plain,(
  ( ! [X2,X0,X3,X1] : (~sP13(X0,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X1) | p(s(bool,X1)) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2086,f948])).
fof(f5102,plain,(
  ( ! [X0,X1] : (sP7(X0,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X1)) ) | $spl158),
  inference(cnf_transformation,[],[f5102_D])).
fof(f5102_D,plain,(
  ( ! [X0,X1] : (sP7(X0,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X1)) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f11310,plain,(
  ( ! [X14,X12,X13,X11] : (sP7(X11,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X12) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X13),X14),s(fun(sK16,bool),sK20),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1041,f678])).
fof(f1041,plain,(
  ( ! [X146,X147] : (sP7(X146,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)),X147)) )),
  inference(resolution,[],[f647,f775])).
fof(f11309,plain,(
  ( ! [X10,X8,X7,X9] : (sP7(X7,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X8) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK16),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) )),
  inference(superposition,[],[f1041,f680])).
fof(f5101,plain,(
  ( ! [X0,X1] : (sP7(X0,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)),X1)) ) | $spl156),
  inference(cnf_transformation,[],[f5101_D])).
fof(f5101_D,plain,(
  ( ! [X0,X1] : (sP7(X0,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)),X1)) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f11303,plain,(
  ( ! [X14,X12,X13,X11] : (sP7(X11,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)),X12) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1041,f678])).
fof(f11302,plain,(
  ( ! [X10,X8,X7,X9] : (sP7(X7,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)),X8) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) )),
  inference(superposition,[],[f1041,f680])).
fof(f11301,plain,(
  ( ! [X2,X0,X3,X1] : (sP7(X0,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)),X1) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1041,f948])).
fof(f7251,plain,(
  ( ! [X0,X1] : (~sP12('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0,X1) | p(s(bool,X0))) ) | $spl250),
  inference(cnf_transformation,[],[f7251_D])).
fof(f7251_D,plain,(
  ( ! [X0,X1] : (~sP12('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0,X1) | p(s(bool,X0))) ) <=> ~$spl250),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl250])])).
fof(f11288,plain,(
  ( ! [X14,X12,X13,X11] : (~sP12('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X11,X12) | p(s(bool,X11)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2081,f678])).
fof(f2081,plain,(
  ( ! [X171,X172] : (~sP12('const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)),X171,X172) | p(s(bool,X171))) ) | $spl5),
  inference(resolution,[],[f1135,f798])).
fof(f11287,plain,(
  ( ! [X10,X8,X7,X9] : (~sP12('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X7,X8) | p(s(bool,X7)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f2081,f680])).
fof(f11286,plain,(
  ( ! [X2,X0,X3,X1] : (~sP12('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0,X1) | p(s(bool,X0)) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2081,f948])).
fof(f7237,plain,(
  ( ! [X0,X1] : (~sP12(X0,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X1) | p(s(bool,X0))) ) | $spl248),
  inference(cnf_transformation,[],[f7237_D])).
fof(f7237_D,plain,(
  ( ! [X0,X1] : (~sP12(X0,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X1) | p(s(bool,X0))) ) <=> ~$spl248),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl248])])).
fof(f11275,plain,(
  ( ! [X14,X12,X13,X11] : (~sP12(X11,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X12) | p(s(bool,X11)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2080,f678])).
fof(f2080,plain,(
  ( ! [X169,X170] : (~sP12(X169,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)),X170) | p(s(bool,X169))) ) | $spl5),
  inference(resolution,[],[f1135,f798])).
fof(f11274,plain,(
  ( ! [X10,X8,X7,X9] : (~sP12(X7,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X8) | p(s(bool,X7)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f2080,f680])).
fof(f11273,plain,(
  ( ! [X2,X0,X3,X1] : (~sP12(X0,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X1) | p(s(bool,X0)) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2080,f948])).
fof(f7064,plain,(
  ( ! [X0,X1] : (~sP12(X0,X1,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X0))) ) | $spl246),
  inference(cnf_transformation,[],[f7064_D])).
fof(f7064_D,plain,(
  ( ! [X0,X1] : (~sP12(X0,X1,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X0))) ) <=> ~$spl246),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl246])])).
fof(f11262,plain,(
  ( ! [X14,X12,X13,X11] : (~sP12(X11,X12,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X11)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2079,f678])).
fof(f2079,plain,(
  ( ! [X167,X168] : (~sP12(X167,X168,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | p(s(bool,X167))) ) | $spl5),
  inference(resolution,[],[f1135,f797])).
fof(f11261,plain,(
  ( ! [X10,X8,X7,X9] : (~sP12(X7,X8,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X7)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f2079,f680])).
fof(f11260,plain,(
  ( ! [X2,X0,X3,X1] : (~sP12(X0,X1,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X0)) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2079,f948])).
fof(f7049,plain,(
  ( ! [X0,X1] : (~sP12('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0,X1) | p(s(bool,X1))) ) | $spl244),
  inference(cnf_transformation,[],[f7049_D])).
fof(f7049_D,plain,(
  ( ! [X0,X1] : (~sP12('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0,X1) | p(s(bool,X1))) ) <=> ~$spl244),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl244])])).
fof(f11249,plain,(
  ( ! [X14,X12,X13,X11] : (~sP12('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X11,X12) | p(s(bool,X12)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2078,f678])).
fof(f2078,plain,(
  ( ! [X165,X166] : (~sP12('const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)),X166,X165) | p(s(bool,X165))) ) | $spl5),
  inference(resolution,[],[f1135,f797])).
fof(f11248,plain,(
  ( ! [X10,X8,X7,X9] : (~sP12('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X7,X8) | p(s(bool,X8)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f2078,f680])).
fof(f11247,plain,(
  ( ! [X2,X0,X3,X1] : (~sP12('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0,X1) | p(s(bool,X1)) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2078,f948])).
fof(f6839,plain,(
  ( ! [X0,X1] : (sP11(X0,X1,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X1))) ) | $spl240),
  inference(cnf_transformation,[],[f6839_D])).
fof(f6839_D,plain,(
  ( ! [X0,X1] : (sP11(X0,X1,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X1))) ) <=> ~$spl240),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl240])])).
fof(f11236,plain,(
  ( ! [X14,X12,X13,X11] : (sP11(X11,X12,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X12)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2071,f678])).
fof(f2071,plain,(
  ( ! [X152,X151] : (sP11(X151,X152,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | p(s(bool,X152))) ) | $spl5),
  inference(resolution,[],[f1135,f783])).
fof(f11235,plain,(
  ( ! [X10,X8,X7,X9] : (sP11(X7,X8,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X8)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f2071,f680])).
fof(f11234,plain,(
  ( ! [X2,X0,X3,X1] : (sP11(X0,X1,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X1)) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2071,f948])).
fof(f4995,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) | $spl154),
  inference(cnf_transformation,[],[f4995_D])).
fof(f4995_D,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f11221,plain,(
  ( ! [X14,X12,X13,X11] : (sP7(X11,X12,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X13),X14),s(fun(sK16,bool),sK20),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1040,f678])).
fof(f1040,plain,(
  ( ! [X144,X145] : (sP7(X144,X145,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) )),
  inference(resolution,[],[f647,f774])).
fof(f11220,plain,(
  ( ! [X10,X8,X7,X9] : (sP7(X7,X8,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK16),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) )),
  inference(superposition,[],[f1040,f680])).
fof(f4994,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)))) ) | $spl152),
  inference(cnf_transformation,[],[f4994_D])).
fof(f4994_D,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)))) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f11214,plain,(
  ( ! [X14,X12,X13,X11] : (sP7(X11,X12,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1040,f678])).
fof(f11213,plain,(
  ( ! [X10,X8,X7,X9] : (sP7(X7,X8,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) )),
  inference(superposition,[],[f1040,f680])).
fof(f11212,plain,(
  ( ! [X2,X0,X3,X1] : (sP7(X0,X1,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1040,f948])).
fof(f6825,plain,(
  ( ! [X0,X1] : (sP11(X0,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X1) | p(s(bool,X1))) ) | $spl238),
  inference(cnf_transformation,[],[f6825_D])).
fof(f6825_D,plain,(
  ( ! [X0,X1] : (sP11(X0,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X1) | p(s(bool,X1))) ) <=> ~$spl238),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl238])])).
fof(f11199,plain,(
  ( ! [X14,X12,X13,X11] : (sP11(X11,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X12) | p(s(bool,X12)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2070,f678])).
fof(f2070,plain,(
  ( ! [X149,X150] : (sP11(X149,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)),X150) | p(s(bool,X150))) ) | $spl5),
  inference(resolution,[],[f1135,f783])).
fof(f11198,plain,(
  ( ! [X10,X8,X7,X9] : (sP11(X7,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X8) | p(s(bool,X8)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f2070,f680])).
fof(f11197,plain,(
  ( ! [X2,X0,X3,X1] : (sP11(X0,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X1) | p(s(bool,X1)) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2070,f948])).
fof(f6811,plain,(
  ( ! [X0,X1] : (~sP11('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0,X1) | ~p(s(bool,X0))) ) | $spl236),
  inference(cnf_transformation,[],[f6811_D])).
fof(f6811_D,plain,(
  ( ! [X0,X1] : (~sP11('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0,X1) | ~p(s(bool,X0))) ) <=> ~$spl236),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl236])])).
fof(f11186,plain,(
  ( ! [X14,X12,X13,X11] : (~sP11('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X11,X12) | ~p(s(bool,X11)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2069,f678])).
fof(f2069,plain,(
  ( ! [X147,X148] : (~sP11('const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)),X147,X148) | ~p(s(bool,X147))) ) | $spl5),
  inference(resolution,[],[f1135,f782])).
fof(f11185,plain,(
  ( ! [X10,X8,X7,X9] : (~sP11('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X7,X8) | ~p(s(bool,X7)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f2069,f680])).
fof(f11184,plain,(
  ( ! [X2,X0,X3,X1] : (~sP11('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2069,f948])).
fof(f6618,plain,(
  ( ! [X0,X1] : (~sP11('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0,X1) | ~p(s(bool,X1))) ) | $spl234),
  inference(cnf_transformation,[],[f6618_D])).
fof(f6618_D,plain,(
  ( ! [X0,X1] : (~sP11('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0,X1) | ~p(s(bool,X1))) ) <=> ~$spl234),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl234])])).
fof(f11173,plain,(
  ( ! [X14,X12,X13,X11] : (~sP11('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X11,X12) | ~p(s(bool,X12)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2068,f678])).
fof(f2068,plain,(
  ( ! [X146,X145] : (~sP11('const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)),X146,X145) | ~p(s(bool,X145))) ) | $spl5),
  inference(resolution,[],[f1135,f781])).
fof(f11172,plain,(
  ( ! [X10,X8,X7,X9] : (~sP11('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X7,X8) | ~p(s(bool,X8)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f2068,f680])).
fof(f11171,plain,(
  ( ! [X2,X0,X3,X1] : (~sP11('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2068,f948])).
fof(f6604,plain,(
  ( ! [X0,X1] : (sP6('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0,X1) | p(s(bool,X0))) ) | $spl232),
  inference(cnf_transformation,[],[f6604_D])).
fof(f6604_D,plain,(
  ( ! [X0,X1] : (sP6('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0,X1) | p(s(bool,X0))) ) <=> ~$spl232),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl232])])).
fof(f11158,plain,(
  ( ! [X14,X12,X13,X11] : (sP6('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X11,X12) | p(s(bool,X11)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2050,f678])).
fof(f2050,plain,(
  ( ! [X109,X110] : (sP6('const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)),X109,X110) | p(s(bool,X109))) ) | $spl5),
  inference(resolution,[],[f1135,f753])).
fof(f11157,plain,(
  ( ! [X10,X8,X7,X9] : (sP6('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X7,X8) | p(s(bool,X7)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f2050,f680])).
fof(f11156,plain,(
  ( ! [X2,X0,X3,X1] : (sP6('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0,X1) | p(s(bool,X0)) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2050,f948])).
fof(f6410,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X1) | p(s(bool,X0))) ) | $spl220),
  inference(cnf_transformation,[],[f6410_D])).
fof(f6410_D,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X1) | p(s(bool,X0))) ) <=> ~$spl220),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl220])])).
fof(f11142,plain,(
  ( ! [X14,X12,X13,X11] : (sP6(X11,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X12) | p(s(bool,X11)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2049,f678])).
fof(f2049,plain,(
  ( ! [X107,X108] : (sP6(X107,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)),X108) | p(s(bool,X107))) ) | $spl5),
  inference(resolution,[],[f1135,f753])).
fof(f11141,plain,(
  ( ! [X10,X8,X7,X9] : (sP6(X7,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X8) | p(s(bool,X7)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f2049,f680])).
fof(f11140,plain,(
  ( ! [X2,X0,X3,X1] : (sP6(X0,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X1) | p(s(bool,X0)) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2049,f948])).
fof(f4953,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) | $spl150),
  inference(cnf_transformation,[],[f4953_D])).
fof(f4953_D,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f11126,plain,(
  ( ! [X14,X12,X13,X11] : (sP8(X11,X12,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X13),X14),s(fun(sK16,bool),sK20),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1039,f678])).
fof(f1039,plain,(
  ( ! [X142,X143] : (sP8(X142,X143,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) )),
  inference(resolution,[],[f647,f770])).
fof(f11125,plain,(
  ( ! [X10,X8,X7,X9] : (sP8(X7,X8,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK16),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) )),
  inference(superposition,[],[f1039,f680])).
fof(f4952,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)))) ) | $spl148),
  inference(cnf_transformation,[],[f4952_D])).
fof(f4952_D,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)))) ) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f11119,plain,(
  ( ! [X14,X12,X13,X11] : (sP8(X11,X12,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1039,f678])).
fof(f11118,plain,(
  ( ! [X10,X8,X7,X9] : (sP8(X7,X8,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) )),
  inference(superposition,[],[f1039,f680])).
fof(f11117,plain,(
  ( ! [X2,X0,X3,X1] : (sP8(X0,X1,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1039,f948])).
fof(f6393,plain,(
  ( ! [X0,X1] : (sP6('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0,X1) | p(s(bool,X1))) ) | $spl218),
  inference(cnf_transformation,[],[f6393_D])).
fof(f6393_D,plain,(
  ( ! [X0,X1] : (sP6('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0,X1) | p(s(bool,X1))) ) <=> ~$spl218),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl218])])).
fof(f11102,plain,(
  ( ! [X14,X12,X13,X11] : (sP6('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X11,X12) | p(s(bool,X12)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2048,f678])).
fof(f2048,plain,(
  ( ! [X105,X106] : (sP6('const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)),X105,X106) | p(s(bool,X106))) ) | $spl5),
  inference(resolution,[],[f1135,f752])).
fof(f11101,plain,(
  ( ! [X10,X8,X7,X9] : (sP6('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X7,X8) | p(s(bool,X8)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f2048,f680])).
fof(f11100,plain,(
  ( ! [X2,X0,X3,X1] : (sP6('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0,X1) | p(s(bool,X1)) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2048,f948])).
fof(f6375,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X0))) ) | $spl216),
  inference(cnf_transformation,[],[f6375_D])).
fof(f6375_D,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X0))) ) <=> ~$spl216),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl216])])).
fof(f11086,plain,(
  ( ! [X14,X12,X13,X11] : (sP6(X11,X12,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X11)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2047,f678])).
fof(f2047,plain,(
  ( ! [X103,X104] : (sP6(X103,X104,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | p(s(bool,X103))) ) | $spl5),
  inference(resolution,[],[f1135,f752])).
fof(f11085,plain,(
  ( ! [X10,X8,X7,X9] : (sP6(X7,X8,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X7)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f2047,f680])).
fof(f11084,plain,(
  ( ! [X2,X0,X3,X1] : (sP6(X0,X1,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X0)) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2047,f948])).
fof(f6358,plain,(
  ( ! [X0,X1] : (~sP6(X0,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X1) | ~p(s(bool,X0))) ) | $spl214),
  inference(cnf_transformation,[],[f6358_D])).
fof(f6358_D,plain,(
  ( ! [X0,X1] : (~sP6(X0,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X1) | ~p(s(bool,X0))) ) <=> ~$spl214),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl214])])).
fof(f11069,plain,(
  ( ! [X14,X12,X13,X11] : (~sP6(X11,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X12) | ~p(s(bool,X11)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2041,f678])).
fof(f2041,plain,(
  ( ! [X92,X91] : (~sP6(X91,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)),X92) | ~p(s(bool,X91))) ) | $spl5),
  inference(resolution,[],[f1135,f743])).
fof(f11068,plain,(
  ( ! [X10,X8,X7,X9] : (~sP6(X7,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X8) | ~p(s(bool,X7)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f2041,f680])).
fof(f11067,plain,(
  ( ! [X2,X0,X3,X1] : (~sP6(X0,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X1) | ~p(s(bool,X0)) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2041,f948])).
fof(f6341,plain,(
  ( ! [X0,X1] : (~sP6(X0,X1,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X0))) ) | $spl212),
  inference(cnf_transformation,[],[f6341_D])).
fof(f6341_D,plain,(
  ( ! [X0,X1] : (~sP6(X0,X1,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X0))) ) <=> ~$spl212),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl212])])).
fof(f11053,plain,(
  ( ! [X14,X12,X13,X11] : (~sP6(X11,X12,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X11)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2040,f678])).
fof(f2040,plain,(
  ( ! [X90,X89] : (~sP6(X89,X90,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | ~p(s(bool,X89))) ) | $spl5),
  inference(resolution,[],[f1135,f742])).
fof(f11052,plain,(
  ( ! [X10,X8,X7,X9] : (~sP6(X7,X8,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X7)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f2040,f680])).
fof(f11051,plain,(
  ( ! [X2,X0,X3,X1] : (~sP6(X0,X1,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X0)) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2040,f948])).
fof(f6137,plain,(
  ( ! [X0,X1] : (~sP4('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0,X1) | ~p(s(bool,X0))) ) | $spl210),
  inference(cnf_transformation,[],[f6137_D])).
fof(f6137_D,plain,(
  ( ! [X0,X1] : (~sP4('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0,X1) | ~p(s(bool,X0))) ) <=> ~$spl210),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl210])])).
fof(f11036,plain,(
  ( ! [X14,X12,X13,X11] : (~sP4('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X11,X12) | ~p(s(bool,X11)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2025,f678])).
fof(f2025,plain,(
  ( ! [X59,X60] : (~sP4('const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)),X59,X60) | ~p(s(bool,X59))) ) | $spl5),
  inference(resolution,[],[f1135,f725])).
fof(f11035,plain,(
  ( ! [X10,X8,X7,X9] : (~sP4('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X7,X8) | ~p(s(bool,X7)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f2025,f680])).
fof(f11034,plain,(
  ( ! [X2,X0,X3,X1] : (~sP4('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2025,f948])).
fof(f4928,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) | $spl146),
  inference(cnf_transformation,[],[f4928_D])).
fof(f4928_D,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f11021,plain,(
  ( ! [X14,X12,X13,X11] : (sP9(X11,X12,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X13),X14),s(fun(sK16,bool),sK20),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1036,f678])).
fof(f1036,plain,(
  ( ! [X136,X137] : (sP9(X136,X137,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) )),
  inference(resolution,[],[f647,f766])).
fof(f11020,plain,(
  ( ! [X10,X8,X7,X9] : (sP9(X7,X8,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK16),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) )),
  inference(superposition,[],[f1036,f680])).
fof(f4927,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)))) ) | $spl144),
  inference(cnf_transformation,[],[f4927_D])).
fof(f4927_D,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)))) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f11014,plain,(
  ( ! [X14,X12,X13,X11] : (sP9(X11,X12,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1036,f678])).
fof(f11013,plain,(
  ( ! [X10,X8,X7,X9] : (sP9(X7,X8,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) )),
  inference(superposition,[],[f1036,f680])).
fof(f11012,plain,(
  ( ! [X2,X0,X3,X1] : (sP9(X0,X1,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1036,f948])).
fof(f6119,plain,(
  ( ! [X0,X1] : (~sP4('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0,X1) | ~p(s(bool,X1))) ) | $spl208),
  inference(cnf_transformation,[],[f6119_D])).
fof(f6119_D,plain,(
  ( ! [X0,X1] : (~sP4('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0,X1) | ~p(s(bool,X1))) ) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f10996,plain,(
  ( ! [X14,X12,X13,X11] : (~sP4('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X11,X12) | ~p(s(bool,X12)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2024,f678])).
fof(f2024,plain,(
  ( ! [X57,X58] : (~sP4('const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)),X58,X57) | ~p(s(bool,X57))) ) | $spl5),
  inference(resolution,[],[f1135,f724])).
fof(f10995,plain,(
  ( ! [X10,X8,X7,X9] : (~sP4('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X7,X8) | ~p(s(bool,X8)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f2024,f680])).
fof(f10994,plain,(
  ( ! [X2,X0,X3,X1] : (~sP4('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2024,f948])).
fof(f5943,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X0))) ) | $spl206),
  inference(cnf_transformation,[],[f5943_D])).
fof(f5943_D,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X0))) ) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f10981,plain,(
  ( ! [X14,X12,X13,X11] : (sP2(X11,X12,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X11)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2019,f678])).
fof(f2019,plain,(
  ( ! [X47,X48] : (sP2(X47,X48,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | ~p(s(bool,X47))) ) | $spl5),
  inference(resolution,[],[f1135,f719])).
fof(f10980,plain,(
  ( ! [X10,X8,X7,X9] : (sP2(X7,X8,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X7)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f2019,f680])).
fof(f10979,plain,(
  ( ! [X2,X0,X3,X1] : (sP2(X0,X1,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X0)) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2019,f948])).
fof(f5925,plain,(
  ( ! [X0,X1] : (~sP2('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0,X1) | ~p(s(bool,X0))) ) | $spl204),
  inference(cnf_transformation,[],[f5925_D])).
fof(f5925_D,plain,(
  ( ! [X0,X1] : (~sP2('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0,X1) | ~p(s(bool,X0))) ) <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f10960,plain,(
  ( ! [X14,X12,X13,X11] : (~sP2('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X11,X12) | ~p(s(bool,X11)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2011,f678])).
fof(f2011,plain,(
  ( ! [X31,X32] : (~sP2('const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)),X31,X32) | ~p(s(bool,X31))) ) | $spl5),
  inference(resolution,[],[f1135,f708])).
fof(f708,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | ~sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f509])).
fof(f10959,plain,(
  ( ! [X10,X8,X7,X9] : (~sP2('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X7,X8) | ~p(s(bool,X7)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f2011,f680])).
fof(f10958,plain,(
  ( ! [X2,X0,X3,X1] : (~sP2('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2011,f948])).
fof(f5907,plain,(
  ( ! [X0,X1] : (~sP2(X0,X1,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X0))) ) | $spl202),
  inference(cnf_transformation,[],[f5907_D])).
fof(f5907_D,plain,(
  ( ! [X0,X1] : (~sP2(X0,X1,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X0))) ) <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f10944,plain,(
  ( ! [X14,X12,X13,X11] : (~sP2(X11,X12,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X11)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2010,f678])).
fof(f2010,plain,(
  ( ! [X30,X29] : (~sP2(X29,X30,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | p(s(bool,X29))) ) | $spl5),
  inference(resolution,[],[f1135,f707])).
fof(f707,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | p(s(bool,X0)) | ~sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f509])).
fof(f10943,plain,(
  ( ! [X10,X8,X7,X9] : (~sP2(X7,X8,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X7)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f2010,f680])).
fof(f10942,plain,(
  ( ! [X2,X0,X3,X1] : (~sP2(X0,X1,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X0)) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2010,f948])).
fof(f5711,plain,(
  ( ! [X0,X1] : (~sP2('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0,X1) | p(s(bool,X1))) ) | $spl198),
  inference(cnf_transformation,[],[f5711_D])).
fof(f5711_D,plain,(
  ( ! [X0,X1] : (~sP2('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0,X1) | p(s(bool,X1))) ) <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f10929,plain,(
  ( ! [X14,X12,X13,X11] : (~sP2('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X11,X12) | p(s(bool,X12)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2009,f678])).
fof(f2009,plain,(
  ( ! [X28,X27] : (~sP2('const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)),X28,X27) | p(s(bool,X27))) ) | $spl5),
  inference(resolution,[],[f1135,f707])).
fof(f10928,plain,(
  ( ! [X10,X8,X7,X9] : (~sP2('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X7,X8) | p(s(bool,X8)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f2009,f680])).
fof(f10927,plain,(
  ( ! [X2,X0,X3,X1] : (~sP2('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0,X1) | p(s(bool,X1)) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2009,f948])).
fof(f4820,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) | $spl142),
  inference(cnf_transformation,[],[f4820_D])).
fof(f4820_D,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f10914,plain,(
  ( ! [X14,X12,X13,X11] : (sP5(X11,X12,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X13),X14),s(fun(sK16,bool),sK20),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1023,f678])).
fof(f1023,plain,(
  ( ! [X111,X110] : (sP5(X110,X111,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) )),
  inference(resolution,[],[f647,f749])).
fof(f10913,plain,(
  ( ! [X10,X8,X7,X9] : (sP5(X7,X8,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK16),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) )),
  inference(superposition,[],[f1023,f680])).
fof(f4819,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)))) ) | $spl140),
  inference(cnf_transformation,[],[f4819_D])).
fof(f4819_D,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)))) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f10907,plain,(
  ( ! [X14,X12,X13,X11] : (sP5(X11,X12,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1023,f678])).
fof(f10906,plain,(
  ( ! [X10,X8,X7,X9] : (sP5(X7,X8,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) )),
  inference(superposition,[],[f1023,f680])).
fof(f10905,plain,(
  ( ! [X2,X0,X3,X1] : (sP5(X0,X1,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1023,f948])).
fof(f10871,plain,(
  ( ! [X103,X105] : (p(s(bool,X103)) | s(bool,X103) != s(bool,'const.bool.IN_2'(s(sK17,X105),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | $spl344),
  inference(cnf_transformation,[],[f10871_D])).
fof(f10871_D,plain,(
  ( ! [X103,X105] : (p(s(bool,X103)) | s(bool,X103) != s(bool,'const.bool.IN_2'(s(sK17,X105),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) <=> ~$spl344),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl344])])).
fof(f10826,plain,(
  ( ! [X453,X455,X454] : (p(s(bool,X453)) | p(s(bool,X454)) | s(bool,X454) != s(bool,'const.bool.IN_2'(s(sK17,X455),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | $spl286),
  inference(resolution,[],[f8902,f957])).
fof(f10719,plain,(
  ( ! [X132,X130,X131,X129] : (p(s(bool,X129)) | ~p(s(bool,X130)) | s(bool,X129) != s(bool,'const.bool.IN_2'(s(sK17,X131),s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | s(bool,X130) != s(bool,X132)) ) | $spl286),
  inference(resolution,[],[f8902,f706])).
fof(f10716,plain,(
  ( ! [X121,X118,X120,X119] : (p(s(bool,X118)) | ~p(s(bool,X119)) | p(s(bool,X120)) | s(bool,X120) != s(bool,'const.bool.IN_2'(s(sK17,X121),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | $spl286),
  inference(resolution,[],[f8902,f705])).
fof(f10715,plain,(
  ( ! [X116,X114,X117,X115] : (p(s(bool,X114)) | p(s(bool,X115)) | s(bool,X114) != s(bool,X116) | s(bool,X115) != s(bool,'const.bool.IN_2'(s(sK17,X117),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | $spl286),
  inference(resolution,[],[f8902,f704])).
fof(f10714,plain,(
  ( ! [X111,X112,X110,X113] : (p(s(bool,X110)) | ~p(s(bool,X111)) | s(bool,X110) != s(bool,'const.bool.IN_2'(s(sK17,X112),s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | s(bool,X111) != s(bool,X113)) ) | $spl286),
  inference(resolution,[],[f8902,f702])).
fof(f10712,plain,(
  ( ! [X103,X105,X102,X104] : (p(s(bool,X102)) | p(s(bool,X103)) | s(bool,X102) != s(bool,X104) | s(bool,X103) != s(bool,'const.bool.IN_2'(s(sK17,X105),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | $spl286),
  inference(resolution,[],[f8902,f700])).
fof(f10673,plain,(
  ( ! [X121,X119] : (~p(s(bool,X119)) | s(bool,X119) != s(bool,'const.bool.IN_2'(s(sK17,X121),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | $spl342),
  inference(cnf_transformation,[],[f10673_D])).
fof(f10673_D,plain,(
  ( ! [X121,X119] : (~p(s(bool,X119)) | s(bool,X119) != s(bool,'const.bool.IN_2'(s(sK17,X121),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) <=> ~$spl342),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl342])])).
fof(f10625,plain,(
  ( ! [X449,X448,X450] : (p(s(bool,X448)) | ~p(s(bool,X449)) | s(bool,X449) != s(bool,'const.bool.IN_2'(s(sK17,X450),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | $spl276),
  inference(resolution,[],[f8860,f957])).
fof(f10523,plain,(
  ( ! [X146,X147,X145,X148] : (~p(s(bool,X145)) | ~p(s(bool,X146)) | s(bool,X146) != s(bool,'const.bool.IN_2'(s(sK17,X147),s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | s(bool,X145) != s(bool,X148)) ) | $spl276),
  inference(resolution,[],[f8860,f706])).
fof(f10521,plain,(
  ( ! [X140,X138,X141,X139] : (p(s(bool,X138)) | ~p(s(bool,X139)) | ~p(s(bool,X140)) | s(bool,X139) != s(bool,'const.bool.IN_2'(s(sK17,X141),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | $spl276),
  inference(resolution,[],[f8860,f705])).
fof(f10519,plain,(
  ( ! [X132,X130,X133,X131] : (p(s(bool,X130)) | ~p(s(bool,X131)) | s(bool,X130) != s(bool,X132) | s(bool,X131) != s(bool,'const.bool.IN_2'(s(sK17,X133),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | $spl276),
  inference(resolution,[],[f8860,f704])).
fof(f10518,plain,(
  ( ! [X127,X128,X126,X129] : (~p(s(bool,X126)) | ~p(s(bool,X127)) | s(bool,X127) != s(bool,'const.bool.IN_2'(s(sK17,X128),s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | s(bool,X126) != s(bool,X129)) ) | $spl276),
  inference(resolution,[],[f8860,f702])).
fof(f10516,plain,(
  ( ! [X121,X118,X120,X119] : (p(s(bool,X118)) | ~p(s(bool,X119)) | s(bool,X118) != s(bool,X120) | s(bool,X119) != s(bool,'const.bool.IN_2'(s(sK17,X121),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | $spl276),
  inference(resolution,[],[f8860,f700])).
fof(f4730,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X1)) ) | $spl138),
  inference(cnf_transformation,[],[f4730_D])).
fof(f4730_D,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X1)) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f10478,plain,(
  ( ! [X14,X12,X13,X11] : (sP3(X11,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X12) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X13),X14),s(fun(sK16,bool),sK20),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1007,f678])).
fof(f1007,plain,(
  ( ! [X78,X79] : (sP3(X78,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)),X79)) )),
  inference(resolution,[],[f647,f733])).
fof(f10477,plain,(
  ( ! [X10,X8,X7,X9] : (sP3(X7,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X8) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK16),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) )),
  inference(superposition,[],[f1007,f680])).
fof(f4729,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)),X1)) ) | $spl136),
  inference(cnf_transformation,[],[f4729_D])).
fof(f4729_D,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)),X1)) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f10471,plain,(
  ( ! [X14,X12,X13,X11] : (sP3(X11,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)),X12) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1007,f678])).
fof(f10470,plain,(
  ( ! [X10,X8,X7,X9] : (sP3(X7,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)),X8) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) )),
  inference(superposition,[],[f1007,f680])).
fof(f10469,plain,(
  ( ! [X2,X0,X3,X1] : (sP3(X0,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)),X1) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1007,f948])).
fof(f4704,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) | $spl134),
  inference(cnf_transformation,[],[f4704_D])).
fof(f4704_D,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f10454,plain,(
  ( ! [X14,X12,X13,X11] : (sP3(X11,X12,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X13),X14),s(fun(sK16,bool),sK20),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1006,f678])).
fof(f1006,plain,(
  ( ! [X76,X77] : (sP3(X76,X77,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) )),
  inference(resolution,[],[f647,f732])).
fof(f10453,plain,(
  ( ! [X10,X8,X7,X9] : (sP3(X7,X8,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK16),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) )),
  inference(superposition,[],[f1006,f680])).
fof(f4703,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)))) ) | $spl132),
  inference(cnf_transformation,[],[f4703_D])).
fof(f4703_D,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)))) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f10447,plain,(
  ( ! [X14,X12,X13,X11] : (sP3(X11,X12,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1006,f678])).
fof(f10446,plain,(
  ( ! [X10,X8,X7,X9] : (sP3(X7,X8,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) )),
  inference(superposition,[],[f1006,f680])).
fof(f10445,plain,(
  ( ! [X2,X0,X3,X1] : (sP3(X0,X1,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1006,f948])).
fof(f4347,plain,(
  ( ! [X0,X1] : (sP1(X0,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X1)) ) | $spl128),
  inference(cnf_transformation,[],[f4347_D])).
fof(f4347_D,plain,(
  ( ! [X0,X1] : (sP1(X0,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X1)) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f10352,plain,(
  ( ! [X14,X12,X13,X11] : (sP1(X11,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X12) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X13),X14),s(fun(sK16,bool),sK20),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f993,f678])).
fof(f993,plain,(
  ( ! [X50,X51] : (sP1(X50,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)),X51)) )),
  inference(resolution,[],[f647,f716])).
fof(f10351,plain,(
  ( ! [X10,X8,X7,X9] : (sP1(X7,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X8) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK16),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) )),
  inference(superposition,[],[f993,f680])).
fof(f4346,plain,(
  ( ! [X0,X1] : (sP1(X0,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)),X1)) ) | $spl126),
  inference(cnf_transformation,[],[f4346_D])).
fof(f4346_D,plain,(
  ( ! [X0,X1] : (sP1(X0,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)),X1)) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f10345,plain,(
  ( ! [X14,X12,X13,X11] : (sP1(X11,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)),X12) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f993,f678])).
fof(f10344,plain,(
  ( ! [X10,X8,X7,X9] : (sP1(X7,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)),X8) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) )),
  inference(superposition,[],[f993,f680])).
fof(f10343,plain,(
  ( ! [X2,X0,X3,X1] : (sP1(X0,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)),X1) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f993,f948])).
fof(f4112,plain,(
  ~sP0(sK16,sK17,sK19,'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0')),sK18) | $spl115),
  inference(cnf_transformation,[],[f4112_D])).
fof(f4112_D,plain,(
  ~sP0(sK16,sK17,sK19,'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0')),sK18) <=> ~$spl115),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl115])])).
fof(f10329,plain,(
  ( ! [X8,X7] : (~sP0(sK16,sK17,sK19,'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0')),sK18) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X7),X8),s(fun(sK17,bool),sK18),s(fun(X7,bool),'const.pred_set.EMPTY_0'))))) ) | $spl41),
  inference(superposition,[],[f1337,f678])).
fof(f10328,plain,(
  ( ! [X6,X5] : (~sP0(sK16,sK17,sK19,'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0')),sK18) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X5,sK17),X6),s(fun(X5,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl41),
  inference(superposition,[],[f1337,f680])).
fof(f10327,plain,(
  ( ! [X0,X1] : (~sP0(sK16,sK17,sK19,'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0')),sK18) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X0),X1),s(fun(sK17,bool),sK18),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl41),
  inference(superposition,[],[f1337,f948])).
fof(f10265,plain,(
  ( ! [X8,X9] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X8),X9),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))),s(fun(X8,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f648,f678])).
fof(f10264,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X6,sK16),X7),s(fun(X6,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))))))) )),
  inference(superposition,[],[f648,f680])).
fof(f10262,plain,(
  ( ! [X2,X1] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK16,X1),X2),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))),s(fun(X1,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f648,f948])).
fof(f1358,plain,(
  ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | $spl51),
  inference(cnf_transformation,[],[f1358_D])).
fof(f1358_D,plain,(
  ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) <=> ~$spl51),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl51])])).
fof(f10298,plain,(
  ( ! [X8,X7] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X7),X8),s(fun(sK17,bool),sK18),s(fun(X7,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(forward_demodulation,[],[f10256,f697])).
fof(f697,plain,(
  ( ! [X2,X0,X1] : (s(fun(X1,bool),'const.pred_set.EMPTY_0') = s(fun(X1,bool),'const.pred_set.IMAGE_2'(s(fun(X0,X1),X2),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(cnf_transformation,[],[f306])).
fof(f306,plain,(
  ! [X0,X1,X2] : s(fun(X1,bool),'const.pred_set.EMPTY_0') = s(fun(X1,bool),'const.pred_set.IMAGE_2'(s(fun(X0,X1),X2),s(fun(X0,bool),'const.pred_set.EMPTY_0')))),
  inference(flattening,[],[f305])).
fof(f305,plain,(
  ! [X0,X1] : ! [X2] : s(fun(X1,bool),'const.pred_set.EMPTY_0') = s(fun(X1,bool),'const.pred_set.IMAGE_2'(s(fun(X0,X1),X2),s(fun(X0,bool),'const.pred_set.EMPTY_0')))),
  inference(rectify,[],[f218])).
fof(f218,axiom,(
  ! [X7,X6] : ! [X9] : s(fun(X6,bool),'const.pred_set.EMPTY_0') = s(fun(X6,bool),'const.pred_set.IMAGE_2'(s(fun(X7,X6),X9),s(fun(X7,bool),'const.pred_set.EMPTY_0')))),
  file('thm.util_prob.INJ_IMAGE_BIJ','thm.pred_set.IMAGE_EMPTY')).
fof(f10256,plain,(
  ( ! [X8,X7] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X7),X8),s(fun(sK17,bool),sK18),s(fun(X7,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f648,f678])).
fof(f10297,plain,(
  ( ! [X6,X5] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X5,sK17),X6),s(fun(X5,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) )),
  inference(forward_demodulation,[],[f10255,f697])).
fof(f10255,plain,(
  ( ! [X6,X5] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X5,sK17),X6),s(fun(X5,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) )),
  inference(superposition,[],[f648,f680])).
fof(f10296,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X0),X1),s(fun(sK17,bool),sK18),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(forward_demodulation,[],[f10254,f697])).
fof(f10254,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X0),X1),s(fun(sK17,bool),sK18),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f648,f948])).
fof(f1340,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | s(bool,X4) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18)))))) ) | $spl44),
  inference(cnf_transformation,[],[f1340_D])).
fof(f1340_D,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | s(bool,X4) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18)))))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f10247,plain,(
  ( ! [X225,X226] : (p(s(bool,X225)) | ~p(s(bool,X226)) | s(bool,X226) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18)))))) )),
  inference(resolution,[],[f648,f957])).
fof(f10145,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))))) | s(bool,X24) != s(bool,X26)) )),
  inference(resolution,[],[f648,f706])).
fof(f10143,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18)))))) )),
  inference(resolution,[],[f648,f705])).
fof(f10141,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18)))))) )),
  inference(resolution,[],[f648,f704])).
fof(f10140,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))))) | s(bool,X9) != s(bool,X11)) )),
  inference(resolution,[],[f648,f702])).
fof(f10138,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18)))))) )),
  inference(resolution,[],[f648,f700])).
fof(f1339,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK16,happ(s(fun(sK17,sK16),sK19),s(sK17,sK25('const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18)),sK18,sK19,sK16,sK17)))),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18)))))) | $spl43),
  inference(cnf_transformation,[],[f1339_D])).
fof(f1339_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK16,happ(s(fun(sK17,sK16),sK19),s(sK17,sK25('const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18)),sK18,sK19,sK16,sK17)))),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18)))))) <=> ~$spl43),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl43])])).
fof(f1337,plain,(
  ~sP0(sK16,sK17,sK19,'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18)),sK18) | $spl41),
  inference(cnf_transformation,[],[f1337_D])).
fof(f1337_D,plain,(
  ~sP0(sK16,sK17,sK19,'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18)),sK18) <=> ~$spl41),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl41])])).
fof(f10134,plain,(
  ~sP0(sK16,sK17,sK19,'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18)),sK18) | ~p(s(bool,'const.bool.IN_2'(s(sK16,happ(s(fun(sK17,sK16),sK19),s(sK17,sK25('const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18)),sK18,sK19,sK16,sK17)))),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))))))),
  inference(resolution,[],[f648,f672])).
fof(f672,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | ~sP0(X1,X0,X2,X4,X3) | ~p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,sK25(X4,X3,X2,X1,X0)))),s(fun(X1,bool),X4))))) )),
  inference(cnf_transformation,[],[f493])).
fof(f493,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | (! [X6] : (~p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X6))),s(fun(X1,bool),X4))))) & sP0(X1,X0,X2,X4,X3))) & ((p(s(bool,'const.bool.IN_2'(s(X0,sK25(X4,X3,X2,X1,X0)),s(fun(X0,bool),X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,sK25(X4,X3,X2,X1,X0)))),s(fun(X1,bool),X4))))) | ~sP0(X1,X0,X2,X4,X3) | p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK25])],[f492])).
fof(f492,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | (! [X6] : (~p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X6))),s(fun(X1,bool),X4))))) & sP0(X1,X0,X2,X4,X3))) & (? [X5] : (p(s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X5))),s(fun(X1,bool),X4))))) | ~sP0(X1,X0,X2,X4,X3) | p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))))),
  inference(rectify,[],[f491])).
fof(f491,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | (! [X9] : (~p(s(bool,'const.bool.IN_2'(s(X0,X9),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X9))),s(fun(X1,bool),X4))))) & sP0(X1,X0,X2,X4,X3))) & (? [X9] : (p(s(bool,'const.bool.IN_2'(s(X0,X9),s(fun(X0,bool),X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X9))),s(fun(X1,bool),X4))))) | ~sP0(X1,X0,X2,X4,X3) | p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))))),
  inference(flattening,[],[f490])).
fof(f490,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | (! [X9] : (~p(s(bool,'const.bool.IN_2'(s(X0,X9),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X9))),s(fun(X1,bool),X4))))) & sP0(X1,X0,X2,X4,X3))) & ((? [X9] : (p(s(bool,'const.bool.IN_2'(s(X0,X9),s(fun(X0,bool),X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X9))),s(fun(X1,bool),X4))))) | ~sP0(X1,X0,X2,X4,X3)) | p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))))),
  inference(nnf_transformation,[],[f459])).
fof(f459,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) <=> (! [X9] : (~p(s(bool,'const.bool.IN_2'(s(X0,X9),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X9))),s(fun(X1,bool),X4))))) & sP0(X1,X0,X2,X4,X3)))),
  inference(definition_folding,[],[f408,f458])).
fof(f458,plain,(
  ! [X1,X0,X2,X4,X3] : (sP0(X1,X0,X2,X4,X3) <=> ? [X5] : (! [X8] : (~p(s(bool,'const.bool.IN_2'(s(X1,X8),s(fun(X1,bool),X4)))) | p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X5),s(X1,X8))),s(fun(X0,bool),X3))))) & ! [X7] : (~p(s(bool,'const.bool.IN_2'(s(X0,X7),s(fun(X0,bool),X3)))) | s(X0,X7) = s(X0,happ(s(fun(X1,X0),X5),s(X1,happ(s(fun(X0,X1),X2),s(X0,X7)))))) & ! [X6] : (~p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X4)))) | s(X1,X6) = s(X1,happ(s(fun(X0,X1),X2),s(X0,happ(s(fun(X1,X0),X5),s(X1,X6))))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f408,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) <=> (! [X9] : (~p(s(bool,'const.bool.IN_2'(s(X0,X9),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X9))),s(fun(X1,bool),X4))))) & ? [X5] : (! [X8] : (~p(s(bool,'const.bool.IN_2'(s(X1,X8),s(fun(X1,bool),X4)))) | p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X5),s(X1,X8))),s(fun(X0,bool),X3))))) & ! [X7] : (~p(s(bool,'const.bool.IN_2'(s(X0,X7),s(fun(X0,bool),X3)))) | s(X0,X7) = s(X0,happ(s(fun(X1,X0),X5),s(X1,happ(s(fun(X0,X1),X2),s(X0,X7)))))) & ! [X6] : (~p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X4)))) | s(X1,X6) = s(X1,happ(s(fun(X0,X1),X2),s(X0,happ(s(fun(X1,X0),X5),s(X1,X6)))))))))),
  inference(ennf_transformation,[],[f270])).
fof(f270,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) <=> (! [X9] : (p(s(bool,'const.bool.IN_2'(s(X0,X9),s(fun(X0,bool),X3)))) => p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X9))),s(fun(X1,bool),X4))))) & ? [X5] : (! [X8] : (p(s(bool,'const.bool.IN_2'(s(X1,X8),s(fun(X1,bool),X4)))) => p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X5),s(X1,X8))),s(fun(X0,bool),X3))))) & ! [X7] : (p(s(bool,'const.bool.IN_2'(s(X0,X7),s(fun(X0,bool),X3)))) => s(X0,X7) = s(X0,happ(s(fun(X1,X0),X5),s(X1,happ(s(fun(X0,X1),X2),s(X0,X7)))))) & ! [X6] : (p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X4)))) => s(X1,X6) = s(X1,happ(s(fun(X0,X1),X2),s(X0,happ(s(fun(X1,X0),X5),s(X1,X6)))))))))),
  inference(flattening,[],[f269])).
fof(f269,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) <=> (! [X9] : (p(s(bool,'const.bool.IN_2'(s(X0,X9),s(fun(X0,bool),X3)))) => p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X9))),s(fun(X1,bool),X4))))) & ? [X5] : (! [X8] : (p(s(bool,'const.bool.IN_2'(s(X1,X8),s(fun(X1,bool),X4)))) => p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X5),s(X1,X8))),s(fun(X0,bool),X3))))) & ! [X7] : (p(s(bool,'const.bool.IN_2'(s(X0,X7),s(fun(X0,bool),X3)))) => s(X0,X7) = s(X0,happ(s(fun(X1,X0),X5),s(X1,happ(s(fun(X0,X1),X2),s(X0,X7)))))) & ! [X6] : (p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X4)))) => s(X1,X6) = s(X1,happ(s(fun(X0,X1),X2),s(X0,happ(s(fun(X1,X0),X5),s(X1,X6)))))))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X7,X6] : ! [X9,X8,X10] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(X7,X6),X9),s(fun(X7,bool),X8),s(fun(X6,bool),X10)))) <=> (! [X13] : (p(s(bool,'const.bool.IN_2'(s(X7,X13),s(fun(X7,bool),X8)))) => p(s(bool,'const.bool.IN_2'(s(X6,happ(s(fun(X7,X6),X9),s(X7,X13))),s(fun(X6,bool),X10))))) & ? [X11] : (! [X13] : (p(s(bool,'const.bool.IN_2'(s(X6,X13),s(fun(X6,bool),X10)))) => p(s(bool,'const.bool.IN_2'(s(X7,happ(s(fun(X6,X7),X11),s(X6,X13))),s(fun(X7,bool),X8))))) & ! [X13] : (p(s(bool,'const.bool.IN_2'(s(X7,X13),s(fun(X7,bool),X8)))) => s(X7,X13) = s(X7,happ(s(fun(X6,X7),X11),s(X6,happ(s(fun(X7,X6),X9),s(X7,X13)))))) & ! [X13] : (p(s(bool,'const.bool.IN_2'(s(X6,X13),s(fun(X6,bool),X10)))) => s(X6,X13) = s(X6,happ(s(fun(X7,X6),X9),s(X7,happ(s(fun(X6,X7),X11),s(X6,X13)))))))))),
  file('thm.util_prob.INJ_IMAGE_BIJ','thm.pred_set.BIJ_IFF_INV')).
fof(f1335,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK17,sK25('const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18)),sK18,sK19,sK16,sK17)),s(fun(sK17,bool),sK18)))) | $spl38),
  inference(cnf_transformation,[],[f1335_D])).
fof(f1335_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK17,sK25('const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18)),sK18,sK19,sK16,sK17)),s(fun(sK17,bool),sK18)))) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f10133,plain,(
  ~sP0(sK16,sK17,sK19,'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18)),sK18) | p(s(bool,'const.bool.IN_2'(s(sK17,sK25('const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18)),sK18,sK19,sK16,sK17)),s(fun(sK17,bool),sK18))))),
  inference(resolution,[],[f648,f671])).
fof(f671,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | ~sP0(X1,X0,X2,X4,X3) | p(s(bool,'const.bool.IN_2'(s(X0,sK25(X4,X3,X2,X1,X0)),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f493])).
fof(f1331,plain,(
  ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18)))))) | $spl35),
  inference(cnf_transformation,[],[f1331_D])).
fof(f1331_D,plain,(
  ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18)))))) <=> ~$spl35),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl35])])).
fof(f1333,plain,(
  ~p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18)))))) | $spl37),
  inference(cnf_transformation,[],[f1333_D])).
fof(f1333_D,plain,(
  ~p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18)))))) <=> ~$spl37),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl37])])).
fof(f10132,plain,(
  ~p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18)))))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))))))),
  inference(resolution,[],[f648,f675])).
fof(f675,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | ~p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))) )),
  inference(cnf_transformation,[],[f495])).
fof(f9496,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK16,happ(s(fun(sK17,sK16),sK19),s(sK17,sK61('const.pred_set.EMPTY_0',sK18,sK19,sK16,sK17)))),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | $spl335),
  inference(cnf_transformation,[],[f9496_D])).
fof(f9496_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK16,happ(s(fun(sK17,sK16),sK19),s(sK17,sK61('const.pred_set.EMPTY_0',sK18,sK19,sK16,sK17)))),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) <=> ~$spl335),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl335])])).
fof(f9502,plain,(
  s(sK17,sK62('const.pred_set.EMPTY_0',sK18,sK19,sK16,sK17)) != s(sK17,sK63('const.pred_set.EMPTY_0',sK18,sK19,sK16,sK17)) | $spl341),
  inference(cnf_transformation,[],[f9502_D])).
fof(f9502_D,plain,(
  s(sK17,sK62('const.pred_set.EMPTY_0',sK18,sK19,sK16,sK17)) != s(sK17,sK63('const.pred_set.EMPTY_0',sK18,sK19,sK16,sK17)) <=> ~$spl341),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl341])])).
fof(f10076,plain,(
  s(sK17,sK62('const.pred_set.EMPTY_0',sK18,sK19,sK16,sK17)) != s(sK17,sK63('const.pred_set.EMPTY_0',sK18,sK19,sK16,sK17)) | ~p(s(bool,'const.bool.IN_2'(s(sK16,happ(s(fun(sK17,sK16),sK19),s(sK17,sK61('const.pred_set.EMPTY_0',sK18,sK19,sK16,sK17)))),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | $spl29),
  inference(resolution,[],[f7943,f943])).
fof(f943,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP15(X0,X1,X2,X3,X4) | s(X0,sK62(X4,X3,X2,X1,X0)) != s(X0,sK63(X4,X3,X2,X1,X0)) | ~p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,sK61(X4,X3,X2,X1,X0)))),s(fun(X1,bool),X4))))) )),
  inference(cnf_transformation,[],[f643])).
fof(f643,plain,(
  ! [X0,X1,X2,X3,X4] : ((~sP15(X0,X1,X2,X3,X4) | (! [X10] : (~p(s(bool,'const.bool.IN_2'(s(X0,X10),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X10))),s(fun(X1,bool),X4))))) & ! [X8,X9] : (~p(s(bool,'const.bool.IN_2'(s(X0,X8),s(fun(X0,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X9),s(fun(X0,bool),X3)))) | s(X1,happ(s(fun(X0,X1),X2),s(X0,X8))) != s(X1,happ(s(fun(X0,X1),X2),s(X0,X9))) | s(X0,X8) = s(X0,X9)))) & ((p(s(bool,'const.bool.IN_2'(s(X0,sK61(X4,X3,X2,X1,X0)),s(fun(X0,bool),X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,sK61(X4,X3,X2,X1,X0)))),s(fun(X1,bool),X4))))) | (p(s(bool,'const.bool.IN_2'(s(X0,sK62(X4,X3,X2,X1,X0)),s(fun(X0,bool),X3)))) & p(s(bool,'const.bool.IN_2'(s(X0,sK63(X4,X3,X2,X1,X0)),s(fun(X0,bool),X3)))) & s(X1,happ(s(fun(X0,X1),X2),s(X0,sK62(X4,X3,X2,X1,X0)))) = s(X1,happ(s(fun(X0,X1),X2),s(X0,sK63(X4,X3,X2,X1,X0)))) & s(X0,sK62(X4,X3,X2,X1,X0)) != s(X0,sK63(X4,X3,X2,X1,X0))) | sP15(X0,X1,X2,X3,X4)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK61,sK62,sK63])],[f642])).
fof(f642,plain,(
  ! [X0,X1,X2,X3,X4] : ((~sP15(X0,X1,X2,X3,X4) | (! [X10] : (~p(s(bool,'const.bool.IN_2'(s(X0,X10),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X10))),s(fun(X1,bool),X4))))) & ! [X8,X9] : (~p(s(bool,'const.bool.IN_2'(s(X0,X8),s(fun(X0,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X9),s(fun(X0,bool),X3)))) | s(X1,happ(s(fun(X0,X1),X2),s(X0,X8))) != s(X1,happ(s(fun(X0,X1),X2),s(X0,X9))) | s(X0,X8) = s(X0,X9)))) & (? [X7] : (p(s(bool,'const.bool.IN_2'(s(X0,X7),s(fun(X0,bool),X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X7))),s(fun(X1,bool),X4))))) | ? [X5,X6] : (p(s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),X3)))) & p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X3)))) & s(X1,happ(s(fun(X0,X1),X2),s(X0,X5))) = s(X1,happ(s(fun(X0,X1),X2),s(X0,X6))) & s(X0,X5) != s(X0,X6)) | sP15(X0,X1,X2,X3,X4)))),
  inference(rectify,[],[f641])).
fof(f641,plain,(
  ! [X1,X0,X2,X3,X4] : ((~sP15(X1,X0,X2,X3,X4) | (! [X7] : (~p(s(bool,'const.bool.IN_2'(s(X1,X7),s(fun(X1,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))),s(fun(X0,bool),X4))))) & ! [X5,X6] : (~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X3)))) | s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) | s(X1,X5) = s(X1,X6)))) & (? [X7] : (p(s(bool,'const.bool.IN_2'(s(X1,X7),s(fun(X1,bool),X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))),s(fun(X0,bool),X4))))) | ? [X5,X6] : (p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) & p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X3)))) & s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) & s(X1,X5) != s(X1,X6)) | sP15(X1,X0,X2,X3,X4)))),
  inference(flattening,[],[f640])).
fof(f640,plain,(
  ! [X1,X0,X2,X3,X4] : ((~sP15(X1,X0,X2,X3,X4) | (! [X7] : (~p(s(bool,'const.bool.IN_2'(s(X1,X7),s(fun(X1,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))),s(fun(X0,bool),X4))))) & ! [X5,X6] : (~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X3)))) | s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) | s(X1,X5) = s(X1,X6)))) & ((? [X7] : (p(s(bool,'const.bool.IN_2'(s(X1,X7),s(fun(X1,bool),X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))),s(fun(X0,bool),X4))))) | ? [X5,X6] : (p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) & p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X3)))) & s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) & s(X1,X5) != s(X1,X6))) | sP15(X1,X0,X2,X3,X4)))),
  inference(nnf_transformation,[],[f482])).
fof(f482,plain,(
  ! [X1,X0,X2,X3,X4] : (sP15(X1,X0,X2,X3,X4) <=> (! [X7] : (~p(s(bool,'const.bool.IN_2'(s(X1,X7),s(fun(X1,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))),s(fun(X0,bool),X4))))) & ! [X5,X6] : (~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X3)))) | s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) | s(X1,X5) = s(X1,X6))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP15])])).
fof(f7943,plain,(
  ~sP15(sK17,sK16,sK19,sK18,'const.pred_set.EMPTY_0') | $spl29),
  inference(resolution,[],[f1174,f945])).
fof(f945,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4)))) | ~sP15(X1,X0,X2,X3,X4)) )),
  inference(cnf_transformation,[],[f644])).
fof(f644,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4)))) | sP15(X1,X0,X2,X3,X4)) & (~sP15(X1,X0,X2,X3,X4) | p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4))))))),
  inference(nnf_transformation,[],[f483])).
fof(f483,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4)))) <=> sP15(X1,X0,X2,X3,X4))),
  inference(definition_folding,[],[f454,f482])).
fof(f454,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4)))) <=> (! [X7] : (~p(s(bool,'const.bool.IN_2'(s(X1,X7),s(fun(X1,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))),s(fun(X0,bool),X4))))) & ! [X5,X6] : (~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X3)))) | s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) | s(X1,X5) = s(X1,X6))))),
  inference(flattening,[],[f453])).
fof(f453,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4)))) <=> (! [X7] : (~p(s(bool,'const.bool.IN_2'(s(X1,X7),s(fun(X1,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))),s(fun(X0,bool),X4))))) & ! [X5,X6] : ((~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X3))))) | (s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) | s(X1,X5) = s(X1,X6)))))),
  inference(ennf_transformation,[],[f386])).
fof(f386,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4)))) <=> (! [X7] : (p(s(bool,'const.bool.IN_2'(s(X1,X7),s(fun(X1,bool),X3)))) => p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))),s(fun(X0,bool),X4))))) & ! [X5,X6] : ((p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) & p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X3))))) => (s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) => s(X1,X5) = s(X1,X6)))))),
  inference(flattening,[],[f385])).
fof(f385,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4)))) <=> (! [X7] : (p(s(bool,'const.bool.IN_2'(s(X1,X7),s(fun(X1,bool),X3)))) => p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))),s(fun(X0,bool),X4))))) & ! [X5,X6] : ((p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) & p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X3))))) => (s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) => s(X1,X5) = s(X1,X6)))))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X6,X7] : ! [X9,X8,X10] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X7,X6),X9),s(fun(X7,bool),X8),s(fun(X6,bool),X10)))) <=> (! [X13] : (p(s(bool,'const.bool.IN_2'(s(X7,X13),s(fun(X7,bool),X8)))) => p(s(bool,'const.bool.IN_2'(s(X6,happ(s(fun(X7,X6),X9),s(X7,X13))),s(fun(X6,bool),X10))))) & ! [X13,X24] : ((p(s(bool,'const.bool.IN_2'(s(X7,X13),s(fun(X7,bool),X8)))) & p(s(bool,'const.bool.IN_2'(s(X7,X24),s(fun(X7,bool),X8))))) => (s(X6,happ(s(fun(X7,X6),X9),s(X7,X13))) = s(X6,happ(s(fun(X7,X6),X9),s(X7,X24))) => s(X7,X13) = s(X7,X24)))))),
  file('thm.util_prob.INJ_IMAGE_BIJ','thm.pred_set.INJ_DEF')).
fof(f9492,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK17,sK61('const.pred_set.EMPTY_0',sK18,sK19,sK16,sK17)),s(fun(sK17,bool),sK18)))) | $spl330),
  inference(cnf_transformation,[],[f9492_D])).
fof(f9492_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK17,sK61('const.pred_set.EMPTY_0',sK18,sK19,sK16,sK17)),s(fun(sK17,bool),sK18)))) <=> ~$spl330),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl330])])).
fof(f10075,plain,(
  s(sK17,sK62('const.pred_set.EMPTY_0',sK18,sK19,sK16,sK17)) != s(sK17,sK63('const.pred_set.EMPTY_0',sK18,sK19,sK16,sK17)) | p(s(bool,'const.bool.IN_2'(s(sK17,sK61('const.pred_set.EMPTY_0',sK18,sK19,sK16,sK17)),s(fun(sK17,bool),sK18)))) | $spl29),
  inference(resolution,[],[f7943,f942])).
fof(f942,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP15(X0,X1,X2,X3,X4) | s(X0,sK62(X4,X3,X2,X1,X0)) != s(X0,sK63(X4,X3,X2,X1,X0)) | p(s(bool,'const.bool.IN_2'(s(X0,sK61(X4,X3,X2,X1,X0)),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f643])).
fof(f9500,plain,(
  s(sK16,happ(s(fun(sK17,sK16),sK19),s(sK17,sK62('const.pred_set.EMPTY_0',sK18,sK19,sK16,sK17)))) = s(sK16,happ(s(fun(sK17,sK16),sK19),s(sK17,sK63('const.pred_set.EMPTY_0',sK18,sK19,sK16,sK17)))) | $spl338),
  inference(cnf_transformation,[],[f9500_D])).
fof(f9500_D,plain,(
  s(sK16,happ(s(fun(sK17,sK16),sK19),s(sK17,sK62('const.pred_set.EMPTY_0',sK18,sK19,sK16,sK17)))) = s(sK16,happ(s(fun(sK17,sK16),sK19),s(sK17,sK63('const.pred_set.EMPTY_0',sK18,sK19,sK16,sK17)))) <=> ~$spl338),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl338])])).
fof(f10074,plain,(
  s(sK16,happ(s(fun(sK17,sK16),sK19),s(sK17,sK62('const.pred_set.EMPTY_0',sK18,sK19,sK16,sK17)))) = s(sK16,happ(s(fun(sK17,sK16),sK19),s(sK17,sK63('const.pred_set.EMPTY_0',sK18,sK19,sK16,sK17)))) | ~p(s(bool,'const.bool.IN_2'(s(sK16,happ(s(fun(sK17,sK16),sK19),s(sK17,sK61('const.pred_set.EMPTY_0',sK18,sK19,sK16,sK17)))),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | $spl29),
  inference(resolution,[],[f7943,f941])).
fof(f941,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP15(X0,X1,X2,X3,X4) | s(X1,happ(s(fun(X0,X1),X2),s(X0,sK62(X4,X3,X2,X1,X0)))) = s(X1,happ(s(fun(X0,X1),X2),s(X0,sK63(X4,X3,X2,X1,X0)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,sK61(X4,X3,X2,X1,X0)))),s(fun(X1,bool),X4))))) )),
  inference(cnf_transformation,[],[f643])).
fof(f10073,plain,(
  s(sK16,happ(s(fun(sK17,sK16),sK19),s(sK17,sK62('const.pred_set.EMPTY_0',sK18,sK19,sK16,sK17)))) = s(sK16,happ(s(fun(sK17,sK16),sK19),s(sK17,sK63('const.pred_set.EMPTY_0',sK18,sK19,sK16,sK17)))) | p(s(bool,'const.bool.IN_2'(s(sK17,sK61('const.pred_set.EMPTY_0',sK18,sK19,sK16,sK17)),s(fun(sK17,bool),sK18)))) | $spl29),
  inference(resolution,[],[f7943,f940])).
fof(f940,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP15(X0,X1,X2,X3,X4) | s(X1,happ(s(fun(X0,X1),X2),s(X0,sK62(X4,X3,X2,X1,X0)))) = s(X1,happ(s(fun(X0,X1),X2),s(X0,sK63(X4,X3,X2,X1,X0)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK61(X4,X3,X2,X1,X0)),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f643])).
fof(f9498,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK17,sK63('const.pred_set.EMPTY_0',sK18,sK19,sK16,sK17)),s(fun(sK17,bool),sK18)))) | $spl336),
  inference(cnf_transformation,[],[f9498_D])).
fof(f9498_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK17,sK63('const.pred_set.EMPTY_0',sK18,sK19,sK16,sK17)),s(fun(sK17,bool),sK18)))) <=> ~$spl336),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl336])])).
fof(f10072,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK17,sK63('const.pred_set.EMPTY_0',sK18,sK19,sK16,sK17)),s(fun(sK17,bool),sK18)))) | ~p(s(bool,'const.bool.IN_2'(s(sK16,happ(s(fun(sK17,sK16),sK19),s(sK17,sK61('const.pred_set.EMPTY_0',sK18,sK19,sK16,sK17)))),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | $spl29),
  inference(resolution,[],[f7943,f939])).
fof(f939,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP15(X0,X1,X2,X3,X4) | p(s(bool,'const.bool.IN_2'(s(X0,sK63(X4,X3,X2,X1,X0)),s(fun(X0,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,sK61(X4,X3,X2,X1,X0)))),s(fun(X1,bool),X4))))) )),
  inference(cnf_transformation,[],[f643])).
fof(f10071,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK17,sK63('const.pred_set.EMPTY_0',sK18,sK19,sK16,sK17)),s(fun(sK17,bool),sK18)))) | p(s(bool,'const.bool.IN_2'(s(sK17,sK61('const.pred_set.EMPTY_0',sK18,sK19,sK16,sK17)),s(fun(sK17,bool),sK18)))) | $spl29),
  inference(resolution,[],[f7943,f938])).
fof(f938,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP15(X0,X1,X2,X3,X4) | p(s(bool,'const.bool.IN_2'(s(X0,sK63(X4,X3,X2,X1,X0)),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK61(X4,X3,X2,X1,X0)),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f643])).
fof(f9494,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK17,sK62('const.pred_set.EMPTY_0',sK18,sK19,sK16,sK17)),s(fun(sK17,bool),sK18)))) | $spl332),
  inference(cnf_transformation,[],[f9494_D])).
fof(f9494_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK17,sK62('const.pred_set.EMPTY_0',sK18,sK19,sK16,sK17)),s(fun(sK17,bool),sK18)))) <=> ~$spl332),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl332])])).
fof(f10070,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK17,sK62('const.pred_set.EMPTY_0',sK18,sK19,sK16,sK17)),s(fun(sK17,bool),sK18)))) | ~p(s(bool,'const.bool.IN_2'(s(sK16,happ(s(fun(sK17,sK16),sK19),s(sK17,sK61('const.pred_set.EMPTY_0',sK18,sK19,sK16,sK17)))),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | $spl29),
  inference(resolution,[],[f7943,f937])).
fof(f937,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP15(X0,X1,X2,X3,X4) | p(s(bool,'const.bool.IN_2'(s(X0,sK62(X4,X3,X2,X1,X0)),s(fun(X0,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,sK61(X4,X3,X2,X1,X0)))),s(fun(X1,bool),X4))))) )),
  inference(cnf_transformation,[],[f643])).
fof(f10069,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK17,sK62('const.pred_set.EMPTY_0',sK18,sK19,sK16,sK17)),s(fun(sK17,bool),sK18)))) | p(s(bool,'const.bool.IN_2'(s(sK17,sK61('const.pred_set.EMPTY_0',sK18,sK19,sK16,sK17)),s(fun(sK17,bool),sK18)))) | $spl29),
  inference(resolution,[],[f7943,f936])).
fof(f936,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP15(X0,X1,X2,X3,X4) | p(s(bool,'const.bool.IN_2'(s(X0,sK62(X4,X3,X2,X1,X0)),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK61(X4,X3,X2,X1,X0)),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f643])).
fof(f10055,plain,(
  $false | $spl28),
  inference(trivial_inequality_removal,[],[f9896])).
fof(f9896,plain,(
  s(fun(sK16,bool),'const.pred_set.EMPTY_0') != s(fun(sK16,bool),'const.pred_set.EMPTY_0') | $spl28),
  inference(resolution,[],[f8672,f681])).
fof(f8672,plain,(
  ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | $spl28),
  inference(forward_demodulation,[],[f8114,f697])).
fof(f8114,plain,(
  ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))))) | $spl28),
  inference(backward_demodulation,[],[f7676,f648])).
fof(f7676,plain,(
  s(fun(sK17,bool),'const.pred_set.EMPTY_0') = s(fun(sK17,bool),sK18) | $spl28),
  inference(resolution,[],[f1175,f948])).
fof(f10054,plain,(
  $false | $spl28),
  inference(trivial_inequality_removal,[],[f9897])).
fof(f9897,plain,(
  s(fun(sK17,bool),'const.pred_set.EMPTY_0') != s(fun(sK17,bool),'const.pred_set.EMPTY_0') | $spl28),
  inference(resolution,[],[f8672,f679])).
fof(f9807,plain,(
  ( ! [X225,X226] : (p(s(bool,X225)) | ~p(s(bool,X226)) | s(bool,X226) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) ) | $spl5),
  inference(resolution,[],[f1135,f957])).
fof(f9705,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | s(bool,X24) != s(bool,X26)) ) | $spl5),
  inference(resolution,[],[f1135,f706])).
fof(f9703,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) ) | $spl5),
  inference(resolution,[],[f1135,f705])).
fof(f9701,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) ) | $spl5),
  inference(resolution,[],[f1135,f704])).
fof(f9700,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | s(bool,X9) != s(bool,X11)) ) | $spl5),
  inference(resolution,[],[f1135,f702])).
fof(f9698,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) ) | $spl5),
  inference(resolution,[],[f1135,f700])).
fof(f1175,plain,(
  p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | $spl28),
  inference(cnf_transformation,[],[f1175_D])).
fof(f1175_D,plain,(
  p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f9654,plain,(
  ( ! [X8,X7] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X7),X8),s(fun(sK16,bool),sK20),s(fun(X7,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f647,f678])).
fof(f9653,plain,(
  ( ! [X6,X5] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X5,sK16),X6),s(fun(X5,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) )),
  inference(superposition,[],[f647,f680])).
fof(f1142,plain,(
  ( ! [X13] : (s(bool,X13) != s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20))) | p(s(bool,X13))) ) | $spl12),
  inference(cnf_transformation,[],[f1142_D])).
fof(f1142_D,plain,(
  ( ! [X13] : (s(bool,X13) != s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20))) | p(s(bool,X13))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f9635,plain,(
  ( ! [X251,X250] : (p(s(bool,X250)) | p(s(bool,X251)) | s(bool,X251) != s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) )),
  inference(resolution,[],[f647,f957])).
fof(f9528,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X33) != s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20))) | s(bool,X34) != s(bool,X35)) )),
  inference(resolution,[],[f647,f706])).
fof(f9525,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X26) != s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) )),
  inference(resolution,[],[f647,f705])).
fof(f9524,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) )),
  inference(resolution,[],[f647,f704])).
fof(f9523,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20))) | s(bool,X19) != s(bool,X20)) )),
  inference(resolution,[],[f647,f702])).
fof(f9521,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) )),
  inference(resolution,[],[f647,f700])).
fof(f1137,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))))) | $spl6),
  inference(cnf_transformation,[],[f1137_D])).
fof(f1137_D,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f9510,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))))),
  inference(resolution,[],[f647,f692])).
fof(f692,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(X1,bool),X3))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X1,X0),X4),s(fun(X1,bool),X3))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X1,bool),X3)))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X4),s(fun(X1,bool),X3),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f419])).
fof(f419,plain,(
  ! [X0,X1,X2,X3,X4] : (~p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X4),s(fun(X1,bool),X3),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X1,bool),X3)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(X1,bool),X3))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X1,X0),X4),s(fun(X1,bool),X3))))))),
  inference(flattening,[],[f418])).
fof(f418,plain,(
  ! [X0,X1,X2,X3,X4] : (~p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X4),s(fun(X1,bool),X3),s(fun(X0,bool),X2)))) | (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X1,bool),X3)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(X1,bool),X3))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X1,X0),X4),s(fun(X1,bool),X3)))))))),
  inference(ennf_transformation,[],[f296])).
fof(f296,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X4),s(fun(X1,bool),X3),s(fun(X0,bool),X2)))) => (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X1,bool),X3)))) => s('type.num.num','const.pred_set.CARD_1'(s(fun(X1,bool),X3))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X1,X0),X4),s(fun(X1,bool),X3)))))))),
  inference(flattening,[],[f295])).
fof(f295,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X4),s(fun(X1,bool),X3),s(fun(X0,bool),X2)))) => (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X1,bool),X3)))) => s('type.num.num','const.pred_set.CARD_1'(s(fun(X1,bool),X3))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X1,X0),X4),s(fun(X1,bool),X3)))))))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X6,X7] : ! [X10,X8,X9] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X7,X6),X9),s(fun(X7,bool),X8),s(fun(X6,bool),X10)))) => (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X7,bool),X8)))) => s('type.num.num','const.pred_set.CARD_1'(s(fun(X7,bool),X8))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(X6,bool),'const.pred_set.IMAGE_2'(s(fun(X7,X6),X9),s(fun(X7,bool),X8)))))))),
  file('thm.util_prob.INJ_IMAGE_BIJ','thm.pred_set.INJ_CARD_IMAGE_EQ')).
fof(f1131,plain,(
  ~p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) | $spl1),
  inference(cnf_transformation,[],[f1131_D])).
fof(f1131_D,plain,(
  ~p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f1133,plain,(
  p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) | $spl2),
  inference(cnf_transformation,[],[f1133_D])).
fof(f1133_D,plain,(
  p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f9509,plain,(
  p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) | ~p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20))))),
  inference(resolution,[],[f647,f675])).
fof(f9489,plain,(
  s(sK17,sK62('const.pred_set.EMPTY_0',sK18,sK19,sK16,sK17)) != s(sK17,sK63('const.pred_set.EMPTY_0',sK18,sK19,sK16,sK17)) | ~p(s(bool,'const.bool.IN_2'(s(sK16,happ(s(fun(sK17,sK16),sK19),s(sK17,sK61('const.pred_set.EMPTY_0',sK18,sK19,sK16,sK17)))),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | $spl29),
  inference(resolution,[],[f7943,f943])).
fof(f9488,plain,(
  s(sK17,sK62('const.pred_set.EMPTY_0',sK18,sK19,sK16,sK17)) != s(sK17,sK63('const.pred_set.EMPTY_0',sK18,sK19,sK16,sK17)) | p(s(bool,'const.bool.IN_2'(s(sK17,sK61('const.pred_set.EMPTY_0',sK18,sK19,sK16,sK17)),s(fun(sK17,bool),sK18)))) | $spl29),
  inference(resolution,[],[f7943,f942])).
fof(f9487,plain,(
  s(sK16,happ(s(fun(sK17,sK16),sK19),s(sK17,sK62('const.pred_set.EMPTY_0',sK18,sK19,sK16,sK17)))) = s(sK16,happ(s(fun(sK17,sK16),sK19),s(sK17,sK63('const.pred_set.EMPTY_0',sK18,sK19,sK16,sK17)))) | ~p(s(bool,'const.bool.IN_2'(s(sK16,happ(s(fun(sK17,sK16),sK19),s(sK17,sK61('const.pred_set.EMPTY_0',sK18,sK19,sK16,sK17)))),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | $spl29),
  inference(resolution,[],[f7943,f941])).
fof(f9486,plain,(
  s(sK16,happ(s(fun(sK17,sK16),sK19),s(sK17,sK62('const.pred_set.EMPTY_0',sK18,sK19,sK16,sK17)))) = s(sK16,happ(s(fun(sK17,sK16),sK19),s(sK17,sK63('const.pred_set.EMPTY_0',sK18,sK19,sK16,sK17)))) | p(s(bool,'const.bool.IN_2'(s(sK17,sK61('const.pred_set.EMPTY_0',sK18,sK19,sK16,sK17)),s(fun(sK17,bool),sK18)))) | $spl29),
  inference(resolution,[],[f7943,f940])).
fof(f9485,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK17,sK63('const.pred_set.EMPTY_0',sK18,sK19,sK16,sK17)),s(fun(sK17,bool),sK18)))) | ~p(s(bool,'const.bool.IN_2'(s(sK16,happ(s(fun(sK17,sK16),sK19),s(sK17,sK61('const.pred_set.EMPTY_0',sK18,sK19,sK16,sK17)))),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | $spl29),
  inference(resolution,[],[f7943,f939])).
fof(f9484,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK17,sK63('const.pred_set.EMPTY_0',sK18,sK19,sK16,sK17)),s(fun(sK17,bool),sK18)))) | p(s(bool,'const.bool.IN_2'(s(sK17,sK61('const.pred_set.EMPTY_0',sK18,sK19,sK16,sK17)),s(fun(sK17,bool),sK18)))) | $spl29),
  inference(resolution,[],[f7943,f938])).
fof(f9483,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK17,sK62('const.pred_set.EMPTY_0',sK18,sK19,sK16,sK17)),s(fun(sK17,bool),sK18)))) | ~p(s(bool,'const.bool.IN_2'(s(sK16,happ(s(fun(sK17,sK16),sK19),s(sK17,sK61('const.pred_set.EMPTY_0',sK18,sK19,sK16,sK17)))),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | $spl29),
  inference(resolution,[],[f7943,f937])).
fof(f9482,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK17,sK62('const.pred_set.EMPTY_0',sK18,sK19,sK16,sK17)),s(fun(sK17,bool),sK18)))) | p(s(bool,'const.bool.IN_2'(s(sK17,sK61('const.pred_set.EMPTY_0',sK18,sK19,sK16,sK17)),s(fun(sK17,bool),sK18)))) | $spl29),
  inference(resolution,[],[f7943,f936])).
fof(f8902,plain,(
  ( ! [X520] : (p(s(bool,'const.bool.IN_2'(s(sK17,X520),s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl286),
  inference(cnf_transformation,[],[f8902_D])).
fof(f8902_D,plain,(
  ( ! [X520] : (p(s(bool,'const.bool.IN_2'(s(sK17,X520),s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) <=> ~$spl286),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl286])])).
fof(f9073,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK17,sK59(sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | $spl329),
  inference(cnf_transformation,[],[f9073_D])).
fof(f9073_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK17,sK59(sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) <=> ~$spl329),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl329])])).
fof(f9076,plain,(
  ( ! [X606] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK59(sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,'const.bool.IN_2'(s(sK17,X606),s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl28),
  inference(forward_demodulation,[],[f9075,f7676])).
fof(f9075,plain,(
  ( ! [X606] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK59(sK18,sK17)),s(fun(sK17,bool),sK18)))) | p(s(bool,'const.bool.IN_2'(s(sK17,X606),s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl28),
  inference(forward_demodulation,[],[f9074,f956])).
fof(f956,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2))) = s(bool,happ(s(fun(X0,bool),X2),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f404])).
fof(f404,plain,(
  ! [X0,X1,X2] : s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2))) = s(bool,happ(s(fun(X0,bool),X2),s(X0,X1)))),
  inference(flattening,[],[f403])).
fof(f403,plain,(
  ! [X0] : ! [X1,X2] : s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2))) = s(bool,happ(s(fun(X0,bool),X2),s(X0,X1)))),
  inference(rectify,[],[f254])).
fof(f254,axiom,(
  ! [X7] : ! [X5,X48] : s(bool,'const.bool.IN_2'(s(X7,X5),s(fun(X7,bool),X48))) = s(bool,happ(s(fun(X7,bool),X48),s(X7,X5)))),
  file('thm.util_prob.INJ_IMAGE_BIJ','thm.bool.IN_DEF')).
fof(f9074,plain,(
  ( ! [X606] : (p(s(bool,'const.bool.IN_2'(s(sK17,X606),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,sK59(sK18,sK17)))))) ) | $spl28),
  inference(forward_demodulation,[],[f8377,f956])).
fof(f8377,plain,(
  ( ! [X606] : (p(s(bool,happ(s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(sK17,X606)))) | ~p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,sK59(sK18,sK17)))))) ) | $spl28),
  inference(superposition,[],[f915,f7676])).
fof(f915,plain,(
  ( ! [X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK59(X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f630])).
fof(f630,plain,(
  ! [X0,X1] : ((! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK58(X1,X0)))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK59(X1,X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK58,sK59])],[f629])).
fof(f629,plain,(
  ! [X0,X1] : ((! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(rectify,[],[f628])).
fof(f628,plain,(
  ! [X0,X1] : ((! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(nnf_transformation,[],[f445])).
fof(f445,plain,(
  ! [X0,X1] : (? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(ennf_transformation,[],[f370])).
fof(f370,plain,(
  ! [X0,X1] : (~! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(flattening,[],[f369])).
fof(f369,plain,(
  ! [X0] : ! [X1] : (~! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(rectify,[],[f148])).
fof(f148,axiom,(
  ! [X7] : ! [X30] : (~! [X13] : p(s(bool,happ(s(fun(X7,bool),X30),s(X7,X13)))) <=> ? [X13] : ~p(s(bool,happ(s(fun(X7,bool),X30),s(X7,X13)))))),
  file('thm.util_prob.INJ_IMAGE_BIJ','thm.bool.NOT_FORALL_THM')).
fof(f9071,plain,(
  ( ! [X605] : (p(s(bool,'const.bool.IN_2'(s(sK17,X605),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,sK59(sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl28),
  inference(forward_demodulation,[],[f9070,f7676])).
fof(f9070,plain,(
  ( ! [X605] : (p(s(bool,'const.bool.IN_2'(s(sK17,X605),s(fun(sK17,bool),sK18)))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,sK59(sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl28),
  inference(forward_demodulation,[],[f9069,f956])).
fof(f9069,plain,(
  ( ! [X605] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK59(sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,X605))))) ) | $spl28),
  inference(forward_demodulation,[],[f8376,f956])).
fof(f8376,plain,(
  ( ! [X605] : (~p(s(bool,happ(s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(sK17,sK59(sK18,sK17))))) | p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,X605))))) ) | $spl28),
  inference(superposition,[],[f915,f7676])).
fof(f9065,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK17,sK58(sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | $spl327),
  inference(cnf_transformation,[],[f9065_D])).
fof(f9065_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK17,sK58(sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) <=> ~$spl327),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl327])])).
fof(f9068,plain,(
  ( ! [X604] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK58(sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,'const.bool.IN_2'(s(sK17,X604),s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl28),
  inference(forward_demodulation,[],[f9067,f7676])).
fof(f9067,plain,(
  ( ! [X604] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK58(sK18,sK17)),s(fun(sK17,bool),sK18)))) | p(s(bool,'const.bool.IN_2'(s(sK17,X604),s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl28),
  inference(forward_demodulation,[],[f9066,f956])).
fof(f9066,plain,(
  ( ! [X604] : (p(s(bool,'const.bool.IN_2'(s(sK17,X604),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,sK58(sK18,sK17)))))) ) | $spl28),
  inference(forward_demodulation,[],[f8375,f956])).
fof(f8375,plain,(
  ( ! [X604] : (p(s(bool,happ(s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(sK17,X604)))) | ~p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,sK58(sK18,sK17)))))) ) | $spl28),
  inference(superposition,[],[f914,f7676])).
fof(f914,plain,(
  ( ! [X0,X5,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK58(X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))) )),
  inference(cnf_transformation,[],[f630])).
fof(f9063,plain,(
  ( ! [X603] : (p(s(bool,'const.bool.IN_2'(s(sK17,X603),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,sK58(sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl28),
  inference(forward_demodulation,[],[f9062,f7676])).
fof(f9062,plain,(
  ( ! [X603] : (p(s(bool,'const.bool.IN_2'(s(sK17,X603),s(fun(sK17,bool),sK18)))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,sK58(sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl28),
  inference(forward_demodulation,[],[f9061,f956])).
fof(f9061,plain,(
  ( ! [X603] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK58(sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,X603))))) ) | $spl28),
  inference(forward_demodulation,[],[f8374,f956])).
fof(f8374,plain,(
  ( ! [X603] : (~p(s(bool,happ(s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(sK17,sK58(sK18,sK17))))) | p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,X603))))) ) | $spl28),
  inference(superposition,[],[f914,f7676])).
fof(f8860,plain,(
  ( ! [X497] : (~p(s(bool,'const.bool.IN_2'(s(sK17,X497),s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl276),
  inference(cnf_transformation,[],[f8860_D])).
fof(f8860_D,plain,(
  ( ! [X497] : (~p(s(bool,'const.bool.IN_2'(s(sK17,X497),s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) <=> ~$spl276),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl276])])).
fof(f9057,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK17,sK57(sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | $spl324),
  inference(cnf_transformation,[],[f9057_D])).
fof(f9057_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK17,sK57(sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) <=> ~$spl324),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl324])])).
fof(f9060,plain,(
  ( ! [X602] : (p(s(bool,'const.bool.IN_2'(s(sK17,sK57(sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,X602),s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl28),
  inference(forward_demodulation,[],[f9059,f7676])).
fof(f9059,plain,(
  ( ! [X602] : (p(s(bool,'const.bool.IN_2'(s(sK17,sK57(sK18,sK17)),s(fun(sK17,bool),sK18)))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,X602),s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl28),
  inference(forward_demodulation,[],[f9058,f956])).
fof(f9058,plain,(
  ( ! [X602] : (~p(s(bool,'const.bool.IN_2'(s(sK17,X602),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,sK57(sK18,sK17)))))) ) | $spl28),
  inference(forward_demodulation,[],[f8373,f956])).
fof(f8373,plain,(
  ( ! [X602] : (~p(s(bool,happ(s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(sK17,X602)))) | p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,sK57(sK18,sK17)))))) ) | $spl28),
  inference(superposition,[],[f913,f7676])).
fof(f913,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK57(X1,X0)))))) )),
  inference(cnf_transformation,[],[f627])).
fof(f627,plain,(
  ! [X0,X1] : ((p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK56(X1,X0))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) & (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK57(X1,X0))))) | ! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK56,sK57])],[f626])).
fof(f626,plain,(
  ! [X0,X1] : ((? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) & (? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(rectify,[],[f625])).
fof(f625,plain,(
  ! [X0,X1] : ((? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(nnf_transformation,[],[f444])).
fof(f444,plain,(
  ! [X0,X1] : (! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(ennf_transformation,[],[f368])).
fof(f368,plain,(
  ! [X0,X1] : (~? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(flattening,[],[f367])).
fof(f367,plain,(
  ! [X0] : ! [X1] : (~? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(rectify,[],[f139])).
fof(f139,axiom,(
  ! [X7] : ! [X30] : (~? [X13] : p(s(bool,happ(s(fun(X7,bool),X30),s(X7,X13)))) <=> ! [X13] : ~p(s(bool,happ(s(fun(X7,bool),X30),s(X7,X13)))))),
  file('thm.util_prob.INJ_IMAGE_BIJ','thm.bool.NOT_EXISTS_THM')).
fof(f9055,plain,(
  ( ! [X601] : (~p(s(bool,'const.bool.IN_2'(s(sK17,X601),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,'const.bool.IN_2'(s(sK17,sK57(sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl28),
  inference(forward_demodulation,[],[f9054,f7676])).
fof(f9054,plain,(
  ( ! [X601] : (~p(s(bool,'const.bool.IN_2'(s(sK17,X601),s(fun(sK17,bool),sK18)))) | p(s(bool,'const.bool.IN_2'(s(sK17,sK57(sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl28),
  inference(forward_demodulation,[],[f9053,f956])).
fof(f9053,plain,(
  ( ! [X601] : (p(s(bool,'const.bool.IN_2'(s(sK17,sK57(sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,X601))))) ) | $spl28),
  inference(forward_demodulation,[],[f8372,f956])).
fof(f8372,plain,(
  ( ! [X601] : (p(s(bool,happ(s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(sK17,sK57(sK18,sK17))))) | ~p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,X601))))) ) | $spl28),
  inference(superposition,[],[f913,f7676])).
fof(f9049,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK17,sK56(sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | $spl322),
  inference(cnf_transformation,[],[f9049_D])).
fof(f9049_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK17,sK56(sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) <=> ~$spl322),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl322])])).
fof(f9052,plain,(
  ( ! [X600] : (p(s(bool,'const.bool.IN_2'(s(sK17,sK56(sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,X600),s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl28),
  inference(forward_demodulation,[],[f9051,f7676])).
fof(f9051,plain,(
  ( ! [X600] : (p(s(bool,'const.bool.IN_2'(s(sK17,sK56(sK18,sK17)),s(fun(sK17,bool),sK18)))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,X600),s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl28),
  inference(forward_demodulation,[],[f9050,f956])).
fof(f9050,plain,(
  ( ! [X600] : (~p(s(bool,'const.bool.IN_2'(s(sK17,X600),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,sK56(sK18,sK17)))))) ) | $spl28),
  inference(forward_demodulation,[],[f8371,f956])).
fof(f8371,plain,(
  ( ! [X600] : (~p(s(bool,happ(s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(sK17,X600)))) | p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,sK56(sK18,sK17)))))) ) | $spl28),
  inference(superposition,[],[f912,f7676])).
fof(f912,plain,(
  ( ! [X4,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK56(X1,X0)))))) )),
  inference(cnf_transformation,[],[f627])).
fof(f9047,plain,(
  ( ! [X599] : (~p(s(bool,'const.bool.IN_2'(s(sK17,X599),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,'const.bool.IN_2'(s(sK17,sK56(sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl28),
  inference(forward_demodulation,[],[f9046,f7676])).
fof(f9046,plain,(
  ( ! [X599] : (~p(s(bool,'const.bool.IN_2'(s(sK17,X599),s(fun(sK17,bool),sK18)))) | p(s(bool,'const.bool.IN_2'(s(sK17,sK56(sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl28),
  inference(forward_demodulation,[],[f9045,f956])).
fof(f9045,plain,(
  ( ! [X599] : (p(s(bool,'const.bool.IN_2'(s(sK17,sK56(sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,X599))))) ) | $spl28),
  inference(forward_demodulation,[],[f8370,f956])).
fof(f8370,plain,(
  ( ! [X599] : (p(s(bool,happ(s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(sK17,sK56(sK18,sK17))))) | ~p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,X599))))) ) | $spl28),
  inference(superposition,[],[f912,f7676])).
fof(f9041,plain,(
  ( ! [X595] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK55(sK18,X595,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X595))) ) | $spl320),
  inference(cnf_transformation,[],[f9041_D])).
fof(f9041_D,plain,(
  ( ! [X595] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK55(sK18,X595,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X595))) ) <=> ~$spl320),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl320])])).
fof(f9044,plain,(
  ( ! [X598,X597] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK55(sK18,X598,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,'const.bool.IN_2'(s(sK17,X597),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X598))) ) | $spl28),
  inference(forward_demodulation,[],[f9043,f7676])).
fof(f9043,plain,(
  ( ! [X598,X597] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK55(sK18,X598,sK17)),s(fun(sK17,bool),sK18)))) | p(s(bool,'const.bool.IN_2'(s(sK17,X597),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X598))) ) | $spl28),
  inference(forward_demodulation,[],[f9042,f956])).
fof(f9042,plain,(
  ( ! [X598,X597] : (p(s(bool,'const.bool.IN_2'(s(sK17,X597),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,sK55(sK18,X598,sK17))))) | ~p(s(bool,X598))) ) | $spl28),
  inference(forward_demodulation,[],[f8369,f956])).
fof(f8369,plain,(
  ( ! [X598,X597] : (p(s(bool,happ(s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(sK17,X597)))) | ~p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,sK55(sK18,X598,sK17))))) | ~p(s(bool,X598))) ) | $spl28),
  inference(superposition,[],[f907,f7676])).
fof(f907,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK55(X2,X1,X0))))) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f620])).
fof(f620,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK54(X2,X1,X0))))) | ! [X5] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))))) & ((~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK55(X2,X1,X0)))))) | (p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK54,sK55])],[f619])).
fof(f619,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,X1)) | ? [X6] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))) | ! [X5] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))))) & (? [X4] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(rectify,[],[f618])).
fof(f618,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,X1)) | ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | ! [X4] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (? [X4] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(flattening,[],[f617])).
fof(f617,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) | ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ! [X4] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (? [X4] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f364])).
fof(f364,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> ! [X4] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f363])).
fof(f363,plain,(
  ! [X0] : ! [X1,X2] : ((p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> ! [X4] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f94])).
fof(f94,axiom,(
  ! [X7] : ! [X30,X33] : ((p(s(bool,X30)) & ! [X13] : p(s(bool,happ(s(fun(X7,bool),X33),s(X7,X13))))) <=> ! [X13] : (p(s(bool,X30)) & p(s(bool,happ(s(fun(X7,bool),X33),s(X7,X13))))))),
  file('thm.util_prob.INJ_IMAGE_BIJ','thm.bool.RIGHT_AND_FORALL_THM')).
fof(f9040,plain,(
  ( ! [X596,X595] : (p(s(bool,'const.bool.IN_2'(s(sK17,X596),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,sK55(sK18,X595,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X595))) ) | $spl28),
  inference(forward_demodulation,[],[f9039,f7676])).
fof(f9039,plain,(
  ( ! [X596,X595] : (p(s(bool,'const.bool.IN_2'(s(sK17,X596),s(fun(sK17,bool),sK18)))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,sK55(sK18,X595,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X595))) ) | $spl28),
  inference(forward_demodulation,[],[f9038,f956])).
fof(f9038,plain,(
  ( ! [X596,X595] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK55(sK18,X595,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,X596)))) | ~p(s(bool,X595))) ) | $spl28),
  inference(forward_demodulation,[],[f8368,f956])).
fof(f8368,plain,(
  ( ! [X596,X595] : (~p(s(bool,happ(s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(sK17,sK55(sK18,X595,sK17))))) | p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,X596)))) | ~p(s(bool,X595))) ) | $spl28),
  inference(superposition,[],[f907,f7676])).
fof(f9034,plain,(
  ( ! [X591] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK54(sK18,X591,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X591))) ) | $spl318),
  inference(cnf_transformation,[],[f9034_D])).
fof(f9034_D,plain,(
  ( ! [X591] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK54(sK18,X591,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X591))) ) <=> ~$spl318),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl318])])).
fof(f9037,plain,(
  ( ! [X594,X593] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK54(sK18,X594,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,'const.bool.IN_2'(s(sK17,X593),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X594))) ) | $spl28),
  inference(forward_demodulation,[],[f9036,f7676])).
fof(f9036,plain,(
  ( ! [X594,X593] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK54(sK18,X594,sK17)),s(fun(sK17,bool),sK18)))) | p(s(bool,'const.bool.IN_2'(s(sK17,X593),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X594))) ) | $spl28),
  inference(forward_demodulation,[],[f9035,f956])).
fof(f9035,plain,(
  ( ! [X594,X593] : (p(s(bool,'const.bool.IN_2'(s(sK17,X593),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,sK54(sK18,X594,sK17))))) | ~p(s(bool,X594))) ) | $spl28),
  inference(forward_demodulation,[],[f8367,f956])).
fof(f8367,plain,(
  ( ! [X594,X593] : (p(s(bool,happ(s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(sK17,X593)))) | ~p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,sK54(sK18,X594,sK17))))) | ~p(s(bool,X594))) ) | $spl28),
  inference(superposition,[],[f905,f7676])).
fof(f905,plain,(
  ( ! [X2,X0,X5,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK54(X2,X1,X0))))) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f620])).
fof(f9033,plain,(
  ( ! [X592,X591] : (p(s(bool,'const.bool.IN_2'(s(sK17,X592),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,sK54(sK18,X591,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X591))) ) | $spl28),
  inference(forward_demodulation,[],[f9032,f7676])).
fof(f9032,plain,(
  ( ! [X592,X591] : (p(s(bool,'const.bool.IN_2'(s(sK17,X592),s(fun(sK17,bool),sK18)))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,sK54(sK18,X591,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X591))) ) | $spl28),
  inference(forward_demodulation,[],[f9031,f956])).
fof(f9031,plain,(
  ( ! [X592,X591] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK54(sK18,X591,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,X592)))) | ~p(s(bool,X591))) ) | $spl28),
  inference(forward_demodulation,[],[f8366,f956])).
fof(f8366,plain,(
  ( ! [X592,X591] : (~p(s(bool,happ(s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(sK17,sK54(sK18,X591,sK17))))) | p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,X592)))) | ~p(s(bool,X591))) ) | $spl28),
  inference(superposition,[],[f905,f7676])).
fof(f9027,plain,(
  ( ! [X587] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK53(X587,sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X587))) ) | $spl316),
  inference(cnf_transformation,[],[f9027_D])).
fof(f9027_D,plain,(
  ( ! [X587] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK53(X587,sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X587))) ) <=> ~$spl316),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl316])])).
fof(f9030,plain,(
  ( ! [X590,X589] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK53(X590,sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,'const.bool.IN_2'(s(sK17,X589),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X590))) ) | $spl28),
  inference(forward_demodulation,[],[f9029,f7676])).
fof(f9029,plain,(
  ( ! [X590,X589] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK53(X590,sK18,sK17)),s(fun(sK17,bool),sK18)))) | p(s(bool,'const.bool.IN_2'(s(sK17,X589),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X590))) ) | $spl28),
  inference(forward_demodulation,[],[f9028,f956])).
fof(f9028,plain,(
  ( ! [X590,X589] : (p(s(bool,'const.bool.IN_2'(s(sK17,X589),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,sK53(X590,sK18,sK17))))) | ~p(s(bool,X590))) ) | $spl28),
  inference(forward_demodulation,[],[f8365,f956])).
fof(f8365,plain,(
  ( ! [X590,X589] : (p(s(bool,happ(s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(sK17,X589)))) | ~p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,sK53(X590,sK18,sK17))))) | ~p(s(bool,X590))) ) | $spl28),
  inference(superposition,[],[f903,f7676])).
fof(f903,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK53(X2,X1,X0))))) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f616])).
fof(f616,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,X2)) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK52(X2,X1,X0))))) | ! [X5] : (p(s(bool,X2)) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))))) & ((~p(s(bool,X2)) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK53(X2,X1,X0)))))) | (p(s(bool,X2)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK52,sK53])],[f615])).
fof(f615,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,X2)) | ? [X6] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) | ! [X5] : (p(s(bool,X2)) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))))) & (? [X4] : (~p(s(bool,X2)) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) | (p(s(bool,X2)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))))),
  inference(rectify,[],[f614])).
fof(f614,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,X2)) | ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ! [X4] : (p(s(bool,X2)) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))))) & (? [X4] : (~p(s(bool,X2)) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) | (p(s(bool,X2)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))))),
  inference(flattening,[],[f613])).
fof(f613,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X2)) | ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | ! [X4] : (p(s(bool,X2)) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))))) & (? [X4] : (~p(s(bool,X2)) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) | (p(s(bool,X2)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f362])).
fof(f362,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) <=> ! [X4] : (p(s(bool,X2)) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))))),
  inference(flattening,[],[f361])).
fof(f361,plain,(
  ! [X0] : ! [X1,X2] : ((p(s(bool,X2)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) <=> ! [X4] : (p(s(bool,X2)) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))))),
  inference(rectify,[],[f115])).
fof(f115,axiom,(
  ! [X7] : ! [X30,X33] : ((p(s(bool,X33)) & ! [X13] : p(s(bool,happ(s(fun(X7,bool),X30),s(X7,X13))))) <=> ! [X13] : (p(s(bool,X33)) & p(s(bool,happ(s(fun(X7,bool),X30),s(X7,X13))))))),
  file('thm.util_prob.INJ_IMAGE_BIJ','thm.bool.PULL_FORALL|split|2')).
fof(f9026,plain,(
  ( ! [X588,X587] : (p(s(bool,'const.bool.IN_2'(s(sK17,X588),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,sK53(X587,sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X587))) ) | $spl28),
  inference(forward_demodulation,[],[f9025,f7676])).
fof(f9025,plain,(
  ( ! [X588,X587] : (p(s(bool,'const.bool.IN_2'(s(sK17,X588),s(fun(sK17,bool),sK18)))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,sK53(X587,sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X587))) ) | $spl28),
  inference(forward_demodulation,[],[f9024,f956])).
fof(f9024,plain,(
  ( ! [X588,X587] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK53(X587,sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,X588)))) | ~p(s(bool,X587))) ) | $spl28),
  inference(forward_demodulation,[],[f8364,f956])).
fof(f8364,plain,(
  ( ! [X588,X587] : (~p(s(bool,happ(s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(sK17,sK53(X587,sK18,sK17))))) | p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,X588)))) | ~p(s(bool,X587))) ) | $spl28),
  inference(superposition,[],[f903,f7676])).
fof(f9020,plain,(
  ( ! [X583] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK52(X583,sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X583))) ) | $spl314),
  inference(cnf_transformation,[],[f9020_D])).
fof(f9020_D,plain,(
  ( ! [X583] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK52(X583,sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X583))) ) <=> ~$spl314),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl314])])).
fof(f9023,plain,(
  ( ! [X586,X585] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK52(X586,sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,'const.bool.IN_2'(s(sK17,X585),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X586))) ) | $spl28),
  inference(forward_demodulation,[],[f9022,f7676])).
fof(f9022,plain,(
  ( ! [X586,X585] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK52(X586,sK18,sK17)),s(fun(sK17,bool),sK18)))) | p(s(bool,'const.bool.IN_2'(s(sK17,X585),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X586))) ) | $spl28),
  inference(forward_demodulation,[],[f9021,f956])).
fof(f9021,plain,(
  ( ! [X586,X585] : (p(s(bool,'const.bool.IN_2'(s(sK17,X585),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,sK52(X586,sK18,sK17))))) | ~p(s(bool,X586))) ) | $spl28),
  inference(forward_demodulation,[],[f8363,f956])).
fof(f8363,plain,(
  ( ! [X586,X585] : (p(s(bool,happ(s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(sK17,X585)))) | ~p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,sK52(X586,sK18,sK17))))) | ~p(s(bool,X586))) ) | $spl28),
  inference(superposition,[],[f901,f7676])).
fof(f901,plain,(
  ( ! [X2,X0,X5,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK52(X2,X1,X0))))) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f616])).
fof(f9019,plain,(
  ( ! [X584,X583] : (p(s(bool,'const.bool.IN_2'(s(sK17,X584),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,sK52(X583,sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X583))) ) | $spl28),
  inference(forward_demodulation,[],[f9018,f7676])).
fof(f9018,plain,(
  ( ! [X584,X583] : (p(s(bool,'const.bool.IN_2'(s(sK17,X584),s(fun(sK17,bool),sK18)))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,sK52(X583,sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X583))) ) | $spl28),
  inference(forward_demodulation,[],[f9017,f956])).
fof(f9017,plain,(
  ( ! [X584,X583] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK52(X583,sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,X584)))) | ~p(s(bool,X583))) ) | $spl28),
  inference(forward_demodulation,[],[f8362,f956])).
fof(f8362,plain,(
  ( ! [X584,X583] : (~p(s(bool,happ(s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(sK17,sK52(X583,sK18,sK17))))) | p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,X584)))) | ~p(s(bool,X583))) ) | $spl28),
  inference(superposition,[],[f901,f7676])).
fof(f9013,plain,(
  ( ! [X579] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK51(X579,sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X579))) ) | $spl312),
  inference(cnf_transformation,[],[f9013_D])).
fof(f9013_D,plain,(
  ( ! [X579] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK51(X579,sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X579))) ) <=> ~$spl312),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl312])])).
fof(f9016,plain,(
  ( ! [X582,X581] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK51(X582,sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,'const.bool.IN_2'(s(sK17,X581),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X582))) ) | $spl28),
  inference(forward_demodulation,[],[f9015,f7676])).
fof(f9015,plain,(
  ( ! [X582,X581] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK51(X582,sK18,sK17)),s(fun(sK17,bool),sK18)))) | p(s(bool,'const.bool.IN_2'(s(sK17,X581),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X582))) ) | $spl28),
  inference(forward_demodulation,[],[f9014,f956])).
fof(f9014,plain,(
  ( ! [X582,X581] : (p(s(bool,'const.bool.IN_2'(s(sK17,X581),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X582)) | ~p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,sK51(X582,sK18,sK17)))))) ) | $spl28),
  inference(forward_demodulation,[],[f8361,f956])).
fof(f8361,plain,(
  ( ! [X582,X581] : (p(s(bool,happ(s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(sK17,X581)))) | ~p(s(bool,X582)) | ~p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,sK51(X582,sK18,sK17)))))) ) | $spl28),
  inference(superposition,[],[f898,f7676])).
fof(f898,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,X2)) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK51(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f612])).
fof(f612,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK50(X2,X1,X0))))) | ~p(s(bool,X2)) | ! [X5] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & p(s(bool,X2)))) & ((~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK51(X2,X1,X0))))) | ~p(s(bool,X2))) | (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK50,sK51])],[f611])).
fof(f611,plain,(
  ! [X0,X1,X2] : ((? [X6] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) | ~p(s(bool,X2)) | ! [X5] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & p(s(bool,X2)))) & (? [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,X2))) | (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2)))))),
  inference(rectify,[],[f610])).
fof(f610,plain,(
  ! [X0,X1,X2] : ((? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,X2)) | ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & p(s(bool,X2)))) & (? [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,X2))) | (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2)))))),
  inference(flattening,[],[f609])).
fof(f609,plain,(
  ! [X0,X1,X2] : (((? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,X2))) | ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & p(s(bool,X2)))) & (? [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,X2))) | (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2)))))),
  inference(nnf_transformation,[],[f360])).
fof(f360,plain,(
  ! [X0,X1,X2] : ((! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2))) <=> ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & p(s(bool,X2))))),
  inference(flattening,[],[f359])).
fof(f359,plain,(
  ! [X0] : ! [X1,X2] : ((! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2))) <=> ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & p(s(bool,X2))))),
  inference(rectify,[],[f92])).
fof(f92,axiom,(
  ! [X7] : ! [X30,X33] : ((! [X13] : p(s(bool,happ(s(fun(X7,bool),X30),s(X7,X13)))) & p(s(bool,X33))) <=> ! [X13] : (p(s(bool,happ(s(fun(X7,bool),X30),s(X7,X13)))) & p(s(bool,X33))))),
  file('thm.util_prob.INJ_IMAGE_BIJ','thm.bool.LEFT_AND_FORALL_THM')).
fof(f9012,plain,(
  ( ! [X580,X579] : (p(s(bool,'const.bool.IN_2'(s(sK17,X580),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,sK51(X579,sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X579))) ) | $spl28),
  inference(forward_demodulation,[],[f9011,f7676])).
fof(f9011,plain,(
  ( ! [X580,X579] : (p(s(bool,'const.bool.IN_2'(s(sK17,X580),s(fun(sK17,bool),sK18)))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,sK51(X579,sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X579))) ) | $spl28),
  inference(forward_demodulation,[],[f9010,f956])).
fof(f9010,plain,(
  ( ! [X580,X579] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK51(X579,sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,X580)))) | ~p(s(bool,X579))) ) | $spl28),
  inference(forward_demodulation,[],[f8360,f956])).
fof(f8360,plain,(
  ( ! [X580,X579] : (~p(s(bool,happ(s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(sK17,sK51(X579,sK18,sK17))))) | p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,X580)))) | ~p(s(bool,X579))) ) | $spl28),
  inference(superposition,[],[f898,f7676])).
fof(f9006,plain,(
  ( ! [X575] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK50(X575,sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X575))) ) | $spl310),
  inference(cnf_transformation,[],[f9006_D])).
fof(f9006_D,plain,(
  ( ! [X575] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK50(X575,sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X575))) ) <=> ~$spl310),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl310])])).
fof(f9009,plain,(
  ( ! [X578,X577] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK50(X578,sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,'const.bool.IN_2'(s(sK17,X577),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X578))) ) | $spl28),
  inference(forward_demodulation,[],[f9008,f7676])).
fof(f9008,plain,(
  ( ! [X578,X577] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK50(X578,sK18,sK17)),s(fun(sK17,bool),sK18)))) | p(s(bool,'const.bool.IN_2'(s(sK17,X577),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X578))) ) | $spl28),
  inference(forward_demodulation,[],[f9007,f956])).
fof(f9007,plain,(
  ( ! [X578,X577] : (p(s(bool,'const.bool.IN_2'(s(sK17,X577),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X578)) | ~p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,sK50(X578,sK18,sK17)))))) ) | $spl28),
  inference(forward_demodulation,[],[f8359,f956])).
fof(f8359,plain,(
  ( ! [X578,X577] : (p(s(bool,happ(s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(sK17,X577)))) | ~p(s(bool,X578)) | ~p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,sK50(X578,sK18,sK17)))))) ) | $spl28),
  inference(superposition,[],[f896,f7676])).
fof(f896,plain,(
  ( ! [X2,X0,X5,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ~p(s(bool,X2)) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK50(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f612])).
fof(f9005,plain,(
  ( ! [X575,X576] : (p(s(bool,'const.bool.IN_2'(s(sK17,X576),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,sK50(X575,sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X575))) ) | $spl28),
  inference(forward_demodulation,[],[f9004,f7676])).
fof(f9004,plain,(
  ( ! [X575,X576] : (p(s(bool,'const.bool.IN_2'(s(sK17,X576),s(fun(sK17,bool),sK18)))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,sK50(X575,sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X575))) ) | $spl28),
  inference(forward_demodulation,[],[f9003,f956])).
fof(f9003,plain,(
  ( ! [X575,X576] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK50(X575,sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,X576)))) | ~p(s(bool,X575))) ) | $spl28),
  inference(forward_demodulation,[],[f8358,f956])).
fof(f8358,plain,(
  ( ! [X575,X576] : (~p(s(bool,happ(s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(sK17,sK50(X575,sK18,sK17))))) | p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,X576)))) | ~p(s(bool,X575))) ) | $spl28),
  inference(superposition,[],[f896,f7676])).
fof(f8999,plain,(
  ( ! [X571] : (p(s(bool,'const.bool.IN_2'(s(sK17,sK49(sK18,X571,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X571))) ) | $spl308),
  inference(cnf_transformation,[],[f8999_D])).
fof(f8999_D,plain,(
  ( ! [X571] : (p(s(bool,'const.bool.IN_2'(s(sK17,sK49(sK18,X571,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X571))) ) <=> ~$spl308),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl308])])).
fof(f9002,plain,(
  ( ! [X573,X574] : (p(s(bool,'const.bool.IN_2'(s(sK17,sK49(sK18,X574,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,X573),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X574))) ) | $spl28),
  inference(forward_demodulation,[],[f9001,f7676])).
fof(f9001,plain,(
  ( ! [X573,X574] : (p(s(bool,'const.bool.IN_2'(s(sK17,sK49(sK18,X574,sK17)),s(fun(sK17,bool),sK18)))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,X573),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X574))) ) | $spl28),
  inference(forward_demodulation,[],[f9000,f956])).
fof(f9000,plain,(
  ( ! [X573,X574] : (~p(s(bool,'const.bool.IN_2'(s(sK17,X573),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,sK49(sK18,X574,sK17))))) | p(s(bool,X574))) ) | $spl28),
  inference(forward_demodulation,[],[f8357,f956])).
fof(f8357,plain,(
  ( ! [X573,X574] : (~p(s(bool,happ(s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(sK17,X573)))) | p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,sK49(sK18,X574,sK17))))) | p(s(bool,X574))) ) | $spl28),
  inference(superposition,[],[f895,f7676])).
fof(f895,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK49(X2,X1,X0))))) | p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) )),
  inference(cnf_transformation,[],[f608])).
fof(f608,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X6] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK48(X2,X1,X0))))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK49(X2,X1,X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK48,sK49])],[f607])).
fof(f607,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X6] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | ? [X5] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(rectify,[],[f606])).
fof(f606,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(flattening,[],[f605])).
fof(f605,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f358])).
fof(f358,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f357])).
fof(f357,plain,(
  ! [X0] : ! [X1,X2] : ((p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f140])).
fof(f140,axiom,(
  ! [X7] : ! [X30,X33] : ((p(s(bool,X30)) | ? [X13] : p(s(bool,happ(s(fun(X7,bool),X33),s(X7,X13))))) <=> ? [X13] : (p(s(bool,X30)) | p(s(bool,happ(s(fun(X7,bool),X33),s(X7,X13))))))),
  file('thm.util_prob.INJ_IMAGE_BIJ','thm.bool.RIGHT_OR_EXISTS_THM')).
fof(f8998,plain,(
  ( ! [X571,X572] : (~p(s(bool,'const.bool.IN_2'(s(sK17,X572),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,'const.bool.IN_2'(s(sK17,sK49(sK18,X571,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X571))) ) | $spl28),
  inference(forward_demodulation,[],[f8997,f7676])).
fof(f8997,plain,(
  ( ! [X571,X572] : (~p(s(bool,'const.bool.IN_2'(s(sK17,X572),s(fun(sK17,bool),sK18)))) | p(s(bool,'const.bool.IN_2'(s(sK17,sK49(sK18,X571,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X571))) ) | $spl28),
  inference(forward_demodulation,[],[f8996,f956])).
fof(f8996,plain,(
  ( ! [X571,X572] : (p(s(bool,'const.bool.IN_2'(s(sK17,sK49(sK18,X571,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X571)) | ~p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,X572))))) ) | $spl28),
  inference(forward_demodulation,[],[f8356,f956])).
fof(f8356,plain,(
  ( ! [X571,X572] : (p(s(bool,happ(s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(sK17,sK49(sK18,X571,sK17))))) | p(s(bool,X571)) | ~p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,X572))))) ) | $spl28),
  inference(superposition,[],[f895,f7676])).
fof(f8992,plain,(
  ( ! [X567] : (p(s(bool,'const.bool.IN_2'(s(sK17,sK48(sK18,X567,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X567))) ) | $spl306),
  inference(cnf_transformation,[],[f8992_D])).
fof(f8992_D,plain,(
  ( ! [X567] : (p(s(bool,'const.bool.IN_2'(s(sK17,sK48(sK18,X567,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X567))) ) <=> ~$spl306),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl306])])).
fof(f8995,plain,(
  ( ! [X569,X570] : (p(s(bool,'const.bool.IN_2'(s(sK17,sK48(sK18,X570,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,X569),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X570))) ) | $spl28),
  inference(forward_demodulation,[],[f8994,f7676])).
fof(f8994,plain,(
  ( ! [X569,X570] : (p(s(bool,'const.bool.IN_2'(s(sK17,sK48(sK18,X570,sK17)),s(fun(sK17,bool),sK18)))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,X569),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X570))) ) | $spl28),
  inference(forward_demodulation,[],[f8993,f956])).
fof(f8993,plain,(
  ( ! [X569,X570] : (~p(s(bool,'const.bool.IN_2'(s(sK17,X569),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,sK48(sK18,X570,sK17))))) | p(s(bool,X570))) ) | $spl28),
  inference(forward_demodulation,[],[f8355,f956])).
fof(f8355,plain,(
  ( ! [X569,X570] : (~p(s(bool,happ(s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(sK17,X569)))) | p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,sK48(sK18,X570,sK17))))) | p(s(bool,X570))) ) | $spl28),
  inference(superposition,[],[f893,f7676])).
fof(f893,plain,(
  ( ! [X6,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK48(X2,X1,X0))))) | p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) )),
  inference(cnf_transformation,[],[f608])).
fof(f8991,plain,(
  ( ! [X567,X568] : (~p(s(bool,'const.bool.IN_2'(s(sK17,X568),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,'const.bool.IN_2'(s(sK17,sK48(sK18,X567,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X567))) ) | $spl28),
  inference(forward_demodulation,[],[f8990,f7676])).
fof(f8990,plain,(
  ( ! [X567,X568] : (~p(s(bool,'const.bool.IN_2'(s(sK17,X568),s(fun(sK17,bool),sK18)))) | p(s(bool,'const.bool.IN_2'(s(sK17,sK48(sK18,X567,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X567))) ) | $spl28),
  inference(forward_demodulation,[],[f8989,f956])).
fof(f8989,plain,(
  ( ! [X567,X568] : (p(s(bool,'const.bool.IN_2'(s(sK17,sK48(sK18,X567,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X567)) | ~p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,X568))))) ) | $spl28),
  inference(forward_demodulation,[],[f8354,f956])).
fof(f8354,plain,(
  ( ! [X567,X568] : (p(s(bool,happ(s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(sK17,sK48(sK18,X567,sK17))))) | p(s(bool,X567)) | ~p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,X568))))) ) | $spl28),
  inference(superposition,[],[f893,f7676])).
fof(f8985,plain,(
  ( ! [X563] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK47(X563,sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X563))) ) | $spl304),
  inference(cnf_transformation,[],[f8985_D])).
fof(f8985_D,plain,(
  ( ! [X563] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK47(X563,sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X563))) ) <=> ~$spl304),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl304])])).
fof(f8988,plain,(
  ( ! [X565,X566] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK47(X566,sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,'const.bool.IN_2'(s(sK17,X565),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X566))) ) | $spl28),
  inference(forward_demodulation,[],[f8987,f7676])).
fof(f8987,plain,(
  ( ! [X565,X566] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK47(X566,sK18,sK17)),s(fun(sK17,bool),sK18)))) | p(s(bool,'const.bool.IN_2'(s(sK17,X565),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X566))) ) | $spl28),
  inference(forward_demodulation,[],[f8986,f956])).
fof(f8986,plain,(
  ( ! [X565,X566] : (p(s(bool,'const.bool.IN_2'(s(sK17,X565),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X566)) | ~p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,sK47(X566,sK18,sK17)))))) ) | $spl28),
  inference(forward_demodulation,[],[f8353,f956])).
fof(f8353,plain,(
  ( ! [X565,X566] : (p(s(bool,happ(s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(sK17,X565)))) | ~p(s(bool,X566)) | ~p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,sK47(X566,sK18,sK17)))))) ) | $spl28),
  inference(superposition,[],[f891,f7676])).
fof(f891,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,X2)) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK47(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f604])).
fof(f604,plain,(
  ! [X0,X1,X2] : (((p(s(bool,X2)) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK46(X2,X1,X0)))))) | ! [X5] : (~p(s(bool,X2)) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))))) & ((p(s(bool,X2)) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK47(X2,X1,X0)))))) | ~p(s(bool,X2)) | ! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK46,sK47])],[f603])).
fof(f603,plain,(
  ! [X0,X1,X2] : (((p(s(bool,X2)) & ? [X6] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6))))) | ! [X5] : (~p(s(bool,X2)) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))))) & (? [X4] : (p(s(bool,X2)) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) | ~p(s(bool,X2)) | ! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))))),
  inference(rectify,[],[f602])).
fof(f602,plain,(
  ! [X0,X1,X2] : (((p(s(bool,X2)) & ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | ! [X4] : (~p(s(bool,X2)) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))))) & (? [X4] : (p(s(bool,X2)) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) | ~p(s(bool,X2)) | ! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))))),
  inference(flattening,[],[f601])).
fof(f601,plain,(
  ! [X0,X1,X2] : (((p(s(bool,X2)) & ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | ! [X4] : (~p(s(bool,X2)) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))))) & (? [X4] : (p(s(bool,X2)) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) | (~p(s(bool,X2)) | ! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f443])).
fof(f443,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,X2)) | ! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) <=> ! [X4] : (~p(s(bool,X2)) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))))),
  inference(ennf_transformation,[],[f356])).
fof(f356,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) => ! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) <=> ! [X4] : (p(s(bool,X2)) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))))),
  inference(flattening,[],[f355])).
fof(f355,plain,(
  ! [X0] : ! [X1,X2] : ((p(s(bool,X2)) => ! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) <=> ! [X4] : (p(s(bool,X2)) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))))),
  inference(rectify,[],[f117])).
fof(f117,axiom,(
  ! [X7] : ! [X30,X33] : ((p(s(bool,X33)) => ! [X13] : p(s(bool,happ(s(fun(X7,bool),X30),s(X7,X13))))) <=> ! [X13] : (p(s(bool,X33)) => p(s(bool,happ(s(fun(X7,bool),X30),s(X7,X13))))))),
  file('thm.util_prob.INJ_IMAGE_BIJ','thm.bool.PULL_FORALL|split|0')).
fof(f8984,plain,(
  ( ! [X563,X564] : (p(s(bool,'const.bool.IN_2'(s(sK17,X564),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,sK47(X563,sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X563))) ) | $spl28),
  inference(forward_demodulation,[],[f8983,f7676])).
fof(f8983,plain,(
  ( ! [X563,X564] : (p(s(bool,'const.bool.IN_2'(s(sK17,X564),s(fun(sK17,bool),sK18)))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,sK47(X563,sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X563))) ) | $spl28),
  inference(forward_demodulation,[],[f8982,f956])).
fof(f8982,plain,(
  ( ! [X563,X564] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK47(X563,sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,X564)))) | ~p(s(bool,X563))) ) | $spl28),
  inference(forward_demodulation,[],[f8352,f956])).
fof(f8352,plain,(
  ( ! [X563,X564] : (~p(s(bool,happ(s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(sK17,sK47(X563,sK18,sK17))))) | p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,X564)))) | ~p(s(bool,X563))) ) | $spl28),
  inference(superposition,[],[f891,f7676])).
fof(f8978,plain,(
  ( ! [X559] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK46(X559,sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X559))) ) | $spl302),
  inference(cnf_transformation,[],[f8978_D])).
fof(f8978_D,plain,(
  ( ! [X559] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK46(X559,sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X559))) ) <=> ~$spl302),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl302])])).
fof(f8981,plain,(
  ( ! [X561,X562] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK46(X562,sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,'const.bool.IN_2'(s(sK17,X561),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X562))) ) | $spl28),
  inference(forward_demodulation,[],[f8980,f7676])).
fof(f8980,plain,(
  ( ! [X561,X562] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK46(X562,sK18,sK17)),s(fun(sK17,bool),sK18)))) | p(s(bool,'const.bool.IN_2'(s(sK17,X561),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X562))) ) | $spl28),
  inference(forward_demodulation,[],[f8979,f956])).
fof(f8979,plain,(
  ( ! [X561,X562] : (p(s(bool,'const.bool.IN_2'(s(sK17,X561),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X562)) | ~p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,sK46(X562,sK18,sK17)))))) ) | $spl28),
  inference(forward_demodulation,[],[f8351,f956])).
fof(f8351,plain,(
  ( ! [X561,X562] : (p(s(bool,happ(s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(sK17,X561)))) | ~p(s(bool,X562)) | ~p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,sK46(X562,sK18,sK17)))))) ) | $spl28),
  inference(superposition,[],[f889,f7676])).
fof(f889,plain,(
  ( ! [X2,X0,X5,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ~p(s(bool,X2)) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK46(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f604])).
fof(f8977,plain,(
  ( ! [X559,X560] : (p(s(bool,'const.bool.IN_2'(s(sK17,X560),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,sK46(X559,sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X559))) ) | $spl28),
  inference(forward_demodulation,[],[f8976,f7676])).
fof(f8976,plain,(
  ( ! [X559,X560] : (p(s(bool,'const.bool.IN_2'(s(sK17,X560),s(fun(sK17,bool),sK18)))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,sK46(X559,sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X559))) ) | $spl28),
  inference(forward_demodulation,[],[f8975,f956])).
fof(f8975,plain,(
  ( ! [X559,X560] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK46(X559,sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,X560)))) | ~p(s(bool,X559))) ) | $spl28),
  inference(forward_demodulation,[],[f8350,f956])).
fof(f8350,plain,(
  ( ! [X559,X560] : (~p(s(bool,happ(s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(sK17,sK46(X559,sK18,sK17))))) | p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,X560)))) | ~p(s(bool,X559))) ) | $spl28),
  inference(superposition,[],[f889,f7676])).
fof(f8970,plain,(
  ( ! [X555] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK44(sK18,X555,sK17)),s(fun(sK17,bool),X555)))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,sK45(sK18,X555,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl300),
  inference(cnf_transformation,[],[f8970_D])).
fof(f8970_D,plain,(
  ( ! [X555] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK44(sK18,X555,sK17)),s(fun(sK17,bool),X555)))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,sK45(sK18,X555,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) <=> ~$spl300),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl300])])).
fof(f8974,plain,(
  ( ! [X557,X558] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK44(sK18,X558,sK17)),s(fun(sK17,bool),X558)))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,sK45(sK18,X558,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,'const.bool.IN_2'(s(sK17,X557),s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl28),
  inference(forward_demodulation,[],[f8973,f956])).
fof(f8973,plain,(
  ( ! [X557,X558] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK45(sK18,X558,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,'const.bool.IN_2'(s(sK17,X557),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,happ(s(fun(sK17,bool),X558),s(sK17,sK44(sK18,X558,sK17)))))) ) | $spl28),
  inference(forward_demodulation,[],[f8972,f7676])).
fof(f8972,plain,(
  ( ! [X557,X558] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK45(sK18,X558,sK17)),s(fun(sK17,bool),sK18)))) | p(s(bool,'const.bool.IN_2'(s(sK17,X557),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,happ(s(fun(sK17,bool),X558),s(sK17,sK44(sK18,X558,sK17)))))) ) | $spl28),
  inference(forward_demodulation,[],[f8971,f956])).
fof(f8971,plain,(
  ( ! [X557,X558] : (p(s(bool,'const.bool.IN_2'(s(sK17,X557),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,sK45(sK18,X558,sK17))))) | ~p(s(bool,happ(s(fun(sK17,bool),X558),s(sK17,sK44(sK18,X558,sK17)))))) ) | $spl28),
  inference(forward_demodulation,[],[f8349,f956])).
fof(f8349,plain,(
  ( ! [X557,X558] : (p(s(bool,happ(s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(sK17,X557)))) | ~p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,sK45(sK18,X558,sK17))))) | ~p(s(bool,happ(s(fun(sK17,bool),X558),s(sK17,sK44(sK18,X558,sK17)))))) ) | $spl28),
  inference(superposition,[],[f887,f7676])).
fof(f887,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK45(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK44(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f600])).
fof(f600,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK43(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK43(X2,X1,X0)))))) | (! [X7] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X7)))) & ! [X6] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK44(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK45(X2,X1,X0))))) | ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK43,sK44,sK45])],[f599])).
fof(f599,plain,(
  ! [X0,X1,X2] : ((? [X8] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X8)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X8))))) | (! [X7] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X7)))) & ! [X6] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))))) & (? [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(rectify,[],[f598])).
fof(f598,plain,(
  ! [X0,X1,X2] : ((? [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (? [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(flattening,[],[f597])).
fof(f597,plain,(
  ! [X0,X1,X2] : ((? [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & ((? [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f354])).
fof(f354,plain,(
  ! [X0,X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f353])).
fof(f353,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f100])).
fof(f100,axiom,(
  ! [X7] : ! [X30,X33] : (! [X13] : (p(s(bool,happ(s(fun(X7,bool),X30),s(X7,X13)))) & p(s(bool,happ(s(fun(X7,bool),X33),s(X7,X13))))) <=> (! [X13] : p(s(bool,happ(s(fun(X7,bool),X30),s(X7,X13)))) & ! [X13] : p(s(bool,happ(s(fun(X7,bool),X33),s(X7,X13))))))),
  file('thm.util_prob.INJ_IMAGE_BIJ','thm.bool.FORALL_AND_THM')).
fof(f8969,plain,(
  ( ! [X555,X556] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK44(sK18,X555,sK17)),s(fun(sK17,bool),X555)))) | p(s(bool,'const.bool.IN_2'(s(sK17,X556),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,sK45(sK18,X555,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl28),
  inference(forward_demodulation,[],[f8968,f956])).
fof(f8968,plain,(
  ( ! [X555,X556] : (p(s(bool,'const.bool.IN_2'(s(sK17,X556),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,sK45(sK18,X555,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,happ(s(fun(sK17,bool),X555),s(sK17,sK44(sK18,X555,sK17)))))) ) | $spl28),
  inference(forward_demodulation,[],[f8967,f7676])).
fof(f8967,plain,(
  ( ! [X555,X556] : (p(s(bool,'const.bool.IN_2'(s(sK17,X556),s(fun(sK17,bool),sK18)))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,sK45(sK18,X555,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,happ(s(fun(sK17,bool),X555),s(sK17,sK44(sK18,X555,sK17)))))) ) | $spl28),
  inference(forward_demodulation,[],[f8966,f956])).
fof(f8966,plain,(
  ( ! [X555,X556] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK45(sK18,X555,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,X556)))) | ~p(s(bool,happ(s(fun(sK17,bool),X555),s(sK17,sK44(sK18,X555,sK17)))))) ) | $spl28),
  inference(forward_demodulation,[],[f8348,f956])).
fof(f8348,plain,(
  ( ! [X555,X556] : (~p(s(bool,happ(s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(sK17,sK45(sK18,X555,sK17))))) | p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,X556)))) | ~p(s(bool,happ(s(fun(sK17,bool),X555),s(sK17,sK44(sK18,X555,sK17)))))) ) | $spl28),
  inference(superposition,[],[f887,f7676])).
fof(f8955,plain,(
  ( ! [X547] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK45(X547,sK18,sK17)),s(fun(sK17,bool),X547)))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,sK44(X547,sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl298),
  inference(cnf_transformation,[],[f8955_D])).
fof(f8955_D,plain,(
  ( ! [X547] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK45(X547,sK18,sK17)),s(fun(sK17,bool),X547)))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,sK44(X547,sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) <=> ~$spl298),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl298])])).
fof(f8962,plain,(
  ( ! [X551,X552] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK44(X552,sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,sK45(X552,sK18,sK17)),s(fun(sK17,bool),X552)))) | p(s(bool,'const.bool.IN_2'(s(sK17,X551),s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl28),
  inference(forward_demodulation,[],[f8961,f7676])).
fof(f8961,plain,(
  ( ! [X551,X552] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK44(X552,sK18,sK17)),s(fun(sK17,bool),sK18)))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,sK45(X552,sK18,sK17)),s(fun(sK17,bool),X552)))) | p(s(bool,'const.bool.IN_2'(s(sK17,X551),s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl28),
  inference(forward_demodulation,[],[f8960,f956])).
fof(f8960,plain,(
  ( ! [X551,X552] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK45(X552,sK18,sK17)),s(fun(sK17,bool),X552)))) | p(s(bool,'const.bool.IN_2'(s(sK17,X551),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,sK44(X552,sK18,sK17)))))) ) | $spl28),
  inference(forward_demodulation,[],[f8959,f956])).
fof(f8959,plain,(
  ( ! [X551,X552] : (p(s(bool,'const.bool.IN_2'(s(sK17,X551),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,happ(s(fun(sK17,bool),X552),s(sK17,sK45(X552,sK18,sK17))))) | ~p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,sK44(X552,sK18,sK17)))))) ) | $spl28),
  inference(forward_demodulation,[],[f8346,f956])).
fof(f8346,plain,(
  ( ! [X551,X552] : (p(s(bool,happ(s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(sK17,X551)))) | ~p(s(bool,happ(s(fun(sK17,bool),X552),s(sK17,sK45(X552,sK18,sK17))))) | ~p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,sK44(X552,sK18,sK17)))))) ) | $spl28),
  inference(superposition,[],[f886,f7676])).
fof(f886,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK45(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK44(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f600])).
fof(f8954,plain,(
  ( ! [X547,X548] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK45(X547,sK18,sK17)),s(fun(sK17,bool),X547)))) | p(s(bool,'const.bool.IN_2'(s(sK17,X548),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,sK44(X547,sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl28),
  inference(forward_demodulation,[],[f8953,f956])).
fof(f8953,plain,(
  ( ! [X547,X548] : (p(s(bool,'const.bool.IN_2'(s(sK17,X548),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,sK44(X547,sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,happ(s(fun(sK17,bool),X547),s(sK17,sK45(X547,sK18,sK17)))))) ) | $spl28),
  inference(forward_demodulation,[],[f8952,f7676])).
fof(f8952,plain,(
  ( ! [X547,X548] : (p(s(bool,'const.bool.IN_2'(s(sK17,X548),s(fun(sK17,bool),sK18)))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,sK44(X547,sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,happ(s(fun(sK17,bool),X547),s(sK17,sK45(X547,sK18,sK17)))))) ) | $spl28),
  inference(forward_demodulation,[],[f8951,f956])).
fof(f8951,plain,(
  ( ! [X547,X548] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK44(X547,sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,X548)))) | ~p(s(bool,happ(s(fun(sK17,bool),X547),s(sK17,sK45(X547,sK18,sK17)))))) ) | $spl28),
  inference(forward_demodulation,[],[f8344,f956])).
fof(f8344,plain,(
  ( ! [X547,X548] : (~p(s(bool,happ(s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(sK17,sK44(X547,sK18,sK17))))) | p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,X548)))) | ~p(s(bool,happ(s(fun(sK17,bool),X547),s(sK17,sK45(X547,sK18,sK17)))))) ) | $spl28),
  inference(superposition,[],[f886,f7676])).
fof(f8946,plain,(
  ( ! [X543] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK43(sK18,X543,sK17)),s(fun(sK17,bool),X543)))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,sK43(sK18,X543,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl296),
  inference(cnf_transformation,[],[f8946_D])).
fof(f8946_D,plain,(
  ( ! [X543] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK43(sK18,X543,sK17)),s(fun(sK17,bool),X543)))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,sK43(sK18,X543,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) <=> ~$spl296),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl296])])).
fof(f8950,plain,(
  ( ! [X546,X545] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK43(sK18,X546,sK17)),s(fun(sK17,bool),X546)))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,sK43(sK18,X546,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,'const.bool.IN_2'(s(sK17,X545),s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl28),
  inference(forward_demodulation,[],[f8949,f956])).
fof(f8949,plain,(
  ( ! [X546,X545] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK43(sK18,X546,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,'const.bool.IN_2'(s(sK17,X545),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,happ(s(fun(sK17,bool),X546),s(sK17,sK43(sK18,X546,sK17)))))) ) | $spl28),
  inference(forward_demodulation,[],[f8948,f7676])).
fof(f8948,plain,(
  ( ! [X546,X545] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK43(sK18,X546,sK17)),s(fun(sK17,bool),sK18)))) | p(s(bool,'const.bool.IN_2'(s(sK17,X545),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,happ(s(fun(sK17,bool),X546),s(sK17,sK43(sK18,X546,sK17)))))) ) | $spl28),
  inference(forward_demodulation,[],[f8947,f956])).
fof(f8947,plain,(
  ( ! [X546,X545] : (p(s(bool,'const.bool.IN_2'(s(sK17,X545),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,sK43(sK18,X546,sK17))))) | ~p(s(bool,happ(s(fun(sK17,bool),X546),s(sK17,sK43(sK18,X546,sK17)))))) ) | $spl28),
  inference(forward_demodulation,[],[f8343,f956])).
fof(f8343,plain,(
  ( ! [X546,X545] : (p(s(bool,happ(s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(sK17,X545)))) | ~p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,sK43(sK18,X546,sK17))))) | ~p(s(bool,happ(s(fun(sK17,bool),X546),s(sK17,sK43(sK18,X546,sK17)))))) ) | $spl28),
  inference(superposition,[],[f885,f7676])).
fof(f885,plain,(
  ( ! [X6,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK43(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK43(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f600])).
fof(f8945,plain,(
  ( ! [X543,X544] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK43(sK18,X543,sK17)),s(fun(sK17,bool),X543)))) | p(s(bool,'const.bool.IN_2'(s(sK17,X544),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,sK43(sK18,X543,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl28),
  inference(forward_demodulation,[],[f8944,f956])).
fof(f8944,plain,(
  ( ! [X543,X544] : (p(s(bool,'const.bool.IN_2'(s(sK17,X544),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,sK43(sK18,X543,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,happ(s(fun(sK17,bool),X543),s(sK17,sK43(sK18,X543,sK17)))))) ) | $spl28),
  inference(forward_demodulation,[],[f8943,f7676])).
fof(f8943,plain,(
  ( ! [X543,X544] : (p(s(bool,'const.bool.IN_2'(s(sK17,X544),s(fun(sK17,bool),sK18)))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,sK43(sK18,X543,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,happ(s(fun(sK17,bool),X543),s(sK17,sK43(sK18,X543,sK17)))))) ) | $spl28),
  inference(forward_demodulation,[],[f8942,f956])).
fof(f8942,plain,(
  ( ! [X543,X544] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK43(sK18,X543,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,X544)))) | ~p(s(bool,happ(s(fun(sK17,bool),X543),s(sK17,sK43(sK18,X543,sK17)))))) ) | $spl28),
  inference(forward_demodulation,[],[f8342,f956])).
fof(f8342,plain,(
  ( ! [X543,X544] : (~p(s(bool,happ(s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(sK17,sK43(sK18,X543,sK17))))) | p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,X544)))) | ~p(s(bool,happ(s(fun(sK17,bool),X543),s(sK17,sK43(sK18,X543,sK17)))))) ) | $spl28),
  inference(superposition,[],[f885,f7676])).
fof(f8931,plain,(
  ( ! [X535] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK43(X535,sK18,sK17)),s(fun(sK17,bool),X535)))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,sK43(X535,sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl294),
  inference(cnf_transformation,[],[f8931_D])).
fof(f8931_D,plain,(
  ( ! [X535] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK43(X535,sK18,sK17)),s(fun(sK17,bool),X535)))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,sK43(X535,sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) <=> ~$spl294),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl294])])).
fof(f8938,plain,(
  ( ! [X540,X539] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK43(X540,sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,sK43(X540,sK18,sK17)),s(fun(sK17,bool),X540)))) | p(s(bool,'const.bool.IN_2'(s(sK17,X539),s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl28),
  inference(forward_demodulation,[],[f8937,f7676])).
fof(f8937,plain,(
  ( ! [X540,X539] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK43(X540,sK18,sK17)),s(fun(sK17,bool),sK18)))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,sK43(X540,sK18,sK17)),s(fun(sK17,bool),X540)))) | p(s(bool,'const.bool.IN_2'(s(sK17,X539),s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl28),
  inference(forward_demodulation,[],[f8936,f956])).
fof(f8936,plain,(
  ( ! [X540,X539] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK43(X540,sK18,sK17)),s(fun(sK17,bool),X540)))) | p(s(bool,'const.bool.IN_2'(s(sK17,X539),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,sK43(X540,sK18,sK17)))))) ) | $spl28),
  inference(forward_demodulation,[],[f8935,f956])).
fof(f8935,plain,(
  ( ! [X540,X539] : (p(s(bool,'const.bool.IN_2'(s(sK17,X539),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,happ(s(fun(sK17,bool),X540),s(sK17,sK43(X540,sK18,sK17))))) | ~p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,sK43(X540,sK18,sK17)))))) ) | $spl28),
  inference(forward_demodulation,[],[f8340,f956])).
fof(f8340,plain,(
  ( ! [X540,X539] : (p(s(bool,happ(s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(sK17,X539)))) | ~p(s(bool,happ(s(fun(sK17,bool),X540),s(sK17,sK43(X540,sK18,sK17))))) | ~p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,sK43(X540,sK18,sK17)))))) ) | $spl28),
  inference(superposition,[],[f884,f7676])).
fof(f884,plain,(
  ( ! [X2,X0,X7,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X7)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK43(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK43(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f600])).
fof(f8930,plain,(
  ( ! [X536,X535] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK43(X535,sK18,sK17)),s(fun(sK17,bool),X535)))) | p(s(bool,'const.bool.IN_2'(s(sK17,X536),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,sK43(X535,sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0'))))) ) | $spl28),
  inference(forward_demodulation,[],[f8929,f956])).
fof(f8929,plain,(
  ( ! [X536,X535] : (p(s(bool,'const.bool.IN_2'(s(sK17,X536),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,sK43(X535,sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,happ(s(fun(sK17,bool),X535),s(sK17,sK43(X535,sK18,sK17)))))) ) | $spl28),
  inference(forward_demodulation,[],[f8928,f7676])).
fof(f8928,plain,(
  ( ! [X536,X535] : (p(s(bool,'const.bool.IN_2'(s(sK17,X536),s(fun(sK17,bool),sK18)))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,sK43(X535,sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,happ(s(fun(sK17,bool),X535),s(sK17,sK43(X535,sK18,sK17)))))) ) | $spl28),
  inference(forward_demodulation,[],[f8927,f956])).
fof(f8927,plain,(
  ( ! [X536,X535] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK43(X535,sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,X536)))) | ~p(s(bool,happ(s(fun(sK17,bool),X535),s(sK17,sK43(X535,sK18,sK17)))))) ) | $spl28),
  inference(forward_demodulation,[],[f8338,f956])).
fof(f8338,plain,(
  ( ! [X536,X535] : (~p(s(bool,happ(s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(sK17,sK43(X535,sK18,sK17))))) | p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,X536)))) | ~p(s(bool,happ(s(fun(sK17,bool),X535),s(sK17,sK43(X535,sK18,sK17)))))) ) | $spl28),
  inference(superposition,[],[f884,f7676])).
fof(f8923,plain,(
  ( ! [X531] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK42(sK18,X531,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X531))) ) | $spl292),
  inference(cnf_transformation,[],[f8923_D])).
fof(f8923_D,plain,(
  ( ! [X531] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK42(sK18,X531,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X531))) ) <=> ~$spl292),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl292])])).
fof(f8926,plain,(
  ( ! [X534,X533] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK42(sK18,X534,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,'const.bool.IN_2'(s(sK17,X533),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X534))) ) | $spl28),
  inference(forward_demodulation,[],[f8925,f7676])).
fof(f8925,plain,(
  ( ! [X534,X533] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK42(sK18,X534,sK17)),s(fun(sK17,bool),sK18)))) | p(s(bool,'const.bool.IN_2'(s(sK17,X533),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X534))) ) | $spl28),
  inference(forward_demodulation,[],[f8924,f956])).
fof(f8924,plain,(
  ( ! [X534,X533] : (p(s(bool,'const.bool.IN_2'(s(sK17,X533),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X534)) | ~p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,sK42(sK18,X534,sK17)))))) ) | $spl28),
  inference(forward_demodulation,[],[f8337,f956])).
fof(f8337,plain,(
  ( ! [X534,X533] : (p(s(bool,happ(s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(sK17,X533)))) | p(s(bool,X534)) | ~p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,sK42(sK18,X534,sK17)))))) ) | $spl28),
  inference(superposition,[],[f883,f7676])).
fof(f883,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK42(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f596])).
fof(f596,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK41(X2,X1,X0)))))) | p(s(bool,X1)) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5))))) & ((~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK42(X2,X1,X0)))))) | ! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK41,sK42])],[f595])).
fof(f595,plain,(
  ! [X0,X1,X2] : ((? [X6] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | p(s(bool,X1)) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5))))) & ((~p(s(bool,X1)) & ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(rectify,[],[f594])).
fof(f594,plain,(
  ! [X0,X1,X2] : ((? [X3] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | p(s(bool,X1)) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) & ((~p(s(bool,X1)) & ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(flattening,[],[f593])).
fof(f593,plain,(
  ! [X0,X1,X2] : ((? [X3] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (p(s(bool,X1)) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & ((~p(s(bool,X1)) & ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f352])).
fof(f352,plain,(
  ! [X0,X1,X2] : (! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (p(s(bool,X1)) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f351])).
fof(f351,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (p(s(bool,X1)) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f93])).
fof(f93,axiom,(
  ! [X7] : ! [X30,X33] : (! [X13] : (p(s(bool,X30)) | p(s(bool,happ(s(fun(X7,bool),X33),s(X7,X13))))) <=> (p(s(bool,X30)) | ! [X13] : p(s(bool,happ(s(fun(X7,bool),X33),s(X7,X13))))))),
  file('thm.util_prob.INJ_IMAGE_BIJ','thm.bool.RIGHT_FORALL_OR_THM')).
fof(f8922,plain,(
  ( ! [X532,X531] : (p(s(bool,'const.bool.IN_2'(s(sK17,X532),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,sK42(sK18,X531,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X531))) ) | $spl28),
  inference(forward_demodulation,[],[f8921,f7676])).
fof(f8921,plain,(
  ( ! [X532,X531] : (p(s(bool,'const.bool.IN_2'(s(sK17,X532),s(fun(sK17,bool),sK18)))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,sK42(sK18,X531,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X531))) ) | $spl28),
  inference(forward_demodulation,[],[f8920,f956])).
fof(f8920,plain,(
  ( ! [X532,X531] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK42(sK18,X531,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,X532)))) | p(s(bool,X531))) ) | $spl28),
  inference(forward_demodulation,[],[f8336,f956])).
fof(f8336,plain,(
  ( ! [X532,X531] : (~p(s(bool,happ(s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(sK17,sK42(sK18,X531,sK17))))) | p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,X532)))) | p(s(bool,X531))) ) | $spl28),
  inference(superposition,[],[f883,f7676])).
fof(f8916,plain,(
  ( ! [X527] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK41(sK18,X527,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X527))) ) | $spl290),
  inference(cnf_transformation,[],[f8916_D])).
fof(f8916_D,plain,(
  ( ! [X527] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK41(sK18,X527,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X527))) ) <=> ~$spl290),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl290])])).
fof(f8919,plain,(
  ( ! [X530,X529] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK41(sK18,X530,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,'const.bool.IN_2'(s(sK17,X529),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X530))) ) | $spl28),
  inference(forward_demodulation,[],[f8918,f7676])).
fof(f8918,plain,(
  ( ! [X530,X529] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK41(sK18,X530,sK17)),s(fun(sK17,bool),sK18)))) | p(s(bool,'const.bool.IN_2'(s(sK17,X529),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X530))) ) | $spl28),
  inference(forward_demodulation,[],[f8917,f956])).
fof(f8917,plain,(
  ( ! [X530,X529] : (p(s(bool,'const.bool.IN_2'(s(sK17,X529),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X530)) | ~p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,sK41(sK18,X530,sK17)))))) ) | $spl28),
  inference(forward_demodulation,[],[f8335,f956])).
fof(f8335,plain,(
  ( ! [X530,X529] : (p(s(bool,happ(s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(sK17,X529)))) | p(s(bool,X530)) | ~p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,sK41(sK18,X530,sK17)))))) ) | $spl28),
  inference(superposition,[],[f881,f7676])).
fof(f881,plain,(
  ( ! [X2,X0,X5,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) | p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK41(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f596])).
fof(f8915,plain,(
  ( ! [X528,X527] : (p(s(bool,'const.bool.IN_2'(s(sK17,X528),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,sK41(sK18,X527,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X527))) ) | $spl28),
  inference(forward_demodulation,[],[f8914,f7676])).
fof(f8914,plain,(
  ( ! [X528,X527] : (p(s(bool,'const.bool.IN_2'(s(sK17,X528),s(fun(sK17,bool),sK18)))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,sK41(sK18,X527,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X527))) ) | $spl28),
  inference(forward_demodulation,[],[f8913,f956])).
fof(f8913,plain,(
  ( ! [X528,X527] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK41(sK18,X527,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,X528)))) | p(s(bool,X527))) ) | $spl28),
  inference(forward_demodulation,[],[f8334,f956])).
fof(f8334,plain,(
  ( ! [X528,X527] : (~p(s(bool,happ(s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(sK17,sK41(sK18,X527,sK17))))) | p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,X528)))) | p(s(bool,X527))) ) | $spl28),
  inference(superposition,[],[f881,f7676])).
fof(f8909,plain,(
  ( ! [X523] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK40(sK18,X523,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X523))) ) | $spl288),
  inference(cnf_transformation,[],[f8909_D])).
fof(f8909_D,plain,(
  ( ! [X523] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK40(sK18,X523,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X523))) ) <=> ~$spl288),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl288])])).
fof(f8912,plain,(
  ( ! [X526,X525] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK40(sK18,X526,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,'const.bool.IN_2'(s(sK17,X525),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X526))) ) | $spl28),
  inference(forward_demodulation,[],[f8911,f7676])).
fof(f8911,plain,(
  ( ! [X526,X525] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK40(sK18,X526,sK17)),s(fun(sK17,bool),sK18)))) | p(s(bool,'const.bool.IN_2'(s(sK17,X525),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X526))) ) | $spl28),
  inference(forward_demodulation,[],[f8910,f956])).
fof(f8910,plain,(
  ( ! [X526,X525] : (p(s(bool,'const.bool.IN_2'(s(sK17,X525),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X526)) | ~p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,sK40(sK18,X526,sK17)))))) ) | $spl28),
  inference(forward_demodulation,[],[f8333,f956])).
fof(f8333,plain,(
  ( ! [X526,X525] : (p(s(bool,happ(s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(sK17,X525)))) | p(s(bool,X526)) | ~p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,sK40(sK18,X526,sK17)))))) ) | $spl28),
  inference(superposition,[],[f878,f7676])).
fof(f878,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK40(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f592])).
fof(f592,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK39(X2,X1,X0))))) & ~p(s(bool,X1))) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) | p(s(bool,X1))) & ((~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK40(X2,X1,X0))))) & ~p(s(bool,X1))) | ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | p(s(bool,X1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK39,sK40])],[f591])).
fof(f591,plain,(
  ! [X0,X1,X2] : ((? [X6] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))) & ~p(s(bool,X1))) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) | p(s(bool,X1))) & ((? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) & ~p(s(bool,X1))) | ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | p(s(bool,X1)))))),
  inference(rectify,[],[f590])).
fof(f590,plain,(
  ! [X0,X1,X2] : ((? [X3] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) & ~p(s(bool,X1))) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,X1))) & ((? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) & ~p(s(bool,X1))) | ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | p(s(bool,X1)))))),
  inference(flattening,[],[f589])).
fof(f589,plain,(
  ! [X0,X1,X2] : ((? [X3] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) & ~p(s(bool,X1))) | (! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,X1)))) & ((? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) & ~p(s(bool,X1))) | ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | p(s(bool,X1)))))),
  inference(nnf_transformation,[],[f350])).
fof(f350,plain,(
  ! [X0,X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | p(s(bool,X1))) <=> (! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,X1))))),
  inference(flattening,[],[f349])).
fof(f349,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | p(s(bool,X1))) <=> (! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,X1))))),
  inference(rectify,[],[f95])).
fof(f95,axiom,(
  ! [X7] : ! [X33,X30] : (! [X13] : (p(s(bool,happ(s(fun(X7,bool),X30),s(X7,X13)))) | p(s(bool,X33))) <=> (! [X13] : p(s(bool,happ(s(fun(X7,bool),X30),s(X7,X13)))) | p(s(bool,X33))))),
  file('thm.util_prob.INJ_IMAGE_BIJ','thm.bool.LEFT_FORALL_OR_THM')).
fof(f8908,plain,(
  ( ! [X524,X523] : (p(s(bool,'const.bool.IN_2'(s(sK17,X524),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,sK40(sK18,X523,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X523))) ) | $spl28),
  inference(forward_demodulation,[],[f8907,f7676])).
fof(f8907,plain,(
  ( ! [X524,X523] : (p(s(bool,'const.bool.IN_2'(s(sK17,X524),s(fun(sK17,bool),sK18)))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,sK40(sK18,X523,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X523))) ) | $spl28),
  inference(forward_demodulation,[],[f8906,f956])).
fof(f8906,plain,(
  ( ! [X524,X523] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK40(sK18,X523,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X523)) | p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,X524))))) ) | $spl28),
  inference(forward_demodulation,[],[f8332,f956])).
fof(f8332,plain,(
  ( ! [X524,X523] : (~p(s(bool,happ(s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(sK17,sK40(sK18,X523,sK17))))) | p(s(bool,X523)) | p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,X524))))) ) | $spl28),
  inference(superposition,[],[f878,f7676])).
fof(f8901,plain,(
  ( ! [X519] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK39(sK18,X519,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X519))) ) | $spl284),
  inference(cnf_transformation,[],[f8901_D])).
fof(f8901_D,plain,(
  ( ! [X519] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK39(sK18,X519,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X519))) ) <=> ~$spl284),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl284])])).
fof(f8905,plain,(
  ( ! [X522,X521] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK39(sK18,X522,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,'const.bool.IN_2'(s(sK17,X521),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X522))) ) | $spl28),
  inference(forward_demodulation,[],[f8904,f7676])).
fof(f8904,plain,(
  ( ! [X522,X521] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK39(sK18,X522,sK17)),s(fun(sK17,bool),sK18)))) | p(s(bool,'const.bool.IN_2'(s(sK17,X521),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X522))) ) | $spl28),
  inference(forward_demodulation,[],[f8903,f956])).
fof(f8903,plain,(
  ( ! [X522,X521] : (p(s(bool,'const.bool.IN_2'(s(sK17,X521),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X522)) | ~p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,sK39(sK18,X522,sK17)))))) ) | $spl28),
  inference(forward_demodulation,[],[f8331,f956])).
fof(f8331,plain,(
  ( ! [X522,X521] : (p(s(bool,happ(s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(sK17,X521)))) | p(s(bool,X522)) | ~p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,sK39(sK18,X522,sK17)))))) ) | $spl28),
  inference(superposition,[],[f876,f7676])).
fof(f876,plain,(
  ( ! [X2,X0,X5,X1] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK39(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f592])).
fof(f8900,plain,(
  ( ! [X520,X519] : (p(s(bool,'const.bool.IN_2'(s(sK17,X520),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,sK39(sK18,X519,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X519))) ) | $spl28),
  inference(forward_demodulation,[],[f8899,f7676])).
fof(f8899,plain,(
  ( ! [X520,X519] : (p(s(bool,'const.bool.IN_2'(s(sK17,X520),s(fun(sK17,bool),sK18)))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,sK39(sK18,X519,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X519))) ) | $spl28),
  inference(forward_demodulation,[],[f8898,f956])).
fof(f8898,plain,(
  ( ! [X520,X519] : (~p(s(bool,'const.bool.IN_2'(s(sK17,sK39(sK18,X519,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,X519)) | p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,X520))))) ) | $spl28),
  inference(forward_demodulation,[],[f8330,f956])).
fof(f8330,plain,(
  ( ! [X520,X519] : (~p(s(bool,happ(s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(sK17,sK39(sK18,X519,sK17))))) | p(s(bool,X519)) | p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,X520))))) ) | $spl28),
  inference(superposition,[],[f876,f7676])).
fof(f8881,plain,(
  ( ! [X508] : (p(s(bool,'const.bool.IN_2'(s(sK17,sK37(sK18,X508,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X508))) ) | $spl282),
  inference(cnf_transformation,[],[f8881_D])).
fof(f8881_D,plain,(
  ( ! [X508] : (p(s(bool,'const.bool.IN_2'(s(sK17,sK37(sK18,X508,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X508))) ) <=> ~$spl282),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl282])])).
fof(f8884,plain,(
  ( ! [X511,X510] : (p(s(bool,'const.bool.IN_2'(s(sK17,sK37(sK18,X511,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,X510),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X511))) ) | $spl28),
  inference(forward_demodulation,[],[f8883,f7676])).
fof(f8883,plain,(
  ( ! [X511,X510] : (p(s(bool,'const.bool.IN_2'(s(sK17,sK37(sK18,X511,sK17)),s(fun(sK17,bool),sK18)))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,X510),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X511))) ) | $spl28),
  inference(forward_demodulation,[],[f8882,f956])).
fof(f8882,plain,(
  ( ! [X511,X510] : (~p(s(bool,'const.bool.IN_2'(s(sK17,X510),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,sK37(sK18,X511,sK17))))) | ~p(s(bool,X511))) ) | $spl28),
  inference(forward_demodulation,[],[f8324,f956])).
fof(f8324,plain,(
  ( ! [X511,X510] : (~p(s(bool,happ(s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(sK17,X510)))) | p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,sK37(sK18,X511,sK17))))) | ~p(s(bool,X511))) ) | $spl28),
  inference(superposition,[],[f872,f7676])).
fof(f872,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK37(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f585])).
fof(f585,plain,(
  ! [X0,X1,X2] : ((! [X6] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK36(X2,X1,X0))))))) & (~p(s(bool,X1)) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK37(X2,X1,X0))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK36,sK37])],[f584])).
fof(f584,plain,(
  ! [X0,X1,X2] : ((! [X6] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | (p(s(bool,X1)) & ? [X5] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))))) & (~p(s(bool,X1)) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X3] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(rectify,[],[f583])).
fof(f583,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (p(s(bool,X1)) & ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (~p(s(bool,X1)) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X3] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(flattening,[],[f582])).
fof(f582,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (p(s(bool,X1)) & ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & ((~p(s(bool,X1)) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ? [X3] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f346])).
fof(f346,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (p(s(bool,X1)) & ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f345])).
fof(f345,plain,(
  ! [X0] : ! [X1,X2] : (? [X3] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (p(s(bool,X1)) & ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f130])).
fof(f130,axiom,(
  ! [X7] : ! [X30,X33] : (? [X13] : (p(s(bool,X30)) & p(s(bool,happ(s(fun(X7,bool),X33),s(X7,X13))))) <=> (p(s(bool,X30)) & ? [X13] : p(s(bool,happ(s(fun(X7,bool),X33),s(X7,X13))))))),
  file('thm.util_prob.INJ_IMAGE_BIJ','thm.bool.RIGHT_EXISTS_AND_THM')).
fof(f8880,plain,(
  ( ! [X509,X508] : (~p(s(bool,'const.bool.IN_2'(s(sK17,X509),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,'const.bool.IN_2'(s(sK17,sK37(sK18,X508,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X508))) ) | $spl28),
  inference(forward_demodulation,[],[f8879,f7676])).
fof(f8879,plain,(
  ( ! [X509,X508] : (~p(s(bool,'const.bool.IN_2'(s(sK17,X509),s(fun(sK17,bool),sK18)))) | p(s(bool,'const.bool.IN_2'(s(sK17,sK37(sK18,X508,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X508))) ) | $spl28),
  inference(forward_demodulation,[],[f8878,f956])).
fof(f8878,plain,(
  ( ! [X509,X508] : (p(s(bool,'const.bool.IN_2'(s(sK17,sK37(sK18,X508,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,X509)))) | ~p(s(bool,X508))) ) | $spl28),
  inference(forward_demodulation,[],[f8323,f956])).
fof(f8323,plain,(
  ( ! [X509,X508] : (p(s(bool,happ(s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(sK17,sK37(sK18,X508,sK17))))) | ~p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,X509)))) | ~p(s(bool,X508))) ) | $spl28),
  inference(superposition,[],[f872,f7676])).
fof(f8874,plain,(
  ( ! [X504] : (p(s(bool,'const.bool.IN_2'(s(sK17,sK36(sK18,X504,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X504))) ) | $spl280),
  inference(cnf_transformation,[],[f8874_D])).
fof(f8874_D,plain,(
  ( ! [X504] : (p(s(bool,'const.bool.IN_2'(s(sK17,sK36(sK18,X504,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X504))) ) <=> ~$spl280),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl280])])).
fof(f8877,plain,(
  ( ! [X507,X506] : (p(s(bool,'const.bool.IN_2'(s(sK17,sK36(sK18,X507,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,X506),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X507))) ) | $spl28),
  inference(forward_demodulation,[],[f8876,f7676])).
fof(f8876,plain,(
  ( ! [X507,X506] : (p(s(bool,'const.bool.IN_2'(s(sK17,sK36(sK18,X507,sK17)),s(fun(sK17,bool),sK18)))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,X506),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X507))) ) | $spl28),
  inference(forward_demodulation,[],[f8875,f956])).
fof(f8875,plain,(
  ( ! [X507,X506] : (~p(s(bool,'const.bool.IN_2'(s(sK17,X506),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,sK36(sK18,X507,sK17))))) | ~p(s(bool,X507))) ) | $spl28),
  inference(forward_demodulation,[],[f8322,f956])).
fof(f8322,plain,(
  ( ! [X507,X506] : (~p(s(bool,happ(s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(sK17,X506)))) | p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,sK36(sK18,X507,sK17))))) | ~p(s(bool,X507))) ) | $spl28),
  inference(superposition,[],[f870,f7676])).
fof(f870,plain,(
  ( ! [X6,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK36(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f585])).
fof(f8873,plain,(
  ( ! [X505,X504] : (~p(s(bool,'const.bool.IN_2'(s(sK17,X505),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,'const.bool.IN_2'(s(sK17,sK36(sK18,X504,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X504))) ) | $spl28),
  inference(forward_demodulation,[],[f8872,f7676])).
fof(f8872,plain,(
  ( ! [X505,X504] : (~p(s(bool,'const.bool.IN_2'(s(sK17,X505),s(fun(sK17,bool),sK18)))) | p(s(bool,'const.bool.IN_2'(s(sK17,sK36(sK18,X504,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X504))) ) | $spl28),
  inference(forward_demodulation,[],[f8871,f956])).
fof(f8871,plain,(
  ( ! [X505,X504] : (p(s(bool,'const.bool.IN_2'(s(sK17,sK36(sK18,X504,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,X505)))) | ~p(s(bool,X504))) ) | $spl28),
  inference(forward_demodulation,[],[f8321,f956])).
fof(f8321,plain,(
  ( ! [X505,X504] : (p(s(bool,happ(s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(sK17,sK36(sK18,X504,sK17))))) | ~p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,X505)))) | ~p(s(bool,X504))) ) | $spl28),
  inference(superposition,[],[f870,f7676])).
fof(f8867,plain,(
  ( ! [X500] : (p(s(bool,'const.bool.IN_2'(s(sK17,sK35(X500,sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X500))) ) | $spl278),
  inference(cnf_transformation,[],[f8867_D])).
fof(f8867_D,plain,(
  ( ! [X500] : (p(s(bool,'const.bool.IN_2'(s(sK17,sK35(X500,sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X500))) ) <=> ~$spl278),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl278])])).
fof(f8870,plain,(
  ( ! [X502,X503] : (p(s(bool,'const.bool.IN_2'(s(sK17,sK35(X503,sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,X502),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X503))) ) | $spl28),
  inference(forward_demodulation,[],[f8869,f7676])).
fof(f8869,plain,(
  ( ! [X502,X503] : (p(s(bool,'const.bool.IN_2'(s(sK17,sK35(X503,sK18,sK17)),s(fun(sK17,bool),sK18)))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,X502),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X503))) ) | $spl28),
  inference(forward_demodulation,[],[f8868,f956])).
fof(f8868,plain,(
  ( ! [X502,X503] : (~p(s(bool,'const.bool.IN_2'(s(sK17,X502),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,sK35(X503,sK18,sK17))))) | ~p(s(bool,X503))) ) | $spl28),
  inference(forward_demodulation,[],[f8320,f956])).
fof(f8320,plain,(
  ( ! [X502,X503] : (~p(s(bool,happ(s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(sK17,X502)))) | p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,sK35(X503,sK18,sK17))))) | ~p(s(bool,X503))) ) | $spl28),
  inference(superposition,[],[f867,f7676])).
fof(f867,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK35(X2,X1,X0))))) | ~p(s(bool,X2)) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) )),
  inference(cnf_transformation,[],[f581])).
fof(f581,plain,(
  ! [X0,X1,X2] : ((! [X6] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) | ~p(s(bool,X2))) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK34(X2,X1,X0))))) & p(s(bool,X2)))) & (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,X2)) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK35(X2,X1,X0))))) & p(s(bool,X2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK34,sK35])],[f580])).
fof(f580,plain,(
  ! [X0,X1,X2] : ((! [X6] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) | ~p(s(bool,X2))) | (? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & p(s(bool,X2)))) & (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,X2)) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2)))))),
  inference(rectify,[],[f579])).
fof(f579,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,X2))) | (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & p(s(bool,X2)))) & (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,X2)) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2)))))),
  inference(flattening,[],[f578])).
fof(f578,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,X2))) | (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & p(s(bool,X2)))) & ((! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,X2))) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2)))))),
  inference(nnf_transformation,[],[f344])).
fof(f344,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2))) <=> (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & p(s(bool,X2))))),
  inference(flattening,[],[f343])).
fof(f343,plain,(
  ! [X0] : ! [X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2))) <=> (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & p(s(bool,X2))))),
  inference(rectify,[],[f141])).
fof(f141,axiom,(
  ! [X7] : ! [X30,X33] : (? [X13] : (p(s(bool,happ(s(fun(X7,bool),X30),s(X7,X13)))) & p(s(bool,X33))) <=> (? [X13] : p(s(bool,happ(s(fun(X7,bool),X30),s(X7,X13)))) & p(s(bool,X33))))),
  file('thm.util_prob.INJ_IMAGE_BIJ','thm.bool.LEFT_EXISTS_AND_THM')).
fof(f8866,plain,(
  ( ! [X500,X501] : (~p(s(bool,'const.bool.IN_2'(s(sK17,X501),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,'const.bool.IN_2'(s(sK17,sK35(X500,sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X500))) ) | $spl28),
  inference(forward_demodulation,[],[f8865,f7676])).
fof(f8865,plain,(
  ( ! [X500,X501] : (~p(s(bool,'const.bool.IN_2'(s(sK17,X501),s(fun(sK17,bool),sK18)))) | p(s(bool,'const.bool.IN_2'(s(sK17,sK35(X500,sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X500))) ) | $spl28),
  inference(forward_demodulation,[],[f8864,f956])).
fof(f8864,plain,(
  ( ! [X500,X501] : (p(s(bool,'const.bool.IN_2'(s(sK17,sK35(X500,sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X500)) | ~p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,X501))))) ) | $spl28),
  inference(forward_demodulation,[],[f8319,f956])).
fof(f8319,plain,(
  ( ! [X500,X501] : (p(s(bool,happ(s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(sK17,sK35(X500,sK18,sK17))))) | ~p(s(bool,X500)) | ~p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,X501))))) ) | $spl28),
  inference(superposition,[],[f867,f7676])).
fof(f8859,plain,(
  ( ! [X496] : (p(s(bool,'const.bool.IN_2'(s(sK17,sK34(X496,sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X496))) ) | $spl274),
  inference(cnf_transformation,[],[f8859_D])).
fof(f8859_D,plain,(
  ( ! [X496] : (p(s(bool,'const.bool.IN_2'(s(sK17,sK34(X496,sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X496))) ) <=> ~$spl274),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl274])])).
fof(f8863,plain,(
  ( ! [X499,X498] : (p(s(bool,'const.bool.IN_2'(s(sK17,sK34(X499,sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,X498),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X499))) ) | $spl28),
  inference(forward_demodulation,[],[f8862,f7676])).
fof(f8862,plain,(
  ( ! [X499,X498] : (p(s(bool,'const.bool.IN_2'(s(sK17,sK34(X499,sK18,sK17)),s(fun(sK17,bool),sK18)))) | ~p(s(bool,'const.bool.IN_2'(s(sK17,X498),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X499))) ) | $spl28),
  inference(forward_demodulation,[],[f8861,f956])).
fof(f8861,plain,(
  ( ! [X499,X498] : (~p(s(bool,'const.bool.IN_2'(s(sK17,X498),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,sK34(X499,sK18,sK17))))) | ~p(s(bool,X499))) ) | $spl28),
  inference(forward_demodulation,[],[f8318,f956])).
fof(f8318,plain,(
  ( ! [X499,X498] : (~p(s(bool,happ(s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(sK17,X498)))) | p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,sK34(X499,sK18,sK17))))) | ~p(s(bool,X499))) ) | $spl28),
  inference(superposition,[],[f865,f7676])).
fof(f865,plain,(
  ( ! [X6,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK34(X2,X1,X0))))) | ~p(s(bool,X2)) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6))))) )),
  inference(cnf_transformation,[],[f581])).
fof(f8858,plain,(
  ( ! [X496,X497] : (~p(s(bool,'const.bool.IN_2'(s(sK17,X497),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,'const.bool.IN_2'(s(sK17,sK34(X496,sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X496))) ) | $spl28),
  inference(forward_demodulation,[],[f8857,f7676])).
fof(f8857,plain,(
  ( ! [X496,X497] : (~p(s(bool,'const.bool.IN_2'(s(sK17,X497),s(fun(sK17,bool),sK18)))) | p(s(bool,'const.bool.IN_2'(s(sK17,sK34(X496,sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X496))) ) | $spl28),
  inference(forward_demodulation,[],[f8856,f956])).
fof(f8856,plain,(
  ( ! [X496,X497] : (p(s(bool,'const.bool.IN_2'(s(sK17,sK34(X496,sK18,sK17)),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X496)) | ~p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,X497))))) ) | $spl28),
  inference(forward_demodulation,[],[f8317,f956])).
fof(f8317,plain,(
  ( ! [X496,X497] : (p(s(bool,happ(s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(sK17,sK34(X496,sK18,sK17))))) | ~p(s(bool,X496)) | ~p(s(bool,happ(s(fun(sK17,bool),sK18),s(sK17,X497))))) ) | $spl28),
  inference(superposition,[],[f865,f7676])).
fof(f8056,plain,(
  ( ! [X225,X226] : (p(s(bool,X225)) | ~p(s(bool,X226)) | s(bool,X226) != s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) | $spl29),
  inference(resolution,[],[f1174,f957])).
fof(f7954,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | s(bool,X24) != s(bool,X26)) ) | $spl29),
  inference(resolution,[],[f1174,f706])).
fof(f7952,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) | $spl29),
  inference(resolution,[],[f1174,f705])).
fof(f7950,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) | $spl29),
  inference(resolution,[],[f1174,f704])).
fof(f7949,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | s(bool,X9) != s(bool,X11)) ) | $spl29),
  inference(resolution,[],[f1174,f702])).
fof(f7947,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) | $spl29),
  inference(resolution,[],[f1174,f700])).
fof(f7941,plain,(
  ( ! [X8,X7] : (sP2(X7,X8,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X7))) ) | $spl270),
  inference(cnf_transformation,[],[f7941_D])).
fof(f7941_D,plain,(
  ( ! [X8,X7] : (sP2(X7,X8,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X7))) ) <=> ~$spl270),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl270])])).
fof(f7935,plain,(
  ( ! [X14,X12,X13,X11] : (sP2(X11,X12,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X11)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X13),X14),s(fun(sK16,bool),sK20),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f2492,f678])).
fof(f7934,plain,(
  ( ! [X10,X8,X7,X9] : (sP2(X7,X8,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X7)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK16),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) ) | $spl9),
  inference(superposition,[],[f2492,f680])).
fof(f7921,plain,(
  ( ! [X8,X7] : (~sP2('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X7,X8) | ~p(s(bool,X7))) ) | $spl268),
  inference(cnf_transformation,[],[f7921_D])).
fof(f7921_D,plain,(
  ( ! [X8,X7] : (~sP2('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X7,X8) | ~p(s(bool,X7))) ) <=> ~$spl268),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl268])])).
fof(f7913,plain,(
  ( ! [X14,X12,X13,X11] : (~sP2('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X11,X12) | ~p(s(bool,X11)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X13),X14),s(fun(sK16,bool),sK20),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f2484,f678])).
fof(f2484,plain,(
  ( ! [X31,X32] : (~sP2('const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)),X31,X32) | ~p(s(bool,X31))) ) | $spl9),
  inference(resolution,[],[f1139,f708])).
fof(f7912,plain,(
  ( ! [X10,X8,X7,X9] : (~sP2('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X7,X8) | ~p(s(bool,X7)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK16),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) ) | $spl9),
  inference(superposition,[],[f2484,f680])).
fof(f7902,plain,(
  ( ! [X8,X7] : (~sP2(X7,X8,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X7))) ) | $spl266),
  inference(cnf_transformation,[],[f7902_D])).
fof(f7902_D,plain,(
  ( ! [X8,X7] : (~sP2(X7,X8,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X7))) ) <=> ~$spl266),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl266])])).
fof(f7896,plain,(
  ( ! [X14,X12,X13,X11] : (~sP2(X11,X12,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X11)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X13),X14),s(fun(sK16,bool),sK20),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f2483,f678])).
fof(f2483,plain,(
  ( ! [X30,X29] : (~sP2(X29,X30,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) | p(s(bool,X29))) ) | $spl9),
  inference(resolution,[],[f1139,f707])).
fof(f7895,plain,(
  ( ! [X10,X8,X7,X9] : (~sP2(X7,X8,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X7)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK16),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) ) | $spl9),
  inference(superposition,[],[f2483,f680])).
fof(f7886,plain,(
  ( ! [X8,X7] : (~sP2('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X7,X8) | p(s(bool,X8))) ) | $spl264),
  inference(cnf_transformation,[],[f7886_D])).
fof(f7886_D,plain,(
  ( ! [X8,X7] : (~sP2('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X7,X8) | p(s(bool,X8))) ) <=> ~$spl264),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl264])])).
fof(f7880,plain,(
  ( ! [X14,X12,X13,X11] : (~sP2('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X11,X12) | p(s(bool,X12)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X13),X14),s(fun(sK16,bool),sK20),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f2482,f678])).
fof(f2482,plain,(
  ( ! [X28,X27] : (~sP2('const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)),X28,X27) | p(s(bool,X27))) ) | $spl9),
  inference(resolution,[],[f1139,f707])).
fof(f7879,plain,(
  ( ! [X10,X8,X7,X9] : (~sP2('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X7,X8) | p(s(bool,X8)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK16),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) ) | $spl9),
  inference(superposition,[],[f2482,f680])).
fof(f7845,plain,(
  ( ! [X13] : (p(s(bool,X13)) | s(bool,X13) != s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) | $spl262),
  inference(cnf_transformation,[],[f7845_D])).
fof(f7845_D,plain,(
  ( ! [X13] : (p(s(bool,X13)) | s(bool,X13) != s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) <=> ~$spl262),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl262])])).
fof(f7803,plain,(
  ( ! [X251,X250] : (p(s(bool,X250)) | p(s(bool,X251)) | s(bool,X251) != s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) | $spl28),
  inference(resolution,[],[f1175,f957])).
fof(f7696,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X33) != s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | s(bool,X34) != s(bool,X35)) ) | $spl28),
  inference(resolution,[],[f1175,f706])).
fof(f7693,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X26) != s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) | $spl28),
  inference(resolution,[],[f1175,f705])).
fof(f7692,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) | $spl28),
  inference(resolution,[],[f1175,f704])).
fof(f7691,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | s(bool,X19) != s(bool,X20)) ) | $spl28),
  inference(resolution,[],[f1175,f702])).
fof(f7689,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) | $spl28),
  inference(resolution,[],[f1175,f700])).
fof(f1365,plain,(
  p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | $spl58),
  inference(cnf_transformation,[],[f1365_D])).
fof(f1365_D,plain,(
  p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f7677,plain,(
  p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | $spl28),
  inference(resolution,[],[f1175,f675])).
fof(f7642,plain,(
  ( ! [X19,X17,X20] : (~p(s(bool,X17)) | s(bool,X17) != s(bool,'const.pred_set.INJ_3'(s(fun(sK16,X19),X20),s(fun(sK16,bool),sK20),s(fun(X19,bool),'const.pred_set.EMPTY_0')))) ) | $spl260),
  inference(cnf_transformation,[],[f7642_D])).
fof(f7642_D,plain,(
  ( ! [X19,X17,X20] : (~p(s(bool,X17)) | s(bool,X17) != s(bool,'const.pred_set.INJ_3'(s(fun(sK16,X19),X20),s(fun(sK16,bool),sK20),s(fun(X19,bool),'const.pred_set.EMPTY_0')))) ) <=> ~$spl260),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl260])])).
fof(f7587,plain,(
  ( ! [X455,X457,X454,X456] : (p(s(bool,X454)) | ~p(s(bool,X455)) | s(bool,X455) != s(bool,'const.pred_set.INJ_3'(s(fun(sK16,X456),X457),s(fun(sK16,bool),sK20),s(fun(X456,bool),'const.pred_set.EMPTY_0')))) ) | $spl26),
  inference(resolution,[],[f1173,f957])).
fof(f7485,plain,(
  ( ! [X52,X50,X53,X51,X49] : (~p(s(bool,X49)) | ~p(s(bool,X50)) | s(bool,X50) != s(bool,'const.pred_set.INJ_3'(s(fun(sK16,X51),X52),s(fun(sK16,bool),sK20),s(fun(X51,bool),'const.pred_set.EMPTY_0'))) | s(bool,X49) != s(bool,X53)) ) | $spl26),
  inference(resolution,[],[f1173,f706])).
fof(f7483,plain,(
  ( ! [X45,X43,X41,X44,X42] : (p(s(bool,X41)) | ~p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X42) != s(bool,'const.pred_set.INJ_3'(s(fun(sK16,X44),X45),s(fun(sK16,bool),sK20),s(fun(X44,bool),'const.pred_set.EMPTY_0')))) ) | $spl26),
  inference(resolution,[],[f1173,f705])).
fof(f7481,plain,(
  ( ! [X35,X33,X31,X34,X32] : (p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,X33) | s(bool,X32) != s(bool,'const.pred_set.INJ_3'(s(fun(sK16,X34),X35),s(fun(sK16,bool),sK20),s(fun(X34,bool),'const.pred_set.EMPTY_0')))) ) | $spl26),
  inference(resolution,[],[f1173,f704])).
fof(f7480,plain,(
  ( ! [X30,X28,X26,X29,X27] : (~p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.pred_set.INJ_3'(s(fun(sK16,X28),X29),s(fun(sK16,bool),sK20),s(fun(X28,bool),'const.pred_set.EMPTY_0'))) | s(bool,X26) != s(bool,X30)) ) | $spl26),
  inference(resolution,[],[f1173,f702])).
fof(f7478,plain,(
  ( ! [X19,X17,X20,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,X18) | s(bool,X17) != s(bool,'const.pred_set.INJ_3'(s(fun(sK16,X19),X20),s(fun(sK16,bool),sK20),s(fun(X19,bool),'const.pred_set.EMPTY_0')))) ) | $spl26),
  inference(resolution,[],[f1173,f700])).
fof(f7466,plain,(
  ( ! [X14,X12,X13,X11] : (~sP13('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X11,X12) | p(s(bool,X11)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2089,f678])).
fof(f7465,plain,(
  ( ! [X10,X8,X7,X9] : (~sP13('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X7,X8) | p(s(bool,X7)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f2089,f680])).
fof(f7464,plain,(
  ( ! [X2,X0,X3,X1] : (~sP13('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0,X1) | p(s(bool,X0)) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2089,f948])).
fof(f7451,plain,(
  ( ! [X14,X12,X13,X11] : (~sP13(X11,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X12) | p(s(bool,X11)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2088,f678])).
fof(f7450,plain,(
  ( ! [X10,X8,X7,X9] : (~sP13(X7,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X8) | p(s(bool,X7)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f2088,f680])).
fof(f7449,plain,(
  ( ! [X2,X0,X3,X1] : (~sP13(X0,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X1) | p(s(bool,X0)) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2088,f948])).
fof(f7437,plain,(
  ( ! [X14,X12,X13,X11] : (~sP13(X11,X12,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X12)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2087,f678])).
fof(f7436,plain,(
  ( ! [X10,X8,X7,X9] : (~sP13(X7,X8,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X8)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f2087,f680])).
fof(f7435,plain,(
  ( ! [X2,X0,X3,X1] : (~sP13(X0,X1,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X1)) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2087,f948])).
fof(f7422,plain,(
  ( ! [X14,X12,X13,X11] : (~sP13(X11,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X12) | p(s(bool,X12)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2086,f678])).
fof(f7421,plain,(
  ( ! [X10,X8,X7,X9] : (~sP13(X7,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X8) | p(s(bool,X8)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f2086,f680])).
fof(f7420,plain,(
  ( ! [X2,X0,X3,X1] : (~sP13(X0,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X1) | p(s(bool,X1)) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2086,f948])).
fof(f7383,plain,(
  ( ! [X8,X7] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X7),X8),s(fun(sK16,bool),sK20),s(fun(X7,bool),'const.pred_set.EMPTY_0'))))) ) | $spl3),
  inference(superposition,[],[f1132,f678])).
fof(f7382,plain,(
  ( ! [X6,X5] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X5,sK16),X6),s(fun(X5,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) ) | $spl3),
  inference(superposition,[],[f1132,f680])).
fof(f1173,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK16,X0),X1),s(fun(sK16,bool),sK20),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl26),
  inference(cnf_transformation,[],[f1173_D])).
fof(f1173_D,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK16,X0),X1),s(fun(sK16,bool),sK20),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f7381,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK16,X0),X1),s(fun(sK16,bool),sK20),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl3),
  inference(superposition,[],[f1132,f948])).
fof(f7376,plain,(
  ( ! [X8,X7] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X7),X8),s(fun(sK17,bool),sK18),s(fun(X7,bool),'const.pred_set.EMPTY_0'))))) ) | $spl3),
  inference(superposition,[],[f1132,f678])).
fof(f7375,plain,(
  ( ! [X6,X5] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X5,sK17),X6),s(fun(X5,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl3),
  inference(superposition,[],[f1132,f680])).
fof(f7374,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X0),X1),s(fun(sK17,bool),sK18),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl3),
  inference(superposition,[],[f1132,f948])).
fof(f7367,plain,(
  ( ! [X225,X226] : (p(s(bool,X225)) | ~p(s(bool,X226)) | s(bool,X226) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) ) | $spl3),
  inference(resolution,[],[f1132,f957])).
fof(f7265,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20))) | s(bool,X24) != s(bool,X26)) ) | $spl3),
  inference(resolution,[],[f1132,f706])).
fof(f7263,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) ) | $spl3),
  inference(resolution,[],[f1132,f705])).
fof(f7261,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) ) | $spl3),
  inference(resolution,[],[f1132,f704])).
fof(f7260,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20))) | s(bool,X9) != s(bool,X11)) ) | $spl3),
  inference(resolution,[],[f1132,f702])).
fof(f7258,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) ) | $spl3),
  inference(resolution,[],[f1132,f700])).
fof(f7246,plain,(
  ( ! [X14,X12,X13,X11] : (~sP12('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X11,X12) | p(s(bool,X11)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2081,f678])).
fof(f7245,plain,(
  ( ! [X10,X8,X7,X9] : (~sP12('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X7,X8) | p(s(bool,X7)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f2081,f680])).
fof(f7244,plain,(
  ( ! [X2,X0,X3,X1] : (~sP12('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0,X1) | p(s(bool,X0)) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2081,f948])).
fof(f7232,plain,(
  ( ! [X14,X12,X13,X11] : (~sP12(X11,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X12) | p(s(bool,X11)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2080,f678])).
fof(f7231,plain,(
  ( ! [X10,X8,X7,X9] : (~sP12(X7,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X8) | p(s(bool,X7)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f2080,f680])).
fof(f7230,plain,(
  ( ! [X2,X0,X3,X1] : (~sP12(X0,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X1) | p(s(bool,X0)) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2080,f948])).
fof(f7192,plain,(
  ( ! [X8,X7] : (~p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X7),X8),s(fun(sK16,bool),sK20),s(fun(X7,bool),'const.pred_set.EMPTY_0'))))) ) | $spl1),
  inference(superposition,[],[f1131,f678])).
fof(f7191,plain,(
  ( ! [X6,X5] : (~p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X5,sK16),X6),s(fun(X5,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) ) | $spl1),
  inference(superposition,[],[f1131,f680])).
fof(f7190,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK16,X0),X1),s(fun(sK16,bool),sK20),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl1),
  inference(superposition,[],[f1131,f948])).
fof(f7185,plain,(
  ( ! [X8,X7] : (~p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X7),X8),s(fun(sK17,bool),sK18),s(fun(X7,bool),'const.pred_set.EMPTY_0'))))) ) | $spl1),
  inference(superposition,[],[f1131,f678])).
fof(f7184,plain,(
  ( ! [X6,X5] : (~p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X5,sK17),X6),s(fun(X5,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl1),
  inference(superposition,[],[f1131,f680])).
fof(f7183,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X0),X1),s(fun(sK17,bool),sK18),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl1),
  inference(superposition,[],[f1131,f948])).
fof(f7176,plain,(
  ( ! [X222,X223] : (p(s(bool,X222)) | ~p(s(bool,X223)) | s(bool,X223) != s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) ) | $spl1),
  inference(resolution,[],[f1131,f957])).
fof(f7074,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20))) | s(bool,X21) != s(bool,X23)) ) | $spl1),
  inference(resolution,[],[f1131,f706])).
fof(f7072,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) ) | $spl1),
  inference(resolution,[],[f1131,f705])).
fof(f7070,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) ) | $spl1),
  inference(resolution,[],[f1131,f704])).
fof(f7069,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20))) | s(bool,X6) != s(bool,X8)) ) | $spl1),
  inference(resolution,[],[f1131,f702])).
fof(f7067,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) ) | $spl1),
  inference(resolution,[],[f1131,f700])).
fof(f7059,plain,(
  ( ! [X14,X12,X13,X11] : (~sP12(X11,X12,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X11)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2079,f678])).
fof(f7058,plain,(
  ( ! [X10,X8,X7,X9] : (~sP12(X7,X8,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X7)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f2079,f680])).
fof(f7057,plain,(
  ( ! [X2,X0,X3,X1] : (~sP12(X0,X1,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X0)) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2079,f948])).
fof(f7044,plain,(
  ( ! [X14,X12,X13,X11] : (~sP12('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X11,X12) | p(s(bool,X12)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2078,f678])).
fof(f7043,plain,(
  ( ! [X10,X8,X7,X9] : (~sP12('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X7,X8) | p(s(bool,X8)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f2078,f680])).
fof(f7042,plain,(
  ( ! [X2,X0,X3,X1] : (~sP12('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0,X1) | p(s(bool,X1)) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2078,f948])).
fof(f6983,plain,(
  ( ! [X8,X7] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X7),X8),s(fun(sK16,bool),sK20),s(fun(X7,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f1133,f678])).
fof(f6982,plain,(
  ( ! [X6,X5] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X5,sK16),X6),s(fun(X5,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) ) | $spl2),
  inference(superposition,[],[f1133,f680])).
fof(f6981,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK16,X0),X1),s(fun(sK16,bool),sK20),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f1133,f948])).
fof(f6580,plain,(
  p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)))) | $spl230),
  inference(cnf_transformation,[],[f6580_D])).
fof(f6580_D,plain,(
  p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)))) <=> ~$spl230),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl230])])).
fof(f6976,plain,(
  ( ! [X8,X7] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X7),X8),s(fun(sK17,bool),sK18),s(fun(X7,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f1133,f678])).
fof(f6975,plain,(
  ( ! [X6,X5] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X5,sK17),X6),s(fun(X5,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl2),
  inference(superposition,[],[f1133,f680])).
fof(f6974,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X0),X1),s(fun(sK17,bool),sK18),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f1133,f948])).
fof(f7002,plain,(
  ( ! [X10] : (p(s(bool,X10)) | s(bool,X10) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) ) | $spl242),
  inference(cnf_transformation,[],[f7002_D])).
fof(f7002_D,plain,(
  ( ! [X10] : (p(s(bool,X10)) | s(bool,X10) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) ) <=> ~$spl242),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl242])])).
fof(f6964,plain,(
  ( ! [X248,X247] : (p(s(bool,X247)) | p(s(bool,X248)) | s(bool,X248) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) ) | $spl2),
  inference(resolution,[],[f1133,f957])).
fof(f6857,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X30) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20))) | s(bool,X31) != s(bool,X32)) ) | $spl2),
  inference(resolution,[],[f1133,f706])).
fof(f6854,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | p(s(bool,X23)) | s(bool,X23) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) ) | $spl2),
  inference(resolution,[],[f1133,f705])).
fof(f6853,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X18) != s(bool,X20) | s(bool,X19) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) ) | $spl2),
  inference(resolution,[],[f1133,f704])).
fof(f6852,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20))) | s(bool,X16) != s(bool,X17)) ) | $spl2),
  inference(resolution,[],[f1133,f702])).
fof(f6850,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) ) | $spl2),
  inference(resolution,[],[f1133,f700])).
fof(f6834,plain,(
  ( ! [X14,X12,X13,X11] : (sP11(X11,X12,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X12)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2071,f678])).
fof(f6833,plain,(
  ( ! [X10,X8,X7,X9] : (sP11(X7,X8,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X8)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f2071,f680])).
fof(f6832,plain,(
  ( ! [X2,X0,X3,X1] : (sP11(X0,X1,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X1)) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2071,f948])).
fof(f6820,plain,(
  ( ! [X14,X12,X13,X11] : (sP11(X11,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X12) | p(s(bool,X12)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2070,f678])).
fof(f6819,plain,(
  ( ! [X10,X8,X7,X9] : (sP11(X7,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X8) | p(s(bool,X8)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f2070,f680])).
fof(f6818,plain,(
  ( ! [X2,X0,X3,X1] : (sP11(X0,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X1) | p(s(bool,X1)) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2070,f948])).
fof(f6806,plain,(
  ( ! [X14,X12,X13,X11] : (~sP11('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X11,X12) | ~p(s(bool,X11)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2069,f678])).
fof(f6805,plain,(
  ( ! [X10,X8,X7,X9] : (~sP11('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X7,X8) | ~p(s(bool,X7)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f2069,f680])).
fof(f6804,plain,(
  ( ! [X2,X0,X3,X1] : (~sP11('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2069,f948])).
fof(f5694,plain,(
  p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | $spl196),
  inference(cnf_transformation,[],[f5694_D])).
fof(f5694_D,plain,(
  p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f6754,plain,(
  ( ! [X8,X7] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X7),X8),s(fun(sK16,bool),sK20),s(fun(X7,bool),'const.pred_set.EMPTY_0'))))) ) | $spl0),
  inference(superposition,[],[f1130,f678])).
fof(f1130,plain,(
  p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) | $spl0),
  inference(cnf_transformation,[],[f1130_D])).
fof(f1130_D,plain,(
  p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f6753,plain,(
  ( ! [X6,X5] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X5,sK16),X6),s(fun(X5,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) ) | $spl0),
  inference(superposition,[],[f1130,f680])).
fof(f6752,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK16,X0),X1),s(fun(sK16,bool),sK20),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl0),
  inference(superposition,[],[f1130,f948])).
fof(f5692,plain,(
  p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)))) | $spl194),
  inference(cnf_transformation,[],[f5692_D])).
fof(f5692_D,plain,(
  p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)))) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f6747,plain,(
  ( ! [X8,X7] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X7),X8),s(fun(sK17,bool),sK18),s(fun(X7,bool),'const.pred_set.EMPTY_0'))))) ) | $spl0),
  inference(superposition,[],[f1130,f678])).
fof(f6746,plain,(
  ( ! [X6,X5] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X5,sK17),X6),s(fun(X5,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl0),
  inference(superposition,[],[f1130,f680])).
fof(f6745,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X0),X1),s(fun(sK17,bool),sK18),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl0),
  inference(superposition,[],[f1130,f948])).
fof(f5867,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) ) | $spl200),
  inference(cnf_transformation,[],[f5867_D])).
fof(f5867_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) ) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f6735,plain,(
  ( ! [X239,X238] : (p(s(bool,X238)) | p(s(bool,X239)) | s(bool,X239) != s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) ) | $spl0),
  inference(resolution,[],[f1130,f957])).
fof(f6628,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20))) | s(bool,X22) != s(bool,X23)) ) | $spl0),
  inference(resolution,[],[f1130,f706])).
fof(f6625,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) ) | $spl0),
  inference(resolution,[],[f1130,f705])).
fof(f6624,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) ) | $spl0),
  inference(resolution,[],[f1130,f704])).
fof(f6623,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20))) | s(bool,X7) != s(bool,X8)) ) | $spl0),
  inference(resolution,[],[f1130,f702])).
fof(f6621,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) ) | $spl0),
  inference(resolution,[],[f1130,f700])).
fof(f6613,plain,(
  ( ! [X14,X12,X13,X11] : (~sP11('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X11,X12) | ~p(s(bool,X12)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2068,f678])).
fof(f6612,plain,(
  ( ! [X10,X8,X7,X9] : (~sP11('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X7,X8) | ~p(s(bool,X8)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f2068,f680])).
fof(f6611,plain,(
  ( ! [X2,X0,X3,X1] : (~sP11('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2068,f948])).
fof(f6597,plain,(
  ( ! [X14,X12,X13,X11] : (sP6('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X11,X12) | p(s(bool,X11)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2050,f678])).
fof(f6596,plain,(
  ( ! [X10,X8,X7,X9] : (sP6('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X7,X8) | p(s(bool,X7)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f2050,f680])).
fof(f6595,plain,(
  ( ! [X2,X0,X3,X1] : (sP6('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0,X1) | p(s(bool,X0)) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2050,f948])).
fof(f6542,plain,(
  ( ! [X8,X7] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X7),X8),s(fun(sK16,bool),sK20),s(fun(X7,bool),'const.pred_set.EMPTY_0'))))) ) | $spl3),
  inference(superposition,[],[f1132,f678])).
fof(f6541,plain,(
  ( ! [X6,X5] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X5,sK16),X6),s(fun(X5,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) ) | $spl3),
  inference(superposition,[],[f1132,f680])).
fof(f6540,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK16,X0),X1),s(fun(sK16,bool),sK20),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl3),
  inference(superposition,[],[f1132,f948])).
fof(f6535,plain,(
  ( ! [X8,X7] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X7),X8),s(fun(sK17,bool),sK18),s(fun(X7,bool),'const.pred_set.EMPTY_0'))))) ) | $spl3),
  inference(superposition,[],[f1132,f678])).
fof(f6534,plain,(
  ( ! [X6,X5] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X5,sK17),X6),s(fun(X5,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl3),
  inference(superposition,[],[f1132,f680])).
fof(f6533,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X0),X1),s(fun(sK17,bool),sK18),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl3),
  inference(superposition,[],[f1132,f948])).
fof(f6526,plain,(
  ( ! [X225,X226] : (p(s(bool,X225)) | ~p(s(bool,X226)) | s(bool,X226) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) ) | $spl3),
  inference(resolution,[],[f1132,f957])).
fof(f6424,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20))) | s(bool,X24) != s(bool,X26)) ) | $spl3),
  inference(resolution,[],[f1132,f706])).
fof(f6422,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) ) | $spl3),
  inference(resolution,[],[f1132,f705])).
fof(f6420,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) ) | $spl3),
  inference(resolution,[],[f1132,f704])).
fof(f6419,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20))) | s(bool,X9) != s(bool,X11)) ) | $spl3),
  inference(resolution,[],[f1132,f702])).
fof(f6417,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) ) | $spl3),
  inference(resolution,[],[f1132,f700])).
fof(f6566,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK16,happ(s(fun(sK17,sK16),sK19),s(sK17,sK25(sK20,sK18,sK19,sK16,sK17)))),s(fun(sK16,bool),sK20)))) | $spl227),
  inference(cnf_transformation,[],[f6566_D])).
fof(f6566_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK16,happ(s(fun(sK17,sK16),sK19),s(sK17,sK25(sK20,sK18,sK19,sK16,sK17)))),s(fun(sK16,bool),sK20)))) <=> ~$spl227),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl227])])).
fof(f6564,plain,(
  ~sP0(sK16,sK17,sK19,sK20,sK18) | $spl225),
  inference(cnf_transformation,[],[f6564_D])).
fof(f6564_D,plain,(
  ~sP0(sK16,sK17,sK19,sK20,sK18) <=> ~$spl225),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl225])])).
fof(f6413,plain,(
  ~sP0(sK16,sK17,sK19,sK20,sK18) | ~p(s(bool,'const.bool.IN_2'(s(sK16,happ(s(fun(sK17,sK16),sK19),s(sK17,sK25(sK20,sK18,sK19,sK16,sK17)))),s(fun(sK16,bool),sK20)))) | $spl3),
  inference(resolution,[],[f1132,f672])).
fof(f6562,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK17,sK25(sK20,sK18,sK19,sK16,sK17)),s(fun(sK17,bool),sK18)))) | $spl222),
  inference(cnf_transformation,[],[f6562_D])).
fof(f6562_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK17,sK25(sK20,sK18,sK19,sK16,sK17)),s(fun(sK17,bool),sK18)))) <=> ~$spl222),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl222])])).
fof(f6412,plain,(
  ~sP0(sK16,sK17,sK19,sK20,sK18) | p(s(bool,'const.bool.IN_2'(s(sK17,sK25(sK20,sK18,sK19,sK16,sK17)),s(fun(sK17,bool),sK18)))) | $spl3),
  inference(resolution,[],[f1132,f671])).
fof(f6404,plain,(
  ( ! [X14,X12,X13,X11] : (sP6(X11,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X12) | p(s(bool,X11)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2049,f678])).
fof(f6403,plain,(
  ( ! [X10,X8,X7,X9] : (sP6(X7,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X8) | p(s(bool,X7)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f2049,f680])).
fof(f6402,plain,(
  ( ! [X2,X0,X3,X1] : (sP6(X0,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X1) | p(s(bool,X0)) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2049,f948])).
fof(f6386,plain,(
  ( ! [X14,X12,X13,X11] : (sP6('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X11,X12) | p(s(bool,X12)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2048,f678])).
fof(f6385,plain,(
  ( ! [X10,X8,X7,X9] : (sP6('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X7,X8) | p(s(bool,X8)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f2048,f680])).
fof(f6384,plain,(
  ( ! [X2,X0,X3,X1] : (sP6('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0,X1) | p(s(bool,X1)) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2048,f948])).
fof(f6369,plain,(
  ( ! [X14,X12,X13,X11] : (sP6(X11,X12,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X11)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2047,f678])).
fof(f6368,plain,(
  ( ! [X10,X8,X7,X9] : (sP6(X7,X8,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X7)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f2047,f680])).
fof(f6367,plain,(
  ( ! [X2,X0,X3,X1] : (sP6(X0,X1,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X0)) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2047,f948])).
fof(f6351,plain,(
  ( ! [X14,X12,X13,X11] : (~sP6(X11,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X12) | ~p(s(bool,X11)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2041,f678])).
fof(f6350,plain,(
  ( ! [X10,X8,X7,X9] : (~sP6(X7,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X8) | ~p(s(bool,X7)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f2041,f680])).
fof(f6349,plain,(
  ( ! [X2,X0,X3,X1] : (~sP6(X0,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X1) | ~p(s(bool,X0)) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2041,f948])).
fof(f6334,plain,(
  ( ! [X14,X12,X13,X11] : (~sP6(X11,X12,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X11)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2040,f678])).
fof(f6333,plain,(
  ( ! [X10,X8,X7,X9] : (~sP6(X7,X8,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X7)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f2040,f680])).
fof(f6332,plain,(
  ( ! [X2,X0,X3,X1] : (~sP6(X0,X1,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X0)) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2040,f948])).
fof(f6300,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f6144,f1131])).
fof(f6144,plain,(
  p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) | $spl2),
  inference(resolution,[],[f1133,f674])).
fof(f6129,plain,(
  ( ! [X14,X12,X13,X11] : (~sP4('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X11,X12) | ~p(s(bool,X11)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2025,f678])).
fof(f6128,plain,(
  ( ! [X10,X8,X7,X9] : (~sP4('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X7,X8) | ~p(s(bool,X7)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f2025,f680])).
fof(f6127,plain,(
  ( ! [X2,X0,X3,X1] : (~sP4('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2025,f948])).
fof(f6111,plain,(
  ( ! [X14,X12,X13,X11] : (~sP4('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X11,X12) | ~p(s(bool,X12)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2024,f678])).
fof(f6110,plain,(
  ( ! [X10,X8,X7,X9] : (~sP4('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X7,X8) | ~p(s(bool,X8)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f2024,f680])).
fof(f6109,plain,(
  ( ! [X2,X0,X3,X1] : (~sP4('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2024,f948])).
fof(f6071,plain,(
  ( ! [X8,X7] : (~p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X7),X8),s(fun(sK16,bool),sK20),s(fun(X7,bool),'const.pred_set.EMPTY_0'))))) ) | $spl1),
  inference(superposition,[],[f1131,f678])).
fof(f6070,plain,(
  ( ! [X6,X5] : (~p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X5,sK16),X6),s(fun(X5,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) ) | $spl1),
  inference(superposition,[],[f1131,f680])).
fof(f6069,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK16,X0),X1),s(fun(sK16,bool),sK20),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl1),
  inference(superposition,[],[f1131,f948])).
fof(f6064,plain,(
  ( ! [X8,X7] : (~p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X7),X8),s(fun(sK17,bool),sK18),s(fun(X7,bool),'const.pred_set.EMPTY_0'))))) ) | $spl1),
  inference(superposition,[],[f1131,f678])).
fof(f6063,plain,(
  ( ! [X6,X5] : (~p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X5,sK17),X6),s(fun(X5,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl1),
  inference(superposition,[],[f1131,f680])).
fof(f6062,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X0),X1),s(fun(sK17,bool),sK18),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl1),
  inference(superposition,[],[f1131,f948])).
fof(f6055,plain,(
  ( ! [X222,X223] : (p(s(bool,X222)) | ~p(s(bool,X223)) | s(bool,X223) != s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) ) | $spl1),
  inference(resolution,[],[f1131,f957])).
fof(f5953,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20))) | s(bool,X21) != s(bool,X23)) ) | $spl1),
  inference(resolution,[],[f1131,f706])).
fof(f5951,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) ) | $spl1),
  inference(resolution,[],[f1131,f705])).
fof(f5949,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) ) | $spl1),
  inference(resolution,[],[f1131,f704])).
fof(f5948,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20))) | s(bool,X6) != s(bool,X8)) ) | $spl1),
  inference(resolution,[],[f1131,f702])).
fof(f5946,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) ) | $spl1),
  inference(resolution,[],[f1131,f700])).
fof(f5937,plain,(
  ( ! [X14,X12,X13,X11] : (sP2(X11,X12,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X11)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2019,f678])).
fof(f5936,plain,(
  ( ! [X10,X8,X7,X9] : (sP2(X7,X8,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X7)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f2019,f680])).
fof(f5935,plain,(
  ( ! [X2,X0,X3,X1] : (sP2(X0,X1,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X0)) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2019,f948])).
fof(f5917,plain,(
  ( ! [X14,X12,X13,X11] : (~sP2('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X11,X12) | ~p(s(bool,X11)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2011,f678])).
fof(f5916,plain,(
  ( ! [X10,X8,X7,X9] : (~sP2('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X7,X8) | ~p(s(bool,X7)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f2011,f680])).
fof(f5915,plain,(
  ( ! [X2,X0,X3,X1] : (~sP2('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2011,f948])).
fof(f5901,plain,(
  ( ! [X14,X12,X13,X11] : (~sP2(X11,X12,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X11)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2010,f678])).
fof(f5900,plain,(
  ( ! [X10,X8,X7,X9] : (~sP2(X7,X8,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X7)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f2010,f680])).
fof(f5899,plain,(
  ( ! [X2,X0,X3,X1] : (~sP2(X0,X1,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X0)) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2010,f948])).
fof(f5847,plain,(
  ( ! [X8,X7] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X7),X8),s(fun(sK16,bool),sK20),s(fun(X7,bool),'const.pred_set.EMPTY_0'))))) ) | $spl0),
  inference(superposition,[],[f1130,f678])).
fof(f5846,plain,(
  ( ! [X6,X5] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X5,sK16),X6),s(fun(X5,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) ) | $spl0),
  inference(superposition,[],[f1130,f680])).
fof(f5845,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK16,X0),X1),s(fun(sK16,bool),sK20),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl0),
  inference(superposition,[],[f1130,f948])).
fof(f5840,plain,(
  ( ! [X8,X7] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X7),X8),s(fun(sK17,bool),sK18),s(fun(X7,bool),'const.pred_set.EMPTY_0'))))) ) | $spl0),
  inference(superposition,[],[f1130,f678])).
fof(f5839,plain,(
  ( ! [X6,X5] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X5,sK17),X6),s(fun(X5,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl0),
  inference(superposition,[],[f1130,f680])).
fof(f5838,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X0),X1),s(fun(sK17,bool),sK18),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl0),
  inference(superposition,[],[f1130,f948])).
fof(f5828,plain,(
  ( ! [X239,X238] : (p(s(bool,X238)) | p(s(bool,X239)) | s(bool,X239) != s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) ) | $spl0),
  inference(resolution,[],[f1130,f957])).
fof(f5721,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20))) | s(bool,X22) != s(bool,X23)) ) | $spl0),
  inference(resolution,[],[f1130,f706])).
fof(f5718,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) ) | $spl0),
  inference(resolution,[],[f1130,f705])).
fof(f5717,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) ) | $spl0),
  inference(resolution,[],[f1130,f704])).
fof(f5716,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20))) | s(bool,X7) != s(bool,X8)) ) | $spl0),
  inference(resolution,[],[f1130,f702])).
fof(f5714,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) ) | $spl0),
  inference(resolution,[],[f1130,f700])).
fof(f5705,plain,(
  ( ! [X14,X12,X13,X11] : (~sP2('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X11,X12) | p(s(bool,X12)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2009,f678])).
fof(f5704,plain,(
  ( ! [X10,X8,X7,X9] : (~sP2('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X7,X8) | p(s(bool,X8)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f2009,f680])).
fof(f5703,plain,(
  ( ! [X2,X0,X3,X1] : (~sP2('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0,X1) | p(s(bool,X1)) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2009,f948])).
fof(f5660,plain,(
  ( ! [X8,X7] : (~p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X7),X8),s(fun(sK16,bool),sK20),s(fun(X7,bool),'const.pred_set.EMPTY_0'))))) ) | $spl1),
  inference(superposition,[],[f1131,f678])).
fof(f5659,plain,(
  ( ! [X6,X5] : (~p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X5,sK16),X6),s(fun(X5,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) ) | $spl1),
  inference(superposition,[],[f1131,f680])).
fof(f5658,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK16,X0),X1),s(fun(sK16,bool),sK20),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl1),
  inference(superposition,[],[f1131,f948])).
fof(f5653,plain,(
  ( ! [X8,X7] : (~p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X7),X8),s(fun(sK17,bool),sK18),s(fun(X7,bool),'const.pred_set.EMPTY_0'))))) ) | $spl1),
  inference(superposition,[],[f1131,f678])).
fof(f5652,plain,(
  ( ! [X6,X5] : (~p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X5,sK17),X6),s(fun(X5,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl1),
  inference(superposition,[],[f1131,f680])).
fof(f5651,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X0),X1),s(fun(sK17,bool),sK18),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl1),
  inference(superposition,[],[f1131,f948])).
fof(f5644,plain,(
  ( ! [X222,X223] : (p(s(bool,X222)) | ~p(s(bool,X223)) | s(bool,X223) != s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) ) | $spl1),
  inference(resolution,[],[f1131,f957])).
fof(f5542,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20))) | s(bool,X21) != s(bool,X23)) ) | $spl1),
  inference(resolution,[],[f1131,f706])).
fof(f5540,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) ) | $spl1),
  inference(resolution,[],[f1131,f705])).
fof(f5538,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) ) | $spl1),
  inference(resolution,[],[f1131,f704])).
fof(f5537,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20))) | s(bool,X6) != s(bool,X8)) ) | $spl1),
  inference(resolution,[],[f1131,f702])).
fof(f5535,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) ) | $spl1),
  inference(resolution,[],[f1131,f700])).
fof(f5376,plain,(
  ( ! [X0] : (sP14('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X0) | ~p(s(bool,X0))) ) | $spl190),
  inference(cnf_transformation,[],[f5376_D])).
fof(f5376_D,plain,(
  ( ! [X0] : (sP14('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X0) | ~p(s(bool,X0))) ) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f5369,plain,(
  ( ! [X10,X11,X9] : (sP14('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X9) | ~p(s(bool,X9)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X10),X11),s(fun(sK16,bool),sK20),s(fun(X10,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f2570,f678])).
fof(f2570,plain,(
  ( ! [X200] : (sP14('const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)),X200) | ~p(s(bool,X200))) ) | $spl9),
  inference(resolution,[],[f1139,f844])).
fof(f5368,plain,(
  ( ! [X6,X8,X7] : (sP14('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X6) | ~p(s(bool,X6)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X7,sK16),X8),s(fun(X7,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) ) | $spl9),
  inference(superposition,[],[f2570,f680])).
fof(f5367,plain,(
  ( ! [X2,X0,X1] : (sP14('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X0) | ~p(s(bool,X0)) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK16,X1),X2),s(fun(sK16,bool),sK20),s(fun(X1,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f2570,f948])).
fof(f5359,plain,(
  ( ! [X0] : (sP14(X0,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X0))) ) | $spl188),
  inference(cnf_transformation,[],[f5359_D])).
fof(f5359_D,plain,(
  ( ! [X0] : (sP14(X0,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X0))) ) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f5352,plain,(
  ( ! [X10,X11,X9] : (sP14(X9,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X9)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X10),X11),s(fun(sK16,bool),sK20),s(fun(X10,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f2569,f678])).
fof(f2569,plain,(
  ( ! [X199] : (sP14(X199,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) | ~p(s(bool,X199))) ) | $spl9),
  inference(resolution,[],[f1139,f841])).
fof(f5351,plain,(
  ( ! [X6,X8,X7] : (sP14(X6,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X6)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X7,sK16),X8),s(fun(X7,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) ) | $spl9),
  inference(superposition,[],[f2569,f680])).
fof(f5350,plain,(
  ( ! [X2,X0,X1] : (sP14(X0,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X0)) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK16,X1),X2),s(fun(sK16,bool),sK20),s(fun(X1,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f2569,f948])).
fof(f5342,plain,(
  ( ! [X0] : (~sP14('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X0) | p(s(bool,X0))) ) | $spl186),
  inference(cnf_transformation,[],[f5342_D])).
fof(f5342_D,plain,(
  ( ! [X0] : (~sP14('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X0) | p(s(bool,X0))) ) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f5335,plain,(
  ( ! [X10,X11,X9] : (~sP14('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X9) | p(s(bool,X9)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X10),X11),s(fun(sK16,bool),sK20),s(fun(X10,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f2568,f678])).
fof(f2568,plain,(
  ( ! [X198] : (~sP14('const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)),X198) | p(s(bool,X198))) ) | $spl9),
  inference(resolution,[],[f1139,f839])).
fof(f839,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | ~sP14(X0,X1)) )),
  inference(cnf_transformation,[],[f572])).
fof(f5334,plain,(
  ( ! [X6,X8,X7] : (~sP14('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X6) | p(s(bool,X6)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X7,sK16),X8),s(fun(X7,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) ) | $spl9),
  inference(superposition,[],[f2568,f680])).
fof(f5333,plain,(
  ( ! [X2,X0,X1] : (~sP14('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X0) | p(s(bool,X0)) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK16,X1),X2),s(fun(sK16,bool),sK20),s(fun(X1,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f2568,f948])).
fof(f5327,plain,(
  ( ! [X0] : (~sP14(X0,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X0))) ) | $spl184),
  inference(cnf_transformation,[],[f5327_D])).
fof(f5327_D,plain,(
  ( ! [X0] : (~sP14(X0,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X0))) ) <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f5320,plain,(
  ( ! [X10,X11,X9] : (~sP14(X9,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X9)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X10),X11),s(fun(sK16,bool),sK20),s(fun(X10,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f2567,f678])).
fof(f2567,plain,(
  ( ! [X197] : (~sP14(X197,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) | p(s(bool,X197))) ) | $spl9),
  inference(resolution,[],[f1139,f839])).
fof(f5319,plain,(
  ( ! [X6,X8,X7] : (~sP14(X6,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X6)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X7,sK16),X8),s(fun(X7,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) ) | $spl9),
  inference(superposition,[],[f2567,f680])).
fof(f5318,plain,(
  ( ! [X2,X0,X1] : (~sP14(X0,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X0)) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK16,X1),X2),s(fun(sK16,bool),sK20),s(fun(X1,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f2567,f948])).
fof(f5306,plain,(
  ( ! [X14,X12,X13,X11] : (sP13(X11,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X12) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X13),X14),s(fun(sK16,bool),sK20),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1067,f678])).
fof(f5305,plain,(
  ( ! [X10,X8,X7,X9] : (sP13(X7,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X8) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK16),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) )),
  inference(superposition,[],[f1067,f680])).
fof(f5304,plain,(
  ( ! [X2,X0,X3,X1] : (sP13(X0,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X1) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK16,X2),X3),s(fun(sK16,bool),sK20),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1067,f948])).
fof(f5299,plain,(
  ( ! [X14,X12,X13,X11] : (sP13(X11,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)),X12) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1067,f678])).
fof(f5298,plain,(
  ( ! [X10,X8,X7,X9] : (sP13(X7,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)),X8) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) )),
  inference(superposition,[],[f1067,f680])).
fof(f5297,plain,(
  ( ! [X2,X0,X3,X1] : (sP13(X0,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)),X1) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1067,f948])).
fof(f5288,plain,(
  ( ! [X0] : (sP14('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0) | ~p(s(bool,X0))) ) | $spl178),
  inference(cnf_transformation,[],[f5288_D])).
fof(f5288_D,plain,(
  ( ! [X0] : (sP14('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0) | ~p(s(bool,X0))) ) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f5281,plain,(
  ( ! [X10,X11,X9] : (sP14('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X9) | ~p(s(bool,X9)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X10),X11),s(fun(sK17,bool),sK18),s(fun(X10,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2097,f678])).
fof(f2097,plain,(
  ( ! [X200] : (sP14('const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)),X200) | ~p(s(bool,X200))) ) | $spl5),
  inference(resolution,[],[f1135,f844])).
fof(f5280,plain,(
  ( ! [X6,X8,X7] : (sP14('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X6) | ~p(s(bool,X6)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X7,sK17),X8),s(fun(X7,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f2097,f680])).
fof(f5279,plain,(
  ( ! [X2,X0,X1] : (sP14('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0) | ~p(s(bool,X0)) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X1),X2),s(fun(sK17,bool),sK18),s(fun(X1,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2097,f948])).
fof(f5266,plain,(
  ( ! [X14,X12,X13,X11] : (sP12('const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X11,X12) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X13),X14),s(fun(sK16,bool),sK20),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1060,f678])).
fof(f5265,plain,(
  ( ! [X10,X8,X7,X9] : (sP12('const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X7,X8) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK16),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) )),
  inference(superposition,[],[f1060,f680])).
fof(f5264,plain,(
  ( ! [X2,X0,X3,X1] : (sP12('const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X0,X1) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK16,X2),X3),s(fun(sK16,bool),sK20),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1060,f948])).
fof(f5259,plain,(
  ( ! [X14,X12,X13,X11] : (sP12('const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)),X11,X12) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1060,f678])).
fof(f5258,plain,(
  ( ! [X10,X8,X7,X9] : (sP12('const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)),X7,X8) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) )),
  inference(superposition,[],[f1060,f680])).
fof(f5257,plain,(
  ( ! [X2,X0,X3,X1] : (sP12('const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)),X0,X1) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1060,f948])).
fof(f5248,plain,(
  ( ! [X0] : (sP14(X0,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X0))) ) | $spl172),
  inference(cnf_transformation,[],[f5248_D])).
fof(f5248_D,plain,(
  ( ! [X0] : (sP14(X0,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X0))) ) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f5241,plain,(
  ( ! [X10,X11,X9] : (sP14(X9,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X9)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X10),X11),s(fun(sK17,bool),sK18),s(fun(X10,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2096,f678])).
fof(f2096,plain,(
  ( ! [X199] : (sP14(X199,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | ~p(s(bool,X199))) ) | $spl5),
  inference(resolution,[],[f1135,f841])).
fof(f5240,plain,(
  ( ! [X6,X8,X7] : (sP14(X6,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X6)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X7,sK17),X8),s(fun(X7,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f2096,f680])).
fof(f5239,plain,(
  ( ! [X2,X0,X1] : (sP14(X0,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X0)) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X1),X2),s(fun(sK17,bool),sK18),s(fun(X1,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2096,f948])).
fof(f5232,plain,(
  ( ! [X0] : (~sP14('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0) | p(s(bool,X0))) ) | $spl170),
  inference(cnf_transformation,[],[f5232_D])).
fof(f5232_D,plain,(
  ( ! [X0] : (~sP14('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0) | p(s(bool,X0))) ) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f5225,plain,(
  ( ! [X10,X11,X9] : (~sP14('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X9) | p(s(bool,X9)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X10),X11),s(fun(sK17,bool),sK18),s(fun(X10,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2095,f678])).
fof(f2095,plain,(
  ( ! [X198] : (~sP14('const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)),X198) | p(s(bool,X198))) ) | $spl5),
  inference(resolution,[],[f1135,f839])).
fof(f5224,plain,(
  ( ! [X6,X8,X7] : (~sP14('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X6) | p(s(bool,X6)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X7,sK17),X8),s(fun(X7,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f2095,f680])).
fof(f5223,plain,(
  ( ! [X2,X0,X1] : (~sP14('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0) | p(s(bool,X0)) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X1),X2),s(fun(sK17,bool),sK18),s(fun(X1,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2095,f948])).
fof(f5218,plain,(
  ( ! [X0] : (~sP14(X0,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X0))) ) | $spl168),
  inference(cnf_transformation,[],[f5218_D])).
fof(f5218_D,plain,(
  ( ! [X0] : (~sP14(X0,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X0))) ) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f5211,plain,(
  ( ! [X10,X11,X9] : (~sP14(X9,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X9)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X10),X11),s(fun(sK17,bool),sK18),s(fun(X10,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2094,f678])).
fof(f2094,plain,(
  ( ! [X197] : (~sP14(X197,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | p(s(bool,X197))) ) | $spl5),
  inference(resolution,[],[f1135,f839])).
fof(f5210,plain,(
  ( ! [X6,X8,X7] : (~sP14(X6,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X6)) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X7,sK17),X8),s(fun(X7,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f2094,f680])).
fof(f5209,plain,(
  ( ! [X2,X0,X1] : (~sP14(X0,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X0)) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X1),X2),s(fun(sK17,bool),sK18),s(fun(X1,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f2094,f948])).
fof(f5198,plain,(
  ( ! [X14,X12,X13,X11] : (sP11('const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X11,X12) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X13),X14),s(fun(sK16,bool),sK20),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1049,f678])).
fof(f5197,plain,(
  ( ! [X10,X8,X7,X9] : (sP11('const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X7,X8) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK16),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) )),
  inference(superposition,[],[f1049,f680])).
fof(f5196,plain,(
  ( ! [X2,X0,X3,X1] : (sP11('const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X0,X1) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK16,X2),X3),s(fun(sK16,bool),sK20),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1049,f948])).
fof(f5191,plain,(
  ( ! [X14,X12,X13,X11] : (sP11('const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)),X11,X12) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1049,f678])).
fof(f5190,plain,(
  ( ! [X10,X8,X7,X9] : (sP11('const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)),X7,X8) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) )),
  inference(superposition,[],[f1049,f680])).
fof(f5189,plain,(
  ( ! [X2,X0,X3,X1] : (sP11('const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)),X0,X1) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1049,f948])).
fof(f5171,plain,(
  ( ! [X14,X12,X13,X11] : (sP7('const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X11,X12) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X13),X14),s(fun(sK16,bool),sK20),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1042,f678])).
fof(f5170,plain,(
  ( ! [X10,X8,X7,X9] : (sP7('const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X7,X8) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK16),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) )),
  inference(superposition,[],[f1042,f680])).
fof(f5169,plain,(
  ( ! [X2,X0,X3,X1] : (sP7('const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X0,X1) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK16,X2),X3),s(fun(sK16,bool),sK20),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1042,f948])).
fof(f5164,plain,(
  ( ! [X14,X12,X13,X11] : (sP7('const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)),X11,X12) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1042,f678])).
fof(f5163,plain,(
  ( ! [X10,X8,X7,X9] : (sP7('const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)),X7,X8) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) )),
  inference(superposition,[],[f1042,f680])).
fof(f5162,plain,(
  ( ! [X2,X0,X3,X1] : (sP7('const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)),X0,X1) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1042,f948])).
fof(f5094,plain,(
  ( ! [X14,X12,X13,X11] : (sP7(X11,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X12) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X13),X14),s(fun(sK16,bool),sK20),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1041,f678])).
fof(f5093,plain,(
  ( ! [X10,X8,X7,X9] : (sP7(X7,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X8) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK16),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) )),
  inference(superposition,[],[f1041,f680])).
fof(f5092,plain,(
  ( ! [X2,X0,X3,X1] : (sP7(X0,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X1) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK16,X2),X3),s(fun(sK16,bool),sK20),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1041,f948])).
fof(f5087,plain,(
  ( ! [X14,X12,X13,X11] : (sP7(X11,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)),X12) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1041,f678])).
fof(f5086,plain,(
  ( ! [X10,X8,X7,X9] : (sP7(X7,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)),X8) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) )),
  inference(superposition,[],[f1041,f680])).
fof(f5085,plain,(
  ( ! [X2,X0,X3,X1] : (sP7(X0,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)),X1) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1041,f948])).
fof(f4987,plain,(
  ( ! [X14,X12,X13,X11] : (sP7(X11,X12,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X13),X14),s(fun(sK16,bool),sK20),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1040,f678])).
fof(f4986,plain,(
  ( ! [X10,X8,X7,X9] : (sP7(X7,X8,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK16),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) )),
  inference(superposition,[],[f1040,f680])).
fof(f4985,plain,(
  ( ! [X2,X0,X3,X1] : (sP7(X0,X1,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK16,X2),X3),s(fun(sK16,bool),sK20),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1040,f948])).
fof(f4980,plain,(
  ( ! [X14,X12,X13,X11] : (sP7(X11,X12,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1040,f678])).
fof(f4979,plain,(
  ( ! [X10,X8,X7,X9] : (sP7(X7,X8,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) )),
  inference(superposition,[],[f1040,f680])).
fof(f4978,plain,(
  ( ! [X2,X0,X3,X1] : (sP7(X0,X1,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1040,f948])).
fof(f4946,plain,(
  ( ! [X14,X12,X13,X11] : (sP8(X11,X12,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X13),X14),s(fun(sK16,bool),sK20),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1039,f678])).
fof(f4945,plain,(
  ( ! [X10,X8,X7,X9] : (sP8(X7,X8,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK16),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) )),
  inference(superposition,[],[f1039,f680])).
fof(f4944,plain,(
  ( ! [X2,X0,X3,X1] : (sP8(X0,X1,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK16,X2),X3),s(fun(sK16,bool),sK20),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1039,f948])).
fof(f4939,plain,(
  ( ! [X14,X12,X13,X11] : (sP8(X11,X12,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1039,f678])).
fof(f4938,plain,(
  ( ! [X10,X8,X7,X9] : (sP8(X7,X8,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) )),
  inference(superposition,[],[f1039,f680])).
fof(f4937,plain,(
  ( ! [X2,X0,X3,X1] : (sP8(X0,X1,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1039,f948])).
fof(f4921,plain,(
  ( ! [X14,X12,X13,X11] : (sP9(X11,X12,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X13),X14),s(fun(sK16,bool),sK20),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1036,f678])).
fof(f4920,plain,(
  ( ! [X10,X8,X7,X9] : (sP9(X7,X8,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK16),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) )),
  inference(superposition,[],[f1036,f680])).
fof(f4919,plain,(
  ( ! [X2,X0,X3,X1] : (sP9(X0,X1,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK16,X2),X3),s(fun(sK16,bool),sK20),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1036,f948])).
fof(f4914,plain,(
  ( ! [X14,X12,X13,X11] : (sP9(X11,X12,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1036,f678])).
fof(f4913,plain,(
  ( ! [X10,X8,X7,X9] : (sP9(X7,X8,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) )),
  inference(superposition,[],[f1036,f680])).
fof(f4912,plain,(
  ( ! [X2,X0,X3,X1] : (sP9(X0,X1,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1036,f948])).
fof(f4814,plain,(
  ( ! [X14,X12,X13,X11] : (sP5(X11,X12,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X13),X14),s(fun(sK16,bool),sK20),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1023,f678])).
fof(f4813,plain,(
  ( ! [X10,X8,X7,X9] : (sP5(X7,X8,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK16),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) )),
  inference(superposition,[],[f1023,f680])).
fof(f4812,plain,(
  ( ! [X2,X0,X3,X1] : (sP5(X0,X1,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK16,X2),X3),s(fun(sK16,bool),sK20),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1023,f948])).
fof(f4807,plain,(
  ( ! [X14,X12,X13,X11] : (sP5(X11,X12,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1023,f678])).
fof(f4806,plain,(
  ( ! [X10,X8,X7,X9] : (sP5(X7,X8,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) )),
  inference(superposition,[],[f1023,f680])).
fof(f4805,plain,(
  ( ! [X2,X0,X3,X1] : (sP5(X0,X1,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1023,f948])).
fof(f4723,plain,(
  ( ! [X14,X12,X13,X11] : (sP3(X11,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X12) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X13),X14),s(fun(sK16,bool),sK20),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1007,f678])).
fof(f4722,plain,(
  ( ! [X10,X8,X7,X9] : (sP3(X7,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X8) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK16),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) )),
  inference(superposition,[],[f1007,f680])).
fof(f4721,plain,(
  ( ! [X2,X0,X3,X1] : (sP3(X0,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X1) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK16,X2),X3),s(fun(sK16,bool),sK20),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1007,f948])).
fof(f4716,plain,(
  ( ! [X14,X12,X13,X11] : (sP3(X11,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)),X12) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1007,f678])).
fof(f4715,plain,(
  ( ! [X10,X8,X7,X9] : (sP3(X7,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)),X8) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) )),
  inference(superposition,[],[f1007,f680])).
fof(f4714,plain,(
  ( ! [X2,X0,X3,X1] : (sP3(X0,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)),X1) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1007,f948])).
fof(f4697,plain,(
  ( ! [X14,X12,X13,X11] : (sP3(X11,X12,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X13),X14),s(fun(sK16,bool),sK20),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1006,f678])).
fof(f4696,plain,(
  ( ! [X10,X8,X7,X9] : (sP3(X7,X8,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK16),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) )),
  inference(superposition,[],[f1006,f680])).
fof(f4695,plain,(
  ( ! [X2,X0,X3,X1] : (sP3(X0,X1,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK16,X2),X3),s(fun(sK16,bool),sK20),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1006,f948])).
fof(f4690,plain,(
  ( ! [X14,X12,X13,X11] : (sP3(X11,X12,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1006,f678])).
fof(f4689,plain,(
  ( ! [X10,X8,X7,X9] : (sP3(X7,X8,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) )),
  inference(superposition,[],[f1006,f680])).
fof(f4688,plain,(
  ( ! [X2,X0,X3,X1] : (sP3(X0,X1,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f1006,f948])).
fof(f4309,plain,(
  ( ! [X4] : (s(bool,X4) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X4))) ) | $spl124),
  inference(cnf_transformation,[],[f4309_D])).
fof(f4309_D,plain,(
  ( ! [X4] : (s(bool,X4) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X4))) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f4640,plain,(
  ( ! [X225,X226] : (p(s(bool,X225)) | ~p(s(bool,X226)) | s(bool,X226) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) | $spl77),
  inference(resolution,[],[f2297,f957])).
fof(f4538,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | s(bool,X24) != s(bool,X26)) ) | $spl77),
  inference(resolution,[],[f2297,f706])).
fof(f4536,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) | $spl77),
  inference(resolution,[],[f2297,f705])).
fof(f4534,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) | $spl77),
  inference(resolution,[],[f2297,f704])).
fof(f4533,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | s(bool,X9) != s(bool,X11)) ) | $spl77),
  inference(resolution,[],[f2297,f702])).
fof(f4531,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) | $spl77),
  inference(resolution,[],[f2297,f700])).
fof(f4505,plain,(
  ( ! [X22] : (p(s(bool,X22)) | s(bool,X22) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) | $spl130),
  inference(cnf_transformation,[],[f4505_D])).
fof(f4505_D,plain,(
  ( ! [X22] : (p(s(bool,X22)) | s(bool,X22) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f4471,plain,(
  ( ! [X259,X260] : (p(s(bool,X259)) | p(s(bool,X260)) | s(bool,X260) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) | $spl76),
  inference(resolution,[],[f2296,f957])).
fof(f4364,plain,(
  ( ! [X43,X44,X42] : (p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X42) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | s(bool,X43) != s(bool,X44)) ) | $spl76),
  inference(resolution,[],[f2296,f706])).
fof(f4361,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X33)) | ~p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X35) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) | $spl76),
  inference(resolution,[],[f2296,f705])).
fof(f4360,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | p(s(bool,X31)) | s(bool,X30) != s(bool,X32) | s(bool,X31) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) | $spl76),
  inference(resolution,[],[f2296,f704])).
fof(f4359,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | s(bool,X28) != s(bool,X29)) ) | $spl76),
  inference(resolution,[],[f2296,f702])).
fof(f4357,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) | $spl76),
  inference(resolution,[],[f2296,f700])).
fof(f4340,plain,(
  ( ! [X14,X12,X13,X11] : (sP1(X11,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X12) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X13),X14),s(fun(sK16,bool),sK20),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f993,f678])).
fof(f4339,plain,(
  ( ! [X10,X8,X7,X9] : (sP1(X7,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X8) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK16),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) )),
  inference(superposition,[],[f993,f680])).
fof(f4338,plain,(
  ( ! [X2,X0,X3,X1] : (sP1(X0,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X1) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK16,X2),X3),s(fun(sK16,bool),sK20),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f993,f948])).
fof(f4333,plain,(
  ( ! [X14,X12,X13,X11] : (sP1(X11,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)),X12) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f993,f678])).
fof(f4332,plain,(
  ( ! [X10,X8,X7,X9] : (sP1(X7,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)),X8) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) )),
  inference(superposition,[],[f993,f680])).
fof(f4331,plain,(
  ( ! [X2,X0,X3,X1] : (sP1(X0,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20)),X1) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f993,f948])).
fof(f4282,plain,(
  ( ! [X225,X226] : (p(s(bool,X225)) | ~p(s(bool,X226)) | s(bool,X226) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) | $spl77),
  inference(resolution,[],[f2297,f957])).
fof(f4180,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | s(bool,X24) != s(bool,X26)) ) | $spl77),
  inference(resolution,[],[f2297,f706])).
fof(f4178,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) | $spl77),
  inference(resolution,[],[f2297,f705])).
fof(f4176,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) | $spl77),
  inference(resolution,[],[f2297,f704])).
fof(f4175,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | s(bool,X9) != s(bool,X11)) ) | $spl77),
  inference(resolution,[],[f2297,f702])).
fof(f4173,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) | $spl77),
  inference(resolution,[],[f2297,f700])).
fof(f4171,plain,(
  ( ! [X0,X1] : (sP8('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X0,X1)) ) | $spl122),
  inference(cnf_transformation,[],[f4171_D])).
fof(f4171_D,plain,(
  ( ! [X0,X1] : (sP8('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X0,X1)) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f4163,plain,(
  ( ! [X14,X12,X13,X11] : (sP8('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X11,X12) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X13),X14),s(fun(sK16,bool),sK20),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f2209,f678])).
fof(f2209,plain,(
  ( ! [X134,X133] : (sP8('const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)),X133,X134)) ) | $spl9),
  inference(resolution,[],[f1139,f772])).
fof(f4162,plain,(
  ( ! [X10,X8,X7,X9] : (sP8('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X7,X8) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK16),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) ) | $spl9),
  inference(superposition,[],[f2209,f680])).
fof(f4161,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X0,X1) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK16,X2),X3),s(fun(sK16,bool),sK20),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f2209,f948])).
fof(f4156,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X1)) ) | $spl120),
  inference(cnf_transformation,[],[f4156_D])).
fof(f4156_D,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X1)) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f4148,plain,(
  ( ! [X14,X12,X13,X11] : (sP8(X11,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X12) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X13),X14),s(fun(sK16,bool),sK20),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f2208,f678])).
fof(f2208,plain,(
  ( ! [X132,X131] : (sP8(X131,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)),X132)) ) | $spl9),
  inference(resolution,[],[f1139,f771])).
fof(f4147,plain,(
  ( ! [X10,X8,X7,X9] : (sP8(X7,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X8) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK16),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) ) | $spl9),
  inference(superposition,[],[f2208,f680])).
fof(f4146,plain,(
  ( ! [X2,X0,X3,X1] : (sP8(X0,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X1) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK16,X2),X3),s(fun(sK16,bool),sK20),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f2208,f948])).
fof(f4141,plain,(
  ( ! [X0,X1] : (sP9('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X0,X1)) ) | $spl118),
  inference(cnf_transformation,[],[f4141_D])).
fof(f4141_D,plain,(
  ( ! [X0,X1] : (sP9('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X0,X1)) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f4133,plain,(
  ( ! [X14,X12,X13,X11] : (sP9('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X11,X12) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X13),X14),s(fun(sK16,bool),sK20),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f2206,f678])).
fof(f2206,plain,(
  ( ! [X127,X128] : (sP9('const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)),X127,X128)) ) | $spl9),
  inference(resolution,[],[f1139,f768])).
fof(f4132,plain,(
  ( ! [X10,X8,X7,X9] : (sP9('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X7,X8) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK16),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) ) | $spl9),
  inference(superposition,[],[f2206,f680])).
fof(f4131,plain,(
  ( ! [X2,X0,X3,X1] : (sP9('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X0,X1) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK16,X2),X3),s(fun(sK16,bool),sK20),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f2206,f948])).
fof(f4126,plain,(
  ( ! [X0,X1] : (sP9(X0,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X1)) ) | $spl116),
  inference(cnf_transformation,[],[f4126_D])).
fof(f4126_D,plain,(
  ( ! [X0,X1] : (sP9(X0,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X1)) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f4119,plain,(
  ( ! [X14,X12,X13,X11] : (sP9(X11,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X12) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X13),X14),s(fun(sK16,bool),sK20),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f2205,f678])).
fof(f2205,plain,(
  ( ! [X125,X126] : (sP9(X125,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)),X126)) ) | $spl9),
  inference(resolution,[],[f1139,f767])).
fof(f4118,plain,(
  ( ! [X10,X8,X7,X9] : (sP9(X7,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X8) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK16),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) ) | $spl9),
  inference(superposition,[],[f2205,f680])).
fof(f4117,plain,(
  ( ! [X2,X0,X3,X1] : (sP9(X0,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X1) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK16,X2),X3),s(fun(sK16,bool),sK20),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f2205,f948])).
fof(f4106,plain,(
  ( ! [X8,X7] : (~sP0(sK16,sK17,sK19,'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0')),sK18) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X7),X8),s(fun(sK17,bool),sK18),s(fun(X7,bool),'const.pred_set.EMPTY_0'))))) ) | $spl41),
  inference(superposition,[],[f1337,f678])).
fof(f4105,plain,(
  ( ! [X6,X5] : (~sP0(sK16,sK17,sK19,'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0')),sK18) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X5,sK17),X6),s(fun(X5,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl41),
  inference(superposition,[],[f1337,f680])).
fof(f4104,plain,(
  ( ! [X0,X1] : (~sP0(sK16,sK17,sK19,'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0')),sK18) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X0),X1),s(fun(sK17,bool),sK18),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl41),
  inference(superposition,[],[f1337,f948])).
fof(f4090,plain,(
  ( ! [X0,X1] : (sP5('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X0,X1)) ) | $spl112),
  inference(cnf_transformation,[],[f4090_D])).
fof(f4090_D,plain,(
  ( ! [X0,X1] : (sP5('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X0,X1)) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f4084,plain,(
  ( ! [X14,X12,X13,X11] : (sP5('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X11,X12) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X13),X14),s(fun(sK16,bool),sK20),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f2193,f678])).
fof(f2193,plain,(
  ( ! [X101,X102] : (sP5('const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)),X101,X102)) ) | $spl9),
  inference(resolution,[],[f1139,f751])).
fof(f4083,plain,(
  ( ! [X10,X8,X7,X9] : (sP5('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X7,X8) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK16),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) ) | $spl9),
  inference(superposition,[],[f2193,f680])).
fof(f4082,plain,(
  ( ! [X2,X0,X3,X1] : (sP5('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X0,X1) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK16,X2),X3),s(fun(sK16,bool),sK20),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f2193,f948])).
fof(f4076,plain,(
  ( ! [X0,X1] : (sP5(X0,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X1)) ) | $spl110),
  inference(cnf_transformation,[],[f4076_D])).
fof(f4076_D,plain,(
  ( ! [X0,X1] : (sP5(X0,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X1)) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f4070,plain,(
  ( ! [X14,X12,X13,X11] : (sP5(X11,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X12) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X13),X14),s(fun(sK16,bool),sK20),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f2192,f678])).
fof(f2192,plain,(
  ( ! [X99,X100] : (sP5(X99,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)),X100)) ) | $spl9),
  inference(resolution,[],[f1139,f750])).
fof(f4069,plain,(
  ( ! [X10,X8,X7,X9] : (sP5(X7,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X8) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK16),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) ) | $spl9),
  inference(superposition,[],[f2192,f680])).
fof(f4068,plain,(
  ( ! [X2,X0,X3,X1] : (sP5(X0,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X1) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK16,X2),X3),s(fun(sK16,bool),sK20),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f2192,f948])).
fof(f4062,plain,(
  ( ! [X0,X1] : (sP3('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X0,X1)) ) | $spl108),
  inference(cnf_transformation,[],[f4062_D])).
fof(f4062_D,plain,(
  ( ! [X0,X1] : (sP3('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X0,X1)) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f4055,plain,(
  ( ! [X14,X12,X13,X11] : (sP3('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X11,X12) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X13),X14),s(fun(sK16,bool),sK20),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f2179,f678])).
fof(f2179,plain,(
  ( ! [X74,X73] : (sP3('const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)),X73,X74)) ) | $spl9),
  inference(resolution,[],[f1139,f734])).
fof(f4054,plain,(
  ( ! [X10,X8,X7,X9] : (sP3('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X7,X8) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK16),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) ) | $spl9),
  inference(superposition,[],[f2179,f680])).
fof(f4053,plain,(
  ( ! [X2,X0,X3,X1] : (sP3('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X0,X1) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK16,X2),X3),s(fun(sK16,bool),sK20),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f2179,f948])).
fof(f4047,plain,(
  ( ! [X0,X1] : (sP1('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X0,X1)) ) | $spl106),
  inference(cnf_transformation,[],[f4047_D])).
fof(f4047_D,plain,(
  ( ! [X0,X1] : (sP1('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X0,X1)) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f4041,plain,(
  ( ! [X14,X12,X13,X11] : (sP1('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X11,X12) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X13),X14),s(fun(sK16,bool),sK20),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f2163,f678])).
fof(f2163,plain,(
  ( ! [X41,X42] : (sP1('const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)),X41,X42)) ) | $spl9),
  inference(resolution,[],[f1139,f717])).
fof(f4040,plain,(
  ( ! [X10,X8,X7,X9] : (sP1('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X7,X8) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK16),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) ) | $spl9),
  inference(superposition,[],[f2163,f680])).
fof(f4039,plain,(
  ( ! [X2,X0,X3,X1] : (sP1('const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')),X0,X1) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK16,X2),X3),s(fun(sK16,bool),sK20),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f2163,f948])).
fof(f4033,plain,(
  ( ! [X0,X1] : (sP1(X0,X1,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) | $spl104),
  inference(cnf_transformation,[],[f4033_D])).
fof(f4033_D,plain,(
  ( ! [X0,X1] : (sP1(X0,X1,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f4027,plain,(
  ( ! [X14,X12,X13,X11] : (sP1(X11,X12,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X13),X14),s(fun(sK16,bool),sK20),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f2162,f678])).
fof(f2162,plain,(
  ( ! [X39,X40] : (sP1(X39,X40,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)))) ) | $spl9),
  inference(resolution,[],[f1139,f715])).
fof(f4026,plain,(
  ( ! [X10,X8,X7,X9] : (sP1(X7,X8,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK16),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) ) | $spl9),
  inference(superposition,[],[f2162,f680])).
fof(f4025,plain,(
  ( ! [X2,X0,X3,X1] : (sP1(X0,X1,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK16,X2),X3),s(fun(sK16,bool),sK20),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f2162,f948])).
fof(f3981,plain,(
  ( ! [X225,X226] : (p(s(bool,X225)) | ~p(s(bool,X226)) | s(bool,X226) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | $spl69),
  inference(resolution,[],[f1824,f957])).
fof(f3879,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | s(bool,X24) != s(bool,X26)) ) | $spl69),
  inference(resolution,[],[f1824,f706])).
fof(f3877,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | $spl69),
  inference(resolution,[],[f1824,f705])).
fof(f3875,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | $spl69),
  inference(resolution,[],[f1824,f704])).
fof(f3874,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | s(bool,X9) != s(bool,X11)) ) | $spl69),
  inference(resolution,[],[f1824,f702])).
fof(f3872,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | $spl69),
  inference(resolution,[],[f1824,f700])).
fof(f3741,plain,(
  ( ! [X259,X260] : (p(s(bool,X259)) | p(s(bool,X260)) | s(bool,X260) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | $spl68),
  inference(resolution,[],[f1823,f957])).
fof(f3634,plain,(
  ( ! [X43,X44,X42] : (p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X42) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | s(bool,X43) != s(bool,X44)) ) | $spl68),
  inference(resolution,[],[f1823,f706])).
fof(f3604,plain,(
  ( ! [X8] : (s(bool,X8) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | $spl100),
  inference(cnf_transformation,[],[f3604_D])).
fof(f3604_D,plain,(
  ( ! [X8] : (s(bool,X8) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f1145,plain,(
  ( ! [X15,X16] : (s(bool,X15) != s(bool,X16) | p(s(bool,X15)) | ~p(s(bool,X16))) ) | $spl18),
  inference(cnf_transformation,[],[f1145_D])).
fof(f1145_D,plain,(
  ( ! [X15,X16] : (s(bool,X15) != s(bool,X16) | p(s(bool,X15)) | ~p(s(bool,X16))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f3633,plain,(
  ( ! [X39,X41,X40] : (p(s(bool,X39)) | ~p(s(bool,X40)) | s(bool,X39) != s(bool,X40) | s(bool,X41) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | $spl68),
  inference(resolution,[],[f1823,f706])).
fof(f3631,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X33)) | ~p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X35) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | $spl68),
  inference(resolution,[],[f1823,f705])).
fof(f3630,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | p(s(bool,X31)) | s(bool,X30) != s(bool,X32) | s(bool,X31) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | $spl68),
  inference(resolution,[],[f1823,f704])).
fof(f3629,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | s(bool,X28) != s(bool,X29)) ) | $spl68),
  inference(resolution,[],[f1823,f702])).
fof(f3628,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X25) | s(bool,X26) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | $spl68),
  inference(resolution,[],[f1823,f702])).
fof(f3627,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | $spl68),
  inference(resolution,[],[f1823,f700])).
fof(f3576,plain,(
  ( ! [X225,X226] : (p(s(bool,X225)) | ~p(s(bool,X226)) | s(bool,X226) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | $spl69),
  inference(resolution,[],[f1824,f957])).
fof(f3474,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | s(bool,X24) != s(bool,X26)) ) | $spl69),
  inference(resolution,[],[f1824,f706])).
fof(f3472,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | $spl69),
  inference(resolution,[],[f1824,f705])).
fof(f3471,plain,(
  ( ! [X17,X15,X16] : (~p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X17) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | s(bool,X15) != s(bool,X16)) ) | $spl69),
  inference(resolution,[],[f1824,f704])).
fof(f3470,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | $spl69),
  inference(resolution,[],[f1824,f704])).
fof(f3469,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | s(bool,X9) != s(bool,X11)) ) | $spl69),
  inference(resolution,[],[f1824,f702])).
fof(f3468,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | p(s(bool,X7)) | s(bool,X8) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | s(bool,X6) != s(bool,X7)) ) | $spl69),
  inference(resolution,[],[f1824,f700])).
fof(f3467,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) ) | $spl69),
  inference(resolution,[],[f1824,f700])).
fof(f3457,plain,(
  ( ! [X14,X12,X13,X11] : (sP8('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X11,X12) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f1736,f678])).
fof(f3456,plain,(
  ( ! [X10,X8,X7,X9] : (sP8('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X7,X8) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f1736,f680])).
fof(f3455,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0,X1) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f1736,f948])).
fof(f3442,plain,(
  ( ! [X14,X12,X13,X11] : (sP8(X11,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X12) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f1735,f678])).
fof(f3441,plain,(
  ( ! [X10,X8,X7,X9] : (sP8(X7,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X8) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f1735,f680])).
fof(f3440,plain,(
  ( ! [X2,X0,X3,X1] : (sP8(X0,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X1) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f1735,f948])).
fof(f3254,plain,(
  $false | $spl20),
  inference(resolution,[],[f3120,f950])).
fof(f3120,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl20),
  inference(equality_resolution,[],[f1146])).
fof(f3260,plain,(
  $false | $spl20),
  inference(resolution,[],[f3120,f951])).
fof(f3261,plain,(
  $false | $spl20),
  inference(resolution,[],[f3120,f647])).
fof(f3277,plain,(
  $false | $spl20),
  inference(resolution,[],[f3120,f686])).
fof(f3281,plain,(
  $false | $spl20),
  inference(resolution,[],[f3120,f698])).
fof(f3071,plain,(
  ( ! [X14,X12,X13,X11] : (sP9('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X11,X12) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f1733,f678])).
fof(f3070,plain,(
  ( ! [X10,X8,X7,X9] : (sP9('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X7,X8) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f1733,f680])).
fof(f3069,plain,(
  ( ! [X2,X0,X3,X1] : (sP9('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0,X1) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f1733,f948])).
fof(f3057,plain,(
  ( ! [X14,X12,X13,X11] : (sP9(X11,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X12) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f1732,f678])).
fof(f3056,plain,(
  ( ! [X10,X8,X7,X9] : (sP9(X7,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X8) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f1732,f680])).
fof(f3055,plain,(
  ( ! [X2,X0,X3,X1] : (sP9(X0,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X1) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f1732,f948])).
fof(f3044,plain,(
  ( ! [X14,X12,X13,X11] : (sP5('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X11,X12) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f1720,f678])).
fof(f3043,plain,(
  ( ! [X10,X8,X7,X9] : (sP5('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X7,X8) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f1720,f680])).
fof(f3042,plain,(
  ( ! [X2,X0,X3,X1] : (sP5('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0,X1) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f1720,f948])).
fof(f3030,plain,(
  ( ! [X14,X12,X13,X11] : (sP5(X11,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X12) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f1719,f678])).
fof(f3029,plain,(
  ( ! [X10,X8,X7,X9] : (sP5(X7,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X8) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f1719,f680])).
fof(f3028,plain,(
  ( ! [X2,X0,X3,X1] : (sP5(X0,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X1) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f1719,f948])).
fof(f2875,plain,(
  $false | $spl14),
  inference(resolution,[],[f2703,f648])).
fof(f2703,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(equality_resolution,[],[f1143])).
fof(f2952,plain,(
  $false | ($spl5 | $spl14)),
  inference(resolution,[],[f2703,f1135])).
fof(f2953,plain,(
  $false | ($spl9 | $spl14)),
  inference(resolution,[],[f2703,f1139])).
fof(f2655,plain,(
  ( ! [X14,X12,X13,X11] : (sP3('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X11,X12) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f1706,f678])).
fof(f2654,plain,(
  ( ! [X10,X8,X7,X9] : (sP3('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X7,X8) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f1706,f680])).
fof(f2653,plain,(
  ( ! [X2,X0,X3,X1] : (sP3('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0,X1) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f1706,f948])).
fof(f2641,plain,(
  ( ! [X14,X12,X13,X11] : (sP1('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X11,X12) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f1690,f678])).
fof(f2640,plain,(
  ( ! [X10,X8,X7,X9] : (sP1('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X7,X8) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f1690,f680])).
fof(f2639,plain,(
  ( ! [X2,X0,X3,X1] : (sP1('const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')),X0,X1) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f1690,f948])).
fof(f2627,plain,(
  ( ! [X14,X12,X13,X11] : (sP1(X11,X12,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X13),X14),s(fun(sK17,bool),sK18),s(fun(X13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f1689,f678])).
fof(f2626,plain,(
  ( ! [X10,X8,X7,X9] : (sP1(X7,X8,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X9,sK17),X10),s(fun(X9,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f1689,f680])).
fof(f2625,plain,(
  ( ! [X2,X0,X3,X1] : (sP1(X0,X1,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X2),X3),s(fun(sK17,bool),sK18),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f1689,f948])).
fof(f2297,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | $spl77),
  inference(cnf_transformation,[],[f2297_D])).
fof(f2297_D,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) <=> ~$spl77),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl77])])).
fof(f2588,plain,(
  ( ! [X8,X7] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X7),X8),s(fun(sK16,bool),sK20),s(fun(X7,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f1139,f678])).
fof(f2587,plain,(
  ( ! [X6,X5] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X5,sK16),X6),s(fun(X5,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) ) | $spl9),
  inference(superposition,[],[f1139,f680])).
fof(f2586,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK16,X0),X1),s(fun(sK16,bool),sK20),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f1139,f948])).
fof(f2281,plain,(
  ( ! [X4] : (s(bool,X4) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) | ~p(s(bool,X4))) ) | $spl72),
  inference(cnf_transformation,[],[f2281_D])).
fof(f2281_D,plain,(
  ( ! [X4] : (s(bool,X4) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) | ~p(s(bool,X4))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f2583,plain,(
  ( ! [X225,X226] : (p(s(bool,X225)) | ~p(s(bool,X226)) | s(bool,X226) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)))) ) | $spl9),
  inference(resolution,[],[f1139,f957])).
fof(f2481,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) | s(bool,X24) != s(bool,X26)) ) | $spl9),
  inference(resolution,[],[f1139,f706])).
fof(f2479,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)))) ) | $spl9),
  inference(resolution,[],[f1139,f705])).
fof(f2282,plain,(
  ( ! [X8] : (s(bool,X8) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)))) ) | $spl74),
  inference(cnf_transformation,[],[f2282_D])).
fof(f2282_D,plain,(
  ( ! [X8] : (s(bool,X8) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f2478,plain,(
  ( ! [X17,X15,X16] : (~p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X17) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) | s(bool,X15) != s(bool,X16)) ) | $spl9),
  inference(resolution,[],[f1139,f704])).
fof(f2477,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)))) ) | $spl9),
  inference(resolution,[],[f1139,f704])).
fof(f2476,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) | s(bool,X9) != s(bool,X11)) ) | $spl9),
  inference(resolution,[],[f1139,f702])).
fof(f2475,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | p(s(bool,X7)) | s(bool,X8) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) | s(bool,X6) != s(bool,X7)) ) | $spl9),
  inference(resolution,[],[f1139,f700])).
fof(f2474,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)))) ) | $spl9),
  inference(resolution,[],[f1139,f700])).
fof(f2296,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | $spl76),
  inference(cnf_transformation,[],[f2296_D])).
fof(f2296_D,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f2429,plain,(
  ( ! [X8,X7] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X7),X8),s(fun(sK16,bool),sK20),s(fun(X7,bool),'const.pred_set.EMPTY_0'))))) ) | $spl8),
  inference(superposition,[],[f1138,f678])).
fof(f1138,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)))) | $spl8),
  inference(cnf_transformation,[],[f1138_D])).
fof(f1138_D,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f2428,plain,(
  ( ! [X6,X5] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X5,sK16),X6),s(fun(X5,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) ) | $spl8),
  inference(superposition,[],[f1138,f680])).
fof(f2427,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK16,X0),X1),s(fun(sK16,bool),sK20),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl8),
  inference(superposition,[],[f1138,f948])).
fof(f2448,plain,(
  ( ! [X22] : (p(s(bool,X22)) | s(bool,X22) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)))) ) | $spl78),
  inference(cnf_transformation,[],[f2448_D])).
fof(f2448_D,plain,(
  ( ! [X22] : (p(s(bool,X22)) | s(bool,X22) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f2421,plain,(
  ( ! [X259,X260] : (p(s(bool,X259)) | p(s(bool,X260)) | s(bool,X260) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)))) ) | $spl8),
  inference(resolution,[],[f1138,f957])).
fof(f2314,plain,(
  ( ! [X43,X44,X42] : (p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X42) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) | s(bool,X43) != s(bool,X44)) ) | $spl8),
  inference(resolution,[],[f1138,f706])).
fof(f2313,plain,(
  ( ! [X39,X41,X40] : (p(s(bool,X39)) | ~p(s(bool,X40)) | s(bool,X39) != s(bool,X40) | s(bool,X41) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)))) ) | $spl8),
  inference(resolution,[],[f1138,f706])).
fof(f2311,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X33)) | ~p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X35) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)))) ) | $spl8),
  inference(resolution,[],[f1138,f705])).
fof(f2310,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | p(s(bool,X31)) | s(bool,X30) != s(bool,X32) | s(bool,X31) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)))) ) | $spl8),
  inference(resolution,[],[f1138,f704])).
fof(f2309,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) | s(bool,X28) != s(bool,X29)) ) | $spl8),
  inference(resolution,[],[f1138,f702])).
fof(f2308,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X25) | s(bool,X26) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)))) ) | $spl8),
  inference(resolution,[],[f1138,f702])).
fof(f2307,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)))) ) | $spl8),
  inference(resolution,[],[f1138,f700])).
fof(f2262,plain,(
  ( ! [X8,X7] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X7),X8),s(fun(sK16,bool),sK20),s(fun(X7,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f1139,f678])).
fof(f2261,plain,(
  ( ! [X6,X5] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X5,sK16),X6),s(fun(X5,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) ) | $spl9),
  inference(superposition,[],[f1139,f680])).
fof(f2260,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK16,X0),X1),s(fun(sK16,bool),sK20),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f1139,f948])).
fof(f2257,plain,(
  ( ! [X225,X226] : (p(s(bool,X225)) | ~p(s(bool,X226)) | s(bool,X226) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)))) ) | $spl9),
  inference(resolution,[],[f1139,f957])).
fof(f2155,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) | s(bool,X24) != s(bool,X26)) ) | $spl9),
  inference(resolution,[],[f1139,f706])).
fof(f2153,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)))) ) | $spl9),
  inference(resolution,[],[f1139,f705])).
fof(f2152,plain,(
  ( ! [X17,X15,X16] : (~p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X17) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) | s(bool,X15) != s(bool,X16)) ) | $spl9),
  inference(resolution,[],[f1139,f704])).
fof(f2151,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)))) ) | $spl9),
  inference(resolution,[],[f1139,f704])).
fof(f2150,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) | s(bool,X9) != s(bool,X11)) ) | $spl9),
  inference(resolution,[],[f1139,f702])).
fof(f2149,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | p(s(bool,X7)) | s(bool,X8) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))) | s(bool,X6) != s(bool,X7)) ) | $spl9),
  inference(resolution,[],[f1139,f700])).
fof(f2148,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20)))) ) | $spl9),
  inference(resolution,[],[f1139,f700])).
fof(f2115,plain,(
  ( ! [X8,X7] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X7),X8),s(fun(sK17,bool),sK18),s(fun(X7,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f1135,f678])).
fof(f2114,plain,(
  ( ! [X6,X5] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X5,sK17),X6),s(fun(X5,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f1135,f680])).
fof(f2113,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X0),X1),s(fun(sK17,bool),sK18),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f1135,f948])).
fof(f2110,plain,(
  ( ! [X225,X226] : (p(s(bool,X225)) | ~p(s(bool,X226)) | s(bool,X226) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) ) | $spl5),
  inference(resolution,[],[f1135,f957])).
fof(f2008,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | s(bool,X24) != s(bool,X26)) ) | $spl5),
  inference(resolution,[],[f1135,f706])).
fof(f2006,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) ) | $spl5),
  inference(resolution,[],[f1135,f705])).
fof(f1809,plain,(
  ( ! [X8] : (s(bool,X8) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) ) | $spl66),
  inference(cnf_transformation,[],[f1809_D])).
fof(f1809_D,plain,(
  ( ! [X8] : (s(bool,X8) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f2005,plain,(
  ( ! [X17,X15,X16] : (~p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X17) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | s(bool,X15) != s(bool,X16)) ) | $spl5),
  inference(resolution,[],[f1135,f704])).
fof(f2004,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) ) | $spl5),
  inference(resolution,[],[f1135,f704])).
fof(f2003,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | s(bool,X9) != s(bool,X11)) ) | $spl5),
  inference(resolution,[],[f1135,f702])).
fof(f2002,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | p(s(bool,X7)) | s(bool,X8) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | s(bool,X6) != s(bool,X7)) ) | $spl5),
  inference(resolution,[],[f1135,f700])).
fof(f2001,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) ) | $spl5),
  inference(resolution,[],[f1135,f700])).
fof(f1956,plain,(
  ( ! [X8,X7] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X7),X8),s(fun(sK17,bool),sK18),s(fun(X7,bool),'const.pred_set.EMPTY_0'))))) ) | $spl4),
  inference(superposition,[],[f1134,f678])).
fof(f1955,plain,(
  ( ! [X6,X5] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X5,sK17),X6),s(fun(X5,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl4),
  inference(superposition,[],[f1134,f680])).
fof(f1954,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X0),X1),s(fun(sK17,bool),sK18),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl4),
  inference(superposition,[],[f1134,f948])).
fof(f1948,plain,(
  ( ! [X259,X260] : (p(s(bool,X259)) | p(s(bool,X260)) | s(bool,X260) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) ) | $spl4),
  inference(resolution,[],[f1134,f957])).
fof(f1841,plain,(
  ( ! [X43,X44,X42] : (p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X42) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | s(bool,X43) != s(bool,X44)) ) | $spl4),
  inference(resolution,[],[f1134,f706])).
fof(f1840,plain,(
  ( ! [X39,X41,X40] : (p(s(bool,X39)) | ~p(s(bool,X40)) | s(bool,X39) != s(bool,X40) | s(bool,X41) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) ) | $spl4),
  inference(resolution,[],[f1134,f706])).
fof(f1838,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X33)) | ~p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X35) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) ) | $spl4),
  inference(resolution,[],[f1134,f705])).
fof(f1837,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | p(s(bool,X31)) | s(bool,X30) != s(bool,X32) | s(bool,X31) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) ) | $spl4),
  inference(resolution,[],[f1134,f704])).
fof(f1836,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | s(bool,X28) != s(bool,X29)) ) | $spl4),
  inference(resolution,[],[f1134,f702])).
fof(f1835,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X25) | s(bool,X26) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) ) | $spl4),
  inference(resolution,[],[f1134,f702])).
fof(f1834,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) ) | $spl4),
  inference(resolution,[],[f1134,f700])).
fof(f1789,plain,(
  ( ! [X8,X7] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X7),X8),s(fun(sK17,bool),sK18),s(fun(X7,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f1135,f678])).
fof(f1788,plain,(
  ( ! [X6,X5] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X5,sK17),X6),s(fun(X5,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) ) | $spl5),
  inference(superposition,[],[f1135,f680])).
fof(f1787,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X0),X1),s(fun(sK17,bool),sK18),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl5),
  inference(superposition,[],[f1135,f948])).
fof(f1784,plain,(
  ( ! [X225,X226] : (p(s(bool,X225)) | ~p(s(bool,X226)) | s(bool,X226) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) ) | $spl5),
  inference(resolution,[],[f1135,f957])).
fof(f1682,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | s(bool,X24) != s(bool,X26)) ) | $spl5),
  inference(resolution,[],[f1135,f706])).
fof(f1680,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) ) | $spl5),
  inference(resolution,[],[f1135,f705])).
fof(f1679,plain,(
  ( ! [X17,X15,X16] : (~p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X17) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | s(bool,X15) != s(bool,X16)) ) | $spl5),
  inference(resolution,[],[f1135,f704])).
fof(f1678,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) ) | $spl5),
  inference(resolution,[],[f1135,f704])).
fof(f1677,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | s(bool,X9) != s(bool,X11)) ) | $spl5),
  inference(resolution,[],[f1135,f702])).
fof(f1676,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | p(s(bool,X7)) | s(bool,X8) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))) | s(bool,X6) != s(bool,X7)) ) | $spl5),
  inference(resolution,[],[f1135,f700])).
fof(f1675,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18)))) ) | $spl5),
  inference(resolution,[],[f1135,f700])).
fof(f1525,plain,(
  $false | $spl24),
  inference(resolution,[],[f1148,f648])).
fof(f1311,plain,(
  ( ! [X8,X9] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X8),X9),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))),s(fun(X8,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f648,f678])).
fof(f1310,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X6,sK16),X7),s(fun(X6,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))))))) )),
  inference(superposition,[],[f648,f680])).
fof(f1308,plain,(
  ( ! [X2,X1] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK16,X1),X2),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))),s(fun(X1,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f648,f948])).
fof(f1361,plain,(
  ( ! [X8,X7] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X7),X8),s(fun(sK17,bool),sK18),s(fun(X7,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(forward_demodulation,[],[f1302,f697])).
fof(f1302,plain,(
  ( ! [X8,X7] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,X7),X8),s(fun(sK17,bool),sK18),s(fun(X7,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f648,f678])).
fof(f1359,plain,(
  ( ! [X6,X5] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X5,sK17),X6),s(fun(X5,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) )),
  inference(forward_demodulation,[],[f1301,f697])).
fof(f1301,plain,(
  ( ! [X6,X5] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X5,sK17),X6),s(fun(X5,bool),'const.pred_set.EMPTY_0'),s(fun(sK17,bool),sK18))))) )),
  inference(superposition,[],[f648,f680])).
fof(f1355,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X0),X1),s(fun(sK17,bool),sK18),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(forward_demodulation,[],[f1300,f697])).
fof(f1300,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),'const.pred_set.EMPTY_0')))))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,X0),X1),s(fun(sK17,bool),sK18),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f648,f948])).
fof(f1293,plain,(
  ( ! [X225,X226] : (p(s(bool,X225)) | ~p(s(bool,X226)) | s(bool,X226) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18)))))) )),
  inference(resolution,[],[f648,f957])).
fof(f1191,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))))) | s(bool,X24) != s(bool,X26)) )),
  inference(resolution,[],[f648,f706])).
fof(f1189,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18)))))) )),
  inference(resolution,[],[f648,f705])).
fof(f1341,plain,(
  ( ! [X8] : (s(bool,X8) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18)))))) ) | $spl46),
  inference(cnf_transformation,[],[f1341_D])).
fof(f1341_D,plain,(
  ( ! [X8] : (s(bool,X8) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18)))))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f1188,plain,(
  ( ! [X17,X15,X16] : (~p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X17) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))))) | s(bool,X15) != s(bool,X16)) )),
  inference(resolution,[],[f648,f704])).
fof(f1187,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18)))))) )),
  inference(resolution,[],[f648,f704])).
fof(f1186,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))))) | s(bool,X9) != s(bool,X11)) )),
  inference(resolution,[],[f648,f702])).
fof(f1185,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | p(s(bool,X7)) | s(bool,X8) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))))) | s(bool,X6) != s(bool,X7)) )),
  inference(resolution,[],[f648,f700])).
fof(f1184,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18)))))) )),
  inference(resolution,[],[f648,f700])).
fof(f1180,plain,(
  ~sP0(sK16,sK17,sK19,'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18)),sK18) | ~p(s(bool,'const.bool.IN_2'(s(sK16,happ(s(fun(sK17,sK16),sK19),s(sK17,sK25('const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18)),sK18,sK19,sK16,sK17)))),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))))))),
  inference(resolution,[],[f648,f672])).
fof(f1179,plain,(
  ~sP0(sK16,sK17,sK19,'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18)),sK18) | p(s(bool,'const.bool.IN_2'(s(sK17,sK25('const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18)),sK18,sK19,sK16,sK17)),s(fun(sK17,bool),sK18))))),
  inference(resolution,[],[f648,f671])).
fof(f1178,plain,(
  ~p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18)))))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))))))),
  inference(resolution,[],[f648,f675])).
fof(f1111,plain,(
  ( ! [X8,X7] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK16,X7),X8),s(fun(sK16,bool),sK20),s(fun(X7,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f647,f678])).
fof(f1110,plain,(
  ( ! [X6,X5] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X5,sK16),X6),s(fun(X5,bool),'const.pred_set.EMPTY_0'),s(fun(sK16,bool),sK20))))) )),
  inference(superposition,[],[f647,f680])).
fof(f1109,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK16,X0),X1),s(fun(sK16,bool),sK20),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(superposition,[],[f647,f948])).
fof(f1092,plain,(
  ( ! [X251,X250] : (p(s(bool,X250)) | p(s(bool,X251)) | s(bool,X251) != s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) )),
  inference(resolution,[],[f647,f957])).
fof(f985,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X33) != s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20))) | s(bool,X34) != s(bool,X35)) )),
  inference(resolution,[],[f647,f706])).
fof(f1144,plain,(
  ( ! [X17] : (s(bool,X17) != s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) ) | $spl16),
  inference(cnf_transformation,[],[f1144_D])).
fof(f1144_D,plain,(
  ( ! [X17] : (s(bool,X17) != s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f984,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X30) != s(bool,X31) | s(bool,X32) != s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) )),
  inference(resolution,[],[f647,f706])).
fof(f982,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X26) != s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) )),
  inference(resolution,[],[f647,f705])).
fof(f981,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) )),
  inference(resolution,[],[f647,f704])).
fof(f980,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20))) | s(bool,X19) != s(bool,X20)) )),
  inference(resolution,[],[f647,f702])).
fof(f979,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,X16) | s(bool,X17) != s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) )),
  inference(resolution,[],[f647,f702])).
fof(f978,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.pred_set.INJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) )),
  inference(resolution,[],[f647,f700])).
fof(f1141,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK16,bool),sK20)))))) | $spl10),
  inference(cnf_transformation,[],[f1141_D])).
fof(f1141_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK16,bool),sK20)))))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f976,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK16,bool),sK20)))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK16,bool),sK20))))),
  inference(resolution,[],[f647,f946])).
fof(f967,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK17,bool),sK18))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK16,bool),'const.pred_set.IMAGE_2'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK17,bool),sK18))))),
  inference(resolution,[],[f647,f692])).
fof(f966,plain,(
  p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20)))) | ~p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK17,sK16),sK19),s(fun(sK17,bool),sK18),s(fun(sK16,bool),sK20))))),
  inference(resolution,[],[f647,f675])).
% SZS output end Proof for thm.util_prob
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 1036
% Passive clauses: 14670
% Generated clauses: 34188
% Final active clauses: 643
% Final passive clauses: 7220
% Input formulas: 261
% Initial clauses: 310
% 
% Selected by SInE selection: 78
% SInE iterations: 1
% 
% Duplicate literals: 1155
% Trivial inequalities: 16
% Fw subsumption resolutions: 1019
% Fw demodulations: 3329
% Bw demodulations: 85
% 
% Simple tautologies: 2635
% Equational tautologies: 77
% Forward subsumptions: 9815
% 
% Binary resolution: 17014
% Forward superposition: 8845
% Backward superposition: 1452
% Equality resolution: 16
% 
% Split clauses: 1304
% Split components: 2783
% SAT solver clauses: 1389
% SAT solver unit clauses: 19
% SAT solver binary clauses: 235
% 
% Sat splits: 1304
% Sat splitting refutations: 28
% 
% Memory used [KB]: 14967
% Time elapsed: 0.628 s
% ------------------------------
% ------------------------------
% Success in time 12.37 s
