Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.option
% (110244)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 873
% Passive clauses: 5921
% Generated clauses: 24186
% Final active clauses: 604
% Final passive clauses: 3391
% Input formulas: 48
% Initial clauses: 149
% 
% Duplicate literals: 435
% Trivial inequalities: 140
% Fw subsumption resolutions: 2573
% Fw demodulations: 1968
% Bw demodulations: 70
% 
% Simple tautologies: 215
% Equational tautologies: 66
% Forward subsumptions: 11697
% Fw demodulations to eq. taut.: 5
% 
% Binary resolution: 7962
% Factoring: 1464
% Forward superposition: 4599
% Backward superposition: 4301
% Self superposition: 39
% Equality factoring: 30
% Equality resolution: 28
% 
% Split clauses: 904
% Split components: 1819
% SAT solver clauses: 1514
% SAT solver unit clauses: 22
% SAT solver binary clauses: 242
% SAT solver learnt clauses: 2
% SAT solver learnt literals: 5
% 
% Sat splits: 904
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 1506
% TWLsolver calls for satisfiability: 521
% 
% Memory used [KB]: 10106
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.option
% (110262)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 566
% Passive clauses: 3044
% Generated clauses: 21375
% Final active clauses: 284
% Final passive clauses: 1528
% Input formulas: 48
% Initial clauses: 149
% 
% Duplicate literals: 264
% Trivial inequalities: 207
% Fw subsumption resolutions: 2029
% Fw demodulations: 1116
% Bw demodulations: 811
% Condensations: 607
% Global subsumptions: 18
% 
% Simple tautologies: 244
% Equational tautologies: 29
% Forward subsumptions: 9884
% Fw demodulations to eq. taut.: 5
% Bw demodulations to eq. taut.: 148
% 
% Binary resolution: 4153
% Factoring: 288
% Forward superposition: 3345
% Backward superposition: 2399
% Self superposition: 38
% Equality factoring: 28
% Equality resolution: 34
% 
% Split clauses: 379
% Split components: 759
% SAT solver clauses: 3473
% SAT solver unit clauses: 165
% SAT solver binary clauses: 91
% SAT solver learnt clauses: 40
% SAT solver learnt literals: 19
% 
% Sat splits: 379
% Sat splitting refutations: 74
% 
% TWLsolver clauses: 3397
% TWLsolver calls for satisfiability: 1158
% 
% Memory used [KB]: 13432
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.option
% (110281)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1644
% Passive clauses: 10855
% Generated clauses: 125448
% Final active clauses: 481
% Final passive clauses: 1453
% Input formulas: 48
% Initial clauses: 149
% 
% Duplicate literals: 3185
% Trivial inequalities: 381
% Fw subsumption resolutions: 3448
% Fw demodulations: 9367
% Bw demodulations: 8147
% Condensations: 7650
% Global subsumptions: 2735
% 
% Simple tautologies: 226
% Equational tautologies: 246
% Forward subsumptions: 80549
% Fw demodulations to eq. taut.: 566
% Bw demodulations to eq. taut.: 98
% 
% Unit resulting resolution: 10505
% Factoring: 1
% Forward superposition: 37533
% Backward superposition: 35304
% Self superposition: 130
% Equality factoring: 72
% Equality resolution: 47
% 
% Split clauses: 112
% Split components: 237
% SAT solver clauses: 17457
% SAT solver unit clauses: 1134
% SAT solver binary clauses: 3765
% 
% Sat splits: 112
% Sat splitting refutations: 3640
% 
% Memory used [KB]: 138164
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.option
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.option
% SZS output start Proof for thm.option
fof(f28568,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f215,f214,f369,f365,f213,f364,f370,f371,f372,f373,f374,f376,f377,f422,f468,f469,f470,f471,f472,f474,f475,f520,f616,f613,f658,f659,f660,f688,f689,f690,f691,f692,f694,f695,f737,f786,f787,f788,f789,f790,f895,f792,f793,f837,f838,f840,f841,f842,f843,f844,f845,f846,f855,f857,f912,f890,f889,f903,f915,f916,f917,f918,f919,f921,f922,f964,f998,f999,f1000,f1001,f1002,f1004,f1005,f1050,f1081,f1082,f1083,f1084,f1085,f1087,f1088,f1133,f1211,f1212,f1268,f1269,f1270,f1271,f1272,f1374,f1274,f1275,f1319,f1320,f1337,f1339,f1384,f1385,f1386,f1387,f1388,f1492,f1390,f1391,f1435,f1436,f1438,f1439,f1440,f1441,f1442,f1443,f1444,f1453,f1455,f1508,f1487,f1486,f1500,f1210,f1209,f1208,f1207,f1206,f1205,f1511,f1204,f1203,f1512,f1519,f1520,f1521,f1522,f1523,f1525,f1526,f1571,f1606,f1607,f1608,f1609,f1610,f1612,f454,f1613,f1655,f1716,f1717,f1252,f1816,f1827,f1830,f1839,f2046,f2182,f2183,f2184,f2185,f544,f2186,f2188,f2189,f2234,f543,f2321,f1248,f2406,f2407,f2408,f2409,f894,f2410,f2513,f2412,f2413,f2457,f2458,f893,f2475,f2477,f2523,f2524,f2525,f2526,f1491,f2527,f2630,f2529,f2530,f2574,f2575,f1490,f2592,f2594,f2870,f2872,f2888,f2889,f3107,f3119,f3149,f3158,f3188,f3197,f3242,f3252,f3265,f3277,f3304,f3314,f3319,f3329,f3411,f3412,f3413,f3459,f3415,f3416,f3417,f3465,f3467,f3453,f3452,f3471,f3472,f3473,f3474,f3580,f3475,f3581,f3477,f3478,f3522,f3582,f3523,f3579,f3525,f3526,f3527,f3528,f3529,f3530,f3531,f3540,f3542,f3591,f3595,f3576,f3575,f3588,f3607,f3618,f3689,f3688,f3692,f3693,f2871,f3774,f3767,f3766,f3765,f3830,f3831,f3832,f3833,f3834,f3835,f3891,f3892,f3893,f3954,f3955,f3956,f3965,f3966,f3967,f3968,f3969,f3971,f4067,f3972,f4014,f4021,f4022,f4023,f4072,f4025,f4026,f4027,f4078,f4063,f4062,f4071,f4153,f4198,f4199,f4200,f4311,f4330,f4331,f4332,f4333,f4441,f4334,f4442,f4336,f4337,f4381,f4443,f4382,f4440,f4385,f4386,f4387,f4388,f4389,f4390,f4391,f4400,f4402,f4452,f4456,f4437,f4436,f4449,f4461,f4467,f4484,f4497,f4498,f4499,f4500,f4501,f4502,f4504,f4601,f4505,f4547,f4555,f4556,f4557,f4606,f4559,f4560,f4561,f4612,f4597,f4596,f4605,f4631,f4643,f4685,f4696,f4716,f4727,f4777,f4829,f4870,f5036,f5037,f5039,f5040,f5131,f5133,f5180,f5182,f4798,f5261,f5263,f5313,f5315,f5357,f5358,f5359,f5360,f5469,f5361,f5470,f5363,f5364,f5408,f5471,f5409,f5468,f5428,f5430,f5476,f5483,f5484,f5485,f5486,f5594,f5487,f5595,f5489,f5490,f5534,f5596,f5535,f5593,f5554,f5556,f5601,f5867,f5868,f5872,f5883,f5884,f5888,f3466,f6117,f6119,f3617,f6155,f6157,f6159,f6163,f6166,f6169,f6172,f6175,f6250,f6339,f6353,f6369,f6486,f6574,f6587,f6603,f6717,f6801,f6814,f6830,f6944,f6251,f7028,f6340,f7053,f7069,f7425,f7426,f7427,f7428,f7429,f7431,f7545,f7432,f7474,f7544,f7484,f7485,f7486,f7549,f7488,f7489,f7490,f7555,f7528,f7527,f7548,f7565,f7586,f7585,f7602,f7603,f7666,f7667,f7668,f7669,f7778,f7670,f7779,f7672,f7673,f7717,f7718,f7738,f7740,f7833,f7847,f8096,f8097,f8098,f8099,f8100,f8102,f8223,f8103,f8145,f8222,f8156,f8157,f8158,f8227,f8160,f8161,f8162,f8233,f8200,f8199,f8226,f8285,f8291,f8352,f8358,f8359,f8374,f8464,f8465,f8466,f8467,f8577,f8468,f8578,f8470,f8471,f8515,f8579,f8516,f8576,f8584,f8586,f8894,f8896,f8909,f8910,f8945,f8947,f8959,f8960,f8994,f8996,f9008,f9009,f9044,f9046,f9057,f9058,f9078,f9080,f9154,f9156,f4884,f9174,f9175,f9176,f9177,f9287,f9178,f9288,f9180,f9181,f9225,f9289,f9226,f9286,f9232,f9233,f9234,f9235,f9236,f9237,f9238,f9247,f9249,f9298,f9302,f9283,f9282,f9295,f9621,f9622,f9623,f9624,f9625,f9627,f9753,f9628,f9670,f9752,f9797,f9798,f9799,f9800,f9801,f9803,f9924,f9804,f9846,f9923,f9859,f9860,f3458,f9861,f9927,f9863,f9864,f9865,f9933,f9903,f9902,f9926,f9974,f9975,f9976,f9977,f10089,f9978,f10090,f9980,f9981,f10025,f455,f10026,f10049,f10051,f10116,f10130,f10572,f10583,f10576,f10574,f10612,f10703,f10702,f10701,f10766,f10767,f10768,f10769,f10770,f10771,f11160,f11240,f11250,f11333,f11343,f11394,f11405,f11459,f11470,f11534,f11525,f11537,f11769,f11770,f11771,f11772,f11773,f11774,f11775,f11776,f11777,f11778,f11779,f11781,f11760,f11780,f11759,f11782,f11785,f11787,f11788,f11789,f11756,f11790,f11791,f11792,f11727,f11823,f11734,f11786,f11840,f11847,f11905,f11908,f11911,f11913,f11916,f12037,f12040,f12044,f12339,f12340,f12341,f12342,f12343,f12344,f12345,f12346,f12347,f12348,f12349,f12351,f12330,f12350,f12329,f12352,f12355,f12357,f12358,f12359,f12326,f12360,f12361,f12362,f12303,f12356,f12795,f12796,f12797,f12798,f12799,f12800,f12801,f12802,f12803,f12804,f12805,f12807,f12786,f12806,f12785,f12808,f12811,f12813,f12814,f12815,f12782,f12816,f12817,f12818,f12759,f12812,f13314,f13315,f13316,f13317,f13318,f13319,f13320,f13321,f13322,f13323,f13324,f13326,f13305,f13325,f13304,f13327,f13330,f13332,f13333,f13334,f8292,f13301,f13335,f13336,f13337,f13278,f11824,f13331,f14030,f14031,f14032,f14033,f14034,f14035,f14036,f14037,f14038,f14039,f14040,f14042,f14017,f14041,f14016,f14043,f14364,f14365,f14366,f14367,f14368,f14373,f14374,f14375,f14376,f14377,f14490,f14569,f14568,f14567,f14632,f14633,f14634,f14635,f14636,f14637,f14755,f14756,f14757,f14758,f14759,f14764,f14765,f14766,f14767,f14768,f14979,f14980,f14981,f14982,f14983,f14988,f14989,f14990,f14991,f14992,f15202,f14463,f15203,f14464,f15204,f14465,f15205,f14466,f15206,f14467,f15211,f14468,f15212,f14469,f15213,f14470,f15214,f14471,f15215,f14472,f16070,f16580,f8585,f16714,f16713,f16721,f16732,f16734,f16736,f16738,f16739,f16784,f16807,f16809,f16910,f16912,f16914,f16916,f16917,f16962,f16985,f16987,f17088,f17090,f17092,f17094,f17095,f17140,f17202,f17163,f17165,f17208,f17362,f17364,f17366,f17368,f17369,f17414,f17476,f17437,f17439,f17482,f17583,f17585,f17592,f17593,f17676,f17675,f17681,f17731,f17733,f17739,f17740,f17781,f17783,f17791,f17792,f17834,f17836,f17842,f17843,f17884,f17886,f17896,f17897,f17939,f17941,f17949,f17950,f17992,f17994,f18004,f18005,f18047,f18049,f18057,f18058,f18099,f18101,f18110,f18111,f18153,f18155,f18162,f18163,f18204,f18206,f18215,f18216,f18259,f18261,f18268,f18269,f18310,f18312,f18322,f18323,f18365,f18367,f18375,f18376,f18417,f18419,f18429,f18430,f18472,f18474,f18482,f18483,f18511,f18513,f18535,f18536,f18726,f18728,f18748,f18749,f18777,f18779,f18801,f18802,f18831,f18833,f18853,f18854,f18958,f18467,f18413,f18254,f18200,f18042,f17988,f17829,f18959,f17777,f18960,f11970,f11969,f11968,f11967,f18981,f11953,f15300,f11952,f15077,f11951,f14853,f11950,f18987,f14462,f19062,f19064,f19066,f19068,f19069,f19114,f452,f19408,f19464,f19606,f19663,f19698,f19710,f19890,f19902,f20032,f20033,f20034,f20035,f20038,f20039,f20040,f20041,f20375,f20376,f20377,f20378,f20379,f20380,f20381,f20382,f1504,f20437,f20439,f20441,f20443,f20444,f20489,f20550,f20512,f20514,f20563,f20565,f20567,f20569,f20570,f20615,f20685,f20681,f20679,f21155,f21156,f21157,f21158,f21161,f21162,f21163,f21164,f21167,f21168,f21169,f21170,f21173,f21174,f21175,f21176,f21177,f21178,f21179,f21180,f21218,f21219,f21187,f21220,f21221,f21188,f21224,f21225,f21189,f21226,f21227,f21190,f21230,f21231,f21191,f21232,f21233,f21192,f21236,f21237,f21193,f21238,f21239,f21194,f21240,f21241,f21195,f21242,f7780,f21243,f21196,f21401,f21403,f21407,f21409,f21413,f21415,f21419,f21421,f21423,f21425,f21429,f21454,f21586,f21587,f21588,f21589,f21592,f21593,f21594,f21595,f21746,f21443,f21748,f21444,f21752,f21445,f21754,f21446,f21758,f21447,f21760,f21448,f21764,f21449,f21766,f21450,f21768,f21451,f21770,f21452,f21773,f21787,f1826,f21996,f21997,f21998,f21999,f22000,f22001,f22002,f22003,f908,f22062,f22064,f22066,f22068,f22069,f22114,f22185,f22181,f22179,f22197,f22199,f22201,f22203,f22204,f22249,f22311,f22272,f22274,f20556,f22404,f22405,f22406,f22407,f22410,f22411,f22412,f22413,f22416,f22417,f22418,f22419,f22422,f22423,f22424,f22425,f22426,f22427,f22428,f22429,f22467,f6135,f22468,f22436,f22469,f17790,f22470,f22437,f22473,f17895,f22474,f22438,f22475,f18003,f22476,f22439,f22479,f18109,f22480,f22440,f22481,f18214,f22482,f22441,f22485,f18321,f22486,f22442,f22487,f18428,f22488,f22443,f22489,f18534,f22490,f22444,f22491,f18800,f10091,f22492,f22445,f22676,f2276,f2275,f22703,f366,f2291,f2290,f367,f22718,f22720,f22721,f22722,f22725,f22767,f987,f22798,f22800,f22817,f1249,f1247,f22818,f1253,f1251,f22908,f22963,f22964,f22969,f22970,f22979,f22980,f23465,f23467,f23468,f23469,f23472,f23514,f23563,f23570,f23572,f23573,f23574,f23577,f23619,f23668,f23746,f23756,f368,f23844,f23843,f23842,f23906,f23907,f23908,f23909,f23910,f23911,f24011,f24013,f24014,f24015,f24018,f24060,f4066,f24137,f24139,f24140,f24141,f24144,f24186,f4600,f24353,f24355,f24357,f24359,f24360,f24405,f7777,f24426,f24428,f24560,f24562,f24564,f24566,f24567,f24612,f10088,f24633,f24635,f7786,f24972,f24974,f24975,f24976,f24979,f25021,f767,f25177,f25179,f25180,f25181,f25184,f25226,f25287,f25487,f25489,f25490,f25491,f25494,f25536,f25598,f25730,f25729,f25735,f25734,f25744,f25746,f26094,f26096,f26140,f26141,f26095,f26184,f26186,f26229,f26230,f26185,f26476,f26478,f26479,f26480,f26483,f26525,f26588,f26625,f26627,f26628,f26629,f26632,f26674,f26736,f26824,f26823,f26829,f16720,f27102,f27412,f1493,f27668,f27926,f896,f11161,f27958,f27960,f27962,f27964,f27965,f28010,f16846,f28035,f28037,f16852,f28116,f28117,f28150,f28151,f4842,f28250,f28252,f28254,f453,f28256,f458,f28257,f456,f28302,f17024,f28327,f28329,f17030,f906,f28408,f28409,f4323,f459,f28545])).
fof(f28545,plain,(
  $false | ($spl1 | $spl2 | $spl44)),
  inference(subsumption_resolution,[],[f28488,f22995])).
fof(f22995,plain,(
  s('type.option.option'(sK6),'const.option.NONE_0') != s('type.option.option'(sK6),'const.bool.COND_3'(s(bool,'T_0'),s('type.option.option'(sK6),'const.option.NONE_0'),s('type.option.option'(sK6),sK7))) | ($spl1 | $spl44)),
  inference(backward_demodulation,[],[f1251,f365])).
fof(f28488,plain,(
  s('type.option.option'(sK6),'const.option.NONE_0') = s('type.option.option'(sK6),'const.bool.COND_3'(s(bool,'T_0'),s('type.option.option'(sK6),'const.option.NONE_0'),s('type.option.option'(sK6),sK7))) | ($spl1 | $spl2 | $spl44)),
  inference(superposition,[],[f27929,f1254])).
fof(f1254,plain,(
  s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6)))) | $spl1),
  inference(subsumption_resolution,[],[f1227,f355])).
fof(f355,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f110])).
fof(f110,plain,(
  ! [X0,X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X2)))),
  inference(flattening,[],[f109])).
fof(f109,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X2)))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X6] : ! [X14,X11] : s(X6,X11) = s(X6,'const.bool.COND_3'(s(bool,X14),s(X6,X11),s(X6,X11)))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|1','thm.bool.COND_ID')).
fof(f1227,plain,(
  s('type.option.option'(sK6),'const.option.NONE_0') != s('type.option.option'(sK6),'const.bool.COND_3'(s(bool,sK8),s('type.option.option'(sK6),'const.option.NONE_0'),s('type.option.option'(sK6),'const.option.NONE_0'))) | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6)))) | $spl1),
  inference(superposition,[],[f365,f357])).
fof(f357,plain,(
  ( ! [X0,X1] : (s('type.option.option'(X0),X1) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,sK9(X1,X0)))) | s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X1)) )),
  inference(cnf_transformation,[],[f212])).
fof(f212,plain,(
  ! [X0,X1] : (s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X1) | s('type.option.option'(X0),X1) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,sK9(X1,X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK9])],[f114])).
fof(f114,plain,(
  ! [X0,X1] : (s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X1) | ? [X2] : s('type.option.option'(X0),X1) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2))))),
  inference(flattening,[],[f113])).
fof(f113,plain,(
  ! [X0] : ! [X1] : (s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X1) | ? [X2] : s('type.option.option'(X0),X1) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2))))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X6] : ! [X18] : (s('type.option.option'(X6),'const.option.NONE_0') = s('type.option.option'(X6),X18) | ? [X12] : s('type.option.option'(X6),X18) = s('type.option.option'(X6),'const.option.SOME_1'(s(X6,X12))))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|1','thm.option.option_nchotomy')).
fof(f27929,plain,(
  ( ! [X74,X75] : (s('type.option.option'(X74),'const.option.NONE_0') = s('type.option.option'(X74),'const.bool.COND_3'(s(bool,'T_0'),s('type.option.option'(X74),'const.option.NONE_0'),s('type.option.option'(X74),'const.option.SOME_1'(s(X74,X75)))))) ) | ($spl2 | $spl44)),
  inference(forward_demodulation,[],[f22766,f1251])).
fof(f22766,plain,(
  ( ! [X74,X75] : (s('type.option.option'(X74),'const.option.NONE_0') = s('type.option.option'(X74),'const.bool.COND_3'(s(bool,sK8),s('type.option.option'(X74),'const.option.NONE_0'),s('type.option.option'(X74),'const.option.SOME_1'(s(X74,X75)))))) ) | $spl2),
  inference(resolution,[],[f366,f350])).
fof(f350,plain,(
  ( ! [X2,X0,X1] : (s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),'const.bool.COND_3'(s(bool,X2),s('type.option.option'(X0),'const.option.NONE_0'),s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1))))) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f209])).
fof(f209,plain,(
  ! [X0,X1,X2] : ((s('type.option.option'(X0),'const.option.NONE_0') != s('type.option.option'(X0),'const.bool.COND_3'(s(bool,X2),s('type.option.option'(X0),'const.option.NONE_0'),s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1))))) | p(s(bool,X2))) & (~p(s(bool,X2)) | s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),'const.bool.COND_3'(s(bool,X2),s('type.option.option'(X0),'const.option.NONE_0'),s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1)))))))),
  inference(nnf_transformation,[],[f104])).
fof(f104,plain,(
  ! [X0,X1,X2] : (s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),'const.bool.COND_3'(s(bool,X2),s('type.option.option'(X0),'const.option.NONE_0'),s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1))))) <=> p(s(bool,X2)))),
  inference(flattening,[],[f103])).
fof(f103,plain,(
  ! [X0] : ! [X1,X2] : (s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),'const.bool.COND_3'(s(bool,X2),s('type.option.option'(X0),'const.option.NONE_0'),s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1))))) <=> p(s(bool,X2)))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X6] : ! [X12,X8] : (s('type.option.option'(X6),'const.option.NONE_0') = s('type.option.option'(X6),'const.bool.COND_3'(s(bool,X8),s('type.option.option'(X6),'const.option.NONE_0'),s('type.option.option'(X6),'const.option.SOME_1'(s(X6,X12))))) <=> p(s(bool,X8)))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|1','thm.option.IF_EQUALS_OPTION|split|1')).
fof(f459,plain,(
  ( ! [X18] : (p(s(bool,X18))) ) | $spl18),
  inference(cnf_transformation,[],[f459_D])).
fof(f459_D,plain,(
  ( ! [X18] : (p(s(bool,X18))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f4323,plain,(
  ( ! [X4,X5] : (sP4('T_0',X4) | p(s(bool,X5)) | sP4(X4,X5)) ) | $spl108),
  inference(cnf_transformation,[],[f4323_D])).
fof(f4323_D,plain,(
  ( ! [X4,X5] : (sP4('T_0',X4) | p(s(bool,X5)) | sP4(X4,X5)) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f28409,plain,(
  ( ! [X4,X5,X3] : (sP4(X3,X4) | sP4('T_0',X3) | p(s(bool,X4)) | p(s(bool,X5))) ) | ($spl2 | $spl100)),
  inference(subsumption_resolution,[],[f28378,f307])).
fof(f307,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f191])).
fof(f191,plain,(
  ! [X0] : ((~p(s(bool,X0)) | ~p(s(bool,'F_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,X0)) & p(s(bool,'F_0')))))),
  inference(flattening,[],[f190])).
fof(f190,plain,(
  ! [X0] : (((~p(s(bool,X0)) | ~p(s(bool,'F_0'))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,X0)) & p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0] : ((p(s(bool,X0)) & p(s(bool,'F_0'))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X11] : ((p(s(bool,X11)) & p(s(bool,'F_0'))) <=> p(s(bool,'F_0')))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|1','thm.bool.AND_CLAUSES|split|3')).
fof(f28378,plain,(
  ( ! [X4,X5,X3] : (sP4(X3,X4) | sP4('T_0',X3) | p(s(bool,'F_0')) | p(s(bool,X4)) | p(s(bool,X5))) ) | ($spl2 | $spl100)),
  inference(resolution,[],[f28303,f273])).
fof(f273,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP5(X0,X1)) )),
  inference(cnf_transformation,[],[f174])).
fof(f174,plain,(
  ! [X0,X1] : ((~sP5(X0,X1) | (p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP5(X0,X1)))),
  inference(flattening,[],[f173])).
fof(f173,plain,(
  ! [X0,X1] : ((~sP5(X0,X1) | ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0')))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP5(X0,X1)))),
  inference(nnf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0,X1] : (sP5(X0,X1) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(definition_folding,[],[f124,f144])).
fof(f144,plain,(
  ! [X0,X1] : (sP5(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f124,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f123])).
fof(f123,plain,(
  ! [X0,X1] : (((p(s(bool,X1)) | p(s(bool,X0))) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | (p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(ennf_transformation,[],[f61])).
fof(f61,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(flattening,[],[f60])).
fof(f60,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X19,X6] : ((~(p(s(bool,X6)) | p(s(bool,X19))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X6)) => p(s(bool,'F_0'))) => (~p(s(bool,X19)) => p(s(bool,'F_0')))))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|1','thm.sat.sat_9_prove_ltb')).
fof(f28303,plain,(
  ( ! [X241,X239,X240] : (sP5(X239,X240) | sP4(X241,X239) | sP4('T_0',X241)) ) | ($spl2 | $spl100)),
  inference(resolution,[],[f25346,f3618])).
fof(f25346,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,X4)) | sP5(X5,X6) | sP4(X4,X5)) ) | $spl2),
  inference(resolution,[],[f25208,f22740])).
fof(f22740,plain,(
  ( ! [X52] : (~sP4(sK8,X52) | ~p(s(bool,X52))) ) | $spl2),
  inference(resolution,[],[f366,f254])).
fof(f254,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | ~sP4(X0,X1)) )),
  inference(cnf_transformation,[],[f167])).
fof(f167,plain,(
  ! [X0,X1] : ((~sP4(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1)))),
  inference(rectify,[],[f166])).
fof(f166,plain,(
  ! [X1,X0] : ((~sP4(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP4(X1,X0)))),
  inference(flattening,[],[f165])).
fof(f165,plain,(
  ! [X1,X0] : ((~sP4(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP4(X1,X0)))),
  inference(nnf_transformation,[],[f142])).
fof(f142,plain,(
  ! [X1,X0] : (sP4(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f25208,plain,(
  ( ! [X125,X123,X124] : (sP4(sK8,X123) | sP4(X123,X124) | sP5(X124,X125)) ) | $spl2),
  inference(resolution,[],[f23489,f271])).
fof(f271,plain,(
  ( ! [X0,X1] : (sP5(X0,X1) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f172])).
fof(f172,plain,(
  ! [X0,X1] : ((~sP5(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP5(X0,X1)))),
  inference(flattening,[],[f171])).
fof(f171,plain,(
  ! [X0,X1] : ((~sP5(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP5(X0,X1)))),
  inference(nnf_transformation,[],[f144])).
fof(f23489,plain,(
  ( ! [X78,X79] : (sP4(sK8,X78) | sP4(X78,X79) | p(s(bool,X79))) ) | $spl2),
  inference(resolution,[],[f22742,f255])).
fof(f255,plain,(
  ( ! [X0,X1] : (sP4(X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f167])).
fof(f22742,plain,(
  ( ! [X54] : (p(s(bool,X54)) | sP4(sK8,X54)) ) | $spl2),
  inference(resolution,[],[f366,f255])).
fof(f28408,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1) | sP4('T_0',X0) | p(s(bool,X1)) | p(s(bool,X2))) ) | ($spl2 | $spl100)),
  inference(subsumption_resolution,[],[f28377,f307])).
fof(f28377,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1) | sP4('T_0',X0) | p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,X2))) ) | ($spl2 | $spl100)),
  inference(resolution,[],[f28303,f269])).
fof(f269,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP5(X0,X1)) )),
  inference(cnf_transformation,[],[f172])).
fof(f906,plain,(
  ( ! [X140,X141] : (p(s(bool,X140)) | ~sP5(X140,X141) | p(s(bool,X141))) ) | $spl34),
  inference(cnf_transformation,[],[f906_D])).
fof(f906_D,plain,(
  ( ! [X140,X141] : (p(s(bool,X140)) | ~sP5(X140,X141) | p(s(bool,X141))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f17030,plain,(
  ( ! [X300,X301] : (sP5(X300,X301) | sP4('F_0',X300)) ) | $spl382),
  inference(cnf_transformation,[],[f17030_D])).
fof(f17030_D,plain,(
  ( ! [X300,X301] : (sP5(X300,X301) | sP4('F_0',X300)) ) <=> ~$spl382),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl382])])).
fof(f28329,plain,(
  ( ! [X315,X316,X314,X317] : (sP5(X314,X315) | sP4('F_0',X314) | p(s(bool,X316)) | p(s(bool,X317)) | ~sP5(X316,X317)) ) | $spl2),
  inference(resolution,[],[f25346,f269])).
fof(f28327,plain,(
  ( ! [X308,X306,X309,X307] : (sP5(X306,X307) | sP4('F_0',X306) | p(s(bool,X308)) | p(s(bool,X309)) | ~sP5(X308,X309)) ) | $spl2),
  inference(resolution,[],[f25346,f273])).
fof(f17024,plain,(
  ( ! [X4,X2,X0,X1] : (sP5(X0,X1) | s(bool,X2) != s(bool,X4) | ~p(s(bool,X4)) | sP4(X2,X0)) ) | $spl380),
  inference(cnf_transformation,[],[f17024_D])).
fof(f17024_D,plain,(
  ( ! [X4,X2,X0,X1] : (sP5(X0,X1) | s(bool,X2) != s(bool,X4) | ~p(s(bool,X4)) | sP4(X2,X0)) ) <=> ~$spl380),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl380])])).
fof(f28302,plain,(
  ( ! [X237,X235,X238,X236,X234] : (sP5(X234,X235) | sP4(X236,X234) | p(s(bool,X237)) | ~p(s(bool,X238)) | s(bool,X236) != s(bool,X238)) ) | $spl2),
  inference(resolution,[],[f25346,f360])).
fof(f360,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f220])).
fof(f220,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f150])).
fof(f150,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f149])).
fof(f149,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f117])).
fof(f117,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X12,X16,X13,X17] : ((s(bool,X12) = s(bool,X16) & (p(s(bool,X16)) => s(bool,X13) = s(bool,X17))) => ((p(s(bool,X12)) => p(s(bool,X13))) <=> (p(s(bool,X16)) => p(s(bool,X17)))))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|1','thm.bool.IMP_CONG')).
fof(f456,plain,(
  ( ! [X6,X8] : (s(bool,X6) != s(bool,X8) | ~p(s(bool,X6))) ) | $spl14),
  inference(cnf_transformation,[],[f456_D])).
fof(f456_D,plain,(
  ( ! [X6,X8] : (s(bool,X6) != s(bool,X8) | ~p(s(bool,X6))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f28257,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (sP5(X42,X43) | sP4(X44,X42) | ~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X44) != s(bool,X46) | s(bool,X45) != s(bool,X47)) ) | $spl2),
  inference(resolution,[],[f25346,f223])).
fof(f223,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f150])).
fof(f458,plain,(
  ( ! [X20] : (~p(s(bool,X20))) ) | $spl16),
  inference(cnf_transformation,[],[f458_D])).
fof(f458_D,plain,(
  ( ! [X20] : (~p(s(bool,X20))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f28256,plain,(
  ( ! [X39,X37,X41,X38,X36,X40] : (sP5(X36,X37) | sP4(X38,X36) | p(s(bool,X39)) | ~p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X38) != s(bool,X40)) ) | $spl2),
  inference(resolution,[],[f25346,f222])).
fof(f222,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f150])).
fof(f453,plain,(
  ( ! [X2,X0] : (p(s(bool,X0)) | s(bool,X0) != s(bool,X2)) ) | $spl8),
  inference(cnf_transformation,[],[f453_D])).
fof(f453_D,plain,(
  ( ! [X2,X0] : (p(s(bool,X0)) | s(bool,X0) != s(bool,X2)) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f28254,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (sP5(X24,X25) | sP4(X26,X24) | p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,X29) | s(bool,X26) != s(bool,X28)) ) | $spl2),
  inference(resolution,[],[f25346,f221])).
fof(f221,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f150])).
fof(f28252,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (sP5(X12,X13) | sP4(X14,X12) | ~p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,X17)) ) | $spl2),
  inference(resolution,[],[f25346,f219])).
fof(f219,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f150])).
fof(f28250,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP5(X0,X1) | sP4(X2,X0) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X2) != s(bool,X4)) ) | $spl2),
  inference(resolution,[],[f25346,f217])).
fof(f217,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f150])).
fof(f4842,plain,(
  ( ! [X4,X5] : (sP4(X4,'T_0') | p(s(bool,X5)) | sP4(X4,X5)) ) | $spl138),
  inference(cnf_transformation,[],[f4842_D])).
fof(f4842_D,plain,(
  ( ! [X4,X5] : (sP4(X4,'T_0') | p(s(bool,X5)) | sP4(X4,X5)) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f28151,plain,(
  ( ! [X4,X5,X3] : (sP4(X3,X4) | sP4(X3,'T_0') | p(s(bool,X5)) | p(s(bool,X4))) ) | ($spl2 | $spl120)),
  inference(subsumption_resolution,[],[f28119,f307])).
fof(f28119,plain,(
  ( ! [X4,X5,X3] : (sP4(X3,X4) | sP4(X3,'T_0') | p(s(bool,'F_0')) | p(s(bool,X5)) | p(s(bool,X4))) ) | ($spl2 | $spl120)),
  inference(resolution,[],[f28012,f273])).
fof(f28012,plain,(
  ( ! [X243,X244,X242] : (sP5(X242,X243) | sP4(X244,X243) | sP4(X244,'T_0')) ) | ($spl2 | $spl120)),
  inference(resolution,[],[f25311,f4467])).
fof(f25311,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,X4)) | sP5(X6,X5) | sP4(X4,X5)) ) | $spl2),
  inference(resolution,[],[f25207,f22740])).
fof(f25207,plain,(
  ( ! [X121,X122,X120] : (sP4(sK8,X120) | sP4(X120,X121) | sP5(X122,X121)) ) | $spl2),
  inference(resolution,[],[f23489,f270])).
fof(f270,plain,(
  ( ! [X0,X1] : (sP5(X0,X1) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f172])).
fof(f28150,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1) | sP4(X0,'T_0') | p(s(bool,X2)) | p(s(bool,X1))) ) | ($spl2 | $spl120)),
  inference(subsumption_resolution,[],[f28118,f307])).
fof(f28118,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1) | sP4(X0,'T_0') | p(s(bool,'F_0')) | p(s(bool,X2)) | p(s(bool,X1))) ) | ($spl2 | $spl120)),
  inference(resolution,[],[f28012,f269])).
fof(f28117,plain,(
  ( ! [X4,X5,X3] : (sP4(X3,X4) | sP4('T_0',X3) | p(s(bool,X5)) | p(s(bool,X4))) ) | ($spl2 | $spl100)),
  inference(subsumption_resolution,[],[f28086,f307])).
fof(f28086,plain,(
  ( ! [X4,X5,X3] : (sP4(X3,X4) | sP4('T_0',X3) | p(s(bool,'F_0')) | p(s(bool,X5)) | p(s(bool,X4))) ) | ($spl2 | $spl100)),
  inference(resolution,[],[f28011,f273])).
fof(f28011,plain,(
  ( ! [X241,X239,X240] : (sP5(X239,X240) | sP4(X241,X240) | sP4('T_0',X241)) ) | ($spl2 | $spl100)),
  inference(resolution,[],[f25311,f3618])).
fof(f28116,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1) | sP4('T_0',X0) | p(s(bool,X2)) | p(s(bool,X1))) ) | ($spl2 | $spl100)),
  inference(subsumption_resolution,[],[f28085,f307])).
fof(f28085,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1) | sP4('T_0',X0) | p(s(bool,'F_0')) | p(s(bool,X2)) | p(s(bool,X1))) ) | ($spl2 | $spl100)),
  inference(resolution,[],[f28011,f269])).
fof(f16852,plain,(
  ( ! [X300,X301] : (sP5(X300,X301) | sP4('F_0',X301)) ) | $spl378),
  inference(cnf_transformation,[],[f16852_D])).
fof(f16852_D,plain,(
  ( ! [X300,X301] : (sP5(X300,X301) | sP4('F_0',X301)) ) <=> ~$spl378),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl378])])).
fof(f28037,plain,(
  ( ! [X315,X316,X314,X317] : (sP5(X314,X315) | sP4('F_0',X315) | p(s(bool,X316)) | p(s(bool,X317)) | ~sP5(X316,X317)) ) | $spl2),
  inference(resolution,[],[f25311,f269])).
fof(f28035,plain,(
  ( ! [X308,X306,X309,X307] : (sP5(X306,X307) | sP4('F_0',X307) | p(s(bool,X308)) | p(s(bool,X309)) | ~sP5(X308,X309)) ) | $spl2),
  inference(resolution,[],[f25311,f273])).
fof(f16846,plain,(
  ( ! [X4,X2,X0,X1] : (sP5(X0,X1) | s(bool,X2) != s(bool,X4) | ~p(s(bool,X4)) | sP4(X2,X1)) ) | $spl376),
  inference(cnf_transformation,[],[f16846_D])).
fof(f16846_D,plain,(
  ( ! [X4,X2,X0,X1] : (sP5(X0,X1) | s(bool,X2) != s(bool,X4) | ~p(s(bool,X4)) | sP4(X2,X1)) ) <=> ~$spl376),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl376])])).
fof(f28010,plain,(
  ( ! [X237,X235,X238,X236,X234] : (sP5(X234,X235) | sP4(X236,X235) | p(s(bool,X237)) | ~p(s(bool,X238)) | s(bool,X236) != s(bool,X238)) ) | $spl2),
  inference(resolution,[],[f25311,f360])).
fof(f27965,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (sP5(X42,X43) | sP4(X44,X43) | ~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X44) != s(bool,X46) | s(bool,X45) != s(bool,X47)) ) | $spl2),
  inference(resolution,[],[f25311,f223])).
fof(f27964,plain,(
  ( ! [X39,X37,X41,X38,X36,X40] : (sP5(X36,X37) | sP4(X38,X37) | p(s(bool,X39)) | ~p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X38) != s(bool,X40)) ) | $spl2),
  inference(resolution,[],[f25311,f222])).
fof(f27962,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (sP5(X24,X25) | sP4(X26,X25) | p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,X29) | s(bool,X26) != s(bool,X28)) ) | $spl2),
  inference(resolution,[],[f25311,f221])).
fof(f27960,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (sP5(X12,X13) | sP4(X14,X13) | ~p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,X17)) ) | $spl2),
  inference(resolution,[],[f25311,f219])).
fof(f27958,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP5(X0,X1) | sP4(X2,X1) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X2) != s(bool,X4)) ) | $spl2),
  inference(resolution,[],[f25311,f217])).
fof(f11161,plain,(
  ( ! [X6,X7,X5] : (sP3(X6,X7,X5) | ~sP1(X7,X5,X6) | ~sP2(X6,X5,X7) | ~p(s(bool,X5))) ) | $spl244),
  inference(cnf_transformation,[],[f11161_D])).
fof(f11161_D,plain,(
  ( ! [X6,X7,X5] : (sP3(X6,X7,X5) | ~sP1(X7,X5,X6) | ~sP2(X6,X5,X7) | ~p(s(bool,X5))) ) <=> ~$spl244),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl244])])).
fof(f896,plain,(
  ( ! [X20] : (sP4(X20,sK8)) ) | $spl30),
  inference(cnf_transformation,[],[f896_D])).
fof(f896_D,plain,(
  ( ! [X20] : (sP4(X20,sK8)) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f27926,plain,(
  ( ! [X6,X4,X7,X5] : (sP4(X4,sK8) | sP3(X6,X7,X5) | ~p(s(bool,X5)) | ~sP2(X6,X5,X7) | ~sP1(X7,X5,X6)) ) | $spl2),
  inference(subsumption_resolution,[],[f27913,f23593])).
fof(f23593,plain,(
  ( ! [X76,X77] : (~sP4(X77,X76) | ~p(s(bool,X77)) | sP4(X76,sK8)) ) | $spl2),
  inference(resolution,[],[f22743,f254])).
fof(f22743,plain,(
  ( ! [X55] : (p(s(bool,X55)) | sP4(X55,sK8)) ) | $spl2),
  inference(resolution,[],[f366,f258])).
fof(f258,plain,(
  ( ! [X0,X1] : (sP4(X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f167])).
fof(f27913,plain,(
  ( ! [X6,X4,X7,X5] : (sP4(X4,sK8) | sP4(X5,X4) | sP3(X6,X7,X5) | ~p(s(bool,X5)) | ~sP2(X6,X5,X7) | ~sP1(X7,X5,X6)) ) | $spl2),
  inference(resolution,[],[f26644,f228])).
fof(f228,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X2)) | ~sP2(X0,X2,X1) | ~sP1(X1,X2,X0) | ~sP0(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f153])).
fof(f153,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | (sP0(X2,X1,X0) & sP1(X1,X2,X0) & sP2(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP0(X2,X1,X0) | ~sP1(X1,X2,X0) | ~sP2(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f152])).
fof(f152,plain,(
  ! [X2,X1,X0] : ((~sP3(X2,X1,X0) | (sP0(X0,X1,X2) & sP1(X1,X0,X2) & sP2(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP0(X0,X1,X2) | ~sP1(X1,X0,X2) | ~sP2(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP3(X2,X1,X0)))),
  inference(flattening,[],[f151])).
fof(f151,plain,(
  ! [X2,X1,X0] : ((~sP3(X2,X1,X0) | (sP0(X0,X1,X2) & sP1(X1,X0,X2) & sP2(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP0(X0,X1,X2) | ~sP1(X1,X0,X2) | ~sP2(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP3(X2,X1,X0)))),
  inference(nnf_transformation,[],[f140])).
fof(f140,plain,(
  ! [X2,X1,X0] : (sP3(X2,X1,X0) <=> (sP0(X0,X1,X2) & sP1(X1,X0,X2) & sP2(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f26644,plain,(
  ( ! [X87,X85,X86,X84] : (sP4(X84,sK8) | sP4(X85,X84) | sP0(X85,X86,X87)) ) | $spl2),
  inference(resolution,[],[f23595,f242])).
fof(f242,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f161])).
fof(f161,plain,(
  ! [X0,X1,X2] : ((~sP0(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP0(X0,X1,X2)))),
  inference(flattening,[],[f160])).
fof(f160,plain,(
  ! [X0,X1,X2] : ((~sP0(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP0(X0,X1,X2)))),
  inference(nnf_transformation,[],[f137])).
fof(f137,plain,(
  ! [X0,X1,X2] : (sP0(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f23595,plain,(
  ( ! [X80,X81] : (sP4(X80,sK8) | sP4(X81,X80) | p(s(bool,X81))) ) | $spl2),
  inference(resolution,[],[f22743,f258])).
fof(f27668,plain,(
  ( ! [X6,X4,X7,X5] : (sP4(X4,sK8) | sP3(X6,X7,X5) | ~p(s(bool,X5)) | ~sP2(X6,X5,X7) | ~sP1(X7,X5,X6)) ) | $spl2),
  inference(subsumption_resolution,[],[f27655,f23592])).
fof(f23592,plain,(
  ( ! [X74,X75] : (~sP4(X74,X75) | ~p(s(bool,X75)) | sP4(X74,sK8)) ) | $spl2),
  inference(resolution,[],[f22743,f254])).
fof(f27655,plain,(
  ( ! [X6,X4,X7,X5] : (sP4(X4,sK8) | sP4(X4,X5) | sP3(X6,X7,X5) | ~p(s(bool,X5)) | ~sP2(X6,X5,X7) | ~sP1(X7,X5,X6)) ) | $spl2),
  inference(resolution,[],[f26495,f228])).
fof(f26495,plain,(
  ( ! [X87,X85,X86,X84] : (sP4(X84,sK8) | sP4(X84,X85) | sP0(X85,X86,X87)) ) | $spl2),
  inference(resolution,[],[f23594,f242])).
fof(f23594,plain,(
  ( ! [X78,X79] : (sP4(X78,sK8) | sP4(X78,X79) | p(s(bool,X79))) ) | $spl2),
  inference(resolution,[],[f22743,f255])).
fof(f1493,plain,(
  ( ! [X20] : (sP4(sK8,X20)) ) | $spl52),
  inference(cnf_transformation,[],[f1493_D])).
fof(f1493_D,plain,(
  ( ! [X20] : (sP4(sK8,X20)) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f27412,plain,(
  ( ! [X6,X4,X7,X5] : (sP4(sK8,X4) | sP3(X6,X7,X5) | ~p(s(bool,X5)) | ~sP2(X6,X5,X7) | ~sP1(X7,X5,X6)) ) | $spl2),
  inference(subsumption_resolution,[],[f27397,f23488])).
fof(f23488,plain,(
  ( ! [X76,X77] : (~sP4(X77,X76) | ~p(s(bool,X77)) | sP4(sK8,X76)) ) | $spl2),
  inference(resolution,[],[f22742,f254])).
fof(f27397,plain,(
  ( ! [X6,X4,X7,X5] : (sP4(sK8,X4) | sP4(X5,X4) | sP3(X6,X7,X5) | ~p(s(bool,X5)) | ~sP2(X6,X5,X7) | ~sP1(X7,X5,X6)) ) | $spl2),
  inference(resolution,[],[f25506,f228])).
fof(f25506,plain,(
  ( ! [X87,X85,X86,X84] : (sP4(sK8,X84) | sP4(X85,X84) | sP0(X85,X86,X87)) ) | $spl2),
  inference(resolution,[],[f23490,f242])).
fof(f23490,plain,(
  ( ! [X80,X81] : (sP4(sK8,X80) | sP4(X81,X80) | p(s(bool,X81))) ) | $spl2),
  inference(resolution,[],[f22742,f258])).
fof(f27102,plain,(
  ( ! [X6,X4,X7,X5] : (sP4(sK8,X4) | sP3(X6,X7,X5) | ~p(s(bool,X5)) | ~sP2(X6,X5,X7) | ~sP1(X7,X5,X6)) ) | $spl2),
  inference(subsumption_resolution,[],[f27087,f23487])).
fof(f23487,plain,(
  ( ! [X74,X75] : (~sP4(X74,X75) | ~p(s(bool,X75)) | sP4(sK8,X74)) ) | $spl2),
  inference(resolution,[],[f22742,f254])).
fof(f27087,plain,(
  ( ! [X6,X4,X7,X5] : (sP4(sK8,X4) | sP4(X4,X5) | sP3(X6,X7,X5) | ~p(s(bool,X5)) | ~sP2(X6,X5,X7) | ~sP1(X7,X5,X6)) ) | $spl2),
  inference(resolution,[],[f25196,f228])).
fof(f25196,plain,(
  ( ! [X87,X85,X86,X84] : (sP4(sK8,X84) | sP4(X84,X85) | sP0(X85,X86,X87)) ) | $spl2),
  inference(resolution,[],[f23489,f242])).
fof(f16720,plain,(
  s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | $spl375),
  inference(cnf_transformation,[],[f16720_D])).
fof(f16720_D,plain,(
  s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) <=> ~$spl375),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl375])])).
fof(f26829,plain,(
  ( ! [X25] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | p(s(bool,X25))) ) | ($spl2 | $spl44)),
  inference(subsumption_resolution,[],[f26785,f322])).
fof(f322,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f200])).
fof(f200,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X11] : (s(bool,'T_0') = s(bool,X11) <=> p(s(bool,X11)))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|1','thm.bool.EQ_CLAUSES|split|0')).
fof(f26785,plain,(
  ( ! [X25] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | p(s(bool,X25)) | s(bool,'T_0') = s(bool,X25)) ) | ($spl2 | $spl44)),
  inference(superposition,[],[f26097,f358])).
fof(f358,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|1',aHL_BOOL_CASES)).
fof(f26097,plain,(
  ( ! [X12] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X12))) | p(s(bool,X12))) ) | ($spl2 | $spl44)),
  inference(forward_demodulation,[],[f22782,f1251])).
fof(f22782,plain,(
  ( ! [X12] : (p(s(bool,X12)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,sK8))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X12)))) ) | $spl2),
  inference(superposition,[],[f366,f351])).
fof(f351,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,X2) | s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1))) != s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2)))) )),
  inference(cnf_transformation,[],[f210])).
fof(f210,plain,(
  ! [X0,X1,X2] : ((s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1))) != s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1))) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2)))))),
  inference(nnf_transformation,[],[f106])).
fof(f106,plain,(
  ! [X0,X1,X2] : (s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1))) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2))) <=> s(X0,X1) = s(X0,X2))),
  inference(flattening,[],[f105])).
fof(f105,plain,(
  ! [X0] : ! [X1,X2] : (s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1))) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X6] : ! [X12,X13] : (s('type.option.option'(X6),'const.option.SOME_1'(s(X6,X12))) = s('type.option.option'(X6),'const.option.SOME_1'(s(X6,X13))) <=> s(X6,X12) = s(X6,X13))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|1','thm.option.SOME_|49||49|')).
fof(f26823,plain,(
  ( ! [X21] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | p(s(bool,X21))) ) | ($spl2 | $spl44)),
  inference(duplicate_literal_removal,[],[f26781])).
fof(f26781,plain,(
  ( ! [X21] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | p(s(bool,X21)) | p(s(bool,X21))) ) | ($spl2 | $spl44)),
  inference(superposition,[],[f26097,f321])).
fof(f321,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f199])).
fof(f199,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f82])).
fof(f82,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X11] : (s(bool,'F_0') = s(bool,X11) <=> ~p(s(bool,X11)))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|1','thm.bool.EQ_CLAUSES|split|3')).
fof(f26824,plain,(
  ( ! [X20] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | p(s(bool,X20))) ) | ($spl2 | $spl44)),
  inference(duplicate_literal_removal,[],[f26780])).
fof(f26780,plain,(
  ( ! [X20] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | p(s(bool,X20)) | p(s(bool,X20))) ) | ($spl2 | $spl44)),
  inference(superposition,[],[f26097,f319])).
fof(f319,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f198])).
fof(f198,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f80])).
fof(f80,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X11] : (s(bool,'F_0') = s(bool,X11) <=> ~p(s(bool,X11)))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|1','thm.bool.EQ_CLAUSES|split|2')).
fof(f26736,plain,(
  ( ! [X0,X3,X1] : (sP4(X0,sK8) | s(bool,X1) != s(bool,X3) | p(s(bool,X3)) | sP4(X1,X0)) ) | $spl530),
  inference(cnf_transformation,[],[f26736_D])).
fof(f26736_D,plain,(
  ( ! [X0,X3,X1] : (sP4(X0,sK8) | s(bool,X1) != s(bool,X3) | p(s(bool,X3)) | sP4(X1,X0)) ) <=> ~$spl530),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl530])])).
fof(f26674,plain,(
  ( ! [X177,X175,X176,X174] : (sP4(X174,sK8) | sP4(X175,X174) | p(s(bool,X176)) | p(s(bool,X177)) | s(bool,X175) != s(bool,X177)) ) | $spl2),
  inference(resolution,[],[f23595,f360])).
fof(f26632,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP4(X35,sK8) | sP4(X36,X35) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X37) | s(bool,X38) != s(bool,X39)) ) | $spl2),
  inference(resolution,[],[f23595,f223])).
fof(f26629,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP4(X20,sK8) | sP4(X21,X20) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X21) != s(bool,X24)) ) | $spl2),
  inference(resolution,[],[f23595,f222])).
fof(f26628,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP4(X15,sK8) | sP4(X16,X15) | p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) ) | $spl2),
  inference(resolution,[],[f23595,f221])).
fof(f26627,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP4(X10,sK8) | sP4(X11,X10) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X12) | s(bool,X13) != s(bool,X14)) ) | $spl2),
  inference(resolution,[],[f23595,f219])).
fof(f26625,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP4(X0,sK8) | sP4(X1,X0) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) ) | $spl2),
  inference(resolution,[],[f23595,f217])).
fof(f26588,plain,(
  ( ! [X0,X3,X1] : (sP4(X0,sK8) | s(bool,X1) != s(bool,X3) | p(s(bool,X3)) | sP4(X0,X1)) ) | $spl528),
  inference(cnf_transformation,[],[f26588_D])).
fof(f26588_D,plain,(
  ( ! [X0,X3,X1] : (sP4(X0,sK8) | s(bool,X1) != s(bool,X3) | p(s(bool,X3)) | sP4(X0,X1)) ) <=> ~$spl528),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl528])])).
fof(f26525,plain,(
  ( ! [X177,X175,X176,X174] : (sP4(X174,sK8) | sP4(X174,X175) | p(s(bool,X176)) | p(s(bool,X177)) | s(bool,X175) != s(bool,X177)) ) | $spl2),
  inference(resolution,[],[f23594,f360])).
fof(f26483,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP4(X35,sK8) | sP4(X35,X36) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X37) | s(bool,X38) != s(bool,X39)) ) | $spl2),
  inference(resolution,[],[f23594,f223])).
fof(f26480,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP4(X20,sK8) | sP4(X20,X21) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X21) != s(bool,X24)) ) | $spl2),
  inference(resolution,[],[f23594,f222])).
fof(f26479,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP4(X15,sK8) | sP4(X15,X16) | p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) ) | $spl2),
  inference(resolution,[],[f23594,f221])).
fof(f26478,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP4(X10,sK8) | sP4(X10,X11) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X12) | s(bool,X13) != s(bool,X14)) ) | $spl2),
  inference(resolution,[],[f23594,f219])).
fof(f26476,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP4(X0,sK8) | sP4(X0,X1) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) ) | $spl2),
  inference(resolution,[],[f23594,f217])).
fof(f26185,plain,(
  ( ! [X0,X1] : (sP4(X0,sK8) | p(s(bool,X1)) | sP4(X1,X0)) ) | $spl526),
  inference(cnf_transformation,[],[f26185_D])).
fof(f26185_D,plain,(
  ( ! [X0,X1] : (sP4(X0,sK8) | p(s(bool,X1)) | sP4(X1,X0)) ) <=> ~$spl526),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl526])])).
fof(f26230,plain,(
  ( ! [X4,X5,X3] : (sP4(X3,sK8) | sP4(X4,X3) | p(s(bool,X4)) | p(s(bool,X5))) ) | ($spl2 | $spl44)),
  inference(subsumption_resolution,[],[f26217,f307])).
fof(f26217,plain,(
  ( ! [X4,X5,X3] : (sP4(X3,sK8) | sP4(X4,X3) | p(s(bool,'F_0')) | p(s(bool,X4)) | p(s(bool,X5))) ) | ($spl2 | $spl44)),
  inference(resolution,[],[f26047,f273])).
fof(f26047,plain,(
  ( ! [X78,X79,X77] : (sP4(X77,sK8) | sP4(X78,X77) | sP5(X78,X79)) ) | ($spl2 | $spl44)),
  inference(subsumption_resolution,[],[f26046,f361])).
fof(f361,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f282])).
fof(f282,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f175])).
fof(f175,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f64])).
fof(f64,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f63])).
fof(f63,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X6] : ! [X12] : (s(X6,X12) = s(X6,X12) <=> p(s(bool,'T_0')))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|1','thm.bool.REFL_CLAUSE')).
fof(f26046,plain,(
  ( ! [X78,X79,X77] : (~p(s(bool,'T_0')) | sP4(X77,sK8) | sP4(X78,X77) | sP5(X78,X79)) ) | ($spl2 | $spl44)),
  inference(forward_demodulation,[],[f26013,f1251])).
fof(f26013,plain,(
  ( ! [X78,X79,X77] : (~p(s(bool,sK8)) | sP4(X77,sK8) | sP4(X78,X77) | sP5(X78,X79)) ) | $spl2),
  inference(resolution,[],[f23593,f25518])).
fof(f25518,plain,(
  ( ! [X125,X123,X124] : (sP4(sK8,X123) | sP4(X124,X123) | sP5(X124,X125)) ) | $spl2),
  inference(resolution,[],[f23490,f271])).
fof(f26229,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,sK8) | sP4(X1,X0) | p(s(bool,X1)) | p(s(bool,X2))) ) | ($spl2 | $spl44)),
  inference(subsumption_resolution,[],[f26216,f307])).
fof(f26216,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,sK8) | sP4(X1,X0) | p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,X2))) ) | ($spl2 | $spl44)),
  inference(resolution,[],[f26047,f269])).
fof(f26186,plain,(
  ( ! [X4,X5,X3] : (sP4(X3,sK8) | sP4(X4,X3) | p(s(bool,X5)) | p(s(bool,X4))) ) | ($spl2 | $spl44)),
  inference(subsumption_resolution,[],[f26172,f307])).
fof(f26172,plain,(
  ( ! [X4,X5,X3] : (sP4(X3,sK8) | sP4(X4,X3) | p(s(bool,'F_0')) | p(s(bool,X5)) | p(s(bool,X4))) ) | ($spl2 | $spl44)),
  inference(resolution,[],[f26045,f273])).
fof(f26045,plain,(
  ( ! [X76,X74,X75] : (sP4(X74,sK8) | sP4(X75,X74) | sP5(X76,X75)) ) | ($spl2 | $spl44)),
  inference(subsumption_resolution,[],[f26044,f361])).
fof(f26044,plain,(
  ( ! [X76,X74,X75] : (~p(s(bool,'T_0')) | sP4(X74,sK8) | sP4(X75,X74) | sP5(X76,X75)) ) | ($spl2 | $spl44)),
  inference(forward_demodulation,[],[f26012,f1251])).
fof(f26012,plain,(
  ( ! [X76,X74,X75] : (~p(s(bool,sK8)) | sP4(X74,sK8) | sP4(X75,X74) | sP5(X76,X75)) ) | $spl2),
  inference(resolution,[],[f23593,f25517])).
fof(f25517,plain,(
  ( ! [X121,X122,X120] : (sP4(sK8,X120) | sP4(X121,X120) | sP5(X122,X121)) ) | $spl2),
  inference(resolution,[],[f23490,f270])).
fof(f26184,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,sK8) | sP4(X1,X0) | p(s(bool,X2)) | p(s(bool,X1))) ) | ($spl2 | $spl44)),
  inference(subsumption_resolution,[],[f26171,f307])).
fof(f26171,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,sK8) | sP4(X1,X0) | p(s(bool,'F_0')) | p(s(bool,X2)) | p(s(bool,X1))) ) | ($spl2 | $spl44)),
  inference(resolution,[],[f26045,f269])).
fof(f26095,plain,(
  ( ! [X0,X1] : (sP4(X0,sK8) | p(s(bool,X1)) | sP4(X0,X1)) ) | $spl524),
  inference(cnf_transformation,[],[f26095_D])).
fof(f26095_D,plain,(
  ( ! [X0,X1] : (sP4(X0,sK8) | p(s(bool,X1)) | sP4(X0,X1)) ) <=> ~$spl524),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl524])])).
fof(f26141,plain,(
  ( ! [X4,X5,X3] : (sP4(X3,sK8) | sP4(X3,X4) | p(s(bool,X4)) | p(s(bool,X5))) ) | ($spl2 | $spl44)),
  inference(subsumption_resolution,[],[f26128,f307])).
fof(f26128,plain,(
  ( ! [X4,X5,X3] : (sP4(X3,sK8) | sP4(X3,X4) | p(s(bool,'F_0')) | p(s(bool,X4)) | p(s(bool,X5))) ) | ($spl2 | $spl44)),
  inference(resolution,[],[f26039,f273])).
fof(f26039,plain,(
  ( ! [X68,X69,X67] : (sP4(X67,sK8) | sP4(X67,X68) | sP5(X68,X69)) ) | ($spl2 | $spl44)),
  inference(subsumption_resolution,[],[f26038,f361])).
fof(f26038,plain,(
  ( ! [X68,X69,X67] : (~p(s(bool,'T_0')) | sP4(X67,sK8) | sP4(X67,X68) | sP5(X68,X69)) ) | ($spl2 | $spl44)),
  inference(forward_demodulation,[],[f26009,f1251])).
fof(f26009,plain,(
  ( ! [X68,X69,X67] : (~p(s(bool,sK8)) | sP4(X67,sK8) | sP4(X67,X68) | sP5(X68,X69)) ) | $spl2),
  inference(resolution,[],[f23593,f25208])).
fof(f26140,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,sK8) | sP4(X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) ) | ($spl2 | $spl44)),
  inference(subsumption_resolution,[],[f26127,f307])).
fof(f26127,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,sK8) | sP4(X0,X1) | p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,X2))) ) | ($spl2 | $spl44)),
  inference(resolution,[],[f26039,f269])).
fof(f26096,plain,(
  ( ! [X4,X5,X3] : (sP4(X3,sK8) | sP4(X3,X4) | p(s(bool,X5)) | p(s(bool,X4))) ) | ($spl2 | $spl44)),
  inference(subsumption_resolution,[],[f26082,f307])).
fof(f26082,plain,(
  ( ! [X4,X5,X3] : (sP4(X3,sK8) | sP4(X3,X4) | p(s(bool,'F_0')) | p(s(bool,X5)) | p(s(bool,X4))) ) | ($spl2 | $spl44)),
  inference(resolution,[],[f26037,f273])).
fof(f26037,plain,(
  ( ! [X66,X64,X65] : (sP4(X64,sK8) | sP4(X64,X65) | sP5(X66,X65)) ) | ($spl2 | $spl44)),
  inference(subsumption_resolution,[],[f26036,f361])).
fof(f26036,plain,(
  ( ! [X66,X64,X65] : (~p(s(bool,'T_0')) | sP4(X64,sK8) | sP4(X64,X65) | sP5(X66,X65)) ) | ($spl2 | $spl44)),
  inference(forward_demodulation,[],[f26008,f1251])).
fof(f26008,plain,(
  ( ! [X66,X64,X65] : (~p(s(bool,sK8)) | sP4(X64,sK8) | sP4(X64,X65) | sP5(X66,X65)) ) | $spl2),
  inference(resolution,[],[f23593,f25207])).
fof(f26094,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,sK8) | sP4(X0,X1) | p(s(bool,X2)) | p(s(bool,X1))) ) | ($spl2 | $spl44)),
  inference(subsumption_resolution,[],[f26081,f307])).
fof(f26081,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,sK8) | sP4(X0,X1) | p(s(bool,'F_0')) | p(s(bool,X2)) | p(s(bool,X1))) ) | ($spl2 | $spl44)),
  inference(resolution,[],[f26037,f269])).
fof(f25746,plain,(
  s(bool,'T_0') = s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7))) | $spl522),
  inference(cnf_transformation,[],[f25746_D])).
fof(f25746_D,plain,(
  s(bool,'T_0') = s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7))) <=> ~$spl522),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl522])])).
fof(f25744,plain,(
  ( ! [X25] : (p(s(bool,X25)) | s(bool,'T_0') = s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) | $spl24),
  inference(subsumption_resolution,[],[f25721,f319])).
fof(f25721,plain,(
  ( ! [X25] : (s(bool,'F_0') != s(bool,X25) | p(s(bool,X25)) | s(bool,'T_0') = s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) | $spl24),
  inference(superposition,[],[f767,f358])).
fof(f25734,plain,(
  s(bool,'F_0') != s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7))) | $spl521),
  inference(cnf_transformation,[],[f25734_D])).
fof(f25734_D,plain,(
  s(bool,'F_0') != s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7))) <=> ~$spl521),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl521])])).
fof(f25735,plain,(
  ( ! [X25] : (s(bool,'F_0') != s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7))) | p(s(bool,X25))) ) | $spl24),
  inference(subsumption_resolution,[],[f25690,f322])).
fof(f25690,plain,(
  ( ! [X25] : (s(bool,'F_0') != s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7))) | p(s(bool,X25)) | s(bool,'T_0') = s(bool,X25)) ) | $spl24),
  inference(superposition,[],[f767,f358])).
fof(f25729,plain,(
  ( ! [X21] : (s(bool,'F_0') != s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7))) | p(s(bool,X21))) ) | $spl24),
  inference(duplicate_literal_removal,[],[f25686])).
fof(f25686,plain,(
  ( ! [X21] : (s(bool,'F_0') != s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7))) | p(s(bool,X21)) | p(s(bool,X21))) ) | $spl24),
  inference(superposition,[],[f767,f321])).
fof(f25730,plain,(
  ( ! [X20] : (s(bool,'F_0') != s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7))) | p(s(bool,X20))) ) | $spl24),
  inference(duplicate_literal_removal,[],[f25685])).
fof(f25685,plain,(
  ( ! [X20] : (s(bool,'F_0') != s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7))) | p(s(bool,X20)) | p(s(bool,X20))) ) | $spl24),
  inference(superposition,[],[f767,f319])).
fof(f25598,plain,(
  ( ! [X0,X3,X1] : (sP4(sK8,X0) | s(bool,X1) != s(bool,X3) | p(s(bool,X3)) | sP4(X1,X0)) ) | $spl518),
  inference(cnf_transformation,[],[f25598_D])).
fof(f25598_D,plain,(
  ( ! [X0,X3,X1] : (sP4(sK8,X0) | s(bool,X1) != s(bool,X3) | p(s(bool,X3)) | sP4(X1,X0)) ) <=> ~$spl518),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl518])])).
fof(f25536,plain,(
  ( ! [X177,X175,X176,X174] : (sP4(sK8,X174) | sP4(X175,X174) | p(s(bool,X176)) | p(s(bool,X177)) | s(bool,X175) != s(bool,X177)) ) | $spl2),
  inference(resolution,[],[f23490,f360])).
fof(f25494,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP4(sK8,X35) | sP4(X36,X35) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X37) | s(bool,X38) != s(bool,X39)) ) | $spl2),
  inference(resolution,[],[f23490,f223])).
fof(f25491,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP4(sK8,X20) | sP4(X21,X20) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X21) != s(bool,X24)) ) | $spl2),
  inference(resolution,[],[f23490,f222])).
fof(f25490,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP4(sK8,X15) | sP4(X16,X15) | p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) ) | $spl2),
  inference(resolution,[],[f23490,f221])).
fof(f25489,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP4(sK8,X10) | sP4(X11,X10) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X12) | s(bool,X13) != s(bool,X14)) ) | $spl2),
  inference(resolution,[],[f23490,f219])).
fof(f25487,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP4(sK8,X0) | sP4(X1,X0) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) ) | $spl2),
  inference(resolution,[],[f23490,f217])).
fof(f25287,plain,(
  ( ! [X0,X3,X1] : (sP4(sK8,X0) | s(bool,X1) != s(bool,X3) | p(s(bool,X3)) | sP4(X0,X1)) ) | $spl516),
  inference(cnf_transformation,[],[f25287_D])).
fof(f25287_D,plain,(
  ( ! [X0,X3,X1] : (sP4(sK8,X0) | s(bool,X1) != s(bool,X3) | p(s(bool,X3)) | sP4(X0,X1)) ) <=> ~$spl516),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl516])])).
fof(f25226,plain,(
  ( ! [X177,X175,X176,X174] : (sP4(sK8,X174) | sP4(X174,X175) | p(s(bool,X176)) | p(s(bool,X177)) | s(bool,X175) != s(bool,X177)) ) | $spl2),
  inference(resolution,[],[f23489,f360])).
fof(f25184,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP4(sK8,X35) | sP4(X35,X36) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X37) | s(bool,X38) != s(bool,X39)) ) | $spl2),
  inference(resolution,[],[f23489,f223])).
fof(f25181,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP4(sK8,X20) | sP4(X20,X21) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X21) != s(bool,X24)) ) | $spl2),
  inference(resolution,[],[f23489,f222])).
fof(f25180,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP4(sK8,X15) | sP4(X15,X16) | p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) ) | $spl2),
  inference(resolution,[],[f23489,f221])).
fof(f25179,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP4(sK8,X10) | sP4(X10,X11) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X12) | s(bool,X13) != s(bool,X14)) ) | $spl2),
  inference(resolution,[],[f23489,f219])).
fof(f25177,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP4(sK8,X0) | sP4(X0,X1) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) ) | $spl2),
  inference(resolution,[],[f23489,f217])).
fof(f767,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7))) | p(s(bool,X1))) ) | $spl24),
  inference(cnf_transformation,[],[f767_D])).
fof(f767_D,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7))) | p(s(bool,X1))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f25021,plain,(
  ( ! [X76,X77] : (p(s(bool,X76)) | p(s(bool,X77)) | s(bool,X77) != s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) | $spl275),
  inference(resolution,[],[f24764,f360])).
fof(f24764,plain,(
  p(s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) | $spl275),
  inference(subsumption_resolution,[],[f24754,f361])).
fof(f24754,plain,(
  p(s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) | ~p(s(bool,'T_0')) | $spl275),
  inference(resolution,[],[f11822,f258])).
fof(f11822,plain,(
  ~sP4('const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)),'T_0') | $spl275),
  inference(cnf_transformation,[],[f11822_D])).
fof(f11822_D,plain,(
  ~sP4('const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)),'T_0') <=> ~$spl275),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl275])])).
fof(f24979,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7))) | s(bool,X22) != s(bool,X23)) ) | $spl275),
  inference(resolution,[],[f24764,f223])).
fof(f24976,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) | $spl275),
  inference(resolution,[],[f24764,f222])).
fof(f24975,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) | $spl275),
  inference(resolution,[],[f24764,f221])).
fof(f24974,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7))) | s(bool,X7) != s(bool,X8)) ) | $spl275),
  inference(resolution,[],[f24764,f219])).
fof(f24972,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) | $spl275),
  inference(resolution,[],[f24764,f217])).
fof(f7786,plain,(
  sP4('F_0','F_0') | $spl188),
  inference(cnf_transformation,[],[f7786_D])).
fof(f7786_D,plain,(
  sP4('F_0','F_0') <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f24635,plain,(
  ( ! [X149,X148] : (sP4('F_0','F_0') | p(s(bool,X148)) | p(s(bool,X149)) | ~sP5(X148,X149)) ) | $spl230),
  inference(resolution,[],[f23352,f269])).
fof(f23352,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | sP4('F_0',X1)) ) | $spl230),
  inference(subsumption_resolution,[],[f23342,f361])).
fof(f23342,plain,(
  ( ! [X1] : (sP4('F_0',X1) | ~p(s(bool,X1)) | ~p(s(bool,'T_0'))) ) | $spl230),
  inference(resolution,[],[f9926,f254])).
fof(f24633,plain,(
  ( ! [X144,X145] : (sP4('F_0','F_0') | p(s(bool,X144)) | p(s(bool,X145)) | ~sP5(X144,X145)) ) | $spl230),
  inference(resolution,[],[f23352,f273])).
fof(f10088,plain,(
  ( ! [X2,X0] : (sP4('F_0',X0) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) | $spl232),
  inference(cnf_transformation,[],[f10088_D])).
fof(f10088_D,plain,(
  ( ! [X2,X0] : (sP4('F_0',X0) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) <=> ~$spl232),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl232])])).
fof(f24612,plain,(
  ( ! [X132,X130,X131] : (sP4('F_0',X130) | p(s(bool,X131)) | ~p(s(bool,X132)) | s(bool,X130) != s(bool,X132)) ) | $spl230),
  inference(resolution,[],[f23352,f360])).
fof(f24567,plain,(
  ( ! [X30,X28,X31,X29] : (sP4('F_0',X28) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) ) | $spl230),
  inference(resolution,[],[f23352,f223])).
fof(f24566,plain,(
  ( ! [X26,X24,X27,X25] : (sP4('F_0',X24) | p(s(bool,X25)) | ~p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X24) != s(bool,X26)) ) | $spl230),
  inference(resolution,[],[f23352,f222])).
fof(f24564,plain,(
  ( ! [X19,X17,X18,X16] : (sP4('F_0',X16) | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) ) | $spl230),
  inference(resolution,[],[f23352,f221])).
fof(f24562,plain,(
  ( ! [X10,X8,X11,X9] : (sP4('F_0',X8) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) ) | $spl230),
  inference(resolution,[],[f23352,f219])).
fof(f24560,plain,(
  ( ! [X2,X0,X3,X1] : (sP4('F_0',X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl230),
  inference(resolution,[],[f23352,f217])).
fof(f24428,plain,(
  ( ! [X149,X148] : (sP4('F_0','F_0') | p(s(bool,X148)) | p(s(bool,X149)) | ~sP5(X148,X149)) ) | $spl180),
  inference(resolution,[],[f7586,f269])).
fof(f24426,plain,(
  ( ! [X144,X145] : (sP4('F_0','F_0') | p(s(bool,X144)) | p(s(bool,X145)) | ~sP5(X144,X145)) ) | $spl180),
  inference(resolution,[],[f7586,f273])).
fof(f7777,plain,(
  ( ! [X2,X0] : (sP4(X0,'F_0') | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) | $spl182),
  inference(cnf_transformation,[],[f7777_D])).
fof(f7777_D,plain,(
  ( ! [X2,X0] : (sP4(X0,'F_0') | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f24405,plain,(
  ( ! [X132,X130,X131] : (sP4(X130,'F_0') | p(s(bool,X131)) | ~p(s(bool,X132)) | s(bool,X130) != s(bool,X132)) ) | $spl180),
  inference(resolution,[],[f7586,f360])).
fof(f24360,plain,(
  ( ! [X30,X28,X31,X29] : (sP4(X28,'F_0') | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) ) | $spl180),
  inference(resolution,[],[f7586,f223])).
fof(f24359,plain,(
  ( ! [X26,X24,X27,X25] : (sP4(X24,'F_0') | p(s(bool,X25)) | ~p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X24) != s(bool,X26)) ) | $spl180),
  inference(resolution,[],[f7586,f222])).
fof(f24357,plain,(
  ( ! [X19,X17,X18,X16] : (sP4(X16,'F_0') | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) ) | $spl180),
  inference(resolution,[],[f7586,f221])).
fof(f24355,plain,(
  ( ! [X10,X8,X11,X9] : (sP4(X8,'F_0') | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) ) | $spl180),
  inference(resolution,[],[f7586,f219])).
fof(f24353,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(X0,'F_0') | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl180),
  inference(resolution,[],[f7586,f217])).
fof(f4600,plain,(
  ( ! [X2,X0] : (sP4(X0,'T_0') | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) | $spl124),
  inference(cnf_transformation,[],[f4600_D])).
fof(f4600_D,plain,(
  ( ! [X2,X0] : (sP4(X0,'T_0') | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f24186,plain,(
  ( ! [X127,X125,X126] : (sP4(X125,'T_0') | p(s(bool,X126)) | p(s(bool,X127)) | s(bool,X125) != s(bool,X127)) ) | $spl120),
  inference(resolution,[],[f4467,f360])).
fof(f24144,plain,(
  ( ! [X30,X28,X31,X29] : (sP4(X28,'T_0') | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) ) | $spl120),
  inference(resolution,[],[f4467,f223])).
fof(f24141,plain,(
  ( ! [X19,X17,X18,X16] : (sP4(X16,'T_0') | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) ) | $spl120),
  inference(resolution,[],[f4467,f222])).
fof(f24140,plain,(
  ( ! [X14,X12,X15,X13] : (sP4(X12,'T_0') | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl120),
  inference(resolution,[],[f4467,f221])).
fof(f24139,plain,(
  ( ! [X10,X8,X11,X9] : (sP4(X8,'T_0') | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) ) | $spl120),
  inference(resolution,[],[f4467,f219])).
fof(f24137,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(X0,'T_0') | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl120),
  inference(resolution,[],[f4467,f217])).
fof(f4066,plain,(
  ( ! [X2,X0] : (sP4('T_0',X0) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) | $spl102),
  inference(cnf_transformation,[],[f4066_D])).
fof(f4066_D,plain,(
  ( ! [X2,X0] : (sP4('T_0',X0) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f24060,plain,(
  ( ! [X127,X125,X126] : (sP4('T_0',X125) | p(s(bool,X126)) | p(s(bool,X127)) | s(bool,X125) != s(bool,X127)) ) | $spl100),
  inference(resolution,[],[f3618,f360])).
fof(f24018,plain,(
  ( ! [X30,X28,X31,X29] : (sP4('T_0',X28) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) ) | $spl100),
  inference(resolution,[],[f3618,f223])).
fof(f24015,plain,(
  ( ! [X19,X17,X18,X16] : (sP4('T_0',X16) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) ) | $spl100),
  inference(resolution,[],[f3618,f222])).
fof(f24014,plain,(
  ( ! [X14,X12,X15,X13] : (sP4('T_0',X12) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl100),
  inference(resolution,[],[f3618,f221])).
fof(f24013,plain,(
  ( ! [X10,X8,X11,X9] : (sP4('T_0',X8) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) ) | $spl100),
  inference(resolution,[],[f3618,f219])).
fof(f24011,plain,(
  ( ! [X2,X0,X3,X1] : (sP4('T_0',X0) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl100),
  inference(resolution,[],[f3618,f217])).
fof(f23911,plain,(
  $false | ($spl20 | $spl44)),
  inference(subsumption_resolution,[],[f23841,f23776])).
fof(f23776,plain,(
  ( ! [X1] : (~p(s(bool,X1))) ) | ($spl20 | $spl44)),
  inference(subsumption_resolution,[],[f23775,f325])).
fof(f325,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f201])).
fof(f201,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f85])).
fof(f85,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X11] : (s(bool,'T_0') = s(bool,X11) <=> p(s(bool,X11)))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|1','thm.bool.EQ_CLAUSES|split|1')).
fof(f23775,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,X1) | ~p(s(bool,X1))) ) | ($spl20 | $spl44)),
  inference(forward_demodulation,[],[f543,f1251])).
fof(f23841,plain,(
  p(s(bool,'F_0')) | ($spl20 | $spl44)),
  inference(resolution,[],[f23776,f335])).
fof(f335,plain,(
  p(s(bool,'F_0')) | p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f203])).
fof(f203,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f93])).
fof(f93,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f21])).
fof(f21,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|1','thm.bool.NOT_CLAUSES|split|1')).
fof(f23910,plain,(
  $false | ($spl20 | $spl44)),
  inference(subsumption_resolution,[],[f23840,f23776])).
fof(f23840,plain,(
  p(s(bool,'F_0')) | ($spl20 | $spl44)),
  inference(resolution,[],[f23776,f333])).
fof(f333,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f202])).
fof(f202,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f92])).
fof(f92,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f22])).
fof(f22,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|1','thm.bool.NOT_CLAUSES|split|2')).
fof(f23909,plain,(
  $false | ($spl20 | $spl44)),
  inference(subsumption_resolution,[],[f23836,f23776])).
fof(f23836,plain,(
  ( ! [X129] : (p(s(bool,X129))) ) | ($spl20 | $spl44)),
  inference(resolution,[],[f23776,f300])).
fof(f300,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f187])).
fof(f187,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))))),
  inference(flattening,[],[f186])).
fof(f186,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X11] : ((p(s(bool,X11)) => p(s(bool,X11))) <=> p(s(bool,'T_0')))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|1','thm.bool.IMP_CLAUSES|split|3')).
fof(f23908,plain,(
  $false | ($spl20 | $spl44)),
  inference(subsumption_resolution,[],[f23835,f23776])).
fof(f23835,plain,(
  ( ! [X128] : (p(s(bool,X128))) ) | ($spl20 | $spl44)),
  inference(resolution,[],[f23776,f297])).
fof(f297,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f185])).
fof(f185,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,'T_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f184])).
fof(f184,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (~p(s(bool,'T_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f130])).
fof(f130,plain,(
  ! [X0] : ((~p(s(bool,'T_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(ennf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0] : ((p(s(bool,'T_0')) => p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X11] : ((p(s(bool,'T_0')) => p(s(bool,X11))) <=> p(s(bool,X11)))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|1','thm.bool.IMP_CLAUSES|split|0')).
fof(f23907,plain,(
  $false | ($spl20 | $spl44)),
  inference(subsumption_resolution,[],[f23834,f23776])).
fof(f23834,plain,(
  ( ! [X127] : (p(s(bool,X127))) ) | ($spl20 | $spl44)),
  inference(resolution,[],[f23776,f294])).
fof(f294,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f183])).
fof(f183,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,'T_0'))))),
  inference(flattening,[],[f182])).
fof(f182,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f129])).
fof(f129,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f72])).
fof(f72,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X11] : ((p(s(bool,X11)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|1','thm.bool.IMP_CLAUSES|split|1')).
fof(f23906,plain,(
  $false | ($spl20 | $spl44)),
  inference(subsumption_resolution,[],[f23832,f23776])).
fof(f23832,plain,(
  p(s(bool,'F_0')) | ($spl20 | $spl44)),
  inference(resolution,[],[f23776,f291])).
fof(f291,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f181])).
fof(f181,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f180])).
fof(f180,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f128])).
fof(f128,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X11] : ((p(s(bool,'F_0')) => p(s(bool,X11))) <=> p(s(bool,'T_0')))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|1','thm.bool.IMP_CLAUSES|split|2')).
fof(f23842,plain,(
  $false | ($spl20 | $spl44)),
  inference(resolution,[],[f23776,f337])).
fof(f337,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  p(s(bool,'T_0'))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|1','thm.bool.TRUTH')).
fof(f23843,plain,(
  $false | ($spl20 | $spl44)),
  inference(resolution,[],[f23776,f338])).
fof(f338,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|1',aHL_TRUTH)).
fof(f23844,plain,(
  $false | ($spl20 | $spl44)),
  inference(resolution,[],[f23776,f361])).
fof(f368,plain,(
  ~p(s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) | $spl5),
  inference(cnf_transformation,[],[f368_D])).
fof(f368_D,plain,(
  ~p(s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f23756,plain,(
  ( ! [X2,X3] : (sP3(X2,X3,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7))) | ~sP1(X3,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)),X2) | ~sP2(X2,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)),X3)) ) | $spl514),
  inference(cnf_transformation,[],[f23756_D])).
fof(f23756_D,plain,(
  ( ! [X2,X3] : (sP3(X2,X3,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7))) | ~sP1(X3,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)),X2) | ~sP2(X2,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)),X3)) ) <=> ~$spl514),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl514])])).
fof(f23746,plain,(
  ( ! [X2,X3] : (sP3(X2,X3,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7))) | ~p(s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) | ~sP2(X2,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)),X3) | ~sP1(X3,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)),X2)) ) | $spl4),
  inference(resolution,[],[f707,f228])).
fof(f707,plain,(
  ( ! [X47,X46] : (sP0('const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)),X46,X47)) ) | $spl4),
  inference(resolution,[],[f369,f242])).
fof(f23668,plain,(
  ( ! [X2,X0] : (sP4(X0,sK8) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) | $spl504),
  inference(cnf_transformation,[],[f23668_D])).
fof(f23668_D,plain,(
  ( ! [X2,X0] : (sP4(X0,sK8) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) <=> ~$spl504),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl504])])).
fof(f23619,plain,(
  ( ! [X127,X125,X126] : (sP4(X125,sK8) | p(s(bool,X126)) | p(s(bool,X127)) | s(bool,X125) != s(bool,X127)) ) | $spl2),
  inference(resolution,[],[f22743,f360])).
fof(f23577,plain,(
  ( ! [X30,X28,X31,X29] : (sP4(X28,sK8) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) ) | $spl2),
  inference(resolution,[],[f22743,f223])).
fof(f23574,plain,(
  ( ! [X19,X17,X18,X16] : (sP4(X16,sK8) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) ) | $spl2),
  inference(resolution,[],[f22743,f222])).
fof(f23573,plain,(
  ( ! [X14,X12,X15,X13] : (sP4(X12,sK8) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl2),
  inference(resolution,[],[f22743,f221])).
fof(f23572,plain,(
  ( ! [X10,X8,X11,X9] : (sP4(X8,sK8) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) ) | $spl2),
  inference(resolution,[],[f22743,f219])).
fof(f23570,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(X0,sK8) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl2),
  inference(resolution,[],[f22743,f217])).
fof(f23563,plain,(
  ( ! [X2,X0] : (sP4(sK8,X0) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) | $spl502),
  inference(cnf_transformation,[],[f23563_D])).
fof(f23563_D,plain,(
  ( ! [X2,X0] : (sP4(sK8,X0) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) <=> ~$spl502),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl502])])).
fof(f23514,plain,(
  ( ! [X127,X125,X126] : (sP4(sK8,X125) | p(s(bool,X126)) | p(s(bool,X127)) | s(bool,X125) != s(bool,X127)) ) | $spl2),
  inference(resolution,[],[f22742,f360])).
fof(f23472,plain,(
  ( ! [X30,X28,X31,X29] : (sP4(sK8,X28) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) ) | $spl2),
  inference(resolution,[],[f22742,f223])).
fof(f23469,plain,(
  ( ! [X19,X17,X18,X16] : (sP4(sK8,X16) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) ) | $spl2),
  inference(resolution,[],[f22742,f222])).
fof(f23468,plain,(
  ( ! [X14,X12,X15,X13] : (sP4(sK8,X12) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl2),
  inference(resolution,[],[f22742,f221])).
fof(f23467,plain,(
  ( ! [X10,X8,X11,X9] : (sP4(sK8,X8) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) ) | $spl2),
  inference(resolution,[],[f22742,f219])).
fof(f23465,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(sK8,X0) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl2),
  inference(resolution,[],[f22742,f217])).
fof(f22980,plain,(
  $false | ($spl55 | $spl99 | $spl116)),
  inference(subsumption_resolution,[],[f22976,f1499])).
fof(f1499,plain,(
  ~sP4(sK8,'T_0') | $spl55),
  inference(cnf_transformation,[],[f1499_D])).
fof(f1499_D,plain,(
  ~sP4(sK8,'T_0') <=> ~$spl55),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl55])])).
fof(f22976,plain,(
  sP4(sK8,'T_0') | ($spl99 | $spl116)),
  inference(resolution,[],[f4449,f3616])).
fof(f3616,plain,(
  ~sP4(sK8,sK8) | $spl99),
  inference(cnf_transformation,[],[f3616_D])).
fof(f3616_D,plain,(
  ~sP4(sK8,sK8) <=> ~$spl99),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl99])])).
fof(f22979,plain,(
  $false | ($spl33 | $spl116 | $spl179)),
  inference(subsumption_resolution,[],[f22975,f7584])).
fof(f7584,plain,(
  ~sP4('T_0','T_0') | $spl179),
  inference(cnf_transformation,[],[f7584_D])).
fof(f7584_D,plain,(
  ~sP4('T_0','T_0') <=> ~$spl179),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl179])])).
fof(f22975,plain,(
  sP4('T_0','T_0') | ($spl33 | $spl116)),
  inference(resolution,[],[f4449,f902])).
fof(f902,plain,(
  ~sP4('T_0',sK8) | $spl33),
  inference(cnf_transformation,[],[f902_D])).
fof(f902_D,plain,(
  ~sP4('T_0',sK8) <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f22970,plain,(
  $false | ($spl33 | $spl94 | $spl99)),
  inference(subsumption_resolution,[],[f22962,f902])).
fof(f22962,plain,(
  sP4('T_0',sK8) | ($spl94 | $spl99)),
  inference(resolution,[],[f3588,f3616])).
fof(f22969,plain,(
  $false | ($spl33 | $spl94 | $spl179)),
  inference(subsumption_resolution,[],[f22961,f7584])).
fof(f22961,plain,(
  sP4('T_0','T_0') | ($spl33 | $spl94)),
  inference(resolution,[],[f3588,f902])).
fof(f22964,plain,(
  $false | ($spl33 | $spl94 | $spl99)),
  inference(subsumption_resolution,[],[f22952,f3616])).
fof(f22952,plain,(
  sP4(sK8,sK8) | ($spl33 | $spl94)),
  inference(resolution,[],[f3588,f902])).
fof(f22963,plain,(
  $false | ($spl33 | $spl94 | $spl179)),
  inference(subsumption_resolution,[],[f22951,f902])).
fof(f22951,plain,(
  sP4('T_0',sK8) | ($spl94 | $spl179)),
  inference(resolution,[],[f3588,f7584])).
fof(f22908,plain,(
  $false | ($spl2 | $spl32)),
  inference(subsumption_resolution,[],[f22716,f366])).
fof(f22716,plain,(
  ~p(s(bool,sK8)) | $spl32),
  inference(subsumption_resolution,[],[f22713,f361])).
fof(f22713,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,sK8)) | $spl32),
  inference(resolution,[],[f903,f254])).
fof(f1251,plain,(
  s(bool,'T_0') = s(bool,sK8) | $spl44),
  inference(cnf_transformation,[],[f1251_D])).
fof(f1251_D,plain,(
  s(bool,'T_0') = s(bool,sK8) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f1253,plain,(
  s('type.option.option'(sK6),'const.option.NONE_0') != s('type.option.option'(sK6),'const.bool.COND_3'(s(bool,'F_0'),s('type.option.option'(sK6),'const.option.NONE_0'),s('type.option.option'(sK6),sK7))) | $spl47),
  inference(cnf_transformation,[],[f1253_D])).
fof(f1253_D,plain,(
  s('type.option.option'(sK6),'const.option.NONE_0') != s('type.option.option'(sK6),'const.bool.COND_3'(s(bool,'F_0'),s('type.option.option'(sK6),'const.option.NONE_0'),s('type.option.option'(sK6),sK7))) <=> ~$spl47),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl47])])).
fof(f22818,plain,(
  s('type.option.option'(sK6),'const.option.NONE_0') != s('type.option.option'(sK6),'const.bool.COND_3'(s(bool,'F_0'),s('type.option.option'(sK6),'const.option.NONE_0'),s('type.option.option'(sK6),sK7))) | s(bool,'T_0') = s(bool,sK8) | $spl1),
  inference(superposition,[],[f365,f358])).
fof(f1247,plain,(
  s(bool,'F_0') = s(bool,sK8) | $spl40),
  inference(cnf_transformation,[],[f1247_D])).
fof(f1247_D,plain,(
  s(bool,'F_0') = s(bool,sK8) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f1249,plain,(
  s('type.option.option'(sK6),'const.option.NONE_0') != s('type.option.option'(sK6),'const.bool.COND_3'(s(bool,'T_0'),s('type.option.option'(sK6),'const.option.NONE_0'),s('type.option.option'(sK6),sK7))) | $spl43),
  inference(cnf_transformation,[],[f1249_D])).
fof(f1249_D,plain,(
  s('type.option.option'(sK6),'const.option.NONE_0') != s('type.option.option'(sK6),'const.bool.COND_3'(s(bool,'T_0'),s('type.option.option'(sK6),'const.option.NONE_0'),s('type.option.option'(sK6),sK7))) <=> ~$spl43),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl43])])).
fof(f22817,plain,(
  s('type.option.option'(sK6),'const.option.NONE_0') != s('type.option.option'(sK6),'const.bool.COND_3'(s(bool,'T_0'),s('type.option.option'(sK6),'const.option.NONE_0'),s('type.option.option'(sK6),sK7))) | s(bool,'F_0') = s(bool,sK8) | $spl1),
  inference(superposition,[],[f365,f358])).
fof(f22800,plain,(
  $false | ($spl2 | $spl54)),
  inference(subsumption_resolution,[],[f22799,f366])).
fof(f22799,plain,(
  ~p(s(bool,sK8)) | $spl54),
  inference(subsumption_resolution,[],[f22795,f361])).
fof(f22795,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,sK8)) | $spl54),
  inference(resolution,[],[f1500,f261])).
fof(f261,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | ~sP4(X1,X0)) )),
  inference(cnf_transformation,[],[f168])).
fof(f168,plain,(
  ! [X0,X1] : ((((p(s(bool,X1)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)))) | sP4(X1,X0)) & (~sP4(X1,X0) | ((~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X0)) | p(s(bool,X1))))))),
  inference(nnf_transformation,[],[f143])).
fof(f143,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> sP4(X1,X0))),
  inference(definition_folding,[],[f57,f142])).
fof(f57,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  inference(flattening,[],[f56])).
fof(f56,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  inference(rectify,[],[f49])).
fof(f49,axiom,(
  ! [X21,X22] : ((p(s(bool,X22)) <=> ~p(s(bool,X21))) <=> ((p(s(bool,X22)) | p(s(bool,X21))) & (~p(s(bool,X21)) | ~p(s(bool,X22)))))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|1','thm.sat.sat_29_prove_ltb')).
fof(f22798,plain,(
  $false | ($spl2 | $spl54)),
  inference(subsumption_resolution,[],[f22797,f338])).
fof(f22797,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl54)),
  inference(subsumption_resolution,[],[f22794,f366])).
fof(f22794,plain,(
  ~p(s(bool,sK8)) | ~p(s(bool,'T_0')) | $spl54),
  inference(resolution,[],[f1500,f254])).
fof(f987,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,sK8) != s(bool,X1)) ) | $spl38),
  inference(cnf_transformation,[],[f987_D])).
fof(f987_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,sK8) != s(bool,X1)) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f22767,plain,(
  ( ! [X76,X77] : (p(s(bool,X76)) | p(s(bool,X77)) | s(bool,sK8) != s(bool,X77)) ) | $spl2),
  inference(resolution,[],[f366,f360])).
fof(f22725,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,sK8) != s(bool,X21) | s(bool,X22) != s(bool,X23)) ) | $spl2),
  inference(resolution,[],[f366,f223])).
fof(f22722,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,sK8) != s(bool,X14)) ) | $spl2),
  inference(resolution,[],[f366,f222])).
fof(f22721,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,sK8) != s(bool,X10)) ) | $spl2),
  inference(resolution,[],[f366,f221])).
fof(f22720,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,sK8) != s(bool,X6) | s(bool,X7) != s(bool,X8)) ) | $spl2),
  inference(resolution,[],[f366,f219])).
fof(f22718,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,sK8) != s(bool,X1)) ) | $spl2),
  inference(resolution,[],[f366,f217])).
fof(f367,plain,(
  ~p(s(bool,sK8)) | $spl3),
  inference(cnf_transformation,[],[f367_D])).
fof(f367_D,plain,(
  ~p(s(bool,sK8)) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f2290,plain,(
  ( ! [X10] : (~p(s(bool,sK8)) | ~p(s(bool,X10))) ) | $spl45),
  inference(subsumption_resolution,[],[f2273,f323])).
fof(f323,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f200])).
fof(f2273,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,X10) | ~p(s(bool,sK8)) | ~p(s(bool,X10))) ) | $spl45),
  inference(superposition,[],[f1250,f281])).
fof(f281,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f125])).
fof(f125,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f62])).
fof(f62,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X9,X10] : ((p(s(bool,X9)) => p(s(bool,X10))) => ((p(s(bool,X10)) => p(s(bool,X9))) => s(bool,X9) = s(bool,X10)))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|1','thm.bool.IMP_ANTISYM_AX')).
fof(f1250,plain,(
  s(bool,'T_0') != s(bool,sK8) | $spl45),
  inference(cnf_transformation,[],[f1250_D])).
fof(f1250_D,plain,(
  s(bool,'T_0') != s(bool,sK8) <=> ~$spl45),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl45])])).
fof(f2291,plain,(
  ( ! [X11] : (~p(s(bool,X11)) | ~p(s(bool,sK8))) ) | $spl45),
  inference(subsumption_resolution,[],[f2274,f323])).
fof(f2274,plain,(
  ( ! [X11] : (s(bool,'T_0') != s(bool,X11) | ~p(s(bool,X11)) | ~p(s(bool,sK8))) ) | $spl45),
  inference(superposition,[],[f1250,f281])).
fof(f366,plain,(
  p(s(bool,sK8)) | $spl2),
  inference(cnf_transformation,[],[f366_D])).
fof(f366_D,plain,(
  p(s(bool,sK8)) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f22703,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl501),
  inference(cnf_transformation,[],[f22703_D])).
fof(f22703_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl501),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl501])])).
fof(f2275,plain,(
  s(bool,'T_0') != s(bool,'F_0') | p(s(bool,sK8)) | $spl45),
  inference(superposition,[],[f1250,f319])).
fof(f2276,plain,(
  s(bool,'T_0') != s(bool,'F_0') | p(s(bool,sK8)) | $spl45),
  inference(superposition,[],[f1250,f321])).
fof(f22676,plain,(
  $false | ($spl1 | $spl5)),
  inference(subsumption_resolution,[],[f22604,f19106])).
fof(f19106,plain,(
  s(bool,'T_0') != s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7))) | $spl5),
  inference(resolution,[],[f368,f322])).
fof(f22604,plain,(
  s(bool,'T_0') = s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7))) | $spl1),
  inference(superposition,[],[f339,f1254])).
fof(f339,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.option.IS_SOME_1'(s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1)))))) )),
  inference(cnf_transformation,[],[f95])).
fof(f95,plain,(
  ! [X0,X1] : s(bool,'T_0') = s(bool,'const.option.IS_SOME_1'(s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1)))))),
  inference(flattening,[],[f94])).
fof(f94,plain,(
  ! [X0] : ! [X1] : s(bool,'T_0') = s(bool,'const.option.IS_SOME_1'(s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1)))))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X6] : ! [X12] : s(bool,'T_0') = s(bool,'const.option.IS_SOME_1'(s('type.option.option'(X6),'const.option.SOME_1'(s(X6,X12)))))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|1','thm.option.IS_SOME_DEF|split|0')).
fof(f22445,plain,(
  ( ! [X163,X164] : (sP4('F_0',X163) | sP4(X163,'const.option.IS_NONE_1'(s('type.option.option'(X164),'const.option.NONE_0')))) ) | $spl498),
  inference(cnf_transformation,[],[f22445_D])).
fof(f22445_D,plain,(
  ( ! [X163,X164] : (sP4('F_0',X163) | sP4(X163,'const.option.IS_NONE_1'(s('type.option.option'(X164),'const.option.NONE_0')))) ) <=> ~$spl498),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl498])])).
fof(f22492,plain,(
  ( ! [X167,X165,X166] : (p(s(bool,X165)) | sP4('F_0',X166) | sP4(X166,'const.option.IS_NONE_1'(s('type.option.option'(X167),'const.option.NONE_0')))) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f10093,f14064])).
fof(f14064,plain,(
  ( ! [X246,X247,X245] : (sP4(X245,'const.option.IS_NONE_1'(s('type.option.option'(X246),'const.option.NONE_0'))) | sP5(X245,X247)) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f13972,f7584])).
fof(f13972,plain,(
  ( ! [X246,X247,X245] : (sP4('T_0','T_0') | sP4(X245,'const.option.IS_NONE_1'(s('type.option.option'(X246),'const.option.NONE_0'))) | sP5(X245,X247)) ) | $spl3),
  inference(resolution,[],[f8592,f8127])).
fof(f8127,plain,(
  ( ! [X125,X123,X124] : (sP4('T_0',X123) | sP4(X124,X123) | sP5(X124,X125)) ) | $spl3),
  inference(resolution,[],[f3990,f271])).
fof(f3990,plain,(
  ( ! [X80,X81] : (sP4('T_0',X80) | sP4(X81,X80) | p(s(bool,X81))) ) | $spl3),
  inference(resolution,[],[f3605,f258])).
fof(f3605,plain,(
  ( ! [X0] : (p(s(bool,X0)) | sP4('T_0',X0)) ) | $spl3),
  inference(resolution,[],[f3534,f1541])).
fof(f1541,plain,(
  ( ! [X52] : (~sP4(X52,sK8) | p(s(bool,X52))) ) | $spl3),
  inference(resolution,[],[f367,f253])).
fof(f253,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | ~sP4(X0,X1)) )),
  inference(cnf_transformation,[],[f167])).
fof(f3534,plain,(
  ( ! [X203] : (sP4('T_0',X203) | sP4(X203,sK8)) ) | $spl3),
  inference(resolution,[],[f2430,f337])).
fof(f2430,plain,(
  ( ! [X78,X79] : (~p(s(bool,X79)) | sP4(X79,X78) | sP4(X78,sK8)) ) | $spl3),
  inference(resolution,[],[f1543,f255])).
fof(f1543,plain,(
  ( ! [X54] : (~p(s(bool,X54)) | sP4(X54,sK8)) ) | $spl3),
  inference(resolution,[],[f367,f255])).
fof(f8592,plain,(
  ( ! [X59,X58] : (~sP4(X59,'const.option.IS_NONE_1'(s('type.option.option'(X58),'const.option.NONE_0'))) | sP4(X59,'T_0')) ) | $spl3),
  inference(subsumption_resolution,[],[f8563,f361])).
fof(f8563,plain,(
  ( ! [X59,X58] : (~p(s(bool,'T_0')) | sP4(X59,'T_0') | ~sP4(X59,'const.option.IS_NONE_1'(s('type.option.option'(X58),'const.option.NONE_0')))) ) | $spl3),
  inference(superposition,[],[f4520,f359])).
fof(f359,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.option.IS_NONE_1'(s('type.option.option'(X0),'const.option.NONE_0')))) )),
  inference(cnf_transformation,[],[f115])).
fof(f115,plain,(
  ! [X0] : s(bool,'T_0') = s(bool,'const.option.IS_NONE_1'(s('type.option.option'(X0),'const.option.NONE_0')))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X6] : s(bool,'T_0') = s(bool,'const.option.IS_NONE_1'(s('type.option.option'(X6),'const.option.NONE_0')))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|1','thm.option.IS_NONE_DEF|split|1')).
fof(f4520,plain,(
  ( ! [X74,X75] : (~p(s(bool,X75)) | sP4(X74,'T_0') | ~sP4(X74,X75)) ) | $spl3),
  inference(resolution,[],[f4459,f254])).
fof(f4459,plain,(
  ( ! [X0] : (p(s(bool,X0)) | sP4(X0,'T_0')) ) | $spl3),
  inference(resolution,[],[f4394,f1541])).
fof(f4394,plain,(
  ( ! [X205] : (sP4(X205,sK8) | sP4(X205,'T_0')) ) | $spl3),
  inference(resolution,[],[f2431,f337])).
fof(f2431,plain,(
  ( ! [X80,X81] : (~p(s(bool,X81)) | sP4(X80,X81) | sP4(X80,sK8)) ) | $spl3),
  inference(resolution,[],[f1543,f258])).
fof(f10093,plain,(
  ( ! [X92,X93] : (~sP5(X92,X93) | p(s(bool,X93)) | sP4('F_0',X92)) ) | $spl3),
  inference(subsumption_resolution,[],[f10005,f307])).
fof(f10005,plain,(
  ( ! [X92,X93] : (sP4('F_0',X92) | p(s(bool,'F_0')) | p(s(bool,X93)) | ~sP5(X92,X93)) ) | $spl3),
  inference(resolution,[],[f9969,f269])).
fof(f9969,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | sP4('F_0',X1)) ) | $spl3),
  inference(subsumption_resolution,[],[f9936,f361])).
fof(f9936,plain,(
  ( ! [X1] : (sP4('F_0',X1) | ~p(s(bool,X1)) | ~p(s(bool,'T_0'))) ) | $spl3),
  inference(resolution,[],[f9866,f254])).
fof(f9866,plain,(
  ( ! [X222] : (sP4(X222,'T_0') | sP4('F_0',X222)) ) | $spl3),
  inference(resolution,[],[f4523,f332])).
fof(f332,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f91])).
fof(f91,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|1',aHL_FALSITY)).
fof(f4523,plain,(
  ( ! [X80,X81] : (sP4(X80,'T_0') | sP4(X81,X80) | p(s(bool,X81))) ) | $spl3),
  inference(resolution,[],[f4459,f258])).
fof(f10091,plain,(
  ( ! [X20] : (sP4('F_0',X20)) ) | $spl236),
  inference(cnf_transformation,[],[f10091_D])).
fof(f10091_D,plain,(
  ( ! [X20] : (sP4('F_0',X20)) ) <=> ~$spl236),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl236])])).
fof(f18800,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.option.IS_NONE_1'(s('type.option.option'(X1),'const.option.NONE_0'))) | p(s(bool,X0))) ) | $spl408),
  inference(cnf_transformation,[],[f18800_D])).
fof(f18800_D,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.option.IS_NONE_1'(s('type.option.option'(X1),'const.option.NONE_0'))) | p(s(bool,X0))) ) <=> ~$spl408),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl408])])).
fof(f22491,plain,(
  ( ! [X163,X164,X162] : (p(s(bool,X162)) | sP4('F_0',X163) | sP4(X162,'const.option.IS_NONE_1'(s('type.option.option'(X164),'const.option.NONE_0')))) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f10093,f14063])).
fof(f14063,plain,(
  ( ! [X243,X244,X242] : (sP4(X242,'const.option.IS_NONE_1'(s('type.option.option'(X243),'const.option.NONE_0'))) | sP5(X244,X242)) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f13971,f7584])).
fof(f13971,plain,(
  ( ! [X243,X244,X242] : (sP4('T_0','T_0') | sP4(X242,'const.option.IS_NONE_1'(s('type.option.option'(X243),'const.option.NONE_0'))) | sP5(X244,X242)) ) | $spl3),
  inference(resolution,[],[f8592,f8126])).
fof(f8126,plain,(
  ( ! [X121,X122,X120] : (sP4('T_0',X120) | sP4(X121,X120) | sP5(X122,X121)) ) | $spl3),
  inference(resolution,[],[f3990,f270])).
fof(f22444,plain,(
  ( ! [X158,X157] : (sP4('F_0',X157) | sP4('const.option.IS_NONE_1'(s('type.option.option'(X158),'const.option.NONE_0')),X157)) ) | $spl496),
  inference(cnf_transformation,[],[f22444_D])).
fof(f22444_D,plain,(
  ( ! [X158,X157] : (sP4('F_0',X157) | sP4('const.option.IS_NONE_1'(s('type.option.option'(X158),'const.option.NONE_0')),X157)) ) <=> ~$spl496),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl496])])).
fof(f22490,plain,(
  ( ! [X161,X159,X160] : (p(s(bool,X159)) | sP4('F_0',X160) | sP4('const.option.IS_NONE_1'(s('type.option.option'(X161),'const.option.NONE_0')),X160)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f10093,f14056])).
fof(f14056,plain,(
  ( ! [X218,X216,X217] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(X216),'const.option.NONE_0')),X217) | sP5(X217,X218)) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f13963,f7584])).
fof(f13963,plain,(
  ( ! [X218,X216,X217] : (sP4('T_0','T_0') | sP4('const.option.IS_NONE_1'(s('type.option.option'(X216),'const.option.NONE_0')),X217) | sP5(X217,X218)) ) | $spl3),
  inference(resolution,[],[f8592,f7456])).
fof(f7456,plain,(
  ( ! [X125,X123,X124] : (sP4('T_0',X123) | sP4(X123,X124) | sP5(X124,X125)) ) | $spl3),
  inference(resolution,[],[f3989,f271])).
fof(f3989,plain,(
  ( ! [X78,X79] : (sP4('T_0',X78) | sP4(X78,X79) | p(s(bool,X79))) ) | $spl3),
  inference(resolution,[],[f3605,f255])).
fof(f18534,plain,(
  ( ! [X0,X1] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(X0),'const.option.NONE_0')),X1) | p(s(bool,X1))) ) | $spl406),
  inference(cnf_transformation,[],[f18534_D])).
fof(f18534_D,plain,(
  ( ! [X0,X1] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(X0),'const.option.NONE_0')),X1) | p(s(bool,X1))) ) <=> ~$spl406),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl406])])).
fof(f22489,plain,(
  ( ! [X158,X156,X157] : (p(s(bool,X156)) | sP4('F_0',X157) | sP4('const.option.IS_NONE_1'(s('type.option.option'(X158),'const.option.NONE_0')),X156)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f10093,f14055])).
fof(f14055,plain,(
  ( ! [X215,X213,X214] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(X213),'const.option.NONE_0')),X214) | sP5(X215,X214)) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f13962,f7584])).
fof(f13962,plain,(
  ( ! [X215,X213,X214] : (sP4('T_0','T_0') | sP4('const.option.IS_NONE_1'(s('type.option.option'(X213),'const.option.NONE_0')),X214) | sP5(X215,X214)) ) | $spl3),
  inference(resolution,[],[f8592,f7455])).
fof(f7455,plain,(
  ( ! [X121,X122,X120] : (sP4('T_0',X120) | sP4(X120,X121) | sP5(X122,X121)) ) | $spl3),
  inference(resolution,[],[f3989,f270])).
fof(f22443,plain,(
  ( ! [X149,X147,X150,X148] : (sP4('F_0',X147) | sP4(X147,X148) | sP1(X148,X149,X150)) ) | $spl494),
  inference(cnf_transformation,[],[f22443_D])).
fof(f22443_D,plain,(
  ( ! [X149,X147,X150,X148] : (sP4('F_0',X147) | sP4(X147,X148) | sP1(X148,X149,X150)) ) <=> ~$spl494),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl494])])).
fof(f22488,plain,(
  ( ! [X154,X152,X151,X155,X153] : (p(s(bool,X151)) | sP4('F_0',X152) | sP1(X153,X154,X155) | sP4(X152,X153)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f10093,f13362])).
fof(f13362,plain,(
  ( ! [X600,X602,X598,X601,X599] : (sP1(X598,X599,X600) | sP4(X601,X598) | sP5(X601,X602)) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f13252,f7584])).
fof(f13252,plain,(
  ( ! [X600,X602,X598,X601,X599] : (sP4('T_0','T_0') | sP1(X598,X599,X600) | sP4(X601,X598) | sP5(X601,X602)) ) | $spl3),
  inference(resolution,[],[f8480,f8127])).
fof(f8480,plain,(
  ( ! [X74,X72,X75,X73] : (~sP4(X72,X73) | sP4(X72,'T_0') | sP1(X73,X74,X75)) ) | $spl3),
  inference(resolution,[],[f4520,f238])).
fof(f238,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f159])).
fof(f159,plain,(
  ! [X0,X1,X2] : ((~sP1(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1,X2)))),
  inference(rectify,[],[f158])).
fof(f158,plain,(
  ! [X1,X0,X2] : ((~sP1(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0,X2)))),
  inference(flattening,[],[f157])).
fof(f157,plain,(
  ! [X1,X0,X2] : ((~sP1(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0,X2)))),
  inference(nnf_transformation,[],[f138])).
fof(f138,plain,(
  ! [X1,X0,X2] : (sP1(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f18428,plain,(
  ( ! [X2,X0,X3,X1] : (sP1(X0,X1,X2) | p(s(bool,X3)) | sP4(X3,X0)) ) | $spl404),
  inference(cnf_transformation,[],[f18428_D])).
fof(f18428_D,plain,(
  ( ! [X2,X0,X3,X1] : (sP1(X0,X1,X2) | p(s(bool,X3)) | sP4(X3,X0)) ) <=> ~$spl404),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl404])])).
fof(f22487,plain,(
  ( ! [X146,X149,X147,X150,X148] : (p(s(bool,X146)) | sP4('F_0',X147) | sP1(X148,X149,X150) | sP4(X146,X148)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f10093,f13361])).
fof(f13361,plain,(
  ( ! [X596,X594,X597,X593,X595] : (sP1(X593,X594,X595) | sP4(X596,X593) | sP5(X597,X596)) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f13251,f7584])).
fof(f13251,plain,(
  ( ! [X596,X594,X597,X593,X595] : (sP4('T_0','T_0') | sP1(X593,X594,X595) | sP4(X596,X593) | sP5(X597,X596)) ) | $spl3),
  inference(resolution,[],[f8480,f8126])).
fof(f22442,plain,(
  ( ! [X140,X138,X139,X137] : (sP4('F_0',X137) | sP1(X138,X139,X140) | sP4(X138,X137)) ) | $spl492),
  inference(cnf_transformation,[],[f22442_D])).
fof(f22442_D,plain,(
  ( ! [X140,X138,X139,X137] : (sP4('F_0',X137) | sP1(X138,X139,X140) | sP4(X138,X137)) ) <=> ~$spl492),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl492])])).
fof(f22486,plain,(
  ( ! [X144,X142,X145,X143,X141] : (p(s(bool,X141)) | sP4('F_0',X142) | sP1(X143,X144,X145) | sP4(X143,X142)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f10093,f13354])).
fof(f13354,plain,(
  ( ! [X553,X555,X551,X552,X554] : (sP1(X551,X552,X553) | sP4(X551,X554) | sP5(X554,X555)) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f13243,f7584])).
fof(f13243,plain,(
  ( ! [X553,X555,X551,X552,X554] : (sP4('T_0','T_0') | sP1(X551,X552,X553) | sP4(X551,X554) | sP5(X554,X555)) ) | $spl3),
  inference(resolution,[],[f8480,f7456])).
fof(f18321,plain,(
  ( ! [X2,X0,X3,X1] : (sP1(X0,X1,X2) | p(s(bool,X3)) | sP4(X0,X3)) ) | $spl402),
  inference(cnf_transformation,[],[f18321_D])).
fof(f18321_D,plain,(
  ( ! [X2,X0,X3,X1] : (sP1(X0,X1,X2) | p(s(bool,X3)) | sP4(X0,X3)) ) <=> ~$spl402),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl402])])).
fof(f22485,plain,(
  ( ! [X140,X138,X136,X139,X137] : (p(s(bool,X136)) | sP4('F_0',X137) | sP1(X138,X139,X140) | sP4(X138,X136)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f10093,f13353])).
fof(f13353,plain,(
  ( ! [X546,X549,X547,X548,X550] : (sP1(X546,X547,X548) | sP4(X546,X549) | sP5(X550,X549)) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f13242,f7584])).
fof(f13242,plain,(
  ( ! [X546,X549,X547,X548,X550] : (sP4('T_0','T_0') | sP1(X546,X547,X548) | sP4(X546,X549) | sP5(X550,X549)) ) | $spl3),
  inference(resolution,[],[f8480,f7455])).
fof(f22441,plain,(
  ( ! [X121,X122,X120,X119] : (sP4('F_0',X119) | sP4(X119,X121) | sP1(X120,X121,X122)) ) | $spl490),
  inference(cnf_transformation,[],[f22441_D])).
fof(f22441_D,plain,(
  ( ! [X121,X122,X120,X119] : (sP4('F_0',X119) | sP4(X119,X121) | sP1(X120,X121,X122)) ) <=> ~$spl490),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl490])])).
fof(f22482,plain,(
  ( ! [X127,X125,X123,X126,X124] : (p(s(bool,X123)) | sP4('F_0',X124) | sP1(X125,X126,X127) | sP4(X124,X126)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f10093,f12842])).
fof(f12842,plain,(
  ( ! [X600,X602,X598,X601,X599] : (sP1(X598,X599,X600) | sP4(X601,X599) | sP5(X601,X602)) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f12733,f7584])).
fof(f12733,plain,(
  ( ! [X600,X602,X598,X601,X599] : (sP4('T_0','T_0') | sP1(X598,X599,X600) | sP4(X601,X599) | sP5(X601,X602)) ) | $spl3),
  inference(resolution,[],[f8479,f8127])).
fof(f8479,plain,(
  ( ! [X70,X68,X71,X69] : (~sP4(X68,X69) | sP4(X68,'T_0') | sP1(X70,X69,X71)) ) | $spl3),
  inference(resolution,[],[f4520,f237])).
fof(f237,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f159])).
fof(f18214,plain,(
  ( ! [X2,X0,X3,X1] : (sP1(X0,X1,X2) | p(s(bool,X3)) | sP4(X3,X1)) ) | $spl400),
  inference(cnf_transformation,[],[f18214_D])).
fof(f18214_D,plain,(
  ( ! [X2,X0,X3,X1] : (sP1(X0,X1,X2) | p(s(bool,X3)) | sP4(X3,X1)) ) <=> ~$spl400),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl400])])).
fof(f22481,plain,(
  ( ! [X121,X118,X122,X120,X119] : (p(s(bool,X118)) | sP4('F_0',X119) | sP1(X120,X121,X122) | sP4(X118,X121)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f10093,f12841])).
fof(f12841,plain,(
  ( ! [X596,X594,X597,X593,X595] : (sP1(X593,X594,X595) | sP4(X596,X594) | sP5(X597,X596)) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f12732,f7584])).
fof(f12732,plain,(
  ( ! [X596,X594,X597,X593,X595] : (sP4('T_0','T_0') | sP1(X593,X594,X595) | sP4(X596,X594) | sP5(X597,X596)) ) | $spl3),
  inference(resolution,[],[f8479,f8126])).
fof(f22440,plain,(
  ( ! [X111,X109,X112,X110] : (sP4('F_0',X109) | sP1(X110,X111,X112) | sP4(X111,X109)) ) | $spl488),
  inference(cnf_transformation,[],[f22440_D])).
fof(f22440_D,plain,(
  ( ! [X111,X109,X112,X110] : (sP4('F_0',X109) | sP1(X110,X111,X112) | sP4(X111,X109)) ) <=> ~$spl488),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl488])])).
fof(f22480,plain,(
  ( ! [X116,X114,X117,X115,X113] : (p(s(bool,X113)) | sP4('F_0',X114) | sP1(X115,X116,X117) | sP4(X116,X114)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f10093,f12834])).
fof(f12834,plain,(
  ( ! [X553,X555,X551,X552,X554] : (sP1(X551,X552,X553) | sP4(X552,X554) | sP5(X554,X555)) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f12724,f7584])).
fof(f12724,plain,(
  ( ! [X553,X555,X551,X552,X554] : (sP4('T_0','T_0') | sP1(X551,X552,X553) | sP4(X552,X554) | sP5(X554,X555)) ) | $spl3),
  inference(resolution,[],[f8479,f7456])).
fof(f18109,plain,(
  ( ! [X2,X0,X3,X1] : (sP1(X0,X1,X2) | p(s(bool,X3)) | sP4(X1,X3)) ) | $spl398),
  inference(cnf_transformation,[],[f18109_D])).
fof(f18109_D,plain,(
  ( ! [X2,X0,X3,X1] : (sP1(X0,X1,X2) | p(s(bool,X3)) | sP4(X1,X3)) ) <=> ~$spl398),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl398])])).
fof(f22479,plain,(
  ( ! [X111,X109,X112,X110,X108] : (p(s(bool,X108)) | sP4('F_0',X109) | sP1(X110,X111,X112) | sP4(X111,X108)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f10093,f12833])).
fof(f12833,plain,(
  ( ! [X546,X549,X547,X548,X550] : (sP1(X546,X547,X548) | sP4(X547,X549) | sP5(X550,X549)) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f12723,f7584])).
fof(f12723,plain,(
  ( ! [X546,X549,X547,X548,X550] : (sP4('T_0','T_0') | sP1(X546,X547,X548) | sP4(X547,X549) | sP5(X550,X549)) ) | $spl3),
  inference(resolution,[],[f8479,f7455])).
fof(f22439,plain,(
  ( ! [X94,X92,X93,X91] : (sP4('F_0',X91) | sP4(X91,X92) | sP2(X92,X93,X94)) ) | $spl486),
  inference(cnf_transformation,[],[f22439_D])).
fof(f22439_D,plain,(
  ( ! [X94,X92,X93,X91] : (sP4('F_0',X91) | sP4(X91,X92) | sP2(X92,X93,X94)) ) <=> ~$spl486),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl486])])).
fof(f22476,plain,(
  ( ! [X99,X97,X95,X98,X96] : (p(s(bool,X95)) | sP4('F_0',X96) | sP2(X97,X98,X99) | sP4(X96,X97)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f10093,f12387])).
fof(f12387,plain,(
  ( ! [X600,X602,X598,X601,X599] : (sP2(X598,X599,X600) | sP4(X601,X598) | sP5(X601,X602)) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f12277,f7584])).
fof(f12277,plain,(
  ( ! [X600,X602,X598,X601,X599] : (sP4('T_0','T_0') | sP2(X598,X599,X600) | sP4(X601,X598) | sP5(X601,X602)) ) | $spl3),
  inference(resolution,[],[f8477,f8127])).
fof(f8477,plain,(
  ( ! [X61,X62,X60,X63] : (~sP4(X60,X61) | sP4(X60,'T_0') | sP2(X61,X62,X63)) ) | $spl3),
  inference(resolution,[],[f4520,f234])).
fof(f234,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f156])).
fof(f156,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(rectify,[],[f155])).
fof(f155,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(flattening,[],[f154])).
fof(f154,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(nnf_transformation,[],[f139])).
fof(f139,plain,(
  ! [X2,X0,X1] : (sP2(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f18003,plain,(
  ( ! [X2,X0,X3,X1] : (sP2(X0,X1,X2) | p(s(bool,X3)) | sP4(X3,X0)) ) | $spl396),
  inference(cnf_transformation,[],[f18003_D])).
fof(f18003_D,plain,(
  ( ! [X2,X0,X3,X1] : (sP2(X0,X1,X2) | p(s(bool,X3)) | sP4(X3,X0)) ) <=> ~$spl396),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl396])])).
fof(f22475,plain,(
  ( ! [X94,X92,X90,X93,X91] : (p(s(bool,X90)) | sP4('F_0',X91) | sP2(X92,X93,X94) | sP4(X90,X92)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f10093,f12386])).
fof(f12386,plain,(
  ( ! [X596,X594,X597,X593,X595] : (sP2(X593,X594,X595) | sP4(X596,X593) | sP5(X597,X596)) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f12276,f7584])).
fof(f12276,plain,(
  ( ! [X596,X594,X597,X593,X595] : (sP4('T_0','T_0') | sP2(X593,X594,X595) | sP4(X596,X593) | sP5(X597,X596)) ) | $spl3),
  inference(resolution,[],[f8477,f8126])).
fof(f22438,plain,(
  ( ! [X83,X81,X84,X82] : (sP4('F_0',X81) | sP2(X82,X83,X84) | sP4(X82,X81)) ) | $spl484),
  inference(cnf_transformation,[],[f22438_D])).
fof(f22438_D,plain,(
  ( ! [X83,X81,X84,X82] : (sP4('F_0',X81) | sP2(X82,X83,X84) | sP4(X82,X81)) ) <=> ~$spl484),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl484])])).
fof(f22474,plain,(
  ( ! [X88,X87,X85,X89,X86] : (p(s(bool,X85)) | sP4('F_0',X86) | sP2(X87,X88,X89) | sP4(X87,X86)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f10093,f12379])).
fof(f12379,plain,(
  ( ! [X553,X555,X551,X552,X554] : (sP2(X551,X552,X553) | sP4(X551,X554) | sP5(X554,X555)) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f12268,f7584])).
fof(f12268,plain,(
  ( ! [X553,X555,X551,X552,X554] : (sP4('T_0','T_0') | sP2(X551,X552,X553) | sP4(X551,X554) | sP5(X554,X555)) ) | $spl3),
  inference(resolution,[],[f8477,f7456])).
fof(f17895,plain,(
  ( ! [X2,X0,X3,X1] : (sP2(X0,X1,X2) | p(s(bool,X3)) | sP4(X0,X3)) ) | $spl394),
  inference(cnf_transformation,[],[f17895_D])).
fof(f17895_D,plain,(
  ( ! [X2,X0,X3,X1] : (sP2(X0,X1,X2) | p(s(bool,X3)) | sP4(X0,X3)) ) <=> ~$spl394),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl394])])).
fof(f22473,plain,(
  ( ! [X80,X83,X81,X84,X82] : (p(s(bool,X80)) | sP4('F_0',X81) | sP2(X82,X83,X84) | sP4(X82,X80)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f10093,f12378])).
fof(f12378,plain,(
  ( ! [X546,X549,X547,X548,X550] : (sP2(X546,X547,X548) | sP4(X546,X549) | sP5(X550,X549)) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f12267,f7584])).
fof(f12267,plain,(
  ( ! [X546,X549,X547,X548,X550] : (sP4('T_0','T_0') | sP2(X546,X547,X548) | sP4(X546,X549) | sP5(X550,X549)) ) | $spl3),
  inference(resolution,[],[f8477,f7455])).
fof(f22437,plain,(
  ( ! [X66,X64,X65,X63] : (sP4('F_0',X63) | sP4(X63,X65) | sP2(X64,X65,X66)) ) | $spl482),
  inference(cnf_transformation,[],[f22437_D])).
fof(f22437_D,plain,(
  ( ! [X66,X64,X65,X63] : (sP4('F_0',X63) | sP4(X63,X65) | sP2(X64,X65,X66)) ) <=> ~$spl482),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl482])])).
fof(f22470,plain,(
  ( ! [X70,X68,X71,X69,X67] : (p(s(bool,X67)) | sP4('F_0',X68) | sP2(X69,X70,X71) | sP4(X68,X70)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f10093,f11816])).
fof(f11816,plain,(
  ( ! [X600,X602,X598,X601,X599] : (sP2(X598,X599,X600) | sP4(X601,X599) | sP5(X601,X602)) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f11707,f7584])).
fof(f11707,plain,(
  ( ! [X600,X602,X598,X601,X599] : (sP4('T_0','T_0') | sP2(X598,X599,X600) | sP4(X601,X599) | sP5(X601,X602)) ) | $spl3),
  inference(resolution,[],[f8476,f8127])).
fof(f8476,plain,(
  ( ! [X59,X57,X58,X56] : (~sP4(X56,X57) | sP4(X56,'T_0') | sP2(X58,X57,X59)) ) | $spl3),
  inference(resolution,[],[f4520,f233])).
fof(f233,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f156])).
fof(f17790,plain,(
  ( ! [X2,X0,X3,X1] : (sP2(X0,X1,X2) | p(s(bool,X3)) | sP4(X3,X1)) ) | $spl392),
  inference(cnf_transformation,[],[f17790_D])).
fof(f17790_D,plain,(
  ( ! [X2,X0,X3,X1] : (sP2(X0,X1,X2) | p(s(bool,X3)) | sP4(X3,X1)) ) <=> ~$spl392),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl392])])).
fof(f22469,plain,(
  ( ! [X66,X64,X62,X65,X63] : (p(s(bool,X62)) | sP4('F_0',X63) | sP2(X64,X65,X66) | sP4(X62,X65)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f10093,f11815])).
fof(f11815,plain,(
  ( ! [X596,X594,X597,X593,X595] : (sP2(X593,X594,X595) | sP4(X596,X594) | sP5(X597,X596)) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f11706,f7584])).
fof(f11706,plain,(
  ( ! [X596,X594,X597,X593,X595] : (sP4('T_0','T_0') | sP2(X593,X594,X595) | sP4(X596,X594) | sP5(X597,X596)) ) | $spl3),
  inference(resolution,[],[f8476,f8126])).
fof(f22436,plain,(
  ( ! [X54,X56,X55,X53] : (sP4('F_0',X53) | sP2(X54,X55,X56) | sP4(X55,X53)) ) | $spl480),
  inference(cnf_transformation,[],[f22436_D])).
fof(f22436_D,plain,(
  ( ! [X54,X56,X55,X53] : (sP4('F_0',X53) | sP2(X54,X55,X56) | sP4(X55,X53)) ) <=> ~$spl480),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl480])])).
fof(f22468,plain,(
  ( ! [X61,X59,X57,X60,X58] : (p(s(bool,X57)) | sP4('F_0',X58) | sP2(X59,X60,X61) | sP4(X60,X58)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f10093,f11808])).
fof(f11808,plain,(
  ( ! [X553,X555,X551,X552,X554] : (sP2(X551,X552,X553) | sP4(X552,X554) | sP5(X554,X555)) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f11698,f7584])).
fof(f11698,plain,(
  ( ! [X553,X555,X551,X552,X554] : (sP4('T_0','T_0') | sP2(X551,X552,X553) | sP4(X552,X554) | sP5(X554,X555)) ) | $spl3),
  inference(resolution,[],[f8476,f7456])).
fof(f6135,plain,(
  ( ! [X23,X21,X22,X20] : (sP4(X20,X21) | p(s(bool,X21)) | sP2(X22,X20,X23)) ) | $spl158),
  inference(cnf_transformation,[],[f6135_D])).
fof(f6135_D,plain,(
  ( ! [X23,X21,X22,X20] : (sP4(X20,X21) | p(s(bool,X21)) | sP2(X22,X20,X23)) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f22467,plain,(
  ( ! [X54,X52,X56,X55,X53] : (p(s(bool,X52)) | sP4('F_0',X53) | sP2(X54,X55,X56) | sP4(X55,X52)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f10093,f11807])).
fof(f11807,plain,(
  ( ! [X546,X549,X547,X548,X550] : (sP2(X546,X547,X548) | sP4(X547,X549) | sP5(X550,X549)) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f11697,f7584])).
fof(f11697,plain,(
  ( ! [X546,X549,X547,X548,X550] : (sP4('T_0','T_0') | sP2(X546,X547,X548) | sP4(X547,X549) | sP5(X550,X549)) ) | $spl3),
  inference(resolution,[],[f8476,f7455])).
fof(f22429,plain,(
  ( ! [X167,X165,X166] : (p(s(bool,X165)) | sP4('F_0',X166) | sP4(X165,'const.option.IS_NONE_1'(s('type.option.option'(X167),'const.option.NONE_0')))) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f10092,f14064])).
fof(f10092,plain,(
  ( ! [X90,X91] : (~sP5(X91,X90) | p(s(bool,X91)) | sP4('F_0',X90)) ) | $spl3),
  inference(subsumption_resolution,[],[f10004,f307])).
fof(f10004,plain,(
  ( ! [X90,X91] : (sP4('F_0',X90) | p(s(bool,'F_0')) | p(s(bool,X91)) | ~sP5(X91,X90)) ) | $spl3),
  inference(resolution,[],[f9969,f269])).
fof(f22428,plain,(
  ( ! [X163,X164,X162] : (p(s(bool,X162)) | sP4('F_0',X163) | sP4(X163,'const.option.IS_NONE_1'(s('type.option.option'(X164),'const.option.NONE_0')))) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f10092,f14063])).
fof(f22427,plain,(
  ( ! [X161,X159,X160] : (p(s(bool,X159)) | sP4('F_0',X160) | sP4('const.option.IS_NONE_1'(s('type.option.option'(X161),'const.option.NONE_0')),X159)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f10092,f14056])).
fof(f22426,plain,(
  ( ! [X158,X156,X157] : (p(s(bool,X156)) | sP4('F_0',X157) | sP4('const.option.IS_NONE_1'(s('type.option.option'(X158),'const.option.NONE_0')),X157)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f10092,f14055])).
fof(f22425,plain,(
  ( ! [X154,X152,X151,X155,X153] : (p(s(bool,X151)) | sP4('F_0',X152) | sP1(X153,X154,X155) | sP4(X151,X153)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f10092,f13362])).
fof(f22424,plain,(
  ( ! [X146,X149,X147,X150,X148] : (p(s(bool,X146)) | sP4('F_0',X147) | sP1(X148,X149,X150) | sP4(X147,X148)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f10092,f13361])).
fof(f22423,plain,(
  ( ! [X144,X142,X145,X143,X141] : (p(s(bool,X141)) | sP4('F_0',X142) | sP1(X143,X144,X145) | sP4(X143,X141)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f10092,f13354])).
fof(f22422,plain,(
  ( ! [X140,X138,X136,X139,X137] : (p(s(bool,X136)) | sP4('F_0',X137) | sP1(X138,X139,X140) | sP4(X138,X137)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f10092,f13353])).
fof(f22419,plain,(
  ( ! [X127,X125,X123,X126,X124] : (p(s(bool,X123)) | sP4('F_0',X124) | sP1(X125,X126,X127) | sP4(X123,X126)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f10092,f12842])).
fof(f22418,plain,(
  ( ! [X121,X118,X122,X120,X119] : (p(s(bool,X118)) | sP4('F_0',X119) | sP1(X120,X121,X122) | sP4(X119,X121)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f10092,f12841])).
fof(f22417,plain,(
  ( ! [X116,X114,X117,X115,X113] : (p(s(bool,X113)) | sP4('F_0',X114) | sP1(X115,X116,X117) | sP4(X116,X113)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f10092,f12834])).
fof(f22416,plain,(
  ( ! [X111,X109,X112,X110,X108] : (p(s(bool,X108)) | sP4('F_0',X109) | sP1(X110,X111,X112) | sP4(X111,X109)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f10092,f12833])).
fof(f22413,plain,(
  ( ! [X99,X97,X95,X98,X96] : (p(s(bool,X95)) | sP4('F_0',X96) | sP2(X97,X98,X99) | sP4(X95,X97)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f10092,f12387])).
fof(f22412,plain,(
  ( ! [X94,X92,X90,X93,X91] : (p(s(bool,X90)) | sP4('F_0',X91) | sP2(X92,X93,X94) | sP4(X91,X92)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f10092,f12386])).
fof(f22411,plain,(
  ( ! [X88,X87,X85,X89,X86] : (p(s(bool,X85)) | sP4('F_0',X86) | sP2(X87,X88,X89) | sP4(X87,X85)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f10092,f12379])).
fof(f22410,plain,(
  ( ! [X80,X83,X81,X84,X82] : (p(s(bool,X80)) | sP4('F_0',X81) | sP2(X82,X83,X84) | sP4(X82,X81)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f10092,f12378])).
fof(f22407,plain,(
  ( ! [X70,X68,X71,X69,X67] : (p(s(bool,X67)) | sP4('F_0',X68) | sP2(X69,X70,X71) | sP4(X67,X70)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f10092,f11816])).
fof(f22406,plain,(
  ( ! [X66,X64,X62,X65,X63] : (p(s(bool,X62)) | sP4('F_0',X63) | sP2(X64,X65,X66) | sP4(X63,X65)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f10092,f11815])).
fof(f22405,plain,(
  ( ! [X61,X59,X57,X60,X58] : (p(s(bool,X57)) | sP4('F_0',X58) | sP2(X59,X60,X61) | sP4(X60,X57)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f10092,f11808])).
fof(f22404,plain,(
  ( ! [X54,X52,X56,X55,X53] : (p(s(bool,X52)) | sP4('F_0',X53) | sP2(X54,X55,X56) | sP4(X55,X53)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f10092,f11807])).
fof(f20556,plain,(
  ( ! [X225] : (sP4('F_0',X225) | sP4(X225,'F_0')) ) | $spl428),
  inference(cnf_transformation,[],[f20556_D])).
fof(f20556_D,plain,(
  ( ! [X225] : (sP4('F_0',X225) | sP4(X225,'F_0')) ) <=> ~$spl428),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl428])])).
fof(f22274,plain,(
  ( ! [X231,X233,X232] : (sP4(X231,'F_0') | sP4('F_0',X231) | p(s(bool,X232)) | p(s(bool,X233)) | ~sP5(X232,X233)) ) | $spl3),
  inference(resolution,[],[f9999,f269])).
fof(f9999,plain,(
  ( ! [X80,X81] : (~p(s(bool,X81)) | sP4(X80,X81) | sP4('F_0',X80)) ) | $spl3),
  inference(resolution,[],[f9969,f258])).
fof(f22272,plain,(
  ( ! [X227,X225,X226] : (sP4(X225,'F_0') | sP4('F_0',X225) | p(s(bool,X226)) | p(s(bool,X227)) | ~sP5(X226,X227)) ) | $spl3),
  inference(resolution,[],[f9999,f273])).
fof(f22311,plain,(
  ( ! [X0,X3,X1] : (sP4(X0,X1) | s(bool,X1) != s(bool,X3) | ~p(s(bool,X3)) | sP4('F_0',X0)) ) | $spl478),
  inference(cnf_transformation,[],[f22311_D])).
fof(f22311_D,plain,(
  ( ! [X0,X3,X1] : (sP4(X0,X1) | s(bool,X1) != s(bool,X3) | ~p(s(bool,X3)) | sP4('F_0',X0)) ) <=> ~$spl478),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl478])])).
fof(f22249,plain,(
  ( ! [X185,X182,X184,X183] : (sP4(X182,X183) | sP4('F_0',X182) | p(s(bool,X184)) | ~p(s(bool,X185)) | s(bool,X183) != s(bool,X185)) ) | $spl3),
  inference(resolution,[],[f9999,f360])).
fof(f22204,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP4(X35,X36) | sP4('F_0',X35) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X38) | s(bool,X37) != s(bool,X39)) ) | $spl3),
  inference(resolution,[],[f9999,f223])).
fof(f22203,plain,(
  ( ! [X30,X33,X31,X34,X32] : (sP4(X30,X31) | sP4('F_0',X30) | p(s(bool,X32)) | ~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X31) != s(bool,X33)) ) | $spl3),
  inference(resolution,[],[f9999,f222])).
fof(f22201,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP4(X20,X21) | sP4('F_0',X20) | p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X24) | s(bool,X21) != s(bool,X23)) ) | $spl3),
  inference(resolution,[],[f9999,f221])).
fof(f22199,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP4(X10,X11) | sP4('F_0',X10) | ~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,X14)) ) | $spl3),
  inference(resolution,[],[f9999,f219])).
fof(f22197,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP4(X0,X1) | sP4('F_0',X0) | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) ) | $spl3),
  inference(resolution,[],[f9999,f217])).
fof(f22179,plain,(
  ( ! [X231,X233,X232] : (sP4('F_0',X231) | p(s(bool,X232)) | p(s(bool,X233)) | ~sP5(X232,X233)) ) | $spl3),
  inference(duplicate_literal_removal,[],[f22139])).
fof(f22139,plain,(
  ( ! [X231,X233,X232] : (sP4('F_0',X231) | sP4('F_0',X231) | p(s(bool,X232)) | p(s(bool,X233)) | ~sP5(X232,X233)) ) | $spl3),
  inference(resolution,[],[f9998,f269])).
fof(f9998,plain,(
  ( ! [X78,X79] : (~p(s(bool,X79)) | sP4(X79,X78) | sP4('F_0',X78)) ) | $spl3),
  inference(resolution,[],[f9969,f255])).
fof(f22181,plain,(
  ( ! [X227,X225,X226] : (sP4('F_0',X225) | p(s(bool,X226)) | p(s(bool,X227)) | ~sP5(X226,X227)) ) | $spl3),
  inference(duplicate_literal_removal,[],[f22137])).
fof(f22137,plain,(
  ( ! [X227,X225,X226] : (sP4('F_0',X225) | sP4('F_0',X225) | p(s(bool,X226)) | p(s(bool,X227)) | ~sP5(X226,X227)) ) | $spl3),
  inference(resolution,[],[f9998,f273])).
fof(f22185,plain,(
  ( ! [X0,X3,X1] : (sP4(X0,X1) | s(bool,X0) != s(bool,X3) | ~p(s(bool,X3)) | sP4('F_0',X1)) ) | $spl476),
  inference(cnf_transformation,[],[f22185_D])).
fof(f22185_D,plain,(
  ( ! [X0,X3,X1] : (sP4(X0,X1) | s(bool,X0) != s(bool,X3) | ~p(s(bool,X3)) | sP4('F_0',X1)) ) <=> ~$spl476),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl476])])).
fof(f22114,plain,(
  ( ! [X185,X182,X184,X183] : (sP4(X182,X183) | sP4('F_0',X183) | p(s(bool,X184)) | ~p(s(bool,X185)) | s(bool,X182) != s(bool,X185)) ) | $spl3),
  inference(resolution,[],[f9998,f360])).
fof(f22069,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP4(X35,X36) | sP4('F_0',X36) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X35) != s(bool,X38) | s(bool,X37) != s(bool,X39)) ) | $spl3),
  inference(resolution,[],[f9998,f223])).
fof(f22068,plain,(
  ( ! [X30,X33,X31,X34,X32] : (sP4(X30,X31) | sP4('F_0',X31) | p(s(bool,X32)) | ~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X30) != s(bool,X33)) ) | $spl3),
  inference(resolution,[],[f9998,f222])).
fof(f22066,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP4(X20,X21) | sP4('F_0',X21) | p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X24) | s(bool,X20) != s(bool,X23)) ) | $spl3),
  inference(resolution,[],[f9998,f221])).
fof(f22064,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP4(X10,X11) | sP4('F_0',X11) | ~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X10) != s(bool,X13) | s(bool,X12) != s(bool,X14)) ) | $spl3),
  inference(resolution,[],[f9998,f219])).
fof(f22062,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP4(X0,X1) | sP4('F_0',X1) | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X0) != s(bool,X3)) ) | $spl3),
  inference(resolution,[],[f9998,f217])).
fof(f908,plain,(
  sP4('F_0',sK8) | $spl36),
  inference(cnf_transformation,[],[f908_D])).
fof(f908_D,plain,(
  sP4('F_0',sK8) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f22003,plain,(
  ( ! [X600,X598,X597,X599] : (p(s(bool,X597)) | sP4('F_0',sK8) | sP4(X597,X598) | sP1(X598,X599,X600)) ) | $spl3),
  inference(resolution,[],[f9997,f5499])).
fof(f5499,plain,(
  ( ! [X74,X72,X75,X73] : (sP4(X72,X73) | sP4(sK8,X72) | sP1(X73,X74,X75)) ) | $spl3),
  inference(resolution,[],[f2548,f238])).
fof(f2548,plain,(
  ( ! [X80,X81] : (~p(s(bool,X81)) | sP4(X80,X81) | sP4(sK8,X80)) ) | $spl3),
  inference(resolution,[],[f1544,f258])).
fof(f1544,plain,(
  ( ! [X55] : (~p(s(bool,X55)) | sP4(sK8,X55)) ) | $spl3),
  inference(resolution,[],[f367,f258])).
fof(f9997,plain,(
  ( ! [X76,X77] : (~sP4(X76,X77) | p(s(bool,X77)) | sP4('F_0',X76)) ) | $spl3),
  inference(resolution,[],[f9969,f253])).
fof(f22002,plain,(
  ( ! [X596,X594,X593,X595] : (p(s(bool,X593)) | sP4('F_0',sK8) | sP4(X593,X594) | sP1(X595,X594,X596)) ) | $spl3),
  inference(resolution,[],[f9997,f5498])).
fof(f5498,plain,(
  ( ! [X70,X68,X71,X69] : (sP4(X68,X69) | sP4(sK8,X68) | sP1(X70,X69,X71)) ) | $spl3),
  inference(resolution,[],[f2548,f237])).
fof(f22001,plain,(
  ( ! [X592,X590,X589,X591] : (p(s(bool,X589)) | sP4('F_0',sK8) | sP4(X589,X590) | sP2(X590,X591,X592)) ) | $spl3),
  inference(resolution,[],[f9997,f5496])).
fof(f5496,plain,(
  ( ! [X61,X62,X60,X63] : (sP4(X60,X61) | sP4(sK8,X60) | sP2(X61,X62,X63)) ) | $spl3),
  inference(resolution,[],[f2548,f234])).
fof(f22000,plain,(
  ( ! [X588,X586,X585,X587] : (p(s(bool,X585)) | sP4('F_0',sK8) | sP4(X585,X586) | sP2(X587,X586,X588)) ) | $spl3),
  inference(resolution,[],[f9997,f5495])).
fof(f5495,plain,(
  ( ! [X59,X57,X58,X56] : (sP4(X56,X57) | sP4(sK8,X56) | sP2(X58,X57,X59)) ) | $spl3),
  inference(resolution,[],[f2548,f233])).
fof(f21999,plain,(
  ( ! [X584,X582,X581,X583] : (p(s(bool,X581)) | sP4('F_0',sK8) | sP4(X582,X581) | sP1(X582,X583,X584)) ) | $spl3),
  inference(resolution,[],[f9997,f5373])).
fof(f5373,plain,(
  ( ! [X74,X72,X75,X73] : (sP4(X72,X73) | sP4(sK8,X73) | sP1(X72,X74,X75)) ) | $spl3),
  inference(resolution,[],[f2547,f238])).
fof(f2547,plain,(
  ( ! [X78,X79] : (~p(s(bool,X79)) | sP4(X79,X78) | sP4(sK8,X78)) ) | $spl3),
  inference(resolution,[],[f1544,f255])).
fof(f21998,plain,(
  ( ! [X580,X578,X577,X579] : (p(s(bool,X577)) | sP4('F_0',sK8) | sP4(X578,X577) | sP1(X579,X578,X580)) ) | $spl3),
  inference(resolution,[],[f9997,f5372])).
fof(f5372,plain,(
  ( ! [X70,X68,X71,X69] : (sP4(X68,X69) | sP4(sK8,X69) | sP1(X70,X68,X71)) ) | $spl3),
  inference(resolution,[],[f2547,f237])).
fof(f21997,plain,(
  ( ! [X573,X575,X576,X574] : (p(s(bool,X573)) | sP4('F_0',sK8) | sP4(X574,X573) | sP2(X574,X575,X576)) ) | $spl3),
  inference(resolution,[],[f9997,f5370])).
fof(f5370,plain,(
  ( ! [X61,X62,X60,X63] : (sP4(X60,X61) | sP4(sK8,X61) | sP2(X60,X62,X63)) ) | $spl3),
  inference(resolution,[],[f2547,f234])).
fof(f21996,plain,(
  ( ! [X569,X571,X572,X570] : (p(s(bool,X569)) | sP4('F_0',sK8) | sP4(X570,X569) | sP2(X571,X570,X572)) ) | $spl3),
  inference(resolution,[],[f9997,f5369])).
fof(f5369,plain,(
  ( ! [X59,X57,X58,X56] : (sP4(X56,X57) | sP4(sK8,X57) | sP2(X58,X56,X59)) ) | $spl3),
  inference(resolution,[],[f2547,f233])).
fof(f1826,plain,(
  s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6)))) | $spl62),
  inference(cnf_transformation,[],[f1826_D])).
fof(f1826_D,plain,(
  s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6)))) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f21787,plain,(
  ( ! [X0] : (~sP5('const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0) | p(s(bool,X0))) ) | $spl474),
  inference(cnf_transformation,[],[f21787_D])).
fof(f21787_D,plain,(
  ( ! [X0] : (~sP5('const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0) | p(s(bool,X0))) ) <=> ~$spl474),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl474])])).
fof(f21773,plain,(
  ( ! [X0] : (~sP5('const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0) | p(s(bool,X0)) | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6))))) ) | $spl5),
  inference(superposition,[],[f19152,f357])).
fof(f19152,plain,(
  ( ! [X61] : (~sP5('const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)),X61) | p(s(bool,X61))) ) | $spl5),
  inference(subsumption_resolution,[],[f19093,f307])).
fof(f19093,plain,(
  ( ! [X61] : (p(s(bool,'F_0')) | p(s(bool,X61)) | ~sP5('const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)),X61)) ) | $spl5),
  inference(resolution,[],[f368,f269])).
fof(f21452,plain,(
  ( ! [X119] : (sP4('const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)),'const.option.IS_NONE_1'(s('type.option.option'(X119),'const.option.NONE_0')))) ) | $spl470),
  inference(cnf_transformation,[],[f21452_D])).
fof(f21452_D,plain,(
  ( ! [X119] : (sP4('const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)),'const.option.IS_NONE_1'(s('type.option.option'(X119),'const.option.NONE_0')))) ) <=> ~$spl470),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl470])])).
fof(f21770,plain,(
  ( ! [X121,X120] : (p(s(bool,X120)) | sP4('const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)),'const.option.IS_NONE_1'(s('type.option.option'(X121),'const.option.NONE_0')))) ) | ($spl3 | $spl5 | $spl179)),
  inference(resolution,[],[f19152,f14064])).
fof(f21451,plain,(
  ( ! [X115] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(X115),'const.option.NONE_0')),'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) | $spl468),
  inference(cnf_transformation,[],[f21451_D])).
fof(f21451_D,plain,(
  ( ! [X115] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(X115),'const.option.NONE_0')),'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) <=> ~$spl468),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl468])])).
fof(f21768,plain,(
  ( ! [X116,X117] : (p(s(bool,X116)) | sP4('const.option.IS_NONE_1'(s('type.option.option'(X117),'const.option.NONE_0')),'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) | ($spl3 | $spl5 | $spl179)),
  inference(resolution,[],[f19152,f14056])).
fof(f21450,plain,(
  ( ! [X109,X107,X108] : (sP1(X107,X108,X109) | sP4('const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)),X107)) ) | $spl466),
  inference(cnf_transformation,[],[f21450_D])).
fof(f21450_D,plain,(
  ( ! [X109,X107,X108] : (sP1(X107,X108,X109) | sP4('const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)),X107)) ) <=> ~$spl466),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl466])])).
fof(f21766,plain,(
  ( ! [X111,X112,X110,X113] : (p(s(bool,X110)) | sP1(X111,X112,X113) | sP4('const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)),X111)) ) | ($spl3 | $spl5 | $spl179)),
  inference(resolution,[],[f19152,f13362])).
fof(f21449,plain,(
  ( ! [X101,X99,X100] : (sP1(X99,X100,X101) | sP4(X99,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) | $spl464),
  inference(cnf_transformation,[],[f21449_D])).
fof(f21449_D,plain,(
  ( ! [X101,X99,X100] : (sP1(X99,X100,X101) | sP4(X99,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) <=> ~$spl464),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl464])])).
fof(f21764,plain,(
  ( ! [X103,X105,X102,X104] : (p(s(bool,X102)) | sP1(X103,X104,X105) | sP4(X103,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) | ($spl3 | $spl5 | $spl179)),
  inference(resolution,[],[f19152,f13354])).
fof(f21448,plain,(
  ( ! [X87,X85,X86] : (sP1(X85,X86,X87) | sP4('const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)),X86)) ) | $spl462),
  inference(cnf_transformation,[],[f21448_D])).
fof(f21448_D,plain,(
  ( ! [X87,X85,X86] : (sP1(X85,X86,X87) | sP4('const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)),X86)) ) <=> ~$spl462),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl462])])).
fof(f21760,plain,(
  ( ! [X90,X88,X91,X89] : (p(s(bool,X88)) | sP1(X89,X90,X91) | sP4('const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)),X90)) ) | ($spl3 | $spl5 | $spl179)),
  inference(resolution,[],[f19152,f12842])).
fof(f21447,plain,(
  ( ! [X78,X79,X77] : (sP1(X77,X78,X79) | sP4(X78,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) | $spl460),
  inference(cnf_transformation,[],[f21447_D])).
fof(f21447_D,plain,(
  ( ! [X78,X79,X77] : (sP1(X77,X78,X79) | sP4(X78,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) <=> ~$spl460),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl460])])).
fof(f21758,plain,(
  ( ! [X80,X83,X81,X82] : (p(s(bool,X80)) | sP1(X81,X82,X83) | sP4(X82,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) | ($spl3 | $spl5 | $spl179)),
  inference(resolution,[],[f19152,f12834])).
fof(f21446,plain,(
  ( ! [X64,X65,X63] : (sP2(X63,X64,X65) | sP4('const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)),X63)) ) | $spl458),
  inference(cnf_transformation,[],[f21446_D])).
fof(f21446_D,plain,(
  ( ! [X64,X65,X63] : (sP2(X63,X64,X65) | sP4('const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)),X63)) ) <=> ~$spl458),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl458])])).
fof(f21754,plain,(
  ( ! [X68,X66,X69,X67] : (p(s(bool,X66)) | sP2(X67,X68,X69) | sP4('const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)),X67)) ) | ($spl3 | $spl5 | $spl179)),
  inference(resolution,[],[f19152,f12387])).
fof(f21445,plain,(
  ( ! [X57,X56,X55] : (sP2(X55,X56,X57) | sP4(X55,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) | $spl456),
  inference(cnf_transformation,[],[f21445_D])).
fof(f21445_D,plain,(
  ( ! [X57,X56,X55] : (sP2(X55,X56,X57) | sP4(X55,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) <=> ~$spl456),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl456])])).
fof(f21752,plain,(
  ( ! [X61,X59,X60,X58] : (p(s(bool,X58)) | sP2(X59,X60,X61) | sP4(X59,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) | ($spl3 | $spl5 | $spl179)),
  inference(resolution,[],[f19152,f12379])).
fof(f21444,plain,(
  ( ! [X43,X41,X42] : (sP2(X41,X42,X43) | sP4('const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)),X42)) ) | $spl454),
  inference(cnf_transformation,[],[f21444_D])).
fof(f21444_D,plain,(
  ( ! [X43,X41,X42] : (sP2(X41,X42,X43) | sP4('const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)),X42)) ) <=> ~$spl454),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl454])])).
fof(f21748,plain,(
  ( ! [X47,X45,X46,X44] : (p(s(bool,X44)) | sP2(X45,X46,X47) | sP4('const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)),X46)) ) | ($spl3 | $spl5 | $spl179)),
  inference(resolution,[],[f19152,f11816])).
fof(f21443,plain,(
  ( ! [X35,X33,X34] : (sP2(X33,X34,X35) | sP4(X34,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) | $spl452),
  inference(cnf_transformation,[],[f21443_D])).
fof(f21443_D,plain,(
  ( ! [X35,X33,X34] : (sP2(X33,X34,X35) | sP4(X34,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) <=> ~$spl452),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl452])])).
fof(f21746,plain,(
  ( ! [X39,X37,X38,X36] : (p(s(bool,X36)) | sP2(X37,X38,X39) | sP4(X38,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) | ($spl3 | $spl5 | $spl179)),
  inference(resolution,[],[f19152,f11808])).
fof(f21595,plain,(
  ( ! [X436,X434,X437,X435] : (p(s(bool,X434)) | sP4('F_0',sK8) | sP4(X434,X435) | sP1(X435,X436,X437)) ) | $spl3),
  inference(resolution,[],[f9996,f4346])).
fof(f4346,plain,(
  ( ! [X74,X72,X75,X73] : (sP4(X72,X73) | sP4(X72,sK8) | sP1(X73,X74,X75)) ) | $spl3),
  inference(resolution,[],[f2431,f238])).
fof(f9996,plain,(
  ( ! [X74,X75] : (~sP4(X75,X74) | p(s(bool,X75)) | sP4('F_0',X74)) ) | $spl3),
  inference(resolution,[],[f9969,f253])).
fof(f21594,plain,(
  ( ! [X430,X432,X431,X433] : (p(s(bool,X430)) | sP4('F_0',sK8) | sP4(X430,X431) | sP1(X432,X431,X433)) ) | $spl3),
  inference(resolution,[],[f9996,f4345])).
fof(f4345,plain,(
  ( ! [X70,X68,X71,X69] : (sP4(X68,X69) | sP4(X68,sK8) | sP1(X70,X69,X71)) ) | $spl3),
  inference(resolution,[],[f2431,f237])).
fof(f21593,plain,(
  ( ! [X428,X426,X429,X427] : (p(s(bool,X426)) | sP4('F_0',sK8) | sP4(X426,X427) | sP2(X427,X428,X429)) ) | $spl3),
  inference(resolution,[],[f9996,f4343])).
fof(f4343,plain,(
  ( ! [X61,X62,X60,X63] : (sP4(X60,X61) | sP4(X60,sK8) | sP2(X61,X62,X63)) ) | $spl3),
  inference(resolution,[],[f2431,f234])).
fof(f21592,plain,(
  ( ! [X422,X424,X423,X425] : (p(s(bool,X422)) | sP4('F_0',sK8) | sP4(X422,X423) | sP2(X424,X423,X425)) ) | $spl3),
  inference(resolution,[],[f9996,f4342])).
fof(f4342,plain,(
  ( ! [X59,X57,X58,X56] : (sP4(X56,X57) | sP4(X56,sK8) | sP2(X58,X57,X59)) ) | $spl3),
  inference(resolution,[],[f2431,f233])).
fof(f21589,plain,(
  ( ! [X416,X418,X415,X417] : (p(s(bool,X415)) | sP4('F_0',sK8) | sP4(X416,X415) | sP1(X416,X417,X418)) ) | $spl3),
  inference(resolution,[],[f9996,f3487])).
fof(f3487,plain,(
  ( ! [X74,X72,X75,X73] : (sP4(X72,X73) | sP4(X73,sK8) | sP1(X72,X74,X75)) ) | $spl3),
  inference(resolution,[],[f2430,f238])).
fof(f21588,plain,(
  ( ! [X412,X414,X413,X411] : (p(s(bool,X411)) | sP4('F_0',sK8) | sP4(X412,X411) | sP1(X413,X412,X414)) ) | $spl3),
  inference(resolution,[],[f9996,f3486])).
fof(f3486,plain,(
  ( ! [X70,X68,X71,X69] : (sP4(X68,X69) | sP4(X69,sK8) | sP1(X70,X68,X71)) ) | $spl3),
  inference(resolution,[],[f2430,f237])).
fof(f21587,plain,(
  ( ! [X408,X410,X407,X409] : (p(s(bool,X407)) | sP4('F_0',sK8) | sP4(X408,X407) | sP2(X408,X409,X410)) ) | $spl3),
  inference(resolution,[],[f9996,f3484])).
fof(f3484,plain,(
  ( ! [X61,X62,X60,X63] : (sP4(X60,X61) | sP4(X61,sK8) | sP2(X60,X62,X63)) ) | $spl3),
  inference(resolution,[],[f2430,f234])).
fof(f21586,plain,(
  ( ! [X405,X403,X404,X406] : (p(s(bool,X403)) | sP4('F_0',sK8) | sP4(X404,X403) | sP2(X405,X404,X406)) ) | $spl3),
  inference(resolution,[],[f9996,f3483])).
fof(f3483,plain,(
  ( ! [X59,X57,X58,X56] : (sP4(X56,X57) | sP4(X57,sK8) | sP2(X58,X56,X59)) ) | $spl3),
  inference(resolution,[],[f2430,f233])).
fof(f21454,plain,(
  ( ! [X0] : (~sP5(X0,'const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0'))) | p(s(bool,X0))) ) | $spl472),
  inference(cnf_transformation,[],[f21454_D])).
fof(f21454_D,plain,(
  ( ! [X0] : (~sP5(X0,'const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0'))) | p(s(bool,X0))) ) <=> ~$spl472),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl472])])).
fof(f21429,plain,(
  ( ! [X0] : (~sP5(X0,'const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0'))) | p(s(bool,X0)) | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6))))) ) | $spl5),
  inference(superposition,[],[f19151,f357])).
fof(f19151,plain,(
  ( ! [X60] : (~sP5(X60,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7))) | p(s(bool,X60))) ) | $spl5),
  inference(subsumption_resolution,[],[f19092,f307])).
fof(f19092,plain,(
  ( ! [X60] : (p(s(bool,'F_0')) | p(s(bool,X60)) | ~sP5(X60,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) | $spl5),
  inference(resolution,[],[f368,f269])).
fof(f21425,plain,(
  ( ! [X118,X119] : (p(s(bool,X118)) | sP4('const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)),'const.option.IS_NONE_1'(s('type.option.option'(X119),'const.option.NONE_0')))) ) | ($spl3 | $spl5 | $spl179)),
  inference(resolution,[],[f19151,f14063])).
fof(f21423,plain,(
  ( ! [X114,X115] : (p(s(bool,X114)) | sP4('const.option.IS_NONE_1'(s('type.option.option'(X115),'const.option.NONE_0')),'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) | ($spl3 | $spl5 | $spl179)),
  inference(resolution,[],[f19151,f14055])).
fof(f21421,plain,(
  ( ! [X109,X107,X108,X106] : (p(s(bool,X106)) | sP1(X107,X108,X109) | sP4('const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)),X107)) ) | ($spl3 | $spl5 | $spl179)),
  inference(resolution,[],[f19151,f13361])).
fof(f21419,plain,(
  ( ! [X101,X99,X100,X98] : (p(s(bool,X98)) | sP1(X99,X100,X101) | sP4(X99,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) | ($spl3 | $spl5 | $spl179)),
  inference(resolution,[],[f19151,f13353])).
fof(f21415,plain,(
  ( ! [X87,X85,X86,X84] : (p(s(bool,X84)) | sP1(X85,X86,X87) | sP4('const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)),X86)) ) | ($spl3 | $spl5 | $spl179)),
  inference(resolution,[],[f19151,f12841])).
fof(f21413,plain,(
  ( ! [X78,X76,X79,X77] : (p(s(bool,X76)) | sP1(X77,X78,X79) | sP4(X78,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) | ($spl3 | $spl5 | $spl179)),
  inference(resolution,[],[f19151,f12833])).
fof(f21409,plain,(
  ( ! [X64,X62,X65,X63] : (p(s(bool,X62)) | sP2(X63,X64,X65) | sP4('const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)),X63)) ) | ($spl3 | $spl5 | $spl179)),
  inference(resolution,[],[f19151,f12386])).
fof(f21407,plain,(
  ( ! [X57,X54,X56,X55] : (p(s(bool,X54)) | sP2(X55,X56,X57) | sP4(X55,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) | ($spl3 | $spl5 | $spl179)),
  inference(resolution,[],[f19151,f12378])).
fof(f21403,plain,(
  ( ! [X43,X41,X42,X40] : (p(s(bool,X40)) | sP2(X41,X42,X43) | sP4('const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)),X42)) ) | ($spl3 | $spl5 | $spl179)),
  inference(resolution,[],[f19151,f11815])).
fof(f21401,plain,(
  ( ! [X35,X33,X34,X32] : (p(s(bool,X32)) | sP2(X33,X34,X35) | sP4(X34,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) | ($spl3 | $spl5 | $spl179)),
  inference(resolution,[],[f19151,f11807])).
fof(f21196,plain,(
  ( ! [X163,X164] : (sP4(X163,'F_0') | sP4(X163,'const.option.IS_NONE_1'(s('type.option.option'(X164),'const.option.NONE_0')))) ) | $spl450),
  inference(cnf_transformation,[],[f21196_D])).
fof(f21196_D,plain,(
  ( ! [X163,X164] : (sP4(X163,'F_0') | sP4(X163,'const.option.IS_NONE_1'(s('type.option.option'(X164),'const.option.NONE_0')))) ) <=> ~$spl450),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl450])])).
fof(f21243,plain,(
  ( ! [X167,X165,X166] : (p(s(bool,X165)) | sP4(X166,'F_0') | sP4(X166,'const.option.IS_NONE_1'(s('type.option.option'(X167),'const.option.NONE_0')))) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f7782,f14064])).
fof(f7782,plain,(
  ( ! [X92,X93] : (~sP5(X92,X93) | p(s(bool,X93)) | sP4(X92,'F_0')) ) | $spl3),
  inference(subsumption_resolution,[],[f7697,f307])).
fof(f7697,plain,(
  ( ! [X92,X93] : (sP4(X92,'F_0') | p(s(bool,'F_0')) | p(s(bool,X93)) | ~sP5(X92,X93)) ) | $spl3),
  inference(resolution,[],[f7582,f269])).
fof(f7582,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | sP4(X1,'F_0')) ) | $spl3),
  inference(subsumption_resolution,[],[f7558,f361])).
fof(f7558,plain,(
  ( ! [X1] : (sP4(X1,'F_0') | ~p(s(bool,'T_0')) | ~p(s(bool,X1))) ) | $spl3),
  inference(resolution,[],[f7491,f254])).
fof(f7491,plain,(
  ( ! [X214] : (sP4('T_0',X214) | sP4(X214,'F_0')) ) | $spl3),
  inference(resolution,[],[f3989,f332])).
fof(f7780,plain,(
  ( ! [X20] : (sP4(X20,'F_0')) ) | $spl186),
  inference(cnf_transformation,[],[f7780_D])).
fof(f7780_D,plain,(
  ( ! [X20] : (sP4(X20,'F_0')) ) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f21242,plain,(
  ( ! [X163,X164,X162] : (p(s(bool,X162)) | sP4(X163,'F_0') | sP4(X162,'const.option.IS_NONE_1'(s('type.option.option'(X164),'const.option.NONE_0')))) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f7782,f14063])).
fof(f21195,plain,(
  ( ! [X158,X157] : (sP4(X157,'F_0') | sP4('const.option.IS_NONE_1'(s('type.option.option'(X158),'const.option.NONE_0')),X157)) ) | $spl448),
  inference(cnf_transformation,[],[f21195_D])).
fof(f21195_D,plain,(
  ( ! [X158,X157] : (sP4(X157,'F_0') | sP4('const.option.IS_NONE_1'(s('type.option.option'(X158),'const.option.NONE_0')),X157)) ) <=> ~$spl448),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl448])])).
fof(f21241,plain,(
  ( ! [X161,X159,X160] : (p(s(bool,X159)) | sP4(X160,'F_0') | sP4('const.option.IS_NONE_1'(s('type.option.option'(X161),'const.option.NONE_0')),X160)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f7782,f14056])).
fof(f21240,plain,(
  ( ! [X158,X156,X157] : (p(s(bool,X156)) | sP4(X157,'F_0') | sP4('const.option.IS_NONE_1'(s('type.option.option'(X158),'const.option.NONE_0')),X156)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f7782,f14055])).
fof(f21194,plain,(
  ( ! [X149,X147,X150,X148] : (sP4(X147,'F_0') | sP4(X147,X148) | sP1(X148,X149,X150)) ) | $spl446),
  inference(cnf_transformation,[],[f21194_D])).
fof(f21194_D,plain,(
  ( ! [X149,X147,X150,X148] : (sP4(X147,'F_0') | sP4(X147,X148) | sP1(X148,X149,X150)) ) <=> ~$spl446),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl446])])).
fof(f21239,plain,(
  ( ! [X154,X152,X151,X155,X153] : (p(s(bool,X151)) | sP4(X152,'F_0') | sP1(X153,X154,X155) | sP4(X152,X153)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f7782,f13362])).
fof(f21238,plain,(
  ( ! [X146,X149,X147,X150,X148] : (p(s(bool,X146)) | sP4(X147,'F_0') | sP1(X148,X149,X150) | sP4(X146,X148)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f7782,f13361])).
fof(f21193,plain,(
  ( ! [X140,X138,X139,X137] : (sP4(X137,'F_0') | sP1(X138,X139,X140) | sP4(X138,X137)) ) | $spl444),
  inference(cnf_transformation,[],[f21193_D])).
fof(f21193_D,plain,(
  ( ! [X140,X138,X139,X137] : (sP4(X137,'F_0') | sP1(X138,X139,X140) | sP4(X138,X137)) ) <=> ~$spl444),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl444])])).
fof(f21237,plain,(
  ( ! [X144,X142,X145,X143,X141] : (p(s(bool,X141)) | sP4(X142,'F_0') | sP1(X143,X144,X145) | sP4(X143,X142)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f7782,f13354])).
fof(f21236,plain,(
  ( ! [X140,X138,X136,X139,X137] : (p(s(bool,X136)) | sP4(X137,'F_0') | sP1(X138,X139,X140) | sP4(X138,X136)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f7782,f13353])).
fof(f21192,plain,(
  ( ! [X121,X122,X120,X119] : (sP4(X119,'F_0') | sP4(X119,X121) | sP1(X120,X121,X122)) ) | $spl442),
  inference(cnf_transformation,[],[f21192_D])).
fof(f21192_D,plain,(
  ( ! [X121,X122,X120,X119] : (sP4(X119,'F_0') | sP4(X119,X121) | sP1(X120,X121,X122)) ) <=> ~$spl442),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl442])])).
fof(f21233,plain,(
  ( ! [X127,X125,X123,X126,X124] : (p(s(bool,X123)) | sP4(X124,'F_0') | sP1(X125,X126,X127) | sP4(X124,X126)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f7782,f12842])).
fof(f21232,plain,(
  ( ! [X121,X118,X122,X120,X119] : (p(s(bool,X118)) | sP4(X119,'F_0') | sP1(X120,X121,X122) | sP4(X118,X121)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f7782,f12841])).
fof(f21191,plain,(
  ( ! [X111,X109,X112,X110] : (sP4(X109,'F_0') | sP1(X110,X111,X112) | sP4(X111,X109)) ) | $spl440),
  inference(cnf_transformation,[],[f21191_D])).
fof(f21191_D,plain,(
  ( ! [X111,X109,X112,X110] : (sP4(X109,'F_0') | sP1(X110,X111,X112) | sP4(X111,X109)) ) <=> ~$spl440),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl440])])).
fof(f21231,plain,(
  ( ! [X116,X114,X117,X115,X113] : (p(s(bool,X113)) | sP4(X114,'F_0') | sP1(X115,X116,X117) | sP4(X116,X114)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f7782,f12834])).
fof(f21230,plain,(
  ( ! [X111,X109,X112,X110,X108] : (p(s(bool,X108)) | sP4(X109,'F_0') | sP1(X110,X111,X112) | sP4(X111,X108)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f7782,f12833])).
fof(f21190,plain,(
  ( ! [X94,X92,X93,X91] : (sP4(X91,'F_0') | sP4(X91,X92) | sP2(X92,X93,X94)) ) | $spl438),
  inference(cnf_transformation,[],[f21190_D])).
fof(f21190_D,plain,(
  ( ! [X94,X92,X93,X91] : (sP4(X91,'F_0') | sP4(X91,X92) | sP2(X92,X93,X94)) ) <=> ~$spl438),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl438])])).
fof(f21227,plain,(
  ( ! [X99,X97,X95,X98,X96] : (p(s(bool,X95)) | sP4(X96,'F_0') | sP2(X97,X98,X99) | sP4(X96,X97)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f7782,f12387])).
fof(f21226,plain,(
  ( ! [X94,X92,X90,X93,X91] : (p(s(bool,X90)) | sP4(X91,'F_0') | sP2(X92,X93,X94) | sP4(X90,X92)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f7782,f12386])).
fof(f21189,plain,(
  ( ! [X83,X81,X84,X82] : (sP4(X81,'F_0') | sP2(X82,X83,X84) | sP4(X82,X81)) ) | $spl436),
  inference(cnf_transformation,[],[f21189_D])).
fof(f21189_D,plain,(
  ( ! [X83,X81,X84,X82] : (sP4(X81,'F_0') | sP2(X82,X83,X84) | sP4(X82,X81)) ) <=> ~$spl436),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl436])])).
fof(f21225,plain,(
  ( ! [X88,X87,X85,X89,X86] : (p(s(bool,X85)) | sP4(X86,'F_0') | sP2(X87,X88,X89) | sP4(X87,X86)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f7782,f12379])).
fof(f21224,plain,(
  ( ! [X80,X83,X81,X84,X82] : (p(s(bool,X80)) | sP4(X81,'F_0') | sP2(X82,X83,X84) | sP4(X82,X80)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f7782,f12378])).
fof(f21188,plain,(
  ( ! [X66,X64,X65,X63] : (sP4(X63,'F_0') | sP4(X63,X65) | sP2(X64,X65,X66)) ) | $spl434),
  inference(cnf_transformation,[],[f21188_D])).
fof(f21188_D,plain,(
  ( ! [X66,X64,X65,X63] : (sP4(X63,'F_0') | sP4(X63,X65) | sP2(X64,X65,X66)) ) <=> ~$spl434),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl434])])).
fof(f21221,plain,(
  ( ! [X70,X68,X71,X69,X67] : (p(s(bool,X67)) | sP4(X68,'F_0') | sP2(X69,X70,X71) | sP4(X68,X70)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f7782,f11816])).
fof(f21220,plain,(
  ( ! [X66,X64,X62,X65,X63] : (p(s(bool,X62)) | sP4(X63,'F_0') | sP2(X64,X65,X66) | sP4(X62,X65)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f7782,f11815])).
fof(f21187,plain,(
  ( ! [X54,X56,X55,X53] : (sP4(X53,'F_0') | sP2(X54,X55,X56) | sP4(X55,X53)) ) | $spl432),
  inference(cnf_transformation,[],[f21187_D])).
fof(f21187_D,plain,(
  ( ! [X54,X56,X55,X53] : (sP4(X53,'F_0') | sP2(X54,X55,X56) | sP4(X55,X53)) ) <=> ~$spl432),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl432])])).
fof(f21219,plain,(
  ( ! [X61,X59,X57,X60,X58] : (p(s(bool,X57)) | sP4(X58,'F_0') | sP2(X59,X60,X61) | sP4(X60,X58)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f7782,f11808])).
fof(f21218,plain,(
  ( ! [X54,X52,X56,X55,X53] : (p(s(bool,X52)) | sP4(X53,'F_0') | sP2(X54,X55,X56) | sP4(X55,X52)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f7782,f11807])).
fof(f21180,plain,(
  ( ! [X167,X165,X166] : (p(s(bool,X165)) | sP4(X166,'F_0') | sP4(X165,'const.option.IS_NONE_1'(s('type.option.option'(X167),'const.option.NONE_0')))) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f7781,f14064])).
fof(f7781,plain,(
  ( ! [X90,X91] : (~sP5(X91,X90) | p(s(bool,X91)) | sP4(X90,'F_0')) ) | $spl3),
  inference(subsumption_resolution,[],[f7696,f307])).
fof(f7696,plain,(
  ( ! [X90,X91] : (sP4(X90,'F_0') | p(s(bool,'F_0')) | p(s(bool,X91)) | ~sP5(X91,X90)) ) | $spl3),
  inference(resolution,[],[f7582,f269])).
fof(f21179,plain,(
  ( ! [X163,X164,X162] : (p(s(bool,X162)) | sP4(X163,'F_0') | sP4(X163,'const.option.IS_NONE_1'(s('type.option.option'(X164),'const.option.NONE_0')))) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f7781,f14063])).
fof(f21178,plain,(
  ( ! [X161,X159,X160] : (p(s(bool,X159)) | sP4(X160,'F_0') | sP4('const.option.IS_NONE_1'(s('type.option.option'(X161),'const.option.NONE_0')),X159)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f7781,f14056])).
fof(f21177,plain,(
  ( ! [X158,X156,X157] : (p(s(bool,X156)) | sP4(X157,'F_0') | sP4('const.option.IS_NONE_1'(s('type.option.option'(X158),'const.option.NONE_0')),X157)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f7781,f14055])).
fof(f21176,plain,(
  ( ! [X154,X152,X151,X155,X153] : (p(s(bool,X151)) | sP4(X152,'F_0') | sP1(X153,X154,X155) | sP4(X151,X153)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f7781,f13362])).
fof(f21175,plain,(
  ( ! [X146,X149,X147,X150,X148] : (p(s(bool,X146)) | sP4(X147,'F_0') | sP1(X148,X149,X150) | sP4(X147,X148)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f7781,f13361])).
fof(f21174,plain,(
  ( ! [X144,X142,X145,X143,X141] : (p(s(bool,X141)) | sP4(X142,'F_0') | sP1(X143,X144,X145) | sP4(X143,X141)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f7781,f13354])).
fof(f21173,plain,(
  ( ! [X140,X138,X136,X139,X137] : (p(s(bool,X136)) | sP4(X137,'F_0') | sP1(X138,X139,X140) | sP4(X138,X137)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f7781,f13353])).
fof(f21170,plain,(
  ( ! [X127,X125,X123,X126,X124] : (p(s(bool,X123)) | sP4(X124,'F_0') | sP1(X125,X126,X127) | sP4(X123,X126)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f7781,f12842])).
fof(f21169,plain,(
  ( ! [X121,X118,X122,X120,X119] : (p(s(bool,X118)) | sP4(X119,'F_0') | sP1(X120,X121,X122) | sP4(X119,X121)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f7781,f12841])).
fof(f21168,plain,(
  ( ! [X116,X114,X117,X115,X113] : (p(s(bool,X113)) | sP4(X114,'F_0') | sP1(X115,X116,X117) | sP4(X116,X113)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f7781,f12834])).
fof(f21167,plain,(
  ( ! [X111,X109,X112,X110,X108] : (p(s(bool,X108)) | sP4(X109,'F_0') | sP1(X110,X111,X112) | sP4(X111,X109)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f7781,f12833])).
fof(f21164,plain,(
  ( ! [X99,X97,X95,X98,X96] : (p(s(bool,X95)) | sP4(X96,'F_0') | sP2(X97,X98,X99) | sP4(X95,X97)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f7781,f12387])).
fof(f21163,plain,(
  ( ! [X94,X92,X90,X93,X91] : (p(s(bool,X90)) | sP4(X91,'F_0') | sP2(X92,X93,X94) | sP4(X91,X92)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f7781,f12386])).
fof(f21162,plain,(
  ( ! [X88,X87,X85,X89,X86] : (p(s(bool,X85)) | sP4(X86,'F_0') | sP2(X87,X88,X89) | sP4(X87,X85)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f7781,f12379])).
fof(f21161,plain,(
  ( ! [X80,X83,X81,X84,X82] : (p(s(bool,X80)) | sP4(X81,'F_0') | sP2(X82,X83,X84) | sP4(X82,X81)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f7781,f12378])).
fof(f21158,plain,(
  ( ! [X70,X68,X71,X69,X67] : (p(s(bool,X67)) | sP4(X68,'F_0') | sP2(X69,X70,X71) | sP4(X67,X70)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f7781,f11816])).
fof(f21157,plain,(
  ( ! [X66,X64,X62,X65,X63] : (p(s(bool,X62)) | sP4(X63,'F_0') | sP2(X64,X65,X66) | sP4(X63,X65)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f7781,f11815])).
fof(f21156,plain,(
  ( ! [X61,X59,X57,X60,X58] : (p(s(bool,X57)) | sP4(X58,'F_0') | sP2(X59,X60,X61) | sP4(X60,X57)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f7781,f11808])).
fof(f21155,plain,(
  ( ! [X54,X52,X56,X55,X53] : (p(s(bool,X52)) | sP4(X53,'F_0') | sP2(X54,X55,X56) | sP4(X55,X53)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f7781,f11807])).
fof(f20679,plain,(
  ( ! [X231,X233,X232] : (sP4(X231,'F_0') | p(s(bool,X232)) | p(s(bool,X233)) | ~sP5(X232,X233)) ) | $spl3),
  inference(duplicate_literal_removal,[],[f20640])).
fof(f20640,plain,(
  ( ! [X231,X233,X232] : (sP4(X231,'F_0') | sP4(X231,'F_0') | p(s(bool,X232)) | p(s(bool,X233)) | ~sP5(X232,X233)) ) | $spl3),
  inference(resolution,[],[f7691,f269])).
fof(f7691,plain,(
  ( ! [X80,X81] : (~p(s(bool,X81)) | sP4(X80,X81) | sP4(X80,'F_0')) ) | $spl3),
  inference(resolution,[],[f7582,f258])).
fof(f20681,plain,(
  ( ! [X227,X225,X226] : (sP4(X225,'F_0') | p(s(bool,X226)) | p(s(bool,X227)) | ~sP5(X226,X227)) ) | $spl3),
  inference(duplicate_literal_removal,[],[f20638])).
fof(f20638,plain,(
  ( ! [X227,X225,X226] : (sP4(X225,'F_0') | sP4(X225,'F_0') | p(s(bool,X226)) | p(s(bool,X227)) | ~sP5(X226,X227)) ) | $spl3),
  inference(resolution,[],[f7691,f273])).
fof(f20685,plain,(
  ( ! [X0,X3,X1] : (sP4(X0,X1) | s(bool,X1) != s(bool,X3) | ~p(s(bool,X3)) | sP4(X0,'F_0')) ) | $spl430),
  inference(cnf_transformation,[],[f20685_D])).
fof(f20685_D,plain,(
  ( ! [X0,X3,X1] : (sP4(X0,X1) | s(bool,X1) != s(bool,X3) | ~p(s(bool,X3)) | sP4(X0,'F_0')) ) <=> ~$spl430),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl430])])).
fof(f20615,plain,(
  ( ! [X185,X182,X184,X183] : (sP4(X182,X183) | sP4(X182,'F_0') | p(s(bool,X184)) | ~p(s(bool,X185)) | s(bool,X183) != s(bool,X185)) ) | $spl3),
  inference(resolution,[],[f7691,f360])).
fof(f20570,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP4(X35,X36) | sP4(X35,'F_0') | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X38) | s(bool,X37) != s(bool,X39)) ) | $spl3),
  inference(resolution,[],[f7691,f223])).
fof(f20569,plain,(
  ( ! [X30,X33,X31,X34,X32] : (sP4(X30,X31) | sP4(X30,'F_0') | p(s(bool,X32)) | ~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X31) != s(bool,X33)) ) | $spl3),
  inference(resolution,[],[f7691,f222])).
fof(f20567,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP4(X20,X21) | sP4(X20,'F_0') | p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X24) | s(bool,X21) != s(bool,X23)) ) | $spl3),
  inference(resolution,[],[f7691,f221])).
fof(f20565,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP4(X10,X11) | sP4(X10,'F_0') | ~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,X14)) ) | $spl3),
  inference(resolution,[],[f7691,f219])).
fof(f20563,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP4(X0,X1) | sP4(X0,'F_0') | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) ) | $spl3),
  inference(resolution,[],[f7691,f217])).
fof(f20514,plain,(
  ( ! [X231,X233,X232] : (sP4('F_0',X231) | sP4(X231,'F_0') | p(s(bool,X232)) | p(s(bool,X233)) | ~sP5(X232,X233)) ) | $spl3),
  inference(resolution,[],[f7690,f269])).
fof(f7690,plain,(
  ( ! [X78,X79] : (~p(s(bool,X79)) | sP4(X79,X78) | sP4(X78,'F_0')) ) | $spl3),
  inference(resolution,[],[f7582,f255])).
fof(f20512,plain,(
  ( ! [X227,X225,X226] : (sP4('F_0',X225) | sP4(X225,'F_0') | p(s(bool,X226)) | p(s(bool,X227)) | ~sP5(X226,X227)) ) | $spl3),
  inference(resolution,[],[f7690,f273])).
fof(f20550,plain,(
  ( ! [X0,X3,X1] : (sP4(X0,X1) | s(bool,X0) != s(bool,X3) | ~p(s(bool,X3)) | sP4(X1,'F_0')) ) | $spl426),
  inference(cnf_transformation,[],[f20550_D])).
fof(f20550_D,plain,(
  ( ! [X0,X3,X1] : (sP4(X0,X1) | s(bool,X0) != s(bool,X3) | ~p(s(bool,X3)) | sP4(X1,'F_0')) ) <=> ~$spl426),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl426])])).
fof(f20489,plain,(
  ( ! [X185,X182,X184,X183] : (sP4(X182,X183) | sP4(X183,'F_0') | p(s(bool,X184)) | ~p(s(bool,X185)) | s(bool,X182) != s(bool,X185)) ) | $spl3),
  inference(resolution,[],[f7690,f360])).
fof(f20444,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP4(X35,X36) | sP4(X36,'F_0') | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X35) != s(bool,X38) | s(bool,X37) != s(bool,X39)) ) | $spl3),
  inference(resolution,[],[f7690,f223])).
fof(f20443,plain,(
  ( ! [X30,X33,X31,X34,X32] : (sP4(X30,X31) | sP4(X31,'F_0') | p(s(bool,X32)) | ~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X30) != s(bool,X33)) ) | $spl3),
  inference(resolution,[],[f7690,f222])).
fof(f20441,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP4(X20,X21) | sP4(X21,'F_0') | p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X24) | s(bool,X20) != s(bool,X23)) ) | $spl3),
  inference(resolution,[],[f7690,f221])).
fof(f20439,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP4(X10,X11) | sP4(X11,'F_0') | ~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X10) != s(bool,X13) | s(bool,X12) != s(bool,X14)) ) | $spl3),
  inference(resolution,[],[f7690,f219])).
fof(f20437,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP4(X0,X1) | sP4(X1,'F_0') | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X0) != s(bool,X3)) ) | $spl3),
  inference(resolution,[],[f7690,f217])).
fof(f1504,plain,(
  sP4(sK8,'F_0') | $spl56),
  inference(cnf_transformation,[],[f1504_D])).
fof(f1504_D,plain,(
  sP4(sK8,'F_0') <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f20382,plain,(
  ( ! [X596,X594,X593,X595] : (p(s(bool,X593)) | sP4(sK8,'F_0') | sP4(X593,X594) | sP1(X594,X595,X596)) ) | $spl3),
  inference(resolution,[],[f7689,f5499])).
fof(f7689,plain,(
  ( ! [X76,X77] : (~sP4(X76,X77) | p(s(bool,X77)) | sP4(X76,'F_0')) ) | $spl3),
  inference(resolution,[],[f7582,f253])).
fof(f20381,plain,(
  ( ! [X592,X590,X589,X591] : (p(s(bool,X589)) | sP4(sK8,'F_0') | sP4(X589,X590) | sP1(X591,X590,X592)) ) | $spl3),
  inference(resolution,[],[f7689,f5498])).
fof(f20380,plain,(
  ( ! [X588,X586,X585,X587] : (p(s(bool,X585)) | sP4(sK8,'F_0') | sP4(X585,X586) | sP2(X586,X587,X588)) ) | $spl3),
  inference(resolution,[],[f7689,f5496])).
fof(f20379,plain,(
  ( ! [X584,X582,X581,X583] : (p(s(bool,X581)) | sP4(sK8,'F_0') | sP4(X581,X582) | sP2(X583,X582,X584)) ) | $spl3),
  inference(resolution,[],[f7689,f5495])).
fof(f20378,plain,(
  ( ! [X580,X578,X577,X579] : (p(s(bool,X577)) | sP4(sK8,'F_0') | sP4(X578,X577) | sP1(X578,X579,X580)) ) | $spl3),
  inference(resolution,[],[f7689,f5373])).
fof(f20377,plain,(
  ( ! [X573,X575,X576,X574] : (p(s(bool,X573)) | sP4(sK8,'F_0') | sP4(X574,X573) | sP1(X575,X574,X576)) ) | $spl3),
  inference(resolution,[],[f7689,f5372])).
fof(f20376,plain,(
  ( ! [X569,X571,X572,X570] : (p(s(bool,X569)) | sP4(sK8,'F_0') | sP4(X570,X569) | sP2(X570,X571,X572)) ) | $spl3),
  inference(resolution,[],[f7689,f5370])).
fof(f20375,plain,(
  ( ! [X565,X567,X568,X566] : (p(s(bool,X565)) | sP4(sK8,'F_0') | sP4(X566,X565) | sP2(X567,X566,X568)) ) | $spl3),
  inference(resolution,[],[f7689,f5369])).
fof(f20041,plain,(
  ( ! [X432,X434,X433,X435] : (p(s(bool,X432)) | sP4(sK8,'F_0') | sP4(X432,X433) | sP1(X433,X434,X435)) ) | $spl3),
  inference(resolution,[],[f7688,f4346])).
fof(f7688,plain,(
  ( ! [X74,X75] : (~sP4(X75,X74) | p(s(bool,X75)) | sP4(X74,'F_0')) ) | $spl3),
  inference(resolution,[],[f7582,f253])).
fof(f20040,plain,(
  ( ! [X428,X430,X429,X431] : (p(s(bool,X428)) | sP4(sK8,'F_0') | sP4(X428,X429) | sP1(X430,X429,X431)) ) | $spl3),
  inference(resolution,[],[f7688,f4345])).
fof(f20039,plain,(
  ( ! [X424,X426,X425,X427] : (p(s(bool,X424)) | sP4(sK8,'F_0') | sP4(X424,X425) | sP2(X425,X426,X427)) ) | $spl3),
  inference(resolution,[],[f7688,f4343])).
fof(f20038,plain,(
  ( ! [X420,X422,X421,X423] : (p(s(bool,X420)) | sP4(sK8,'F_0') | sP4(X420,X421) | sP2(X422,X421,X423)) ) | $spl3),
  inference(resolution,[],[f7688,f4342])).
fof(f20035,plain,(
  ( ! [X414,X416,X413,X415] : (p(s(bool,X413)) | sP4(sK8,'F_0') | sP4(X414,X413) | sP1(X414,X415,X416)) ) | $spl3),
  inference(resolution,[],[f7688,f3487])).
fof(f20034,plain,(
  ( ! [X412,X410,X409,X411] : (p(s(bool,X409)) | sP4(sK8,'F_0') | sP4(X410,X409) | sP1(X411,X410,X412)) ) | $spl3),
  inference(resolution,[],[f7688,f3486])).
fof(f20033,plain,(
  ( ! [X408,X405,X407,X406] : (p(s(bool,X405)) | sP4(sK8,'F_0') | sP4(X406,X405) | sP2(X406,X407,X408)) ) | $spl3),
  inference(resolution,[],[f7688,f3484])).
fof(f20032,plain,(
  ( ! [X401,X403,X404,X402] : (p(s(bool,X401)) | sP4(sK8,'F_0') | sP4(X402,X401) | sP2(X403,X402,X404)) ) | $spl3),
  inference(resolution,[],[f7688,f3483])).
fof(f19902,plain,(
  ( ! [X0] : (sP4('const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0) | ~p(s(bool,X0))) ) | $spl424),
  inference(cnf_transformation,[],[f19902_D])).
fof(f19902_D,plain,(
  ( ! [X0] : (sP4('const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0) | ~p(s(bool,X0))) ) <=> ~$spl424),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl424])])).
fof(f19890,plain,(
  ( ! [X0] : (sP4('const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0) | ~p(s(bool,X0)) | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6))))) ) | $spl5),
  inference(superposition,[],[f19087,f357])).
fof(f19087,plain,(
  ( ! [X55] : (sP4('const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)),X55) | ~p(s(bool,X55))) ) | $spl5),
  inference(resolution,[],[f368,f258])).
fof(f19710,plain,(
  ( ! [X0] : (sP4(X0,'const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0'))) | ~p(s(bool,X0))) ) | $spl422),
  inference(cnf_transformation,[],[f19710_D])).
fof(f19710_D,plain,(
  ( ! [X0] : (sP4(X0,'const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0'))) | ~p(s(bool,X0))) ) <=> ~$spl422),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl422])])).
fof(f19698,plain,(
  ( ! [X0] : (sP4(X0,'const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0'))) | ~p(s(bool,X0)) | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6))))) ) | $spl5),
  inference(superposition,[],[f19086,f357])).
fof(f19086,plain,(
  ( ! [X54] : (sP4(X54,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7))) | ~p(s(bool,X54))) ) | $spl5),
  inference(resolution,[],[f368,f255])).
fof(f19663,plain,(
  ( ! [X0] : (~sP4('const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0) | p(s(bool,X0))) ) | $spl420),
  inference(cnf_transformation,[],[f19663_D])).
fof(f19663_D,plain,(
  ( ! [X0] : (~sP4('const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0) | p(s(bool,X0))) ) <=> ~$spl420),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl420])])).
fof(f19606,plain,(
  ( ! [X0] : (~sP4('const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0) | p(s(bool,X0)) | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6))))) ) | $spl5),
  inference(superposition,[],[f19085,f357])).
fof(f19085,plain,(
  ( ! [X53] : (~sP4('const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)),X53) | p(s(bool,X53))) ) | $spl5),
  inference(resolution,[],[f368,f253])).
fof(f19464,plain,(
  ( ! [X0] : (~sP4(X0,'const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0'))) | p(s(bool,X0))) ) | $spl418),
  inference(cnf_transformation,[],[f19464_D])).
fof(f19464_D,plain,(
  ( ! [X0] : (~sP4(X0,'const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0'))) | p(s(bool,X0))) ) <=> ~$spl418),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl418])])).
fof(f19408,plain,(
  ( ! [X0] : (~sP4(X0,'const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0'))) | p(s(bool,X0)) | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6))))) ) | $spl5),
  inference(superposition,[],[f19084,f357])).
fof(f19084,plain,(
  ( ! [X52] : (~sP4(X52,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7))) | p(s(bool,X52))) ) | $spl5),
  inference(resolution,[],[f368,f253])).
fof(f452,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) | $spl6),
  inference(cnf_transformation,[],[f452_D])).
fof(f452_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f19114,plain,(
  ( ! [X78,X79] : (p(s(bool,X78)) | ~p(s(bool,X79)) | s(bool,X79) != s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) | $spl5),
  inference(resolution,[],[f368,f360])).
fof(f19069,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7))) | s(bool,X21) != s(bool,X23)) ) | $spl5),
  inference(resolution,[],[f368,f223])).
fof(f19068,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) | $spl5),
  inference(resolution,[],[f368,f222])).
fof(f19066,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) | $spl5),
  inference(resolution,[],[f368,f221])).
fof(f19064,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7))) | s(bool,X6) != s(bool,X8)) ) | $spl5),
  inference(resolution,[],[f368,f219])).
fof(f19062,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) | $spl5),
  inference(resolution,[],[f368,f217])).
fof(f14462,plain,(
  ( ! [X518,X517] : (sP2(X517,'T_0',X518)) ) | $spl346),
  inference(cnf_transformation,[],[f14462_D])).
fof(f14462_D,plain,(
  ( ! [X518,X517] : (sP2(X517,'T_0',X518)) ) <=> ~$spl346),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl346])])).
fof(f18987,plain,(
  sP4('const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)),sK8) | $spl416),
  inference(cnf_transformation,[],[f18987_D])).
fof(f18987_D,plain,(
  sP4('const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)),sK8) <=> ~$spl416),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl416])])).
fof(f11950,plain,(
  ( ! [X26,X27] : (sP4('const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)),sK8) | sP2(X26,'T_0',X27)) ) | ($spl3 | $spl275)),
  inference(resolution,[],[f11822,f4342])).
fof(f14853,plain,(
  ( ! [X518,X517] : (sP2('T_0',X517,X518)) ) | $spl368),
  inference(cnf_transformation,[],[f14853_D])).
fof(f14853_D,plain,(
  ( ! [X518,X517] : (sP2('T_0',X517,X518)) ) <=> ~$spl368),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl368])])).
fof(f11951,plain,(
  ( ! [X28,X29] : (sP4('const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)),sK8) | sP2('T_0',X28,X29)) ) | ($spl3 | $spl275)),
  inference(resolution,[],[f11822,f4343])).
fof(f15077,plain,(
  ( ! [X518,X517] : (sP1(X517,'T_0',X518)) ) | $spl370),
  inference(cnf_transformation,[],[f15077_D])).
fof(f15077_D,plain,(
  ( ! [X518,X517] : (sP1(X517,'T_0',X518)) ) <=> ~$spl370),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl370])])).
fof(f11952,plain,(
  ( ! [X30,X31] : (sP4('const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)),sK8) | sP1(X30,'T_0',X31)) ) | ($spl3 | $spl275)),
  inference(resolution,[],[f11822,f4345])).
fof(f15300,plain,(
  ( ! [X518,X517] : (sP1('T_0',X517,X518)) ) | $spl372),
  inference(cnf_transformation,[],[f15300_D])).
fof(f15300_D,plain,(
  ( ! [X518,X517] : (sP1('T_0',X517,X518)) ) <=> ~$spl372),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl372])])).
fof(f11953,plain,(
  ( ! [X33,X32] : (sP4('const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)),sK8) | sP1('T_0',X32,X33)) ) | ($spl3 | $spl275)),
  inference(resolution,[],[f11822,f4346])).
fof(f18981,plain,(
  sP4(sK8,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7))) | $spl414),
  inference(cnf_transformation,[],[f18981_D])).
fof(f18981_D,plain,(
  sP4(sK8,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7))) <=> ~$spl414),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl414])])).
fof(f11967,plain,(
  ( ! [X43,X42] : (sP4(sK8,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7))) | sP2(X42,'T_0',X43)) ) | ($spl3 | $spl275)),
  inference(resolution,[],[f11822,f5495])).
fof(f11968,plain,(
  ( ! [X45,X44] : (sP4(sK8,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7))) | sP2('T_0',X44,X45)) ) | ($spl3 | $spl275)),
  inference(resolution,[],[f11822,f5496])).
fof(f11969,plain,(
  ( ! [X47,X46] : (sP4(sK8,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7))) | sP1(X46,'T_0',X47)) ) | ($spl3 | $spl275)),
  inference(resolution,[],[f11822,f5498])).
fof(f11970,plain,(
  ( ! [X48,X49] : (sP4(sK8,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7))) | sP1('T_0',X48,X49)) ) | ($spl3 | $spl275)),
  inference(resolution,[],[f11822,f5499])).
fof(f18960,plain,(
  ( ! [X155] : (sP5(X155,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) | $spl412),
  inference(cnf_transformation,[],[f18960_D])).
fof(f18960_D,plain,(
  ( ! [X155] : (sP5(X155,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) <=> ~$spl412),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl412])])).
fof(f17777,plain,(
  ( ! [X154,X155,X153] : (sP2(X153,'T_0',X154) | sP5(X155,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) | ($spl3 | $spl179 | $spl275)),
  inference(resolution,[],[f11815,f11822])).
fof(f18959,plain,(
  ( ! [X155] : (sP5('const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)),X155)) ) | $spl410),
  inference(cnf_transformation,[],[f18959_D])).
fof(f18959_D,plain,(
  ( ! [X155] : (sP5('const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)),X155)) ) <=> ~$spl410),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl410])])).
fof(f17829,plain,(
  ( ! [X154,X155,X153] : (sP2(X153,'T_0',X154) | sP5('const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)),X155)) ) | ($spl3 | $spl179 | $spl275)),
  inference(resolution,[],[f11816,f11822])).
fof(f17988,plain,(
  ( ! [X154,X155,X153] : (sP2('T_0',X153,X154) | sP5(X155,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) | ($spl3 | $spl179 | $spl275)),
  inference(resolution,[],[f12386,f11822])).
fof(f18042,plain,(
  ( ! [X154,X155,X153] : (sP2('T_0',X153,X154) | sP5('const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)),X155)) ) | ($spl3 | $spl179 | $spl275)),
  inference(resolution,[],[f12387,f11822])).
fof(f18200,plain,(
  ( ! [X154,X155,X153] : (sP1(X153,'T_0',X154) | sP5(X155,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) | ($spl3 | $spl179 | $spl275)),
  inference(resolution,[],[f12841,f11822])).
fof(f18254,plain,(
  ( ! [X154,X155,X153] : (sP1(X153,'T_0',X154) | sP5('const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)),X155)) ) | ($spl3 | $spl179 | $spl275)),
  inference(resolution,[],[f12842,f11822])).
fof(f18413,plain,(
  ( ! [X154,X155,X153] : (sP1('T_0',X153,X154) | sP5(X155,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) | ($spl3 | $spl179 | $spl275)),
  inference(resolution,[],[f13361,f11822])).
fof(f18467,plain,(
  ( ! [X154,X155,X153] : (sP1('T_0',X153,X154) | sP5('const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)),X155)) ) | ($spl3 | $spl179 | $spl275)),
  inference(resolution,[],[f13362,f11822])).
fof(f18958,plain,(
  $false | ($spl0 | $spl3 | $spl4)),
  inference(subsumption_resolution,[],[f18939,f2640])).
fof(f2640,plain,(
  s('type.option.option'(sK6),'const.option.NONE_0') = s('type.option.option'(sK6),'const.bool.COND_3'(s(bool,'F_0'),s('type.option.option'(sK6),'const.option.NONE_0'),s('type.option.option'(sK6),sK7))) | ($spl0 | $spl3)),
  inference(backward_demodulation,[],[f1561,f364])).
fof(f1561,plain,(
  s(bool,'F_0') = s(bool,sK8) | $spl3),
  inference(resolution,[],[f367,f319])).
fof(f18939,plain,(
  s('type.option.option'(sK6),'const.option.NONE_0') != s('type.option.option'(sK6),'const.bool.COND_3'(s(bool,'F_0'),s('type.option.option'(sK6),'const.option.NONE_0'),s('type.option.option'(sK6),sK7))) | ($spl3 | $spl4)),
  inference(superposition,[],[f17541,f1687])).
fof(f1687,plain,(
  s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6)))) | $spl4),
  inference(subsumption_resolution,[],[f1686,f332])).
fof(f1686,plain,(
  p(s(bool,'F_0')) | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6)))) | $spl4),
  inference(forward_demodulation,[],[f1656,f340])).
fof(f340,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.option.IS_SOME_1'(s('type.option.option'(X0),'const.option.NONE_0')))) )),
  inference(cnf_transformation,[],[f96])).
fof(f96,plain,(
  ! [X0] : s(bool,'F_0') = s(bool,'const.option.IS_SOME_1'(s('type.option.option'(X0),'const.option.NONE_0')))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X6] : s(bool,'F_0') = s(bool,'const.option.IS_SOME_1'(s('type.option.option'(X6),'const.option.NONE_0')))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|1','thm.option.IS_SOME_DEF|split|1')).
fof(f1656,plain,(
  p(s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0')))) | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6)))) | $spl4),
  inference(superposition,[],[f369,f357])).
fof(f17541,plain,(
  ( ! [X74,X75] : (s('type.option.option'(X74),'const.option.NONE_0') != s('type.option.option'(X74),'const.bool.COND_3'(s(bool,'F_0'),s('type.option.option'(X74),'const.option.NONE_0'),s('type.option.option'(X74),'const.option.SOME_1'(s(X74,X75)))))) ) | $spl3),
  inference(forward_demodulation,[],[f1569,f1561])).
fof(f1569,plain,(
  ( ! [X74,X75] : (s('type.option.option'(X74),'const.option.NONE_0') != s('type.option.option'(X74),'const.bool.COND_3'(s(bool,sK8),s('type.option.option'(X74),'const.option.NONE_0'),s('type.option.option'(X74),'const.option.SOME_1'(s(X74,X75)))))) ) | $spl3),
  inference(resolution,[],[f367,f349])).
fof(f349,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | s('type.option.option'(X0),'const.option.NONE_0') != s('type.option.option'(X0),'const.bool.COND_3'(s(bool,X2),s('type.option.option'(X0),'const.option.NONE_0'),s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1)))))) )),
  inference(cnf_transformation,[],[f209])).
fof(f18854,plain,(
  ( ! [X17,X15,X16] : (sP4(X15,'const.option.IS_NONE_1'(s('type.option.option'(X16),'const.option.NONE_0'))) | p(s(bool,X15)) | p(s(bool,X17))) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f18835,f307])).
fof(f18835,plain,(
  ( ! [X17,X15,X16] : (sP4(X15,'const.option.IS_NONE_1'(s('type.option.option'(X16),'const.option.NONE_0'))) | p(s(bool,'F_0')) | p(s(bool,X15)) | p(s(bool,X17))) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f14064,f269])).
fof(f18853,plain,(
  ( ! [X14,X12,X13] : (sP4(X12,'const.option.IS_NONE_1'(s('type.option.option'(X13),'const.option.NONE_0'))) | p(s(bool,X12)) | p(s(bool,X14))) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f18834,f307])).
fof(f18834,plain,(
  ( ! [X14,X12,X13] : (sP4(X12,'const.option.IS_NONE_1'(s('type.option.option'(X13),'const.option.NONE_0'))) | p(s(bool,'F_0')) | p(s(bool,X12)) | p(s(bool,X14))) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f14064,f273])).
fof(f18833,plain,(
  ( ! [X10,X11,X9] : (sP4(X9,'const.option.IS_NONE_1'(s('type.option.option'(X10),'const.option.NONE_0'))) | p(s(bool,X9)) | sP4(X11,sK8)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f14064,f2514])).
fof(f2514,plain,(
  ( ! [X90,X91] : (~sP5(X91,X90) | p(s(bool,X91)) | sP4(X90,sK8)) ) | $spl3),
  inference(subsumption_resolution,[],[f2436,f307])).
fof(f2436,plain,(
  ( ! [X90,X91] : (sP4(X90,sK8) | p(s(bool,'F_0')) | p(s(bool,X91)) | ~sP5(X91,X90)) ) | $spl3),
  inference(resolution,[],[f1543,f269])).
fof(f18831,plain,(
  ( ! [X4,X5,X3] : (sP4(X3,'const.option.IS_NONE_1'(s('type.option.option'(X4),'const.option.NONE_0'))) | p(s(bool,X3)) | sP4(sK8,X5)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f14064,f2631])).
fof(f2631,plain,(
  ( ! [X90,X91] : (~sP5(X91,X90) | p(s(bool,X91)) | sP4(sK8,X90)) ) | $spl3),
  inference(subsumption_resolution,[],[f2553,f307])).
fof(f2553,plain,(
  ( ! [X90,X91] : (sP4(sK8,X90) | p(s(bool,'F_0')) | p(s(bool,X91)) | ~sP5(X91,X90)) ) | $spl3),
  inference(resolution,[],[f1544,f269])).
fof(f18802,plain,(
  ( ! [X17,X15,X16] : (sP4(X15,'const.option.IS_NONE_1'(s('type.option.option'(X16),'const.option.NONE_0'))) | p(s(bool,X17)) | p(s(bool,X15))) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f18782,f307])).
fof(f18782,plain,(
  ( ! [X17,X15,X16] : (sP4(X15,'const.option.IS_NONE_1'(s('type.option.option'(X16),'const.option.NONE_0'))) | p(s(bool,'F_0')) | p(s(bool,X17)) | p(s(bool,X15))) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f14063,f269])).
fof(f18801,plain,(
  ( ! [X14,X12,X13] : (sP4(X12,'const.option.IS_NONE_1'(s('type.option.option'(X13),'const.option.NONE_0'))) | p(s(bool,X14)) | p(s(bool,X12))) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f18781,f307])).
fof(f18781,plain,(
  ( ! [X14,X12,X13] : (sP4(X12,'const.option.IS_NONE_1'(s('type.option.option'(X13),'const.option.NONE_0'))) | p(s(bool,'F_0')) | p(s(bool,X14)) | p(s(bool,X12))) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f14063,f273])).
fof(f18779,plain,(
  ( ! [X6,X8,X7] : (sP4(X6,'const.option.IS_NONE_1'(s('type.option.option'(X7),'const.option.NONE_0'))) | p(s(bool,X6)) | sP4(X8,sK8)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f14063,f2515])).
fof(f2515,plain,(
  ( ! [X92,X93] : (~sP5(X92,X93) | p(s(bool,X93)) | sP4(X92,sK8)) ) | $spl3),
  inference(subsumption_resolution,[],[f2437,f307])).
fof(f2437,plain,(
  ( ! [X92,X93] : (sP4(X92,sK8) | p(s(bool,'F_0')) | p(s(bool,X93)) | ~sP5(X92,X93)) ) | $spl3),
  inference(resolution,[],[f1543,f269])).
fof(f18777,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,'const.option.IS_NONE_1'(s('type.option.option'(X1),'const.option.NONE_0'))) | p(s(bool,X0)) | sP4(sK8,X2)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f14063,f2632])).
fof(f2632,plain,(
  ( ! [X92,X93] : (~sP5(X92,X93) | p(s(bool,X93)) | sP4(sK8,X92)) ) | $spl3),
  inference(subsumption_resolution,[],[f2554,f307])).
fof(f2554,plain,(
  ( ! [X92,X93] : (sP4(sK8,X92) | p(s(bool,'F_0')) | p(s(bool,X93)) | ~sP5(X92,X93)) ) | $spl3),
  inference(resolution,[],[f1544,f269])).
fof(f18749,plain,(
  ( ! [X17,X15,X16] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(X15),'const.option.NONE_0')),X16) | p(s(bool,X16)) | p(s(bool,X17))) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f18730,f307])).
fof(f18730,plain,(
  ( ! [X17,X15,X16] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(X15),'const.option.NONE_0')),X16) | p(s(bool,'F_0')) | p(s(bool,X16)) | p(s(bool,X17))) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f14056,f269])).
fof(f18748,plain,(
  ( ! [X14,X12,X13] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(X12),'const.option.NONE_0')),X13) | p(s(bool,X13)) | p(s(bool,X14))) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f18729,f307])).
fof(f18729,plain,(
  ( ! [X14,X12,X13] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(X12),'const.option.NONE_0')),X13) | p(s(bool,'F_0')) | p(s(bool,X13)) | p(s(bool,X14))) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f14056,f273])).
fof(f18728,plain,(
  ( ! [X10,X11,X9] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(X9),'const.option.NONE_0')),X10) | p(s(bool,X10)) | sP4(X11,sK8)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f14056,f2514])).
fof(f18726,plain,(
  ( ! [X4,X5,X3] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(X3),'const.option.NONE_0')),X4) | p(s(bool,X4)) | sP4(sK8,X5)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f14056,f2631])).
fof(f18536,plain,(
  ( ! [X17,X15,X16] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(X15),'const.option.NONE_0')),X16) | p(s(bool,X17)) | p(s(bool,X16))) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f18516,f307])).
fof(f18516,plain,(
  ( ! [X17,X15,X16] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(X15),'const.option.NONE_0')),X16) | p(s(bool,'F_0')) | p(s(bool,X17)) | p(s(bool,X16))) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f14055,f269])).
fof(f18535,plain,(
  ( ! [X14,X12,X13] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(X12),'const.option.NONE_0')),X13) | p(s(bool,X14)) | p(s(bool,X13))) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f18515,f307])).
fof(f18515,plain,(
  ( ! [X14,X12,X13] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(X12),'const.option.NONE_0')),X13) | p(s(bool,'F_0')) | p(s(bool,X14)) | p(s(bool,X13))) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f14055,f273])).
fof(f18513,plain,(
  ( ! [X6,X8,X7] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(X6),'const.option.NONE_0')),X7) | p(s(bool,X7)) | sP4(X8,sK8)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f14055,f2515])).
fof(f18511,plain,(
  ( ! [X2,X0,X1] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(X0),'const.option.NONE_0')),X1) | p(s(bool,X1)) | sP4(sK8,X2)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f14055,f2632])).
fof(f18483,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP1(X25,X26,X27) | sP4(X28,X25) | p(s(bool,X28)) | p(s(bool,X29))) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f18476,f307])).
fof(f18476,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP1(X25,X26,X27) | sP4(X28,X25) | p(s(bool,'F_0')) | p(s(bool,X28)) | p(s(bool,X29))) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f13362,f269])).
fof(f18482,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP1(X20,X21,X22) | sP4(X23,X20) | p(s(bool,X23)) | p(s(bool,X24))) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f18475,f307])).
fof(f18475,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP1(X20,X21,X22) | sP4(X23,X20) | p(s(bool,'F_0')) | p(s(bool,X23)) | p(s(bool,X24))) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f13362,f273])).
fof(f18474,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP1(X15,X16,X17) | sP4(X18,X15) | p(s(bool,X18)) | sP4(X19,sK8)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f13362,f2514])).
fof(f18472,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP1(X5,X6,X7) | sP4(X8,X5) | p(s(bool,X8)) | sP4(sK8,X9)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f13362,f2631])).
fof(f18430,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP1(X25,X26,X27) | sP4(X28,X25) | p(s(bool,X29)) | p(s(bool,X28))) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f18422,f307])).
fof(f18422,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP1(X25,X26,X27) | sP4(X28,X25) | p(s(bool,'F_0')) | p(s(bool,X29)) | p(s(bool,X28))) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f13361,f269])).
fof(f18429,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP1(X20,X21,X22) | sP4(X23,X20) | p(s(bool,X24)) | p(s(bool,X23))) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f18421,f307])).
fof(f18421,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP1(X20,X21,X22) | sP4(X23,X20) | p(s(bool,'F_0')) | p(s(bool,X24)) | p(s(bool,X23))) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f13361,f273])).
fof(f18419,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP1(X10,X11,X12) | sP4(X13,X10) | p(s(bool,X13)) | sP4(X14,sK8)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f13361,f2515])).
fof(f18417,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP1(X0,X1,X2) | sP4(X3,X0) | p(s(bool,X3)) | sP4(sK8,X4)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f13361,f2632])).
fof(f18376,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP1(X25,X26,X27) | sP4(X25,X28) | p(s(bool,X28)) | p(s(bool,X29))) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f18369,f307])).
fof(f18369,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP1(X25,X26,X27) | sP4(X25,X28) | p(s(bool,'F_0')) | p(s(bool,X28)) | p(s(bool,X29))) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f13354,f269])).
fof(f18375,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP1(X20,X21,X22) | sP4(X20,X23) | p(s(bool,X23)) | p(s(bool,X24))) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f18368,f307])).
fof(f18368,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP1(X20,X21,X22) | sP4(X20,X23) | p(s(bool,'F_0')) | p(s(bool,X23)) | p(s(bool,X24))) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f13354,f273])).
fof(f18367,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP1(X15,X16,X17) | sP4(X15,X18) | p(s(bool,X18)) | sP4(X19,sK8)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f13354,f2514])).
fof(f18365,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP1(X5,X6,X7) | sP4(X5,X8) | p(s(bool,X8)) | sP4(sK8,X9)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f13354,f2631])).
fof(f18323,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP1(X25,X26,X27) | sP4(X25,X28) | p(s(bool,X29)) | p(s(bool,X28))) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f18315,f307])).
fof(f18315,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP1(X25,X26,X27) | sP4(X25,X28) | p(s(bool,'F_0')) | p(s(bool,X29)) | p(s(bool,X28))) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f13353,f269])).
fof(f18322,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP1(X20,X21,X22) | sP4(X20,X23) | p(s(bool,X24)) | p(s(bool,X23))) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f18314,f307])).
fof(f18314,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP1(X20,X21,X22) | sP4(X20,X23) | p(s(bool,'F_0')) | p(s(bool,X24)) | p(s(bool,X23))) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f13353,f273])).
fof(f18312,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP1(X10,X11,X12) | sP4(X10,X13) | p(s(bool,X13)) | sP4(X14,sK8)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f13353,f2515])).
fof(f18310,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP1(X0,X1,X2) | sP4(X0,X3) | p(s(bool,X3)) | sP4(sK8,X4)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f13353,f2632])).
fof(f18269,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP1(X25,X26,X27) | sP4(X28,X26) | p(s(bool,X28)) | p(s(bool,X29))) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f18263,f307])).
fof(f18263,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP1(X25,X26,X27) | sP4(X28,X26) | p(s(bool,'F_0')) | p(s(bool,X28)) | p(s(bool,X29))) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f12842,f269])).
fof(f18268,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP1(X20,X21,X22) | sP4(X23,X21) | p(s(bool,X23)) | p(s(bool,X24))) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f18262,f307])).
fof(f18262,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP1(X20,X21,X22) | sP4(X23,X21) | p(s(bool,'F_0')) | p(s(bool,X23)) | p(s(bool,X24))) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f12842,f273])).
fof(f18261,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP1(X15,X16,X17) | sP4(X18,X16) | p(s(bool,X18)) | sP4(X19,sK8)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f12842,f2514])).
fof(f18259,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP1(X5,X6,X7) | sP4(X8,X6) | p(s(bool,X8)) | sP4(sK8,X9)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f12842,f2631])).
fof(f18216,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP1(X25,X26,X27) | sP4(X28,X26) | p(s(bool,X29)) | p(s(bool,X28))) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f18209,f307])).
fof(f18209,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP1(X25,X26,X27) | sP4(X28,X26) | p(s(bool,'F_0')) | p(s(bool,X29)) | p(s(bool,X28))) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f12841,f269])).
fof(f18215,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP1(X20,X21,X22) | sP4(X23,X21) | p(s(bool,X24)) | p(s(bool,X23))) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f18208,f307])).
fof(f18208,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP1(X20,X21,X22) | sP4(X23,X21) | p(s(bool,'F_0')) | p(s(bool,X24)) | p(s(bool,X23))) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f12841,f273])).
fof(f18206,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP1(X10,X11,X12) | sP4(X13,X11) | p(s(bool,X13)) | sP4(X14,sK8)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f12841,f2515])).
fof(f18204,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP1(X0,X1,X2) | sP4(X3,X1) | p(s(bool,X3)) | sP4(sK8,X4)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f12841,f2632])).
fof(f18163,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP1(X25,X26,X27) | sP4(X26,X28) | p(s(bool,X28)) | p(s(bool,X29))) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f18157,f307])).
fof(f18157,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP1(X25,X26,X27) | sP4(X26,X28) | p(s(bool,'F_0')) | p(s(bool,X28)) | p(s(bool,X29))) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f12834,f269])).
fof(f18162,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP1(X20,X21,X22) | sP4(X21,X23) | p(s(bool,X23)) | p(s(bool,X24))) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f18156,f307])).
fof(f18156,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP1(X20,X21,X22) | sP4(X21,X23) | p(s(bool,'F_0')) | p(s(bool,X23)) | p(s(bool,X24))) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f12834,f273])).
fof(f18155,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP1(X15,X16,X17) | sP4(X16,X18) | p(s(bool,X18)) | sP4(X19,sK8)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f12834,f2514])).
fof(f18153,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP1(X5,X6,X7) | sP4(X6,X8) | p(s(bool,X8)) | sP4(sK8,X9)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f12834,f2631])).
fof(f18111,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP1(X25,X26,X27) | sP4(X26,X28) | p(s(bool,X29)) | p(s(bool,X28))) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f18104,f307])).
fof(f18104,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP1(X25,X26,X27) | sP4(X26,X28) | p(s(bool,'F_0')) | p(s(bool,X29)) | p(s(bool,X28))) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f12833,f269])).
fof(f18110,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP1(X20,X21,X22) | sP4(X21,X23) | p(s(bool,X24)) | p(s(bool,X23))) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f18103,f307])).
fof(f18103,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP1(X20,X21,X22) | sP4(X21,X23) | p(s(bool,'F_0')) | p(s(bool,X24)) | p(s(bool,X23))) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f12833,f273])).
fof(f18101,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP1(X10,X11,X12) | sP4(X11,X13) | p(s(bool,X13)) | sP4(X14,sK8)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f12833,f2515])).
fof(f18099,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP1(X0,X1,X2) | sP4(X1,X3) | p(s(bool,X3)) | sP4(sK8,X4)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f12833,f2632])).
fof(f18058,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP2(X25,X26,X27) | sP4(X28,X25) | p(s(bool,X28)) | p(s(bool,X29))) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f18051,f307])).
fof(f18051,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP2(X25,X26,X27) | sP4(X28,X25) | p(s(bool,'F_0')) | p(s(bool,X28)) | p(s(bool,X29))) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f12387,f269])).
fof(f18057,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP2(X20,X21,X22) | sP4(X23,X20) | p(s(bool,X23)) | p(s(bool,X24))) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f18050,f307])).
fof(f18050,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP2(X20,X21,X22) | sP4(X23,X20) | p(s(bool,'F_0')) | p(s(bool,X23)) | p(s(bool,X24))) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f12387,f273])).
fof(f18049,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP2(X15,X16,X17) | sP4(X18,X15) | p(s(bool,X18)) | sP4(X19,sK8)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f12387,f2514])).
fof(f18047,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP2(X5,X6,X7) | sP4(X8,X5) | p(s(bool,X8)) | sP4(sK8,X9)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f12387,f2631])).
fof(f18005,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP2(X25,X26,X27) | sP4(X28,X25) | p(s(bool,X29)) | p(s(bool,X28))) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f17997,f307])).
fof(f17997,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP2(X25,X26,X27) | sP4(X28,X25) | p(s(bool,'F_0')) | p(s(bool,X29)) | p(s(bool,X28))) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f12386,f269])).
fof(f18004,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP2(X20,X21,X22) | sP4(X23,X20) | p(s(bool,X24)) | p(s(bool,X23))) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f17996,f307])).
fof(f17996,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP2(X20,X21,X22) | sP4(X23,X20) | p(s(bool,'F_0')) | p(s(bool,X24)) | p(s(bool,X23))) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f12386,f273])).
fof(f17994,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP2(X10,X11,X12) | sP4(X13,X10) | p(s(bool,X13)) | sP4(X14,sK8)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f12386,f2515])).
fof(f17992,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP2(X0,X1,X2) | sP4(X3,X0) | p(s(bool,X3)) | sP4(sK8,X4)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f12386,f2632])).
fof(f17950,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP2(X25,X26,X27) | sP4(X25,X28) | p(s(bool,X28)) | p(s(bool,X29))) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f17943,f307])).
fof(f17943,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP2(X25,X26,X27) | sP4(X25,X28) | p(s(bool,'F_0')) | p(s(bool,X28)) | p(s(bool,X29))) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f12379,f269])).
fof(f17949,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP2(X20,X21,X22) | sP4(X20,X23) | p(s(bool,X23)) | p(s(bool,X24))) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f17942,f307])).
fof(f17942,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP2(X20,X21,X22) | sP4(X20,X23) | p(s(bool,'F_0')) | p(s(bool,X23)) | p(s(bool,X24))) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f12379,f273])).
fof(f17941,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP2(X15,X16,X17) | sP4(X15,X18) | p(s(bool,X18)) | sP4(X19,sK8)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f12379,f2514])).
fof(f17939,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP2(X5,X6,X7) | sP4(X5,X8) | p(s(bool,X8)) | sP4(sK8,X9)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f12379,f2631])).
fof(f17897,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP2(X25,X26,X27) | sP4(X25,X28) | p(s(bool,X29)) | p(s(bool,X28))) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f17889,f307])).
fof(f17889,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP2(X25,X26,X27) | sP4(X25,X28) | p(s(bool,'F_0')) | p(s(bool,X29)) | p(s(bool,X28))) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f12378,f269])).
fof(f17896,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP2(X20,X21,X22) | sP4(X20,X23) | p(s(bool,X24)) | p(s(bool,X23))) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f17888,f307])).
fof(f17888,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP2(X20,X21,X22) | sP4(X20,X23) | p(s(bool,'F_0')) | p(s(bool,X24)) | p(s(bool,X23))) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f12378,f273])).
fof(f17886,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP2(X10,X11,X12) | sP4(X10,X13) | p(s(bool,X13)) | sP4(X14,sK8)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f12378,f2515])).
fof(f17884,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP2(X0,X1,X2) | sP4(X0,X3) | p(s(bool,X3)) | sP4(sK8,X4)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f12378,f2632])).
fof(f17843,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP2(X25,X26,X27) | sP4(X28,X26) | p(s(bool,X28)) | p(s(bool,X29))) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f17838,f307])).
fof(f17838,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP2(X25,X26,X27) | sP4(X28,X26) | p(s(bool,'F_0')) | p(s(bool,X28)) | p(s(bool,X29))) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f11816,f269])).
fof(f17842,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP2(X20,X21,X22) | sP4(X23,X21) | p(s(bool,X23)) | p(s(bool,X24))) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f17837,f307])).
fof(f17837,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP2(X20,X21,X22) | sP4(X23,X21) | p(s(bool,'F_0')) | p(s(bool,X23)) | p(s(bool,X24))) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f11816,f273])).
fof(f17836,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP2(X15,X16,X17) | sP4(X18,X16) | p(s(bool,X18)) | sP4(X19,sK8)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f11816,f2514])).
fof(f17834,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP2(X5,X6,X7) | sP4(X8,X6) | p(s(bool,X8)) | sP4(sK8,X9)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f11816,f2631])).
fof(f17792,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP2(X25,X26,X27) | sP4(X28,X26) | p(s(bool,X29)) | p(s(bool,X28))) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f17786,f307])).
fof(f17786,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP2(X25,X26,X27) | sP4(X28,X26) | p(s(bool,'F_0')) | p(s(bool,X29)) | p(s(bool,X28))) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f11815,f269])).
fof(f17791,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP2(X20,X21,X22) | sP4(X23,X21) | p(s(bool,X24)) | p(s(bool,X23))) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f17785,f307])).
fof(f17785,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP2(X20,X21,X22) | sP4(X23,X21) | p(s(bool,'F_0')) | p(s(bool,X24)) | p(s(bool,X23))) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f11815,f273])).
fof(f17783,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP2(X10,X11,X12) | sP4(X13,X11) | p(s(bool,X13)) | sP4(X14,sK8)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f11815,f2515])).
fof(f17781,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP2(X0,X1,X2) | sP4(X3,X1) | p(s(bool,X3)) | sP4(sK8,X4)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f11815,f2632])).
fof(f17740,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP2(X25,X26,X27) | sP4(X26,X28) | p(s(bool,X28)) | p(s(bool,X29))) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f17735,f307])).
fof(f17735,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP2(X25,X26,X27) | sP4(X26,X28) | p(s(bool,'F_0')) | p(s(bool,X28)) | p(s(bool,X29))) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f11808,f269])).
fof(f17739,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP2(X20,X21,X22) | sP4(X21,X23) | p(s(bool,X23)) | p(s(bool,X24))) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f17734,f307])).
fof(f17734,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP2(X20,X21,X22) | sP4(X21,X23) | p(s(bool,'F_0')) | p(s(bool,X23)) | p(s(bool,X24))) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f11808,f273])).
fof(f17733,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP2(X15,X16,X17) | sP4(X16,X18) | p(s(bool,X18)) | sP4(X19,sK8)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f11808,f2514])).
fof(f17731,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP2(X5,X6,X7) | sP4(X6,X8) | p(s(bool,X8)) | sP4(sK8,X9)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f11808,f2631])).
fof(f17681,plain,(
  ( ! [X25] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | p(s(bool,X25))) ) | $spl4),
  inference(subsumption_resolution,[],[f17636,f322])).
fof(f17636,plain,(
  ( ! [X25] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | p(s(bool,X25)) | s(bool,'T_0') = s(bool,X25)) ) | $spl4),
  inference(superposition,[],[f17542,f358])).
fof(f17542,plain,(
  ( ! [X12] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X12))) | p(s(bool,X12))) ) | $spl4),
  inference(forward_demodulation,[],[f1677,f1649])).
fof(f1649,plain,(
  s(bool,'T_0') = s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7))) | $spl4),
  inference(resolution,[],[f369,f323])).
fof(f1677,plain,(
  ( ! [X12] : (p(s(bool,X12)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X12))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))))) ) | $spl4),
  inference(superposition,[],[f369,f351])).
fof(f17675,plain,(
  ( ! [X21] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | p(s(bool,X21))) ) | $spl4),
  inference(duplicate_literal_removal,[],[f17632])).
fof(f17632,plain,(
  ( ! [X21] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | p(s(bool,X21)) | p(s(bool,X21))) ) | $spl4),
  inference(superposition,[],[f17542,f321])).
fof(f17676,plain,(
  ( ! [X20] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | p(s(bool,X20))) ) | $spl4),
  inference(duplicate_literal_removal,[],[f17631])).
fof(f17631,plain,(
  ( ! [X20] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | p(s(bool,X20)) | p(s(bool,X20))) ) | $spl4),
  inference(superposition,[],[f17542,f319])).
fof(f17593,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP2(X25,X26,X27) | sP4(X26,X28) | p(s(bool,X29)) | p(s(bool,X28))) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f17588,f307])).
fof(f17588,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP2(X25,X26,X27) | sP4(X26,X28) | p(s(bool,'F_0')) | p(s(bool,X29)) | p(s(bool,X28))) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f11807,f269])).
fof(f17592,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP2(X20,X21,X22) | sP4(X21,X23) | p(s(bool,X24)) | p(s(bool,X23))) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f17587,f307])).
fof(f17587,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP2(X20,X21,X22) | sP4(X21,X23) | p(s(bool,'F_0')) | p(s(bool,X24)) | p(s(bool,X23))) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f11807,f273])).
fof(f17585,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP2(X10,X11,X12) | sP4(X11,X13) | p(s(bool,X13)) | sP4(X14,sK8)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f11807,f2515])).
fof(f17583,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP2(X0,X1,X2) | sP4(X1,X3) | p(s(bool,X3)) | sP4(sK8,X4)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f11807,f2632])).
fof(f17482,plain,(
  ( ! [X300,X301] : (sP5(X300,X301) | sP4(X300,'F_0')) ) | $spl390),
  inference(cnf_transformation,[],[f17482_D])).
fof(f17482_D,plain,(
  ( ! [X300,X301] : (sP5(X300,X301) | sP4(X300,'F_0')) ) <=> ~$spl390),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl390])])).
fof(f17439,plain,(
  ( ! [X308,X310,X309,X311] : (sP5(X308,X309) | sP4(X308,'F_0') | p(s(bool,X310)) | p(s(bool,X311)) | ~sP5(X310,X311)) ) | $spl3),
  inference(resolution,[],[f8354,f269])).
fof(f8354,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,X6)) | sP5(X5,X7) | sP4(X5,X6)) ) | $spl3),
  inference(subsumption_resolution,[],[f8317,f361])).
fof(f8317,plain,(
  ( ! [X6,X7,X5] : (sP4(X5,X6) | sP5(X5,X7) | ~p(s(bool,'T_0')) | ~p(s(bool,X6))) ) | $spl3),
  inference(resolution,[],[f8127,f254])).
fof(f17437,plain,(
  ( ! [X300,X302,X301,X303] : (sP5(X300,X301) | sP4(X300,'F_0') | p(s(bool,X302)) | p(s(bool,X303)) | ~sP5(X302,X303)) ) | $spl3),
  inference(resolution,[],[f8354,f273])).
fof(f17476,plain,(
  ( ! [X4,X2,X0,X1] : (sP5(X0,X1) | s(bool,X2) != s(bool,X4) | ~p(s(bool,X4)) | sP4(X0,X2)) ) | $spl388),
  inference(cnf_transformation,[],[f17476_D])).
fof(f17476_D,plain,(
  ( ! [X4,X2,X0,X1] : (sP5(X0,X1) | s(bool,X2) != s(bool,X4) | ~p(s(bool,X4)) | sP4(X0,X2)) ) <=> ~$spl388),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl388])])).
fof(f17414,plain,(
  ( ! [X237,X235,X238,X236,X234] : (sP5(X234,X235) | sP4(X234,X236) | p(s(bool,X237)) | ~p(s(bool,X238)) | s(bool,X236) != s(bool,X238)) ) | $spl3),
  inference(resolution,[],[f8354,f360])).
fof(f17369,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (sP5(X42,X43) | sP4(X42,X44) | ~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X44) != s(bool,X46) | s(bool,X45) != s(bool,X47)) ) | $spl3),
  inference(resolution,[],[f8354,f223])).
fof(f17368,plain,(
  ( ! [X39,X37,X41,X38,X36,X40] : (sP5(X36,X37) | sP4(X36,X38) | p(s(bool,X39)) | ~p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X38) != s(bool,X40)) ) | $spl3),
  inference(resolution,[],[f8354,f222])).
fof(f17366,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (sP5(X24,X25) | sP4(X24,X26) | p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,X29) | s(bool,X26) != s(bool,X28)) ) | $spl3),
  inference(resolution,[],[f8354,f221])).
fof(f17364,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (sP5(X12,X13) | sP4(X12,X14) | ~p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,X17)) ) | $spl3),
  inference(resolution,[],[f8354,f219])).
fof(f17362,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP5(X0,X1) | sP4(X0,X2) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X2) != s(bool,X4)) ) | $spl3),
  inference(resolution,[],[f8354,f217])).
fof(f17208,plain,(
  ( ! [X300,X301] : (sP5(X300,X301) | sP4(X301,'F_0')) ) | $spl386),
  inference(cnf_transformation,[],[f17208_D])).
fof(f17208_D,plain,(
  ( ! [X300,X301] : (sP5(X300,X301) | sP4(X301,'F_0')) ) <=> ~$spl386),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl386])])).
fof(f17165,plain,(
  ( ! [X308,X310,X309,X311] : (sP5(X308,X309) | sP4(X309,'F_0') | p(s(bool,X310)) | p(s(bool,X311)) | ~sP5(X310,X311)) ) | $spl3),
  inference(resolution,[],[f8287,f269])).
fof(f8287,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,X6)) | sP5(X7,X5) | sP4(X5,X6)) ) | $spl3),
  inference(subsumption_resolution,[],[f8250,f361])).
fof(f8250,plain,(
  ( ! [X6,X7,X5] : (sP4(X5,X6) | sP5(X7,X5) | ~p(s(bool,'T_0')) | ~p(s(bool,X6))) ) | $spl3),
  inference(resolution,[],[f8126,f254])).
fof(f17163,plain,(
  ( ! [X300,X302,X301,X303] : (sP5(X300,X301) | sP4(X301,'F_0') | p(s(bool,X302)) | p(s(bool,X303)) | ~sP5(X302,X303)) ) | $spl3),
  inference(resolution,[],[f8287,f273])).
fof(f17202,plain,(
  ( ! [X4,X2,X0,X1] : (sP5(X0,X1) | s(bool,X2) != s(bool,X4) | ~p(s(bool,X4)) | sP4(X1,X2)) ) | $spl384),
  inference(cnf_transformation,[],[f17202_D])).
fof(f17202_D,plain,(
  ( ! [X4,X2,X0,X1] : (sP5(X0,X1) | s(bool,X2) != s(bool,X4) | ~p(s(bool,X4)) | sP4(X1,X2)) ) <=> ~$spl384),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl384])])).
fof(f17140,plain,(
  ( ! [X237,X235,X238,X236,X234] : (sP5(X234,X235) | sP4(X235,X236) | p(s(bool,X237)) | ~p(s(bool,X238)) | s(bool,X236) != s(bool,X238)) ) | $spl3),
  inference(resolution,[],[f8287,f360])).
fof(f17095,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (sP5(X42,X43) | sP4(X43,X44) | ~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X44) != s(bool,X46) | s(bool,X45) != s(bool,X47)) ) | $spl3),
  inference(resolution,[],[f8287,f223])).
fof(f17094,plain,(
  ( ! [X39,X37,X41,X38,X36,X40] : (sP5(X36,X37) | sP4(X37,X38) | p(s(bool,X39)) | ~p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X38) != s(bool,X40)) ) | $spl3),
  inference(resolution,[],[f8287,f222])).
fof(f17092,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (sP5(X24,X25) | sP4(X25,X26) | p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,X29) | s(bool,X26) != s(bool,X28)) ) | $spl3),
  inference(resolution,[],[f8287,f221])).
fof(f17090,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (sP5(X12,X13) | sP4(X13,X14) | ~p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,X17)) ) | $spl3),
  inference(resolution,[],[f8287,f219])).
fof(f17088,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP5(X0,X1) | sP4(X1,X2) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X2) != s(bool,X4)) ) | $spl3),
  inference(resolution,[],[f8287,f217])).
fof(f16987,plain,(
  ( ! [X308,X310,X309,X311] : (sP5(X308,X309) | sP4('F_0',X308) | p(s(bool,X310)) | p(s(bool,X311)) | ~sP5(X310,X311)) ) | $spl3),
  inference(resolution,[],[f8014,f269])).
fof(f8014,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,X5)) | sP5(X6,X7) | sP4(X5,X6)) ) | $spl3),
  inference(subsumption_resolution,[],[f7977,f361])).
fof(f7977,plain,(
  ( ! [X6,X7,X5] : (sP4(X5,X6) | sP5(X6,X7) | ~p(s(bool,'T_0')) | ~p(s(bool,X5))) ) | $spl3),
  inference(resolution,[],[f7456,f254])).
fof(f16985,plain,(
  ( ! [X300,X302,X301,X303] : (sP5(X300,X301) | sP4('F_0',X300) | p(s(bool,X302)) | p(s(bool,X303)) | ~sP5(X302,X303)) ) | $spl3),
  inference(resolution,[],[f8014,f273])).
fof(f16962,plain,(
  ( ! [X237,X235,X238,X236,X234] : (sP5(X234,X235) | sP4(X236,X234) | p(s(bool,X237)) | ~p(s(bool,X238)) | s(bool,X236) != s(bool,X238)) ) | $spl3),
  inference(resolution,[],[f8014,f360])).
fof(f16917,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (sP5(X42,X43) | sP4(X44,X42) | ~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X44) != s(bool,X46) | s(bool,X45) != s(bool,X47)) ) | $spl3),
  inference(resolution,[],[f8014,f223])).
fof(f16916,plain,(
  ( ! [X39,X37,X41,X38,X36,X40] : (sP5(X36,X37) | sP4(X38,X36) | p(s(bool,X39)) | ~p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X38) != s(bool,X40)) ) | $spl3),
  inference(resolution,[],[f8014,f222])).
fof(f16914,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (sP5(X24,X25) | sP4(X26,X24) | p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,X29) | s(bool,X26) != s(bool,X28)) ) | $spl3),
  inference(resolution,[],[f8014,f221])).
fof(f16912,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (sP5(X12,X13) | sP4(X14,X12) | ~p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,X17)) ) | $spl3),
  inference(resolution,[],[f8014,f219])).
fof(f16910,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP5(X0,X1) | sP4(X2,X0) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X2) != s(bool,X4)) ) | $spl3),
  inference(resolution,[],[f8014,f217])).
fof(f16809,plain,(
  ( ! [X308,X310,X309,X311] : (sP5(X308,X309) | sP4('F_0',X309) | p(s(bool,X310)) | p(s(bool,X311)) | ~sP5(X310,X311)) ) | $spl3),
  inference(resolution,[],[f7956,f269])).
fof(f7956,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,X5)) | sP5(X7,X6) | sP4(X5,X6)) ) | $spl3),
  inference(subsumption_resolution,[],[f7919,f361])).
fof(f7919,plain,(
  ( ! [X6,X7,X5] : (sP4(X5,X6) | sP5(X7,X6) | ~p(s(bool,'T_0')) | ~p(s(bool,X5))) ) | $spl3),
  inference(resolution,[],[f7455,f254])).
fof(f16807,plain,(
  ( ! [X300,X302,X301,X303] : (sP5(X300,X301) | sP4('F_0',X301) | p(s(bool,X302)) | p(s(bool,X303)) | ~sP5(X302,X303)) ) | $spl3),
  inference(resolution,[],[f7956,f273])).
fof(f16784,plain,(
  ( ! [X237,X235,X238,X236,X234] : (sP5(X234,X235) | sP4(X236,X235) | p(s(bool,X237)) | ~p(s(bool,X238)) | s(bool,X236) != s(bool,X238)) ) | $spl3),
  inference(resolution,[],[f7956,f360])).
fof(f16739,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (sP5(X42,X43) | sP4(X44,X43) | ~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X44) != s(bool,X46) | s(bool,X45) != s(bool,X47)) ) | $spl3),
  inference(resolution,[],[f7956,f223])).
fof(f16738,plain,(
  ( ! [X39,X37,X41,X38,X36,X40] : (sP5(X36,X37) | sP4(X38,X37) | p(s(bool,X39)) | ~p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X38) != s(bool,X40)) ) | $spl3),
  inference(resolution,[],[f7956,f222])).
fof(f16736,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (sP5(X24,X25) | sP4(X26,X25) | p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,X29) | s(bool,X26) != s(bool,X28)) ) | $spl3),
  inference(resolution,[],[f7956,f221])).
fof(f16734,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (sP5(X12,X13) | sP4(X14,X13) | ~p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,X17)) ) | $spl3),
  inference(resolution,[],[f7956,f219])).
fof(f16732,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP5(X0,X1) | sP4(X2,X1) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X2) != s(bool,X4)) ) | $spl3),
  inference(resolution,[],[f7956,f217])).
fof(f16721,plain,(
  ( ! [X24] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | ~p(s(bool,X24))) ) | $spl3),
  inference(subsumption_resolution,[],[f16677,f318])).
fof(f318,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f198])).
fof(f16677,plain,(
  ( ! [X24] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | ~p(s(bool,X24)) | s(bool,'F_0') = s(bool,X24)) ) | $spl3),
  inference(superposition,[],[f16073,f358])).
fof(f16073,plain,(
  ( ! [X12] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X12))) | ~p(s(bool,X12))) ) | $spl3),
  inference(forward_demodulation,[],[f1586,f1561])).
fof(f1586,plain,(
  ( ! [X12] : (~p(s(bool,X12)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,sK8))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X12)))) ) | $spl3),
  inference(superposition,[],[f367,f351])).
fof(f16713,plain,(
  ( ! [X23] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | ~p(s(bool,X23))) ) | $spl3),
  inference(duplicate_literal_removal,[],[f16676])).
fof(f16676,plain,(
  ( ! [X23] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | ~p(s(bool,X23)) | ~p(s(bool,X23))) ) | $spl3),
  inference(superposition,[],[f16073,f325])).
fof(f16714,plain,(
  ( ! [X22] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | ~p(s(bool,X22))) ) | $spl3),
  inference(duplicate_literal_removal,[],[f16675])).
fof(f16675,plain,(
  ( ! [X22] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | ~p(s(bool,X22)) | ~p(s(bool,X22))) ) | $spl3),
  inference(superposition,[],[f16073,f323])).
fof(f8585,plain,(
  ( ! [X219] : (sP4(X219,'T_0')) ) | $spl210),
  inference(cnf_transformation,[],[f8585_D])).
fof(f8585_D,plain,(
  ( ! [X219] : (sP4(X219,'T_0')) ) <=> ~$spl210),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl210])])).
fof(f16580,plain,(
  ( ! [X6,X4,X7,X5] : (sP4(X4,'T_0') | sP3(X6,X7,X5) | ~p(s(bool,X5)) | ~sP2(X6,X5,X7) | ~sP1(X7,X5,X6)) ) | $spl3),
  inference(subsumption_resolution,[],[f16552,f4521])).
fof(f4521,plain,(
  ( ! [X76,X77] : (~p(s(bool,X77)) | sP4(X76,'T_0') | ~sP4(X77,X76)) ) | $spl3),
  inference(resolution,[],[f4459,f254])).
fof(f16552,plain,(
  ( ! [X6,X4,X7,X5] : (sP4(X4,'T_0') | sP4(X5,X4) | sP3(X6,X7,X5) | ~p(s(bool,X5)) | ~sP2(X6,X5,X7) | ~sP1(X7,X5,X6)) ) | $spl3),
  inference(resolution,[],[f9816,f228])).
fof(f9816,plain,(
  ( ! [X87,X85,X86,X84] : (sP4(X84,'T_0') | sP4(X85,X84) | sP0(X85,X86,X87)) ) | $spl3),
  inference(resolution,[],[f4523,f242])).
fof(f16070,plain,(
  ( ! [X6,X4,X7,X5] : (sP4(X4,'T_0') | sP3(X6,X7,X5) | ~p(s(bool,X5)) | ~sP2(X6,X5,X7) | ~sP1(X7,X5,X6)) ) | $spl3),
  inference(subsumption_resolution,[],[f16043,f4520])).
fof(f16043,plain,(
  ( ! [X6,X4,X7,X5] : (sP4(X4,'T_0') | sP4(X4,X5) | sP3(X6,X7,X5) | ~p(s(bool,X5)) | ~sP2(X6,X5,X7) | ~sP1(X7,X5,X6)) ) | $spl3),
  inference(resolution,[],[f9640,f228])).
fof(f9640,plain,(
  ( ! [X87,X85,X86,X84] : (sP4(X84,'T_0') | sP4(X84,X85) | sP0(X85,X86,X87)) ) | $spl3),
  inference(resolution,[],[f4522,f242])).
fof(f4522,plain,(
  ( ! [X78,X79] : (sP4(X78,'T_0') | sP4(X78,X79) | p(s(bool,X79))) ) | $spl3),
  inference(resolution,[],[f4459,f255])).
fof(f14472,plain,(
  ( ! [X592,X590,X591,X587] : (sP4(X587,'T_0') | sP0(X590,X591,X592) | sP4(X590,X587)) ) | $spl366),
  inference(cnf_transformation,[],[f14472_D])).
fof(f14472_D,plain,(
  ( ! [X592,X590,X591,X587] : (sP4(X587,'T_0') | sP0(X590,X591,X592) | sP4(X590,X587)) ) <=> ~$spl366),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl366])])).
fof(f15215,plain,(
  ( ! [X592,X588,X590,X589,X591,X587] : (sP4(X587,'T_0') | sP1('T_0',X588,X589) | sP4(X590,X587) | sP0(X590,X591,X592)) ) | $spl3),
  inference(resolution,[],[f9190,f8115])).
fof(f8115,plain,(
  ( ! [X87,X85,X86,X84] : (sP4('T_0',X84) | sP4(X85,X84) | sP0(X85,X86,X87)) ) | $spl3),
  inference(resolution,[],[f3990,f242])).
fof(f9190,plain,(
  ( ! [X74,X72,X75,X73] : (~sP4(X73,X72) | sP4(X72,'T_0') | sP1(X73,X74,X75)) ) | $spl3),
  inference(resolution,[],[f4521,f238])).
fof(f14471,plain,(
  ( ! [X584,X586,X585,X581] : (sP4(X581,'T_0') | sP0(X585,X584,X586) | sP4(X584,X581)) ) | $spl364),
  inference(cnf_transformation,[],[f14471_D])).
fof(f14471_D,plain,(
  ( ! [X584,X586,X585,X581] : (sP4(X581,'T_0') | sP0(X585,X584,X586) | sP4(X584,X581)) ) <=> ~$spl364),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl364])])).
fof(f15214,plain,(
  ( ! [X584,X586,X582,X585,X581,X583] : (sP4(X581,'T_0') | sP1('T_0',X582,X583) | sP4(X584,X581) | sP0(X585,X584,X586)) ) | $spl3),
  inference(resolution,[],[f9190,f8114])).
fof(f8114,plain,(
  ( ! [X80,X83,X81,X82] : (sP4('T_0',X80) | sP4(X81,X80) | sP0(X82,X81,X83)) ) | $spl3),
  inference(resolution,[],[f3990,f241])).
fof(f241,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f161])).
fof(f14470,plain,(
  ( ! [X575,X580,X578,X579] : (sP4(X575,'T_0') | sP0(X579,X580,X578) | sP4(X578,X575)) ) | $spl362),
  inference(cnf_transformation,[],[f14470_D])).
fof(f14470_D,plain,(
  ( ! [X575,X580,X578,X579] : (sP4(X575,'T_0') | sP0(X579,X580,X578) | sP4(X578,X575)) ) <=> ~$spl362),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl362])])).
fof(f15213,plain,(
  ( ! [X575,X580,X576,X578,X577,X579] : (sP4(X575,'T_0') | sP1('T_0',X576,X577) | sP4(X578,X575) | sP0(X579,X580,X578)) ) | $spl3),
  inference(resolution,[],[f9190,f8113])).
fof(f8113,plain,(
  ( ! [X78,X76,X79,X77] : (sP4('T_0',X76) | sP4(X77,X76) | sP0(X78,X79,X77)) ) | $spl3),
  inference(resolution,[],[f3990,f240])).
fof(f240,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f161])).
fof(f14469,plain,(
  ( ! [X573,X569,X572,X574] : (sP4(X569,'T_0') | sP1(X573,X574,X572) | sP4(X572,X569)) ) | $spl360),
  inference(cnf_transformation,[],[f14469_D])).
fof(f14469_D,plain,(
  ( ! [X573,X569,X572,X574] : (sP4(X569,'T_0') | sP1(X573,X574,X572) | sP4(X572,X569)) ) <=> ~$spl360),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl360])])).
fof(f15212,plain,(
  ( ! [X573,X569,X571,X572,X574,X570] : (sP4(X569,'T_0') | sP1('T_0',X570,X571) | sP4(X572,X569) | sP1(X573,X574,X572)) ) | $spl3),
  inference(resolution,[],[f9190,f8112])).
fof(f8112,plain,(
  ( ! [X74,X72,X75,X73] : (sP4('T_0',X72) | sP4(X73,X72) | sP1(X74,X75,X73)) ) | $spl3),
  inference(resolution,[],[f3990,f236])).
fof(f236,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f159])).
fof(f14468,plain,(
  ( ! [X567,X563,X568,X566] : (sP4(X563,'T_0') | sP2(X567,X568,X566) | sP4(X566,X563)) ) | $spl358),
  inference(cnf_transformation,[],[f14468_D])).
fof(f14468_D,plain,(
  ( ! [X567,X563,X568,X566] : (sP4(X563,'T_0') | sP2(X567,X568,X566) | sP4(X566,X563)) ) <=> ~$spl358),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl358])])).
fof(f15211,plain,(
  ( ! [X565,X567,X563,X568,X564,X566] : (sP4(X563,'T_0') | sP1('T_0',X564,X565) | sP4(X566,X563) | sP2(X567,X568,X566)) ) | $spl3),
  inference(resolution,[],[f9190,f8109])).
fof(f8109,plain,(
  ( ! [X61,X62,X60,X63] : (sP4('T_0',X60) | sP4(X61,X60) | sP2(X62,X63,X61)) ) | $spl3),
  inference(resolution,[],[f3990,f232])).
fof(f232,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f156])).
fof(f14467,plain,(
  ( ! [X540,X545,X543,X544] : (sP4(X540,'T_0') | sP0(X543,X544,X545) | sP4(X540,X543)) ) | $spl356),
  inference(cnf_transformation,[],[f14467_D])).
fof(f14467_D,plain,(
  ( ! [X540,X545,X543,X544] : (sP4(X540,'T_0') | sP0(X543,X544,X545) | sP4(X540,X543)) ) <=> ~$spl356),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl356])])).
fof(f15206,plain,(
  ( ! [X540,X542,X545,X541,X543,X544] : (sP4(X540,'T_0') | sP1('T_0',X541,X542) | sP4(X540,X543) | sP0(X543,X544,X545)) ) | $spl3),
  inference(resolution,[],[f9190,f7444])).
fof(f7444,plain,(
  ( ! [X87,X85,X86,X84] : (sP4('T_0',X84) | sP4(X84,X85) | sP0(X85,X86,X87)) ) | $spl3),
  inference(resolution,[],[f3989,f242])).
fof(f14466,plain,(
  ( ! [X538,X534,X537,X539] : (sP4(X534,'T_0') | sP0(X538,X537,X539) | sP4(X534,X537)) ) | $spl354),
  inference(cnf_transformation,[],[f14466_D])).
fof(f14466_D,plain,(
  ( ! [X538,X534,X537,X539] : (sP4(X534,'T_0') | sP0(X538,X537,X539) | sP4(X534,X537)) ) <=> ~$spl354),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl354])])).
fof(f15205,plain,(
  ( ! [X536,X538,X534,X537,X539,X535] : (sP4(X534,'T_0') | sP1('T_0',X535,X536) | sP4(X534,X537) | sP0(X538,X537,X539)) ) | $spl3),
  inference(resolution,[],[f9190,f7443])).
fof(f7443,plain,(
  ( ! [X80,X83,X81,X82] : (sP4('T_0',X80) | sP4(X80,X81) | sP0(X82,X81,X83)) ) | $spl3),
  inference(resolution,[],[f3989,f241])).
fof(f14465,plain,(
  ( ! [X532,X528,X533,X531] : (sP4(X528,'T_0') | sP0(X532,X533,X531) | sP4(X528,X531)) ) | $spl352),
  inference(cnf_transformation,[],[f14465_D])).
fof(f14465_D,plain,(
  ( ! [X532,X528,X533,X531] : (sP4(X528,'T_0') | sP0(X532,X533,X531) | sP4(X528,X531)) ) <=> ~$spl352),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl352])])).
fof(f15204,plain,(
  ( ! [X532,X528,X530,X533,X529,X531] : (sP4(X528,'T_0') | sP1('T_0',X529,X530) | sP4(X528,X531) | sP0(X532,X533,X531)) ) | $spl3),
  inference(resolution,[],[f9190,f7442])).
fof(f7442,plain,(
  ( ! [X78,X76,X79,X77] : (sP4('T_0',X76) | sP4(X76,X77) | sP0(X78,X79,X77)) ) | $spl3),
  inference(resolution,[],[f3989,f240])).
fof(f14464,plain,(
  ( ! [X526,X522,X525,X527] : (sP4(X522,'T_0') | sP1(X526,X527,X525) | sP4(X522,X525)) ) | $spl350),
  inference(cnf_transformation,[],[f14464_D])).
fof(f14464_D,plain,(
  ( ! [X526,X522,X525,X527] : (sP4(X522,'T_0') | sP1(X526,X527,X525) | sP4(X522,X525)) ) <=> ~$spl350),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl350])])).
fof(f15203,plain,(
  ( ! [X524,X526,X522,X525,X527,X523] : (sP4(X522,'T_0') | sP1('T_0',X523,X524) | sP4(X522,X525) | sP1(X526,X527,X525)) ) | $spl3),
  inference(resolution,[],[f9190,f7441])).
fof(f7441,plain,(
  ( ! [X74,X72,X75,X73] : (sP4('T_0',X72) | sP4(X72,X73) | sP1(X74,X75,X73)) ) | $spl3),
  inference(resolution,[],[f3989,f236])).
fof(f14463,plain,(
  ( ! [X520,X516,X521,X519] : (sP4(X516,'T_0') | sP2(X520,X521,X519) | sP4(X516,X519)) ) | $spl348),
  inference(cnf_transformation,[],[f14463_D])).
fof(f14463_D,plain,(
  ( ! [X520,X516,X521,X519] : (sP4(X516,'T_0') | sP2(X520,X521,X519) | sP4(X516,X519)) ) <=> ~$spl348),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl348])])).
fof(f15202,plain,(
  ( ! [X520,X516,X518,X521,X517,X519] : (sP4(X516,'T_0') | sP1('T_0',X517,X518) | sP4(X516,X519) | sP2(X520,X521,X519)) ) | $spl3),
  inference(resolution,[],[f9190,f7438])).
fof(f7438,plain,(
  ( ! [X61,X62,X60,X63] : (sP4('T_0',X60) | sP4(X60,X61) | sP2(X62,X63,X61)) ) | $spl3),
  inference(resolution,[],[f3989,f232])).
fof(f14992,plain,(
  ( ! [X592,X588,X590,X589,X591,X587] : (sP4(X587,'T_0') | sP1(X588,'T_0',X589) | sP4(X590,X587) | sP0(X590,X591,X592)) ) | $spl3),
  inference(resolution,[],[f9189,f8115])).
fof(f9189,plain,(
  ( ! [X70,X68,X71,X69] : (~sP4(X69,X68) | sP4(X68,'T_0') | sP1(X70,X69,X71)) ) | $spl3),
  inference(resolution,[],[f4521,f237])).
fof(f14991,plain,(
  ( ! [X584,X586,X582,X585,X581,X583] : (sP4(X581,'T_0') | sP1(X582,'T_0',X583) | sP4(X584,X581) | sP0(X585,X584,X586)) ) | $spl3),
  inference(resolution,[],[f9189,f8114])).
fof(f14990,plain,(
  ( ! [X575,X580,X576,X578,X577,X579] : (sP4(X575,'T_0') | sP1(X576,'T_0',X577) | sP4(X578,X575) | sP0(X579,X580,X578)) ) | $spl3),
  inference(resolution,[],[f9189,f8113])).
fof(f14989,plain,(
  ( ! [X573,X569,X571,X572,X574,X570] : (sP4(X569,'T_0') | sP1(X570,'T_0',X571) | sP4(X572,X569) | sP1(X573,X574,X572)) ) | $spl3),
  inference(resolution,[],[f9189,f8112])).
fof(f14988,plain,(
  ( ! [X565,X567,X563,X568,X564,X566] : (sP4(X563,'T_0') | sP1(X564,'T_0',X565) | sP4(X566,X563) | sP2(X567,X568,X566)) ) | $spl3),
  inference(resolution,[],[f9189,f8109])).
fof(f14983,plain,(
  ( ! [X540,X542,X545,X541,X543,X544] : (sP4(X540,'T_0') | sP1(X541,'T_0',X542) | sP4(X540,X543) | sP0(X543,X544,X545)) ) | $spl3),
  inference(resolution,[],[f9189,f7444])).
fof(f14982,plain,(
  ( ! [X536,X538,X534,X537,X539,X535] : (sP4(X534,'T_0') | sP1(X535,'T_0',X536) | sP4(X534,X537) | sP0(X538,X537,X539)) ) | $spl3),
  inference(resolution,[],[f9189,f7443])).
fof(f14981,plain,(
  ( ! [X532,X528,X530,X533,X529,X531] : (sP4(X528,'T_0') | sP1(X529,'T_0',X530) | sP4(X528,X531) | sP0(X532,X533,X531)) ) | $spl3),
  inference(resolution,[],[f9189,f7442])).
fof(f14980,plain,(
  ( ! [X524,X526,X522,X525,X527,X523] : (sP4(X522,'T_0') | sP1(X523,'T_0',X524) | sP4(X522,X525) | sP1(X526,X527,X525)) ) | $spl3),
  inference(resolution,[],[f9189,f7441])).
fof(f14979,plain,(
  ( ! [X520,X516,X518,X521,X517,X519] : (sP4(X516,'T_0') | sP1(X517,'T_0',X518) | sP4(X516,X519) | sP2(X520,X521,X519)) ) | $spl3),
  inference(resolution,[],[f9189,f7438])).
fof(f14768,plain,(
  ( ! [X592,X588,X590,X589,X591,X587] : (sP4(X587,'T_0') | sP2('T_0',X588,X589) | sP4(X590,X587) | sP0(X590,X591,X592)) ) | $spl3),
  inference(resolution,[],[f9187,f8115])).
fof(f9187,plain,(
  ( ! [X61,X62,X60,X63] : (~sP4(X61,X60) | sP4(X60,'T_0') | sP2(X61,X62,X63)) ) | $spl3),
  inference(resolution,[],[f4521,f234])).
fof(f14767,plain,(
  ( ! [X584,X586,X582,X585,X581,X583] : (sP4(X581,'T_0') | sP2('T_0',X582,X583) | sP4(X584,X581) | sP0(X585,X584,X586)) ) | $spl3),
  inference(resolution,[],[f9187,f8114])).
fof(f14766,plain,(
  ( ! [X575,X580,X576,X578,X577,X579] : (sP4(X575,'T_0') | sP2('T_0',X576,X577) | sP4(X578,X575) | sP0(X579,X580,X578)) ) | $spl3),
  inference(resolution,[],[f9187,f8113])).
fof(f14765,plain,(
  ( ! [X573,X569,X571,X572,X574,X570] : (sP4(X569,'T_0') | sP2('T_0',X570,X571) | sP4(X572,X569) | sP1(X573,X574,X572)) ) | $spl3),
  inference(resolution,[],[f9187,f8112])).
fof(f14764,plain,(
  ( ! [X565,X567,X563,X568,X564,X566] : (sP4(X563,'T_0') | sP2('T_0',X564,X565) | sP4(X566,X563) | sP2(X567,X568,X566)) ) | $spl3),
  inference(resolution,[],[f9187,f8109])).
fof(f14759,plain,(
  ( ! [X540,X542,X545,X541,X543,X544] : (sP4(X540,'T_0') | sP2('T_0',X541,X542) | sP4(X540,X543) | sP0(X543,X544,X545)) ) | $spl3),
  inference(resolution,[],[f9187,f7444])).
fof(f14758,plain,(
  ( ! [X536,X538,X534,X537,X539,X535] : (sP4(X534,'T_0') | sP2('T_0',X535,X536) | sP4(X534,X537) | sP0(X538,X537,X539)) ) | $spl3),
  inference(resolution,[],[f9187,f7443])).
fof(f14757,plain,(
  ( ! [X532,X528,X530,X533,X529,X531] : (sP4(X528,'T_0') | sP2('T_0',X529,X530) | sP4(X528,X531) | sP0(X532,X533,X531)) ) | $spl3),
  inference(resolution,[],[f9187,f7442])).
fof(f14756,plain,(
  ( ! [X524,X526,X522,X525,X527,X523] : (sP4(X522,'T_0') | sP2('T_0',X523,X524) | sP4(X522,X525) | sP1(X526,X527,X525)) ) | $spl3),
  inference(resolution,[],[f9187,f7441])).
fof(f14755,plain,(
  ( ! [X520,X516,X518,X521,X517,X519] : (sP4(X516,'T_0') | sP2('T_0',X517,X518) | sP4(X516,X519) | sP2(X520,X521,X519)) ) | $spl3),
  inference(resolution,[],[f9187,f7438])).
fof(f14637,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f14566,f458])).
fof(f14566,plain,(
  p(s(bool,'F_0')) | $spl16),
  inference(resolution,[],[f458,f335])).
fof(f14636,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f14565,f458])).
fof(f14565,plain,(
  p(s(bool,'F_0')) | $spl16),
  inference(resolution,[],[f458,f333])).
fof(f14635,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f14561,f458])).
fof(f14561,plain,(
  ( ! [X137] : (p(s(bool,X137))) ) | $spl16),
  inference(resolution,[],[f458,f300])).
fof(f14634,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f14560,f458])).
fof(f14560,plain,(
  ( ! [X136] : (p(s(bool,X136))) ) | $spl16),
  inference(resolution,[],[f458,f297])).
fof(f14633,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f14559,f458])).
fof(f14559,plain,(
  ( ! [X135] : (p(s(bool,X135))) ) | $spl16),
  inference(resolution,[],[f458,f294])).
fof(f14632,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f14557,f458])).
fof(f14557,plain,(
  p(s(bool,'F_0')) | $spl16),
  inference(resolution,[],[f458,f291])).
fof(f14567,plain,(
  $false | $spl16),
  inference(resolution,[],[f458,f337])).
fof(f14568,plain,(
  $false | $spl16),
  inference(resolution,[],[f458,f338])).
fof(f14569,plain,(
  $false | $spl16),
  inference(resolution,[],[f458,f361])).
fof(f14490,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1))) ) | $spl346),
  inference(subsumption_resolution,[],[f14484,f361])).
fof(f14484,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,'T_0')) | p(s(bool,X1))) ) | $spl346),
  inference(resolution,[],[f14462,f231])).
fof(f231,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f156])).
fof(f14377,plain,(
  ( ! [X592,X588,X590,X589,X591,X587] : (sP4(X587,'T_0') | sP2(X588,'T_0',X589) | sP4(X590,X587) | sP0(X590,X591,X592)) ) | $spl3),
  inference(resolution,[],[f9186,f8115])).
fof(f9186,plain,(
  ( ! [X59,X57,X58,X56] : (~sP4(X57,X56) | sP4(X56,'T_0') | sP2(X58,X57,X59)) ) | $spl3),
  inference(resolution,[],[f4521,f233])).
fof(f14376,plain,(
  ( ! [X584,X586,X582,X585,X581,X583] : (sP4(X581,'T_0') | sP2(X582,'T_0',X583) | sP4(X584,X581) | sP0(X585,X584,X586)) ) | $spl3),
  inference(resolution,[],[f9186,f8114])).
fof(f14375,plain,(
  ( ! [X575,X580,X576,X578,X577,X579] : (sP4(X575,'T_0') | sP2(X576,'T_0',X577) | sP4(X578,X575) | sP0(X579,X580,X578)) ) | $spl3),
  inference(resolution,[],[f9186,f8113])).
fof(f14374,plain,(
  ( ! [X573,X569,X571,X572,X574,X570] : (sP4(X569,'T_0') | sP2(X570,'T_0',X571) | sP4(X572,X569) | sP1(X573,X574,X572)) ) | $spl3),
  inference(resolution,[],[f9186,f8112])).
fof(f14373,plain,(
  ( ! [X565,X567,X563,X568,X564,X566] : (sP4(X563,'T_0') | sP2(X564,'T_0',X565) | sP4(X566,X563) | sP2(X567,X568,X566)) ) | $spl3),
  inference(resolution,[],[f9186,f8109])).
fof(f14368,plain,(
  ( ! [X540,X542,X545,X541,X543,X544] : (sP4(X540,'T_0') | sP2(X541,'T_0',X542) | sP4(X540,X543) | sP0(X543,X544,X545)) ) | $spl3),
  inference(resolution,[],[f9186,f7444])).
fof(f14367,plain,(
  ( ! [X536,X538,X534,X537,X539,X535] : (sP4(X534,'T_0') | sP2(X535,'T_0',X536) | sP4(X534,X537) | sP0(X538,X537,X539)) ) | $spl3),
  inference(resolution,[],[f9186,f7443])).
fof(f14366,plain,(
  ( ! [X532,X528,X530,X533,X529,X531] : (sP4(X528,'T_0') | sP2(X529,'T_0',X530) | sP4(X528,X531) | sP0(X532,X533,X531)) ) | $spl3),
  inference(resolution,[],[f9186,f7442])).
fof(f14365,plain,(
  ( ! [X524,X526,X522,X525,X527,X523] : (sP4(X522,'T_0') | sP2(X523,'T_0',X524) | sP4(X522,X525) | sP1(X526,X527,X525)) ) | $spl3),
  inference(resolution,[],[f9186,f7441])).
fof(f14364,plain,(
  ( ! [X520,X516,X518,X521,X517,X519] : (sP4(X516,'T_0') | sP2(X517,'T_0',X518) | sP4(X516,X519) | sP2(X520,X521,X519)) ) | $spl3),
  inference(resolution,[],[f9186,f7438])).
fof(f14043,plain,(
  ( ! [X123,X122] : (sP5('const.option.IS_NONE_1'(s('type.option.option'(X122),'const.option.NONE_0')),X123)) ) | $spl344),
  inference(cnf_transformation,[],[f14043_D])).
fof(f14043_D,plain,(
  ( ! [X123,X122] : (sP5('const.option.IS_NONE_1'(s('type.option.option'(X122),'const.option.NONE_0')),X123)) ) <=> ~$spl344),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl344])])).
fof(f14016,plain,(
  ( ! [X158,X159,X157] : (sP4(X157,'T_0') | sP5('const.option.IS_NONE_1'(s('type.option.option'(X158),'const.option.NONE_0')),X159)) ) | ($spl3 | $spl179)),
  inference(duplicate_literal_removal,[],[f13939])).
fof(f13939,plain,(
  ( ! [X158,X159,X157] : (sP4(X157,'T_0') | sP4(X157,'T_0') | sP5('const.option.IS_NONE_1'(s('type.option.option'(X158),'const.option.NONE_0')),X159)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f8592,f8856])).
fof(f8856,plain,(
  ( ! [X194,X192,X193] : (sP4(X192,'T_0') | sP4(X192,X193) | sP5(X193,X194)) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f8802,f7584])).
fof(f8802,plain,(
  ( ! [X194,X192,X193] : (sP4('T_0','T_0') | sP4(X192,'T_0') | sP4(X192,X193) | sP5(X193,X194)) ) | $spl3),
  inference(resolution,[],[f8520,f7456])).
fof(f8520,plain,(
  ( ! [X194,X195] : (~sP4(X194,X195) | sP4(X194,'T_0') | sP4(X195,'T_0')) ) | $spl3),
  inference(resolution,[],[f4520,f4459])).
fof(f14041,plain,(
  ( ! [X120,X119] : (sP5(X119,'const.option.IS_NONE_1'(s('type.option.option'(X120),'const.option.NONE_0')))) ) | $spl342),
  inference(cnf_transformation,[],[f14041_D])).
fof(f14041_D,plain,(
  ( ! [X120,X119] : (sP5(X119,'const.option.IS_NONE_1'(s('type.option.option'(X120),'const.option.NONE_0')))) ) <=> ~$spl342),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl342])])).
fof(f14017,plain,(
  ( ! [X156,X154,X155] : (sP4(X154,'T_0') | sP5(X155,'const.option.IS_NONE_1'(s('type.option.option'(X156),'const.option.NONE_0')))) ) | ($spl3 | $spl179)),
  inference(duplicate_literal_removal,[],[f13938])).
fof(f13938,plain,(
  ( ! [X156,X154,X155] : (sP4(X154,'T_0') | sP4(X154,'T_0') | sP5(X155,'const.option.IS_NONE_1'(s('type.option.option'(X156),'const.option.NONE_0')))) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f8592,f8855])).
fof(f8855,plain,(
  ( ! [X191,X189,X190] : (sP4(X189,'T_0') | sP4(X189,X190) | sP5(X191,X190)) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f8801,f7584])).
fof(f8801,plain,(
  ( ! [X191,X189,X190] : (sP4('T_0','T_0') | sP4(X189,'T_0') | sP4(X189,X190) | sP5(X191,X190)) ) | $spl3),
  inference(resolution,[],[f8520,f7455])).
fof(f14042,plain,(
  ( ! [X123,X121,X122] : (sP4(X121,'T_0') | sP5('const.option.IS_NONE_1'(s('type.option.option'(X122),'const.option.NONE_0')),X123)) ) | $spl3),
  inference(subsumption_resolution,[],[f13928,f9241])).
fof(f9241,plain,(
  ( ! [X213] : (~sP4('T_0',X213) | sP4(X213,'T_0')) ) | $spl3),
  inference(resolution,[],[f4521,f337])).
fof(f13928,plain,(
  ( ! [X123,X121,X122] : (sP4(X121,'T_0') | sP4('T_0',X121) | sP5('const.option.IS_NONE_1'(s('type.option.option'(X122),'const.option.NONE_0')),X123)) ) | $spl3),
  inference(resolution,[],[f8592,f7456])).
fof(f14040,plain,(
  ( ! [X118,X120,X119] : (sP4(X118,'T_0') | sP5(X119,'const.option.IS_NONE_1'(s('type.option.option'(X120),'const.option.NONE_0')))) ) | $spl3),
  inference(subsumption_resolution,[],[f13927,f9241])).
fof(f13927,plain,(
  ( ! [X118,X120,X119] : (sP4(X118,'T_0') | sP4('T_0',X118) | sP5(X119,'const.option.IS_NONE_1'(s('type.option.option'(X120),'const.option.NONE_0')))) ) | $spl3),
  inference(resolution,[],[f8592,f7455])).
fof(f14039,plain,(
  ( ! [X116,X117,X115] : (sP0('const.option.IS_NONE_1'(s('type.option.option'(X115),'const.option.NONE_0')),X116,X117)) ) | $spl340),
  inference(cnf_transformation,[],[f14039_D])).
fof(f14039_D,plain,(
  ( ! [X116,X117,X115] : (sP0('const.option.IS_NONE_1'(s('type.option.option'(X115),'const.option.NONE_0')),X116,X117)) ) <=> ~$spl340),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl340])])).
fof(f14038,plain,(
  ( ! [X116,X114,X117,X115] : (sP4(X114,'T_0') | sP0('const.option.IS_NONE_1'(s('type.option.option'(X115),'const.option.NONE_0')),X116,X117)) ) | $spl3),
  inference(subsumption_resolution,[],[f13926,f9241])).
fof(f13926,plain,(
  ( ! [X116,X114,X117,X115] : (sP4(X114,'T_0') | sP4('T_0',X114) | sP0('const.option.IS_NONE_1'(s('type.option.option'(X115),'const.option.NONE_0')),X116,X117)) ) | $spl3),
  inference(resolution,[],[f8592,f7444])).
fof(f14037,plain,(
  ( ! [X111,X112,X113] : (sP0(X111,'const.option.IS_NONE_1'(s('type.option.option'(X112),'const.option.NONE_0')),X113)) ) | $spl338),
  inference(cnf_transformation,[],[f14037_D])).
fof(f14037_D,plain,(
  ( ! [X111,X112,X113] : (sP0(X111,'const.option.IS_NONE_1'(s('type.option.option'(X112),'const.option.NONE_0')),X113)) ) <=> ~$spl338),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl338])])).
fof(f14036,plain,(
  ( ! [X111,X112,X110,X113] : (sP4(X110,'T_0') | sP0(X111,'const.option.IS_NONE_1'(s('type.option.option'(X112),'const.option.NONE_0')),X113)) ) | $spl3),
  inference(subsumption_resolution,[],[f13925,f9241])).
fof(f13925,plain,(
  ( ! [X111,X112,X110,X113] : (sP4(X110,'T_0') | sP4('T_0',X110) | sP0(X111,'const.option.IS_NONE_1'(s('type.option.option'(X112),'const.option.NONE_0')),X113)) ) | $spl3),
  inference(resolution,[],[f8592,f7443])).
fof(f14035,plain,(
  ( ! [X109,X107,X108] : (sP0(X107,X108,'const.option.IS_NONE_1'(s('type.option.option'(X109),'const.option.NONE_0')))) ) | $spl336),
  inference(cnf_transformation,[],[f14035_D])).
fof(f14035_D,plain,(
  ( ! [X109,X107,X108] : (sP0(X107,X108,'const.option.IS_NONE_1'(s('type.option.option'(X109),'const.option.NONE_0')))) ) <=> ~$spl336),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl336])])).
fof(f14034,plain,(
  ( ! [X109,X107,X108,X106] : (sP4(X106,'T_0') | sP0(X107,X108,'const.option.IS_NONE_1'(s('type.option.option'(X109),'const.option.NONE_0')))) ) | $spl3),
  inference(subsumption_resolution,[],[f13924,f9241])).
fof(f13924,plain,(
  ( ! [X109,X107,X108,X106] : (sP4(X106,'T_0') | sP4('T_0',X106) | sP0(X107,X108,'const.option.IS_NONE_1'(s('type.option.option'(X109),'const.option.NONE_0')))) ) | $spl3),
  inference(resolution,[],[f8592,f7442])).
fof(f14033,plain,(
  ( ! [X103,X105,X104] : (sP1(X103,X104,'const.option.IS_NONE_1'(s('type.option.option'(X105),'const.option.NONE_0')))) ) | $spl334),
  inference(cnf_transformation,[],[f14033_D])).
fof(f14033_D,plain,(
  ( ! [X103,X105,X104] : (sP1(X103,X104,'const.option.IS_NONE_1'(s('type.option.option'(X105),'const.option.NONE_0')))) ) <=> ~$spl334),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl334])])).
fof(f14032,plain,(
  ( ! [X103,X105,X102,X104] : (sP4(X102,'T_0') | sP1(X103,X104,'const.option.IS_NONE_1'(s('type.option.option'(X105),'const.option.NONE_0')))) ) | $spl3),
  inference(subsumption_resolution,[],[f13923,f9241])).
fof(f13923,plain,(
  ( ! [X103,X105,X102,X104] : (sP4(X102,'T_0') | sP4('T_0',X102) | sP1(X103,X104,'const.option.IS_NONE_1'(s('type.option.option'(X105),'const.option.NONE_0')))) ) | $spl3),
  inference(resolution,[],[f8592,f7441])).
fof(f14031,plain,(
  ( ! [X101,X99,X100] : (sP2(X99,X100,'const.option.IS_NONE_1'(s('type.option.option'(X101),'const.option.NONE_0')))) ) | $spl332),
  inference(cnf_transformation,[],[f14031_D])).
fof(f14031_D,plain,(
  ( ! [X101,X99,X100] : (sP2(X99,X100,'const.option.IS_NONE_1'(s('type.option.option'(X101),'const.option.NONE_0')))) ) <=> ~$spl332),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl332])])).
fof(f14030,plain,(
  ( ! [X101,X99,X100,X98] : (sP4(X98,'T_0') | sP2(X99,X100,'const.option.IS_NONE_1'(s('type.option.option'(X101),'const.option.NONE_0')))) ) | $spl3),
  inference(subsumption_resolution,[],[f13922,f9241])).
fof(f13922,plain,(
  ( ! [X101,X99,X100,X98] : (sP4(X98,'T_0') | sP4('T_0',X98) | sP2(X99,X100,'const.option.IS_NONE_1'(s('type.option.option'(X101),'const.option.NONE_0')))) ) | $spl3),
  inference(resolution,[],[f8592,f7438])).
fof(f13331,plain,(
  ( ! [X366,X365] : (sP1('F_0',X365,X366)) ) | $spl328),
  inference(cnf_transformation,[],[f13331_D])).
fof(f13331_D,plain,(
  ( ! [X366,X365] : (sP1('F_0',X365,X366)) ) <=> ~$spl328),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl328])])).
fof(f11824,plain,(
  ( ! [X686] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(X686),'const.option.NONE_0')),'T_0')) ) | $spl276),
  inference(cnf_transformation,[],[f11824_D])).
fof(f11824_D,plain,(
  ( ! [X686] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(X686),'const.option.NONE_0')),'T_0')) ) <=> ~$spl276),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl276])])).
fof(f13278,plain,(
  ( ! [X685,X683,X684] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(X683),'const.option.NONE_0')),'T_0') | sP1('F_0',X684,X685)) ) | $spl3),
  inference(resolution,[],[f8480,f7792])).
fof(f7792,plain,(
  ( ! [X35] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(X35),'const.option.NONE_0')),'F_0')) ) | $spl3),
  inference(subsumption_resolution,[],[f7764,f361])).
fof(f7764,plain,(
  ( ! [X35] : (~p(s(bool,'T_0')) | sP4('const.option.IS_NONE_1'(s('type.option.option'(X35),'const.option.NONE_0')),'F_0')) ) | $spl3),
  inference(superposition,[],[f7582,f359])).
fof(f13337,plain,(
  ( ! [X397,X396,X398] : (sP1('const.option.IS_SOME_1'(s('type.option.option'(X396),'const.option.NONE_0')),X397,X398)) ) | $spl330),
  inference(cnf_transformation,[],[f13337_D])).
fof(f13337_D,plain,(
  ( ! [X397,X396,X398] : (sP1('const.option.IS_SOME_1'(s('type.option.option'(X396),'const.option.NONE_0')),X397,X398)) ) <=> ~$spl330),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl330])])).
fof(f13336,plain,(
  ( ! [X397,X395,X396,X398] : (sP4(X395,'T_0') | sP1('const.option.IS_SOME_1'(s('type.option.option'(X396),'const.option.NONE_0')),X397,X398)) ) | $spl3),
  inference(subsumption_resolution,[],[f13209,f9241])).
fof(f13209,plain,(
  ( ! [X397,X395,X396,X398] : (sP4(X395,'T_0') | sP1('const.option.IS_SOME_1'(s('type.option.option'(X396),'const.option.NONE_0')),X397,X398) | sP4('T_0',X395)) ) | $spl3),
  inference(resolution,[],[f8480,f7556])).
fof(f7556,plain,(
  ( ! [X56,X55] : (sP4(X56,'const.option.IS_SOME_1'(s('type.option.option'(X55),'const.option.NONE_0'))) | sP4('T_0',X56)) ) | $spl3),
  inference(subsumption_resolution,[],[f7516,f332])).
fof(f7516,plain,(
  ( ! [X56,X55] : (p(s(bool,'F_0')) | sP4('T_0',X56) | sP4(X56,'const.option.IS_SOME_1'(s('type.option.option'(X55),'const.option.NONE_0')))) ) | $spl3),
  inference(superposition,[],[f3989,f340])).
fof(f13335,plain,(
  ( ! [X383,X385,X384] : (sP4(X383,'T_0') | sP1('F_0',X384,X385)) ) | $spl3),
  inference(subsumption_resolution,[],[f13205,f9241])).
fof(f13205,plain,(
  ( ! [X383,X385,X384] : (sP4(X383,'T_0') | sP1('F_0',X384,X385) | sP4('T_0',X383)) ) | $spl3),
  inference(resolution,[],[f8480,f7491])).
fof(f13301,plain,(
  ( ! [X380,X382,X381] : (sP4(X380,'T_0') | sP1('F_0',X381,X382)) ) | $spl3),
  inference(duplicate_literal_removal,[],[f13204])).
fof(f13204,plain,(
  ( ! [X380,X382,X381] : (sP4(X380,'T_0') | sP1('F_0',X381,X382) | sP4(X380,'T_0')) ) | $spl3),
  inference(resolution,[],[f8480,f7721])).
fof(f7721,plain,(
  ( ! [X136] : (sP4(X136,'F_0') | sP4(X136,'T_0')) ) | $spl3),
  inference(resolution,[],[f7582,f4459])).
fof(f8292,plain,(
  ( ! [X0] : (sP4('T_0',X0)) ) | $spl200),
  inference(cnf_transformation,[],[f8292_D])).
fof(f8292_D,plain,(
  ( ! [X0] : (sP4('T_0',X0)) ) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f13334,plain,(
  ( ! [X377,X379,X376,X378] : (sP4(X376,'T_0') | sP1('F_0',X377,X378) | sP4('T_0',X379)) ) | $spl3),
  inference(subsumption_resolution,[],[f13203,f9227])).
fof(f9227,plain,(
  ( ! [X187,X186] : (~sP4(X187,X186) | sP4(X186,'T_0') | sP4('T_0',X187)) ) | $spl3),
  inference(resolution,[],[f4521,f3605])).
fof(f13203,plain,(
  ( ! [X377,X379,X376,X378] : (sP4(X376,'T_0') | sP1('F_0',X377,X378) | sP4('T_0',X379) | sP4(X379,X376)) ) | $spl3),
  inference(resolution,[],[f8480,f7720])).
fof(f7720,plain,(
  ( ! [X134,X135] : (sP4(X134,'F_0') | sP4('T_0',X135) | sP4(X135,X134)) ) | $spl3),
  inference(resolution,[],[f7582,f3989])).
fof(f13333,plain,(
  ( ! [X372,X374,X373,X375] : (sP4(X372,'T_0') | sP1('F_0',X373,X374) | sP4('T_0',X375)) ) | $spl3),
  inference(subsumption_resolution,[],[f13202,f8517])).
fof(f8517,plain,(
  ( ! [X187,X186] : (~sP4(X186,X187) | sP4(X186,'T_0') | sP4('T_0',X187)) ) | $spl3),
  inference(resolution,[],[f4520,f3605])).
fof(f13202,plain,(
  ( ! [X372,X374,X373,X375] : (sP4(X372,'T_0') | sP1('F_0',X373,X374) | sP4('T_0',X375) | sP4(X372,X375)) ) | $spl3),
  inference(resolution,[],[f8480,f8152])).
fof(f8152,plain,(
  ( ! [X194,X195] : (sP4('T_0',X194) | sP4(X195,X194) | sP4(X195,'F_0')) ) | $spl3),
  inference(resolution,[],[f3990,f7582])).
fof(f13332,plain,(
  ( ! [X368,X370,X369,X371] : (sP4(X368,'T_0') | sP1('F_0',X369,X370) | sP4(X371,'T_0')) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f13201,f8520])).
fof(f13201,plain,(
  ( ! [X368,X370,X369,X371] : (sP4(X368,'T_0') | sP1('F_0',X369,X370) | sP4(X371,'T_0') | sP4(X371,X368)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f8480,f8857])).
fof(f8857,plain,(
  ( ! [X196,X197] : (sP4(X196,'T_0') | sP4(X197,'F_0') | sP4(X196,X197)) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f8804,f7584])).
fof(f8804,plain,(
  ( ! [X196,X197] : (sP4('T_0','T_0') | sP4(X196,'T_0') | sP4(X197,'F_0') | sP4(X196,X197)) ) | $spl3),
  inference(resolution,[],[f8520,f7720])).
fof(f13330,plain,(
  ( ! [X364,X366,X365,X367] : (sP4(X364,'T_0') | sP1('F_0',X365,X366) | sP4(X367,'T_0')) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f13200,f8520])).
fof(f13200,plain,(
  ( ! [X364,X366,X365,X367] : (sP4(X364,'T_0') | sP1('F_0',X365,X366) | sP4(X367,'T_0') | sP4(X364,X367)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f8480,f8860])).
fof(f8860,plain,(
  ( ! [X204,X205] : (sP4(X204,'T_0') | sP4(X205,X204) | sP4(X205,'F_0')) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f8807,f7584])).
fof(f8807,plain,(
  ( ! [X204,X205] : (sP4('T_0','T_0') | sP4(X204,'T_0') | sP4(X205,X204) | sP4(X205,'F_0')) ) | $spl3),
  inference(resolution,[],[f8520,f8152])).
fof(f13327,plain,(
  ( ! [X198,X196,X199,X197] : (sP1(X196,X197,X198) | sP5(X196,X199)) ) | $spl326),
  inference(cnf_transformation,[],[f13327_D])).
fof(f13327_D,plain,(
  ( ! [X198,X196,X199,X197] : (sP1(X196,X197,X198) | sP5(X196,X199)) ) <=> ~$spl326),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl326])])).
fof(f13304,plain,(
  ( ! [X255,X253,X257,X254,X256] : (sP4(X253,'T_0') | sP1(X254,X255,X256) | sP5(X254,X257)) ) | ($spl3 | $spl179)),
  inference(duplicate_literal_removal,[],[f13173])).
fof(f13173,plain,(
  ( ! [X255,X253,X257,X254,X256] : (sP4(X253,'T_0') | sP1(X254,X255,X256) | sP4(X253,'T_0') | sP5(X254,X257)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f8480,f8856])).
fof(f13325,plain,(
  ( ! [X191,X194,X192,X193] : (sP1(X191,X192,X193) | sP5(X194,X191)) ) | $spl324),
  inference(cnf_transformation,[],[f13325_D])).
fof(f13325_D,plain,(
  ( ! [X191,X194,X192,X193] : (sP1(X191,X192,X193) | sP5(X194,X191)) ) <=> ~$spl324),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl324])])).
fof(f13305,plain,(
  ( ! [X251,X249,X252,X250,X248] : (sP4(X248,'T_0') | sP1(X249,X250,X251) | sP5(X252,X249)) ) | ($spl3 | $spl179)),
  inference(duplicate_literal_removal,[],[f13172])).
fof(f13172,plain,(
  ( ! [X251,X249,X252,X250,X248] : (sP4(X248,'T_0') | sP1(X249,X250,X251) | sP4(X248,'T_0') | sP5(X252,X249)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f8480,f8855])).
fof(f13326,plain,(
  ( ! [X198,X196,X199,X197,X195] : (sP4(X195,'T_0') | sP1(X196,X197,X198) | sP5(X196,X199)) ) | $spl3),
  inference(subsumption_resolution,[],[f13162,f9241])).
fof(f13162,plain,(
  ( ! [X198,X196,X199,X197,X195] : (sP4(X195,'T_0') | sP1(X196,X197,X198) | sP4('T_0',X195) | sP5(X196,X199)) ) | $spl3),
  inference(resolution,[],[f8480,f7456])).
fof(f13324,plain,(
  ( ! [X191,X194,X192,X190,X193] : (sP4(X190,'T_0') | sP1(X191,X192,X193) | sP5(X194,X191)) ) | $spl3),
  inference(subsumption_resolution,[],[f13161,f9241])).
fof(f13161,plain,(
  ( ! [X191,X194,X192,X190,X193] : (sP4(X190,'T_0') | sP1(X191,X192,X193) | sP4('T_0',X190) | sP5(X194,X191)) ) | $spl3),
  inference(resolution,[],[f8480,f7455])).
fof(f13323,plain,(
  ( ! [X189,X187,X185,X188,X186] : (sP1(X185,X186,X187) | sP0(X185,X188,X189)) ) | $spl322),
  inference(cnf_transformation,[],[f13323_D])).
fof(f13323_D,plain,(
  ( ! [X189,X187,X185,X188,X186] : (sP1(X185,X186,X187) | sP0(X185,X188,X189)) ) <=> ~$spl322),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl322])])).
fof(f13322,plain,(
  ( ! [X189,X187,X185,X188,X186,X184] : (sP4(X184,'T_0') | sP1(X185,X186,X187) | sP0(X185,X188,X189)) ) | $spl3),
  inference(subsumption_resolution,[],[f13160,f9241])).
fof(f13160,plain,(
  ( ! [X189,X187,X185,X188,X186,X184] : (sP4(X184,'T_0') | sP1(X185,X186,X187) | sP4('T_0',X184) | sP0(X185,X188,X189)) ) | $spl3),
  inference(resolution,[],[f8480,f7444])).
fof(f13321,plain,(
  ( ! [X182,X180,X183,X181,X179] : (sP1(X179,X180,X181) | sP0(X182,X179,X183)) ) | $spl320),
  inference(cnf_transformation,[],[f13321_D])).
fof(f13321_D,plain,(
  ( ! [X182,X180,X183,X181,X179] : (sP1(X179,X180,X181) | sP0(X182,X179,X183)) ) <=> ~$spl320),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl320])])).
fof(f13320,plain,(
  ( ! [X182,X180,X178,X183,X181,X179] : (sP4(X178,'T_0') | sP1(X179,X180,X181) | sP0(X182,X179,X183)) ) | $spl3),
  inference(subsumption_resolution,[],[f13159,f9241])).
fof(f13159,plain,(
  ( ! [X182,X180,X178,X183,X181,X179] : (sP4(X178,'T_0') | sP1(X179,X180,X181) | sP4('T_0',X178) | sP0(X182,X179,X183)) ) | $spl3),
  inference(resolution,[],[f8480,f7443])).
fof(f13319,plain,(
  ( ! [X177,X175,X173,X176,X174] : (sP1(X173,X174,X175) | sP0(X176,X177,X173)) ) | $spl318),
  inference(cnf_transformation,[],[f13319_D])).
fof(f13319_D,plain,(
  ( ! [X177,X175,X173,X176,X174] : (sP1(X173,X174,X175) | sP0(X176,X177,X173)) ) <=> ~$spl318),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl318])])).
fof(f13318,plain,(
  ( ! [X177,X175,X173,X176,X174,X172] : (sP4(X172,'T_0') | sP1(X173,X174,X175) | sP0(X176,X177,X173)) ) | $spl3),
  inference(subsumption_resolution,[],[f13158,f9241])).
fof(f13158,plain,(
  ( ! [X177,X175,X173,X176,X174,X172] : (sP4(X172,'T_0') | sP1(X173,X174,X175) | sP4('T_0',X172) | sP0(X176,X177,X173)) ) | $spl3),
  inference(resolution,[],[f8480,f7442])).
fof(f13317,plain,(
  ( ! [X167,X171,X169,X170,X168] : (sP1(X167,X168,X169) | sP1(X170,X171,X167)) ) | $spl316),
  inference(cnf_transformation,[],[f13317_D])).
fof(f13317_D,plain,(
  ( ! [X167,X171,X169,X170,X168] : (sP1(X167,X168,X169) | sP1(X170,X171,X167)) ) <=> ~$spl316),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl316])])).
fof(f13316,plain,(
  ( ! [X167,X171,X169,X166,X170,X168] : (sP4(X166,'T_0') | sP1(X167,X168,X169) | sP1(X170,X171,X167)) ) | $spl3),
  inference(subsumption_resolution,[],[f13157,f9241])).
fof(f13157,plain,(
  ( ! [X167,X171,X169,X166,X170,X168] : (sP4(X166,'T_0') | sP1(X167,X168,X169) | sP4('T_0',X166) | sP1(X170,X171,X167)) ) | $spl3),
  inference(resolution,[],[f8480,f7441])).
fof(f13315,plain,(
  ( ! [X165,X163,X161,X164,X162] : (sP1(X161,X162,X163) | sP2(X164,X165,X161)) ) | $spl314),
  inference(cnf_transformation,[],[f13315_D])).
fof(f13315_D,plain,(
  ( ! [X165,X163,X161,X164,X162] : (sP1(X161,X162,X163) | sP2(X164,X165,X161)) ) <=> ~$spl314),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl314])])).
fof(f13314,plain,(
  ( ! [X165,X163,X161,X164,X162,X160] : (sP4(X160,'T_0') | sP1(X161,X162,X163) | sP2(X164,X165,X161)) ) | $spl3),
  inference(subsumption_resolution,[],[f13156,f9241])).
fof(f13156,plain,(
  ( ! [X165,X163,X161,X164,X162,X160] : (sP4(X160,'T_0') | sP1(X161,X162,X163) | sP4('T_0',X160) | sP2(X164,X165,X161)) ) | $spl3),
  inference(resolution,[],[f8480,f7438])).
fof(f12812,plain,(
  ( ! [X366,X365] : (sP1(X365,'F_0',X366)) ) | $spl310),
  inference(cnf_transformation,[],[f12812_D])).
fof(f12812_D,plain,(
  ( ! [X366,X365] : (sP1(X365,'F_0',X366)) ) <=> ~$spl310),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl310])])).
fof(f12759,plain,(
  ( ! [X685,X683,X684] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(X683),'const.option.NONE_0')),'T_0') | sP1(X684,'F_0',X685)) ) | $spl3),
  inference(resolution,[],[f8479,f7792])).
fof(f12818,plain,(
  ( ! [X397,X396,X398] : (sP1(X396,'const.option.IS_SOME_1'(s('type.option.option'(X397),'const.option.NONE_0')),X398)) ) | $spl312),
  inference(cnf_transformation,[],[f12818_D])).
fof(f12818_D,plain,(
  ( ! [X397,X396,X398] : (sP1(X396,'const.option.IS_SOME_1'(s('type.option.option'(X397),'const.option.NONE_0')),X398)) ) <=> ~$spl312),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl312])])).
fof(f12817,plain,(
  ( ! [X397,X395,X396,X398] : (sP4(X395,'T_0') | sP1(X396,'const.option.IS_SOME_1'(s('type.option.option'(X397),'const.option.NONE_0')),X398)) ) | $spl3),
  inference(subsumption_resolution,[],[f12690,f9241])).
fof(f12690,plain,(
  ( ! [X397,X395,X396,X398] : (sP4(X395,'T_0') | sP1(X396,'const.option.IS_SOME_1'(s('type.option.option'(X397),'const.option.NONE_0')),X398) | sP4('T_0',X395)) ) | $spl3),
  inference(resolution,[],[f8479,f7556])).
fof(f12816,plain,(
  ( ! [X383,X385,X384] : (sP4(X383,'T_0') | sP1(X384,'F_0',X385)) ) | $spl3),
  inference(subsumption_resolution,[],[f12686,f9241])).
fof(f12686,plain,(
  ( ! [X383,X385,X384] : (sP4(X383,'T_0') | sP1(X384,'F_0',X385) | sP4('T_0',X383)) ) | $spl3),
  inference(resolution,[],[f8479,f7491])).
fof(f12782,plain,(
  ( ! [X380,X382,X381] : (sP4(X380,'T_0') | sP1(X381,'F_0',X382)) ) | $spl3),
  inference(duplicate_literal_removal,[],[f12685])).
fof(f12685,plain,(
  ( ! [X380,X382,X381] : (sP4(X380,'T_0') | sP1(X381,'F_0',X382) | sP4(X380,'T_0')) ) | $spl3),
  inference(resolution,[],[f8479,f7721])).
fof(f12815,plain,(
  ( ! [X377,X379,X376,X378] : (sP4(X376,'T_0') | sP1(X377,'F_0',X378) | sP4('T_0',X379)) ) | $spl3),
  inference(subsumption_resolution,[],[f12684,f9227])).
fof(f12684,plain,(
  ( ! [X377,X379,X376,X378] : (sP4(X376,'T_0') | sP1(X377,'F_0',X378) | sP4('T_0',X379) | sP4(X379,X376)) ) | $spl3),
  inference(resolution,[],[f8479,f7720])).
fof(f12814,plain,(
  ( ! [X372,X374,X373,X375] : (sP4(X372,'T_0') | sP1(X373,'F_0',X374) | sP4('T_0',X375)) ) | $spl3),
  inference(subsumption_resolution,[],[f12683,f8517])).
fof(f12683,plain,(
  ( ! [X372,X374,X373,X375] : (sP4(X372,'T_0') | sP1(X373,'F_0',X374) | sP4('T_0',X375) | sP4(X372,X375)) ) | $spl3),
  inference(resolution,[],[f8479,f8152])).
fof(f12813,plain,(
  ( ! [X368,X370,X369,X371] : (sP4(X368,'T_0') | sP1(X369,'F_0',X370) | sP4(X371,'T_0')) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f12682,f8520])).
fof(f12682,plain,(
  ( ! [X368,X370,X369,X371] : (sP4(X368,'T_0') | sP1(X369,'F_0',X370) | sP4(X371,'T_0') | sP4(X371,X368)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f8479,f8857])).
fof(f12811,plain,(
  ( ! [X364,X366,X365,X367] : (sP4(X364,'T_0') | sP1(X365,'F_0',X366) | sP4(X367,'T_0')) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f12681,f8520])).
fof(f12681,plain,(
  ( ! [X364,X366,X365,X367] : (sP4(X364,'T_0') | sP1(X365,'F_0',X366) | sP4(X367,'T_0') | sP4(X364,X367)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f8479,f8860])).
fof(f12808,plain,(
  ( ! [X198,X196,X199,X197] : (sP1(X196,X197,X198) | sP5(X197,X199)) ) | $spl308),
  inference(cnf_transformation,[],[f12808_D])).
fof(f12808_D,plain,(
  ( ! [X198,X196,X199,X197] : (sP1(X196,X197,X198) | sP5(X197,X199)) ) <=> ~$spl308),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl308])])).
fof(f12785,plain,(
  ( ! [X255,X253,X257,X254,X256] : (sP4(X253,'T_0') | sP1(X254,X255,X256) | sP5(X255,X257)) ) | ($spl3 | $spl179)),
  inference(duplicate_literal_removal,[],[f12654])).
fof(f12654,plain,(
  ( ! [X255,X253,X257,X254,X256] : (sP4(X253,'T_0') | sP1(X254,X255,X256) | sP4(X253,'T_0') | sP5(X255,X257)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f8479,f8856])).
fof(f12806,plain,(
  ( ! [X191,X194,X192,X193] : (sP1(X191,X192,X193) | sP5(X194,X192)) ) | $spl306),
  inference(cnf_transformation,[],[f12806_D])).
fof(f12806_D,plain,(
  ( ! [X191,X194,X192,X193] : (sP1(X191,X192,X193) | sP5(X194,X192)) ) <=> ~$spl306),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl306])])).
fof(f12786,plain,(
  ( ! [X251,X249,X252,X250,X248] : (sP4(X248,'T_0') | sP1(X249,X250,X251) | sP5(X252,X250)) ) | ($spl3 | $spl179)),
  inference(duplicate_literal_removal,[],[f12653])).
fof(f12653,plain,(
  ( ! [X251,X249,X252,X250,X248] : (sP4(X248,'T_0') | sP1(X249,X250,X251) | sP4(X248,'T_0') | sP5(X252,X250)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f8479,f8855])).
fof(f12807,plain,(
  ( ! [X198,X196,X199,X197,X195] : (sP4(X195,'T_0') | sP1(X196,X197,X198) | sP5(X197,X199)) ) | $spl3),
  inference(subsumption_resolution,[],[f12643,f9241])).
fof(f12643,plain,(
  ( ! [X198,X196,X199,X197,X195] : (sP4(X195,'T_0') | sP1(X196,X197,X198) | sP4('T_0',X195) | sP5(X197,X199)) ) | $spl3),
  inference(resolution,[],[f8479,f7456])).
fof(f12805,plain,(
  ( ! [X191,X194,X192,X190,X193] : (sP4(X190,'T_0') | sP1(X191,X192,X193) | sP5(X194,X192)) ) | $spl3),
  inference(subsumption_resolution,[],[f12642,f9241])).
fof(f12642,plain,(
  ( ! [X191,X194,X192,X190,X193] : (sP4(X190,'T_0') | sP1(X191,X192,X193) | sP4('T_0',X190) | sP5(X194,X192)) ) | $spl3),
  inference(resolution,[],[f8479,f7455])).
fof(f12804,plain,(
  ( ! [X189,X187,X185,X188,X186] : (sP1(X185,X186,X187) | sP0(X186,X188,X189)) ) | $spl304),
  inference(cnf_transformation,[],[f12804_D])).
fof(f12804_D,plain,(
  ( ! [X189,X187,X185,X188,X186] : (sP1(X185,X186,X187) | sP0(X186,X188,X189)) ) <=> ~$spl304),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl304])])).
fof(f12803,plain,(
  ( ! [X189,X187,X185,X188,X186,X184] : (sP4(X184,'T_0') | sP1(X185,X186,X187) | sP0(X186,X188,X189)) ) | $spl3),
  inference(subsumption_resolution,[],[f12641,f9241])).
fof(f12641,plain,(
  ( ! [X189,X187,X185,X188,X186,X184] : (sP4(X184,'T_0') | sP1(X185,X186,X187) | sP4('T_0',X184) | sP0(X186,X188,X189)) ) | $spl3),
  inference(resolution,[],[f8479,f7444])).
fof(f12802,plain,(
  ( ! [X182,X180,X183,X181,X179] : (sP1(X179,X180,X181) | sP0(X182,X180,X183)) ) | $spl302),
  inference(cnf_transformation,[],[f12802_D])).
fof(f12802_D,plain,(
  ( ! [X182,X180,X183,X181,X179] : (sP1(X179,X180,X181) | sP0(X182,X180,X183)) ) <=> ~$spl302),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl302])])).
fof(f12801,plain,(
  ( ! [X182,X180,X178,X183,X181,X179] : (sP4(X178,'T_0') | sP1(X179,X180,X181) | sP0(X182,X180,X183)) ) | $spl3),
  inference(subsumption_resolution,[],[f12640,f9241])).
fof(f12640,plain,(
  ( ! [X182,X180,X178,X183,X181,X179] : (sP4(X178,'T_0') | sP1(X179,X180,X181) | sP4('T_0',X178) | sP0(X182,X180,X183)) ) | $spl3),
  inference(resolution,[],[f8479,f7443])).
fof(f12800,plain,(
  ( ! [X177,X175,X173,X176,X174] : (sP1(X173,X174,X175) | sP0(X176,X177,X174)) ) | $spl300),
  inference(cnf_transformation,[],[f12800_D])).
fof(f12800_D,plain,(
  ( ! [X177,X175,X173,X176,X174] : (sP1(X173,X174,X175) | sP0(X176,X177,X174)) ) <=> ~$spl300),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl300])])).
fof(f12799,plain,(
  ( ! [X177,X175,X173,X176,X174,X172] : (sP4(X172,'T_0') | sP1(X173,X174,X175) | sP0(X176,X177,X174)) ) | $spl3),
  inference(subsumption_resolution,[],[f12639,f9241])).
fof(f12639,plain,(
  ( ! [X177,X175,X173,X176,X174,X172] : (sP4(X172,'T_0') | sP1(X173,X174,X175) | sP4('T_0',X172) | sP0(X176,X177,X174)) ) | $spl3),
  inference(resolution,[],[f8479,f7442])).
fof(f12798,plain,(
  ( ! [X167,X171,X169,X170,X168] : (sP1(X167,X168,X169) | sP1(X170,X171,X168)) ) | $spl298),
  inference(cnf_transformation,[],[f12798_D])).
fof(f12798_D,plain,(
  ( ! [X167,X171,X169,X170,X168] : (sP1(X167,X168,X169) | sP1(X170,X171,X168)) ) <=> ~$spl298),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl298])])).
fof(f12797,plain,(
  ( ! [X167,X171,X169,X166,X170,X168] : (sP4(X166,'T_0') | sP1(X167,X168,X169) | sP1(X170,X171,X168)) ) | $spl3),
  inference(subsumption_resolution,[],[f12638,f9241])).
fof(f12638,plain,(
  ( ! [X167,X171,X169,X166,X170,X168] : (sP4(X166,'T_0') | sP1(X167,X168,X169) | sP4('T_0',X166) | sP1(X170,X171,X168)) ) | $spl3),
  inference(resolution,[],[f8479,f7441])).
fof(f12796,plain,(
  ( ! [X165,X163,X161,X164,X162] : (sP1(X161,X162,X163) | sP2(X164,X165,X162)) ) | $spl296),
  inference(cnf_transformation,[],[f12796_D])).
fof(f12796_D,plain,(
  ( ! [X165,X163,X161,X164,X162] : (sP1(X161,X162,X163) | sP2(X164,X165,X162)) ) <=> ~$spl296),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl296])])).
fof(f12795,plain,(
  ( ! [X165,X163,X161,X164,X162,X160] : (sP4(X160,'T_0') | sP1(X161,X162,X163) | sP2(X164,X165,X162)) ) | $spl3),
  inference(subsumption_resolution,[],[f12637,f9241])).
fof(f12637,plain,(
  ( ! [X165,X163,X161,X164,X162,X160] : (sP4(X160,'T_0') | sP1(X161,X162,X163) | sP4('T_0',X160) | sP2(X164,X165,X162)) ) | $spl3),
  inference(resolution,[],[f8479,f7438])).
fof(f12356,plain,(
  ( ! [X366,X365] : (sP2('F_0',X365,X366)) ) | $spl292),
  inference(cnf_transformation,[],[f12356_D])).
fof(f12356_D,plain,(
  ( ! [X366,X365] : (sP2('F_0',X365,X366)) ) <=> ~$spl292),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl292])])).
fof(f12303,plain,(
  ( ! [X685,X683,X684] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(X683),'const.option.NONE_0')),'T_0') | sP2('F_0',X684,X685)) ) | $spl3),
  inference(resolution,[],[f8477,f7792])).
fof(f12362,plain,(
  ( ! [X397,X396,X398] : (sP2('const.option.IS_SOME_1'(s('type.option.option'(X396),'const.option.NONE_0')),X397,X398)) ) | $spl294),
  inference(cnf_transformation,[],[f12362_D])).
fof(f12362_D,plain,(
  ( ! [X397,X396,X398] : (sP2('const.option.IS_SOME_1'(s('type.option.option'(X396),'const.option.NONE_0')),X397,X398)) ) <=> ~$spl294),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl294])])).
fof(f12361,plain,(
  ( ! [X397,X395,X396,X398] : (sP4(X395,'T_0') | sP2('const.option.IS_SOME_1'(s('type.option.option'(X396),'const.option.NONE_0')),X397,X398)) ) | $spl3),
  inference(subsumption_resolution,[],[f12234,f9241])).
fof(f12234,plain,(
  ( ! [X397,X395,X396,X398] : (sP4(X395,'T_0') | sP2('const.option.IS_SOME_1'(s('type.option.option'(X396),'const.option.NONE_0')),X397,X398) | sP4('T_0',X395)) ) | $spl3),
  inference(resolution,[],[f8477,f7556])).
fof(f12360,plain,(
  ( ! [X383,X385,X384] : (sP4(X383,'T_0') | sP2('F_0',X384,X385)) ) | $spl3),
  inference(subsumption_resolution,[],[f12230,f9241])).
fof(f12230,plain,(
  ( ! [X383,X385,X384] : (sP4(X383,'T_0') | sP2('F_0',X384,X385) | sP4('T_0',X383)) ) | $spl3),
  inference(resolution,[],[f8477,f7491])).
fof(f12326,plain,(
  ( ! [X380,X382,X381] : (sP4(X380,'T_0') | sP2('F_0',X381,X382)) ) | $spl3),
  inference(duplicate_literal_removal,[],[f12229])).
fof(f12229,plain,(
  ( ! [X380,X382,X381] : (sP4(X380,'T_0') | sP2('F_0',X381,X382) | sP4(X380,'T_0')) ) | $spl3),
  inference(resolution,[],[f8477,f7721])).
fof(f12359,plain,(
  ( ! [X377,X379,X376,X378] : (sP4(X376,'T_0') | sP2('F_0',X377,X378) | sP4('T_0',X379)) ) | $spl3),
  inference(subsumption_resolution,[],[f12228,f9227])).
fof(f12228,plain,(
  ( ! [X377,X379,X376,X378] : (sP4(X376,'T_0') | sP2('F_0',X377,X378) | sP4('T_0',X379) | sP4(X379,X376)) ) | $spl3),
  inference(resolution,[],[f8477,f7720])).
fof(f12358,plain,(
  ( ! [X372,X374,X373,X375] : (sP4(X372,'T_0') | sP2('F_0',X373,X374) | sP4('T_0',X375)) ) | $spl3),
  inference(subsumption_resolution,[],[f12227,f8517])).
fof(f12227,plain,(
  ( ! [X372,X374,X373,X375] : (sP4(X372,'T_0') | sP2('F_0',X373,X374) | sP4('T_0',X375) | sP4(X372,X375)) ) | $spl3),
  inference(resolution,[],[f8477,f8152])).
fof(f12357,plain,(
  ( ! [X368,X370,X369,X371] : (sP4(X368,'T_0') | sP2('F_0',X369,X370) | sP4(X371,'T_0')) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f12226,f8520])).
fof(f12226,plain,(
  ( ! [X368,X370,X369,X371] : (sP4(X368,'T_0') | sP2('F_0',X369,X370) | sP4(X371,'T_0') | sP4(X371,X368)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f8477,f8857])).
fof(f12355,plain,(
  ( ! [X364,X366,X365,X367] : (sP4(X364,'T_0') | sP2('F_0',X365,X366) | sP4(X367,'T_0')) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f12225,f8520])).
fof(f12225,plain,(
  ( ! [X364,X366,X365,X367] : (sP4(X364,'T_0') | sP2('F_0',X365,X366) | sP4(X367,'T_0') | sP4(X364,X367)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f8477,f8860])).
fof(f12352,plain,(
  ( ! [X198,X196,X199,X197] : (sP2(X196,X197,X198) | sP5(X196,X199)) ) | $spl290),
  inference(cnf_transformation,[],[f12352_D])).
fof(f12352_D,plain,(
  ( ! [X198,X196,X199,X197] : (sP2(X196,X197,X198) | sP5(X196,X199)) ) <=> ~$spl290),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl290])])).
fof(f12329,plain,(
  ( ! [X255,X253,X257,X254,X256] : (sP4(X253,'T_0') | sP2(X254,X255,X256) | sP5(X254,X257)) ) | ($spl3 | $spl179)),
  inference(duplicate_literal_removal,[],[f12198])).
fof(f12198,plain,(
  ( ! [X255,X253,X257,X254,X256] : (sP4(X253,'T_0') | sP2(X254,X255,X256) | sP4(X253,'T_0') | sP5(X254,X257)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f8477,f8856])).
fof(f12350,plain,(
  ( ! [X191,X194,X192,X193] : (sP2(X191,X192,X193) | sP5(X194,X191)) ) | $spl288),
  inference(cnf_transformation,[],[f12350_D])).
fof(f12350_D,plain,(
  ( ! [X191,X194,X192,X193] : (sP2(X191,X192,X193) | sP5(X194,X191)) ) <=> ~$spl288),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl288])])).
fof(f12330,plain,(
  ( ! [X251,X249,X252,X250,X248] : (sP4(X248,'T_0') | sP2(X249,X250,X251) | sP5(X252,X249)) ) | ($spl3 | $spl179)),
  inference(duplicate_literal_removal,[],[f12197])).
fof(f12197,plain,(
  ( ! [X251,X249,X252,X250,X248] : (sP4(X248,'T_0') | sP2(X249,X250,X251) | sP4(X248,'T_0') | sP5(X252,X249)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f8477,f8855])).
fof(f12351,plain,(
  ( ! [X198,X196,X199,X197,X195] : (sP4(X195,'T_0') | sP2(X196,X197,X198) | sP5(X196,X199)) ) | $spl3),
  inference(subsumption_resolution,[],[f12187,f9241])).
fof(f12187,plain,(
  ( ! [X198,X196,X199,X197,X195] : (sP4(X195,'T_0') | sP2(X196,X197,X198) | sP4('T_0',X195) | sP5(X196,X199)) ) | $spl3),
  inference(resolution,[],[f8477,f7456])).
fof(f12349,plain,(
  ( ! [X191,X194,X192,X190,X193] : (sP4(X190,'T_0') | sP2(X191,X192,X193) | sP5(X194,X191)) ) | $spl3),
  inference(subsumption_resolution,[],[f12186,f9241])).
fof(f12186,plain,(
  ( ! [X191,X194,X192,X190,X193] : (sP4(X190,'T_0') | sP2(X191,X192,X193) | sP4('T_0',X190) | sP5(X194,X191)) ) | $spl3),
  inference(resolution,[],[f8477,f7455])).
fof(f12348,plain,(
  ( ! [X189,X187,X185,X188,X186] : (sP2(X185,X186,X187) | sP0(X185,X188,X189)) ) | $spl286),
  inference(cnf_transformation,[],[f12348_D])).
fof(f12348_D,plain,(
  ( ! [X189,X187,X185,X188,X186] : (sP2(X185,X186,X187) | sP0(X185,X188,X189)) ) <=> ~$spl286),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl286])])).
fof(f12347,plain,(
  ( ! [X189,X187,X185,X188,X186,X184] : (sP4(X184,'T_0') | sP2(X185,X186,X187) | sP0(X185,X188,X189)) ) | $spl3),
  inference(subsumption_resolution,[],[f12185,f9241])).
fof(f12185,plain,(
  ( ! [X189,X187,X185,X188,X186,X184] : (sP4(X184,'T_0') | sP2(X185,X186,X187) | sP4('T_0',X184) | sP0(X185,X188,X189)) ) | $spl3),
  inference(resolution,[],[f8477,f7444])).
fof(f12346,plain,(
  ( ! [X182,X180,X183,X181,X179] : (sP2(X179,X180,X181) | sP0(X182,X179,X183)) ) | $spl284),
  inference(cnf_transformation,[],[f12346_D])).
fof(f12346_D,plain,(
  ( ! [X182,X180,X183,X181,X179] : (sP2(X179,X180,X181) | sP0(X182,X179,X183)) ) <=> ~$spl284),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl284])])).
fof(f12345,plain,(
  ( ! [X182,X180,X178,X183,X181,X179] : (sP4(X178,'T_0') | sP2(X179,X180,X181) | sP0(X182,X179,X183)) ) | $spl3),
  inference(subsumption_resolution,[],[f12184,f9241])).
fof(f12184,plain,(
  ( ! [X182,X180,X178,X183,X181,X179] : (sP4(X178,'T_0') | sP2(X179,X180,X181) | sP4('T_0',X178) | sP0(X182,X179,X183)) ) | $spl3),
  inference(resolution,[],[f8477,f7443])).
fof(f12344,plain,(
  ( ! [X177,X175,X173,X176,X174] : (sP2(X173,X174,X175) | sP0(X176,X177,X173)) ) | $spl282),
  inference(cnf_transformation,[],[f12344_D])).
fof(f12344_D,plain,(
  ( ! [X177,X175,X173,X176,X174] : (sP2(X173,X174,X175) | sP0(X176,X177,X173)) ) <=> ~$spl282),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl282])])).
fof(f12343,plain,(
  ( ! [X177,X175,X173,X176,X174,X172] : (sP4(X172,'T_0') | sP2(X173,X174,X175) | sP0(X176,X177,X173)) ) | $spl3),
  inference(subsumption_resolution,[],[f12183,f9241])).
fof(f12183,plain,(
  ( ! [X177,X175,X173,X176,X174,X172] : (sP4(X172,'T_0') | sP2(X173,X174,X175) | sP4('T_0',X172) | sP0(X176,X177,X173)) ) | $spl3),
  inference(resolution,[],[f8477,f7442])).
fof(f12342,plain,(
  ( ! [X167,X171,X169,X170,X168] : (sP2(X167,X168,X169) | sP1(X170,X171,X167)) ) | $spl280),
  inference(cnf_transformation,[],[f12342_D])).
fof(f12342_D,plain,(
  ( ! [X167,X171,X169,X170,X168] : (sP2(X167,X168,X169) | sP1(X170,X171,X167)) ) <=> ~$spl280),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl280])])).
fof(f12341,plain,(
  ( ! [X167,X171,X169,X166,X170,X168] : (sP4(X166,'T_0') | sP2(X167,X168,X169) | sP1(X170,X171,X167)) ) | $spl3),
  inference(subsumption_resolution,[],[f12182,f9241])).
fof(f12182,plain,(
  ( ! [X167,X171,X169,X166,X170,X168] : (sP4(X166,'T_0') | sP2(X167,X168,X169) | sP4('T_0',X166) | sP1(X170,X171,X167)) ) | $spl3),
  inference(resolution,[],[f8477,f7441])).
fof(f12340,plain,(
  ( ! [X165,X163,X161,X164,X162] : (sP2(X161,X162,X163) | sP2(X164,X165,X161)) ) | $spl278),
  inference(cnf_transformation,[],[f12340_D])).
fof(f12340_D,plain,(
  ( ! [X165,X163,X161,X164,X162] : (sP2(X161,X162,X163) | sP2(X164,X165,X161)) ) <=> ~$spl278),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl278])])).
fof(f12339,plain,(
  ( ! [X165,X163,X161,X164,X162,X160] : (sP4(X160,'T_0') | sP2(X161,X162,X163) | sP2(X164,X165,X161)) ) | $spl3),
  inference(subsumption_resolution,[],[f12181,f9241])).
fof(f12181,plain,(
  ( ! [X165,X163,X161,X164,X162,X160] : (sP4(X160,'T_0') | sP2(X161,X162,X163) | sP4('T_0',X160) | sP2(X164,X165,X161)) ) | $spl3),
  inference(resolution,[],[f8477,f7438])).
fof(f12044,plain,(
  $false | ($spl3 | $spl179 | $spl276)),
  inference(subsumption_resolution,[],[f12043,f338])).
fof(f12043,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl179 | $spl276)),
  inference(forward_demodulation,[],[f12042,f359])).
fof(f12042,plain,(
  ( ! [X9] : (~p(s(bool,'const.option.IS_NONE_1'(s('type.option.option'(X9),'const.option.NONE_0'))))) ) | ($spl3 | $spl179 | $spl276)),
  inference(subsumption_resolution,[],[f12022,f7584])).
fof(f12022,plain,(
  ( ! [X9] : (~p(s(bool,'const.option.IS_NONE_1'(s('type.option.option'(X9),'const.option.NONE_0')))) | sP4('T_0','T_0')) ) | ($spl3 | $spl276)),
  inference(resolution,[],[f11824,f3988])).
fof(f3988,plain,(
  ( ! [X76,X77] : (~sP4(X77,X76) | ~p(s(bool,X77)) | sP4('T_0',X76)) ) | $spl3),
  inference(resolution,[],[f3605,f254])).
fof(f12040,plain,(
  $false | $spl276),
  inference(subsumption_resolution,[],[f12039,f338])).
fof(f12039,plain,(
  ~p(s(bool,'T_0')) | $spl276),
  inference(forward_demodulation,[],[f12038,f359])).
fof(f12038,plain,(
  ( ! [X2] : (~p(s(bool,'const.option.IS_NONE_1'(s('type.option.option'(X2),'const.option.NONE_0'))))) ) | $spl276),
  inference(subsumption_resolution,[],[f12015,f361])).
fof(f12015,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.option.IS_NONE_1'(s('type.option.option'(X2),'const.option.NONE_0'))))) ) | $spl276),
  inference(resolution,[],[f11824,f261])).
fof(f12037,plain,(
  $false | $spl276),
  inference(subsumption_resolution,[],[f12036,f338])).
fof(f12036,plain,(
  ~p(s(bool,'T_0')) | $spl276),
  inference(forward_demodulation,[],[f12035,f359])).
fof(f12035,plain,(
  ( ! [X1] : (~p(s(bool,'const.option.IS_NONE_1'(s('type.option.option'(X1),'const.option.NONE_0'))))) ) | $spl276),
  inference(subsumption_resolution,[],[f12014,f361])).
fof(f12014,plain,(
  ( ! [X1] : (~p(s(bool,'const.option.IS_NONE_1'(s('type.option.option'(X1),'const.option.NONE_0')))) | ~p(s(bool,'T_0'))) ) | $spl276),
  inference(resolution,[],[f11824,f254])).
fof(f11916,plain,(
  $false | ($spl3 | $spl4 | $spl179 | $spl274)),
  inference(subsumption_resolution,[],[f11915,f338])).
fof(f11915,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl4 | $spl179 | $spl274)),
  inference(forward_demodulation,[],[f11914,f1649])).
fof(f11914,plain,(
  ~p(s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) | ($spl3 | $spl179 | $spl274)),
  inference(subsumption_resolution,[],[f11893,f7584])).
fof(f11893,plain,(
  ~p(s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) | sP4('T_0','T_0') | ($spl3 | $spl274)),
  inference(resolution,[],[f11823,f3988])).
fof(f11913,plain,(
  $false | ($spl3 | $spl4 | $spl274)),
  inference(subsumption_resolution,[],[f11912,f338])).
fof(f11912,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl4 | $spl274)),
  inference(subsumption_resolution,[],[f11892,f1629])).
fof(f1629,plain,(
  ( ! [X53] : (~sP4(X53,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7))) | ~p(s(bool,X53))) ) | $spl4),
  inference(resolution,[],[f369,f254])).
fof(f11892,plain,(
  ~p(s(bool,'T_0')) | sP4('T_0','const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7))) | ($spl3 | $spl274)),
  inference(resolution,[],[f11823,f3987])).
fof(f3987,plain,(
  ( ! [X74,X75] : (~sP4(X74,X75) | ~p(s(bool,X75)) | sP4('T_0',X74)) ) | $spl3),
  inference(resolution,[],[f3605,f254])).
fof(f11911,plain,(
  $false | ($spl4 | $spl274)),
  inference(subsumption_resolution,[],[f11910,f338])).
fof(f11910,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl274)),
  inference(forward_demodulation,[],[f11909,f1649])).
fof(f11909,plain,(
  ~p(s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) | $spl274),
  inference(subsumption_resolution,[],[f11886,f361])).
fof(f11886,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) | $spl274),
  inference(resolution,[],[f11823,f261])).
fof(f11908,plain,(
  $false | ($spl4 | $spl274)),
  inference(subsumption_resolution,[],[f11907,f338])).
fof(f11907,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl274)),
  inference(forward_demodulation,[],[f11906,f1649])).
fof(f11906,plain,(
  ~p(s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) | $spl274),
  inference(subsumption_resolution,[],[f11885,f361])).
fof(f11885,plain,(
  ~p(s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) | ~p(s(bool,'T_0')) | $spl274),
  inference(resolution,[],[f11823,f254])).
fof(f11905,plain,(
  $false | ($spl4 | $spl274)),
  inference(subsumption_resolution,[],[f11883,f338])).
fof(f11883,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl274)),
  inference(resolution,[],[f11823,f1628])).
fof(f1628,plain,(
  ( ! [X52] : (~sP4('const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)),X52) | ~p(s(bool,X52))) ) | $spl4),
  inference(resolution,[],[f369,f254])).
fof(f11847,plain,(
  $false | ($spl4 | $spl210)),
  inference(subsumption_resolution,[],[f11842,f338])).
fof(f11842,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl210)),
  inference(resolution,[],[f8585,f1628])).
fof(f11840,plain,(
  $false | ($spl179 | $spl210)),
  inference(resolution,[],[f8585,f7584])).
fof(f11786,plain,(
  ( ! [X366,X365] : (sP2(X365,'F_0',X366)) ) | $spl270),
  inference(cnf_transformation,[],[f11786_D])).
fof(f11786_D,plain,(
  ( ! [X366,X365] : (sP2(X365,'F_0',X366)) ) <=> ~$spl270),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl270])])).
fof(f11734,plain,(
  ( ! [X687,X688,X686] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(X686),'const.option.NONE_0')),'T_0') | sP2(X687,'F_0',X688)) ) | $spl3),
  inference(resolution,[],[f8476,f7792])).
fof(f11823,plain,(
  sP4('const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)),'T_0') | $spl274),
  inference(cnf_transformation,[],[f11823_D])).
fof(f11823_D,plain,(
  sP4('const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)),'T_0') <=> ~$spl274),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl274])])).
fof(f11727,plain,(
  ( ! [X662,X663] : (sP4('const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)),'T_0') | sP2(X662,'F_0',X663)) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f8476,f7587])).
fof(f7587,plain,(
  sP4('const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)),'F_0') | ($spl3 | $spl4)),
  inference(subsumption_resolution,[],[f7568,f361])).
fof(f7568,plain,(
  sP4('const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)),'F_0') | ~p(s(bool,'T_0')) | ($spl3 | $spl4)),
  inference(resolution,[],[f7491,f1629])).
fof(f11792,plain,(
  ( ! [X397,X396,X398] : (sP2(X396,'const.option.IS_SOME_1'(s('type.option.option'(X397),'const.option.NONE_0')),X398)) ) | $spl272),
  inference(cnf_transformation,[],[f11792_D])).
fof(f11792_D,plain,(
  ( ! [X397,X396,X398] : (sP2(X396,'const.option.IS_SOME_1'(s('type.option.option'(X397),'const.option.NONE_0')),X398)) ) <=> ~$spl272),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl272])])).
fof(f11791,plain,(
  ( ! [X397,X395,X396,X398] : (sP4(X395,'T_0') | sP2(X396,'const.option.IS_SOME_1'(s('type.option.option'(X397),'const.option.NONE_0')),X398)) ) | $spl3),
  inference(subsumption_resolution,[],[f11664,f9241])).
fof(f11664,plain,(
  ( ! [X397,X395,X396,X398] : (sP4(X395,'T_0') | sP2(X396,'const.option.IS_SOME_1'(s('type.option.option'(X397),'const.option.NONE_0')),X398) | sP4('T_0',X395)) ) | $spl3),
  inference(resolution,[],[f8476,f7556])).
fof(f11790,plain,(
  ( ! [X383,X385,X384] : (sP4(X383,'T_0') | sP2(X384,'F_0',X385)) ) | $spl3),
  inference(subsumption_resolution,[],[f11660,f9241])).
fof(f11660,plain,(
  ( ! [X383,X385,X384] : (sP4(X383,'T_0') | sP2(X384,'F_0',X385) | sP4('T_0',X383)) ) | $spl3),
  inference(resolution,[],[f8476,f7491])).
fof(f11756,plain,(
  ( ! [X380,X382,X381] : (sP4(X380,'T_0') | sP2(X381,'F_0',X382)) ) | $spl3),
  inference(duplicate_literal_removal,[],[f11659])).
fof(f11659,plain,(
  ( ! [X380,X382,X381] : (sP4(X380,'T_0') | sP2(X381,'F_0',X382) | sP4(X380,'T_0')) ) | $spl3),
  inference(resolution,[],[f8476,f7721])).
fof(f11789,plain,(
  ( ! [X377,X379,X376,X378] : (sP4(X376,'T_0') | sP2(X377,'F_0',X378) | sP4('T_0',X379)) ) | $spl3),
  inference(subsumption_resolution,[],[f11658,f9227])).
fof(f11658,plain,(
  ( ! [X377,X379,X376,X378] : (sP4(X376,'T_0') | sP2(X377,'F_0',X378) | sP4('T_0',X379) | sP4(X379,X376)) ) | $spl3),
  inference(resolution,[],[f8476,f7720])).
fof(f11788,plain,(
  ( ! [X372,X374,X373,X375] : (sP4(X372,'T_0') | sP2(X373,'F_0',X374) | sP4('T_0',X375)) ) | $spl3),
  inference(subsumption_resolution,[],[f11657,f8517])).
fof(f11657,plain,(
  ( ! [X372,X374,X373,X375] : (sP4(X372,'T_0') | sP2(X373,'F_0',X374) | sP4('T_0',X375) | sP4(X372,X375)) ) | $spl3),
  inference(resolution,[],[f8476,f8152])).
fof(f11787,plain,(
  ( ! [X368,X370,X369,X371] : (sP4(X368,'T_0') | sP2(X369,'F_0',X370) | sP4(X371,'T_0')) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f11656,f8520])).
fof(f11656,plain,(
  ( ! [X368,X370,X369,X371] : (sP4(X368,'T_0') | sP2(X369,'F_0',X370) | sP4(X371,'T_0') | sP4(X371,X368)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f8476,f8857])).
fof(f11785,plain,(
  ( ! [X364,X366,X365,X367] : (sP4(X364,'T_0') | sP2(X365,'F_0',X366) | sP4(X367,'T_0')) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f11655,f8520])).
fof(f11655,plain,(
  ( ! [X364,X366,X365,X367] : (sP4(X364,'T_0') | sP2(X365,'F_0',X366) | sP4(X367,'T_0') | sP4(X364,X367)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f8476,f8860])).
fof(f11782,plain,(
  ( ! [X198,X196,X199,X197] : (sP2(X196,X197,X198) | sP5(X197,X199)) ) | $spl268),
  inference(cnf_transformation,[],[f11782_D])).
fof(f11782_D,plain,(
  ( ! [X198,X196,X199,X197] : (sP2(X196,X197,X198) | sP5(X197,X199)) ) <=> ~$spl268),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl268])])).
fof(f11759,plain,(
  ( ! [X255,X253,X257,X254,X256] : (sP4(X253,'T_0') | sP2(X254,X255,X256) | sP5(X255,X257)) ) | ($spl3 | $spl179)),
  inference(duplicate_literal_removal,[],[f11628])).
fof(f11628,plain,(
  ( ! [X255,X253,X257,X254,X256] : (sP4(X253,'T_0') | sP2(X254,X255,X256) | sP4(X253,'T_0') | sP5(X255,X257)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f8476,f8856])).
fof(f11780,plain,(
  ( ! [X191,X194,X192,X193] : (sP2(X191,X192,X193) | sP5(X194,X192)) ) | $spl266),
  inference(cnf_transformation,[],[f11780_D])).
fof(f11780_D,plain,(
  ( ! [X191,X194,X192,X193] : (sP2(X191,X192,X193) | sP5(X194,X192)) ) <=> ~$spl266),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl266])])).
fof(f11760,plain,(
  ( ! [X251,X249,X252,X250,X248] : (sP4(X248,'T_0') | sP2(X249,X250,X251) | sP5(X252,X250)) ) | ($spl3 | $spl179)),
  inference(duplicate_literal_removal,[],[f11627])).
fof(f11627,plain,(
  ( ! [X251,X249,X252,X250,X248] : (sP4(X248,'T_0') | sP2(X249,X250,X251) | sP4(X248,'T_0') | sP5(X252,X250)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f8476,f8855])).
fof(f11781,plain,(
  ( ! [X198,X196,X199,X197,X195] : (sP4(X195,'T_0') | sP2(X196,X197,X198) | sP5(X197,X199)) ) | $spl3),
  inference(subsumption_resolution,[],[f11617,f9241])).
fof(f11617,plain,(
  ( ! [X198,X196,X199,X197,X195] : (sP4(X195,'T_0') | sP2(X196,X197,X198) | sP4('T_0',X195) | sP5(X197,X199)) ) | $spl3),
  inference(resolution,[],[f8476,f7456])).
fof(f11779,plain,(
  ( ! [X191,X194,X192,X190,X193] : (sP4(X190,'T_0') | sP2(X191,X192,X193) | sP5(X194,X192)) ) | $spl3),
  inference(subsumption_resolution,[],[f11616,f9241])).
fof(f11616,plain,(
  ( ! [X191,X194,X192,X190,X193] : (sP4(X190,'T_0') | sP2(X191,X192,X193) | sP4('T_0',X190) | sP5(X194,X192)) ) | $spl3),
  inference(resolution,[],[f8476,f7455])).
fof(f11778,plain,(
  ( ! [X189,X187,X185,X188,X186] : (sP2(X185,X186,X187) | sP0(X186,X188,X189)) ) | $spl264),
  inference(cnf_transformation,[],[f11778_D])).
fof(f11778_D,plain,(
  ( ! [X189,X187,X185,X188,X186] : (sP2(X185,X186,X187) | sP0(X186,X188,X189)) ) <=> ~$spl264),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl264])])).
fof(f11777,plain,(
  ( ! [X189,X187,X185,X188,X186,X184] : (sP4(X184,'T_0') | sP2(X185,X186,X187) | sP0(X186,X188,X189)) ) | $spl3),
  inference(subsumption_resolution,[],[f11615,f9241])).
fof(f11615,plain,(
  ( ! [X189,X187,X185,X188,X186,X184] : (sP4(X184,'T_0') | sP2(X185,X186,X187) | sP4('T_0',X184) | sP0(X186,X188,X189)) ) | $spl3),
  inference(resolution,[],[f8476,f7444])).
fof(f11776,plain,(
  ( ! [X182,X180,X183,X181,X179] : (sP2(X179,X180,X181) | sP0(X182,X180,X183)) ) | $spl262),
  inference(cnf_transformation,[],[f11776_D])).
fof(f11776_D,plain,(
  ( ! [X182,X180,X183,X181,X179] : (sP2(X179,X180,X181) | sP0(X182,X180,X183)) ) <=> ~$spl262),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl262])])).
fof(f11775,plain,(
  ( ! [X182,X180,X178,X183,X181,X179] : (sP4(X178,'T_0') | sP2(X179,X180,X181) | sP0(X182,X180,X183)) ) | $spl3),
  inference(subsumption_resolution,[],[f11614,f9241])).
fof(f11614,plain,(
  ( ! [X182,X180,X178,X183,X181,X179] : (sP4(X178,'T_0') | sP2(X179,X180,X181) | sP4('T_0',X178) | sP0(X182,X180,X183)) ) | $spl3),
  inference(resolution,[],[f8476,f7443])).
fof(f11774,plain,(
  ( ! [X177,X175,X173,X176,X174] : (sP2(X173,X174,X175) | sP0(X176,X177,X174)) ) | $spl260),
  inference(cnf_transformation,[],[f11774_D])).
fof(f11774_D,plain,(
  ( ! [X177,X175,X173,X176,X174] : (sP2(X173,X174,X175) | sP0(X176,X177,X174)) ) <=> ~$spl260),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl260])])).
fof(f11773,plain,(
  ( ! [X177,X175,X173,X176,X174,X172] : (sP4(X172,'T_0') | sP2(X173,X174,X175) | sP0(X176,X177,X174)) ) | $spl3),
  inference(subsumption_resolution,[],[f11613,f9241])).
fof(f11613,plain,(
  ( ! [X177,X175,X173,X176,X174,X172] : (sP4(X172,'T_0') | sP2(X173,X174,X175) | sP4('T_0',X172) | sP0(X176,X177,X174)) ) | $spl3),
  inference(resolution,[],[f8476,f7442])).
fof(f11772,plain,(
  ( ! [X167,X171,X169,X170,X168] : (sP2(X167,X168,X169) | sP1(X170,X171,X168)) ) | $spl258),
  inference(cnf_transformation,[],[f11772_D])).
fof(f11772_D,plain,(
  ( ! [X167,X171,X169,X170,X168] : (sP2(X167,X168,X169) | sP1(X170,X171,X168)) ) <=> ~$spl258),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl258])])).
fof(f11771,plain,(
  ( ! [X167,X171,X169,X166,X170,X168] : (sP4(X166,'T_0') | sP2(X167,X168,X169) | sP1(X170,X171,X168)) ) | $spl3),
  inference(subsumption_resolution,[],[f11612,f9241])).
fof(f11612,plain,(
  ( ! [X167,X171,X169,X166,X170,X168] : (sP4(X166,'T_0') | sP2(X167,X168,X169) | sP4('T_0',X166) | sP1(X170,X171,X168)) ) | $spl3),
  inference(resolution,[],[f8476,f7441])).
fof(f11770,plain,(
  ( ! [X165,X163,X161,X164,X162] : (sP2(X161,X162,X163) | sP2(X164,X165,X162)) ) | $spl256),
  inference(cnf_transformation,[],[f11770_D])).
fof(f11770_D,plain,(
  ( ! [X165,X163,X161,X164,X162] : (sP2(X161,X162,X163) | sP2(X164,X165,X162)) ) <=> ~$spl256),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl256])])).
fof(f11769,plain,(
  ( ! [X165,X163,X161,X164,X162,X160] : (sP4(X160,'T_0') | sP2(X161,X162,X163) | sP2(X164,X165,X162)) ) | $spl3),
  inference(subsumption_resolution,[],[f11611,f9241])).
fof(f11611,plain,(
  ( ! [X165,X163,X161,X164,X162,X160] : (sP4(X160,'T_0') | sP2(X161,X162,X163) | sP4('T_0',X160) | sP2(X164,X165,X162)) ) | $spl3),
  inference(resolution,[],[f8476,f7438])).
fof(f11537,plain,(
  ( ! [X2,X1] : (sP0('T_0',X1,X2)) ) | $spl254),
  inference(cnf_transformation,[],[f11537_D])).
fof(f11537_D,plain,(
  ( ! [X2,X1] : (sP0('T_0',X1,X2)) ) <=> ~$spl254),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl254])])).
fof(f11525,plain,(
  ( ! [X2,X0,X1] : (sP4('T_0',X0) | sP0('T_0',X1,X2)) ) | $spl3),
  inference(factoring,[],[f8115])).
fof(f11534,plain,(
  ( ! [X6,X4,X7,X5] : (sP4('T_0',X4) | sP3(X6,X7,X5) | ~p(s(bool,X5)) | ~sP2(X6,X5,X7) | ~sP1(X7,X5,X6)) ) | $spl3),
  inference(subsumption_resolution,[],[f11519,f3988])).
fof(f11519,plain,(
  ( ! [X6,X4,X7,X5] : (sP4('T_0',X4) | sP4(X5,X4) | sP3(X6,X7,X5) | ~p(s(bool,X5)) | ~sP2(X6,X5,X7) | ~sP1(X7,X5,X6)) ) | $spl3),
  inference(resolution,[],[f8115,f228])).
fof(f11470,plain,(
  ( ! [X2,X1] : (sP0(X1,'T_0',X2)) ) | $spl252),
  inference(cnf_transformation,[],[f11470_D])).
fof(f11470_D,plain,(
  ( ! [X2,X1] : (sP0(X1,'T_0',X2)) ) <=> ~$spl252),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl252])])).
fof(f11459,plain,(
  ( ! [X2,X0,X1] : (sP4('T_0',X0) | sP0(X1,'T_0',X2)) ) | $spl3),
  inference(factoring,[],[f8114])).
fof(f11405,plain,(
  ( ! [X2,X1] : (sP0(X1,X2,'T_0')) ) | $spl250),
  inference(cnf_transformation,[],[f11405_D])).
fof(f11405_D,plain,(
  ( ! [X2,X1] : (sP0(X1,X2,'T_0')) ) <=> ~$spl250),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl250])])).
fof(f11394,plain,(
  ( ! [X2,X0,X1] : (sP4('T_0',X0) | sP0(X1,X2,'T_0')) ) | $spl3),
  inference(factoring,[],[f8113])).
fof(f11343,plain,(
  ( ! [X2,X1] : (sP1(X1,X2,'T_0')) ) | $spl248),
  inference(cnf_transformation,[],[f11343_D])).
fof(f11343_D,plain,(
  ( ! [X2,X1] : (sP1(X1,X2,'T_0')) ) <=> ~$spl248),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl248])])).
fof(f11333,plain,(
  ( ! [X2,X0,X1] : (sP4('T_0',X0) | sP1(X1,X2,'T_0')) ) | $spl3),
  inference(factoring,[],[f8112])).
fof(f11250,plain,(
  ( ! [X2,X1] : (sP2(X1,X2,'T_0')) ) | $spl246),
  inference(cnf_transformation,[],[f11250_D])).
fof(f11250_D,plain,(
  ( ! [X2,X1] : (sP2(X1,X2,'T_0')) ) <=> ~$spl246),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl246])])).
fof(f11240,plain,(
  ( ! [X2,X0,X1] : (sP4('T_0',X0) | sP2(X1,X2,'T_0')) ) | $spl3),
  inference(factoring,[],[f8109])).
fof(f11160,plain,(
  ( ! [X6,X4,X7,X5] : (sP4('T_0',X4) | sP3(X6,X7,X5) | ~p(s(bool,X5)) | ~sP2(X6,X5,X7) | ~sP1(X7,X5,X6)) ) | $spl3),
  inference(subsumption_resolution,[],[f11145,f3987])).
fof(f11145,plain,(
  ( ! [X6,X4,X7,X5] : (sP4('T_0',X4) | sP4(X4,X5) | sP3(X6,X7,X5) | ~p(s(bool,X5)) | ~sP2(X6,X5,X7) | ~sP1(X7,X5,X6)) ) | $spl3),
  inference(resolution,[],[f7444,f228])).
fof(f10771,plain,(
  $false | ($spl4 | $spl6)),
  inference(subsumption_resolution,[],[f10700,f10631])).
fof(f10631,plain,(
  ( ! [X1] : (~p(s(bool,X1))) ) | ($spl4 | $spl6)),
  inference(subsumption_resolution,[],[f10630,f325])).
fof(f10630,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,X1) | ~p(s(bool,X1))) ) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f452,f1649])).
fof(f10700,plain,(
  p(s(bool,'F_0')) | ($spl4 | $spl6)),
  inference(resolution,[],[f10631,f335])).
fof(f10770,plain,(
  $false | ($spl4 | $spl6)),
  inference(subsumption_resolution,[],[f10699,f10631])).
fof(f10699,plain,(
  p(s(bool,'F_0')) | ($spl4 | $spl6)),
  inference(resolution,[],[f10631,f333])).
fof(f10769,plain,(
  $false | ($spl4 | $spl6)),
  inference(subsumption_resolution,[],[f10695,f10631])).
fof(f10695,plain,(
  ( ! [X137] : (p(s(bool,X137))) ) | ($spl4 | $spl6)),
  inference(resolution,[],[f10631,f300])).
fof(f10768,plain,(
  $false | ($spl4 | $spl6)),
  inference(subsumption_resolution,[],[f10694,f10631])).
fof(f10694,plain,(
  ( ! [X136] : (p(s(bool,X136))) ) | ($spl4 | $spl6)),
  inference(resolution,[],[f10631,f297])).
fof(f10767,plain,(
  $false | ($spl4 | $spl6)),
  inference(subsumption_resolution,[],[f10693,f10631])).
fof(f10693,plain,(
  ( ! [X135] : (p(s(bool,X135))) ) | ($spl4 | $spl6)),
  inference(resolution,[],[f10631,f294])).
fof(f10766,plain,(
  $false | ($spl4 | $spl6)),
  inference(subsumption_resolution,[],[f10691,f10631])).
fof(f10691,plain,(
  p(s(bool,'F_0')) | ($spl4 | $spl6)),
  inference(resolution,[],[f10631,f291])).
fof(f10701,plain,(
  $false | ($spl4 | $spl6)),
  inference(resolution,[],[f10631,f337])).
fof(f10702,plain,(
  $false | ($spl4 | $spl6)),
  inference(resolution,[],[f10631,f338])).
fof(f10703,plain,(
  $false | ($spl4 | $spl6)),
  inference(resolution,[],[f10631,f361])).
fof(f10612,plain,(
  $false | ($spl4 | $spl240)),
  inference(subsumption_resolution,[],[f10574,f10531])).
fof(f10531,plain,(
  s('type.option.option'(sK6),'const.option.NONE_0') != s('type.option.option'(sK6),sK7) | $spl4),
  inference(superposition,[],[f356,f1687])).
fof(f356,plain,(
  ( ! [X0,X1] : (s('type.option.option'(X0),'const.option.NONE_0') != s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1)))) )),
  inference(cnf_transformation,[],[f112])).
fof(f112,plain,(
  ! [X0,X1] : s('type.option.option'(X0),'const.option.NONE_0') != s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1)))),
  inference(flattening,[],[f111])).
fof(f111,plain,(
  ! [X0] : ! [X1] : ~s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1)))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X6] : ! [X12] : ~s('type.option.option'(X6),'const.option.NONE_0') = s('type.option.option'(X6),'const.option.SOME_1'(s(X6,X12)))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|1','thm.option.NOT_NONE_SOME')).
fof(f10574,plain,(
  s('type.option.option'(sK6),'const.option.NONE_0') = s('type.option.option'(sK6),sK7) | $spl240),
  inference(cnf_transformation,[],[f10574_D])).
fof(f10574_D,plain,(
  s('type.option.option'(sK6),'const.option.NONE_0') = s('type.option.option'(sK6),sK7) <=> ~$spl240),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl240])])).
fof(f10576,plain,(
  s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9('const.option.SOME_1'(s(sK6,sK9(sK7,sK6))),sK6)))) | $spl242),
  inference(cnf_transformation,[],[f10576_D])).
fof(f10576_D,plain,(
  s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9('const.option.SOME_1'(s(sK6,sK9(sK7,sK6))),sK6)))) <=> ~$spl242),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl242])])).
fof(f10583,plain,(
  s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9('const.option.SOME_1'(s(sK6,sK9(sK7,sK6))),sK6)))) | s('type.option.option'(sK6),'const.option.NONE_0') = s('type.option.option'(sK6),sK7) | $spl4),
  inference(forward_demodulation,[],[f10534,f1687])).
fof(f10534,plain,(
  s('type.option.option'(sK6),'const.option.NONE_0') = s('type.option.option'(sK6),sK7) | s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6)))) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9('const.option.SOME_1'(s(sK6,sK9(sK7,sK6))),sK6)))) | $spl4),
  inference(superposition,[],[f357,f1687])).
fof(f10572,plain,(
  s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9('const.option.SOME_1'(s(sK6,sK9(sK7,sK6))),sK6)))) | s('type.option.option'(sK6),'const.option.NONE_0') = s('type.option.option'(sK6),sK7) | $spl4),
  inference(forward_demodulation,[],[f10509,f1687])).
fof(f10509,plain,(
  s('type.option.option'(sK6),'const.option.NONE_0') = s('type.option.option'(sK6),sK7) | s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6)))) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9('const.option.SOME_1'(s(sK6,sK9(sK7,sK6))),sK6)))) | $spl4),
  inference(superposition,[],[f1687,f357])).
fof(f10130,plain,(
  sP4('F_0','const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0'))) | $spl238),
  inference(cnf_transformation,[],[f10130_D])).
fof(f10130_D,plain,(
  sP4('F_0','const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0'))) <=> ~$spl238),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl238])])).
fof(f10116,plain,(
  sP4('F_0','const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0'))) | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6)))) | ($spl3 | $spl4)),
  inference(superposition,[],[f9972,f357])).
fof(f9972,plain,(
  sP4('F_0','const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7))) | ($spl3 | $spl4)),
  inference(subsumption_resolution,[],[f9950,f361])).
fof(f9950,plain,(
  sP4('F_0','const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7))) | ~p(s(bool,'T_0')) | ($spl3 | $spl4)),
  inference(resolution,[],[f9866,f1628])).
fof(f10051,plain,(
  ( ! [X154,X155] : (sP4('F_0','F_0') | p(s(bool,X154)) | p(s(bool,X155)) | ~sP5(X154,X155)) ) | $spl3),
  inference(resolution,[],[f9969,f269])).
fof(f10049,plain,(
  ( ! [X151,X150] : (sP4('F_0','F_0') | p(s(bool,X150)) | p(s(bool,X151)) | ~sP5(X150,X151)) ) | $spl3),
  inference(resolution,[],[f9969,f273])).
fof(f10026,plain,(
  ( ! [X132,X130,X131] : (sP4('F_0',X130) | p(s(bool,X131)) | ~p(s(bool,X132)) | s(bool,X130) != s(bool,X132)) ) | $spl3),
  inference(resolution,[],[f9969,f360])).
fof(f455,plain,(
  ( ! [X4,X3] : (s(bool,X3) != s(bool,X4) | ~p(s(bool,X3)) | p(s(bool,X4))) ) | $spl12),
  inference(cnf_transformation,[],[f455_D])).
fof(f455_D,plain,(
  ( ! [X4,X3] : (s(bool,X3) != s(bool,X4) | ~p(s(bool,X3)) | p(s(bool,X4))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f10025,plain,(
  ( ! [X127,X128,X129] : (sP4('F_0',X127) | ~p(s(bool,X128)) | p(s(bool,X129)) | s(bool,X128) != s(bool,X129)) ) | $spl3),
  inference(resolution,[],[f9969,f360])).
fof(f9981,plain,(
  ( ! [X30,X28,X31,X29] : (sP4('F_0',X28) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) ) | $spl3),
  inference(resolution,[],[f9969,f223])).
fof(f9980,plain,(
  ( ! [X26,X24,X27,X25] : (sP4('F_0',X24) | p(s(bool,X25)) | ~p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X24) != s(bool,X26)) ) | $spl3),
  inference(resolution,[],[f9969,f222])).
fof(f10090,plain,(
  ( ! [X23,X21,X20] : (sP4('F_0',X20) | ~p(s(bool,X21)) | p(s(bool,X23)) | s(bool,X21) != s(bool,X23)) ) | $spl3),
  inference(subsumption_resolution,[],[f9979,f360])).
fof(f9979,plain,(
  ( ! [X23,X21,X22,X20] : (sP4('F_0',X20) | ~p(s(bool,X21)) | ~p(s(bool,X22)) | p(s(bool,X23)) | s(bool,X21) != s(bool,X23)) ) | $spl3),
  inference(resolution,[],[f9969,f222])).
fof(f9978,plain,(
  ( ! [X19,X17,X18,X16] : (sP4('F_0',X16) | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) ) | $spl3),
  inference(resolution,[],[f9969,f221])).
fof(f10089,plain,(
  ( ! [X4,X7] : (sP4('F_0',X4) | s(bool,X4) != s(bool,X7)) ) | $spl234),
  inference(cnf_transformation,[],[f10089_D])).
fof(f10089_D,plain,(
  ( ! [X4,X7] : (sP4('F_0',X4) | s(bool,X4) != s(bool,X7)) ) <=> ~$spl234),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl234])])).
fof(f9977,plain,(
  ( ! [X14,X12,X15,X13] : (sP4('F_0',X12) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X12) != s(bool,X15) | s(bool,X13) != s(bool,X14)) ) | $spl3),
  inference(resolution,[],[f9969,f221])).
fof(f9976,plain,(
  ( ! [X10,X8,X11,X9] : (sP4('F_0',X8) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) ) | $spl3),
  inference(resolution,[],[f9969,f219])).
fof(f9975,plain,(
  ( ! [X6,X4,X7,X5] : (sP4('F_0',X4) | ~p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X4) != s(bool,X7) | s(bool,X5) != s(bool,X6)) ) | $spl3),
  inference(resolution,[],[f9969,f217])).
fof(f9974,plain,(
  ( ! [X2,X0,X3,X1] : (sP4('F_0',X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl3),
  inference(resolution,[],[f9969,f217])).
fof(f9926,plain,(
  ( ! [X206] : (sP4(X206,'T_0') | sP4('F_0',X206)) ) | $spl230),
  inference(cnf_transformation,[],[f9926_D])).
fof(f9926_D,plain,(
  ( ! [X206] : (sP4(X206,'T_0') | sP4('F_0',X206)) ) <=> ~$spl230),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl230])])).
fof(f9902,plain,(
  ( ! [X50,X51] : (p(s(bool,X50)) | sP4(X51,'T_0') | sP4('F_0',X51)) ) | $spl3),
  inference(duplicate_literal_removal,[],[f9889])).
fof(f9889,plain,(
  ( ! [X50,X51] : (p(s(bool,X50)) | sP4(X51,'T_0') | sP4('F_0',X51) | p(s(bool,X50))) ) | $spl3),
  inference(superposition,[],[f4523,f319])).
fof(f9903,plain,(
  ( ! [X48,X49] : (p(s(bool,X48)) | sP4(X49,'T_0') | sP4('F_0',X49)) ) | $spl3),
  inference(duplicate_literal_removal,[],[f9888])).
fof(f9888,plain,(
  ( ! [X48,X49] : (p(s(bool,X48)) | sP4(X49,'T_0') | sP4('F_0',X49) | p(s(bool,X48))) ) | $spl3),
  inference(superposition,[],[f4523,f321])).
fof(f9933,plain,(
  ( ! [X47,X46] : (p(s(bool,X46)) | sP4(X47,'T_0') | sP4('F_0',X47)) ) | $spl3),
  inference(subsumption_resolution,[],[f9887,f322])).
fof(f9887,plain,(
  ( ! [X47,X46] : (p(s(bool,X46)) | sP4(X47,'T_0') | sP4('F_0',X47) | s(bool,'T_0') = s(bool,X46)) ) | $spl3),
  inference(superposition,[],[f4523,f358])).
fof(f9865,plain,(
  ( ! [X220,X221] : (sP4(X220,'T_0') | sP4('F_0',X220) | p(s(bool,X221))) ) | $spl3),
  inference(resolution,[],[f4523,f331])).
fof(f331,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f136])).
fof(f136,plain,(
  ! [X0] : (~p(s(bool,'F_0')) | p(s(bool,X0)))),
  inference(ennf_transformation,[],[f90])).
fof(f90,plain,(
  ! [X0] : (p(s(bool,'F_0')) => p(s(bool,X0)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X11] : (p(s(bool,'F_0')) => p(s(bool,X11)))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|1','thm.bool.FALSITY')).
fof(f9864,plain,(
  ( ! [X218,X219] : (sP4(X218,'T_0') | sP4('F_0',X218) | p(s(bool,X219))) ) | $spl3),
  inference(resolution,[],[f4523,f314])).
fof(f314,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | ~p(s(bool,X0)) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,'F_0')) & p(s(bool,X0)))))),
  inference(flattening,[],[f194])).
fof(f194,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) | ~p(s(bool,X0))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,'F_0')) & p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0] : ((p(s(bool,'F_0')) & p(s(bool,X0))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X11] : ((p(s(bool,'F_0')) & p(s(bool,X11))) <=> p(s(bool,'F_0')))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|1','thm.bool.AND_CLAUSES|split|2')).
fof(f9863,plain,(
  ( ! [X216,X217] : (sP4(X216,'T_0') | sP4('F_0',X216) | p(s(bool,X217))) ) | $spl3),
  inference(resolution,[],[f4523,f307])).
fof(f9927,plain,(
  ( ! [X215,X214] : (sP4(X214,'T_0') | sP4('F_0',X214) | p(s(bool,X215))) ) | $spl3),
  inference(subsumption_resolution,[],[f9862,f361])).
fof(f9862,plain,(
  ( ! [X215,X214] : (sP4(X214,'T_0') | sP4('F_0',X214) | p(s(bool,X215)) | ~p(s(bool,'T_0'))) ) | $spl3),
  inference(resolution,[],[f4523,f293])).
fof(f293,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f181])).
fof(f9861,plain,(
  ( ! [X212,X213] : (sP4(X212,'T_0') | sP4('F_0',X212) | ~p(s(bool,X213))) ) | $spl3),
  inference(resolution,[],[f4523,f289])).
fof(f289,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f179])).
fof(f179,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f178])).
fof(f178,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f127])).
fof(f127,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f70])).
fof(f70,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f69])).
fof(f69,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X11] : ((p(s(bool,X11)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X11)))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|1','thm.bool.IMP_CLAUSES|split|4')).
fof(f3458,plain,(
  ( ! [X187,X188] : (sP5(X187,X188)) ) | $spl84),
  inference(cnf_transformation,[],[f3458_D])).
fof(f3458_D,plain,(
  ( ! [X187,X188] : (sP5(X187,X188)) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f9860,plain,(
  ( ! [X210,X211,X209] : (sP4(X209,'T_0') | sP4('F_0',X209) | sP5(X210,X211)) ) | $spl3),
  inference(resolution,[],[f4523,f277])).
fof(f277,plain,(
  ( ! [X0,X1] : (sP5(X0,X1) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f174])).
fof(f9859,plain,(
  ( ! [X208,X206,X207] : (sP4(X206,'T_0') | sP4('F_0',X206) | sP5(X207,X208)) ) | $spl3),
  inference(resolution,[],[f4523,f272])).
fof(f272,plain,(
  ( ! [X0,X1] : (sP5(X0,X1) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f172])).
fof(f9923,plain,(
  ( ! [X0,X3,X1] : (sP4(X0,'T_0') | s(bool,X1) != s(bool,X3) | p(s(bool,X3)) | sP4(X1,X0)) ) | $spl226),
  inference(cnf_transformation,[],[f9923_D])).
fof(f9923_D,plain,(
  ( ! [X0,X3,X1] : (sP4(X0,'T_0') | s(bool,X1) != s(bool,X3) | p(s(bool,X3)) | sP4(X1,X0)) ) <=> ~$spl226),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl226])])).
fof(f9846,plain,(
  ( ! [X177,X175,X176,X174] : (sP4(X174,'T_0') | sP4(X175,X174) | p(s(bool,X176)) | p(s(bool,X177)) | s(bool,X175) != s(bool,X177)) ) | $spl3),
  inference(resolution,[],[f4523,f360])).
fof(f9804,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP4(X35,'T_0') | sP4(X36,X35) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X37) | s(bool,X38) != s(bool,X39)) ) | $spl3),
  inference(resolution,[],[f4523,f223])).
fof(f9924,plain,(
  ( ! [X6,X5,X9] : (sP4(X5,'T_0') | s(bool,X6) != s(bool,X9) | sP4(X6,X5)) ) | $spl228),
  inference(cnf_transformation,[],[f9924_D])).
fof(f9924_D,plain,(
  ( ! [X6,X5,X9] : (sP4(X5,'T_0') | s(bool,X6) != s(bool,X9) | sP4(X6,X5)) ) <=> ~$spl228),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl228])])).
fof(f9803,plain,(
  ( ! [X30,X33,X31,X34,X32] : (sP4(X30,'T_0') | sP4(X31,X30) | p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,X33) | s(bool,X31) != s(bool,X34)) ) | $spl3),
  inference(resolution,[],[f4523,f223])).
fof(f9801,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP4(X20,'T_0') | sP4(X21,X20) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X21) != s(bool,X24)) ) | $spl3),
  inference(resolution,[],[f4523,f222])).
fof(f9800,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP4(X15,'T_0') | sP4(X16,X15) | p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) ) | $spl3),
  inference(resolution,[],[f4523,f221])).
fof(f9799,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP4(X10,'T_0') | sP4(X11,X10) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X12) | s(bool,X13) != s(bool,X14)) ) | $spl3),
  inference(resolution,[],[f4523,f219])).
fof(f9798,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP4(X5,'T_0') | sP4(X6,X5) | p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,X8) | s(bool,X6) != s(bool,X9)) ) | $spl3),
  inference(resolution,[],[f4523,f219])).
fof(f9797,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP4(X0,'T_0') | sP4(X1,X0) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) ) | $spl3),
  inference(resolution,[],[f4523,f217])).
fof(f9752,plain,(
  ( ! [X0,X3,X1] : (sP4(X0,'T_0') | s(bool,X1) != s(bool,X3) | p(s(bool,X3)) | sP4(X0,X1)) ) | $spl222),
  inference(cnf_transformation,[],[f9752_D])).
fof(f9752_D,plain,(
  ( ! [X0,X3,X1] : (sP4(X0,'T_0') | s(bool,X1) != s(bool,X3) | p(s(bool,X3)) | sP4(X0,X1)) ) <=> ~$spl222),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl222])])).
fof(f9670,plain,(
  ( ! [X177,X175,X176,X174] : (sP4(X174,'T_0') | sP4(X174,X175) | p(s(bool,X176)) | p(s(bool,X177)) | s(bool,X175) != s(bool,X177)) ) | $spl3),
  inference(resolution,[],[f4522,f360])).
fof(f9628,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP4(X35,'T_0') | sP4(X35,X36) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X37) | s(bool,X38) != s(bool,X39)) ) | $spl3),
  inference(resolution,[],[f4522,f223])).
fof(f9753,plain,(
  ( ! [X6,X5,X9] : (sP4(X5,'T_0') | s(bool,X6) != s(bool,X9) | sP4(X5,X6)) ) | $spl224),
  inference(cnf_transformation,[],[f9753_D])).
fof(f9753_D,plain,(
  ( ! [X6,X5,X9] : (sP4(X5,'T_0') | s(bool,X6) != s(bool,X9) | sP4(X5,X6)) ) <=> ~$spl224),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl224])])).
fof(f9627,plain,(
  ( ! [X30,X33,X31,X34,X32] : (sP4(X30,'T_0') | sP4(X30,X31) | p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,X33) | s(bool,X31) != s(bool,X34)) ) | $spl3),
  inference(resolution,[],[f4522,f223])).
fof(f9625,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP4(X20,'T_0') | sP4(X20,X21) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X21) != s(bool,X24)) ) | $spl3),
  inference(resolution,[],[f4522,f222])).
fof(f9624,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP4(X15,'T_0') | sP4(X15,X16) | p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) ) | $spl3),
  inference(resolution,[],[f4522,f221])).
fof(f9623,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP4(X10,'T_0') | sP4(X10,X11) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X12) | s(bool,X13) != s(bool,X14)) ) | $spl3),
  inference(resolution,[],[f4522,f219])).
fof(f9622,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP4(X5,'T_0') | sP4(X5,X6) | p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,X8) | s(bool,X6) != s(bool,X9)) ) | $spl3),
  inference(resolution,[],[f4522,f219])).
fof(f9621,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP4(X0,'T_0') | sP4(X0,X1) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) ) | $spl3),
  inference(resolution,[],[f4522,f217])).
fof(f9295,plain,(
  ( ! [X197] : (~sP4('T_0',X197) | sP4(X197,'T_0')) ) | $spl218),
  inference(cnf_transformation,[],[f9295_D])).
fof(f9295_D,plain,(
  ( ! [X197] : (~sP4('T_0',X197) | sP4(X197,'T_0')) ) <=> ~$spl218),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl218])])).
fof(f9282,plain,(
  ( ! [X45,X44] : (~p(s(bool,X44)) | sP4(X45,'T_0') | ~sP4('T_0',X45)) ) | $spl3),
  inference(duplicate_literal_removal,[],[f9266])).
fof(f9266,plain,(
  ( ! [X45,X44] : (~p(s(bool,X44)) | sP4(X45,'T_0') | ~sP4('T_0',X45) | ~p(s(bool,X44))) ) | $spl3),
  inference(superposition,[],[f4521,f323])).
fof(f9283,plain,(
  ( ! [X43,X42] : (~p(s(bool,X42)) | sP4(X43,'T_0') | ~sP4('T_0',X43)) ) | $spl3),
  inference(duplicate_literal_removal,[],[f9265])).
fof(f9265,plain,(
  ( ! [X43,X42] : (~p(s(bool,X42)) | sP4(X43,'T_0') | ~sP4('T_0',X43) | ~p(s(bool,X42))) ) | $spl3),
  inference(superposition,[],[f4521,f325])).
fof(f9302,plain,(
  ( ! [X41,X40] : (~p(s(bool,X40)) | sP4(X41,'T_0') | ~sP4('T_0',X41)) ) | $spl3),
  inference(subsumption_resolution,[],[f9264,f318])).
fof(f9264,plain,(
  ( ! [X41,X40] : (~p(s(bool,X40)) | sP4(X41,'T_0') | ~sP4('T_0',X41) | s(bool,'F_0') = s(bool,X40)) ) | $spl3),
  inference(superposition,[],[f4521,f358])).
fof(f9298,plain,(
  ( ! [X219] : (sP4(X219,'T_0') | ~sP4('F_0',X219)) ) | $spl220),
  inference(cnf_transformation,[],[f9298_D])).
fof(f9298_D,plain,(
  ( ! [X219] : (sP4(X219,'T_0') | ~sP4('F_0',X219)) ) <=> ~$spl220),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl220])])).
fof(f9249,plain,(
  ( ! [X227,X225,X226] : (sP4(X225,'T_0') | ~sP4('F_0',X225) | p(s(bool,X226)) | p(s(bool,X227)) | ~sP5(X226,X227)) ) | $spl3),
  inference(resolution,[],[f4521,f269])).
fof(f9247,plain,(
  ( ! [X220,X219,X221] : (sP4(X219,'T_0') | ~sP4('F_0',X219) | p(s(bool,X220)) | p(s(bool,X221)) | ~sP5(X220,X221)) ) | $spl3),
  inference(resolution,[],[f4521,f273])).
fof(f9238,plain,(
  ( ! [X210,X209] : (sP4(X209,'T_0') | ~sP4('T_0',X209) | ~p(s(bool,X210))) ) | $spl3),
  inference(resolution,[],[f4521,f316])).
fof(f316,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f197])).
fof(f197,plain,(
  ! [X0] : ((~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,'T_0')) & p(s(bool,X0)))))),
  inference(flattening,[],[f196])).
fof(f196,plain,(
  ! [X0] : (((~p(s(bool,'T_0')) | ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,'T_0')) & p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0] : ((p(s(bool,'T_0')) & p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X11] : ((p(s(bool,'T_0')) & p(s(bool,X11))) <=> p(s(bool,X11)))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|1','thm.bool.AND_CLAUSES|split|0')).
fof(f9237,plain,(
  ( ! [X208,X207] : (sP4(X207,'T_0') | ~sP4('T_0',X207) | ~p(s(bool,X208))) ) | $spl3),
  inference(resolution,[],[f4521,f311])).
fof(f311,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f193])).
fof(f193,plain,(
  ! [X0] : ((~p(s(bool,X0)) | ~p(s(bool,'T_0')) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,X0)) & p(s(bool,'T_0')))))),
  inference(flattening,[],[f192])).
fof(f192,plain,(
  ! [X0] : (((~p(s(bool,X0)) | ~p(s(bool,'T_0'))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,X0)) & p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0] : ((p(s(bool,X0)) & p(s(bool,'T_0'))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X11] : ((p(s(bool,X11)) & p(s(bool,'T_0'))) <=> p(s(bool,X11)))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|1','thm.bool.AND_CLAUSES|split|1')).
fof(f9236,plain,(
  ( ! [X206,X205] : (sP4(X205,'T_0') | ~sP4('T_0',X205) | ~p(s(bool,X206))) ) | $spl3),
  inference(resolution,[],[f4521,f301])).
fof(f301,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f187])).
fof(f9235,plain,(
  ( ! [X204,X203] : (sP4(X203,'T_0') | ~sP4('T_0',X203) | p(s(bool,X204))) ) | $spl3),
  inference(resolution,[],[f4521,f300])).
fof(f9234,plain,(
  ( ! [X202,X201] : (sP4(X201,'T_0') | ~sP4('T_0',X201) | p(s(bool,X202))) ) | $spl3),
  inference(resolution,[],[f4521,f297])).
fof(f9233,plain,(
  ( ! [X200,X199] : (sP4(X199,'T_0') | ~sP4('T_0',X199) | p(s(bool,X200))) ) | $spl3),
  inference(resolution,[],[f4521,f294])).
fof(f9232,plain,(
  ( ! [X198,X197] : (sP4(X197,'T_0') | ~sP4('T_0',X197) | ~p(s(bool,X198))) ) | $spl3),
  inference(resolution,[],[f4521,f292])).
fof(f292,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f181])).
fof(f9286,plain,(
  ( ! [X0,X3,X1] : (sP4(X0,'T_0') | s(bool,X1) != s(bool,X3) | ~p(s(bool,X3)) | ~sP4(X1,X0)) ) | $spl212),
  inference(cnf_transformation,[],[f9286_D])).
fof(f9286_D,plain,(
  ( ! [X0,X3,X1] : (sP4(X0,'T_0') | s(bool,X1) != s(bool,X3) | ~p(s(bool,X3)) | ~sP4(X1,X0)) ) <=> ~$spl212),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl212])])).
fof(f9226,plain,(
  ( ! [X185,X182,X184,X183] : (sP4(X182,'T_0') | ~sP4(X183,X182) | p(s(bool,X184)) | ~p(s(bool,X185)) | s(bool,X183) != s(bool,X185)) ) | $spl3),
  inference(resolution,[],[f4521,f360])).
fof(f9289,plain,(
  ( ! [X26,X25] : (sP4(X25,'T_0') | ~sP4(X26,X25)) ) | $spl216),
  inference(cnf_transformation,[],[f9289_D])).
fof(f9289_D,plain,(
  ( ! [X26,X25] : (sP4(X25,'T_0') | ~sP4(X26,X25)) ) <=> ~$spl216),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl216])])).
fof(f9225,plain,(
  ( ! [X180,X178,X181,X179] : (sP4(X178,'T_0') | ~sP4(X179,X178) | ~p(s(bool,X180)) | p(s(bool,X181)) | s(bool,X180) != s(bool,X181)) ) | $spl3),
  inference(resolution,[],[f4521,f360])).
fof(f9181,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP4(X35,'T_0') | ~sP4(X36,X35) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X38) | s(bool,X37) != s(bool,X39)) ) | $spl3),
  inference(resolution,[],[f4521,f223])).
fof(f9180,plain,(
  ( ! [X30,X33,X31,X34,X32] : (sP4(X30,'T_0') | ~sP4(X31,X30) | p(s(bool,X32)) | ~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X31) != s(bool,X33)) ) | $spl3),
  inference(resolution,[],[f4521,f222])).
fof(f9288,plain,(
  ( ! [X26,X29,X27,X25] : (sP4(X25,'T_0') | ~sP4(X26,X25) | ~p(s(bool,X27)) | p(s(bool,X29)) | s(bool,X27) != s(bool,X29)) ) | $spl3),
  inference(subsumption_resolution,[],[f9179,f360])).
fof(f9179,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP4(X25,'T_0') | ~sP4(X26,X25) | ~p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X27) != s(bool,X29)) ) | $spl3),
  inference(resolution,[],[f4521,f222])).
fof(f9178,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP4(X20,'T_0') | ~sP4(X21,X20) | p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X24) | s(bool,X21) != s(bool,X23)) ) | $spl3),
  inference(resolution,[],[f4521,f221])).
fof(f9287,plain,(
  ( ! [X6,X5,X9] : (sP4(X5,'T_0') | s(bool,X6) != s(bool,X9) | ~sP4(X6,X5)) ) | $spl214),
  inference(cnf_transformation,[],[f9287_D])).
fof(f9287_D,plain,(
  ( ! [X6,X5,X9] : (sP4(X5,'T_0') | s(bool,X6) != s(bool,X9) | ~sP4(X6,X5)) ) <=> ~$spl214),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl214])])).
fof(f9177,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP4(X15,'T_0') | ~sP4(X16,X15) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X16) != s(bool,X19) | s(bool,X17) != s(bool,X18)) ) | $spl3),
  inference(resolution,[],[f4521,f221])).
fof(f9176,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP4(X10,'T_0') | ~sP4(X11,X10) | ~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,X14)) ) | $spl3),
  inference(resolution,[],[f4521,f219])).
fof(f9175,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP4(X5,'T_0') | ~sP4(X6,X5) | ~p(s(bool,X7)) | p(s(bool,X8)) | s(bool,X6) != s(bool,X9) | s(bool,X7) != s(bool,X8)) ) | $spl3),
  inference(resolution,[],[f4521,f217])).
fof(f9174,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP4(X0,'T_0') | ~sP4(X1,X0) | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) ) | $spl3),
  inference(resolution,[],[f4521,f217])).
fof(f4884,plain,(
  ( ! [X4,X5] : (sP4(X4,'T_0') | p(s(bool,X5)) | sP4(X5,X4)) ) | $spl140),
  inference(cnf_transformation,[],[f4884_D])).
fof(f4884_D,plain,(
  ( ! [X4,X5] : (sP4(X4,'T_0') | p(s(bool,X5)) | sP4(X5,X4)) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f9156,plain,(
  ( ! [X12,X13] : (sP4(X12,'T_0') | sP4(X13,X12) | p(s(bool,X13)) | sP4(sK8,'F_0')) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f8860,f2545])).
fof(f2545,plain,(
  ( ! [X74,X75] : (~sP4(X75,X74) | p(s(bool,X75)) | sP4(sK8,X74)) ) | $spl3),
  inference(resolution,[],[f1544,f253])).
fof(f9154,plain,(
  ( ! [X8,X9] : (sP4(X8,'T_0') | sP4(X9,X8) | p(s(bool,X9)) | sP4('F_0',sK8)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f8860,f2428])).
fof(f2428,plain,(
  ( ! [X74,X75] : (~sP4(X75,X74) | p(s(bool,X75)) | sP4(X74,sK8)) ) | $spl3),
  inference(resolution,[],[f1543,f253])).
fof(f9080,plain,(
  ( ! [X12,X13] : (sP4(X12,'T_0') | sP4(X12,X13) | p(s(bool,X13)) | sP4(sK8,'F_0')) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f8857,f2545])).
fof(f9078,plain,(
  ( ! [X8,X9] : (sP4(X8,'T_0') | sP4(X8,X9) | p(s(bool,X9)) | sP4('F_0',sK8)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f8857,f2428])).
fof(f9058,plain,(
  ( ! [X17,X15,X16] : (sP4(X15,'T_0') | sP4(X16,X15) | p(s(bool,X16)) | p(s(bool,X17))) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f9048,f307])).
fof(f9048,plain,(
  ( ! [X17,X15,X16] : (sP4(X15,'T_0') | sP4(X16,X15) | p(s(bool,'F_0')) | p(s(bool,X16)) | p(s(bool,X17))) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f8859,f269])).
fof(f8859,plain,(
  ( ! [X202,X203,X201] : (sP4(X201,'T_0') | sP4(X202,X201) | sP5(X202,X203)) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f8806,f7584])).
fof(f8806,plain,(
  ( ! [X202,X203,X201] : (sP4('T_0','T_0') | sP4(X201,'T_0') | sP4(X202,X201) | sP5(X202,X203)) ) | $spl3),
  inference(resolution,[],[f8520,f8127])).
fof(f9057,plain,(
  ( ! [X14,X12,X13] : (sP4(X12,'T_0') | sP4(X13,X12) | p(s(bool,X13)) | p(s(bool,X14))) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f9047,f307])).
fof(f9047,plain,(
  ( ! [X14,X12,X13] : (sP4(X12,'T_0') | sP4(X13,X12) | p(s(bool,'F_0')) | p(s(bool,X13)) | p(s(bool,X14))) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f8859,f273])).
fof(f9046,plain,(
  ( ! [X10,X11,X9] : (sP4(X9,'T_0') | sP4(X10,X9) | p(s(bool,X10)) | sP4(X11,sK8)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f8859,f2514])).
fof(f9044,plain,(
  ( ! [X4,X5,X3] : (sP4(X3,'T_0') | sP4(X4,X3) | p(s(bool,X4)) | sP4(sK8,X5)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f8859,f2631])).
fof(f9009,plain,(
  ( ! [X17,X15,X16] : (sP4(X15,'T_0') | sP4(X16,X15) | p(s(bool,X17)) | p(s(bool,X16))) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f8999,f307])).
fof(f8999,plain,(
  ( ! [X17,X15,X16] : (sP4(X15,'T_0') | sP4(X16,X15) | p(s(bool,'F_0')) | p(s(bool,X17)) | p(s(bool,X16))) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f8858,f269])).
fof(f8858,plain,(
  ( ! [X198,X200,X199] : (sP4(X198,'T_0') | sP4(X199,X198) | sP5(X200,X199)) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f8805,f7584])).
fof(f8805,plain,(
  ( ! [X198,X200,X199] : (sP4('T_0','T_0') | sP4(X198,'T_0') | sP4(X199,X198) | sP5(X200,X199)) ) | $spl3),
  inference(resolution,[],[f8520,f8126])).
fof(f9008,plain,(
  ( ! [X14,X12,X13] : (sP4(X12,'T_0') | sP4(X13,X12) | p(s(bool,X14)) | p(s(bool,X13))) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f8998,f307])).
fof(f8998,plain,(
  ( ! [X14,X12,X13] : (sP4(X12,'T_0') | sP4(X13,X12) | p(s(bool,'F_0')) | p(s(bool,X14)) | p(s(bool,X13))) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f8858,f273])).
fof(f8996,plain,(
  ( ! [X6,X8,X7] : (sP4(X6,'T_0') | sP4(X7,X6) | p(s(bool,X7)) | sP4(X8,sK8)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f8858,f2515])).
fof(f8994,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,'T_0') | sP4(X1,X0) | p(s(bool,X1)) | sP4(sK8,X2)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f8858,f2632])).
fof(f8960,plain,(
  ( ! [X17,X15,X16] : (sP4(X15,'T_0') | sP4(X15,X16) | p(s(bool,X16)) | p(s(bool,X17))) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f8949,f307])).
fof(f8949,plain,(
  ( ! [X17,X15,X16] : (sP4(X15,'T_0') | sP4(X15,X16) | p(s(bool,'F_0')) | p(s(bool,X16)) | p(s(bool,X17))) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f8856,f269])).
fof(f8959,plain,(
  ( ! [X14,X12,X13] : (sP4(X12,'T_0') | sP4(X12,X13) | p(s(bool,X13)) | p(s(bool,X14))) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f8948,f307])).
fof(f8948,plain,(
  ( ! [X14,X12,X13] : (sP4(X12,'T_0') | sP4(X12,X13) | p(s(bool,'F_0')) | p(s(bool,X13)) | p(s(bool,X14))) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f8856,f273])).
fof(f8947,plain,(
  ( ! [X10,X11,X9] : (sP4(X9,'T_0') | sP4(X9,X10) | p(s(bool,X10)) | sP4(X11,sK8)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f8856,f2514])).
fof(f8945,plain,(
  ( ! [X4,X5,X3] : (sP4(X3,'T_0') | sP4(X3,X4) | p(s(bool,X4)) | sP4(sK8,X5)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f8856,f2631])).
fof(f8910,plain,(
  ( ! [X17,X15,X16] : (sP4(X15,'T_0') | sP4(X15,X16) | p(s(bool,X17)) | p(s(bool,X16))) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f8899,f307])).
fof(f8899,plain,(
  ( ! [X17,X15,X16] : (sP4(X15,'T_0') | sP4(X15,X16) | p(s(bool,'F_0')) | p(s(bool,X17)) | p(s(bool,X16))) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f8855,f269])).
fof(f8909,plain,(
  ( ! [X14,X12,X13] : (sP4(X12,'T_0') | sP4(X12,X13) | p(s(bool,X14)) | p(s(bool,X13))) ) | ($spl3 | $spl179)),
  inference(subsumption_resolution,[],[f8898,f307])).
fof(f8898,plain,(
  ( ! [X14,X12,X13] : (sP4(X12,'T_0') | sP4(X12,X13) | p(s(bool,'F_0')) | p(s(bool,X14)) | p(s(bool,X13))) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f8855,f273])).
fof(f8896,plain,(
  ( ! [X6,X8,X7] : (sP4(X6,'T_0') | sP4(X6,X7) | p(s(bool,X7)) | sP4(X8,sK8)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f8855,f2515])).
fof(f8894,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,'T_0') | sP4(X0,X1) | p(s(bool,X1)) | sP4(sK8,X2)) ) | ($spl3 | $spl179)),
  inference(resolution,[],[f8855,f2632])).
fof(f8586,plain,(
  ( ! [X227,X225,X226] : (sP4(X225,'T_0') | p(s(bool,X226)) | p(s(bool,X227)) | ~sP5(X226,X227)) ) | $spl3),
  inference(subsumption_resolution,[],[f8539,f7721])).
fof(f8539,plain,(
  ( ! [X227,X225,X226] : (sP4(X225,'T_0') | ~sP4(X225,'F_0') | p(s(bool,X226)) | p(s(bool,X227)) | ~sP5(X226,X227)) ) | $spl3),
  inference(resolution,[],[f4520,f269])).
fof(f8584,plain,(
  ( ! [X220,X219,X221] : (sP4(X219,'T_0') | p(s(bool,X220)) | p(s(bool,X221)) | ~sP5(X220,X221)) ) | $spl3),
  inference(subsumption_resolution,[],[f8537,f7721])).
fof(f8537,plain,(
  ( ! [X220,X219,X221] : (sP4(X219,'T_0') | ~sP4(X219,'F_0') | p(s(bool,X220)) | p(s(bool,X221)) | ~sP5(X220,X221)) ) | $spl3),
  inference(resolution,[],[f4520,f273])).
fof(f8576,plain,(
  ( ! [X0,X3,X1] : (sP4(X0,'T_0') | s(bool,X1) != s(bool,X3) | ~p(s(bool,X3)) | ~sP4(X0,X1)) ) | $spl204),
  inference(cnf_transformation,[],[f8576_D])).
fof(f8576_D,plain,(
  ( ! [X0,X3,X1] : (sP4(X0,'T_0') | s(bool,X1) != s(bool,X3) | ~p(s(bool,X3)) | ~sP4(X0,X1)) ) <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f8516,plain,(
  ( ! [X185,X182,X184,X183] : (sP4(X182,'T_0') | ~sP4(X182,X183) | p(s(bool,X184)) | ~p(s(bool,X185)) | s(bool,X183) != s(bool,X185)) ) | $spl3),
  inference(resolution,[],[f4520,f360])).
fof(f8579,plain,(
  ( ! [X26,X25] : (sP4(X25,'T_0') | ~sP4(X25,X26)) ) | $spl208),
  inference(cnf_transformation,[],[f8579_D])).
fof(f8579_D,plain,(
  ( ! [X26,X25] : (sP4(X25,'T_0') | ~sP4(X25,X26)) ) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f8515,plain,(
  ( ! [X180,X178,X181,X179] : (sP4(X178,'T_0') | ~sP4(X178,X179) | ~p(s(bool,X180)) | p(s(bool,X181)) | s(bool,X180) != s(bool,X181)) ) | $spl3),
  inference(resolution,[],[f4520,f360])).
fof(f8471,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP4(X35,'T_0') | ~sP4(X35,X36) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X38) | s(bool,X37) != s(bool,X39)) ) | $spl3),
  inference(resolution,[],[f4520,f223])).
fof(f8470,plain,(
  ( ! [X30,X33,X31,X34,X32] : (sP4(X30,'T_0') | ~sP4(X30,X31) | p(s(bool,X32)) | ~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X31) != s(bool,X33)) ) | $spl3),
  inference(resolution,[],[f4520,f222])).
fof(f8578,plain,(
  ( ! [X26,X29,X27,X25] : (sP4(X25,'T_0') | ~sP4(X25,X26) | ~p(s(bool,X27)) | p(s(bool,X29)) | s(bool,X27) != s(bool,X29)) ) | $spl3),
  inference(subsumption_resolution,[],[f8469,f360])).
fof(f8469,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP4(X25,'T_0') | ~sP4(X25,X26) | ~p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X27) != s(bool,X29)) ) | $spl3),
  inference(resolution,[],[f4520,f222])).
fof(f8468,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP4(X20,'T_0') | ~sP4(X20,X21) | p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X24) | s(bool,X21) != s(bool,X23)) ) | $spl3),
  inference(resolution,[],[f4520,f221])).
fof(f8577,plain,(
  ( ! [X6,X5,X9] : (sP4(X5,'T_0') | s(bool,X6) != s(bool,X9) | ~sP4(X5,X6)) ) | $spl206),
  inference(cnf_transformation,[],[f8577_D])).
fof(f8577_D,plain,(
  ( ! [X6,X5,X9] : (sP4(X5,'T_0') | s(bool,X6) != s(bool,X9) | ~sP4(X5,X6)) ) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f8467,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP4(X15,'T_0') | ~sP4(X15,X16) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X16) != s(bool,X19) | s(bool,X17) != s(bool,X18)) ) | $spl3),
  inference(resolution,[],[f4520,f221])).
fof(f8466,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP4(X10,'T_0') | ~sP4(X10,X11) | ~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,X14)) ) | $spl3),
  inference(resolution,[],[f4520,f219])).
fof(f8465,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP4(X5,'T_0') | ~sP4(X5,X6) | ~p(s(bool,X7)) | p(s(bool,X8)) | s(bool,X6) != s(bool,X9) | s(bool,X7) != s(bool,X8)) ) | $spl3),
  inference(resolution,[],[f4520,f217])).
fof(f8464,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP4(X0,'T_0') | ~sP4(X0,X1) | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) ) | $spl3),
  inference(resolution,[],[f4520,f217])).
fof(f8374,plain,(
  $false | ($spl4 | $spl200)),
  inference(subsumption_resolution,[],[f8371,f338])).
fof(f8371,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl200)),
  inference(resolution,[],[f8292,f1629])).
fof(f8359,plain,(
  $false | ($spl179 | $spl200)),
  inference(resolution,[],[f8292,f7584])).
fof(f8358,plain,(
  ( ! [X1] : (sP5('T_0',X1)) ) | $spl202),
  inference(cnf_transformation,[],[f8358_D])).
fof(f8358_D,plain,(
  ( ! [X1] : (sP5('T_0',X1)) ) <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f8352,plain,(
  ( ! [X0,X1] : (sP4('T_0',X0) | sP5('T_0',X1)) ) | $spl3),
  inference(factoring,[],[f8127])).
fof(f8291,plain,(
  ( ! [X1] : (sP5(X1,'T_0')) ) | $spl198),
  inference(cnf_transformation,[],[f8291_D])).
fof(f8291_D,plain,(
  ( ! [X1] : (sP5(X1,'T_0')) ) <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f8285,plain,(
  ( ! [X0,X1] : (sP4('T_0',X0) | sP5(X1,'T_0')) ) | $spl3),
  inference(factoring,[],[f8126])).
fof(f8226,plain,(
  ( ! [X200] : (sP4('F_0',X200) | sP4('T_0',X200)) ) | $spl196),
  inference(cnf_transformation,[],[f8226_D])).
fof(f8226_D,plain,(
  ( ! [X200] : (sP4('F_0',X200) | sP4('T_0',X200)) ) <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f8199,plain,(
  ( ! [X50,X51] : (p(s(bool,X50)) | sP4('T_0',X51) | sP4('F_0',X51)) ) | $spl3),
  inference(duplicate_literal_removal,[],[f8186])).
fof(f8186,plain,(
  ( ! [X50,X51] : (p(s(bool,X50)) | sP4('T_0',X51) | sP4('F_0',X51) | p(s(bool,X50))) ) | $spl3),
  inference(superposition,[],[f3990,f319])).
fof(f8200,plain,(
  ( ! [X48,X49] : (p(s(bool,X48)) | sP4('T_0',X49) | sP4('F_0',X49)) ) | $spl3),
  inference(duplicate_literal_removal,[],[f8185])).
fof(f8185,plain,(
  ( ! [X48,X49] : (p(s(bool,X48)) | sP4('T_0',X49) | sP4('F_0',X49) | p(s(bool,X48))) ) | $spl3),
  inference(superposition,[],[f3990,f321])).
fof(f8233,plain,(
  ( ! [X47,X46] : (p(s(bool,X46)) | sP4('T_0',X47) | sP4('F_0',X47)) ) | $spl3),
  inference(subsumption_resolution,[],[f8184,f322])).
fof(f8184,plain,(
  ( ! [X47,X46] : (p(s(bool,X46)) | sP4('T_0',X47) | sP4('F_0',X47) | s(bool,'T_0') = s(bool,X46)) ) | $spl3),
  inference(superposition,[],[f3990,f358])).
fof(f8162,plain,(
  ( ! [X215,X214] : (sP4('T_0',X214) | sP4('F_0',X214) | p(s(bool,X215))) ) | $spl3),
  inference(resolution,[],[f3990,f331])).
fof(f8161,plain,(
  ( ! [X212,X213] : (sP4('T_0',X212) | sP4('F_0',X212) | p(s(bool,X213))) ) | $spl3),
  inference(resolution,[],[f3990,f314])).
fof(f8160,plain,(
  ( ! [X210,X211] : (sP4('T_0',X210) | sP4('F_0',X210) | p(s(bool,X211))) ) | $spl3),
  inference(resolution,[],[f3990,f307])).
fof(f8227,plain,(
  ( ! [X208,X209] : (sP4('T_0',X208) | sP4('F_0',X208) | p(s(bool,X209))) ) | $spl3),
  inference(subsumption_resolution,[],[f8159,f361])).
fof(f8159,plain,(
  ( ! [X208,X209] : (sP4('T_0',X208) | sP4('F_0',X208) | p(s(bool,X209)) | ~p(s(bool,'T_0'))) ) | $spl3),
  inference(resolution,[],[f3990,f293])).
fof(f8158,plain,(
  ( ! [X206,X207] : (sP4('T_0',X206) | sP4('F_0',X206) | ~p(s(bool,X207))) ) | $spl3),
  inference(resolution,[],[f3990,f289])).
fof(f8157,plain,(
  ( ! [X204,X205,X203] : (sP4('T_0',X203) | sP4('F_0',X203) | sP5(X204,X205)) ) | $spl3),
  inference(resolution,[],[f3990,f277])).
fof(f8156,plain,(
  ( ! [X202,X200,X201] : (sP4('T_0',X200) | sP4('F_0',X200) | sP5(X201,X202)) ) | $spl3),
  inference(resolution,[],[f3990,f272])).
fof(f8222,plain,(
  ( ! [X0,X3,X1] : (sP4('T_0',X0) | s(bool,X1) != s(bool,X3) | p(s(bool,X3)) | sP4(X1,X0)) ) | $spl192),
  inference(cnf_transformation,[],[f8222_D])).
fof(f8222_D,plain,(
  ( ! [X0,X3,X1] : (sP4('T_0',X0) | s(bool,X1) != s(bool,X3) | p(s(bool,X3)) | sP4(X1,X0)) ) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f8145,plain,(
  ( ! [X177,X175,X176,X174] : (sP4('T_0',X174) | sP4(X175,X174) | p(s(bool,X176)) | p(s(bool,X177)) | s(bool,X175) != s(bool,X177)) ) | $spl3),
  inference(resolution,[],[f3990,f360])).
fof(f8103,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP4('T_0',X35) | sP4(X36,X35) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X37) | s(bool,X38) != s(bool,X39)) ) | $spl3),
  inference(resolution,[],[f3990,f223])).
fof(f8223,plain,(
  ( ! [X6,X5,X9] : (sP4('T_0',X5) | s(bool,X6) != s(bool,X9) | sP4(X6,X5)) ) | $spl194),
  inference(cnf_transformation,[],[f8223_D])).
fof(f8223_D,plain,(
  ( ! [X6,X5,X9] : (sP4('T_0',X5) | s(bool,X6) != s(bool,X9) | sP4(X6,X5)) ) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f8102,plain,(
  ( ! [X30,X33,X31,X34,X32] : (sP4('T_0',X30) | sP4(X31,X30) | p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,X33) | s(bool,X31) != s(bool,X34)) ) | $spl3),
  inference(resolution,[],[f3990,f223])).
fof(f8100,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP4('T_0',X20) | sP4(X21,X20) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X21) != s(bool,X24)) ) | $spl3),
  inference(resolution,[],[f3990,f222])).
fof(f8099,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP4('T_0',X15) | sP4(X16,X15) | p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) ) | $spl3),
  inference(resolution,[],[f3990,f221])).
fof(f8098,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP4('T_0',X10) | sP4(X11,X10) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X12) | s(bool,X13) != s(bool,X14)) ) | $spl3),
  inference(resolution,[],[f3990,f219])).
fof(f8097,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP4('T_0',X5) | sP4(X6,X5) | p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,X8) | s(bool,X6) != s(bool,X9)) ) | $spl3),
  inference(resolution,[],[f3990,f219])).
fof(f8096,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP4('T_0',X0) | sP4(X1,X0) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) ) | $spl3),
  inference(resolution,[],[f3990,f217])).
fof(f7847,plain,(
  sP4('const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0')),'F_0') | $spl190),
  inference(cnf_transformation,[],[f7847_D])).
fof(f7847_D,plain,(
  sP4('const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0')),'F_0') <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f7833,plain,(
  sP4('const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0')),'F_0') | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6)))) | ($spl3 | $spl4)),
  inference(superposition,[],[f7587,f357])).
fof(f7740,plain,(
  ( ! [X149,X148] : (sP4('F_0','F_0') | p(s(bool,X148)) | p(s(bool,X149)) | ~sP5(X148,X149)) ) | $spl3),
  inference(resolution,[],[f7582,f269])).
fof(f7738,plain,(
  ( ! [X144,X145] : (sP4('F_0','F_0') | p(s(bool,X144)) | p(s(bool,X145)) | ~sP5(X144,X145)) ) | $spl3),
  inference(resolution,[],[f7582,f273])).
fof(f7718,plain,(
  ( ! [X132,X130,X131] : (sP4(X130,'F_0') | p(s(bool,X131)) | ~p(s(bool,X132)) | s(bool,X130) != s(bool,X132)) ) | $spl3),
  inference(resolution,[],[f7582,f360])).
fof(f7717,plain,(
  ( ! [X127,X128,X129] : (sP4(X127,'F_0') | ~p(s(bool,X128)) | p(s(bool,X129)) | s(bool,X128) != s(bool,X129)) ) | $spl3),
  inference(resolution,[],[f7582,f360])).
fof(f7673,plain,(
  ( ! [X30,X28,X31,X29] : (sP4(X28,'F_0') | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) ) | $spl3),
  inference(resolution,[],[f7582,f223])).
fof(f7672,plain,(
  ( ! [X26,X24,X27,X25] : (sP4(X24,'F_0') | p(s(bool,X25)) | ~p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X24) != s(bool,X26)) ) | $spl3),
  inference(resolution,[],[f7582,f222])).
fof(f7779,plain,(
  ( ! [X23,X21,X20] : (sP4(X20,'F_0') | ~p(s(bool,X21)) | p(s(bool,X23)) | s(bool,X21) != s(bool,X23)) ) | $spl3),
  inference(subsumption_resolution,[],[f7671,f360])).
fof(f7671,plain,(
  ( ! [X23,X21,X22,X20] : (sP4(X20,'F_0') | ~p(s(bool,X21)) | ~p(s(bool,X22)) | p(s(bool,X23)) | s(bool,X21) != s(bool,X23)) ) | $spl3),
  inference(resolution,[],[f7582,f222])).
fof(f7670,plain,(
  ( ! [X19,X17,X18,X16] : (sP4(X16,'F_0') | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) ) | $spl3),
  inference(resolution,[],[f7582,f221])).
fof(f7778,plain,(
  ( ! [X4,X7] : (sP4(X4,'F_0') | s(bool,X4) != s(bool,X7)) ) | $spl184),
  inference(cnf_transformation,[],[f7778_D])).
fof(f7778_D,plain,(
  ( ! [X4,X7] : (sP4(X4,'F_0') | s(bool,X4) != s(bool,X7)) ) <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f7669,plain,(
  ( ! [X14,X12,X15,X13] : (sP4(X12,'F_0') | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X12) != s(bool,X15) | s(bool,X13) != s(bool,X14)) ) | $spl3),
  inference(resolution,[],[f7582,f221])).
fof(f7668,plain,(
  ( ! [X10,X8,X11,X9] : (sP4(X8,'F_0') | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) ) | $spl3),
  inference(resolution,[],[f7582,f219])).
fof(f7667,plain,(
  ( ! [X6,X4,X7,X5] : (sP4(X4,'F_0') | ~p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X4) != s(bool,X7) | s(bool,X5) != s(bool,X6)) ) | $spl3),
  inference(resolution,[],[f7582,f217])).
fof(f7666,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(X0,'F_0') | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl3),
  inference(resolution,[],[f7582,f217])).
fof(f7603,plain,(
  $false | $spl178),
  inference(subsumption_resolution,[],[f7599,f338])).
fof(f7599,plain,(
  ~p(s(bool,'T_0')) | $spl178),
  inference(duplicate_literal_removal,[],[f7590])).
fof(f7590,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'T_0')) | $spl178),
  inference(resolution,[],[f7585,f261])).
fof(f7602,plain,(
  $false | $spl178),
  inference(subsumption_resolution,[],[f7600,f338])).
fof(f7600,plain,(
  ~p(s(bool,'T_0')) | $spl178),
  inference(duplicate_literal_removal,[],[f7589])).
fof(f7589,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'T_0')) | $spl178),
  inference(resolution,[],[f7585,f254])).
fof(f7585,plain,(
  sP4('T_0','T_0') | $spl178),
  inference(cnf_transformation,[],[f7585_D])).
fof(f7585_D,plain,(
  sP4('T_0','T_0') <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f7586,plain,(
  ( ! [X8] : (~p(s(bool,X8)) | sP4(X8,'F_0')) ) | $spl180),
  inference(cnf_transformation,[],[f7586_D])).
fof(f7586_D,plain,(
  ( ! [X8] : (~p(s(bool,X8)) | sP4(X8,'F_0')) ) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f7565,plain,(
  ( ! [X8] : (sP4(X8,'F_0') | ~p(s(bool,X8)) | sP4('T_0','T_0')) ) | $spl3),
  inference(resolution,[],[f7491,f3987])).
fof(f7548,plain,(
  ( ! [X198] : (sP4('T_0',X198) | sP4(X198,'F_0')) ) | $spl176),
  inference(cnf_transformation,[],[f7548_D])).
fof(f7548_D,plain,(
  ( ! [X198] : (sP4('T_0',X198) | sP4(X198,'F_0')) ) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f7527,plain,(
  ( ! [X50,X51] : (p(s(bool,X50)) | sP4('T_0',X51) | sP4(X51,'F_0')) ) | $spl3),
  inference(duplicate_literal_removal,[],[f7514])).
fof(f7514,plain,(
  ( ! [X50,X51] : (p(s(bool,X50)) | sP4('T_0',X51) | sP4(X51,'F_0') | p(s(bool,X50))) ) | $spl3),
  inference(superposition,[],[f3989,f319])).
fof(f7528,plain,(
  ( ! [X48,X49] : (p(s(bool,X48)) | sP4('T_0',X49) | sP4(X49,'F_0')) ) | $spl3),
  inference(duplicate_literal_removal,[],[f7513])).
fof(f7513,plain,(
  ( ! [X48,X49] : (p(s(bool,X48)) | sP4('T_0',X49) | sP4(X49,'F_0') | p(s(bool,X48))) ) | $spl3),
  inference(superposition,[],[f3989,f321])).
fof(f7555,plain,(
  ( ! [X47,X46] : (p(s(bool,X46)) | sP4('T_0',X47) | sP4(X47,'F_0')) ) | $spl3),
  inference(subsumption_resolution,[],[f7512,f322])).
fof(f7512,plain,(
  ( ! [X47,X46] : (p(s(bool,X46)) | sP4('T_0',X47) | sP4(X47,'F_0') | s(bool,'T_0') = s(bool,X46)) ) | $spl3),
  inference(superposition,[],[f3989,f358])).
fof(f7490,plain,(
  ( ! [X212,X213] : (sP4('T_0',X212) | sP4(X212,'F_0') | p(s(bool,X213))) ) | $spl3),
  inference(resolution,[],[f3989,f331])).
fof(f7489,plain,(
  ( ! [X210,X211] : (sP4('T_0',X210) | sP4(X210,'F_0') | p(s(bool,X211))) ) | $spl3),
  inference(resolution,[],[f3989,f314])).
fof(f7488,plain,(
  ( ! [X208,X209] : (sP4('T_0',X208) | sP4(X208,'F_0') | p(s(bool,X209))) ) | $spl3),
  inference(resolution,[],[f3989,f307])).
fof(f7549,plain,(
  ( ! [X206,X207] : (sP4('T_0',X206) | sP4(X206,'F_0') | p(s(bool,X207))) ) | $spl3),
  inference(subsumption_resolution,[],[f7487,f361])).
fof(f7487,plain,(
  ( ! [X206,X207] : (sP4('T_0',X206) | sP4(X206,'F_0') | p(s(bool,X207)) | ~p(s(bool,'T_0'))) ) | $spl3),
  inference(resolution,[],[f3989,f293])).
fof(f7486,plain,(
  ( ! [X204,X205] : (sP4('T_0',X204) | sP4(X204,'F_0') | ~p(s(bool,X205))) ) | $spl3),
  inference(resolution,[],[f3989,f289])).
fof(f7485,plain,(
  ( ! [X202,X203,X201] : (sP4('T_0',X201) | sP4(X201,'F_0') | sP5(X202,X203)) ) | $spl3),
  inference(resolution,[],[f3989,f277])).
fof(f7484,plain,(
  ( ! [X198,X200,X199] : (sP4('T_0',X198) | sP4(X198,'F_0') | sP5(X199,X200)) ) | $spl3),
  inference(resolution,[],[f3989,f272])).
fof(f7544,plain,(
  ( ! [X0,X3,X1] : (sP4('T_0',X0) | s(bool,X1) != s(bool,X3) | p(s(bool,X3)) | sP4(X0,X1)) ) | $spl172),
  inference(cnf_transformation,[],[f7544_D])).
fof(f7544_D,plain,(
  ( ! [X0,X3,X1] : (sP4('T_0',X0) | s(bool,X1) != s(bool,X3) | p(s(bool,X3)) | sP4(X0,X1)) ) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f7474,plain,(
  ( ! [X177,X175,X176,X174] : (sP4('T_0',X174) | sP4(X174,X175) | p(s(bool,X176)) | p(s(bool,X177)) | s(bool,X175) != s(bool,X177)) ) | $spl3),
  inference(resolution,[],[f3989,f360])).
fof(f7432,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP4('T_0',X35) | sP4(X35,X36) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X37) | s(bool,X38) != s(bool,X39)) ) | $spl3),
  inference(resolution,[],[f3989,f223])).
fof(f7545,plain,(
  ( ! [X6,X5,X9] : (sP4('T_0',X5) | s(bool,X6) != s(bool,X9) | sP4(X5,X6)) ) | $spl174),
  inference(cnf_transformation,[],[f7545_D])).
fof(f7545_D,plain,(
  ( ! [X6,X5,X9] : (sP4('T_0',X5) | s(bool,X6) != s(bool,X9) | sP4(X5,X6)) ) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f7431,plain,(
  ( ! [X30,X33,X31,X34,X32] : (sP4('T_0',X30) | sP4(X30,X31) | p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,X33) | s(bool,X31) != s(bool,X34)) ) | $spl3),
  inference(resolution,[],[f3989,f223])).
fof(f7429,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP4('T_0',X20) | sP4(X20,X21) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X21) != s(bool,X24)) ) | $spl3),
  inference(resolution,[],[f3989,f222])).
fof(f7428,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP4('T_0',X15) | sP4(X15,X16) | p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) ) | $spl3),
  inference(resolution,[],[f3989,f221])).
fof(f7427,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP4('T_0',X10) | sP4(X10,X11) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X12) | s(bool,X13) != s(bool,X14)) ) | $spl3),
  inference(resolution,[],[f3989,f219])).
fof(f7426,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP4('T_0',X5) | sP4(X5,X6) | p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,X8) | s(bool,X6) != s(bool,X9)) ) | $spl3),
  inference(resolution,[],[f3989,f219])).
fof(f7425,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP4('T_0',X0) | sP4(X0,X1) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) ) | $spl3),
  inference(resolution,[],[f3989,f217])).
fof(f7069,plain,(
  ( ! [X0] : (sP4(X0,'const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0'))) | sP4(sK8,X0)) ) | $spl170),
  inference(cnf_transformation,[],[f7069_D])).
fof(f7069_D,plain,(
  ( ! [X0] : (sP4(X0,'const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0'))) | sP4(sK8,X0)) ) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f7053,plain,(
  ( ! [X0] : (sP4(X0,'const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0'))) | sP4(sK8,X0) | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6))))) ) | ($spl3 | $spl4)),
  inference(superposition,[],[f5557,f357])).
fof(f5557,plain,(
  ( ! [X222] : (sP4(X222,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7))) | sP4(sK8,X222)) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f2548,f369])).
fof(f6340,plain,(
  ( ! [X10,X8,X11] : (sP3(X10,X8,X11) | ~sP0(X11,X8,X10) | ~sP2(X10,X11,X8) | p(s(bool,X8))) ) | $spl162),
  inference(cnf_transformation,[],[f6340_D])).
fof(f6340_D,plain,(
  ( ! [X10,X8,X11] : (sP3(X10,X8,X11) | ~sP0(X11,X8,X10) | ~sP2(X10,X11,X8) | p(s(bool,X8))) ) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f7028,plain,(
  ( ! [X10,X8,X11,X9] : (sP4(sK8,X8) | sP3(X10,X9,X11) | p(s(bool,X9)) | ~sP2(X10,X11,X9) | ~sP0(X11,X9,X10)) ) | $spl3),
  inference(subsumption_resolution,[],[f7014,f2546])).
fof(f2546,plain,(
  ( ! [X76,X77] : (~sP4(X76,X77) | p(s(bool,X77)) | sP4(sK8,X76)) ) | $spl3),
  inference(resolution,[],[f1544,f253])).
fof(f7014,plain,(
  ( ! [X10,X8,X11,X9] : (sP4(X8,X9) | sP4(sK8,X8) | sP3(X10,X9,X11) | p(s(bool,X9)) | ~sP2(X10,X11,X9) | ~sP0(X11,X9,X10)) ) | $spl3),
  inference(resolution,[],[f5499,f229])).
fof(f229,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X1)) | ~sP2(X0,X2,X1) | ~sP1(X1,X2,X0) | ~sP0(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f153])).
fof(f6251,plain,(
  ( ! [X14,X12,X15] : (sP3(X12,X14,X15) | ~sP0(X15,X14,X12) | ~sP1(X14,X15,X12) | p(s(bool,X12))) ) | $spl160),
  inference(cnf_transformation,[],[f6251_D])).
fof(f6251_D,plain,(
  ( ! [X14,X12,X15] : (sP3(X12,X14,X15) | ~sP0(X15,X14,X12) | ~sP1(X14,X15,X12) | p(s(bool,X12))) ) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f6944,plain,(
  ( ! [X14,X12,X15,X13] : (sP4(sK8,X12) | sP3(X13,X14,X15) | p(s(bool,X13)) | ~sP1(X14,X15,X13) | ~sP0(X15,X14,X13)) ) | $spl3),
  inference(subsumption_resolution,[],[f6930,f2546])).
fof(f6930,plain,(
  ( ! [X14,X12,X15,X13] : (sP4(X12,X13) | sP4(sK8,X12) | sP3(X13,X14,X15) | p(s(bool,X13)) | ~sP1(X14,X15,X13) | ~sP0(X15,X14,X13)) ) | $spl3),
  inference(resolution,[],[f5496,f230])).
fof(f230,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X0)) | ~sP2(X0,X2,X1) | ~sP1(X1,X2,X0) | ~sP0(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f153])).
fof(f6830,plain,(
  ( ! [X0] : (sP4('const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0) | sP4(sK8,X0)) ) | $spl168),
  inference(cnf_transformation,[],[f6830_D])).
fof(f6830_D,plain,(
  ( ! [X0] : (sP4('const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0) | sP4(sK8,X0)) ) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f6814,plain,(
  ( ! [X0] : (sP4('const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0) | sP4(sK8,X0) | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6))))) ) | ($spl3 | $spl4)),
  inference(superposition,[],[f5431,f357])).
fof(f5431,plain,(
  ( ! [X222] : (sP4('const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)),X222) | sP4(sK8,X222)) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f2547,f369])).
fof(f6801,plain,(
  ( ! [X10,X8,X11,X9] : (sP4(sK8,X9) | sP3(X10,X8,X11) | p(s(bool,X8)) | ~sP2(X10,X11,X8) | ~sP0(X11,X8,X10)) ) | $spl3),
  inference(subsumption_resolution,[],[f6787,f2545])).
fof(f6787,plain,(
  ( ! [X10,X8,X11,X9] : (sP4(X8,X9) | sP4(sK8,X9) | sP3(X10,X8,X11) | p(s(bool,X8)) | ~sP2(X10,X11,X8) | ~sP0(X11,X8,X10)) ) | $spl3),
  inference(resolution,[],[f5373,f229])).
fof(f6717,plain,(
  ( ! [X14,X12,X15,X13] : (sP4(sK8,X13) | sP3(X12,X14,X15) | p(s(bool,X12)) | ~sP1(X14,X15,X12) | ~sP0(X15,X14,X12)) ) | $spl3),
  inference(subsumption_resolution,[],[f6703,f2545])).
fof(f6703,plain,(
  ( ! [X14,X12,X15,X13] : (sP4(X12,X13) | sP4(sK8,X13) | sP3(X12,X14,X15) | p(s(bool,X12)) | ~sP1(X14,X15,X12) | ~sP0(X15,X14,X12)) ) | $spl3),
  inference(resolution,[],[f5370,f230])).
fof(f6603,plain,(
  ( ! [X0] : (sP4(X0,'const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0'))) | sP4(X0,sK8)) ) | $spl166),
  inference(cnf_transformation,[],[f6603_D])).
fof(f6603_D,plain,(
  ( ! [X0] : (sP4(X0,'const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0'))) | sP4(X0,sK8)) ) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f6587,plain,(
  ( ! [X0] : (sP4(X0,'const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0'))) | sP4(X0,sK8) | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6))))) ) | ($spl3 | $spl4)),
  inference(superposition,[],[f4403,f357])).
fof(f4403,plain,(
  ( ! [X220] : (sP4(X220,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7))) | sP4(X220,sK8)) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f2431,f369])).
fof(f6574,plain,(
  ( ! [X10,X8,X11,X9] : (sP4(X8,sK8) | sP3(X10,X9,X11) | p(s(bool,X9)) | ~sP2(X10,X11,X9) | ~sP0(X11,X9,X10)) ) | $spl3),
  inference(subsumption_resolution,[],[f6558,f2429])).
fof(f2429,plain,(
  ( ! [X76,X77] : (~sP4(X76,X77) | p(s(bool,X77)) | sP4(X76,sK8)) ) | $spl3),
  inference(resolution,[],[f1543,f253])).
fof(f6558,plain,(
  ( ! [X10,X8,X11,X9] : (sP4(X8,X9) | sP4(X8,sK8) | sP3(X10,X9,X11) | p(s(bool,X9)) | ~sP2(X10,X11,X9) | ~sP0(X11,X9,X10)) ) | $spl3),
  inference(resolution,[],[f4346,f229])).
fof(f6486,plain,(
  ( ! [X14,X12,X15,X13] : (sP4(X12,sK8) | sP3(X13,X14,X15) | p(s(bool,X13)) | ~sP1(X14,X15,X13) | ~sP0(X15,X14,X13)) ) | $spl3),
  inference(subsumption_resolution,[],[f6470,f2429])).
fof(f6470,plain,(
  ( ! [X14,X12,X15,X13] : (sP4(X12,X13) | sP4(X12,sK8) | sP3(X13,X14,X15) | p(s(bool,X13)) | ~sP1(X14,X15,X13) | ~sP0(X15,X14,X13)) ) | $spl3),
  inference(resolution,[],[f4343,f230])).
fof(f6369,plain,(
  ( ! [X0] : (sP4('const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0) | sP4(X0,sK8)) ) | $spl164),
  inference(cnf_transformation,[],[f6369_D])).
fof(f6369_D,plain,(
  ( ! [X0] : (sP4('const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0) | sP4(X0,sK8)) ) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f6353,plain,(
  ( ! [X0] : (sP4('const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0) | sP4(X0,sK8) | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6))))) ) | ($spl3 | $spl4)),
  inference(superposition,[],[f3543,f357])).
fof(f3543,plain,(
  ( ! [X218] : (sP4('const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)),X218) | sP4(X218,sK8)) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f2430,f369])).
fof(f6339,plain,(
  ( ! [X10,X8,X11,X9] : (sP4(X9,sK8) | sP3(X10,X8,X11) | p(s(bool,X8)) | ~sP2(X10,X11,X8) | ~sP0(X11,X8,X10)) ) | $spl3),
  inference(subsumption_resolution,[],[f6323,f2428])).
fof(f6323,plain,(
  ( ! [X10,X8,X11,X9] : (sP4(X8,X9) | sP4(X9,sK8) | sP3(X10,X8,X11) | p(s(bool,X8)) | ~sP2(X10,X11,X8) | ~sP0(X11,X8,X10)) ) | $spl3),
  inference(resolution,[],[f3487,f229])).
fof(f6250,plain,(
  ( ! [X14,X12,X15,X13] : (sP4(X13,sK8) | sP3(X12,X14,X15) | p(s(bool,X12)) | ~sP1(X14,X15,X12) | ~sP0(X15,X14,X12)) ) | $spl3),
  inference(subsumption_resolution,[],[f6234,f2428])).
fof(f6234,plain,(
  ( ! [X14,X12,X15,X13] : (sP4(X12,X13) | sP4(X13,sK8) | sP3(X12,X14,X15) | p(s(bool,X12)) | ~sP1(X14,X15,X12) | ~sP0(X15,X14,X12)) ) | $spl3),
  inference(resolution,[],[f3484,f230])).
fof(f6175,plain,(
  $false | ($spl3 | $spl98)),
  inference(subsumption_resolution,[],[f6174,f332])).
fof(f6174,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl98)),
  inference(forward_demodulation,[],[f6173,f1561])).
fof(f6173,plain,(
  p(s(bool,sK8)) | ($spl3 | $spl98)),
  inference(subsumption_resolution,[],[f6149,f1541])).
fof(f6149,plain,(
  p(s(bool,sK8)) | sP4(sK8,sK8) | ($spl3 | $spl98)),
  inference(resolution,[],[f3617,f2546])).
fof(f6172,plain,(
  $false | ($spl3 | $spl98)),
  inference(subsumption_resolution,[],[f6171,f332])).
fof(f6171,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl98)),
  inference(forward_demodulation,[],[f6170,f1561])).
fof(f6170,plain,(
  p(s(bool,sK8)) | ($spl3 | $spl98)),
  inference(subsumption_resolution,[],[f6148,f1541])).
fof(f6148,plain,(
  p(s(bool,sK8)) | sP4(sK8,sK8) | ($spl3 | $spl98)),
  inference(resolution,[],[f3617,f2545])).
fof(f6169,plain,(
  $false | ($spl3 | $spl98)),
  inference(subsumption_resolution,[],[f6168,f332])).
fof(f6168,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl98)),
  inference(forward_demodulation,[],[f6167,f1561])).
fof(f6167,plain,(
  p(s(bool,sK8)) | ($spl3 | $spl98)),
  inference(subsumption_resolution,[],[f6147,f1541])).
fof(f6147,plain,(
  p(s(bool,sK8)) | sP4(sK8,sK8) | ($spl3 | $spl98)),
  inference(resolution,[],[f3617,f2429])).
fof(f6166,plain,(
  $false | ($spl3 | $spl98)),
  inference(subsumption_resolution,[],[f6165,f332])).
fof(f6165,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl98)),
  inference(forward_demodulation,[],[f6164,f1561])).
fof(f6164,plain,(
  p(s(bool,sK8)) | ($spl3 | $spl98)),
  inference(subsumption_resolution,[],[f6146,f1541])).
fof(f6146,plain,(
  p(s(bool,sK8)) | sP4(sK8,sK8) | ($spl3 | $spl98)),
  inference(resolution,[],[f3617,f2428])).
fof(f6163,plain,(
  $false | ($spl3 | $spl98)),
  inference(subsumption_resolution,[],[f6162,f332])).
fof(f6162,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl98)),
  inference(forward_demodulation,[],[f6150,f1561])).
fof(f6150,plain,(
  p(s(bool,sK8)) | $spl98),
  inference(duplicate_literal_removal,[],[f6145])).
fof(f6145,plain,(
  p(s(bool,sK8)) | p(s(bool,sK8)) | $spl98),
  inference(resolution,[],[f3617,f262])).
fof(f262,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | ~sP4(X1,X0)) )),
  inference(cnf_transformation,[],[f168])).
fof(f6159,plain,(
  $false | ($spl3 | $spl98)),
  inference(subsumption_resolution,[],[f6158,f332])).
fof(f6158,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl98)),
  inference(forward_demodulation,[],[f6153,f1561])).
fof(f6153,plain,(
  p(s(bool,sK8)) | $spl98),
  inference(duplicate_literal_removal,[],[f6142])).
fof(f6142,plain,(
  p(s(bool,sK8)) | p(s(bool,sK8)) | $spl98),
  inference(resolution,[],[f3617,f253])).
fof(f6157,plain,(
  $false | ($spl3 | $spl98)),
  inference(subsumption_resolution,[],[f6156,f332])).
fof(f6156,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl98)),
  inference(forward_demodulation,[],[f6141,f1561])).
fof(f6141,plain,(
  p(s(bool,sK8)) | ($spl3 | $spl98)),
  inference(resolution,[],[f3617,f1541])).
fof(f6155,plain,(
  $false | ($spl3 | $spl98)),
  inference(subsumption_resolution,[],[f6154,f332])).
fof(f6154,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl98)),
  inference(forward_demodulation,[],[f6140,f1561])).
fof(f6140,plain,(
  p(s(bool,sK8)) | ($spl3 | $spl98)),
  inference(resolution,[],[f3617,f1542])).
fof(f1542,plain,(
  ( ! [X53] : (~sP4(sK8,X53) | p(s(bool,X53))) ) | $spl3),
  inference(resolution,[],[f367,f253])).
fof(f3617,plain,(
  sP4(sK8,sK8) | $spl98),
  inference(cnf_transformation,[],[f3617_D])).
fof(f3617_D,plain,(
  sP4(sK8,sK8) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f6119,plain,(
  ( ! [X30,X28,X31,X29] : (sP4(X28,X29) | sP2(X30,X28,X31) | p(s(bool,X29)) | sP4(sK8,sK8)) ) | $spl3),
  inference(resolution,[],[f3483,f2545])).
fof(f6117,plain,(
  ( ! [X23,X21,X22,X20] : (sP4(X20,X21) | sP2(X22,X20,X23) | p(s(bool,X21)) | sP4(sK8,sK8)) ) | $spl3),
  inference(resolution,[],[f3483,f2428])).
fof(f3466,plain,(
  ( ! [X41] : (s(bool,'T_0') != s(bool,X41)) ) | $spl86),
  inference(cnf_transformation,[],[f3466_D])).
fof(f3466_D,plain,(
  ( ! [X41] : (s(bool,'T_0') != s(bool,X41)) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f5888,plain,(
  ( ! [X97,X98,X96] : (s(bool,'T_0') != s(bool,X98) | p(s(bool,X96)) | ~p(s(bool,X97)) | s(bool,X96) != s(bool,X97)) ) | $spl4),
  inference(forward_demodulation,[],[f5887,f1649])).
fof(f5887,plain,(
  ( ! [X97,X98,X96] : (p(s(bool,X96)) | ~p(s(bool,X97)) | s(bool,X96) != s(bool,X97) | s(bool,X98) != s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) | $spl4),
  inference(subsumption_resolution,[],[f5683,f360])).
fof(f5683,plain,(
  ( ! [X97,X98,X96] : (~p(s(bool,'T_0')) | p(s(bool,X96)) | ~p(s(bool,X97)) | s(bool,X96) != s(bool,X97) | s(bool,X98) != s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) | $spl4),
  inference(superposition,[],[f223,f1649])).
fof(f5884,plain,(
  ( ! [X85,X86,X84] : (s(bool,'T_0') != s(bool,X84) | p(s(bool,X85)) | ~p(s(bool,X86)) | s(bool,X85) != s(bool,X86)) ) | $spl4),
  inference(subsumption_resolution,[],[f5679,f322])).
fof(f5679,plain,(
  ( ! [X85,X86,X84] : (s(bool,'T_0') != s(bool,X84) | p(s(bool,X85)) | ~p(s(bool,X84)) | ~p(s(bool,X86)) | s(bool,X85) != s(bool,X86)) ) | $spl4),
  inference(superposition,[],[f223,f1649])).
fof(f5883,plain,(
  ( ! [X83,X81,X82] : (s(bool,'T_0') != s(bool,X81) | p(s(bool,X82)) | ~p(s(bool,X83)) | s(bool,X82) != s(bool,X83)) ) | $spl4),
  inference(subsumption_resolution,[],[f5678,f360])).
fof(f5678,plain,(
  ( ! [X83,X81,X82] : (s(bool,'T_0') != s(bool,X81) | p(s(bool,X82)) | ~p(s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) | ~p(s(bool,X83)) | s(bool,X82) != s(bool,X83)) ) | $spl4),
  inference(superposition,[],[f223,f1649])).
fof(f5872,plain,(
  ( ! [X37,X38,X36] : (s(bool,'T_0') != s(bool,X38) | p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,X37)) ) | $spl4),
  inference(forward_demodulation,[],[f5871,f1649])).
fof(f5871,plain,(
  ( ! [X37,X38,X36] : (p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,X37) | s(bool,X38) != s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) | $spl4),
  inference(subsumption_resolution,[],[f5663,f360])).
fof(f5663,plain,(
  ( ! [X37,X38,X36] : (~p(s(bool,'T_0')) | p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,X37) | s(bool,X38) != s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) | $spl4),
  inference(superposition,[],[f219,f1649])).
fof(f5868,plain,(
  ( ! [X26,X24,X25] : (s(bool,'T_0') != s(bool,X24) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X26)) ) | $spl4),
  inference(subsumption_resolution,[],[f5659,f360])).
fof(f5659,plain,(
  ( ! [X26,X24,X25] : (s(bool,'T_0') != s(bool,X24) | p(s(bool,X25)) | ~p(s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X26)) ) | $spl4),
  inference(superposition,[],[f219,f1649])).
fof(f5867,plain,(
  ( ! [X23,X21,X22] : (s(bool,'T_0') != s(bool,X21) | p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X23)) ) | $spl4),
  inference(subsumption_resolution,[],[f5658,f322])).
fof(f5658,plain,(
  ( ! [X23,X21,X22] : (s(bool,'T_0') != s(bool,X21) | p(s(bool,X22)) | ~p(s(bool,X21)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X23)) ) | $spl4),
  inference(superposition,[],[f219,f1649])).
fof(f5601,plain,(
  ( ! [X213] : (sP4(X213,'F_0') | sP4(sK8,X213)) ) | $spl156),
  inference(cnf_transformation,[],[f5601_D])).
fof(f5601_D,plain,(
  ( ! [X213] : (sP4(X213,'F_0') | sP4(sK8,X213)) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f5556,plain,(
  ( ! [X220,X219,X221] : (sP4(X219,'F_0') | sP4(sK8,X219) | p(s(bool,X220)) | p(s(bool,X221)) | ~sP5(X220,X221)) ) | $spl3),
  inference(resolution,[],[f2548,f269])).
fof(f5554,plain,(
  ( ! [X215,X213,X214] : (sP4(X213,'F_0') | sP4(sK8,X213) | p(s(bool,X214)) | p(s(bool,X215)) | ~sP5(X214,X215)) ) | $spl3),
  inference(resolution,[],[f2548,f273])).
fof(f5593,plain,(
  ( ! [X0,X3,X1] : (sP4(X0,X1) | s(bool,X1) != s(bool,X3) | ~p(s(bool,X3)) | sP4(sK8,X0)) ) | $spl150),
  inference(cnf_transformation,[],[f5593_D])).
fof(f5593_D,plain,(
  ( ! [X0,X3,X1] : (sP4(X0,X1) | s(bool,X1) != s(bool,X3) | ~p(s(bool,X3)) | sP4(sK8,X0)) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f5535,plain,(
  ( ! [X185,X182,X184,X183] : (sP4(X182,X183) | sP4(sK8,X182) | p(s(bool,X184)) | ~p(s(bool,X185)) | s(bool,X183) != s(bool,X185)) ) | $spl3),
  inference(resolution,[],[f2548,f360])).
fof(f5596,plain,(
  ( ! [X26,X25] : (sP4(X25,X26) | sP4(sK8,X25)) ) | $spl154),
  inference(cnf_transformation,[],[f5596_D])).
fof(f5596_D,plain,(
  ( ! [X26,X25] : (sP4(X25,X26) | sP4(sK8,X25)) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f5534,plain,(
  ( ! [X180,X178,X181,X179] : (sP4(X178,X179) | sP4(sK8,X178) | ~p(s(bool,X180)) | p(s(bool,X181)) | s(bool,X180) != s(bool,X181)) ) | $spl3),
  inference(resolution,[],[f2548,f360])).
fof(f5490,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP4(X35,X36) | sP4(sK8,X35) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X38) | s(bool,X37) != s(bool,X39)) ) | $spl3),
  inference(resolution,[],[f2548,f223])).
fof(f5489,plain,(
  ( ! [X30,X33,X31,X34,X32] : (sP4(X30,X31) | sP4(sK8,X30) | p(s(bool,X32)) | ~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X31) != s(bool,X33)) ) | $spl3),
  inference(resolution,[],[f2548,f222])).
fof(f5595,plain,(
  ( ! [X26,X29,X27,X25] : (sP4(X25,X26) | sP4(sK8,X25) | ~p(s(bool,X27)) | p(s(bool,X29)) | s(bool,X27) != s(bool,X29)) ) | $spl3),
  inference(subsumption_resolution,[],[f5488,f360])).
fof(f5488,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP4(X25,X26) | sP4(sK8,X25) | ~p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X27) != s(bool,X29)) ) | $spl3),
  inference(resolution,[],[f2548,f222])).
fof(f5487,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP4(X20,X21) | sP4(sK8,X20) | p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X24) | s(bool,X21) != s(bool,X23)) ) | $spl3),
  inference(resolution,[],[f2548,f221])).
fof(f5594,plain,(
  ( ! [X6,X5,X9] : (sP4(X5,X6) | s(bool,X6) != s(bool,X9) | sP4(sK8,X5)) ) | $spl152),
  inference(cnf_transformation,[],[f5594_D])).
fof(f5594_D,plain,(
  ( ! [X6,X5,X9] : (sP4(X5,X6) | s(bool,X6) != s(bool,X9) | sP4(sK8,X5)) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f5486,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP4(X15,X16) | sP4(sK8,X15) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X16) != s(bool,X19) | s(bool,X17) != s(bool,X18)) ) | $spl3),
  inference(resolution,[],[f2548,f221])).
fof(f5485,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP4(X10,X11) | sP4(sK8,X10) | ~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,X14)) ) | $spl3),
  inference(resolution,[],[f2548,f219])).
fof(f5484,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP4(X5,X6) | sP4(sK8,X5) | ~p(s(bool,X7)) | p(s(bool,X8)) | s(bool,X6) != s(bool,X9) | s(bool,X7) != s(bool,X8)) ) | $spl3),
  inference(resolution,[],[f2548,f217])).
fof(f5483,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP4(X0,X1) | sP4(sK8,X0) | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) ) | $spl3),
  inference(resolution,[],[f2548,f217])).
fof(f5476,plain,(
  ( ! [X213] : (sP4('F_0',X213) | sP4(sK8,X213)) ) | $spl148),
  inference(cnf_transformation,[],[f5476_D])).
fof(f5476_D,plain,(
  ( ! [X213] : (sP4('F_0',X213) | sP4(sK8,X213)) ) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f5430,plain,(
  ( ! [X220,X219,X221] : (sP4('F_0',X219) | sP4(sK8,X219) | p(s(bool,X220)) | p(s(bool,X221)) | ~sP5(X220,X221)) ) | $spl3),
  inference(resolution,[],[f2547,f269])).
fof(f5428,plain,(
  ( ! [X215,X213,X214] : (sP4('F_0',X213) | sP4(sK8,X213) | p(s(bool,X214)) | p(s(bool,X215)) | ~sP5(X214,X215)) ) | $spl3),
  inference(resolution,[],[f2547,f273])).
fof(f5468,plain,(
  ( ! [X0,X3,X1] : (sP4(X0,X1) | s(bool,X0) != s(bool,X3) | ~p(s(bool,X3)) | sP4(sK8,X1)) ) | $spl142),
  inference(cnf_transformation,[],[f5468_D])).
fof(f5468_D,plain,(
  ( ! [X0,X3,X1] : (sP4(X0,X1) | s(bool,X0) != s(bool,X3) | ~p(s(bool,X3)) | sP4(sK8,X1)) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f5409,plain,(
  ( ! [X185,X182,X184,X183] : (sP4(X182,X183) | sP4(sK8,X183) | p(s(bool,X184)) | ~p(s(bool,X185)) | s(bool,X182) != s(bool,X185)) ) | $spl3),
  inference(resolution,[],[f2547,f360])).
fof(f5471,plain,(
  ( ! [X26,X25] : (sP4(X25,X26) | sP4(sK8,X26)) ) | $spl146),
  inference(cnf_transformation,[],[f5471_D])).
fof(f5471_D,plain,(
  ( ! [X26,X25] : (sP4(X25,X26) | sP4(sK8,X26)) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f5408,plain,(
  ( ! [X180,X178,X181,X179] : (sP4(X178,X179) | sP4(sK8,X179) | ~p(s(bool,X180)) | p(s(bool,X181)) | s(bool,X180) != s(bool,X181)) ) | $spl3),
  inference(resolution,[],[f2547,f360])).
fof(f5364,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP4(X35,X36) | sP4(sK8,X36) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X35) != s(bool,X38) | s(bool,X37) != s(bool,X39)) ) | $spl3),
  inference(resolution,[],[f2547,f223])).
fof(f5363,plain,(
  ( ! [X30,X33,X31,X34,X32] : (sP4(X30,X31) | sP4(sK8,X31) | p(s(bool,X32)) | ~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X30) != s(bool,X33)) ) | $spl3),
  inference(resolution,[],[f2547,f222])).
fof(f5470,plain,(
  ( ! [X26,X29,X27,X25] : (sP4(X25,X26) | sP4(sK8,X26) | ~p(s(bool,X27)) | p(s(bool,X29)) | s(bool,X27) != s(bool,X29)) ) | $spl3),
  inference(subsumption_resolution,[],[f5362,f360])).
fof(f5362,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP4(X25,X26) | sP4(sK8,X26) | ~p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X27) != s(bool,X29)) ) | $spl3),
  inference(resolution,[],[f2547,f222])).
fof(f5361,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP4(X20,X21) | sP4(sK8,X21) | p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X24) | s(bool,X20) != s(bool,X23)) ) | $spl3),
  inference(resolution,[],[f2547,f221])).
fof(f5469,plain,(
  ( ! [X6,X5,X9] : (sP4(X5,X6) | s(bool,X5) != s(bool,X9) | sP4(sK8,X6)) ) | $spl144),
  inference(cnf_transformation,[],[f5469_D])).
fof(f5469_D,plain,(
  ( ! [X6,X5,X9] : (sP4(X5,X6) | s(bool,X5) != s(bool,X9) | sP4(sK8,X6)) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f5360,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP4(X15,X16) | sP4(sK8,X16) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X15) != s(bool,X19) | s(bool,X17) != s(bool,X18)) ) | $spl3),
  inference(resolution,[],[f2547,f221])).
fof(f5359,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP4(X10,X11) | sP4(sK8,X11) | ~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X10) != s(bool,X13) | s(bool,X12) != s(bool,X14)) ) | $spl3),
  inference(resolution,[],[f2547,f219])).
fof(f5358,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP4(X5,X6) | sP4(sK8,X6) | ~p(s(bool,X7)) | p(s(bool,X8)) | s(bool,X5) != s(bool,X9) | s(bool,X7) != s(bool,X8)) ) | $spl3),
  inference(resolution,[],[f2547,f217])).
fof(f5357,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP4(X0,X1) | sP4(sK8,X1) | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X0) != s(bool,X3)) ) | $spl3),
  inference(resolution,[],[f2547,f217])).
fof(f5315,plain,(
  ( ! [X17,X16] : (sP4(X16,'T_0') | sP4(X17,X16) | p(s(bool,X17)) | sP4(sK8,sK8)) ) | $spl3),
  inference(resolution,[],[f5121,f2546])).
fof(f5121,plain,(
  ( ! [X33,X32] : (sP4(sK8,X32) | sP4(X33,'T_0') | sP4(X32,X33)) ) | $spl3),
  inference(subsumption_resolution,[],[f5120,f332])).
fof(f5120,plain,(
  ( ! [X33,X32] : (p(s(bool,'F_0')) | sP4(sK8,X32) | sP4(X33,'T_0') | sP4(X32,X33)) ) | $spl3),
  inference(forward_demodulation,[],[f5085,f1561])).
fof(f5085,plain,(
  ( ! [X33,X32] : (p(s(bool,sK8)) | sP4(sK8,X32) | sP4(X33,'T_0') | sP4(X32,X33)) ) | $spl3),
  inference(resolution,[],[f2546,f4551])).
fof(f4551,plain,(
  ( ! [X132,X133] : (sP4(X132,'T_0') | sP4(X133,X132) | sP4(X133,sK8)) ) | $spl3),
  inference(resolution,[],[f4459,f2431])).
fof(f5313,plain,(
  ( ! [X12,X13] : (sP4(X12,'T_0') | sP4(X13,X12) | p(s(bool,X13)) | sP4(sK8,sK8)) ) | $spl3),
  inference(resolution,[],[f5121,f2429])).
fof(f5263,plain,(
  ( ! [X17,X16] : (sP4(X16,'T_0') | sP4(X16,X17) | p(s(bool,X17)) | sP4(sK8,sK8)) ) | $spl3),
  inference(resolution,[],[f5119,f2546])).
fof(f5119,plain,(
  ( ! [X30,X31] : (sP4(sK8,X30) | sP4(X31,'T_0') | sP4(X31,X30)) ) | $spl3),
  inference(subsumption_resolution,[],[f5118,f332])).
fof(f5118,plain,(
  ( ! [X30,X31] : (p(s(bool,'F_0')) | sP4(sK8,X30) | sP4(X31,'T_0') | sP4(X31,X30)) ) | $spl3),
  inference(forward_demodulation,[],[f5084,f1561])).
fof(f5084,plain,(
  ( ! [X30,X31] : (p(s(bool,sK8)) | sP4(sK8,X30) | sP4(X31,'T_0') | sP4(X31,X30)) ) | $spl3),
  inference(resolution,[],[f2546,f4550])).
fof(f4550,plain,(
  ( ! [X130,X131] : (sP4(X130,'T_0') | sP4(X130,X131) | sP4(X131,sK8)) ) | $spl3),
  inference(resolution,[],[f4459,f2430])).
fof(f5261,plain,(
  ( ! [X12,X13] : (sP4(X12,'T_0') | sP4(X12,X13) | p(s(bool,X13)) | sP4(sK8,sK8)) ) | $spl3),
  inference(resolution,[],[f5119,f2429])).
fof(f4798,plain,(
  ( ! [X4,X5] : (sP4(X4,X5) | p(s(bool,X4)) | sP4('T_0',X5)) ) | $spl136),
  inference(cnf_transformation,[],[f4798_D])).
fof(f4798_D,plain,(
  ( ! [X4,X5] : (sP4(X4,X5) | p(s(bool,X4)) | sP4('T_0',X5)) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f5182,plain,(
  ( ! [X17,X16] : (sP4(X16,X17) | sP4('T_0',X17) | p(s(bool,X16)) | sP4(sK8,sK8)) ) | $spl3),
  inference(resolution,[],[f5117,f2546])).
fof(f5117,plain,(
  ( ! [X28,X27] : (sP4(sK8,X27) | sP4(X27,X28) | sP4('T_0',X28)) ) | $spl3),
  inference(subsumption_resolution,[],[f5116,f332])).
fof(f5116,plain,(
  ( ! [X28,X27] : (p(s(bool,'F_0')) | sP4(sK8,X27) | sP4(X27,X28) | sP4('T_0',X28)) ) | $spl3),
  inference(forward_demodulation,[],[f5082,f1561])).
fof(f5082,plain,(
  ( ! [X28,X27] : (p(s(bool,sK8)) | sP4(sK8,X27) | sP4(X27,X28) | sP4('T_0',X28)) ) | $spl3),
  inference(resolution,[],[f2546,f4383])).
fof(f4383,plain,(
  ( ! [X187,X186] : (sP4(X186,X187) | sP4(X186,sK8) | sP4('T_0',X187)) ) | $spl3),
  inference(resolution,[],[f2431,f3605])).
fof(f5180,plain,(
  ( ! [X12,X13] : (sP4(X12,X13) | sP4('T_0',X13) | p(s(bool,X12)) | sP4(sK8,sK8)) ) | $spl3),
  inference(resolution,[],[f5117,f2429])).
fof(f5133,plain,(
  ( ! [X17,X16] : (sP4('T_0',X16) | sP4(X16,X17) | p(s(bool,X17)) | sP4(sK8,sK8)) ) | $spl3),
  inference(resolution,[],[f5115,f2546])).
fof(f5115,plain,(
  ( ! [X26,X25] : (sP4(sK8,X25) | sP4('T_0',X26) | sP4(X26,X25)) ) | $spl3),
  inference(subsumption_resolution,[],[f5114,f332])).
fof(f5114,plain,(
  ( ! [X26,X25] : (p(s(bool,'F_0')) | sP4(sK8,X25) | sP4('T_0',X26) | sP4(X26,X25)) ) | $spl3),
  inference(forward_demodulation,[],[f5081,f1561])).
fof(f5081,plain,(
  ( ! [X26,X25] : (p(s(bool,sK8)) | sP4(sK8,X25) | sP4('T_0',X26) | sP4(X26,X25)) ) | $spl3),
  inference(resolution,[],[f2546,f4017])).
fof(f4017,plain,(
  ( ! [X130,X131] : (sP4('T_0',X130) | sP4(X130,X131) | sP4(X131,sK8)) ) | $spl3),
  inference(resolution,[],[f3605,f2430])).
fof(f5131,plain,(
  ( ! [X12,X13] : (sP4('T_0',X12) | sP4(X12,X13) | p(s(bool,X13)) | sP4(sK8,sK8)) ) | $spl3),
  inference(resolution,[],[f5115,f2429])).
fof(f5040,plain,(
  ( ! [X33,X32] : (p(s(bool,X32)) | sP4(sK8,sK8) | sP4(X33,'T_0') | sP4(X32,X33)) ) | $spl3),
  inference(resolution,[],[f2545,f4551])).
fof(f5039,plain,(
  ( ! [X30,X31] : (p(s(bool,X30)) | sP4(sK8,sK8) | sP4(X31,'T_0') | sP4(X31,X30)) ) | $spl3),
  inference(resolution,[],[f2545,f4550])).
fof(f5037,plain,(
  ( ! [X28,X27] : (p(s(bool,X27)) | sP4(sK8,sK8) | sP4(X27,X28) | sP4('T_0',X28)) ) | $spl3),
  inference(resolution,[],[f2545,f4383])).
fof(f5036,plain,(
  ( ! [X26,X25] : (p(s(bool,X25)) | sP4(sK8,sK8) | sP4('T_0',X26) | sP4(X26,X25)) ) | $spl3),
  inference(resolution,[],[f2545,f4017])).
fof(f4870,plain,(
  ( ! [X4,X5] : (sP4(X4,'T_0') | sP4(X5,X4) | p(s(bool,X5)) | sP4(sK8,sK8)) ) | $spl3),
  inference(resolution,[],[f4551,f2428])).
fof(f4829,plain,(
  ( ! [X4,X5] : (sP4(X4,'T_0') | sP4(X4,X5) | p(s(bool,X5)) | sP4(sK8,sK8)) ) | $spl3),
  inference(resolution,[],[f4550,f2428])).
fof(f4777,plain,(
  ( ! [X4,X5] : (sP4(X4,X5) | sP4('T_0',X5) | p(s(bool,X4)) | sP4(sK8,sK8)) ) | $spl3),
  inference(resolution,[],[f4383,f2428])).
fof(f4727,plain,(
  ( ! [X0] : (sP4(X0,'const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0'))) | p(s(bool,X0))) ) | $spl134),
  inference(cnf_transformation,[],[f4727_D])).
fof(f4727_D,plain,(
  ( ! [X0] : (sP4(X0,'const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0'))) | p(s(bool,X0))) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f4716,plain,(
  ( ! [X0] : (sP4(X0,'const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0'))) | p(s(bool,X0)) | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6))))) ) | $spl4),
  inference(superposition,[],[f1631,f357])).
fof(f1631,plain,(
  ( ! [X55] : (sP4(X55,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7))) | p(s(bool,X55))) ) | $spl4),
  inference(resolution,[],[f369,f258])).
fof(f4696,plain,(
  ( ! [X0] : (sP4('const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0) | p(s(bool,X0))) ) | $spl132),
  inference(cnf_transformation,[],[f4696_D])).
fof(f4696_D,plain,(
  ( ! [X0] : (sP4('const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0) | p(s(bool,X0))) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f4685,plain,(
  ( ! [X0] : (sP4('const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0) | p(s(bool,X0)) | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6))))) ) | $spl4),
  inference(superposition,[],[f1630,f357])).
fof(f1630,plain,(
  ( ! [X54] : (sP4('const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)),X54) | p(s(bool,X54))) ) | $spl4),
  inference(resolution,[],[f369,f255])).
fof(f4643,plain,(
  ( ! [X0] : (~sP4(X0,'const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0'))) | ~p(s(bool,X0))) ) | $spl130),
  inference(cnf_transformation,[],[f4643_D])).
fof(f4643_D,plain,(
  ( ! [X0] : (~sP4(X0,'const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0'))) | ~p(s(bool,X0))) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f4631,plain,(
  ( ! [X0] : (~sP4(X0,'const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0'))) | ~p(s(bool,X0)) | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6))))) ) | $spl4),
  inference(superposition,[],[f1629,f357])).
fof(f4605,plain,(
  sP4('F_0','T_0') | $spl128),
  inference(cnf_transformation,[],[f4605_D])).
fof(f4605_D,plain,(
  sP4('F_0','T_0') <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f4596,plain,(
  ( ! [X31] : (p(s(bool,X31)) | sP4('F_0','T_0')) ) | $spl3),
  inference(duplicate_literal_removal,[],[f4584])).
fof(f4584,plain,(
  ( ! [X31] : (p(s(bool,X31)) | sP4('F_0','T_0') | p(s(bool,X31))) ) | $spl3),
  inference(superposition,[],[f4459,f319])).
fof(f4597,plain,(
  ( ! [X30] : (p(s(bool,X30)) | sP4('F_0','T_0')) ) | $spl3),
  inference(duplicate_literal_removal,[],[f4583])).
fof(f4583,plain,(
  ( ! [X30] : (p(s(bool,X30)) | sP4('F_0','T_0') | p(s(bool,X30))) ) | $spl3),
  inference(superposition,[],[f4459,f321])).
fof(f4612,plain,(
  ( ! [X29] : (p(s(bool,X29)) | sP4('F_0','T_0')) ) | $spl3),
  inference(subsumption_resolution,[],[f4582,f322])).
fof(f4582,plain,(
  ( ! [X29] : (p(s(bool,X29)) | sP4('F_0','T_0') | s(bool,'T_0') = s(bool,X29)) ) | $spl3),
  inference(superposition,[],[f4459,f358])).
fof(f4561,plain,(
  ( ! [X143] : (sP4('F_0','T_0') | p(s(bool,X143))) ) | $spl3),
  inference(resolution,[],[f4459,f331])).
fof(f4560,plain,(
  ( ! [X142] : (sP4('F_0','T_0') | p(s(bool,X142))) ) | $spl3),
  inference(resolution,[],[f4459,f314])).
fof(f4559,plain,(
  ( ! [X141] : (sP4('F_0','T_0') | p(s(bool,X141))) ) | $spl3),
  inference(resolution,[],[f4459,f307])).
fof(f4606,plain,(
  ( ! [X140] : (sP4('F_0','T_0') | p(s(bool,X140))) ) | $spl3),
  inference(subsumption_resolution,[],[f4558,f361])).
fof(f4558,plain,(
  ( ! [X140] : (sP4('F_0','T_0') | p(s(bool,X140)) | ~p(s(bool,'T_0'))) ) | $spl3),
  inference(resolution,[],[f4459,f293])).
fof(f4557,plain,(
  ( ! [X139] : (sP4('F_0','T_0') | ~p(s(bool,X139))) ) | $spl3),
  inference(resolution,[],[f4459,f289])).
fof(f4556,plain,(
  ( ! [X138,X137] : (sP4('F_0','T_0') | sP5(X137,X138)) ) | $spl3),
  inference(resolution,[],[f4459,f277])).
fof(f4555,plain,(
  ( ! [X136,X135] : (sP4('F_0','T_0') | sP5(X135,X136)) ) | $spl3),
  inference(resolution,[],[f4459,f272])).
fof(f4547,plain,(
  ( ! [X127,X125,X126] : (sP4(X125,'T_0') | p(s(bool,X126)) | p(s(bool,X127)) | s(bool,X125) != s(bool,X127)) ) | $spl3),
  inference(resolution,[],[f4459,f360])).
fof(f4505,plain,(
  ( ! [X30,X28,X31,X29] : (sP4(X28,'T_0') | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) ) | $spl3),
  inference(resolution,[],[f4459,f223])).
fof(f4601,plain,(
  ( ! [X4,X7] : (sP4(X4,'T_0') | s(bool,X4) != s(bool,X7)) ) | $spl126),
  inference(cnf_transformation,[],[f4601_D])).
fof(f4601_D,plain,(
  ( ! [X4,X7] : (sP4(X4,'T_0') | s(bool,X4) != s(bool,X7)) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f4504,plain,(
  ( ! [X26,X24,X27,X25] : (sP4(X24,'T_0') | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X26) | s(bool,X24) != s(bool,X27)) ) | $spl3),
  inference(resolution,[],[f4459,f223])).
fof(f4502,plain,(
  ( ! [X19,X17,X18,X16] : (sP4(X16,'T_0') | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) ) | $spl3),
  inference(resolution,[],[f4459,f222])).
fof(f4501,plain,(
  ( ! [X14,X12,X15,X13] : (sP4(X12,'T_0') | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl3),
  inference(resolution,[],[f4459,f221])).
fof(f4500,plain,(
  ( ! [X10,X8,X11,X9] : (sP4(X8,'T_0') | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) ) | $spl3),
  inference(resolution,[],[f4459,f219])).
fof(f4499,plain,(
  ( ! [X6,X4,X7,X5] : (sP4(X4,'T_0') | p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X5) != s(bool,X6) | s(bool,X4) != s(bool,X7)) ) | $spl3),
  inference(resolution,[],[f4459,f219])).
fof(f4498,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(X0,'T_0') | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl3),
  inference(resolution,[],[f4459,f217])).
fof(f4497,plain,(
  ( ! [X0] : (~sP4('const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0) | ~p(s(bool,X0))) ) | $spl122),
  inference(cnf_transformation,[],[f4497_D])).
fof(f4497_D,plain,(
  ( ! [X0] : (~sP4('const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0) | ~p(s(bool,X0))) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f4484,plain,(
  ( ! [X0] : (~sP4('const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0) | ~p(s(bool,X0)) | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6))))) ) | $spl4),
  inference(superposition,[],[f1628,f357])).
fof(f4467,plain,(
  ( ! [X2] : (p(s(bool,X2)) | sP4(X2,'T_0')) ) | $spl120),
  inference(cnf_transformation,[],[f4467_D])).
fof(f4467_D,plain,(
  ( ! [X2] : (p(s(bool,X2)) | sP4(X2,'T_0')) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f4461,plain,(
  ( ! [X2] : (sP4(X2,'T_0') | p(s(bool,X2)) | sP4(sK8,sK8)) ) | $spl3),
  inference(resolution,[],[f4394,f2428])).
fof(f4449,plain,(
  ( ! [X189] : (sP4(X189,sK8) | sP4(X189,'T_0')) ) | $spl116),
  inference(cnf_transformation,[],[f4449_D])).
fof(f4449_D,plain,(
  ( ! [X189] : (sP4(X189,sK8) | sP4(X189,'T_0')) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f4436,plain,(
  ( ! [X45,X44] : (~p(s(bool,X44)) | sP4(X45,'T_0') | sP4(X45,sK8)) ) | $spl3),
  inference(duplicate_literal_removal,[],[f4420])).
fof(f4420,plain,(
  ( ! [X45,X44] : (~p(s(bool,X44)) | sP4(X45,'T_0') | sP4(X45,sK8) | ~p(s(bool,X44))) ) | $spl3),
  inference(superposition,[],[f2431,f323])).
fof(f4437,plain,(
  ( ! [X43,X42] : (~p(s(bool,X42)) | sP4(X43,'T_0') | sP4(X43,sK8)) ) | $spl3),
  inference(duplicate_literal_removal,[],[f4419])).
fof(f4419,plain,(
  ( ! [X43,X42] : (~p(s(bool,X42)) | sP4(X43,'T_0') | sP4(X43,sK8) | ~p(s(bool,X42))) ) | $spl3),
  inference(superposition,[],[f2431,f325])).
fof(f4456,plain,(
  ( ! [X41,X40] : (~p(s(bool,X40)) | sP4(X41,'T_0') | sP4(X41,sK8)) ) | $spl3),
  inference(subsumption_resolution,[],[f4418,f318])).
fof(f4418,plain,(
  ( ! [X41,X40] : (~p(s(bool,X40)) | sP4(X41,'T_0') | sP4(X41,sK8) | s(bool,'F_0') = s(bool,X40)) ) | $spl3),
  inference(superposition,[],[f2431,f358])).
fof(f4452,plain,(
  ( ! [X211] : (sP4(X211,'F_0') | sP4(X211,sK8)) ) | $spl118),
  inference(cnf_transformation,[],[f4452_D])).
fof(f4452_D,plain,(
  ( ! [X211] : (sP4(X211,'F_0') | sP4(X211,sK8)) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f4402,plain,(
  ( ! [X218,X219,X217] : (sP4(X217,'F_0') | sP4(X217,sK8) | p(s(bool,X218)) | p(s(bool,X219)) | ~sP5(X218,X219)) ) | $spl3),
  inference(resolution,[],[f2431,f269])).
fof(f4400,plain,(
  ( ! [X212,X213,X211] : (sP4(X211,'F_0') | sP4(X211,sK8) | p(s(bool,X212)) | p(s(bool,X213)) | ~sP5(X212,X213)) ) | $spl3),
  inference(resolution,[],[f2431,f273])).
fof(f4391,plain,(
  ( ! [X202,X201] : (sP4(X201,'T_0') | sP4(X201,sK8) | ~p(s(bool,X202))) ) | $spl3),
  inference(resolution,[],[f2431,f316])).
fof(f4390,plain,(
  ( ! [X200,X199] : (sP4(X199,'T_0') | sP4(X199,sK8) | ~p(s(bool,X200))) ) | $spl3),
  inference(resolution,[],[f2431,f311])).
fof(f4389,plain,(
  ( ! [X198,X197] : (sP4(X197,'T_0') | sP4(X197,sK8) | ~p(s(bool,X198))) ) | $spl3),
  inference(resolution,[],[f2431,f301])).
fof(f4388,plain,(
  ( ! [X196,X195] : (sP4(X195,'T_0') | sP4(X195,sK8) | p(s(bool,X196))) ) | $spl3),
  inference(resolution,[],[f2431,f300])).
fof(f4387,plain,(
  ( ! [X194,X193] : (sP4(X193,'T_0') | sP4(X193,sK8) | p(s(bool,X194))) ) | $spl3),
  inference(resolution,[],[f2431,f297])).
fof(f4386,plain,(
  ( ! [X191,X192] : (sP4(X191,'T_0') | sP4(X191,sK8) | p(s(bool,X192))) ) | $spl3),
  inference(resolution,[],[f2431,f294])).
fof(f4385,plain,(
  ( ! [X189,X190] : (sP4(X189,'T_0') | sP4(X189,sK8) | ~p(s(bool,X190))) ) | $spl3),
  inference(resolution,[],[f2431,f292])).
fof(f4440,plain,(
  ( ! [X0,X3,X1] : (sP4(X0,X1) | s(bool,X1) != s(bool,X3) | ~p(s(bool,X3)) | sP4(X0,sK8)) ) | $spl110),
  inference(cnf_transformation,[],[f4440_D])).
fof(f4440_D,plain,(
  ( ! [X0,X3,X1] : (sP4(X0,X1) | s(bool,X1) != s(bool,X3) | ~p(s(bool,X3)) | sP4(X0,sK8)) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f4382,plain,(
  ( ! [X185,X182,X184,X183] : (sP4(X182,X183) | sP4(X182,sK8) | p(s(bool,X184)) | ~p(s(bool,X185)) | s(bool,X183) != s(bool,X185)) ) | $spl3),
  inference(resolution,[],[f2431,f360])).
fof(f4443,plain,(
  ( ! [X26,X25] : (sP4(X25,X26) | sP4(X25,sK8)) ) | $spl114),
  inference(cnf_transformation,[],[f4443_D])).
fof(f4443_D,plain,(
  ( ! [X26,X25] : (sP4(X25,X26) | sP4(X25,sK8)) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f4381,plain,(
  ( ! [X180,X178,X181,X179] : (sP4(X178,X179) | sP4(X178,sK8) | ~p(s(bool,X180)) | p(s(bool,X181)) | s(bool,X180) != s(bool,X181)) ) | $spl3),
  inference(resolution,[],[f2431,f360])).
fof(f4337,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP4(X35,X36) | sP4(X35,sK8) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X38) | s(bool,X37) != s(bool,X39)) ) | $spl3),
  inference(resolution,[],[f2431,f223])).
fof(f4336,plain,(
  ( ! [X30,X33,X31,X34,X32] : (sP4(X30,X31) | sP4(X30,sK8) | p(s(bool,X32)) | ~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X31) != s(bool,X33)) ) | $spl3),
  inference(resolution,[],[f2431,f222])).
fof(f4442,plain,(
  ( ! [X26,X29,X27,X25] : (sP4(X25,X26) | sP4(X25,sK8) | ~p(s(bool,X27)) | p(s(bool,X29)) | s(bool,X27) != s(bool,X29)) ) | $spl3),
  inference(subsumption_resolution,[],[f4335,f360])).
fof(f4335,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP4(X25,X26) | sP4(X25,sK8) | ~p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X27) != s(bool,X29)) ) | $spl3),
  inference(resolution,[],[f2431,f222])).
fof(f4334,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP4(X20,X21) | sP4(X20,sK8) | p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X24) | s(bool,X21) != s(bool,X23)) ) | $spl3),
  inference(resolution,[],[f2431,f221])).
fof(f4441,plain,(
  ( ! [X6,X5,X9] : (sP4(X5,X6) | s(bool,X6) != s(bool,X9) | sP4(X5,sK8)) ) | $spl112),
  inference(cnf_transformation,[],[f4441_D])).
fof(f4441_D,plain,(
  ( ! [X6,X5,X9] : (sP4(X5,X6) | s(bool,X6) != s(bool,X9) | sP4(X5,sK8)) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f4333,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP4(X15,X16) | sP4(X15,sK8) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X16) != s(bool,X19) | s(bool,X17) != s(bool,X18)) ) | $spl3),
  inference(resolution,[],[f2431,f221])).
fof(f4332,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP4(X10,X11) | sP4(X10,sK8) | ~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,X14)) ) | $spl3),
  inference(resolution,[],[f2431,f219])).
fof(f4331,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP4(X5,X6) | sP4(X5,sK8) | ~p(s(bool,X7)) | p(s(bool,X8)) | s(bool,X6) != s(bool,X9) | s(bool,X7) != s(bool,X8)) ) | $spl3),
  inference(resolution,[],[f2431,f217])).
fof(f4330,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP4(X0,X1) | sP4(X0,sK8) | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) ) | $spl3),
  inference(resolution,[],[f2431,f217])).
fof(f4311,plain,(
  ( ! [X4,X5] : (sP4('T_0',X4) | sP4(X4,X5) | p(s(bool,X5)) | sP4(sK8,sK8)) ) | $spl3),
  inference(resolution,[],[f4017,f2428])).
fof(f4200,plain,(
  $false | ($spl3 | $spl38)),
  inference(subsumption_resolution,[],[f4155,f4089])).
fof(f4089,plain,(
  ( ! [X1] : (p(s(bool,X1))) ) | ($spl3 | $spl38)),
  inference(subsumption_resolution,[],[f4088,f321])).
fof(f4088,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,X1))) ) | ($spl3 | $spl38)),
  inference(forward_demodulation,[],[f987,f1561])).
fof(f4155,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl38)),
  inference(resolution,[],[f4089,f336])).
fof(f336,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f203])).
fof(f4199,plain,(
  $false | ($spl3 | $spl38)),
  inference(subsumption_resolution,[],[f4154,f4089])).
fof(f4154,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl38)),
  inference(resolution,[],[f4089,f334])).
fof(f334,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f202])).
fof(f4198,plain,(
  $false | ($spl3 | $spl38)),
  inference(subsumption_resolution,[],[f4148,f4089])).
fof(f4148,plain,(
  ( ! [X131] : (~p(s(bool,X131))) ) | ($spl3 | $spl38)),
  inference(resolution,[],[f4089,f289])).
fof(f4153,plain,(
  $false | ($spl3 | $spl38)),
  inference(resolution,[],[f4089,f332])).
fof(f4071,plain,(
  sP4('T_0','F_0') | $spl106),
  inference(cnf_transformation,[],[f4071_D])).
fof(f4071_D,plain,(
  sP4('T_0','F_0') <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f4062,plain,(
  ( ! [X31] : (p(s(bool,X31)) | sP4('T_0','F_0')) ) | $spl3),
  inference(duplicate_literal_removal,[],[f4050])).
fof(f4050,plain,(
  ( ! [X31] : (p(s(bool,X31)) | sP4('T_0','F_0') | p(s(bool,X31))) ) | $spl3),
  inference(superposition,[],[f3605,f319])).
fof(f4063,plain,(
  ( ! [X30] : (p(s(bool,X30)) | sP4('T_0','F_0')) ) | $spl3),
  inference(duplicate_literal_removal,[],[f4049])).
fof(f4049,plain,(
  ( ! [X30] : (p(s(bool,X30)) | sP4('T_0','F_0') | p(s(bool,X30))) ) | $spl3),
  inference(superposition,[],[f3605,f321])).
fof(f4078,plain,(
  ( ! [X29] : (p(s(bool,X29)) | sP4('T_0','F_0')) ) | $spl3),
  inference(subsumption_resolution,[],[f4048,f322])).
fof(f4048,plain,(
  ( ! [X29] : (p(s(bool,X29)) | sP4('T_0','F_0') | s(bool,'T_0') = s(bool,X29)) ) | $spl3),
  inference(superposition,[],[f3605,f358])).
fof(f4027,plain,(
  ( ! [X141] : (sP4('T_0','F_0') | p(s(bool,X141))) ) | $spl3),
  inference(resolution,[],[f3605,f331])).
fof(f4026,plain,(
  ( ! [X140] : (sP4('T_0','F_0') | p(s(bool,X140))) ) | $spl3),
  inference(resolution,[],[f3605,f314])).
fof(f4025,plain,(
  ( ! [X139] : (sP4('T_0','F_0') | p(s(bool,X139))) ) | $spl3),
  inference(resolution,[],[f3605,f307])).
fof(f4072,plain,(
  ( ! [X138] : (sP4('T_0','F_0') | p(s(bool,X138))) ) | $spl3),
  inference(subsumption_resolution,[],[f4024,f361])).
fof(f4024,plain,(
  ( ! [X138] : (sP4('T_0','F_0') | p(s(bool,X138)) | ~p(s(bool,'T_0'))) ) | $spl3),
  inference(resolution,[],[f3605,f293])).
fof(f4023,plain,(
  ( ! [X137] : (sP4('T_0','F_0') | ~p(s(bool,X137))) ) | $spl3),
  inference(resolution,[],[f3605,f289])).
fof(f4022,plain,(
  ( ! [X136,X135] : (sP4('T_0','F_0') | sP5(X135,X136)) ) | $spl3),
  inference(resolution,[],[f3605,f277])).
fof(f4021,plain,(
  ( ! [X134,X133] : (sP4('T_0','F_0') | sP5(X133,X134)) ) | $spl3),
  inference(resolution,[],[f3605,f272])).
fof(f4014,plain,(
  ( ! [X127,X125,X126] : (sP4('T_0',X125) | p(s(bool,X126)) | p(s(bool,X127)) | s(bool,X125) != s(bool,X127)) ) | $spl3),
  inference(resolution,[],[f3605,f360])).
fof(f3972,plain,(
  ( ! [X30,X28,X31,X29] : (sP4('T_0',X28) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) ) | $spl3),
  inference(resolution,[],[f3605,f223])).
fof(f4067,plain,(
  ( ! [X4,X7] : (sP4('T_0',X4) | s(bool,X4) != s(bool,X7)) ) | $spl104),
  inference(cnf_transformation,[],[f4067_D])).
fof(f4067_D,plain,(
  ( ! [X4,X7] : (sP4('T_0',X4) | s(bool,X4) != s(bool,X7)) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f3971,plain,(
  ( ! [X26,X24,X27,X25] : (sP4('T_0',X24) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X26) | s(bool,X24) != s(bool,X27)) ) | $spl3),
  inference(resolution,[],[f3605,f223])).
fof(f3969,plain,(
  ( ! [X19,X17,X18,X16] : (sP4('T_0',X16) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) ) | $spl3),
  inference(resolution,[],[f3605,f222])).
fof(f3968,plain,(
  ( ! [X14,X12,X15,X13] : (sP4('T_0',X12) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl3),
  inference(resolution,[],[f3605,f221])).
fof(f3967,plain,(
  ( ! [X10,X8,X11,X9] : (sP4('T_0',X8) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) ) | $spl3),
  inference(resolution,[],[f3605,f219])).
fof(f3966,plain,(
  ( ! [X6,X4,X7,X5] : (sP4('T_0',X4) | p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X5) != s(bool,X6) | s(bool,X4) != s(bool,X7)) ) | $spl3),
  inference(resolution,[],[f3605,f219])).
fof(f3965,plain,(
  ( ! [X2,X0,X3,X1] : (sP4('T_0',X0) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl3),
  inference(resolution,[],[f3605,f217])).
fof(f3956,plain,(
  $false | $spl86),
  inference(trivial_inequality_removal,[],[f3943])).
fof(f3943,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl86),
  inference(superposition,[],[f3466,f339])).
fof(f3955,plain,(
  $false | $spl86),
  inference(trivial_inequality_removal,[],[f3945])).
fof(f3945,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl86),
  inference(superposition,[],[f3466,f359])).
fof(f3954,plain,(
  $false | $spl86),
  inference(equality_resolution,[],[f3466])).
fof(f3893,plain,(
  $false | $spl68),
  inference(trivial_inequality_removal,[],[f3881])).
fof(f3881,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl68),
  inference(superposition,[],[f2871,f340])).
fof(f3892,plain,(
  $false | ($spl3 | $spl68)),
  inference(trivial_inequality_removal,[],[f3883])).
fof(f3883,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl3 | $spl68)),
  inference(superposition,[],[f2871,f1561])).
fof(f3891,plain,(
  $false | $spl68),
  inference(equality_resolution,[],[f2871])).
fof(f3835,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f3764,f3687])).
fof(f3687,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl14),
  inference(equality_resolution,[],[f456])).
fof(f3764,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(resolution,[],[f3687,f335])).
fof(f3834,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f3763,f3687])).
fof(f3763,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(resolution,[],[f3687,f333])).
fof(f3833,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f3759,f3687])).
fof(f3759,plain,(
  ( ! [X127] : (p(s(bool,X127))) ) | $spl14),
  inference(resolution,[],[f3687,f300])).
fof(f3832,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f3758,f3687])).
fof(f3758,plain,(
  ( ! [X126] : (p(s(bool,X126))) ) | $spl14),
  inference(resolution,[],[f3687,f297])).
fof(f3831,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f3757,f3687])).
fof(f3757,plain,(
  ( ! [X125] : (p(s(bool,X125))) ) | $spl14),
  inference(resolution,[],[f3687,f294])).
fof(f3830,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f3755,f3687])).
fof(f3755,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(resolution,[],[f3687,f291])).
fof(f3765,plain,(
  $false | $spl14),
  inference(resolution,[],[f3687,f337])).
fof(f3766,plain,(
  $false | $spl14),
  inference(resolution,[],[f3687,f338])).
fof(f3767,plain,(
  $false | $spl14),
  inference(resolution,[],[f3687,f361])).
fof(f3774,plain,(
  $false | ($spl4 | $spl14)),
  inference(resolution,[],[f3687,f369])).
fof(f2871,plain,(
  ( ! [X21] : (s(bool,'F_0') != s(bool,X21)) ) | $spl68),
  inference(cnf_transformation,[],[f2871_D])).
fof(f2871_D,plain,(
  ( ! [X21] : (s(bool,'F_0') != s(bool,X21)) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f3693,plain,(
  ( ! [X39,X38] : (s(bool,'F_0') != s(bool,X39) | ~p(s(bool,X38))) ) | $spl14),
  inference(subsumption_resolution,[],[f3638,f456])).
fof(f3638,plain,(
  ( ! [X39,X38] : (s(bool,'F_0') != s(bool,X39) | ~p(s(bool,X38)) | s(bool,'T_0') = s(bool,X38)) ) | $spl14),
  inference(superposition,[],[f456,f358])).
fof(f3692,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') != s(bool,X37) | ~p(s(bool,X36))) ) | $spl14),
  inference(subsumption_resolution,[],[f3637,f456])).
fof(f3637,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') != s(bool,X37) | ~p(s(bool,X36)) | s(bool,'F_0') = s(bool,X36)) ) | $spl14),
  inference(superposition,[],[f456,f358])).
fof(f3688,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') != s(bool,X35) | ~p(s(bool,X34))) ) | $spl14),
  inference(duplicate_literal_removal,[],[f3636])).
fof(f3636,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') != s(bool,X35) | ~p(s(bool,X34)) | ~p(s(bool,X34))) ) | $spl14),
  inference(superposition,[],[f456,f325])).
fof(f3689,plain,(
  ( ! [X33,X32] : (s(bool,'T_0') != s(bool,X33) | ~p(s(bool,X32))) ) | $spl14),
  inference(duplicate_literal_removal,[],[f3635])).
fof(f3635,plain,(
  ( ! [X33,X32] : (s(bool,'T_0') != s(bool,X33) | ~p(s(bool,X32)) | ~p(s(bool,X32))) ) | $spl14),
  inference(superposition,[],[f456,f323])).
fof(f3618,plain,(
  ( ! [X2] : (p(s(bool,X2)) | sP4('T_0',X2)) ) | $spl100),
  inference(cnf_transformation,[],[f3618_D])).
fof(f3618_D,plain,(
  ( ! [X2] : (p(s(bool,X2)) | sP4('T_0',X2)) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f3607,plain,(
  ( ! [X2] : (sP4('T_0',X2) | p(s(bool,X2)) | sP4(sK8,sK8)) ) | $spl3),
  inference(resolution,[],[f3534,f2428])).
fof(f3588,plain,(
  ( ! [X187] : (sP4('T_0',X187) | sP4(X187,sK8)) ) | $spl94),
  inference(cnf_transformation,[],[f3588_D])).
fof(f3588_D,plain,(
  ( ! [X187] : (sP4('T_0',X187) | sP4(X187,sK8)) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f3575,plain,(
  ( ! [X45,X44] : (~p(s(bool,X44)) | sP4('T_0',X45) | sP4(X45,sK8)) ) | $spl3),
  inference(duplicate_literal_removal,[],[f3560])).
fof(f3560,plain,(
  ( ! [X45,X44] : (~p(s(bool,X44)) | sP4('T_0',X45) | sP4(X45,sK8) | ~p(s(bool,X44))) ) | $spl3),
  inference(superposition,[],[f2430,f323])).
fof(f3576,plain,(
  ( ! [X43,X42] : (~p(s(bool,X42)) | sP4('T_0',X43) | sP4(X43,sK8)) ) | $spl3),
  inference(duplicate_literal_removal,[],[f3559])).
fof(f3559,plain,(
  ( ! [X43,X42] : (~p(s(bool,X42)) | sP4('T_0',X43) | sP4(X43,sK8) | ~p(s(bool,X42))) ) | $spl3),
  inference(superposition,[],[f2430,f325])).
fof(f3595,plain,(
  ( ! [X41,X40] : (~p(s(bool,X40)) | sP4('T_0',X41) | sP4(X41,sK8)) ) | $spl3),
  inference(subsumption_resolution,[],[f3558,f318])).
fof(f3558,plain,(
  ( ! [X41,X40] : (~p(s(bool,X40)) | sP4('T_0',X41) | sP4(X41,sK8) | s(bool,'F_0') = s(bool,X40)) ) | $spl3),
  inference(superposition,[],[f2430,f358])).
fof(f3591,plain,(
  ( ! [X209] : (sP4('F_0',X209) | sP4(X209,sK8)) ) | $spl96),
  inference(cnf_transformation,[],[f3591_D])).
fof(f3591_D,plain,(
  ( ! [X209] : (sP4('F_0',X209) | sP4(X209,sK8)) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f3542,plain,(
  ( ! [X216,X215,X217] : (sP4('F_0',X215) | sP4(X215,sK8) | p(s(bool,X216)) | p(s(bool,X217)) | ~sP5(X216,X217)) ) | $spl3),
  inference(resolution,[],[f2430,f269])).
fof(f3540,plain,(
  ( ! [X210,X211,X209] : (sP4('F_0',X209) | sP4(X209,sK8) | p(s(bool,X210)) | p(s(bool,X211)) | ~sP5(X210,X211)) ) | $spl3),
  inference(resolution,[],[f2430,f273])).
fof(f3531,plain,(
  ( ! [X200,X199] : (sP4('T_0',X199) | sP4(X199,sK8) | ~p(s(bool,X200))) ) | $spl3),
  inference(resolution,[],[f2430,f316])).
fof(f3530,plain,(
  ( ! [X198,X197] : (sP4('T_0',X197) | sP4(X197,sK8) | ~p(s(bool,X198))) ) | $spl3),
  inference(resolution,[],[f2430,f311])).
fof(f3529,plain,(
  ( ! [X196,X195] : (sP4('T_0',X195) | sP4(X195,sK8) | ~p(s(bool,X196))) ) | $spl3),
  inference(resolution,[],[f2430,f301])).
fof(f3528,plain,(
  ( ! [X194,X193] : (sP4('T_0',X193) | sP4(X193,sK8) | p(s(bool,X194))) ) | $spl3),
  inference(resolution,[],[f2430,f300])).
fof(f3527,plain,(
  ( ! [X191,X192] : (sP4('T_0',X191) | sP4(X191,sK8) | p(s(bool,X192))) ) | $spl3),
  inference(resolution,[],[f2430,f297])).
fof(f3526,plain,(
  ( ! [X189,X190] : (sP4('T_0',X189) | sP4(X189,sK8) | p(s(bool,X190))) ) | $spl3),
  inference(resolution,[],[f2430,f294])).
fof(f3525,plain,(
  ( ! [X187,X188] : (sP4('T_0',X187) | sP4(X187,sK8) | ~p(s(bool,X188))) ) | $spl3),
  inference(resolution,[],[f2430,f292])).
fof(f3579,plain,(
  ( ! [X0,X3,X1] : (sP4(X0,X1) | s(bool,X0) != s(bool,X3) | ~p(s(bool,X3)) | sP4(X1,sK8)) ) | $spl88),
  inference(cnf_transformation,[],[f3579_D])).
fof(f3579_D,plain,(
  ( ! [X0,X3,X1] : (sP4(X0,X1) | s(bool,X0) != s(bool,X3) | ~p(s(bool,X3)) | sP4(X1,sK8)) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f3523,plain,(
  ( ! [X185,X182,X184,X183] : (sP4(X182,X183) | sP4(X183,sK8) | p(s(bool,X184)) | ~p(s(bool,X185)) | s(bool,X182) != s(bool,X185)) ) | $spl3),
  inference(resolution,[],[f2430,f360])).
fof(f3582,plain,(
  ( ! [X26,X25] : (sP4(X25,X26) | sP4(X26,sK8)) ) | $spl92),
  inference(cnf_transformation,[],[f3582_D])).
fof(f3582_D,plain,(
  ( ! [X26,X25] : (sP4(X25,X26) | sP4(X26,sK8)) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f3522,plain,(
  ( ! [X180,X178,X181,X179] : (sP4(X178,X179) | sP4(X179,sK8) | ~p(s(bool,X180)) | p(s(bool,X181)) | s(bool,X180) != s(bool,X181)) ) | $spl3),
  inference(resolution,[],[f2430,f360])).
fof(f3478,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP4(X35,X36) | sP4(X36,sK8) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X35) != s(bool,X38) | s(bool,X37) != s(bool,X39)) ) | $spl3),
  inference(resolution,[],[f2430,f223])).
fof(f3477,plain,(
  ( ! [X30,X33,X31,X34,X32] : (sP4(X30,X31) | sP4(X31,sK8) | p(s(bool,X32)) | ~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X30) != s(bool,X33)) ) | $spl3),
  inference(resolution,[],[f2430,f222])).
fof(f3581,plain,(
  ( ! [X26,X29,X27,X25] : (sP4(X25,X26) | sP4(X26,sK8) | ~p(s(bool,X27)) | p(s(bool,X29)) | s(bool,X27) != s(bool,X29)) ) | $spl3),
  inference(subsumption_resolution,[],[f3476,f360])).
fof(f3476,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP4(X25,X26) | sP4(X26,sK8) | ~p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X27) != s(bool,X29)) ) | $spl3),
  inference(resolution,[],[f2430,f222])).
fof(f3475,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP4(X20,X21) | sP4(X21,sK8) | p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X24) | s(bool,X20) != s(bool,X23)) ) | $spl3),
  inference(resolution,[],[f2430,f221])).
fof(f3580,plain,(
  ( ! [X6,X5,X9] : (sP4(X5,X6) | s(bool,X5) != s(bool,X9) | sP4(X6,sK8)) ) | $spl90),
  inference(cnf_transformation,[],[f3580_D])).
fof(f3580_D,plain,(
  ( ! [X6,X5,X9] : (sP4(X5,X6) | s(bool,X5) != s(bool,X9) | sP4(X6,sK8)) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f3474,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP4(X15,X16) | sP4(X16,sK8) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X15) != s(bool,X19) | s(bool,X17) != s(bool,X18)) ) | $spl3),
  inference(resolution,[],[f2430,f221])).
fof(f3473,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP4(X10,X11) | sP4(X11,sK8) | ~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X10) != s(bool,X13) | s(bool,X12) != s(bool,X14)) ) | $spl3),
  inference(resolution,[],[f2430,f219])).
fof(f3472,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP4(X5,X6) | sP4(X6,sK8) | ~p(s(bool,X7)) | p(s(bool,X8)) | s(bool,X5) != s(bool,X9) | s(bool,X7) != s(bool,X8)) ) | $spl3),
  inference(resolution,[],[f2430,f217])).
fof(f3471,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP4(X0,X1) | sP4(X1,sK8) | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X0) != s(bool,X3)) ) | $spl3),
  inference(resolution,[],[f2430,f217])).
fof(f3452,plain,(
  ( ! [X50,X51] : (p(s(bool,X50)) | s(bool,'F_0') != s(bool,X51)) ) | $spl8),
  inference(duplicate_literal_removal,[],[f3440])).
fof(f3440,plain,(
  ( ! [X50,X51] : (p(s(bool,X50)) | s(bool,'F_0') != s(bool,X51) | p(s(bool,X50))) ) | $spl8),
  inference(superposition,[],[f453,f319])).
fof(f3453,plain,(
  ( ! [X48,X49] : (p(s(bool,X48)) | s(bool,'F_0') != s(bool,X49)) ) | $spl8),
  inference(duplicate_literal_removal,[],[f3439])).
fof(f3439,plain,(
  ( ! [X48,X49] : (p(s(bool,X48)) | s(bool,'F_0') != s(bool,X49) | p(s(bool,X48))) ) | $spl8),
  inference(superposition,[],[f453,f321])).
fof(f3467,plain,(
  ( ! [X47,X46] : (p(s(bool,X46)) | s(bool,'F_0') != s(bool,X47)) ) | $spl8),
  inference(subsumption_resolution,[],[f3438,f453])).
fof(f3438,plain,(
  ( ! [X47,X46] : (p(s(bool,X46)) | s(bool,'F_0') != s(bool,X47) | s(bool,'T_0') = s(bool,X46)) ) | $spl8),
  inference(superposition,[],[f453,f358])).
fof(f3465,plain,(
  ( ! [X41,X40] : (p(s(bool,X40)) | s(bool,'T_0') != s(bool,X41)) ) | $spl8),
  inference(subsumption_resolution,[],[f3435,f453])).
fof(f3435,plain,(
  ( ! [X41,X40] : (p(s(bool,X40)) | s(bool,'T_0') != s(bool,X41) | s(bool,'F_0') = s(bool,X40)) ) | $spl8),
  inference(superposition,[],[f453,f358])).
fof(f3417,plain,(
  ( ! [X200,X201] : (s(bool,'F_0') != s(bool,X200) | p(s(bool,X201))) ) | $spl8),
  inference(resolution,[],[f453,f331])).
fof(f3416,plain,(
  ( ! [X198,X199] : (s(bool,'F_0') != s(bool,X198) | p(s(bool,X199))) ) | $spl8),
  inference(resolution,[],[f453,f314])).
fof(f3415,plain,(
  ( ! [X196,X197] : (s(bool,'F_0') != s(bool,X196) | p(s(bool,X197))) ) | $spl8),
  inference(resolution,[],[f453,f307])).
fof(f3459,plain,(
  ( ! [X194,X195] : (s(bool,'F_0') != s(bool,X194) | p(s(bool,X195))) ) | $spl8),
  inference(subsumption_resolution,[],[f3414,f361])).
fof(f3414,plain,(
  ( ! [X194,X195] : (s(bool,'F_0') != s(bool,X194) | p(s(bool,X195)) | ~p(s(bool,'T_0'))) ) | $spl8),
  inference(resolution,[],[f453,f293])).
fof(f3413,plain,(
  ( ! [X192,X193] : (s(bool,'F_0') != s(bool,X192) | ~p(s(bool,X193))) ) | $spl8),
  inference(resolution,[],[f453,f289])).
fof(f3412,plain,(
  ( ! [X191,X189,X190] : (s(bool,'F_0') != s(bool,X189) | sP5(X190,X191)) ) | $spl8),
  inference(resolution,[],[f453,f277])).
fof(f3411,plain,(
  ( ! [X187,X188,X186] : (s(bool,'F_0') != s(bool,X186) | sP5(X187,X188)) ) | $spl8),
  inference(resolution,[],[f453,f272])).
fof(f3329,plain,(
  ( ! [X0,X1] : (sP0('const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0,X1)) ) | $spl82),
  inference(cnf_transformation,[],[f3329_D])).
fof(f3329_D,plain,(
  ( ! [X0,X1] : (sP0('const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0,X1)) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f3319,plain,(
  ( ! [X0,X1] : (sP0('const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0,X1) | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6))))) ) | $spl4),
  inference(superposition,[],[f707,f357])).
fof(f3314,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X1)) ) | $spl80),
  inference(cnf_transformation,[],[f3314_D])).
fof(f3314_D,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X1)) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f3304,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X1) | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6))))) ) | $spl4),
  inference(superposition,[],[f706,f357])).
fof(f706,plain,(
  ( ! [X45,X44] : (sP0(X44,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)),X45)) ) | $spl4),
  inference(resolution,[],[f369,f241])).
fof(f3277,plain,(
  sP4(sK8,'const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0'))) | $spl78),
  inference(cnf_transformation,[],[f3277_D])).
fof(f3277_D,plain,(
  sP4(sK8,'const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0'))) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f3265,plain,(
  sP4(sK8,'const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0'))) | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6)))) | ($spl3 | $spl4)),
  inference(superposition,[],[f2595,f357])).
fof(f2595,plain,(
  sP4(sK8,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7))) | ($spl3 | $spl4)),
  inference(resolution,[],[f1544,f369])).
fof(f3252,plain,(
  ( ! [X0,X1] : (sP0(X0,X1,'const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0')))) ) | $spl76),
  inference(cnf_transformation,[],[f3252_D])).
fof(f3252_D,plain,(
  ( ! [X0,X1] : (sP0(X0,X1,'const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0')))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f3242,plain,(
  ( ! [X0,X1] : (sP0(X0,X1,'const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0'))) | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6))))) ) | $spl4),
  inference(superposition,[],[f705,f357])).
fof(f705,plain,(
  ( ! [X43,X42] : (sP0(X42,X43,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) | $spl4),
  inference(resolution,[],[f369,f240])).
fof(f3197,plain,(
  ( ! [X0,X1] : (sP1(X0,X1,'const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0')))) ) | $spl74),
  inference(cnf_transformation,[],[f3197_D])).
fof(f3197_D,plain,(
  ( ! [X0,X1] : (sP1(X0,X1,'const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0')))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f3188,plain,(
  ( ! [X0,X1] : (sP1(X0,X1,'const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0'))) | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6))))) ) | $spl4),
  inference(superposition,[],[f704,f357])).
fof(f704,plain,(
  ( ! [X41,X40] : (sP1(X40,X41,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) | $spl4),
  inference(resolution,[],[f369,f236])).
fof(f3158,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0')))) ) | $spl72),
  inference(cnf_transformation,[],[f3158_D])).
fof(f3158_D,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0')))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f3149,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0'))) | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6))))) ) | $spl4),
  inference(superposition,[],[f701,f357])).
fof(f701,plain,(
  ( ! [X35,X34] : (sP2(X34,X35,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) | $spl4),
  inference(resolution,[],[f369,f232])).
fof(f3119,plain,(
  sP4('const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0')),sK8) | $spl70),
  inference(cnf_transformation,[],[f3119_D])).
fof(f3119_D,plain,(
  sP4('const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0')),sK8) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f3107,plain,(
  sP4('const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0')),sK8) | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6)))) | ($spl3 | $spl4)),
  inference(superposition,[],[f858,f357])).
fof(f858,plain,(
  sP4('const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)),sK8) | ($spl3 | $spl4)),
  inference(resolution,[],[f492,f369])).
fof(f492,plain,(
  ( ! [X54] : (~p(s(bool,X54)) | sP4(X54,sK8)) ) | $spl3),
  inference(resolution,[],[f367,f255])).
fof(f2889,plain,(
  ( ! [X85,X86,X84] : (s(bool,'F_0') != s(bool,X84) | p(s(bool,X85)) | ~p(s(bool,X86)) | s(bool,X85) != s(bool,X86)) ) | $spl3),
  inference(subsumption_resolution,[],[f2688,f221])).
fof(f2688,plain,(
  ( ! [X85,X86,X84] : (s(bool,'F_0') != s(bool,X84) | p(s(bool,X85)) | ~p(s(bool,X84)) | ~p(s(bool,X86)) | s(bool,X85) != s(bool,X86)) ) | $spl3),
  inference(superposition,[],[f223,f1561])).
fof(f2888,plain,(
  ( ! [X83,X81,X82] : (s(bool,'F_0') != s(bool,X81) | p(s(bool,X82)) | ~p(s(bool,X83)) | s(bool,X82) != s(bool,X83)) ) | $spl3),
  inference(subsumption_resolution,[],[f2687,f360])).
fof(f2687,plain,(
  ( ! [X83,X81,X82] : (s(bool,'F_0') != s(bool,X81) | p(s(bool,X82)) | ~p(s(bool,sK8)) | ~p(s(bool,X83)) | s(bool,X82) != s(bool,X83)) ) | $spl3),
  inference(superposition,[],[f223,f1561])).
fof(f2872,plain,(
  ( ! [X26,X24,X25] : (s(bool,'F_0') != s(bool,X24) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X26)) ) | $spl3),
  inference(subsumption_resolution,[],[f2668,f360])).
fof(f2668,plain,(
  ( ! [X26,X24,X25] : (s(bool,'F_0') != s(bool,X24) | p(s(bool,X25)) | ~p(s(bool,sK8)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X26)) ) | $spl3),
  inference(superposition,[],[f219,f1561])).
fof(f2870,plain,(
  ( ! [X23,X21,X22] : (s(bool,'F_0') != s(bool,X21) | p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X23)) ) | $spl3),
  inference(subsumption_resolution,[],[f2667,f221])).
fof(f2667,plain,(
  ( ! [X23,X21,X22] : (s(bool,'F_0') != s(bool,X21) | p(s(bool,X22)) | ~p(s(bool,X21)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X23)) ) | $spl3),
  inference(superposition,[],[f219,f1561])).
fof(f2594,plain,(
  ( ! [X144,X145] : (sP4(sK8,'F_0') | p(s(bool,X144)) | p(s(bool,X145)) | ~sP5(X144,X145)) ) | $spl3),
  inference(resolution,[],[f1544,f269])).
fof(f2592,plain,(
  ( ! [X140,X141] : (sP4(sK8,'F_0') | p(s(bool,X140)) | p(s(bool,X141)) | ~sP5(X140,X141)) ) | $spl3),
  inference(resolution,[],[f1544,f273])).
fof(f1490,plain,(
  ( ! [X2,X0] : (sP4(sK8,X0) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) | $spl48),
  inference(cnf_transformation,[],[f1490_D])).
fof(f1490_D,plain,(
  ( ! [X2,X0] : (sP4(sK8,X0) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f2575,plain,(
  ( ! [X132,X130,X131] : (sP4(sK8,X130) | p(s(bool,X131)) | ~p(s(bool,X132)) | s(bool,X130) != s(bool,X132)) ) | $spl3),
  inference(resolution,[],[f1544,f360])).
fof(f2574,plain,(
  ( ! [X127,X128,X129] : (sP4(sK8,X127) | ~p(s(bool,X128)) | p(s(bool,X129)) | s(bool,X128) != s(bool,X129)) ) | $spl3),
  inference(resolution,[],[f1544,f360])).
fof(f2530,plain,(
  ( ! [X30,X28,X31,X29] : (sP4(sK8,X28) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) ) | $spl3),
  inference(resolution,[],[f1544,f223])).
fof(f2529,plain,(
  ( ! [X26,X24,X27,X25] : (sP4(sK8,X24) | p(s(bool,X25)) | ~p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X24) != s(bool,X26)) ) | $spl3),
  inference(resolution,[],[f1544,f222])).
fof(f2630,plain,(
  ( ! [X23,X21,X20] : (sP4(sK8,X20) | ~p(s(bool,X21)) | p(s(bool,X23)) | s(bool,X21) != s(bool,X23)) ) | $spl3),
  inference(subsumption_resolution,[],[f2528,f360])).
fof(f2528,plain,(
  ( ! [X23,X21,X22,X20] : (sP4(sK8,X20) | ~p(s(bool,X21)) | ~p(s(bool,X22)) | p(s(bool,X23)) | s(bool,X21) != s(bool,X23)) ) | $spl3),
  inference(resolution,[],[f1544,f222])).
fof(f2527,plain,(
  ( ! [X19,X17,X18,X16] : (sP4(sK8,X16) | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) ) | $spl3),
  inference(resolution,[],[f1544,f221])).
fof(f1491,plain,(
  ( ! [X4,X7] : (sP4(sK8,X4) | s(bool,X4) != s(bool,X7)) ) | $spl50),
  inference(cnf_transformation,[],[f1491_D])).
fof(f1491_D,plain,(
  ( ! [X4,X7] : (sP4(sK8,X4) | s(bool,X4) != s(bool,X7)) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f2526,plain,(
  ( ! [X14,X12,X15,X13] : (sP4(sK8,X12) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X12) != s(bool,X15) | s(bool,X13) != s(bool,X14)) ) | $spl3),
  inference(resolution,[],[f1544,f221])).
fof(f2525,plain,(
  ( ! [X10,X8,X11,X9] : (sP4(sK8,X8) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) ) | $spl3),
  inference(resolution,[],[f1544,f219])).
fof(f2524,plain,(
  ( ! [X6,X4,X7,X5] : (sP4(sK8,X4) | ~p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X4) != s(bool,X7) | s(bool,X5) != s(bool,X6)) ) | $spl3),
  inference(resolution,[],[f1544,f217])).
fof(f2523,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(sK8,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl3),
  inference(resolution,[],[f1544,f217])).
fof(f2477,plain,(
  ( ! [X144,X145] : (sP4('F_0',sK8) | p(s(bool,X144)) | p(s(bool,X145)) | ~sP5(X144,X145)) ) | $spl3),
  inference(resolution,[],[f1543,f269])).
fof(f2475,plain,(
  ( ! [X140,X141] : (sP4('F_0',sK8) | p(s(bool,X140)) | p(s(bool,X141)) | ~sP5(X140,X141)) ) | $spl3),
  inference(resolution,[],[f1543,f273])).
fof(f893,plain,(
  ( ! [X2,X0] : (sP4(X0,sK8) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) | $spl26),
  inference(cnf_transformation,[],[f893_D])).
fof(f893_D,plain,(
  ( ! [X2,X0] : (sP4(X0,sK8) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f2458,plain,(
  ( ! [X132,X130,X131] : (sP4(X130,sK8) | p(s(bool,X131)) | ~p(s(bool,X132)) | s(bool,X130) != s(bool,X132)) ) | $spl3),
  inference(resolution,[],[f1543,f360])).
fof(f2457,plain,(
  ( ! [X127,X128,X129] : (sP4(X127,sK8) | ~p(s(bool,X128)) | p(s(bool,X129)) | s(bool,X128) != s(bool,X129)) ) | $spl3),
  inference(resolution,[],[f1543,f360])).
fof(f2413,plain,(
  ( ! [X30,X28,X31,X29] : (sP4(X28,sK8) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) ) | $spl3),
  inference(resolution,[],[f1543,f223])).
fof(f2412,plain,(
  ( ! [X26,X24,X27,X25] : (sP4(X24,sK8) | p(s(bool,X25)) | ~p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X24) != s(bool,X26)) ) | $spl3),
  inference(resolution,[],[f1543,f222])).
fof(f2513,plain,(
  ( ! [X23,X21,X20] : (sP4(X20,sK8) | ~p(s(bool,X21)) | p(s(bool,X23)) | s(bool,X21) != s(bool,X23)) ) | $spl3),
  inference(subsumption_resolution,[],[f2411,f360])).
fof(f2411,plain,(
  ( ! [X23,X21,X22,X20] : (sP4(X20,sK8) | ~p(s(bool,X21)) | ~p(s(bool,X22)) | p(s(bool,X23)) | s(bool,X21) != s(bool,X23)) ) | $spl3),
  inference(resolution,[],[f1543,f222])).
fof(f2410,plain,(
  ( ! [X19,X17,X18,X16] : (sP4(X16,sK8) | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) ) | $spl3),
  inference(resolution,[],[f1543,f221])).
fof(f894,plain,(
  ( ! [X4,X7] : (sP4(X4,sK8) | s(bool,X4) != s(bool,X7)) ) | $spl28),
  inference(cnf_transformation,[],[f894_D])).
fof(f894_D,plain,(
  ( ! [X4,X7] : (sP4(X4,sK8) | s(bool,X4) != s(bool,X7)) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f2409,plain,(
  ( ! [X14,X12,X15,X13] : (sP4(X12,sK8) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X12) != s(bool,X15) | s(bool,X13) != s(bool,X14)) ) | $spl3),
  inference(resolution,[],[f1543,f221])).
fof(f2408,plain,(
  ( ! [X10,X8,X11,X9] : (sP4(X8,sK8) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) ) | $spl3),
  inference(resolution,[],[f1543,f219])).
fof(f2407,plain,(
  ( ! [X6,X4,X7,X5] : (sP4(X4,sK8) | ~p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X4) != s(bool,X7) | s(bool,X5) != s(bool,X6)) ) | $spl3),
  inference(resolution,[],[f1543,f217])).
fof(f2406,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(X0,sK8) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl3),
  inference(resolution,[],[f1543,f217])).
fof(f1248,plain,(
  s('type.option.option'(sK6),'const.option.NONE_0') = s('type.option.option'(sK6),'const.bool.COND_3'(s(bool,'T_0'),s('type.option.option'(sK6),'const.option.NONE_0'),s('type.option.option'(sK6),sK7))) | $spl42),
  inference(cnf_transformation,[],[f1248_D])).
fof(f1248_D,plain,(
  s('type.option.option'(sK6),'const.option.NONE_0') = s('type.option.option'(sK6),'const.bool.COND_3'(s(bool,'T_0'),s('type.option.option'(sK6),'const.option.NONE_0'),s('type.option.option'(sK6),sK7))) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f2321,plain,(
  s('type.option.option'(sK6),'const.option.NONE_0') = s('type.option.option'(sK6),'const.bool.COND_3'(s(bool,'T_0'),s('type.option.option'(sK6),'const.option.NONE_0'),s('type.option.option'(sK6),sK7))) | s(bool,'F_0') = s(bool,sK8) | $spl0),
  inference(superposition,[],[f364,f358])).
fof(f543,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,sK8) != s(bool,X1)) ) | $spl20),
  inference(cnf_transformation,[],[f543_D])).
fof(f543_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,sK8) != s(bool,X1)) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f2234,plain,(
  ( ! [X78,X79] : (p(s(bool,X78)) | ~p(s(bool,X79)) | s(bool,sK8) != s(bool,X79)) ) | $spl3),
  inference(resolution,[],[f367,f360])).
fof(f2189,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,sK8) != s(bool,X22) | s(bool,X21) != s(bool,X23)) ) | $spl3),
  inference(resolution,[],[f367,f223])).
fof(f2188,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,sK8) != s(bool,X19)) ) | $spl3),
  inference(resolution,[],[f367,f222])).
fof(f2186,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,sK8) != s(bool,X13)) ) | $spl3),
  inference(resolution,[],[f367,f221])).
fof(f544,plain,(
  ( ! [X5] : (s(bool,sK8) != s(bool,X5)) ) | $spl22),
  inference(cnf_transformation,[],[f544_D])).
fof(f544_D,plain,(
  ( ! [X5] : (s(bool,sK8) != s(bool,X5)) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f2185,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | p(s(bool,X10)) | s(bool,sK8) != s(bool,X11) | s(bool,X9) != s(bool,X10)) ) | $spl3),
  inference(resolution,[],[f367,f221])).
fof(f2184,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,sK8) != s(bool,X7) | s(bool,X6) != s(bool,X8)) ) | $spl3),
  inference(resolution,[],[f367,f219])).
fof(f2183,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,sK8) != s(bool,X5) | s(bool,X3) != s(bool,X4)) ) | $spl3),
  inference(resolution,[],[f367,f217])).
fof(f2182,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,sK8) != s(bool,X1)) ) | $spl3),
  inference(resolution,[],[f367,f217])).
fof(f2046,plain,(
  $false | ($spl3 | $spl44)),
  inference(subsumption_resolution,[],[f1841,f338])).
fof(f1841,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl44)),
  inference(backward_demodulation,[],[f1251,f367])).
fof(f1839,plain,(
  ( ! [X0] : (sP5('const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0)) ) | $spl66),
  inference(cnf_transformation,[],[f1839_D])).
fof(f1839_D,plain,(
  ( ! [X0] : (sP5('const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0)) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f1830,plain,(
  ( ! [X0] : (sP5('const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0) | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6))))) ) | $spl4),
  inference(superposition,[],[f719,f357])).
fof(f719,plain,(
  ( ! [X61] : (sP5('const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)),X61)) ) | $spl4),
  inference(resolution,[],[f369,f271])).
fof(f1827,plain,(
  ( ! [X0] : (sP5(X0,'const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0')))) ) | $spl64),
  inference(cnf_transformation,[],[f1827_D])).
fof(f1827_D,plain,(
  ( ! [X0] : (sP5(X0,'const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0')))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f1816,plain,(
  ( ! [X0] : (sP5(X0,'const.option.IS_SOME_1'(s('type.option.option'(sK6),'const.option.NONE_0'))) | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6))))) ) | $spl4),
  inference(superposition,[],[f718,f357])).
fof(f718,plain,(
  ( ! [X60] : (sP5(X60,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) | $spl4),
  inference(resolution,[],[f369,f270])).
fof(f1252,plain,(
  s('type.option.option'(sK6),'const.option.NONE_0') = s('type.option.option'(sK6),'const.bool.COND_3'(s(bool,'F_0'),s('type.option.option'(sK6),'const.option.NONE_0'),s('type.option.option'(sK6),sK7))) | $spl46),
  inference(cnf_transformation,[],[f1252_D])).
fof(f1252_D,plain,(
  s('type.option.option'(sK6),'const.option.NONE_0') = s('type.option.option'(sK6),'const.bool.COND_3'(s(bool,'F_0'),s('type.option.option'(sK6),'const.option.NONE_0'),s('type.option.option'(sK6),sK7))) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f1717,plain,(
  s('type.option.option'(sK6),'const.option.NONE_0') = s('type.option.option'(sK6),'const.bool.COND_3'(s(bool,'F_0'),s('type.option.option'(sK6),'const.option.NONE_0'),s('type.option.option'(sK6),sK7))) | s(bool,'T_0') = s(bool,sK8) | $spl0),
  inference(superposition,[],[f364,f358])).
fof(f1716,plain,(
  s('type.option.option'(sK6),'const.option.NONE_0') = s('type.option.option'(sK6),'const.bool.COND_3'(s(bool,'T_0'),s('type.option.option'(sK6),'const.option.NONE_0'),s('type.option.option'(sK6),sK7))) | s(bool,'F_0') = s(bool,sK8) | $spl0),
  inference(superposition,[],[f364,f358])).
fof(f1655,plain,(
  ( ! [X76,X77] : (p(s(bool,X76)) | p(s(bool,X77)) | s(bool,X77) != s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) | $spl4),
  inference(resolution,[],[f369,f360])).
fof(f1613,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7))) | s(bool,X22) != s(bool,X23)) ) | $spl4),
  inference(resolution,[],[f369,f223])).
fof(f454,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) | $spl10),
  inference(cnf_transformation,[],[f454_D])).
fof(f454_D,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f1612,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,X19) | s(bool,X20) != s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) | $spl4),
  inference(resolution,[],[f369,f223])).
fof(f1610,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) | $spl4),
  inference(resolution,[],[f369,f222])).
fof(f1609,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) | $spl4),
  inference(resolution,[],[f369,f221])).
fof(f1608,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7))) | s(bool,X7) != s(bool,X8)) ) | $spl4),
  inference(resolution,[],[f369,f219])).
fof(f1607,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X4) | s(bool,X5) != s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) | $spl4),
  inference(resolution,[],[f369,f219])).
fof(f1606,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) | $spl4),
  inference(resolution,[],[f369,f217])).
fof(f1571,plain,(
  ( ! [X78,X79] : (p(s(bool,X78)) | ~p(s(bool,X79)) | s(bool,sK8) != s(bool,X79)) ) | $spl3),
  inference(resolution,[],[f367,f360])).
fof(f1526,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,sK8) != s(bool,X22) | s(bool,X21) != s(bool,X23)) ) | $spl3),
  inference(resolution,[],[f367,f223])).
fof(f1525,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,sK8) != s(bool,X19)) ) | $spl3),
  inference(resolution,[],[f367,f222])).
fof(f1523,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,sK8) != s(bool,X13)) ) | $spl3),
  inference(resolution,[],[f367,f221])).
fof(f1522,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | p(s(bool,X10)) | s(bool,sK8) != s(bool,X11) | s(bool,X9) != s(bool,X10)) ) | $spl3),
  inference(resolution,[],[f367,f221])).
fof(f1521,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,sK8) != s(bool,X7) | s(bool,X6) != s(bool,X8)) ) | $spl3),
  inference(resolution,[],[f367,f219])).
fof(f1520,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,sK8) != s(bool,X5) | s(bool,X3) != s(bool,X4)) ) | $spl3),
  inference(resolution,[],[f367,f217])).
fof(f1519,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,sK8) != s(bool,X1)) ) | $spl3),
  inference(resolution,[],[f367,f217])).
fof(f1512,plain,(
  ( ! [X9] : (s('type.option.option'(sK6),'const.option.NONE_0') != s('type.option.option'(sK6),'const.bool.COND_3'(s(bool,X9),s('type.option.option'(sK6),'const.option.NONE_0'),s('type.option.option'(sK6),sK7))) | p(s(bool,X9))) ) | $spl60),
  inference(cnf_transformation,[],[f1512_D])).
fof(f1512_D,plain,(
  ( ! [X9] : (s('type.option.option'(sK6),'const.option.NONE_0') != s('type.option.option'(sK6),'const.bool.COND_3'(s(bool,X9),s('type.option.option'(sK6),'const.option.NONE_0'),s('type.option.option'(sK6),sK7))) | p(s(bool,X9))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f1203,plain,(
  ( ! [X8] : (s('type.option.option'(sK6),'const.option.NONE_0') != s('type.option.option'(sK6),'const.bool.COND_3'(s(bool,X8),s('type.option.option'(sK6),'const.option.NONE_0'),s('type.option.option'(sK6),sK7))) | p(s(bool,X8)) | p(s(bool,sK8))) ) | $spl1),
  inference(superposition,[],[f365,f278])).
fof(f278,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f126])).
fof(f1204,plain,(
  ( ! [X9] : (s('type.option.option'(sK6),'const.option.NONE_0') != s('type.option.option'(sK6),'const.bool.COND_3'(s(bool,X9),s('type.option.option'(sK6),'const.option.NONE_0'),s('type.option.option'(sK6),sK7))) | p(s(bool,sK8)) | p(s(bool,X9))) ) | $spl1),
  inference(superposition,[],[f365,f278])).
fof(f1511,plain,(
  ( ! [X11] : (s('type.option.option'(sK6),'const.option.NONE_0') != s('type.option.option'(sK6),'const.bool.COND_3'(s(bool,X11),s('type.option.option'(sK6),'const.option.NONE_0'),s('type.option.option'(sK6),sK7))) | ~p(s(bool,X11))) ) | $spl58),
  inference(cnf_transformation,[],[f1511_D])).
fof(f1511_D,plain,(
  ( ! [X11] : (s('type.option.option'(sK6),'const.option.NONE_0') != s('type.option.option'(sK6),'const.bool.COND_3'(s(bool,X11),s('type.option.option'(sK6),'const.option.NONE_0'),s('type.option.option'(sK6),sK7))) | ~p(s(bool,X11))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f1205,plain,(
  ( ! [X10] : (s('type.option.option'(sK6),'const.option.NONE_0') != s('type.option.option'(sK6),'const.bool.COND_3'(s(bool,X10),s('type.option.option'(sK6),'const.option.NONE_0'),s('type.option.option'(sK6),sK7))) | ~p(s(bool,sK8)) | ~p(s(bool,X10))) ) | $spl1),
  inference(superposition,[],[f365,f281])).
fof(f1206,plain,(
  ( ! [X11] : (s('type.option.option'(sK6),'const.option.NONE_0') != s('type.option.option'(sK6),'const.bool.COND_3'(s(bool,X11),s('type.option.option'(sK6),'const.option.NONE_0'),s('type.option.option'(sK6),sK7))) | ~p(s(bool,X11)) | ~p(s(bool,sK8))) ) | $spl1),
  inference(superposition,[],[f365,f281])).
fof(f1207,plain,(
  s('type.option.option'(sK6),'const.option.NONE_0') != s('type.option.option'(sK6),'const.bool.COND_3'(s(bool,'F_0'),s('type.option.option'(sK6),'const.option.NONE_0'),s('type.option.option'(sK6),sK7))) | p(s(bool,sK8)) | $spl1),
  inference(superposition,[],[f365,f319])).
fof(f1208,plain,(
  s('type.option.option'(sK6),'const.option.NONE_0') != s('type.option.option'(sK6),'const.bool.COND_3'(s(bool,'F_0'),s('type.option.option'(sK6),'const.option.NONE_0'),s('type.option.option'(sK6),sK7))) | p(s(bool,sK8)) | $spl1),
  inference(superposition,[],[f365,f321])).
fof(f1209,plain,(
  s('type.option.option'(sK6),'const.option.NONE_0') != s('type.option.option'(sK6),'const.bool.COND_3'(s(bool,'T_0'),s('type.option.option'(sK6),'const.option.NONE_0'),s('type.option.option'(sK6),sK7))) | ~p(s(bool,sK8)) | $spl1),
  inference(superposition,[],[f365,f323])).
fof(f1210,plain,(
  s('type.option.option'(sK6),'const.option.NONE_0') != s('type.option.option'(sK6),'const.bool.COND_3'(s(bool,'T_0'),s('type.option.option'(sK6),'const.option.NONE_0'),s('type.option.option'(sK6),sK7))) | ~p(s(bool,sK8)) | $spl1),
  inference(superposition,[],[f365,f325])).
fof(f1500,plain,(
  sP4(sK8,'T_0') | $spl54),
  inference(cnf_transformation,[],[f1500_D])).
fof(f1500_D,plain,(
  sP4(sK8,'T_0') <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f1486,plain,(
  ( ! [X28] : (~p(s(bool,X28)) | sP4(sK8,'T_0')) ) | $spl3),
  inference(duplicate_literal_removal,[],[f1472])).
fof(f1472,plain,(
  ( ! [X28] : (~p(s(bool,X28)) | sP4(sK8,'T_0') | ~p(s(bool,X28))) ) | $spl3),
  inference(superposition,[],[f1023,f323])).
fof(f1023,plain,(
  ( ! [X55] : (~p(s(bool,X55)) | sP4(sK8,X55)) ) | $spl3),
  inference(resolution,[],[f367,f258])).
fof(f1487,plain,(
  ( ! [X27] : (~p(s(bool,X27)) | sP4(sK8,'T_0')) ) | $spl3),
  inference(duplicate_literal_removal,[],[f1471])).
fof(f1471,plain,(
  ( ! [X27] : (~p(s(bool,X27)) | sP4(sK8,'T_0') | ~p(s(bool,X27))) ) | $spl3),
  inference(superposition,[],[f1023,f325])).
fof(f1508,plain,(
  ( ! [X26] : (~p(s(bool,X26)) | sP4(sK8,'T_0')) ) | $spl3),
  inference(subsumption_resolution,[],[f1470,f318])).
fof(f1470,plain,(
  ( ! [X26] : (~p(s(bool,X26)) | sP4(sK8,'T_0') | s(bool,'F_0') = s(bool,X26)) ) | $spl3),
  inference(superposition,[],[f1023,f358])).
fof(f1455,plain,(
  ( ! [X144,X145] : (sP4(sK8,'F_0') | p(s(bool,X144)) | p(s(bool,X145)) | ~sP5(X144,X145)) ) | $spl3),
  inference(resolution,[],[f1023,f269])).
fof(f1453,plain,(
  ( ! [X140,X141] : (sP4(sK8,'F_0') | p(s(bool,X140)) | p(s(bool,X141)) | ~sP5(X140,X141)) ) | $spl3),
  inference(resolution,[],[f1023,f273])).
fof(f1444,plain,(
  ( ! [X139] : (sP4(sK8,'T_0') | ~p(s(bool,X139))) ) | $spl3),
  inference(resolution,[],[f1023,f316])).
fof(f1443,plain,(
  ( ! [X138] : (sP4(sK8,'T_0') | ~p(s(bool,X138))) ) | $spl3),
  inference(resolution,[],[f1023,f311])).
fof(f1442,plain,(
  ( ! [X137] : (sP4(sK8,'T_0') | ~p(s(bool,X137))) ) | $spl3),
  inference(resolution,[],[f1023,f301])).
fof(f1441,plain,(
  ( ! [X136] : (sP4(sK8,'T_0') | p(s(bool,X136))) ) | $spl3),
  inference(resolution,[],[f1023,f300])).
fof(f1440,plain,(
  ( ! [X135] : (sP4(sK8,'T_0') | p(s(bool,X135))) ) | $spl3),
  inference(resolution,[],[f1023,f297])).
fof(f1439,plain,(
  ( ! [X134] : (sP4(sK8,'T_0') | p(s(bool,X134))) ) | $spl3),
  inference(resolution,[],[f1023,f294])).
fof(f1438,plain,(
  ( ! [X133] : (sP4(sK8,'T_0') | ~p(s(bool,X133))) ) | $spl3),
  inference(resolution,[],[f1023,f292])).
fof(f1436,plain,(
  ( ! [X132,X130,X131] : (sP4(sK8,X130) | p(s(bool,X131)) | ~p(s(bool,X132)) | s(bool,X130) != s(bool,X132)) ) | $spl3),
  inference(resolution,[],[f1023,f360])).
fof(f1435,plain,(
  ( ! [X127,X128,X129] : (sP4(sK8,X127) | ~p(s(bool,X128)) | p(s(bool,X129)) | s(bool,X128) != s(bool,X129)) ) | $spl3),
  inference(resolution,[],[f1023,f360])).
fof(f1391,plain,(
  ( ! [X30,X28,X31,X29] : (sP4(sK8,X28) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) ) | $spl3),
  inference(resolution,[],[f1023,f223])).
fof(f1390,plain,(
  ( ! [X26,X24,X27,X25] : (sP4(sK8,X24) | p(s(bool,X25)) | ~p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X24) != s(bool,X26)) ) | $spl3),
  inference(resolution,[],[f1023,f222])).
fof(f1492,plain,(
  ( ! [X23,X21,X20] : (sP4(sK8,X20) | ~p(s(bool,X21)) | p(s(bool,X23)) | s(bool,X21) != s(bool,X23)) ) | $spl3),
  inference(subsumption_resolution,[],[f1389,f360])).
fof(f1389,plain,(
  ( ! [X23,X21,X22,X20] : (sP4(sK8,X20) | ~p(s(bool,X21)) | ~p(s(bool,X22)) | p(s(bool,X23)) | s(bool,X21) != s(bool,X23)) ) | $spl3),
  inference(resolution,[],[f1023,f222])).
fof(f1388,plain,(
  ( ! [X19,X17,X18,X16] : (sP4(sK8,X16) | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) ) | $spl3),
  inference(resolution,[],[f1023,f221])).
fof(f1387,plain,(
  ( ! [X14,X12,X15,X13] : (sP4(sK8,X12) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X12) != s(bool,X15) | s(bool,X13) != s(bool,X14)) ) | $spl3),
  inference(resolution,[],[f1023,f221])).
fof(f1386,plain,(
  ( ! [X10,X8,X11,X9] : (sP4(sK8,X8) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) ) | $spl3),
  inference(resolution,[],[f1023,f219])).
fof(f1385,plain,(
  ( ! [X6,X4,X7,X5] : (sP4(sK8,X4) | ~p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X4) != s(bool,X7) | s(bool,X5) != s(bool,X6)) ) | $spl3),
  inference(resolution,[],[f1023,f217])).
fof(f1384,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(sK8,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl3),
  inference(resolution,[],[f1023,f217])).
fof(f1339,plain,(
  ( ! [X144,X145] : (sP4('F_0',sK8) | p(s(bool,X144)) | p(s(bool,X145)) | ~sP5(X144,X145)) ) | $spl3),
  inference(resolution,[],[f1022,f269])).
fof(f1022,plain,(
  ( ! [X54] : (~p(s(bool,X54)) | sP4(X54,sK8)) ) | $spl3),
  inference(resolution,[],[f367,f255])).
fof(f1337,plain,(
  ( ! [X140,X141] : (sP4('F_0',sK8) | p(s(bool,X140)) | p(s(bool,X141)) | ~sP5(X140,X141)) ) | $spl3),
  inference(resolution,[],[f1022,f273])).
fof(f1320,plain,(
  ( ! [X132,X130,X131] : (sP4(X130,sK8) | p(s(bool,X131)) | ~p(s(bool,X132)) | s(bool,X130) != s(bool,X132)) ) | $spl3),
  inference(resolution,[],[f1022,f360])).
fof(f1319,plain,(
  ( ! [X127,X128,X129] : (sP4(X127,sK8) | ~p(s(bool,X128)) | p(s(bool,X129)) | s(bool,X128) != s(bool,X129)) ) | $spl3),
  inference(resolution,[],[f1022,f360])).
fof(f1275,plain,(
  ( ! [X30,X28,X31,X29] : (sP4(X28,sK8) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) ) | $spl3),
  inference(resolution,[],[f1022,f223])).
fof(f1274,plain,(
  ( ! [X26,X24,X27,X25] : (sP4(X24,sK8) | p(s(bool,X25)) | ~p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X24) != s(bool,X26)) ) | $spl3),
  inference(resolution,[],[f1022,f222])).
fof(f1374,plain,(
  ( ! [X23,X21,X20] : (sP4(X20,sK8) | ~p(s(bool,X21)) | p(s(bool,X23)) | s(bool,X21) != s(bool,X23)) ) | $spl3),
  inference(subsumption_resolution,[],[f1273,f360])).
fof(f1273,plain,(
  ( ! [X23,X21,X22,X20] : (sP4(X20,sK8) | ~p(s(bool,X21)) | ~p(s(bool,X22)) | p(s(bool,X23)) | s(bool,X21) != s(bool,X23)) ) | $spl3),
  inference(resolution,[],[f1022,f222])).
fof(f1272,plain,(
  ( ! [X19,X17,X18,X16] : (sP4(X16,sK8) | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) ) | $spl3),
  inference(resolution,[],[f1022,f221])).
fof(f1271,plain,(
  ( ! [X14,X12,X15,X13] : (sP4(X12,sK8) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X12) != s(bool,X15) | s(bool,X13) != s(bool,X14)) ) | $spl3),
  inference(resolution,[],[f1022,f221])).
fof(f1270,plain,(
  ( ! [X10,X8,X11,X9] : (sP4(X8,sK8) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) ) | $spl3),
  inference(resolution,[],[f1022,f219])).
fof(f1269,plain,(
  ( ! [X6,X4,X7,X5] : (sP4(X4,sK8) | ~p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X4) != s(bool,X7) | s(bool,X5) != s(bool,X6)) ) | $spl3),
  inference(resolution,[],[f1022,f217])).
fof(f1268,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(X0,sK8) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl3),
  inference(resolution,[],[f1022,f217])).
fof(f1212,plain,(
  s('type.option.option'(sK6),'const.option.NONE_0') != s('type.option.option'(sK6),'const.bool.COND_3'(s(bool,'F_0'),s('type.option.option'(sK6),'const.option.NONE_0'),s('type.option.option'(sK6),sK7))) | s(bool,'T_0') = s(bool,sK8) | $spl1),
  inference(superposition,[],[f365,f358])).
fof(f1211,plain,(
  s('type.option.option'(sK6),'const.option.NONE_0') != s('type.option.option'(sK6),'const.bool.COND_3'(s(bool,'T_0'),s('type.option.option'(sK6),'const.option.NONE_0'),s('type.option.option'(sK6),sK7))) | s(bool,'F_0') = s(bool,sK8) | $spl1),
  inference(superposition,[],[f365,f358])).
fof(f1133,plain,(
  ( ! [X78,X79] : (p(s(bool,X78)) | ~p(s(bool,X79)) | s(bool,X79) != s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) | $spl5),
  inference(resolution,[],[f368,f360])).
fof(f1088,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7))) | s(bool,X21) != s(bool,X23)) ) | $spl5),
  inference(resolution,[],[f368,f223])).
fof(f1087,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) | $spl5),
  inference(resolution,[],[f368,f222])).
fof(f1085,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) | $spl5),
  inference(resolution,[],[f368,f221])).
fof(f1084,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X11) != s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7))) | s(bool,X9) != s(bool,X10)) ) | $spl5),
  inference(resolution,[],[f368,f221])).
fof(f1083,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7))) | s(bool,X6) != s(bool,X8)) ) | $spl5),
  inference(resolution,[],[f368,f219])).
fof(f1082,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7))) | s(bool,X3) != s(bool,X4)) ) | $spl5),
  inference(resolution,[],[f368,f217])).
fof(f1081,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) | $spl5),
  inference(resolution,[],[f368,f217])).
fof(f1050,plain,(
  ( ! [X78,X79] : (p(s(bool,X78)) | ~p(s(bool,X79)) | s(bool,sK8) != s(bool,X79)) ) | $spl3),
  inference(resolution,[],[f367,f360])).
fof(f1005,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,sK8) != s(bool,X22) | s(bool,X21) != s(bool,X23)) ) | $spl3),
  inference(resolution,[],[f367,f223])).
fof(f1004,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,sK8) != s(bool,X19)) ) | $spl3),
  inference(resolution,[],[f367,f222])).
fof(f1002,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,sK8) != s(bool,X13)) ) | $spl3),
  inference(resolution,[],[f367,f221])).
fof(f1001,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | p(s(bool,X10)) | s(bool,sK8) != s(bool,X11) | s(bool,X9) != s(bool,X10)) ) | $spl3),
  inference(resolution,[],[f367,f221])).
fof(f1000,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,sK8) != s(bool,X7) | s(bool,X6) != s(bool,X8)) ) | $spl3),
  inference(resolution,[],[f367,f219])).
fof(f999,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,sK8) != s(bool,X5) | s(bool,X3) != s(bool,X4)) ) | $spl3),
  inference(resolution,[],[f367,f217])).
fof(f998,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,sK8) != s(bool,X1)) ) | $spl3),
  inference(resolution,[],[f367,f217])).
fof(f964,plain,(
  ( ! [X76,X77] : (p(s(bool,X76)) | p(s(bool,X77)) | s(bool,sK8) != s(bool,X77)) ) | $spl2),
  inference(resolution,[],[f366,f360])).
fof(f922,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,sK8) != s(bool,X21) | s(bool,X22) != s(bool,X23)) ) | $spl2),
  inference(resolution,[],[f366,f223])).
fof(f921,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,X19) | s(bool,sK8) != s(bool,X20)) ) | $spl2),
  inference(resolution,[],[f366,f223])).
fof(f919,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,sK8) != s(bool,X14)) ) | $spl2),
  inference(resolution,[],[f366,f222])).
fof(f918,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,sK8) != s(bool,X10)) ) | $spl2),
  inference(resolution,[],[f366,f221])).
fof(f917,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,sK8) != s(bool,X6) | s(bool,X7) != s(bool,X8)) ) | $spl2),
  inference(resolution,[],[f366,f219])).
fof(f916,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X4) | s(bool,sK8) != s(bool,X5)) ) | $spl2),
  inference(resolution,[],[f366,f219])).
fof(f915,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,sK8) != s(bool,X1)) ) | $spl2),
  inference(resolution,[],[f366,f217])).
fof(f903,plain,(
  sP4('T_0',sK8) | $spl32),
  inference(cnf_transformation,[],[f903_D])).
fof(f903_D,plain,(
  sP4('T_0',sK8) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f889,plain,(
  ( ! [X28] : (~p(s(bool,X28)) | sP4('T_0',sK8)) ) | $spl3),
  inference(duplicate_literal_removal,[],[f875])).
fof(f875,plain,(
  ( ! [X28] : (~p(s(bool,X28)) | sP4('T_0',sK8) | ~p(s(bool,X28))) ) | $spl3),
  inference(superposition,[],[f492,f323])).
fof(f890,plain,(
  ( ! [X27] : (~p(s(bool,X27)) | sP4('T_0',sK8)) ) | $spl3),
  inference(duplicate_literal_removal,[],[f874])).
fof(f874,plain,(
  ( ! [X27] : (~p(s(bool,X27)) | sP4('T_0',sK8) | ~p(s(bool,X27))) ) | $spl3),
  inference(superposition,[],[f492,f325])).
fof(f912,plain,(
  ( ! [X26] : (~p(s(bool,X26)) | sP4('T_0',sK8)) ) | $spl3),
  inference(subsumption_resolution,[],[f873,f318])).
fof(f873,plain,(
  ( ! [X26] : (~p(s(bool,X26)) | sP4('T_0',sK8) | s(bool,'F_0') = s(bool,X26)) ) | $spl3),
  inference(superposition,[],[f492,f358])).
fof(f857,plain,(
  ( ! [X144,X145] : (sP4('F_0',sK8) | p(s(bool,X144)) | p(s(bool,X145)) | ~sP5(X144,X145)) ) | $spl3),
  inference(resolution,[],[f492,f269])).
fof(f855,plain,(
  ( ! [X140,X141] : (sP4('F_0',sK8) | p(s(bool,X140)) | p(s(bool,X141)) | ~sP5(X140,X141)) ) | $spl3),
  inference(resolution,[],[f492,f273])).
fof(f846,plain,(
  ( ! [X139] : (sP4('T_0',sK8) | ~p(s(bool,X139))) ) | $spl3),
  inference(resolution,[],[f492,f316])).
fof(f845,plain,(
  ( ! [X138] : (sP4('T_0',sK8) | ~p(s(bool,X138))) ) | $spl3),
  inference(resolution,[],[f492,f311])).
fof(f844,plain,(
  ( ! [X137] : (sP4('T_0',sK8) | ~p(s(bool,X137))) ) | $spl3),
  inference(resolution,[],[f492,f301])).
fof(f843,plain,(
  ( ! [X136] : (sP4('T_0',sK8) | p(s(bool,X136))) ) | $spl3),
  inference(resolution,[],[f492,f300])).
fof(f842,plain,(
  ( ! [X135] : (sP4('T_0',sK8) | p(s(bool,X135))) ) | $spl3),
  inference(resolution,[],[f492,f297])).
fof(f841,plain,(
  ( ! [X134] : (sP4('T_0',sK8) | p(s(bool,X134))) ) | $spl3),
  inference(resolution,[],[f492,f294])).
fof(f840,plain,(
  ( ! [X133] : (sP4('T_0',sK8) | ~p(s(bool,X133))) ) | $spl3),
  inference(resolution,[],[f492,f292])).
fof(f838,plain,(
  ( ! [X132,X130,X131] : (sP4(X130,sK8) | p(s(bool,X131)) | ~p(s(bool,X132)) | s(bool,X130) != s(bool,X132)) ) | $spl3),
  inference(resolution,[],[f492,f360])).
fof(f837,plain,(
  ( ! [X127,X128,X129] : (sP4(X127,sK8) | ~p(s(bool,X128)) | p(s(bool,X129)) | s(bool,X128) != s(bool,X129)) ) | $spl3),
  inference(resolution,[],[f492,f360])).
fof(f793,plain,(
  ( ! [X30,X28,X31,X29] : (sP4(X28,sK8) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) ) | $spl3),
  inference(resolution,[],[f492,f223])).
fof(f792,plain,(
  ( ! [X26,X24,X27,X25] : (sP4(X24,sK8) | p(s(bool,X25)) | ~p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X24) != s(bool,X26)) ) | $spl3),
  inference(resolution,[],[f492,f222])).
fof(f895,plain,(
  ( ! [X23,X21,X20] : (sP4(X20,sK8) | ~p(s(bool,X21)) | p(s(bool,X23)) | s(bool,X21) != s(bool,X23)) ) | $spl3),
  inference(subsumption_resolution,[],[f791,f360])).
fof(f791,plain,(
  ( ! [X23,X21,X22,X20] : (sP4(X20,sK8) | ~p(s(bool,X21)) | ~p(s(bool,X22)) | p(s(bool,X23)) | s(bool,X21) != s(bool,X23)) ) | $spl3),
  inference(resolution,[],[f492,f222])).
fof(f790,plain,(
  ( ! [X19,X17,X18,X16] : (sP4(X16,sK8) | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) ) | $spl3),
  inference(resolution,[],[f492,f221])).
fof(f789,plain,(
  ( ! [X14,X12,X15,X13] : (sP4(X12,sK8) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X12) != s(bool,X15) | s(bool,X13) != s(bool,X14)) ) | $spl3),
  inference(resolution,[],[f492,f221])).
fof(f788,plain,(
  ( ! [X10,X8,X11,X9] : (sP4(X8,sK8) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) ) | $spl3),
  inference(resolution,[],[f492,f219])).
fof(f787,plain,(
  ( ! [X6,X4,X7,X5] : (sP4(X4,sK8) | ~p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X4) != s(bool,X7) | s(bool,X5) != s(bool,X6)) ) | $spl3),
  inference(resolution,[],[f492,f217])).
fof(f786,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(X0,sK8) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl3),
  inference(resolution,[],[f492,f217])).
fof(f737,plain,(
  ( ! [X76,X77] : (p(s(bool,X76)) | p(s(bool,X77)) | s(bool,X77) != s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) | $spl4),
  inference(resolution,[],[f369,f360])).
fof(f695,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7))) | s(bool,X22) != s(bool,X23)) ) | $spl4),
  inference(resolution,[],[f369,f223])).
fof(f694,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,X19) | s(bool,X20) != s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) | $spl4),
  inference(resolution,[],[f369,f223])).
fof(f692,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) | $spl4),
  inference(resolution,[],[f369,f222])).
fof(f691,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) | $spl4),
  inference(resolution,[],[f369,f221])).
fof(f690,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7))) | s(bool,X7) != s(bool,X8)) ) | $spl4),
  inference(resolution,[],[f369,f219])).
fof(f689,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X4) | s(bool,X5) != s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) | $spl4),
  inference(resolution,[],[f369,f219])).
fof(f688,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) | $spl4),
  inference(resolution,[],[f369,f217])).
fof(f660,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f615,f459])).
fof(f615,plain,(
  ~p(s(bool,'T_0')) | $spl18),
  inference(resolution,[],[f459,f336])).
fof(f659,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f614,f459])).
fof(f614,plain,(
  ~p(s(bool,'T_0')) | $spl18),
  inference(resolution,[],[f459,f334])).
fof(f658,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f608,f459])).
fof(f608,plain,(
  ( ! [X127] : (~p(s(bool,X127))) ) | $spl18),
  inference(resolution,[],[f459,f289])).
fof(f613,plain,(
  $false | $spl18),
  inference(resolution,[],[f459,f332])).
fof(f616,plain,(
  $false | ($spl3 | $spl18)),
  inference(resolution,[],[f459,f367])).
fof(f520,plain,(
  ( ! [X78,X79] : (p(s(bool,X78)) | ~p(s(bool,X79)) | s(bool,sK8) != s(bool,X79)) ) | $spl3),
  inference(resolution,[],[f367,f360])).
fof(f475,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,sK8) != s(bool,X22) | s(bool,X21) != s(bool,X23)) ) | $spl3),
  inference(resolution,[],[f367,f223])).
fof(f474,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,sK8) != s(bool,X19)) ) | $spl3),
  inference(resolution,[],[f367,f222])).
fof(f472,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,sK8) != s(bool,X13)) ) | $spl3),
  inference(resolution,[],[f367,f221])).
fof(f471,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | p(s(bool,X10)) | s(bool,sK8) != s(bool,X11) | s(bool,X9) != s(bool,X10)) ) | $spl3),
  inference(resolution,[],[f367,f221])).
fof(f470,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,sK8) != s(bool,X7) | s(bool,X6) != s(bool,X8)) ) | $spl3),
  inference(resolution,[],[f367,f219])).
fof(f469,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,sK8) != s(bool,X5) | s(bool,X3) != s(bool,X4)) ) | $spl3),
  inference(resolution,[],[f367,f217])).
fof(f468,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,sK8) != s(bool,X1)) ) | $spl3),
  inference(resolution,[],[f367,f217])).
fof(f422,plain,(
  ( ! [X78,X79] : (p(s(bool,X78)) | ~p(s(bool,X79)) | s(bool,X79) != s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) | $spl5),
  inference(resolution,[],[f368,f360])).
fof(f377,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7))) | s(bool,X21) != s(bool,X23)) ) | $spl5),
  inference(resolution,[],[f368,f223])).
fof(f376,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) | $spl5),
  inference(resolution,[],[f368,f222])).
fof(f374,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) | $spl5),
  inference(resolution,[],[f368,f221])).
fof(f373,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X11) != s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7))) | s(bool,X9) != s(bool,X10)) ) | $spl5),
  inference(resolution,[],[f368,f221])).
fof(f372,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7))) | s(bool,X6) != s(bool,X8)) ) | $spl5),
  inference(resolution,[],[f368,f219])).
fof(f371,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7))) | s(bool,X3) != s(bool,X4)) ) | $spl5),
  inference(resolution,[],[f368,f217])).
fof(f370,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) ) | $spl5),
  inference(resolution,[],[f368,f217])).
fof(f364,plain,(
  s('type.option.option'(sK6),'const.option.NONE_0') = s('type.option.option'(sK6),'const.bool.COND_3'(s(bool,sK8),s('type.option.option'(sK6),'const.option.NONE_0'),s('type.option.option'(sK6),sK7))) | $spl0),
  inference(cnf_transformation,[],[f364_D])).
fof(f364_D,plain,(
  s('type.option.option'(sK6),'const.option.NONE_0') = s('type.option.option'(sK6),'const.bool.COND_3'(s(bool,sK8),s('type.option.option'(sK6),'const.option.NONE_0'),s('type.option.option'(sK6),sK7))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f213,plain,(
  p(s(bool,sK8)) | ~p(s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) | s('type.option.option'(sK6),'const.option.NONE_0') = s('type.option.option'(sK6),'const.bool.COND_3'(s(bool,sK8),s('type.option.option'(sK6),'const.option.NONE_0'),s('type.option.option'(sK6),sK7)))),
  inference(cnf_transformation,[],[f148])).
fof(f148,plain,(
  (s('type.option.option'(sK6),'const.option.NONE_0') = s('type.option.option'(sK6),'const.bool.COND_3'(s(bool,sK8),s('type.option.option'(sK6),'const.option.NONE_0'),s('type.option.option'(sK6),sK7))) | ~p(s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) | p(s(bool,sK8))) & (s('type.option.option'(sK6),'const.option.NONE_0') != s('type.option.option'(sK6),'const.bool.COND_3'(s(bool,sK8),s('type.option.option'(sK6),'const.option.NONE_0'),s('type.option.option'(sK6),sK7))) | (p(s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) & ~p(s(bool,sK8))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK6,sK7,sK8])],[f147])).
fof(f147,plain,(
  ? [X0,X1,X2] : ((s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),'const.bool.COND_3'(s(bool,X2),s('type.option.option'(X0),'const.option.NONE_0'),s('type.option.option'(X0),X1))) | ~p(s(bool,'const.option.IS_SOME_1'(s('type.option.option'(X0),X1)))) | p(s(bool,X2))) & (s('type.option.option'(X0),'const.option.NONE_0') != s('type.option.option'(X0),'const.bool.COND_3'(s(bool,X2),s('type.option.option'(X0),'const.option.NONE_0'),s('type.option.option'(X0),X1))) | (p(s(bool,'const.option.IS_SOME_1'(s('type.option.option'(X0),X1)))) & ~p(s(bool,X2)))))),
  inference(flattening,[],[f146])).
fof(f146,plain,(
  ? [X0,X1,X2] : ((s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),'const.bool.COND_3'(s(bool,X2),s('type.option.option'(X0),'const.option.NONE_0'),s('type.option.option'(X0),X1))) | (~p(s(bool,'const.option.IS_SOME_1'(s('type.option.option'(X0),X1)))) | p(s(bool,X2)))) & (s('type.option.option'(X0),'const.option.NONE_0') != s('type.option.option'(X0),'const.bool.COND_3'(s(bool,X2),s('type.option.option'(X0),'const.option.NONE_0'),s('type.option.option'(X0),X1))) | (p(s(bool,'const.option.IS_SOME_1'(s('type.option.option'(X0),X1)))) & ~p(s(bool,X2)))))),
  inference(nnf_transformation,[],[f116])).
fof(f116,plain,(
  ? [X0,X1,X2] : (s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),'const.bool.COND_3'(s(bool,X2),s('type.option.option'(X0),'const.option.NONE_0'),s('type.option.option'(X0),X1))) <~> (~p(s(bool,'const.option.IS_SOME_1'(s('type.option.option'(X0),X1)))) | p(s(bool,X2))))),
  inference(ennf_transformation,[],[f51])).
fof(f51,plain,(
  ~! [X0,X1,X2] : (s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),'const.bool.COND_3'(s(bool,X2),s('type.option.option'(X0),'const.option.NONE_0'),s('type.option.option'(X0),X1))) <=> (p(s(bool,'const.option.IS_SOME_1'(s('type.option.option'(X0),X1)))) => p(s(bool,X2))))),
  inference(flattening,[],[f50])).
fof(f50,plain,(
  ~! [X0] : ! [X1,X2] : (s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),'const.bool.COND_3'(s(bool,X2),s('type.option.option'(X0),'const.option.NONE_0'),s('type.option.option'(X0),X1))) <=> (p(s(bool,'const.option.IS_SOME_1'(s('type.option.option'(X0),X1)))) => p(s(bool,X2))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : (s('type.option.option'(X6),'const.option.NONE_0') = s('type.option.option'(X6),'const.bool.COND_3'(s(bool,X8),s('type.option.option'(X6),'const.option.NONE_0'),s('type.option.option'(X6),X7))) <=> (p(s(bool,'const.option.IS_SOME_1'(s('type.option.option'(X6),X7)))) => p(s(bool,X8))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : (s('type.option.option'(X6),'const.option.NONE_0') = s('type.option.option'(X6),'const.bool.COND_3'(s(bool,X8),s('type.option.option'(X6),'const.option.NONE_0'),s('type.option.option'(X6),X7))) <=> (p(s(bool,'const.option.IS_SOME_1'(s('type.option.option'(X6),X7)))) => p(s(bool,X8))))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|1','thm.option.IF_NONE_EQUALS_OPTION|split|1')).
fof(f365,plain,(
  s('type.option.option'(sK6),'const.option.NONE_0') != s('type.option.option'(sK6),'const.bool.COND_3'(s(bool,sK8),s('type.option.option'(sK6),'const.option.NONE_0'),s('type.option.option'(sK6),sK7))) | $spl1),
  inference(cnf_transformation,[],[f365_D])).
fof(f365_D,plain,(
  s('type.option.option'(sK6),'const.option.NONE_0') != s('type.option.option'(sK6),'const.bool.COND_3'(s(bool,sK8),s('type.option.option'(sK6),'const.option.NONE_0'),s('type.option.option'(sK6),sK7))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f369,plain,(
  p(s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) | $spl4),
  inference(cnf_transformation,[],[f369_D])).
fof(f369_D,plain,(
  p(s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f214,plain,(
  p(s(bool,'const.option.IS_SOME_1'(s('type.option.option'(sK6),sK7)))) | s('type.option.option'(sK6),'const.option.NONE_0') != s('type.option.option'(sK6),'const.bool.COND_3'(s(bool,sK8),s('type.option.option'(sK6),'const.option.NONE_0'),s('type.option.option'(sK6),sK7)))),
  inference(cnf_transformation,[],[f148])).
fof(f215,plain,(
  ~p(s(bool,sK8)) | s('type.option.option'(sK6),'const.option.NONE_0') != s('type.option.option'(sK6),'const.bool.COND_3'(s(bool,sK8),s('type.option.option'(sK6),'const.option.NONE_0'),s('type.option.option'(sK6),sK7)))),
  inference(cnf_transformation,[],[f148])).
% SZS output end Proof for thm.option
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 670
% Passive clauses: 6316
% Generated clauses: 28362
% Final active clauses: 315
% Final passive clauses: 1657
% Input formulas: 48
% Initial clauses: 147
% 
% Selected by SInE selection: 46
% SInE iterations: 1
% 
% Duplicate literals: 877
% Trivial inequalities: 76
% Fw subsumption resolutions: 3448
% Fw demodulations: 1586
% Bw demodulations: 9
% 
% Simple tautologies: 603
% Equational tautologies: 7
% Forward subsumptions: 14526
% 
% Binary resolution: 16249
% Factoring: 147
% Forward superposition: 4241
% Backward superposition: 1260
% Self superposition: 2
% Equality resolution: 8
% 
% Split clauses: 1072
% Split components: 2217
% SAT solver clauses: 1170
% SAT solver unit clauses: 31
% SAT solver binary clauses: 49
% 
% Sat splits: 1072
% Sat splitting refutations: 88
% 
% Memory used [KB]: 8059
% Time elapsed: 0.448 s
% ------------------------------
% ------------------------------
% Success in time 12.141 s
