Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.arithmetic
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.arithmetic
% SZS output start Proof for thm.arithmetic
fof(f13206,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f13201,f227,f340,f342,f204,f343])).
fof(f343,plain,(
  ~p(s(bool,'F_0')) | $spl3),
  inference(cnf_transformation,[],[f343_D])).
fof(f343_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f204,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f41])).
fof(f41,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.arithmetic.STRICTLY_INCREASING_UNBOUNDED',aHL_FALSITY)).
fof(f342,plain,(
  p(s(bool,'F_0')) | $spl2),
  inference(cnf_transformation,[],[f342_D])).
fof(f342_D,plain,(
  p(s(bool,'F_0')) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f340,plain,(
  p(s(bool,'T_0')) | $spl0),
  inference(cnf_transformation,[],[f340_D])).
fof(f340_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f227,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f144])).
fof(f144,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f99])).
fof(f99,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f56])).
fof(f56,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X14] : ((p(s(bool,'F_0')) => p(s(bool,X14))) <=> p(s(bool,'T_0')))),
  file('thm.arithmetic.STRICTLY_INCREASING_UNBOUNDED','thm.bool.IMP_CLAUSES|split|2')).
fof(f13201,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f13200,f353])).
fof(f353,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl0),
  inference(resolution,[],[f217,f340])).
fof(f217,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f138])).
fof(f138,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f50])).
fof(f50,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X14] : (s(bool,'F_0') = s(bool,X14) <=> ~p(s(bool,X14)))),
  file('thm.arithmetic.STRICTLY_INCREASING_UNBOUNDED','thm.bool.EQ_CLAUSES|split|2')).
fof(f13200,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl3),
  inference(forward_demodulation,[],[f13199,f4153])).
fof(f4153,plain,(
  s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.num.num','type.num.num'),sK10)))),
  inference(resolution,[],[f4127,f214])).
fof(f214,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f136])).
fof(f136,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f48])).
fof(f48,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X14] : (s(bool,'T_0') = s(bool,X14) <=> p(s(bool,X14)))),
  file('thm.arithmetic.STRICTLY_INCREASING_UNBOUNDED','thm.bool.EQ_CLAUSES|split|0')).
fof(f4127,plain,(
  p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.num.num','type.num.num'),sK10))))),
  inference(resolution,[],[f212,f196])).
fof(f196,plain,(
  ( ! [X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK10),s('type.num.num',X3))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK10),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))))))))) )),
  inference(cnf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X3] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK10),s('type.num.num',X3))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK10),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))))) & ! [X2] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK11),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK10),s('type.num.num',X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK10,sK11])],[f130])).
fof(f130,plain,(
  ? [X0] : (! [X3] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X3))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))))) & ? [X1] : ! [X2] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X2)))))))),
  inference(rectify,[],[f90])).
fof(f90,plain,(
  ? [X0] : (! [X1] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X1))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))))) & ? [X2] : ! [X3] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X3)))))))),
  inference(ennf_transformation,[],[f38])).
fof(f38,plain,(
  ~! [X0] : (! [X1] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X1))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))))) => ! [X2] : ? [X3] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X3)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : (! [X7] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X6),s('type.num.num',X7))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X6),s('type.num.num','const.num.SUC_1'(s('type.num.num',X7)))))))) => ! [X8] : ? [X7] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X8),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X6),s('type.num.num',X7)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : (! [X7] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X6),s('type.num.num',X7))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X6),s('type.num.num','const.num.SUC_1'(s('type.num.num',X7)))))))) => ! [X8] : ? [X7] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X8),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X6),s('type.num.num',X7)))))))),
  file('thm.arithmetic.STRICTLY_INCREASING_UNBOUNDED','thm.arithmetic.STRICTLY_INCREASING_UNBOUNDED')).
fof(f212,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',sK13(X0)))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK13(X0))))))))) | p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.num.num','type.num.num'),X0))))) )),
  inference(cnf_transformation,[],[f135])).
fof(f135,plain,(
  ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',sK13(X0)))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK13(X0))))))))) | p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.num.num','type.num.num'),X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13])],[f96])).
fof(f96,plain,(
  ! [X0] : (? [X1] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X1))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))))) | p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.num.num','type.num.num'),X0)))))),
  inference(ennf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0] : (! [X1] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X1))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))))) => p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.num.num','type.num.num'),X0)))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6] : (! [X7] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X6),s('type.num.num',X7))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X6),s('type.num.num','const.num.SUC_1'(s('type.num.num',X7)))))))) => p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.num.num','type.num.num'),X6)))))),
  file('thm.arithmetic.STRICTLY_INCREASING_UNBOUNDED','thm.arithmetic.STRICTLY_INCREASING_ONE_ONE')).
fof(f13199,plain,(
  s(bool,'F_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.num.num','type.num.num'),sK10))) | $spl3),
  inference(subsumption_resolution,[],[f13193,f343])).
fof(f13193,plain,(
  p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.num.num','type.num.num'),sK10)))),
  inference(superposition,[],[f3827,f2776])).
fof(f2776,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK11),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK10),s('type.num.num',X0)))))) )),
  inference(resolution,[],[f197,f218])).
fof(f218,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f138])).
fof(f197,plain,(
  ( ! [X2] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK11),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK10),s('type.num.num',X2))))))) )),
  inference(cnf_transformation,[],[f131])).
fof(f3827,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X1),s('type.num.num',sK12(X0,X1))))))) | s(bool,'F_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.num.num','type.num.num'),X1)))) )),
  inference(resolution,[],[f207,f218])).
fof(f207,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.num.num','type.num.num'),X0)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',sK12(X1,X0)))))))) )),
  inference(cnf_transformation,[],[f134])).
fof(f134,plain,(
  ! [X0] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.num.num','type.num.num'),X0)))) | ! [X1] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',sK12(X1,X0))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12])],[f93])).
fof(f93,plain,(
  ! [X0] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.num.num','type.num.num'),X0)))) | ! [X1] : ? [X2] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X2)))))))),
  inference(ennf_transformation,[],[f44])).
fof(f44,plain,(
  ! [X0] : (p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.num.num','type.num.num'),X0)))) => ! [X1] : ? [X2] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X2)))))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : (p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.num.num','type.num.num'),X6)))) => ! [X8] : ? [X7] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X8),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X6),s('type.num.num',X7)))))))),
  file('thm.arithmetic.STRICTLY_INCREASING_UNBOUNDED','thm.arithmetic.ONE_ONE_UNBOUNDED')).
% SZS output end Proof for thm.arithmetic
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 544
% Passive clauses: 6103
% Generated clauses: 13001
% Final active clauses: 508
% Final passive clauses: 5112
% Input formulas: 36
% Initial clauses: 141
% 
% Duplicate literals: 256
% Trivial inequalities: 80
% Fw subsumption resolutions: 1209
% Fw demodulations: 134
% Bw demodulations: 35
% 
% Simple tautologies: 179
% Equational tautologies: 60
% Forward subsumptions: 4911
% 
% Binary resolution: 7231
% Factoring: 974
% Forward superposition: 1185
% Backward superposition: 1559
% Self superposition: 24
% Equality factoring: 26
% Equality resolution: 24
% 
% Split clauses: 237
% Split components: 475
% SAT solver clauses: 481
% SAT solver unit clauses: 20
% SAT solver binary clauses: 104
% SAT solver learnt clauses: 1
% 
% Sat splits: 237
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 475
% TWLsolver calls for satisfiability: 296
% 
% Memory used [KB]: 5756
% Time elapsed: 0.326 s
% ------------------------------
% ------------------------------
% Success in time 0.329 s
