Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.arithmetic
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.arithmetic
% SZS output start Proof for thm.arithmetic
fof(f139,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f136,f115,f48,f99,f49,f75,f50,f79,f77,f124,f51,f80,f78,f76,f74,f83])).
fof(f83,plain,(
  p(s(bool,'T_0')) | $spl8),
  inference(cnf_transformation,[],[f83_D])).
fof(f83_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f74,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f71])).
fof(f71,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f46])).
fof(f46,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f32])).
fof(f32,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f31])).
fof(f31,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X9] : ! [X6] : (s(X9,X6) = s(X9,X6) <=> p(s(bool,'T_0')))),
  file('thm.arithmetic.arithmetic_2|48|9_prove_ltb','thm.bool.REFL_CLAUSE')).
fof(f76,plain,(
  s(bool,sK0) != s(bool,sK1) | $spl1),
  inference(cnf_transformation,[],[f76_D])).
fof(f76_D,plain,(
  s(bool,sK0) != s(bool,sK1) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f78,plain,(
  p(s(bool,sK0)) | $spl2),
  inference(cnf_transformation,[],[f78_D])).
fof(f78_D,plain,(
  p(s(bool,sK0)) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f80,plain,(
  p(s(bool,sK1)) | $spl4),
  inference(cnf_transformation,[],[f80_D])).
fof(f80_D,plain,(
  p(s(bool,sK1)) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f51,plain,(
  p(s(bool,sK1)) | p(s(bool,sK0)) | s(bool,sK0) != s(bool,sK1)),
  inference(cnf_transformation,[],[f38])).
fof(f38,plain,(
  (s(bool,sK0) = s(bool,sK1) | ((p(s(bool,sK0)) | ~p(s(bool,sK1))) & (p(s(bool,sK1)) | ~p(s(bool,sK0))))) & (s(bool,sK0) != s(bool,sK1) | ((~p(s(bool,sK0)) | ~p(s(bool,sK1))) & (p(s(bool,sK0)) | p(s(bool,sK1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f37])).
fof(f37,plain,(
  ? [X0,X1] : ((s(bool,X0) = s(bool,X1) | ((p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (s(bool,X0) != s(bool,X1) | ((~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X1))))))),
  inference(nnf_transformation,[],[f35])).
fof(f35,plain,(
  ? [X0,X1] : (s(bool,X0) = s(bool,X1) <~> (~p(s(bool,X0)) <=> ~p(s(bool,X1))))),
  inference(ennf_transformation,[],[f18])).
fof(f18,plain,(
  ~! [X0,X1] : (s(bool,X0) = s(bool,X1) <=> (~p(s(bool,X0)) <=> ~p(s(bool,X1))))),
  inference(flattening,[],[f17])).
fof(f17,plain,(
  ~! [X0,X1] : (s(bool,X0) = s(bool,X1) <=> (~p(s(bool,X0)) <=> ~p(s(bool,X1))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : (s(bool,X6) = s(bool,X7) <=> (~p(s(bool,X6)) <=> ~p(s(bool,X7))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : (s(bool,X6) = s(bool,X7) <=> (~p(s(bool,X6)) <=> ~p(s(bool,X7))))),
  file('thm.arithmetic.arithmetic_2|48|9_prove_ltb','thm.arithmetic.arithmetic_2|48|9_prove_ltb')).
fof(f124,plain,(
  $false | ($spl0 | $spl2 | $spl5 | $spl8)),
  inference(subsumption_resolution,[],[f118,f83])).
fof(f118,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f116,f79])).
fof(f116,plain,(
  s(bool,'T_0') = s(bool,sK1) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f75,f101])).
fof(f101,plain,(
  s(bool,'T_0') = s(bool,sK0) | $spl2),
  inference(resolution,[],[f78,f62])).
fof(f62,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f41])).
fof(f41,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f23])).
fof(f23,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X8] : (s(bool,'T_0') = s(bool,X8) <=> p(s(bool,X8)))),
  file('thm.arithmetic.arithmetic_2|48|9_prove_ltb','thm.bool.EQ_CLAUSES|split|0')).
fof(f77,plain,(
  ~p(s(bool,sK0)) | $spl3),
  inference(cnf_transformation,[],[f77_D])).
fof(f77_D,plain,(
  ~p(s(bool,sK0)) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f79,plain,(
  ~p(s(bool,sK1)) | $spl5),
  inference(cnf_transformation,[],[f79_D])).
fof(f79_D,plain,(
  ~p(s(bool,sK1)) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f50,plain,(
  ~p(s(bool,sK1)) | ~p(s(bool,sK0)) | s(bool,sK0) != s(bool,sK1)),
  inference(cnf_transformation,[],[f38])).
fof(f75,plain,(
  s(bool,sK0) = s(bool,sK1) | $spl0),
  inference(cnf_transformation,[],[f75_D])).
fof(f75_D,plain,(
  s(bool,sK0) = s(bool,sK1) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f49,plain,(
  ~p(s(bool,sK0)) | p(s(bool,sK1)) | s(bool,sK0) = s(bool,sK1)),
  inference(cnf_transformation,[],[f38])).
fof(f99,plain,(
  $false | ($spl1 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f92,f93])).
fof(f93,plain,(
  s(bool,'F_0') != s(bool,sK1) | ($spl1 | $spl3)),
  inference(backward_demodulation,[],[f91,f76])).
fof(f91,plain,(
  s(bool,'F_0') = s(bool,sK0) | $spl3),
  inference(resolution,[],[f66,f77])).
fof(f66,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f43])).
fof(f43,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f26])).
fof(f26,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f25])).
fof(f25,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X8] : (s(bool,'F_0') = s(bool,X8) <=> ~p(s(bool,X8)))),
  file('thm.arithmetic.arithmetic_2|48|9_prove_ltb','thm.bool.EQ_CLAUSES|split|3')).
fof(f92,plain,(
  s(bool,'F_0') = s(bool,sK1) | $spl5),
  inference(resolution,[],[f66,f79])).
fof(f48,plain,(
  ~p(s(bool,sK1)) | p(s(bool,sK0)) | s(bool,sK0) = s(bool,sK1)),
  inference(cnf_transformation,[],[f38])).
fof(f115,plain,(
  $false | ($spl1 | $spl2 | $spl4)),
  inference(subsumption_resolution,[],[f103,f105])).
fof(f105,plain,(
  s(bool,'T_0') != s(bool,sK1) | ($spl1 | $spl2)),
  inference(backward_demodulation,[],[f101,f76])).
fof(f103,plain,(
  s(bool,'T_0') = s(bool,sK1) | $spl4),
  inference(resolution,[],[f80,f62])).
fof(f136,plain,(
  $false | ($spl0 | $spl3 | $spl4)),
  inference(subsumption_resolution,[],[f131,f77])).
fof(f131,plain,(
  p(s(bool,sK0)) | ($spl0 | $spl4)),
  inference(backward_demodulation,[],[f75,f80])).
% SZS output end Proof for thm.arithmetic
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 24
% Passive clauses: 41
% Generated clauses: 94
% Final active clauses: 9
% Final passive clauses: 10
% Input formulas: 15
% Initial clauses: 26
% 
% Trivial inequalities: 3
% Fw subsumption resolutions: 7
% Fw demodulations: 6
% Bw demodulations: 7
% 
% Simple tautologies: 3
% Equational tautologies: 3
% Forward subsumptions: 21
% 
% Binary resolution: 15
% Backward superposition: 16
% Equality resolution: 1
% 
% Split clauses: 8
% Split components: 20
% SAT solver clauses: 20
% SAT solver unit clauses: 6
% SAT solver binary clauses: 5
% SAT solver learnt clauses: 4
% SAT solver learnt literals: 4
% 
% Sat splits: 8
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 16
% TWLsolver calls for satisfiability: 6
% 
% Memory used [KB]: 639
% Time elapsed: 0.027 s
% ------------------------------
% ------------------------------
% Success in time 0.03 s
