Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.nets
% (113224)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 174
% Passive clauses: 2162
% Generated clauses: 23716
% Final active clauses: 150
% Final passive clauses: 1789
% Input formulas: 21
% Initial clauses: 57
% 
% Duplicate literals: 1303
% Trivial inequalities: 273
% Fw subsumption resolutions: 3596
% Fw demodulations: 3398
% Bw demodulations: 21
% 
% Simple tautologies: 496
% Equational tautologies: 47
% Forward subsumptions: 9264
% 
% Binary resolution: 436
% Forward superposition: 5998
% Backward superposition: 7958
% Self superposition: 204
% Equality factoring: 151
% Equality resolution: 41
% 
% Split clauses: 1288
% Split components: 2659
% SAT solver clauses: 2027
% SAT solver unit clauses: 25
% SAT solver binary clauses: 305
% 
% Sat splits: 1288
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 2024
% TWLsolver calls for satisfiability: 194
% 
% Memory used [KB]: 12153
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 294 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.nets
% (113242)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 125
% Passive clauses: 1073
% Generated clauses: 10223
% Final active clauses: 112
% Final passive clauses: 928
% Input formulas: 21
% Initial clauses: 57
% 
% Duplicate literals: 703
% Trivial inequalities: 74
% Fw subsumption resolutions: 822
% Fw demodulations: 1625
% Bw demodulations: 71
% Condensations: 425
% Global subsumptions: 8
% 
% Simple tautologies: 166
% Equational tautologies: 27
% Forward subsumptions: 3711
% 
% Binary resolution: 326
% Factoring: 4
% Forward superposition: 2548
% Backward superposition: 2435
% Self superposition: 50
% Equality factoring: 50
% Equality resolution: 12
% 
% Split clauses: 138
% Split components: 281
% SAT solver clauses: 1671
% SAT solver unit clauses: 73
% SAT solver binary clauses: 32
% SAT solver learnt clauses: 16
% 
% Sat splits: 138
% 
% TWLsolver clauses: 1639
% TWLsolver calls for satisfiability: 374
% 
% Memory used [KB]: 12665
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.nets
% (113255)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 787
% Passive clauses: 4691
% Generated clauses: 164621
% Final active clauses: 241
% Final passive clauses: 1197
% Input formulas: 21
% Initial clauses: 57
% 
% Duplicate literals: 7443
% Trivial inequalities: 1398
% Fw subsumption resolutions: 19877
% Fw demodulations: 27323
% Bw demodulations: 2275
% Condensations: 13643
% Global subsumptions: 69
% 
% Simple tautologies: 1730
% Equational tautologies: 1057
% Forward subsumptions: 86846
% Fw demodulations to eq. taut.: 59
% Bw demodulations to eq. taut.: 345
% 
% Unit resulting resolution: 5551
% Forward superposition: 49584
% Backward superposition: 33077
% Self superposition: 679
% Equality factoring: 280
% Equality resolution: 127
% 
% Split clauses: 314
% Split components: 637
% SAT solver clauses: 5637
% SAT solver unit clauses: 76
% SAT solver binary clauses: 491
% 
% Sat splits: 314
% Sat splitting refutations: 247
% 
% Memory used [KB]: 60254
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 183 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.nets
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.nets
% SZS output start Proof for thm.nets
fof(f10463,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f1790,f1794,f1792,f1859,f1920,f1860,f1924,f1861,f1926,f1862,f1928,f1890,f1922,f1945,f1993,f2025,f2023,f2021,f2488,f2492,f2490,f2973,f2977,f2975,f3232,f3236,f3234,f5279,f2022,f6695,f6732,f6697,f6735,f6734,f6737,f6699,f6700,f6701,f6702,f6746,f6804,f6876,f6874,f6809,f5281,f6810,f6880,f6811,f6882,f6812,f6884,f6813,f6886,f6814,f6888,f6818,f6892,f6819,f6894,f6820,f6896,f6847,f6731,f6910,f6945,f6743,f6946,f6748,f6947,f6745,f6752,f6948,f6986,f6750,f6754,f7212,f7216,f7217,f7218,f7219,f7221,f7222,f7225,f7226,f7227,f7254,f6742,f7382,f7283,f7383,f7284,f7384,f7286,f7385,f7288,f7387,f7292,f7388,f7294,f7391,f7296,f7392,f7298,f7393,f7300,f7396,f6873,f6875,f7423,f6747,f7314,f7873,f7874,f7903,f7904,f8005,f7893,f8006,f7895,f7897,f8093,f8094,f8223,f8311,f7923,f8312,f7925,f7927,f8490,f8113,f8491,f8115,f8117,f8702,f8244,f8242,f8983,f8985,f8989,f9337,f9338,f9367,f9368,f9397,f9424,f10057,f9357,f10058,f9359,f9361,f10155,f9387,f10156,f9389,f9391,f10277,f9418,f9416,f10373,f9445,f9443,f10448,f10450,f10456,f10461,f10462])).
fof(f10462,plain,(
  $false | $spl109),
  inference(subsumption_resolution,[],[f7923,f8401])).
fof(f8401,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK7(sK5,sK4,sK1,sK0)))))),
  inference(resolution,[],[f8386,f136])).
fof(f136,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f42])).
fof(f42,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X12] : (s(bool,'T_0') = s(bool,X12) <=> p(s(bool,X12)))),
  file('thm.nets.DORDER_LEMMA','thm.bool.EQ_CLAUSES|split|1')).
fof(f8386,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK7(sK5,sK4,sK1,sK0))))))),
  inference(subsumption_resolution,[],[f8366,f140])).
fof(f140,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  p(s(bool,'T_0'))),
  file('thm.nets.DORDER_LEMMA','thm.bool.TRUTH')).
fof(f8366,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK7(sK5,sK4,sK1,sK0))))))),
  inference(superposition,[],[f4485,f185])).
fof(f185,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4))),s(sK0,sK4)))),
  inference(resolution,[],[f88,f136])).
fof(f88,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4))),s(sK0,sK4))))),
  inference(cnf_transformation,[],[f58])).
fof(f58,plain,(
  p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),sK1)))) & ((p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4))),s(sK0,sK4)))) & ! [X9] : (~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X9))),s(sK0,sK4)))) | p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,X9)))))) & (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5))),s(sK0,sK5)))) & ! [X7] : (~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X7))),s(sK0,sK5)))) | p(s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,X7)))))) & ! [X4] : (~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X4))),s(sK0,X4)))) | (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(X4)))),s(sK0,X4)))) & (~p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK6(X4))))) | ~p(s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(X4)))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5,sK6])],[f57])).
fof(f57,plain,(
  ? [X0,X1] : (p(s(bool,'const.nets.dorder_1'(s(fun(X0,fun(X0,bool)),X1)))) & ? [X2,X3] : (? [X8] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X8))),s(X0,X8)))) & ! [X9] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X9))),s(X0,X8)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X9)))))) & ? [X6] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X6))),s(X0,X6)))) & ! [X7] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X7))),s(X0,X6)))) | p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X7)))))) & ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X4)))) | ? [X5] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X4)))) & (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X5)))))))))),
  inference(rectify,[],[f48])).
fof(f48,plain,(
  ? [X0,X1] : (p(s(bool,'const.nets.dorder_1'(s(fun(X0,fun(X0,bool)),X1)))) & ? [X2,X3] : (? [X6] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X6))),s(X0,X6)))) & ! [X7] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X7))),s(X0,X6)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X7)))))) & ? [X4] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X4)))) & ! [X5] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X5)))))) & ! [X8] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X8))),s(X0,X8)))) | ? [X9] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X9))),s(X0,X8)))) & (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X9)))) | ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X9)))))))))),
  inference(flattening,[],[f47])).
fof(f47,plain,(
  ? [X0,X1] : (p(s(bool,'const.nets.dorder_1'(s(fun(X0,fun(X0,bool)),X1)))) & ? [X2,X3] : ((? [X6] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X6))),s(X0,X6)))) & ! [X7] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X7))),s(X0,X6)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X7)))))) & ? [X4] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X4)))) & ! [X5] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X5))))))) & ! [X8] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X8))),s(X0,X8)))) | ? [X9] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X9))),s(X0,X8)))) & (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X9)))) | ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X9)))))))))),
  inference(ennf_transformation,[],[f24])).
fof(f24,plain,(
  ~! [X0,X1] : (p(s(bool,'const.nets.dorder_1'(s(fun(X0,fun(X0,bool)),X1)))) => ! [X2,X3] : ((? [X6] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X6))),s(X0,X6)))) & ! [X7] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X7))),s(X0,X6)))) => p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X7)))))) & ? [X4] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X4)))) & ! [X5] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X4)))) => p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X5))))))) => ? [X8] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X8))),s(X0,X8)))) & ! [X9] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X9))),s(X0,X8)))) => (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X9)))) & p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X9)))))))))),
  inference(flattening,[],[f23])).
fof(f23,plain,(
  ~! [X0] : ! [X1] : (p(s(bool,'const.nets.dorder_1'(s(fun(X0,fun(X0,bool)),X1)))) => ! [X2,X3] : ((? [X6] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X6))),s(X0,X6)))) & ! [X7] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X7))),s(X0,X6)))) => p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X7)))))) & ? [X4] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X4)))) & ! [X5] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X4)))) => p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X5))))))) => ? [X8] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X8))),s(X0,X8)))) & ! [X9] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X9))),s(X0,X8)))) => (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X9)))) & p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X9)))))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7] : (p(s(bool,'const.nets.dorder_1'(s(fun(X6,fun(X6,bool)),X7)))) => ! [X8,X9] : ((? [X10] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X10))),s(X6,X10)))) & ! [X11] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X11))),s(X6,X10)))) => p(s(bool,happ(s(fun(X6,bool),X8),s(X6,X11)))))) & ? [X10] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X10))),s(X6,X10)))) & ! [X11] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X11))),s(X6,X10)))) => p(s(bool,happ(s(fun(X6,bool),X9),s(X6,X11))))))) => ? [X10] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X10))),s(X6,X10)))) & ! [X11] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X11))),s(X6,X10)))) => (p(s(bool,happ(s(fun(X6,bool),X8),s(X6,X11)))) & p(s(bool,happ(s(fun(X6,bool),X9),s(X6,X11)))))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7] : (p(s(bool,'const.nets.dorder_1'(s(fun(X6,fun(X6,bool)),X7)))) => ! [X8,X9] : ((? [X10] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X10))),s(X6,X10)))) & ! [X11] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X11))),s(X6,X10)))) => p(s(bool,happ(s(fun(X6,bool),X8),s(X6,X11)))))) & ? [X10] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X10))),s(X6,X10)))) & ! [X11] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X11))),s(X6,X10)))) => p(s(bool,happ(s(fun(X6,bool),X9),s(X6,X11))))))) => ? [X10] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X10))),s(X6,X10)))) & ! [X11] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X11))),s(X6,X10)))) => (p(s(bool,happ(s(fun(X6,bool),X8),s(X6,X11)))) & p(s(bool,happ(s(fun(X6,bool),X9),s(X6,X11)))))))))),
  file('thm.nets.DORDER_LEMMA','thm.nets.DORDER_LEMMA')).
fof(f4485,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,X0)))) | p(s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK7(sK5,X0,sK1,sK0))))))) )),
  inference(subsumption_resolution,[],[f4454,f260])).
fof(f260,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(sK5,X0,sK1,sK0)))),s(sK0,sK7(sK5,X0,sK1,sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,X0))))) )),
  inference(subsumption_resolution,[],[f233,f87])).
fof(f87,plain,(
  p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),sK1))))),
  inference(cnf_transformation,[],[f58])).
fof(f233,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(sK5,X0,sK1,sK0)))),s(sK0,sK7(sK5,X0,sK1,sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,X0)))) | ~p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),sK1))))) )),
  inference(resolution,[],[f90,f94])).
fof(f94,plain,(
  ( ! [X6,X0,X7,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK7(X7,X6,X1,X0)))),s(X0,sK7(X7,X6,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X7))),s(X0,X7)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X6))),s(X0,X6)))) | ~p(s(bool,'const.nets.dorder_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(cnf_transformation,[],[f61])).
fof(f61,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.nets.dorder_1'(s(fun(X0,fun(X0,bool)),X1)))) | ! [X6,X7] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X6))),s(X0,X6)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X7))),s(X0,X7)))) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK7(X7,X6,X1,X0)))),s(X0,sK7(X7,X6,X1,X0))))) & ! [X9] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X9))),s(X0,sK7(X7,X6,X1,X0))))) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X9))),s(X0,X6)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X9))),s(X0,X7))))))))) & ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK8(X1,X0)))),s(X0,sK8(X1,X0))))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK9(X1,X0)))),s(X0,sK9(X1,X0))))) & ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X4)))) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK10(X4,X1,X0)))),s(X0,X4)))) & (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK10(X4,X1,X0)))),s(X0,sK8(X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK10(X4,X1,X0)))),s(X0,sK9(X1,X0))))))))) | p(s(bool,'const.nets.dorder_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK7,sK8,sK9,sK10])],[f60])).
fof(f60,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.nets.dorder_1'(s(fun(X0,fun(X0,bool)),X1)))) | ! [X6,X7] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X6))),s(X0,X6)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X7))),s(X0,X7)))) | ? [X8] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X8))),s(X0,X8)))) & ! [X9] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X9))),s(X0,X8)))) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X9))),s(X0,X6)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X9))),s(X0,X7))))))))) & (? [X2,X3] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X2)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X3)))) & ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X4)))) | ? [X5] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X4)))) & (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X2)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X3)))))))) | p(s(bool,'const.nets.dorder_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(rectify,[],[f59])).
fof(f59,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.nets.dorder_1'(s(fun(X0,fun(X0,bool)),X1)))) | ! [X2,X3] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X2)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X3)))) | ? [X4] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X4)))) & ! [X5] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X4)))) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X2)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X3))))))))) & (? [X2,X3] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X2)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X3)))) & ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X4)))) | ? [X5] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X4)))) & (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X2)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X3)))))))) | p(s(bool,'const.nets.dorder_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(nnf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0,X1] : (p(s(bool,'const.nets.dorder_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> ! [X2,X3] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X2)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X3)))) | ? [X4] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X4)))) & ! [X5] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X4)))) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X2)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X3)))))))))),
  inference(flattening,[],[f49])).
fof(f49,plain,(
  ! [X0,X1] : (p(s(bool,'const.nets.dorder_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> ! [X2,X3] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X2)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X3))))) | ? [X4] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X4)))) & ! [X5] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X4)))) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X2)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X3)))))))))),
  inference(ennf_transformation,[],[f26])).
fof(f26,plain,(
  ! [X0,X1] : (p(s(bool,'const.nets.dorder_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> ! [X2,X3] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X2)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X3))))) => ? [X4] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X4)))) & ! [X5] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X4)))) => (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X2)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X3)))))))))),
  inference(flattening,[],[f25])).
fof(f25,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.nets.dorder_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> ! [X2,X3] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X2)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X3))))) => ? [X4] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X4)))) & ! [X5] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X4)))) => (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X2)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X3)))))))))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X6] : ! [X7] : (p(s(bool,'const.nets.dorder_1'(s(fun(X6,fun(X6,bool)),X7)))) <=> ! [X13,X14] : ((p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X13))),s(X6,X13)))) & p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X14))),s(X6,X14))))) => ? [X15] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X15))),s(X6,X15)))) & ! [X16] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X16))),s(X6,X15)))) => (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X16))),s(X6,X13)))) & p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X16))),s(X6,X14)))))))))),
  file('thm.nets.DORDER_LEMMA','thm.nets.dorder')).
fof(f90,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5))),s(sK0,sK5))))),
  inference(cnf_transformation,[],[f58])).
fof(f4454,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK7(sK5,X0,sK1,sK0)))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,X0)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(sK5,X0,sK1,sK0)))),s(sK0,sK7(sK5,X0,sK1,sK0)))))) )),
  inference(resolution,[],[f558,f92])).
fof(f92,plain,(
  ( ! [X4] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(X4)))),s(sK0,X4)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X4))),s(sK0,X4))))) )),
  inference(cnf_transformation,[],[f58])).
fof(f558,plain,(
  ( ! [X2,X3] : (~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X2))),s(sK0,sK7(sK5,X3,sK1,sK0))))) | p(s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,X2)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X3))),s(sK0,X3))))) )),
  inference(subsumption_resolution,[],[f557,f108])).
fof(f108,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f67])).
fof(f67,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f66])).
fof(f66,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f30])).
fof(f30,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X12] : ((p(s(bool,'F_0')) => p(s(bool,X12))) <=> p(s(bool,'T_0')))),
  file('thm.nets.DORDER_LEMMA','thm.bool.IMP_CLAUSES|split|2')).
fof(f557,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,X2)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X2))),s(sK0,sK7(sK5,X3,sK1,sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X3))),s(sK0,X3))))) )),
  inference(forward_demodulation,[],[f556,f155])).
fof(f155,plain,(
  s(bool,'T_0') = s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),sK1)))),
  inference(resolution,[],[f87,f136])).
fof(f556,plain,(
  ( ! [X2,X3] : (p(s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,X2)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X2))),s(sK0,sK7(sK5,X3,sK1,sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X3))),s(sK0,X3)))) | ~p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),sK1))))) )),
  inference(subsumption_resolution,[],[f531,f90])).
fof(f531,plain,(
  ( ! [X2,X3] : (p(s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,X2)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X2))),s(sK0,sK7(sK5,X3,sK1,sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5))),s(sK0,sK5)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X3))),s(sK0,X3)))) | ~p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),sK1))))) )),
  inference(resolution,[],[f91,f96])).
fof(f96,plain,(
  ( ! [X6,X0,X7,X1,X9] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X9))),s(X0,X7)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X9))),s(X0,sK7(X7,X6,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X7))),s(X0,X7)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X6))),s(X0,X6)))) | ~p(s(bool,'const.nets.dorder_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(cnf_transformation,[],[f61])).
fof(f91,plain,(
  ( ! [X7] : (~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X7))),s(sK0,sK5)))) | p(s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,X7))))) )),
  inference(cnf_transformation,[],[f58])).
fof(f10461,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f10460,f219])).
fof(f219,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(subsumption_resolution,[],[f210,f153])).
fof(f153,plain,(
  s(bool,'F_0') != s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),sK1)))),
  inference(resolution,[],[f87,f131])).
fof(f131,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f39])).
fof(f39,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f38])).
fof(f38,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X12] : (s(bool,'F_0') = s(bool,X12) <=> ~p(s(bool,X12)))),
  file('thm.nets.DORDER_LEMMA','thm.bool.EQ_CLAUSES|split|3')).
fof(f210,plain,(
  s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),sK1)))),
  inference(superposition,[],[f153,f142])).
fof(f142,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.nets.DORDER_LEMMA',aHL_BOOL_CASES)).
fof(f10460,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl112),
  inference(forward_demodulation,[],[f7927,f8401])).
fof(f10456,plain,(
  $false | $spl111),
  inference(subsumption_resolution,[],[f10455,f141])).
fof(f141,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.nets.DORDER_LEMMA',aHL_TRUTH)).
fof(f10455,plain,(
  ~p(s(bool,'T_0')) | $spl111),
  inference(forward_demodulation,[],[f10454,f155])).
fof(f10454,plain,(
  ~p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),sK1)))) | $spl111),
  inference(subsumption_resolution,[],[f10453,f129])).
fof(f129,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0] : ((~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,'T_0')) & p(s(bool,X0)))))),
  inference(flattening,[],[f80])).
fof(f80,plain,(
  ! [X0] : (((~p(s(bool,'T_0')) | ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,'T_0')) & p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f37])).
fof(f37,plain,(
  ! [X0] : ((p(s(bool,'T_0')) & p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X12] : ((p(s(bool,'T_0')) & p(s(bool,X12))) <=> p(s(bool,X12)))),
  file('thm.nets.DORDER_LEMMA','thm.bool.AND_CLAUSES|split|0')).
fof(f10453,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),sK1)))) | $spl111),
  inference(forward_demodulation,[],[f10452,f185])).
fof(f10452,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4))),s(sK0,sK4)))) | ~p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),sK1)))) | $spl111),
  inference(subsumption_resolution,[],[f10451,f108])).
fof(f10451,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4))),s(sK0,sK4)))) | ~p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),sK1)))) | $spl111),
  inference(forward_demodulation,[],[f10424,f248])).
fof(f248,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5))),s(sK0,sK5)))),
  inference(resolution,[],[f90,f136])).
fof(f10424,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5))),s(sK0,sK5)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4))),s(sK0,sK4)))) | ~p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),sK1)))) | $spl111),
  inference(resolution,[],[f7925,f94])).
fof(f10450,plain,(
  $false | $spl111),
  inference(subsumption_resolution,[],[f10449,f141])).
fof(f10449,plain,(
  ~p(s(bool,'T_0')) | $spl111),
  inference(forward_demodulation,[],[f10423,f248])).
fof(f10423,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5))),s(sK0,sK5)))) | $spl111),
  inference(resolution,[],[f7925,f198])).
fof(f198,plain,(
  ( ! [X1] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(X1,sK4,sK1,sK0)))),s(sK0,sK7(X1,sK4,sK1,sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X1))),s(sK0,X1))))) )),
  inference(subsumption_resolution,[],[f171,f87])).
fof(f171,plain,(
  ( ! [X1] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(X1,sK4,sK1,sK0)))),s(sK0,sK7(X1,sK4,sK1,sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X1))),s(sK0,X1)))) | ~p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),sK1))))) )),
  inference(resolution,[],[f88,f94])).
fof(f10448,plain,(
  $false | $spl111),
  inference(subsumption_resolution,[],[f10447,f141])).
fof(f10447,plain,(
  ~p(s(bool,'T_0')) | $spl111),
  inference(forward_demodulation,[],[f10422,f185])).
fof(f10422,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4))),s(sK0,sK4)))) | $spl111),
  inference(resolution,[],[f7925,f260])).
fof(f9443,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(sK5,sK7(sK4,sK4,sK1,sK0),sK1,sK0)))),s(sK0,sK7(sK5,sK7(sK4,sK4,sK1,sK0),sK1,sK0))))) | $spl141),
  inference(cnf_transformation,[],[f9443_D])).
fof(f9443_D,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(sK5,sK7(sK4,sK4,sK1,sK0),sK1,sK0)))),s(sK0,sK7(sK5,sK7(sK4,sK4,sK1,sK0),sK1,sK0))))) <=> ~$spl141),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl141])])).
fof(f9445,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK6(sK7(sK5,sK7(sK4,sK4,sK1,sK0),sK1,sK0)))))) | $spl143),
  inference(cnf_transformation,[],[f9445_D])).
fof(f9445_D,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK6(sK7(sK5,sK7(sK4,sK4,sK1,sK0),sK1,sK0)))))) <=> ~$spl143),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl143])])).
fof(f10373,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK6(sK7(sK5,sK7(sK4,sK4,sK1,sK0),sK1,sK0)))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(sK5,sK7(sK4,sK4,sK1,sK0),sK1,sK0)))),s(sK0,sK7(sK5,sK7(sK4,sK4,sK1,sK0),sK1,sK0))))) | $spl104),
  inference(subsumption_resolution,[],[f10318,f108])).
fof(f10318,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK6(sK7(sK5,sK7(sK4,sK4,sK1,sK0),sK1,sK0)))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(sK5,sK7(sK4,sK4,sK1,sK0),sK1,sK0)))),s(sK0,sK7(sK5,sK7(sK4,sK4,sK1,sK0),sK1,sK0))))) | $spl104),
  inference(superposition,[],[f93,f9432])).
fof(f9432,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK7(sK5,sK7(sK4,sK4,sK1,sK0),sK1,sK0))))) | $spl104),
  inference(resolution,[],[f9150,f136])).
fof(f9150,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK7(sK5,sK7(sK4,sK4,sK1,sK0),sK1,sK0)))))) | $spl104),
  inference(resolution,[],[f7894,f4485])).
fof(f7894,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(sK4,sK4,sK1,sK0)))),s(sK0,sK7(sK4,sK4,sK1,sK0))))) | $spl104),
  inference(cnf_transformation,[],[f7894_D])).
fof(f7894_D,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(sK4,sK4,sK1,sK0)))),s(sK0,sK7(sK4,sK4,sK1,sK0))))) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f93,plain,(
  ( ! [X4] : (~p(s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(X4))))) | ~p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK6(X4))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X4))),s(sK0,X4))))) )),
  inference(cnf_transformation,[],[f58])).
fof(f9416,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(sK7(sK4,sK4,sK1,sK0),sK5,sK1,sK0)))),s(sK0,sK7(sK7(sK4,sK4,sK1,sK0),sK5,sK1,sK0))))) | $spl137),
  inference(cnf_transformation,[],[f9416_D])).
fof(f9416_D,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(sK7(sK4,sK4,sK1,sK0),sK5,sK1,sK0)))),s(sK0,sK7(sK7(sK4,sK4,sK1,sK0),sK5,sK1,sK0))))) <=> ~$spl137),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl137])])).
fof(f9418,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK6(sK7(sK7(sK4,sK4,sK1,sK0),sK5,sK1,sK0)))))) | $spl139),
  inference(cnf_transformation,[],[f9418_D])).
fof(f9418_D,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK6(sK7(sK7(sK4,sK4,sK1,sK0),sK5,sK1,sK0)))))) <=> ~$spl139),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl139])])).
fof(f10277,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK6(sK7(sK7(sK4,sK4,sK1,sK0),sK5,sK1,sK0)))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(sK7(sK4,sK4,sK1,sK0),sK5,sK1,sK0)))),s(sK0,sK7(sK7(sK4,sK4,sK1,sK0),sK5,sK1,sK0))))) | $spl104),
  inference(subsumption_resolution,[],[f10222,f108])).
fof(f10222,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK6(sK7(sK7(sK4,sK4,sK1,sK0),sK5,sK1,sK0)))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(sK7(sK4,sK4,sK1,sK0),sK5,sK1,sK0)))),s(sK0,sK7(sK7(sK4,sK4,sK1,sK0),sK5,sK1,sK0))))) | $spl104),
  inference(superposition,[],[f93,f9405])).
fof(f9405,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK7(sK7(sK4,sK4,sK1,sK0),sK5,sK1,sK0))))) | $spl104),
  inference(resolution,[],[f9148,f136])).
fof(f9148,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK7(sK7(sK4,sK4,sK1,sK0),sK5,sK1,sK0)))))) | $spl104),
  inference(resolution,[],[f7894,f4305])).
fof(f4305,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,X0)))) | p(s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK7(X0,sK5,sK1,sK0))))))) )),
  inference(subsumption_resolution,[],[f4274,f261])).
fof(f261,plain,(
  ( ! [X1] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(X1,sK5,sK1,sK0)))),s(sK0,sK7(X1,sK5,sK1,sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X1))),s(sK0,X1))))) )),
  inference(subsumption_resolution,[],[f234,f87])).
fof(f234,plain,(
  ( ! [X1] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(X1,sK5,sK1,sK0)))),s(sK0,sK7(X1,sK5,sK1,sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X1))),s(sK0,X1)))) | ~p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),sK1))))) )),
  inference(resolution,[],[f90,f94])).
fof(f4274,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK7(X0,sK5,sK1,sK0)))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,X0)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(X0,sK5,sK1,sK0)))),s(sK0,sK7(X0,sK5,sK1,sK0)))))) )),
  inference(resolution,[],[f555,f92])).
fof(f555,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK7(X1,sK5,sK1,sK0))))) | p(s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,X0)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X1))),s(sK0,X1))))) )),
  inference(subsumption_resolution,[],[f554,f108])).
fof(f554,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,X0)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK7(X1,sK5,sK1,sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X1))),s(sK0,X1))))) )),
  inference(forward_demodulation,[],[f553,f155])).
fof(f553,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,X0)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK7(X1,sK5,sK1,sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X1))),s(sK0,X1)))) | ~p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),sK1))))) )),
  inference(subsumption_resolution,[],[f530,f90])).
fof(f530,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,X0)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK7(X1,sK5,sK1,sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X1))),s(sK0,X1)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5))),s(sK0,sK5)))) | ~p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),sK1))))) )),
  inference(resolution,[],[f91,f95])).
fof(f95,plain,(
  ( ! [X6,X0,X7,X1,X9] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X9))),s(X0,X6)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X9))),s(X0,sK7(X7,X6,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X7))),s(X0,X7)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X6))),s(X0,X6)))) | ~p(s(bool,'const.nets.dorder_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(cnf_transformation,[],[f61])).
fof(f9391,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK7(sK4,sK7(sK4,sK4,sK1,sK0),sK1,sK0))))) | $spl134),
  inference(cnf_transformation,[],[f9391_D])).
fof(f9391_D,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK7(sK4,sK7(sK4,sK4,sK1,sK0),sK1,sK0))))) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f9389,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(sK4,sK7(sK4,sK4,sK1,sK0),sK1,sK0)))),s(sK0,sK7(sK4,sK7(sK4,sK4,sK1,sK0),sK1,sK0))))) | $spl133),
  inference(cnf_transformation,[],[f9389_D])).
fof(f9389_D,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(sK4,sK7(sK4,sK4,sK1,sK0),sK1,sK0)))),s(sK0,sK7(sK4,sK7(sK4,sK4,sK1,sK0),sK1,sK0))))) <=> ~$spl133),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl133])])).
fof(f10156,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(sK4,sK7(sK4,sK4,sK1,sK0),sK1,sK0)))),s(sK0,sK7(sK4,sK7(sK4,sK4,sK1,sK0),sK1,sK0))))) | s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK7(sK4,sK7(sK4,sK4,sK1,sK0),sK1,sK0))))) | $spl104),
  inference(subsumption_resolution,[],[f10100,f108])).
fof(f10100,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(sK4,sK7(sK4,sK4,sK1,sK0),sK1,sK0)))),s(sK0,sK7(sK4,sK7(sK4,sK4,sK1,sK0),sK1,sK0))))) | s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK7(sK4,sK7(sK4,sK4,sK1,sK0),sK1,sK0))))) | $spl104),
  inference(superposition,[],[f775,f9376])).
fof(f9376,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK6(sK7(sK4,sK7(sK4,sK4,sK1,sK0),sK1,sK0))))) | $spl104),
  inference(resolution,[],[f9146,f136])).
fof(f9146,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK6(sK7(sK4,sK7(sK4,sK4,sK1,sK0),sK1,sK0)))))) | $spl104),
  inference(resolution,[],[f7894,f4217])).
fof(f4217,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,X0)))) | p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK6(sK7(sK4,X0,sK1,sK0))))))) )),
  inference(subsumption_resolution,[],[f4186,f197])).
fof(f197,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(sK4,X0,sK1,sK0)))),s(sK0,sK7(sK4,X0,sK1,sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,X0))))) )),
  inference(subsumption_resolution,[],[f170,f87])).
fof(f170,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(sK4,X0,sK1,sK0)))),s(sK0,sK7(sK4,X0,sK1,sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,X0)))) | ~p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),sK1))))) )),
  inference(resolution,[],[f88,f94])).
fof(f4186,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK6(sK7(sK4,X0,sK1,sK0)))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,X0)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(sK4,X0,sK1,sK0)))),s(sK0,sK7(sK4,X0,sK1,sK0)))))) )),
  inference(resolution,[],[f406,f92])).
fof(f406,plain,(
  ( ! [X2,X3] : (~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X2))),s(sK0,sK7(sK4,X3,sK1,sK0))))) | p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,X2)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X3))),s(sK0,X3))))) )),
  inference(subsumption_resolution,[],[f405,f108])).
fof(f405,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,X2)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X2))),s(sK0,sK7(sK4,X3,sK1,sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X3))),s(sK0,X3))))) )),
  inference(forward_demodulation,[],[f404,f155])).
fof(f404,plain,(
  ( ! [X2,X3] : (p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,X2)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X2))),s(sK0,sK7(sK4,X3,sK1,sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X3))),s(sK0,X3)))) | ~p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),sK1))))) )),
  inference(subsumption_resolution,[],[f379,f88])).
fof(f379,plain,(
  ( ! [X2,X3] : (p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,X2)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X2))),s(sK0,sK7(sK4,X3,sK1,sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4))),s(sK0,sK4)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X3))),s(sK0,X3)))) | ~p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),sK1))))) )),
  inference(resolution,[],[f89,f96])).
fof(f89,plain,(
  ( ! [X9] : (~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X9))),s(sK0,sK4)))) | p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,X9))))) )),
  inference(cnf_transformation,[],[f58])).
fof(f775,plain,(
  ( ! [X6] : (~p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK6(X6))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X6))),s(sK0,X6)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(X6))))) )),
  inference(resolution,[],[f93,f132])).
fof(f132,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f82])).
fof(f9387,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK7(sK4,sK7(sK4,sK4,sK1,sK0),sK1,sK0))))) | $spl131),
  inference(cnf_transformation,[],[f9387_D])).
fof(f9387_D,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK7(sK4,sK7(sK4,sK4,sK1,sK0),sK1,sK0))))) <=> ~$spl131),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl131])])).
fof(f10155,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(sK4,sK7(sK4,sK4,sK1,sK0),sK1,sK0)))),s(sK0,sK7(sK4,sK7(sK4,sK4,sK1,sK0),sK1,sK0))))) | s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK7(sK4,sK7(sK4,sK4,sK1,sK0),sK1,sK0))))) | $spl104),
  inference(subsumption_resolution,[],[f10099,f108])).
fof(f10099,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(sK4,sK7(sK4,sK4,sK1,sK0),sK1,sK0)))),s(sK0,sK7(sK4,sK7(sK4,sK4,sK1,sK0),sK1,sK0))))) | s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK7(sK4,sK7(sK4,sK4,sK1,sK0),sK1,sK0))))) | $spl104),
  inference(superposition,[],[f777,f9376])).
fof(f777,plain,(
  ( ! [X8] : (~p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK6(X8))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X8))),s(sK0,X8)))) | s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(X8))))) )),
  inference(resolution,[],[f93,f135])).
fof(f135,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f84])).
fof(f9361,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK7(sK7(sK4,sK4,sK1,sK0),sK4,sK1,sK0))))) | $spl128),
  inference(cnf_transformation,[],[f9361_D])).
fof(f9361_D,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK7(sK7(sK4,sK4,sK1,sK0),sK4,sK1,sK0))))) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f9359,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(sK7(sK4,sK4,sK1,sK0),sK4,sK1,sK0)))),s(sK0,sK7(sK7(sK4,sK4,sK1,sK0),sK4,sK1,sK0))))) | $spl127),
  inference(cnf_transformation,[],[f9359_D])).
fof(f9359_D,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(sK7(sK4,sK4,sK1,sK0),sK4,sK1,sK0)))),s(sK0,sK7(sK7(sK4,sK4,sK1,sK0),sK4,sK1,sK0))))) <=> ~$spl127),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl127])])).
fof(f10058,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(sK7(sK4,sK4,sK1,sK0),sK4,sK1,sK0)))),s(sK0,sK7(sK7(sK4,sK4,sK1,sK0),sK4,sK1,sK0))))) | s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK7(sK7(sK4,sK4,sK1,sK0),sK4,sK1,sK0))))) | $spl104),
  inference(subsumption_resolution,[],[f10002,f108])).
fof(f10002,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(sK7(sK4,sK4,sK1,sK0),sK4,sK1,sK0)))),s(sK0,sK7(sK7(sK4,sK4,sK1,sK0),sK4,sK1,sK0))))) | s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK7(sK7(sK4,sK4,sK1,sK0),sK4,sK1,sK0))))) | $spl104),
  inference(superposition,[],[f775,f9346])).
fof(f9346,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK6(sK7(sK7(sK4,sK4,sK1,sK0),sK4,sK1,sK0))))) | $spl104),
  inference(resolution,[],[f9144,f136])).
fof(f9144,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK6(sK7(sK7(sK4,sK4,sK1,sK0),sK4,sK1,sK0)))))) | $spl104),
  inference(resolution,[],[f7894,f4026])).
fof(f4026,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,X0)))) | p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK6(sK7(X0,sK4,sK1,sK0))))))) )),
  inference(subsumption_resolution,[],[f3995,f198])).
fof(f3995,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK6(sK7(X0,sK4,sK1,sK0)))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,X0)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(X0,sK4,sK1,sK0)))),s(sK0,sK7(X0,sK4,sK1,sK0)))))) )),
  inference(resolution,[],[f403,f92])).
fof(f403,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK7(X1,sK4,sK1,sK0))))) | p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,X0)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X1))),s(sK0,X1))))) )),
  inference(subsumption_resolution,[],[f402,f108])).
fof(f402,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,X0)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK7(X1,sK4,sK1,sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X1))),s(sK0,X1))))) )),
  inference(forward_demodulation,[],[f401,f155])).
fof(f401,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,X0)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK7(X1,sK4,sK1,sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X1))),s(sK0,X1)))) | ~p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),sK1))))) )),
  inference(subsumption_resolution,[],[f378,f88])).
fof(f378,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,X0)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK7(X1,sK4,sK1,sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X1))),s(sK0,X1)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4))),s(sK0,sK4)))) | ~p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),sK1))))) )),
  inference(resolution,[],[f89,f95])).
fof(f9357,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK7(sK7(sK4,sK4,sK1,sK0),sK4,sK1,sK0))))) | $spl125),
  inference(cnf_transformation,[],[f9357_D])).
fof(f9357_D,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK7(sK7(sK4,sK4,sK1,sK0),sK4,sK1,sK0))))) <=> ~$spl125),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl125])])).
fof(f10057,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(sK7(sK4,sK4,sK1,sK0),sK4,sK1,sK0)))),s(sK0,sK7(sK7(sK4,sK4,sK1,sK0),sK4,sK1,sK0))))) | s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK7(sK7(sK4,sK4,sK1,sK0),sK4,sK1,sK0))))) | $spl104),
  inference(subsumption_resolution,[],[f10001,f108])).
fof(f10001,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(sK7(sK4,sK4,sK1,sK0),sK4,sK1,sK0)))),s(sK0,sK7(sK7(sK4,sK4,sK1,sK0),sK4,sK1,sK0))))) | s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK7(sK7(sK4,sK4,sK1,sK0),sK4,sK1,sK0))))) | $spl104),
  inference(superposition,[],[f777,f9346])).
fof(f9424,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK6(sK7(sK5,sK7(sK4,sK4,sK1,sK0),sK1,sK0)))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(sK5,sK7(sK4,sK4,sK1,sK0),sK1,sK0)))),s(sK0,sK7(sK5,sK7(sK4,sK4,sK1,sK0),sK1,sK0))))) | $spl104),
  inference(resolution,[],[f9150,f93])).
fof(f9397,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK6(sK7(sK7(sK4,sK4,sK1,sK0),sK5,sK1,sK0)))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(sK7(sK4,sK4,sK1,sK0),sK5,sK1,sK0)))),s(sK0,sK7(sK7(sK4,sK4,sK1,sK0),sK5,sK1,sK0))))) | $spl104),
  inference(resolution,[],[f9148,f93])).
fof(f9368,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(sK4,sK7(sK4,sK4,sK1,sK0),sK1,sK0)))),s(sK0,sK7(sK4,sK7(sK4,sK4,sK1,sK0),sK1,sK0))))) | s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK7(sK4,sK7(sK4,sK4,sK1,sK0),sK1,sK0))))) | $spl104),
  inference(resolution,[],[f9146,f775])).
fof(f9367,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(sK4,sK7(sK4,sK4,sK1,sK0),sK1,sK0)))),s(sK0,sK7(sK4,sK7(sK4,sK4,sK1,sK0),sK1,sK0))))) | s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK7(sK4,sK7(sK4,sK4,sK1,sK0),sK1,sK0))))) | $spl104),
  inference(resolution,[],[f9146,f777])).
fof(f9338,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(sK7(sK4,sK4,sK1,sK0),sK4,sK1,sK0)))),s(sK0,sK7(sK7(sK4,sK4,sK1,sK0),sK4,sK1,sK0))))) | s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK7(sK7(sK4,sK4,sK1,sK0),sK4,sK1,sK0))))) | $spl104),
  inference(resolution,[],[f9144,f775])).
fof(f9337,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(sK7(sK4,sK4,sK1,sK0),sK4,sK1,sK0)))),s(sK0,sK7(sK7(sK4,sK4,sK1,sK0),sK4,sK1,sK0))))) | s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK7(sK7(sK4,sK4,sK1,sK0),sK4,sK1,sK0))))) | $spl104),
  inference(resolution,[],[f9144,f777])).
fof(f8989,plain,(
  $false | $spl105),
  inference(subsumption_resolution,[],[f8988,f141])).
fof(f8988,plain,(
  ~p(s(bool,'T_0')) | $spl105),
  inference(forward_demodulation,[],[f8987,f155])).
fof(f8987,plain,(
  ~p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),sK1)))) | $spl105),
  inference(subsumption_resolution,[],[f8986,f129])).
fof(f8986,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),sK1)))) | $spl105),
  inference(forward_demodulation,[],[f8981,f185])).
fof(f8981,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4))),s(sK0,sK4)))) | ~p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),sK1)))) | $spl105),
  inference(duplicate_literal_removal,[],[f8958])).
fof(f8958,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4))),s(sK0,sK4)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4))),s(sK0,sK4)))) | ~p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),sK1)))) | $spl105),
  inference(resolution,[],[f7895,f94])).
fof(f8985,plain,(
  $false | $spl105),
  inference(subsumption_resolution,[],[f8984,f141])).
fof(f8984,plain,(
  ~p(s(bool,'T_0')) | $spl105),
  inference(forward_demodulation,[],[f8957,f185])).
fof(f8957,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4))),s(sK0,sK4)))) | $spl105),
  inference(resolution,[],[f7895,f198])).
fof(f8983,plain,(
  $false | $spl105),
  inference(subsumption_resolution,[],[f8982,f141])).
fof(f8982,plain,(
  ~p(s(bool,'T_0')) | $spl105),
  inference(forward_demodulation,[],[f8956,f185])).
fof(f8956,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4))),s(sK0,sK4)))) | $spl105),
  inference(resolution,[],[f7895,f197])).
fof(f8242,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(sK5,sK5,sK1,sK0)))),s(sK0,sK7(sK5,sK5,sK1,sK0))))) | $spl121),
  inference(cnf_transformation,[],[f8242_D])).
fof(f8242_D,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(sK5,sK5,sK1,sK0)))),s(sK0,sK7(sK5,sK5,sK1,sK0))))) <=> ~$spl121),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl121])])).
fof(f8244,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK6(sK7(sK5,sK5,sK1,sK0)))))) | $spl123),
  inference(cnf_transformation,[],[f8244_D])).
fof(f8244_D,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK6(sK7(sK5,sK5,sK1,sK0)))))) <=> ~$spl123),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl123])])).
fof(f8702,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK6(sK7(sK5,sK5,sK1,sK0)))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(sK5,sK5,sK1,sK0)))),s(sK0,sK7(sK5,sK5,sK1,sK0)))))),
  inference(subsumption_resolution,[],[f8649,f108])).
fof(f8649,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK6(sK7(sK5,sK5,sK1,sK0)))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(sK5,sK5,sK1,sK0)))),s(sK0,sK7(sK5,sK5,sK1,sK0)))))),
  inference(superposition,[],[f93,f8231])).
fof(f8231,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK7(sK5,sK5,sK1,sK0)))))),
  inference(resolution,[],[f8192,f136])).
fof(f8192,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK7(sK5,sK5,sK1,sK0))))))),
  inference(subsumption_resolution,[],[f8166,f140])).
fof(f8166,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK7(sK5,sK5,sK1,sK0))))))),
  inference(superposition,[],[f4305,f248])).
fof(f8117,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK7(sK4,sK5,sK1,sK0))))) | $spl118),
  inference(cnf_transformation,[],[f8117_D])).
fof(f8117_D,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK7(sK4,sK5,sK1,sK0))))) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f8115,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(sK4,sK5,sK1,sK0)))),s(sK0,sK7(sK4,sK5,sK1,sK0))))) | $spl117),
  inference(cnf_transformation,[],[f8115_D])).
fof(f8115_D,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(sK4,sK5,sK1,sK0)))),s(sK0,sK7(sK4,sK5,sK1,sK0))))) <=> ~$spl117),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl117])])).
fof(f8491,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(sK4,sK5,sK1,sK0)))),s(sK0,sK7(sK4,sK5,sK1,sK0))))) | s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK7(sK4,sK5,sK1,sK0)))))),
  inference(subsumption_resolution,[],[f8437,f108])).
fof(f8437,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(sK4,sK5,sK1,sK0)))),s(sK0,sK7(sK4,sK5,sK1,sK0))))) | s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK7(sK4,sK5,sK1,sK0)))))),
  inference(superposition,[],[f775,f8102])).
fof(f8102,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK6(sK7(sK4,sK5,sK1,sK0)))))),
  inference(resolution,[],[f8086,f136])).
fof(f8086,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK6(sK7(sK4,sK5,sK1,sK0))))))),
  inference(subsumption_resolution,[],[f8061,f140])).
fof(f8061,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK6(sK7(sK4,sK5,sK1,sK0))))))),
  inference(superposition,[],[f4217,f248])).
fof(f8113,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK7(sK4,sK5,sK1,sK0))))) | $spl115),
  inference(cnf_transformation,[],[f8113_D])).
fof(f8113_D,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK7(sK4,sK5,sK1,sK0))))) <=> ~$spl115),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl115])])).
fof(f8490,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(sK4,sK5,sK1,sK0)))),s(sK0,sK7(sK4,sK5,sK1,sK0))))) | s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK7(sK4,sK5,sK1,sK0)))))),
  inference(subsumption_resolution,[],[f8436,f108])).
fof(f8436,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(sK4,sK5,sK1,sK0)))),s(sK0,sK7(sK4,sK5,sK1,sK0))))) | s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK7(sK4,sK5,sK1,sK0)))))),
  inference(superposition,[],[f777,f8102])).
fof(f7927,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK7(sK5,sK4,sK1,sK0))))) | $spl112),
  inference(cnf_transformation,[],[f7927_D])).
fof(f7927_D,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK7(sK5,sK4,sK1,sK0))))) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f7925,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(sK5,sK4,sK1,sK0)))),s(sK0,sK7(sK5,sK4,sK1,sK0))))) | $spl111),
  inference(cnf_transformation,[],[f7925_D])).
fof(f7925_D,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(sK5,sK4,sK1,sK0)))),s(sK0,sK7(sK5,sK4,sK1,sK0))))) <=> ~$spl111),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl111])])).
fof(f8312,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(sK5,sK4,sK1,sK0)))),s(sK0,sK7(sK5,sK4,sK1,sK0))))) | s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK7(sK5,sK4,sK1,sK0)))))),
  inference(subsumption_resolution,[],[f8258,f108])).
fof(f8258,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(sK5,sK4,sK1,sK0)))),s(sK0,sK7(sK5,sK4,sK1,sK0))))) | s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK7(sK5,sK4,sK1,sK0)))))),
  inference(superposition,[],[f775,f7912])).
fof(f7912,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK6(sK7(sK5,sK4,sK1,sK0)))))),
  inference(resolution,[],[f7866,f136])).
fof(f7866,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK6(sK7(sK5,sK4,sK1,sK0))))))),
  inference(subsumption_resolution,[],[f7840,f140])).
fof(f7840,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK6(sK7(sK5,sK4,sK1,sK0))))))),
  inference(superposition,[],[f4026,f248])).
fof(f7923,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK7(sK5,sK4,sK1,sK0))))) | $spl109),
  inference(cnf_transformation,[],[f7923_D])).
fof(f7923_D,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK7(sK5,sK4,sK1,sK0))))) <=> ~$spl109),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl109])])).
fof(f8311,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(sK5,sK4,sK1,sK0)))),s(sK0,sK7(sK5,sK4,sK1,sK0))))) | s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK7(sK5,sK4,sK1,sK0)))))),
  inference(subsumption_resolution,[],[f8257,f108])).
fof(f8257,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(sK5,sK4,sK1,sK0)))),s(sK0,sK7(sK5,sK4,sK1,sK0))))) | s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK7(sK5,sK4,sK1,sK0)))))),
  inference(superposition,[],[f777,f7912])).
fof(f8223,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK6(sK7(sK5,sK5,sK1,sK0)))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(sK5,sK5,sK1,sK0)))),s(sK0,sK7(sK5,sK5,sK1,sK0)))))),
  inference(resolution,[],[f8192,f93])).
fof(f8094,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(sK4,sK5,sK1,sK0)))),s(sK0,sK7(sK4,sK5,sK1,sK0))))) | s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK7(sK4,sK5,sK1,sK0)))))),
  inference(resolution,[],[f8086,f775])).
fof(f8093,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(sK4,sK5,sK1,sK0)))),s(sK0,sK7(sK4,sK5,sK1,sK0))))) | s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK7(sK4,sK5,sK1,sK0)))))),
  inference(resolution,[],[f8086,f777])).
fof(f7897,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK7(sK4,sK4,sK1,sK0))))) | $spl106),
  inference(cnf_transformation,[],[f7897_D])).
fof(f7897_D,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK7(sK4,sK4,sK1,sK0))))) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f7895,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(sK4,sK4,sK1,sK0)))),s(sK0,sK7(sK4,sK4,sK1,sK0))))) | $spl105),
  inference(cnf_transformation,[],[f7895_D])).
fof(f7895_D,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(sK4,sK4,sK1,sK0)))),s(sK0,sK7(sK4,sK4,sK1,sK0))))) <=> ~$spl105),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl105])])).
fof(f8006,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(sK4,sK4,sK1,sK0)))),s(sK0,sK7(sK4,sK4,sK1,sK0))))) | s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK7(sK4,sK4,sK1,sK0)))))),
  inference(subsumption_resolution,[],[f7952,f108])).
fof(f7952,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(sK4,sK4,sK1,sK0)))),s(sK0,sK7(sK4,sK4,sK1,sK0))))) | s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK7(sK4,sK4,sK1,sK0)))))),
  inference(superposition,[],[f775,f7882])).
fof(f7882,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK6(sK7(sK4,sK4,sK1,sK0)))))),
  inference(resolution,[],[f7865,f136])).
fof(f7865,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK6(sK7(sK4,sK4,sK1,sK0))))))),
  inference(subsumption_resolution,[],[f7839,f140])).
fof(f7839,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK6(sK7(sK4,sK4,sK1,sK0))))))),
  inference(superposition,[],[f4026,f185])).
fof(f7893,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK7(sK4,sK4,sK1,sK0))))) | $spl103),
  inference(cnf_transformation,[],[f7893_D])).
fof(f7893_D,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK7(sK4,sK4,sK1,sK0))))) <=> ~$spl103),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl103])])).
fof(f8005,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(sK4,sK4,sK1,sK0)))),s(sK0,sK7(sK4,sK4,sK1,sK0))))) | s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK7(sK4,sK4,sK1,sK0)))))),
  inference(subsumption_resolution,[],[f7951,f108])).
fof(f7951,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(sK4,sK4,sK1,sK0)))),s(sK0,sK7(sK4,sK4,sK1,sK0))))) | s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK7(sK4,sK4,sK1,sK0)))))),
  inference(superposition,[],[f777,f7882])).
fof(f7904,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(sK5,sK4,sK1,sK0)))),s(sK0,sK7(sK5,sK4,sK1,sK0))))) | s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK7(sK5,sK4,sK1,sK0)))))),
  inference(resolution,[],[f7866,f775])).
fof(f7903,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(sK5,sK4,sK1,sK0)))),s(sK0,sK7(sK5,sK4,sK1,sK0))))) | s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK7(sK5,sK4,sK1,sK0)))))),
  inference(resolution,[],[f7866,f777])).
fof(f7874,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(sK4,sK4,sK1,sK0)))),s(sK0,sK7(sK4,sK4,sK1,sK0))))) | s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK7(sK4,sK4,sK1,sK0)))))),
  inference(resolution,[],[f7865,f775])).
fof(f7873,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(sK4,sK4,sK1,sK0)))),s(sK0,sK7(sK4,sK4,sK1,sK0))))) | s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK7(sK4,sK4,sK1,sK0)))))),
  inference(resolution,[],[f7865,f777])).
fof(f7314,plain,(
  s(fun(sK0,bool),sK3) = s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5))) | $spl100),
  inference(cnf_transformation,[],[f7314_D])).
fof(f7314_D,plain,(
  s(fun(sK0,bool),sK3) = s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5))) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f6747,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK11(sK3,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5)),bool,sK0)))) | $spl51),
  inference(cnf_transformation,[],[f6747_D])).
fof(f6747_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK11(sK3,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5)),bool,sK0)))) <=> ~$spl51),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl51])])).
fof(f7423,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK11(sK3,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5)),bool,sK0)))) | s(fun(sK0,bool),sK3) = s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5))) | $spl52),
  inference(superposition,[],[f143,f7025])).
fof(f7025,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK11(sK3,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5)),bool,sK0)))) | $spl52),
  inference(resolution,[],[f6750,f136])).
fof(f143,plain,(
  ( ! [X2,X0,X3,X1] : (s(fun(X0,X1),X2) = s(fun(X0,X1),X3) | s(X1,f(s(X0,sK11(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK11(X3,X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f86])).
fof(f86,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK11(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK11(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK11])],[f56])).
fof(f56,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f46])).
fof(f46,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f45])).
fof(f45,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.nets.DORDER_LEMMA',aHL_EXT)).
fof(f6875,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK4)))) | $spl61),
  inference(cnf_transformation,[],[f6875_D])).
fof(f6875_D,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK4)))) <=> ~$spl61),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl61])])).
fof(f6873,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK5)))) | $spl58),
  inference(cnf_transformation,[],[f6873_D])).
fof(f6873_D,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK5)))) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f7396,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK11(sK3,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5)),bool,sK0)))) | p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK5)))) | ~p(s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK4)))) | $spl52),
  inference(superposition,[],[f394,f7025])).
fof(f394,plain,(
  ( ! [X2,X3] : (s(bool,f(s(sK0,sK11(X3,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X2)),bool,sK0)))) != s(bool,g(s(sK0,sK11(X3,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X2)),bool,sK0)))) | p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,X2)))) | ~p(s(bool,happ(s(fun(sK0,bool),X3),s(sK0,sK4))))) )),
  inference(superposition,[],[f89,f143])).
fof(f7300,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5))),s(sK0,sK7(sK5,sK4,sK1,sK0))))) | $spl98),
  inference(cnf_transformation,[],[f7300_D])).
fof(f7300_D,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5))),s(sK0,sK7(sK5,sK4,sK1,sK0))))) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f7393,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK11(sK3,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5)),bool,sK0)))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5))),s(sK0,sK7(sK5,sK4,sK1,sK0))))) | $spl52),
  inference(superposition,[],[f4516,f7025])).
fof(f4516,plain,(
  ( ! [X0] : (s(bool,f(s(sK0,sK11(sK3,X0,bool,sK0)))) != s(bool,g(s(sK0,sK11(sK3,X0,bool,sK0)))) | p(s(bool,happ(s(fun(sK0,bool),X0),s(sK0,sK7(sK5,sK4,sK1,sK0)))))) )),
  inference(superposition,[],[f4489,f143])).
fof(f4489,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK7(sK5,sK4,sK1,sK0)))))),
  inference(subsumption_resolution,[],[f4488,f140])).
fof(f4488,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK7(sK5,sK4,sK1,sK0)))))),
  inference(forward_demodulation,[],[f4487,f248])).
fof(f4487,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK7(sK5,sK4,sK1,sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5))),s(sK0,sK5))))),
  inference(subsumption_resolution,[],[f4486,f108])).
fof(f4486,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK7(sK5,sK4,sK1,sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5))),s(sK0,sK5))))),
  inference(forward_demodulation,[],[f4455,f185])).
fof(f4455,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK7(sK5,sK4,sK1,sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4))),s(sK0,sK4)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5))),s(sK0,sK5))))),
  inference(resolution,[],[f558,f198])).
fof(f7298,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5))),s(sK0,sK7(sK5,sK5,sK1,sK0))))) | $spl96),
  inference(cnf_transformation,[],[f7298_D])).
fof(f7298_D,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5))),s(sK0,sK7(sK5,sK5,sK1,sK0))))) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f7392,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK11(sK3,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5)),bool,sK0)))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5))),s(sK0,sK7(sK5,sK5,sK1,sK0))))) | $spl52),
  inference(superposition,[],[f4360,f7025])).
fof(f4360,plain,(
  ( ! [X0] : (s(bool,f(s(sK0,sK11(sK3,X0,bool,sK0)))) != s(bool,g(s(sK0,sK11(sK3,X0,bool,sK0)))) | p(s(bool,happ(s(fun(sK0,bool),X0),s(sK0,sK7(sK5,sK5,sK1,sK0)))))) )),
  inference(superposition,[],[f4311,f143])).
fof(f4311,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK7(sK5,sK5,sK1,sK0)))))),
  inference(subsumption_resolution,[],[f4310,f140])).
fof(f4310,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK7(sK5,sK5,sK1,sK0)))))),
  inference(forward_demodulation,[],[f4303,f248])).
fof(f4303,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK7(sK5,sK5,sK1,sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5))),s(sK0,sK5))))),
  inference(duplicate_literal_removal,[],[f4277])).
fof(f4277,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK7(sK5,sK5,sK1,sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5))),s(sK0,sK5)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5))),s(sK0,sK5))))),
  inference(resolution,[],[f555,f260])).
fof(f7296,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5))),s(sK0,sK7(sK4,sK5,sK1,sK0))))) | $spl94),
  inference(cnf_transformation,[],[f7296_D])).
fof(f7296_D,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5))),s(sK0,sK7(sK4,sK5,sK1,sK0))))) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f7391,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK11(sK3,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5)),bool,sK0)))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5))),s(sK0,sK7(sK4,sK5,sK1,sK0))))) | $spl52),
  inference(superposition,[],[f4338,f7025])).
fof(f4338,plain,(
  ( ! [X0] : (s(bool,f(s(sK0,sK11(sK3,X0,bool,sK0)))) != s(bool,g(s(sK0,sK11(sK3,X0,bool,sK0)))) | p(s(bool,happ(s(fun(sK0,bool),X0),s(sK0,sK7(sK4,sK5,sK1,sK0)))))) )),
  inference(superposition,[],[f4309,f143])).
fof(f4309,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK7(sK4,sK5,sK1,sK0)))))),
  inference(subsumption_resolution,[],[f4308,f140])).
fof(f4308,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK7(sK4,sK5,sK1,sK0)))))),
  inference(forward_demodulation,[],[f4307,f248])).
fof(f4307,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK7(sK4,sK5,sK1,sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5))),s(sK0,sK5))))),
  inference(subsumption_resolution,[],[f4306,f108])).
fof(f4306,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK7(sK4,sK5,sK1,sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5))),s(sK0,sK5))))),
  inference(forward_demodulation,[],[f4275,f185])).
fof(f4275,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK7(sK4,sK5,sK1,sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4))),s(sK0,sK4)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5))),s(sK0,sK5))))),
  inference(resolution,[],[f555,f197])).
fof(f7294,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5))),s(sK0,sK6(sK4)))) | $spl93),
  inference(cnf_transformation,[],[f7294_D])).
fof(f7294_D,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5))),s(sK0,sK6(sK4)))) <=> ~$spl93),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl93])])).
fof(f7388,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK11(sK3,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5)),bool,sK0)))) | s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5))),s(sK0,sK6(sK4)))) | $spl52),
  inference(superposition,[],[f2751,f7025])).
fof(f2751,plain,(
  ( ! [X0] : (s(bool,f(s(sK0,sK11(sK3,X0,bool,sK0)))) != s(bool,g(s(sK0,sK11(sK3,X0,bool,sK0)))) | s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),X0),s(sK0,sK6(sK4))))) )),
  inference(superposition,[],[f2735,f143])).
fof(f2735,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK4))))),
  inference(subsumption_resolution,[],[f2694,f139])).
fof(f139,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f44])).
fof(f44,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.nets.DORDER_LEMMA',aHL_FALSITY)).
fof(f2694,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK4))))),
  inference(superposition,[],[f135,f2666])).
fof(f2666,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK4))))),
  inference(subsumption_resolution,[],[f2665,f140])).
fof(f2665,plain,(
  ~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK4))))),
  inference(forward_demodulation,[],[f2664,f185])).
fof(f2664,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4))),s(sK0,sK4)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK4))))),
  inference(subsumption_resolution,[],[f2653,f108])).
fof(f2653,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4))),s(sK0,sK4)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK4))))),
  inference(superposition,[],[f775,f731])).
fof(f731,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK6(sK4))))),
  inference(resolution,[],[f704,f136])).
fof(f704,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK6(sK4)))))),
  inference(subsumption_resolution,[],[f681,f88])).
fof(f681,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4))),s(sK0,sK4)))) | p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK6(sK4)))))),
  inference(resolution,[],[f92,f89])).
fof(f7292,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5))),s(sK0,sK6(sK4)))) | $spl90),
  inference(cnf_transformation,[],[f7292_D])).
fof(f7292_D,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5))),s(sK0,sK6(sK4)))) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f7387,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK11(sK3,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5)),bool,sK0)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5))),s(sK0,sK6(sK4)))) | $spl52),
  inference(superposition,[],[f2668,f7025])).
fof(f2668,plain,(
  ( ! [X0] : (s(bool,f(s(sK0,sK11(sK3,X0,bool,sK0)))) != s(bool,g(s(sK0,sK11(sK3,X0,bool,sK0)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),X0),s(sK0,sK6(sK4))))) )),
  inference(superposition,[],[f2666,f143])).
fof(f7288,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5))),s(sK0,sK6(sK5)))) | $spl88),
  inference(cnf_transformation,[],[f7288_D])).
fof(f7288_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5))),s(sK0,sK6(sK5)))) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f7385,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK11(sK3,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5)),bool,sK0)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5))),s(sK0,sK6(sK5)))) | $spl52),
  inference(superposition,[],[f919,f7025])).
fof(f919,plain,(
  ( ! [X0] : (s(bool,f(s(sK0,sK11(sK3,X0,bool,sK0)))) != s(bool,g(s(sK0,sK11(sK3,X0,bool,sK0)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),X0),s(sK0,sK6(sK5))))) )),
  inference(superposition,[],[f753,f143])).
fof(f753,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK5))))),
  inference(resolution,[],[f705,f136])).
fof(f705,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK5)))))),
  inference(subsumption_resolution,[],[f682,f90])).
fof(f682,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5))),s(sK0,sK5)))) | p(s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK5)))))),
  inference(resolution,[],[f92,f91])).
fof(f7286,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5))),s(sK0,sK6(sK5)))) | $spl87),
  inference(cnf_transformation,[],[f7286_D])).
fof(f7286_D,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5))),s(sK0,sK6(sK5)))) <=> ~$spl87),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl87])])).
fof(f7384,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK11(sK3,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5)),bool,sK0)))) | s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5))),s(sK0,sK6(sK5)))) | $spl52),
  inference(superposition,[],[f906,f7025])).
fof(f906,plain,(
  ( ! [X0] : (s(bool,f(s(sK0,sK11(sK3,X0,bool,sK0)))) != s(bool,g(s(sK0,sK11(sK3,X0,bool,sK0)))) | s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),X0),s(sK0,sK6(sK5))))) )),
  inference(superposition,[],[f751,f143])).
fof(f751,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK6(sK5))))),
  inference(resolution,[],[f705,f131])).
fof(f7284,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK6(X0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5))),s(sK0,sK6(X0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,X0))))) ) | $spl84),
  inference(cnf_transformation,[],[f7284_D])).
fof(f7284_D,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK6(X0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5))),s(sK0,sK6(X0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,X0))))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f7383,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(sK0,sK11(sK3,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5)),bool,sK0)))) | ~p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK6(X0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,X0)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5))),s(sK0,sK6(X0)))))) ) | $spl52),
  inference(superposition,[],[f779,f7025])).
fof(f779,plain,(
  ( ! [X0,X1] : (s(bool,f(s(sK0,sK11(sK3,X0,bool,sK0)))) != s(bool,g(s(sK0,sK11(sK3,X0,bool,sK0)))) | ~p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK6(X1))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X1))),s(sK0,X1)))) | ~p(s(bool,happ(s(fun(sK0,bool),X0),s(sK0,sK6(X1)))))) )),
  inference(superposition,[],[f93,f143])).
fof(f7283,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5))),s(sK0,sK6(sK5))))) | $spl82),
  inference(cnf_transformation,[],[f7283_D])).
fof(f7283_D,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5))),s(sK0,sK6(sK5))))) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f7382,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK11(sK3,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5)),bool,sK0)))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5))),s(sK0,sK6(sK5))))) | $spl52),
  inference(superposition,[],[f755,f7025])).
fof(f755,plain,(
  ( ! [X0] : (s(bool,f(s(sK0,sK11(sK3,X0,bool,sK0)))) != s(bool,g(s(sK0,sK11(sK3,X0,bool,sK0)))) | p(s(bool,happ(s(fun(sK0,bool),X0),s(sK0,sK6(sK5)))))) )),
  inference(superposition,[],[f705,f143])).
fof(f6742,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5)),sK3,bool,sK0)))) | $spl45),
  inference(cnf_transformation,[],[f6742_D])).
fof(f6742_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5)),sK3,bool,sK0)))) <=> ~$spl45),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl45])])).
fof(f7254,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5)),sK3,bool,sK0)))) | s(fun(sK0,bool),sK3) = s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5))) | $spl46),
  inference(superposition,[],[f143,f7001])).
fof(f7001,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5)),sK3,bool,sK0)))) | $spl46),
  inference(resolution,[],[f6745,f136])).
fof(f7227,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5)),sK3,bool,sK0)))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5))),s(sK0,sK7(sK5,sK4,sK1,sK0))))) | $spl46),
  inference(superposition,[],[f4517,f7001])).
fof(f4517,plain,(
  ( ! [X1] : (s(bool,f(s(sK0,sK11(X1,sK3,bool,sK0)))) != s(bool,g(s(sK0,sK11(X1,sK3,bool,sK0)))) | p(s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK7(sK5,sK4,sK1,sK0)))))) )),
  inference(superposition,[],[f4489,f143])).
fof(f7226,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5)),sK3,bool,sK0)))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5))),s(sK0,sK7(sK5,sK5,sK1,sK0))))) | $spl46),
  inference(superposition,[],[f4361,f7001])).
fof(f4361,plain,(
  ( ! [X1] : (s(bool,f(s(sK0,sK11(X1,sK3,bool,sK0)))) != s(bool,g(s(sK0,sK11(X1,sK3,bool,sK0)))) | p(s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK7(sK5,sK5,sK1,sK0)))))) )),
  inference(superposition,[],[f4311,f143])).
fof(f7225,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5)),sK3,bool,sK0)))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5))),s(sK0,sK7(sK4,sK5,sK1,sK0))))) | $spl46),
  inference(superposition,[],[f4339,f7001])).
fof(f4339,plain,(
  ( ! [X1] : (s(bool,f(s(sK0,sK11(X1,sK3,bool,sK0)))) != s(bool,g(s(sK0,sK11(X1,sK3,bool,sK0)))) | p(s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK7(sK4,sK5,sK1,sK0)))))) )),
  inference(superposition,[],[f4309,f143])).
fof(f7222,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5)),sK3,bool,sK0)))) | s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5))),s(sK0,sK6(sK4)))) | $spl46),
  inference(superposition,[],[f2752,f7001])).
fof(f2752,plain,(
  ( ! [X1] : (s(bool,f(s(sK0,sK11(X1,sK3,bool,sK0)))) != s(bool,g(s(sK0,sK11(X1,sK3,bool,sK0)))) | s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK6(sK4))))) )),
  inference(superposition,[],[f2735,f143])).
fof(f7221,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5)),sK3,bool,sK0)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5))),s(sK0,sK6(sK4)))) | $spl46),
  inference(superposition,[],[f2669,f7001])).
fof(f2669,plain,(
  ( ! [X1] : (s(bool,f(s(sK0,sK11(X1,sK3,bool,sK0)))) != s(bool,g(s(sK0,sK11(X1,sK3,bool,sK0)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK6(sK4))))) )),
  inference(superposition,[],[f2666,f143])).
fof(f7219,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5)),sK3,bool,sK0)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5))),s(sK0,sK6(sK5)))) | $spl46),
  inference(superposition,[],[f920,f7001])).
fof(f920,plain,(
  ( ! [X1] : (s(bool,f(s(sK0,sK11(X1,sK3,bool,sK0)))) != s(bool,g(s(sK0,sK11(X1,sK3,bool,sK0)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK6(sK5))))) )),
  inference(superposition,[],[f753,f143])).
fof(f7218,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5)),sK3,bool,sK0)))) | s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5))),s(sK0,sK6(sK5)))) | $spl46),
  inference(superposition,[],[f907,f7001])).
fof(f907,plain,(
  ( ! [X1] : (s(bool,f(s(sK0,sK11(X1,sK3,bool,sK0)))) != s(bool,g(s(sK0,sK11(X1,sK3,bool,sK0)))) | s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK6(sK5))))) )),
  inference(superposition,[],[f751,f143])).
fof(f7217,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5)),sK3,bool,sK0)))) | ~p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK6(X0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,X0)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5))),s(sK0,sK6(X0)))))) ) | $spl46),
  inference(superposition,[],[f780,f7001])).
fof(f780,plain,(
  ( ! [X2,X3] : (s(bool,f(s(sK0,sK11(X2,sK3,bool,sK0)))) != s(bool,g(s(sK0,sK11(X2,sK3,bool,sK0)))) | ~p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK6(X3))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X3))),s(sK0,X3)))) | ~p(s(bool,happ(s(fun(sK0,bool),X2),s(sK0,sK6(X3)))))) )),
  inference(superposition,[],[f93,f143])).
fof(f7216,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5)),sK3,bool,sK0)))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5))),s(sK0,sK6(sK5))))) | $spl46),
  inference(superposition,[],[f756,f7001])).
fof(f756,plain,(
  ( ! [X1] : (s(bool,f(s(sK0,sK11(X1,sK3,bool,sK0)))) != s(bool,g(s(sK0,sK11(X1,sK3,bool,sK0)))) | p(s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK6(sK5)))))) )),
  inference(superposition,[],[f705,f143])).
fof(f7212,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5)),sK3,bool,sK0)))) | p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK5)))) | ~p(s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK4)))) | $spl46),
  inference(superposition,[],[f393,f7001])).
fof(f393,plain,(
  ( ! [X0,X1] : (s(bool,f(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0)),X1,bool,sK0)))) != s(bool,g(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0)),X1,bool,sK0)))) | p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,X0)))) | ~p(s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK4))))) )),
  inference(superposition,[],[f89,f143])).
fof(f6754,plain,(
  p(s(bool,f(s(sK0,sK11(sK3,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5)),bool,sK0))))) | $spl56),
  inference(cnf_transformation,[],[f6754_D])).
fof(f6754_D,plain,(
  p(s(bool,f(s(sK0,sK11(sK3,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5)),bool,sK0))))) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f6750,plain,(
  p(s(bool,g(s(sK0,sK11(sK3,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5)),bool,sK0))))) | $spl52),
  inference(cnf_transformation,[],[f6750_D])).
fof(f6750_D,plain,(
  p(s(bool,g(s(sK0,sK11(sK3,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5)),bool,sK0))))) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f6986,plain,(
  ( ! [X23] : (~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X23))),s(sK0,X23)))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK7(X23,sK5,sK1,sK0))))),s(sK0,X23))))) ) | $spl80),
  inference(cnf_transformation,[],[f6986_D])).
fof(f6986_D,plain,(
  ( ! [X23] : (~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X23))),s(sK0,X23)))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK7(X23,sK5,sK1,sK0))))),s(sK0,X23))))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f6948,plain,(
  ( ! [X26] : (~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X26))),s(sK0,X26)))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK7(X26,sK5,sK1,sK0))))),s(sK0,X26)))) | p(s(bool,g(s(sK0,sK11(sK3,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5)),bool,sK0))))) | p(s(bool,f(s(sK0,sK11(sK3,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5)),bool,sK0)))))) )),
  inference(resolution,[],[f711,f2937])).
fof(f2937,plain,(
  ( ! [X5] : (p(s(bool,happ(s(fun(sK0,bool),X5),s(sK0,sK5)))) | p(s(bool,g(s(sK0,sK11(sK3,X5,bool,sK0))))) | p(s(bool,f(s(sK0,sK11(sK3,X5,bool,sK0)))))) )),
  inference(trivial_inequality_removal,[],[f2936])).
fof(f2936,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,'F_0') | p(s(bool,happ(s(fun(sK0,bool),X5),s(sK0,sK5)))) | p(s(bool,g(s(sK0,sK11(sK3,X5,bool,sK0))))) | p(s(bool,f(s(sK0,sK11(sK3,X5,bool,sK0)))))) )),
  inference(superposition,[],[f1371,f132])).
fof(f1371,plain,(
  ( ! [X4] : (s(bool,'F_0') != s(bool,f(s(sK0,sK11(sK3,X4,bool,sK0)))) | p(s(bool,happ(s(fun(sK0,bool),X4),s(sK0,sK5)))) | p(s(bool,g(s(sK0,sK11(sK3,X4,bool,sK0)))))) )),
  inference(superposition,[],[f578,f134])).
fof(f134,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f41])).
fof(f41,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f40])).
fof(f40,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X12] : (s(bool,'F_0') = s(bool,X12) <=> ~p(s(bool,X12)))),
  file('thm.nets.DORDER_LEMMA','thm.bool.EQ_CLAUSES|split|2')).
fof(f578,plain,(
  ( ! [X0] : (s(bool,f(s(sK0,sK11(sK3,X0,bool,sK0)))) != s(bool,g(s(sK0,sK11(sK3,X0,bool,sK0)))) | p(s(bool,happ(s(fun(sK0,bool),X0),s(sK0,sK5))))) )),
  inference(superposition,[],[f529,f143])).
fof(f529,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK5))))),
  inference(resolution,[],[f91,f90])).
fof(f711,plain,(
  ( ! [X2,X3] : (~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X3))),s(sK0,X3)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X2))),s(sK0,X2)))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK7(X2,X3,sK1,sK0))))),s(sK0,X2))))) )),
  inference(subsumption_resolution,[],[f710,f108])).
fof(f710,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK7(X2,X3,sK1,sK0))))),s(sK0,X2)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X2))),s(sK0,X2)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X3))),s(sK0,X3))))) )),
  inference(forward_demodulation,[],[f709,f155])).
fof(f709,plain,(
  ( ! [X2,X3] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK7(X2,X3,sK1,sK0))))),s(sK0,X2)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X2))),s(sK0,X2)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X3))),s(sK0,X3)))) | ~p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),sK1))))) )),
  inference(subsumption_resolution,[],[f684,f94])).
fof(f684,plain,(
  ( ! [X2,X3] : (~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(X2,X3,sK1,sK0)))),s(sK0,sK7(X2,X3,sK1,sK0))))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK7(X2,X3,sK1,sK0))))),s(sK0,X2)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X2))),s(sK0,X2)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X3))),s(sK0,X3)))) | ~p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),sK1))))) )),
  inference(resolution,[],[f92,f96])).
fof(f6752,plain,(
  p(s(bool,f(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5)),sK3,bool,sK0))))) | $spl54),
  inference(cnf_transformation,[],[f6752_D])).
fof(f6752_D,plain,(
  p(s(bool,f(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5)),sK3,bool,sK0))))) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f6745,plain,(
  p(s(bool,g(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5)),sK3,bool,sK0))))) | $spl46),
  inference(cnf_transformation,[],[f6745_D])).
fof(f6745_D,plain,(
  p(s(bool,g(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5)),sK3,bool,sK0))))) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f6947,plain,(
  ( ! [X25] : (~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X25))),s(sK0,X25)))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK7(X25,sK5,sK1,sK0))))),s(sK0,X25)))) | p(s(bool,g(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5)),sK3,bool,sK0))))) | p(s(bool,f(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5)),sK3,bool,sK0)))))) )),
  inference(resolution,[],[f711,f3126])).
fof(f3126,plain,(
  ( ! [X5] : (p(s(bool,happ(s(fun(sK0,bool),X5),s(sK0,sK5)))) | p(s(bool,g(s(sK0,sK11(X5,sK3,bool,sK0))))) | p(s(bool,f(s(sK0,sK11(X5,sK3,bool,sK0)))))) )),
  inference(trivial_inequality_removal,[],[f3125])).
fof(f3125,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,'F_0') | p(s(bool,happ(s(fun(sK0,bool),X5),s(sK0,sK5)))) | p(s(bool,g(s(sK0,sK11(X5,sK3,bool,sK0))))) | p(s(bool,f(s(sK0,sK11(X5,sK3,bool,sK0)))))) )),
  inference(superposition,[],[f1513,f132])).
fof(f1513,plain,(
  ( ! [X4] : (s(bool,'F_0') != s(bool,f(s(sK0,sK11(X4,sK3,bool,sK0)))) | p(s(bool,happ(s(fun(sK0,bool),X4),s(sK0,sK5)))) | p(s(bool,g(s(sK0,sK11(X4,sK3,bool,sK0)))))) )),
  inference(superposition,[],[f579,f134])).
fof(f579,plain,(
  ( ! [X1] : (s(bool,f(s(sK0,sK11(X1,sK3,bool,sK0)))) != s(bool,g(s(sK0,sK11(X1,sK3,bool,sK0)))) | p(s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK5))))) )),
  inference(superposition,[],[f529,f143])).
fof(f6748,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK11(sK3,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5)),bool,sK0)))) | $spl50),
  inference(cnf_transformation,[],[f6748_D])).
fof(f6748_D,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK11(sK3,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5)),bool,sK0)))) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f6946,plain,(
  ( ! [X24] : (~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X24))),s(sK0,X24)))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK7(X24,sK5,sK1,sK0))))),s(sK0,X24)))) | p(s(bool,g(s(sK0,sK11(sK3,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5)),bool,sK0))))) | s(bool,'T_0') = s(bool,f(s(sK0,sK11(sK3,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5)),bool,sK0))))) )),
  inference(resolution,[],[f711,f2939])).
fof(f2939,plain,(
  ( ! [X1] : (p(s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK5)))) | p(s(bool,g(s(sK0,sK11(sK3,X1,bool,sK0))))) | s(bool,'T_0') = s(bool,f(s(sK0,sK11(sK3,X1,bool,sK0))))) )),
  inference(trivial_inequality_removal,[],[f2932])).
fof(f2932,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,'F_0') | p(s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK5)))) | p(s(bool,g(s(sK0,sK11(sK3,X1,bool,sK0))))) | s(bool,'T_0') = s(bool,f(s(sK0,sK11(sK3,X1,bool,sK0))))) )),
  inference(superposition,[],[f1371,f142])).
fof(f6743,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5)),sK3,bool,sK0)))) | $spl44),
  inference(cnf_transformation,[],[f6743_D])).
fof(f6743_D,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5)),sK3,bool,sK0)))) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f6945,plain,(
  ( ! [X23] : (~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X23))),s(sK0,X23)))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK7(X23,sK5,sK1,sK0))))),s(sK0,X23)))) | p(s(bool,g(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5)),sK3,bool,sK0))))) | s(bool,'T_0') = s(bool,f(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5)),sK3,bool,sK0))))) )),
  inference(resolution,[],[f711,f3128])).
fof(f3128,plain,(
  ( ! [X1] : (p(s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK5)))) | p(s(bool,g(s(sK0,sK11(X1,sK3,bool,sK0))))) | s(bool,'T_0') = s(bool,f(s(sK0,sK11(X1,sK3,bool,sK0))))) )),
  inference(trivial_inequality_removal,[],[f3121])).
fof(f3121,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,'F_0') | p(s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK5)))) | p(s(bool,g(s(sK0,sK11(X1,sK3,bool,sK0))))) | s(bool,'T_0') = s(bool,f(s(sK0,sK11(X1,sK3,bool,sK0))))) )),
  inference(superposition,[],[f1513,f142])).
fof(f6910,plain,(
  s(fun(sK0,bool),sK2) = s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4))) | $spl78),
  inference(cnf_transformation,[],[f6910_D])).
fof(f6910_D,plain,(
  s(fun(sK0,bool),sK2) = s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4))) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f6731,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4)),sK2,bool,sK0)))) | $spl37),
  inference(cnf_transformation,[],[f6731_D])).
fof(f6731_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4)),sK2,bool,sK0)))) <=> ~$spl37),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl37])])).
fof(f6847,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4)),sK2,bool,sK0)))) | s(fun(sK0,bool),sK2) = s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4))) | $spl38),
  inference(superposition,[],[f143,f6769])).
fof(f6769,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4)),sK2,bool,sK0)))) | $spl38),
  inference(resolution,[],[f6734,f136])).
fof(f6896,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4))),s(sK0,sK7(sK4,sK5,sK1,sK0))))) | $spl76),
  inference(cnf_transformation,[],[f6896_D])).
fof(f6896_D,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4))),s(sK0,sK7(sK4,sK5,sK1,sK0))))) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f6820,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4)),sK2,bool,sK0)))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4))),s(sK0,sK7(sK4,sK5,sK1,sK0))))) | $spl38),
  inference(superposition,[],[f4251,f6769])).
fof(f4251,plain,(
  ( ! [X1] : (s(bool,f(s(sK0,sK11(X1,sK2,bool,sK0)))) != s(bool,g(s(sK0,sK11(X1,sK2,bool,sK0)))) | p(s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK7(sK4,sK5,sK1,sK0)))))) )),
  inference(superposition,[],[f4223,f143])).
fof(f4223,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK7(sK4,sK5,sK1,sK0)))))),
  inference(subsumption_resolution,[],[f4222,f140])).
fof(f4222,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK7(sK4,sK5,sK1,sK0)))))),
  inference(forward_demodulation,[],[f4221,f185])).
fof(f4221,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK7(sK4,sK5,sK1,sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4))),s(sK0,sK4))))),
  inference(subsumption_resolution,[],[f4220,f108])).
fof(f4220,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK7(sK4,sK5,sK1,sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4))),s(sK0,sK4))))),
  inference(forward_demodulation,[],[f4189,f248])).
fof(f4189,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK7(sK4,sK5,sK1,sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5))),s(sK0,sK5)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4))),s(sK0,sK4))))),
  inference(resolution,[],[f406,f261])).
fof(f6894,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4))),s(sK0,sK7(sK5,sK4,sK1,sK0))))) | $spl74),
  inference(cnf_transformation,[],[f6894_D])).
fof(f6894_D,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4))),s(sK0,sK7(sK5,sK4,sK1,sK0))))) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f6819,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4)),sK2,bool,sK0)))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4))),s(sK0,sK7(sK5,sK4,sK1,sK0))))) | $spl38),
  inference(superposition,[],[f4082,f6769])).
fof(f4082,plain,(
  ( ! [X1] : (s(bool,f(s(sK0,sK11(X1,sK2,bool,sK0)))) != s(bool,g(s(sK0,sK11(X1,sK2,bool,sK0)))) | p(s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK7(sK5,sK4,sK1,sK0)))))) )),
  inference(superposition,[],[f4032,f143])).
fof(f4032,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK7(sK5,sK4,sK1,sK0)))))),
  inference(subsumption_resolution,[],[f4031,f140])).
fof(f4031,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK7(sK5,sK4,sK1,sK0)))))),
  inference(forward_demodulation,[],[f4030,f185])).
fof(f4030,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK7(sK5,sK4,sK1,sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4))),s(sK0,sK4))))),
  inference(subsumption_resolution,[],[f4029,f108])).
fof(f4029,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK7(sK5,sK4,sK1,sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4))),s(sK0,sK4))))),
  inference(forward_demodulation,[],[f3998,f248])).
fof(f3998,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK7(sK5,sK4,sK1,sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5))),s(sK0,sK5)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4))),s(sK0,sK4))))),
  inference(resolution,[],[f403,f260])).
fof(f6892,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4))),s(sK0,sK7(sK4,sK4,sK1,sK0))))) | $spl72),
  inference(cnf_transformation,[],[f6892_D])).
fof(f6892_D,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4))),s(sK0,sK7(sK4,sK4,sK1,sK0))))) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f6818,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4)),sK2,bool,sK0)))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4))),s(sK0,sK7(sK4,sK4,sK1,sK0))))) | $spl38),
  inference(superposition,[],[f4060,f6769])).
fof(f4060,plain,(
  ( ! [X1] : (s(bool,f(s(sK0,sK11(X1,sK2,bool,sK0)))) != s(bool,g(s(sK0,sK11(X1,sK2,bool,sK0)))) | p(s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK7(sK4,sK4,sK1,sK0)))))) )),
  inference(superposition,[],[f4028,f143])).
fof(f4028,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK7(sK4,sK4,sK1,sK0)))))),
  inference(subsumption_resolution,[],[f4027,f140])).
fof(f4027,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK7(sK4,sK4,sK1,sK0)))))),
  inference(forward_demodulation,[],[f4025,f185])).
fof(f4025,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK7(sK4,sK4,sK1,sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4))),s(sK0,sK4))))),
  inference(duplicate_literal_removal,[],[f3996])).
fof(f3996,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK7(sK4,sK4,sK1,sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4))),s(sK0,sK4)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4))),s(sK0,sK4))))),
  inference(resolution,[],[f403,f197])).
fof(f6888,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4))),s(sK0,sK6(sK5)))) | $spl71),
  inference(cnf_transformation,[],[f6888_D])).
fof(f6888_D,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4))),s(sK0,sK6(sK5)))) <=> ~$spl71),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl71])])).
fof(f6814,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4)),sK2,bool,sK0)))) | s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4))),s(sK0,sK6(sK5)))) | $spl38),
  inference(superposition,[],[f1203,f6769])).
fof(f1203,plain,(
  ( ! [X1] : (s(bool,f(s(sK0,sK11(X1,sK2,bool,sK0)))) != s(bool,g(s(sK0,sK11(X1,sK2,bool,sK0)))) | s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK6(sK5))))) )),
  inference(superposition,[],[f803,f143])).
fof(f803,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK6(sK5))))),
  inference(resolution,[],[f787,f135])).
fof(f787,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK6(sK5)))))),
  inference(subsumption_resolution,[],[f768,f90])).
fof(f768,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK6(sK5))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5))),s(sK0,sK5))))),
  inference(resolution,[],[f93,f705])).
fof(f6886,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4))),s(sK0,sK6(sK5)))) | $spl68),
  inference(cnf_transformation,[],[f6886_D])).
fof(f6886_D,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4))),s(sK0,sK6(sK5)))) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f6813,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4)),sK2,bool,sK0)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4))),s(sK0,sK6(sK5)))) | $spl38),
  inference(superposition,[],[f1125,f6769])).
fof(f1125,plain,(
  ( ! [X1] : (s(bool,f(s(sK0,sK11(X1,sK2,bool,sK0)))) != s(bool,g(s(sK0,sK11(X1,sK2,bool,sK0)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK6(sK5))))) )),
  inference(superposition,[],[f801,f143])).
fof(f801,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK6(sK5))))),
  inference(resolution,[],[f787,f132])).
fof(f6884,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4))),s(sK0,sK6(sK4)))) | $spl66),
  inference(cnf_transformation,[],[f6884_D])).
fof(f6884_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4))),s(sK0,sK6(sK4)))) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f6812,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4)),sK2,bool,sK0)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4))),s(sK0,sK6(sK4)))) | $spl38),
  inference(superposition,[],[f830,f6769])).
fof(f830,plain,(
  ( ! [X1] : (s(bool,f(s(sK0,sK11(X1,sK2,bool,sK0)))) != s(bool,g(s(sK0,sK11(X1,sK2,bool,sK0)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK6(sK4))))) )),
  inference(superposition,[],[f731,f143])).
fof(f6882,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4))),s(sK0,sK6(sK4)))) | $spl65),
  inference(cnf_transformation,[],[f6882_D])).
fof(f6882_D,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4))),s(sK0,sK6(sK4)))) <=> ~$spl65),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl65])])).
fof(f6811,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4)),sK2,bool,sK0)))) | s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4))),s(sK0,sK6(sK4)))) | $spl38),
  inference(superposition,[],[f817,f6769])).
fof(f817,plain,(
  ( ! [X1] : (s(bool,f(s(sK0,sK11(X1,sK2,bool,sK0)))) != s(bool,g(s(sK0,sK11(X1,sK2,bool,sK0)))) | s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK6(sK4))))) )),
  inference(superposition,[],[f729,f143])).
fof(f729,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK6(sK4))))),
  inference(resolution,[],[f704,f131])).
fof(f6880,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4))),s(sK0,sK6(sK5))))) | $spl63),
  inference(cnf_transformation,[],[f6880_D])).
fof(f6880_D,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4))),s(sK0,sK6(sK5))))) <=> ~$spl63),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl63])])).
fof(f6810,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4)),sK2,bool,sK0)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4))),s(sK0,sK6(sK5))))) | $spl38),
  inference(superposition,[],[f806,f6769])).
fof(f806,plain,(
  ( ! [X1] : (s(bool,f(s(sK0,sK11(X1,sK2,bool,sK0)))) != s(bool,g(s(sK0,sK11(X1,sK2,bool,sK0)))) | ~p(s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK6(sK5)))))) )),
  inference(superposition,[],[f787,f143])).
fof(f5281,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4))),s(sK0,sK6(sK4))))) | $spl34),
  inference(cnf_transformation,[],[f5281_D])).
fof(f5281_D,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4))),s(sK0,sK6(sK4))))) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f6809,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4)),sK2,bool,sK0)))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4))),s(sK0,sK6(sK4))))) | $spl38),
  inference(superposition,[],[f734,f6769])).
fof(f734,plain,(
  ( ! [X1] : (s(bool,f(s(sK0,sK11(X1,sK2,bool,sK0)))) != s(bool,g(s(sK0,sK11(X1,sK2,bool,sK0)))) | p(s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK6(sK4)))))) )),
  inference(superposition,[],[f704,f143])).
fof(f6874,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK5)))) | $spl59),
  inference(cnf_transformation,[],[f6874_D])).
fof(f6874_D,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK5)))) <=> ~$spl59),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl59])])).
fof(f6876,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK4)))) | $spl60),
  inference(cnf_transformation,[],[f6876_D])).
fof(f6876_D,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK4)))) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f6804,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4)),sK2,bool,sK0)))) | p(s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK4)))) | ~p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK5)))) | $spl38),
  inference(superposition,[],[f545,f6769])).
fof(f545,plain,(
  ( ! [X0,X1] : (s(bool,f(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0)),X1,bool,sK0)))) != s(bool,g(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0)),X1,bool,sK0)))) | p(s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,X0)))) | ~p(s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK5))))) )),
  inference(superposition,[],[f91,f143])).
fof(f6746,plain,(
  ( ! [X23] : (~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X23))),s(sK0,X23)))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK7(X23,sK5,sK1,sK0))))),s(sK0,sK5))))) ) | $spl48),
  inference(cnf_transformation,[],[f6746_D])).
fof(f6746_D,plain,(
  ( ! [X23] : (~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X23))),s(sK0,X23)))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK7(X23,sK5,sK1,sK0))))),s(sK0,sK5))))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f6702,plain,(
  ( ! [X26] : (~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X26))),s(sK0,X26)))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK7(X26,sK5,sK1,sK0))))),s(sK0,sK5)))) | p(s(bool,g(s(sK0,sK11(sK3,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5)),bool,sK0))))) | p(s(bool,f(s(sK0,sK11(sK3,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5)),bool,sK0)))))) )),
  inference(resolution,[],[f708,f2937])).
fof(f708,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X1))),s(sK0,X1)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,X0)))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK7(X0,X1,sK1,sK0))))),s(sK0,X1))))) )),
  inference(subsumption_resolution,[],[f707,f108])).
fof(f707,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK7(X0,X1,sK1,sK0))))),s(sK0,X1)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,X0)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X1))),s(sK0,X1))))) )),
  inference(forward_demodulation,[],[f706,f155])).
fof(f706,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK7(X0,X1,sK1,sK0))))),s(sK0,X1)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,X0)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X1))),s(sK0,X1)))) | ~p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),sK1))))) )),
  inference(subsumption_resolution,[],[f683,f94])).
fof(f683,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK7(X0,X1,sK1,sK0)))),s(sK0,sK7(X0,X1,sK1,sK0))))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK7(X0,X1,sK1,sK0))))),s(sK0,X1)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,X0)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X1))),s(sK0,X1)))) | ~p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),sK1))))) )),
  inference(resolution,[],[f92,f95])).
fof(f6701,plain,(
  ( ! [X25] : (~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X25))),s(sK0,X25)))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK7(X25,sK5,sK1,sK0))))),s(sK0,sK5)))) | p(s(bool,g(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5)),sK3,bool,sK0))))) | p(s(bool,f(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5)),sK3,bool,sK0)))))) )),
  inference(resolution,[],[f708,f3126])).
fof(f6700,plain,(
  ( ! [X24] : (~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X24))),s(sK0,X24)))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK7(X24,sK5,sK1,sK0))))),s(sK0,sK5)))) | p(s(bool,g(s(sK0,sK11(sK3,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5)),bool,sK0))))) | s(bool,'T_0') = s(bool,f(s(sK0,sK11(sK3,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5)),bool,sK0))))) )),
  inference(resolution,[],[f708,f2939])).
fof(f6699,plain,(
  ( ! [X23] : (~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X23))),s(sK0,X23)))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK7(X23,sK5,sK1,sK0))))),s(sK0,sK5)))) | p(s(bool,g(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5)),sK3,bool,sK0))))) | s(bool,'T_0') = s(bool,f(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5)),sK3,bool,sK0))))) )),
  inference(resolution,[],[f708,f3128])).
fof(f6737,plain,(
  p(s(bool,f(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4)),sK2,bool,sK0))))) | $spl42),
  inference(cnf_transformation,[],[f6737_D])).
fof(f6737_D,plain,(
  p(s(bool,f(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4)),sK2,bool,sK0))))) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f6734,plain,(
  p(s(bool,g(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4)),sK2,bool,sK0))))) | $spl38),
  inference(cnf_transformation,[],[f6734_D])).
fof(f6734_D,plain,(
  p(s(bool,g(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4)),sK2,bool,sK0))))) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f6735,plain,(
  ( ! [X19] : (~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X19))),s(sK0,X19)))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK7(X19,sK4,sK1,sK0))))),s(sK0,sK4))))) ) | $spl40),
  inference(cnf_transformation,[],[f6735_D])).
fof(f6735_D,plain,(
  ( ! [X19] : (~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X19))),s(sK0,X19)))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK7(X19,sK4,sK1,sK0))))),s(sK0,sK4))))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f6697,plain,(
  ( ! [X21] : (~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X21))),s(sK0,X21)))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK7(X21,sK4,sK1,sK0))))),s(sK0,sK4)))) | p(s(bool,g(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4)),sK2,bool,sK0))))) | p(s(bool,f(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4)),sK2,bool,sK0)))))) )),
  inference(resolution,[],[f708,f2452])).
fof(f2452,plain,(
  ( ! [X5] : (p(s(bool,happ(s(fun(sK0,bool),X5),s(sK0,sK4)))) | p(s(bool,g(s(sK0,sK11(X5,sK2,bool,sK0))))) | p(s(bool,f(s(sK0,sK11(X5,sK2,bool,sK0)))))) )),
  inference(trivial_inequality_removal,[],[f2451])).
fof(f2451,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,'F_0') | p(s(bool,happ(s(fun(sK0,bool),X5),s(sK0,sK4)))) | p(s(bool,g(s(sK0,sK11(X5,sK2,bool,sK0))))) | p(s(bool,f(s(sK0,sK11(X5,sK2,bool,sK0)))))) )),
  inference(superposition,[],[f1365,f132])).
fof(f1365,plain,(
  ( ! [X4] : (s(bool,'F_0') != s(bool,f(s(sK0,sK11(X4,sK2,bool,sK0)))) | p(s(bool,happ(s(fun(sK0,bool),X4),s(sK0,sK4)))) | p(s(bool,g(s(sK0,sK11(X4,sK2,bool,sK0)))))) )),
  inference(superposition,[],[f427,f134])).
fof(f427,plain,(
  ( ! [X1] : (s(bool,f(s(sK0,sK11(X1,sK2,bool,sK0)))) != s(bool,g(s(sK0,sK11(X1,sK2,bool,sK0)))) | p(s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK4))))) )),
  inference(superposition,[],[f377,f143])).
fof(f377,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK4))))),
  inference(resolution,[],[f89,f88])).
fof(f6732,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4)),sK2,bool,sK0)))) | $spl36),
  inference(cnf_transformation,[],[f6732_D])).
fof(f6732_D,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4)),sK2,bool,sK0)))) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f6695,plain,(
  ( ! [X19] : (~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X19))),s(sK0,X19)))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK7(X19,sK4,sK1,sK0))))),s(sK0,sK4)))) | p(s(bool,g(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4)),sK2,bool,sK0))))) | s(bool,'T_0') = s(bool,f(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4)),sK2,bool,sK0))))) )),
  inference(resolution,[],[f708,f2454])).
fof(f2454,plain,(
  ( ! [X1] : (p(s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK4)))) | p(s(bool,g(s(sK0,sK11(X1,sK2,bool,sK0))))) | s(bool,'T_0') = s(bool,f(s(sK0,sK11(X1,sK2,bool,sK0))))) )),
  inference(trivial_inequality_removal,[],[f2447])).
fof(f2447,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,'F_0') | p(s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK4)))) | p(s(bool,g(s(sK0,sK11(X1,sK2,bool,sK0))))) | s(bool,'T_0') = s(bool,f(s(sK0,sK11(X1,sK2,bool,sK0))))) )),
  inference(superposition,[],[f1365,f142])).
fof(f2022,plain,(
  ~p(s(bool,g(s(sK0,sK11(sK2,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4)),bool,sK0))))) | $spl19),
  inference(cnf_transformation,[],[f2022_D])).
fof(f2022_D,plain,(
  ~p(s(bool,g(s(sK0,sK11(sK2,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4)),bool,sK0))))) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f5279,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4))),s(sK0,sK6(sK4))))) | ~p(s(bool,g(s(sK0,sK11(sK2,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4)),bool,sK0))))) | $spl16),
  inference(trivial_inequality_removal,[],[f5268])).
fof(f5268,plain,(
  s(bool,'T_0') != s(bool,'T_0') | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4))),s(sK0,sK6(sK4))))) | ~p(s(bool,g(s(sK0,sK11(sK2,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4)),bool,sK0))))) | $spl16),
  inference(superposition,[],[f1517,f2038])).
fof(f2038,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK11(sK2,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4)),bool,sK0)))) | $spl16),
  inference(resolution,[],[f2021,f136])).
fof(f1517,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,f(s(sK0,sK11(sK2,X2,bool,sK0)))) | p(s(bool,happ(s(fun(sK0,bool),X2),s(sK0,sK6(sK4))))) | ~p(s(bool,g(s(sK0,sK11(sK2,X2,bool,sK0)))))) )),
  inference(superposition,[],[f733,f138])).
fof(f138,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f85])).
fof(f85,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f43])).
fof(f43,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X12] : (s(bool,'T_0') = s(bool,X12) <=> p(s(bool,X12)))),
  file('thm.nets.DORDER_LEMMA','thm.bool.EQ_CLAUSES|split|0')).
fof(f733,plain,(
  ( ! [X0] : (s(bool,f(s(sK0,sK11(sK2,X0,bool,sK0)))) != s(bool,g(s(sK0,sK11(sK2,X0,bool,sK0)))) | p(s(bool,happ(s(fun(sK0,bool),X0),s(sK0,sK6(sK4)))))) )),
  inference(superposition,[],[f704,f143])).
fof(f3234,plain,(
  p(s(bool,f(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK4))),sK3,bool,sK0))))) | $spl30),
  inference(cnf_transformation,[],[f3234_D])).
fof(f3234_D,plain,(
  p(s(bool,f(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK4))),sK3,bool,sK0))))) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f3236,plain,(
  p(s(bool,g(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK4))),sK3,bool,sK0))))) | $spl32),
  inference(cnf_transformation,[],[f3236_D])).
fof(f3236_D,plain,(
  p(s(bool,g(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK4))),sK3,bool,sK0))))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f3232,plain,(
  p(s(bool,g(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK4))),sK3,bool,sK0))))) | p(s(bool,f(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK4))),sK3,bool,sK0)))))),
  inference(subsumption_resolution,[],[f3223,f124])).
fof(f124,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | ~p(s(bool,X0)) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,'F_0')) & p(s(bool,X0)))))),
  inference(flattening,[],[f76])).
fof(f76,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) | ~p(s(bool,X0))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,'F_0')) & p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f35])).
fof(f35,plain,(
  ! [X0] : ((p(s(bool,'F_0')) & p(s(bool,X0))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X12] : ((p(s(bool,'F_0')) & p(s(bool,X12))) <=> p(s(bool,'F_0')))),
  file('thm.nets.DORDER_LEMMA','thm.bool.AND_CLAUSES|split|2')).
fof(f3223,plain,(
  p(s(bool,'F_0')) | p(s(bool,g(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK4))),sK3,bool,sK0))))) | p(s(bool,f(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK4))),sK3,bool,sK0)))))),
  inference(superposition,[],[f3126,f2725])).
fof(f2725,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK4)))),s(sK0,sK5)))),
  inference(subsumption_resolution,[],[f2675,f139])).
fof(f2675,plain,(
  p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK4)))),s(sK0,sK5)))),
  inference(superposition,[],[f539,f2666])).
fof(f539,plain,(
  ( ! [X10] : (p(s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,X10)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X10))),s(sK0,sK5)))) )),
  inference(resolution,[],[f91,f132])).
fof(f2975,plain,(
  p(s(bool,f(s(sK0,sK11(sK3,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK4))),bool,sK0))))) | $spl26),
  inference(cnf_transformation,[],[f2975_D])).
fof(f2975_D,plain,(
  p(s(bool,f(s(sK0,sK11(sK3,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK4))),bool,sK0))))) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f2977,plain,(
  p(s(bool,g(s(sK0,sK11(sK3,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK4))),bool,sK0))))) | $spl28),
  inference(cnf_transformation,[],[f2977_D])).
fof(f2977_D,plain,(
  p(s(bool,g(s(sK0,sK11(sK3,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK4))),bool,sK0))))) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f2973,plain,(
  p(s(bool,g(s(sK0,sK11(sK3,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK4))),bool,sK0))))) | p(s(bool,f(s(sK0,sK11(sK3,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK4))),bool,sK0)))))),
  inference(subsumption_resolution,[],[f2964,f124])).
fof(f2964,plain,(
  p(s(bool,'F_0')) | p(s(bool,g(s(sK0,sK11(sK3,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK4))),bool,sK0))))) | p(s(bool,f(s(sK0,sK11(sK3,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK4))),bool,sK0)))))),
  inference(superposition,[],[f2937,f2725])).
fof(f2490,plain,(
  p(s(bool,f(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK5))),sK2,bool,sK0))))) | $spl22),
  inference(cnf_transformation,[],[f2490_D])).
fof(f2490_D,plain,(
  p(s(bool,f(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK5))),sK2,bool,sK0))))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f2492,plain,(
  p(s(bool,g(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK5))),sK2,bool,sK0))))) | $spl24),
  inference(cnf_transformation,[],[f2492_D])).
fof(f2492_D,plain,(
  p(s(bool,g(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK5))),sK2,bool,sK0))))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f2488,plain,(
  p(s(bool,g(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK5))),sK2,bool,sK0))))) | p(s(bool,f(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK5))),sK2,bool,sK0)))))),
  inference(subsumption_resolution,[],[f2478,f124])).
fof(f2478,plain,(
  p(s(bool,'F_0')) | p(s(bool,g(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK5))),sK2,bool,sK0))))) | p(s(bool,f(s(sK0,sK11(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK5))),sK2,bool,sK0)))))),
  inference(superposition,[],[f2452,f1393])).
fof(f1393,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK5)))),s(sK0,sK4)))),
  inference(subsumption_resolution,[],[f1386,f139])).
fof(f1386,plain,(
  p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK5)))),s(sK0,sK4)))),
  inference(superposition,[],[f387,f801])).
fof(f387,plain,(
  ( ! [X10] : (p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,X10)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X10))),s(sK0,sK4)))) )),
  inference(resolution,[],[f89,f132])).
fof(f2021,plain,(
  p(s(bool,f(s(sK0,sK11(sK2,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4)),bool,sK0))))) | $spl16),
  inference(cnf_transformation,[],[f2021_D])).
fof(f2021_D,plain,(
  p(s(bool,f(s(sK0,sK11(sK2,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4)),bool,sK0))))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f2023,plain,(
  p(s(bool,g(s(sK0,sK11(sK2,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4)),bool,sK0))))) | $spl18),
  inference(cnf_transformation,[],[f2023_D])).
fof(f2023_D,plain,(
  p(s(bool,g(s(sK0,sK11(sK2,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4)),bool,sK0))))) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f2025,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK4)))),s(sK0,sK4))) | $spl20),
  inference(cnf_transformation,[],[f2025_D])).
fof(f2025_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK4)))),s(sK0,sK4))) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f1993,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK4)))),s(sK0,sK4))) | p(s(bool,g(s(sK0,sK11(sK2,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4)),bool,sK0))))) | p(s(bool,f(s(sK0,sK11(sK2,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4)),bool,sK0)))))),
  inference(resolution,[],[f692,f1756])).
fof(f1756,plain,(
  ( ! [X5] : (p(s(bool,happ(s(fun(sK0,bool),X5),s(sK0,sK4)))) | p(s(bool,g(s(sK0,sK11(sK2,X5,bool,sK0))))) | p(s(bool,f(s(sK0,sK11(sK2,X5,bool,sK0)))))) )),
  inference(trivial_inequality_removal,[],[f1755])).
fof(f1755,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,'F_0') | p(s(bool,happ(s(fun(sK0,bool),X5),s(sK0,sK4)))) | p(s(bool,g(s(sK0,sK11(sK2,X5,bool,sK0))))) | p(s(bool,f(s(sK0,sK11(sK2,X5,bool,sK0)))))) )),
  inference(superposition,[],[f1359,f132])).
fof(f1359,plain,(
  ( ! [X4] : (s(bool,'F_0') != s(bool,f(s(sK0,sK11(sK2,X4,bool,sK0)))) | p(s(bool,happ(s(fun(sK0,bool),X4),s(sK0,sK4)))) | p(s(bool,g(s(sK0,sK11(sK2,X4,bool,sK0)))))) )),
  inference(superposition,[],[f426,f134])).
fof(f426,plain,(
  ( ! [X0] : (s(bool,f(s(sK0,sK11(sK2,X0,bool,sK0)))) != s(bool,g(s(sK0,sK11(sK2,X0,bool,sK0)))) | p(s(bool,happ(s(fun(sK0,bool),X0),s(sK0,sK4))))) )),
  inference(superposition,[],[f377,f143])).
fof(f692,plain,(
  ( ! [X11] : (~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X11))),s(sK0,X11)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(X11)))),s(sK0,X11)))) )),
  inference(resolution,[],[f92,f136])).
fof(f1945,plain,(
  s(fun(sK0,bool),sK2) = s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK5)))) | $spl14),
  inference(cnf_transformation,[],[f1945_D])).
fof(f1945_D,plain,(
  s(fun(sK0,bool),sK2) = s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK5)))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f1922,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK11(sK2,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK5))),bool,sK0)))) | $spl7),
  inference(cnf_transformation,[],[f1922_D])).
fof(f1922_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK11(sK2,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK5))),bool,sK0)))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f1890,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK11(sK2,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK5))),bool,sK0)))) | s(fun(sK0,bool),sK2) = s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK5)))) | $spl2),
  inference(superposition,[],[f143,f1821])).
fof(f1821,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK11(sK2,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK5))),bool,sK0)))) | $spl2),
  inference(resolution,[],[f1794,f136])).
fof(f1928,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK5)))),s(sK0,sK6(sK4)))) | $spl12),
  inference(cnf_transformation,[],[f1928_D])).
fof(f1928_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK5)))),s(sK0,sK6(sK4)))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f1862,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK11(sK2,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK5))),bool,sK0)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK5)))),s(sK0,sK6(sK4)))) | $spl2),
  inference(superposition,[],[f829,f1821])).
fof(f829,plain,(
  ( ! [X0] : (s(bool,f(s(sK0,sK11(sK2,X0,bool,sK0)))) != s(bool,g(s(sK0,sK11(sK2,X0,bool,sK0)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),X0),s(sK0,sK6(sK4))))) )),
  inference(superposition,[],[f731,f143])).
fof(f1926,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK5)))),s(sK0,sK6(sK4)))) | $spl11),
  inference(cnf_transformation,[],[f1926_D])).
fof(f1926_D,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK5)))),s(sK0,sK6(sK4)))) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f1861,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK11(sK2,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK5))),bool,sK0)))) | s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK5)))),s(sK0,sK6(sK4)))) | $spl2),
  inference(superposition,[],[f816,f1821])).
fof(f816,plain,(
  ( ! [X0] : (s(bool,f(s(sK0,sK11(sK2,X0,bool,sK0)))) != s(bool,g(s(sK0,sK11(sK2,X0,bool,sK0)))) | s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),X0),s(sK0,sK6(sK4))))) )),
  inference(superposition,[],[f729,f143])).
fof(f1924,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK5)))),s(sK0,sK6(sK5))))) | $spl9),
  inference(cnf_transformation,[],[f1924_D])).
fof(f1924_D,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK5)))),s(sK0,sK6(sK5))))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f1860,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK11(sK2,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK5))),bool,sK0)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK5)))),s(sK0,sK6(sK5))))) | $spl2),
  inference(superposition,[],[f805,f1821])).
fof(f805,plain,(
  ( ! [X0] : (s(bool,f(s(sK0,sK11(sK2,X0,bool,sK0)))) != s(bool,g(s(sK0,sK11(sK2,X0,bool,sK0)))) | ~p(s(bool,happ(s(fun(sK0,bool),X0),s(sK0,sK6(sK5)))))) )),
  inference(superposition,[],[f787,f143])).
fof(f1920,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK5)))),s(sK0,sK6(sK4))))) | $spl4),
  inference(cnf_transformation,[],[f1920_D])).
fof(f1920_D,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK5)))),s(sK0,sK6(sK4))))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f1859,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK11(sK2,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK5))),bool,sK0)))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK5)))),s(sK0,sK6(sK4))))) | $spl2),
  inference(superposition,[],[f733,f1821])).
fof(f1792,plain,(
  p(s(bool,f(s(sK0,sK11(sK2,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK5))),bool,sK0))))) | $spl0),
  inference(cnf_transformation,[],[f1792_D])).
fof(f1792_D,plain,(
  p(s(bool,f(s(sK0,sK11(sK2,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK5))),bool,sK0))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f1794,plain,(
  p(s(bool,g(s(sK0,sK11(sK2,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK5))),bool,sK0))))) | $spl2),
  inference(cnf_transformation,[],[f1794_D])).
fof(f1794_D,plain,(
  p(s(bool,g(s(sK0,sK11(sK2,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK5))),bool,sK0))))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f1790,plain,(
  p(s(bool,g(s(sK0,sK11(sK2,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK5))),bool,sK0))))) | p(s(bool,f(s(sK0,sK11(sK2,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK5))),bool,sK0)))))),
  inference(subsumption_resolution,[],[f1781,f124])).
fof(f1781,plain,(
  p(s(bool,'F_0')) | p(s(bool,g(s(sK0,sK11(sK2,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK5))),bool,sK0))))) | p(s(bool,f(s(sK0,sK11(sK2,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK5))),bool,sK0)))))),
  inference(superposition,[],[f1756,f1393])).
% SZS output end Proof for thm.nets
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 333
% Passive clauses: 2611
% Generated clauses: 10212
% Final active clauses: 241
% Final passive clauses: 2050
% Input formulas: 21
% Initial clauses: 57
% 
% Selected by SInE selection: 21
% SInE iterations: 1
% 
% Duplicate literals: 205
% Trivial inequalities: 324
% Fw subsumption resolutions: 1692
% Fw demodulations: 1329
% Bw demodulations: 90
% 
% Simple tautologies: 275
% Forward subsumptions: 3950
% 
% Binary resolution: 1616
% Forward superposition: 2334
% Backward superposition: 2590
% 
% Split clauses: 80
% Split components: 174
% SAT solver clauses: 89
% SAT solver unit clauses: 8
% SAT solver binary clauses: 18
% 
% Sat splits: 80
% Sat splitting refutations: 8
% 
% Memory used [KB]: 4605
% Time elapsed: 0.254 s
% ------------------------------
% ------------------------------
% Success in time 11.929 s
