Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.pred_set
% (83359)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 156
% Passive clauses: 2929
% Generated clauses: 25854
% Final active clauses: 130
% Final passive clauses: 2417
% Input formulas: 15
% Initial clauses: 28
% 
% Duplicate literals: 1069
% Trivial inequalities: 308
% Fw subsumption resolutions: 4014
% Fw demodulations: 3931
% Bw demodulations: 23
% 
% Simple tautologies: 484
% Equational tautologies: 52
% Forward subsumptions: 11938
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 623
% Factoring: 4
% Forward superposition: 8700
% Backward superposition: 7159
% Self superposition: 223
% Equality factoring: 124
% Equality resolution: 37
% 
% Split clauses: 609
% Split components: 1231
% SAT solver clauses: 1428
% SAT solver unit clauses: 22
% SAT solver binary clauses: 176
% 
% Sat splits: 609
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 1425
% TWLsolver calls for satisfiability: 183
% 
% Memory used [KB]: 10618
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.pred_set
% (83369)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 143
% Passive clauses: 764
% Generated clauses: 16386
% Final active clauses: 127
% Final passive clauses: 604
% Input formulas: 15
% Initial clauses: 28
% 
% Duplicate literals: 566
% Trivial inequalities: 237
% Fw subsumption resolutions: 2006
% Fw demodulations: 1991
% Bw demodulations: 24
% Condensations: 1253
% Global subsumptions: 1
% 
% Simple tautologies: 277
% Equational tautologies: 40
% Forward subsumptions: 7836
% 
% Binary resolution: 411
% Factoring: 6
% Forward superposition: 4289
% Backward superposition: 4454
% Self superposition: 149
% Equality factoring: 94
% Equality resolution: 26
% 
% Split clauses: 92
% Split components: 187
% SAT solver clauses: 1310
% SAT solver unit clauses: 29
% SAT solver binary clauses: 41
% SAT solver learnt clauses: 2
% 
% Sat splits: 92
% Sat splitting refutations: 14
% 
% TWLsolver clauses: 1306
% TWLsolver calls for satisfiability: 240
% 
% Memory used [KB]: 12153
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f5591,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f59,f88,f58,f90,f86,f57,f91,f85,f1275,f1277,f1291,f1292,f1305,f1309,f1311,f1313,f1315,f1317,f1319,f1323,f1325,f1327,f1329,f1331,f1333,f1335,f1337,f1339,f1344,f1407,f1454,f1456,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1470,f1471,f1472,f1796,f1799,f1800,f1798,f1789,f1801,f1803,f1804,f1782,f1806,f1808,f1735,f1815,f1737,f1818,f1805,f1922,f1924,f1926,f1930,f1933,f2571,f2582,f2570,f2583,f2585,f2586,f2560,f2588,f2590,f2552,f2497,f2596,f2502,f2597,f2503,f1813,f2504,f1816,f2599,f2600,f2601,f2603,f2587,f2614,f1276,f2594,f3012,f3021,f3013,f3014,f3015,f3025,f3018,f3016,f3024,f3019,f3022,f3128,f3131,f3150,f3148,f3278,f3288,f3277,f3289,f1797,f3315,f3338,f3660,f5584,f5586])).
fof(f5586,plain,(
  $false | ($spl1 | $spl22)),
  inference(subsumption_resolution,[],[f5585,f61])).
fof(f61,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f19])).
fof(f19,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.pred_set.IN_DISJOINT',aHL_FALSITY)).
fof(f5585,plain,(
  p(s(bool,'F_0')) | ($spl1 | $spl22)),
  inference(backward_demodulation,[],[f5559,f5546])).
fof(f5546,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1)))) | $spl1),
  inference(subsumption_resolution,[],[f5506,f60])).
fof(f60,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pred_set.IN_DISJOINT',aHL_TRUTH)).
fof(f5506,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1)))) | $spl1),
  inference(superposition,[],[f77,f4438])).
fof(f4438,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl1),
  inference(unit_resulting_resolution,[],[f3663,f110,f2691])).
fof(f2691,plain,(
  ( ! [X14,X15,X13] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(X13,sK4(X14,X15,X13)),s(fun(X13,bool),X14))) | s(fun(X13,bool),X14) = s(fun(X13,bool),X15) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X13,sK4(X14,X15,X13)),s(fun(X13,bool),X15)))) )),
  inference(superposition,[],[f74,f62])).
fof(f62,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.pred_set.IN_DISJOINT',aHL_BOOL_CASES)).
fof(f74,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) )),
  inference(cnf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f51])).
fof(f51,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(rectify,[],[f50])).
fof(f50,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(nnf_transformation,[],[f30])).
fof(f30,plain,(
  ! [X0,X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(flattening,[],[f29])).
fof(f29,plain,(
  ! [X0] : ! [X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X6] : ! [X7,X8] : (s(fun(X6,bool),X7) = s(fun(X6,bool),X8) <=> ! [X9] : s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X7))) = s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X8))))),
  file('thm.pred_set.IN_DISJOINT','thm.pred_set.EXTENSION')).
fof(f110,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f72,f69])).
fof(f69,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f48])).
fof(f48,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f24])).
fof(f24,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f23])).
fof(f23,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X8] : (s(bool,'F_0') = s(bool,X8) <=> ~p(s(bool,X8)))),
  file('thm.pred_set.IN_DISJOINT','thm.bool.EQ_CLAUSES|split|3')).
fof(f72,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(cnf_transformation,[],[f28])).
fof(f28,plain,(
  ! [X0,X1] : ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))),
  inference(flattening,[],[f27])).
fof(f27,plain,(
  ! [X0] : ! [X1] : ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X6] : ! [X9] : ~p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),'const.pred_set.EMPTY_0'))))),
  file('thm.pred_set.IN_DISJOINT','thm.pred_set.NOT_IN_EMPTY')).
fof(f3663,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | $spl1),
  inference(unit_resulting_resolution,[],[f86,f76])).
fof(f76,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),'const.pred_set.EMPTY_0') != s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) & (s(fun(X0,bool),'const.pred_set.EMPTY_0') != s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(nnf_transformation,[],[f32])).
fof(f32,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))),
  inference(flattening,[],[f31])).
fof(f31,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X6] : ! [X7,X8] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))) <=> s(fun(X6,bool),'const.pred_set.EMPTY_0') = s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8))))),
  file('thm.pred_set.IN_DISJOINT','thm.pred_set.DISJOINT_DEF')).
fof(f77,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))) & (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))))),
  inference(flattening,[],[f54])).
fof(f54,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))) & ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))))),
  inference(nnf_transformation,[],[f34])).
fof(f34,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(flattening,[],[f33])).
fof(f33,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X6] : ! [X7,X8,X9] : (p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X7)))) & p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X8))))))),
  file('thm.pred_set.IN_DISJOINT','thm.pred_set.IN_INTER')).
fof(f5559,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | ($spl1 | $spl22)),
  inference(unit_resulting_resolution,[],[f5545,f1805])).
fof(f5545,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2)))) | $spl1),
  inference(subsumption_resolution,[],[f5505,f60])).
fof(f5505,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2)))) | $spl1),
  inference(superposition,[],[f78,f4438])).
fof(f78,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f55])).
fof(f5584,plain,(
  $false | ($spl1 | $spl6)),
  inference(subsumption_resolution,[],[f5560,f5546])).
fof(f5560,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1)))) | ($spl1 | $spl6)),
  inference(unit_resulting_resolution,[],[f5545,f91])).
fof(f3660,plain,(
  $false | ($spl0 | $spl2 | $spl4)),
  inference(subsumption_resolution,[],[f3659,f60])).
fof(f3659,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl2 | $spl4)),
  inference(forward_demodulation,[],[f3658,f3328])).
fof(f3328,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),sK1))) | $spl4),
  inference(unit_resulting_resolution,[],[f90,f65])).
fof(f65,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f46])).
fof(f46,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f21])).
fof(f21,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X8] : (s(bool,'T_0') = s(bool,X8) <=> p(s(bool,X8)))),
  file('thm.pred_set.IN_DISJOINT','thm.bool.EQ_CLAUSES|split|1')).
fof(f3658,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),sK1)))) | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f3651,f60])).
fof(f3651,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),sK1)))) | ($spl0 | $spl2)),
  inference(superposition,[],[f3581,f3305])).
fof(f3305,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))) | $spl2),
  inference(unit_resulting_resolution,[],[f88,f65])).
fof(f3581,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK2)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK1))))) ) | $spl0),
  inference(subsumption_resolution,[],[f3580,f61])).
fof(f3580,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK2)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK1))))) ) | $spl0),
  inference(forward_demodulation,[],[f3554,f110])).
fof(f3554,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK2)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK1))))) ) | $spl0),
  inference(superposition,[],[f79,f3343])).
fof(f3343,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | $spl0),
  inference(unit_resulting_resolution,[],[f85,f75])).
fof(f75,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) )),
  inference(cnf_transformation,[],[f53])).
fof(f79,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f55])).
fof(f3338,plain,(
  $false | ($spl2 | $spl4 | $spl34)),
  inference(subsumption_resolution,[],[f3337,f92])).
fof(f92,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f61,f64])).
fof(f64,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f46])).
fof(f3337,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl2 | $spl4 | $spl34)),
  inference(forward_demodulation,[],[f3320,f3305])).
fof(f3320,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))) | ($spl4 | $spl34)),
  inference(unit_resulting_resolution,[],[f90,f2587])).
fof(f3315,plain,(
  $false | ($spl2 | $spl4 | $spl22)),
  inference(subsumption_resolution,[],[f3314,f61])).
fof(f3314,plain,(
  p(s(bool,'F_0')) | ($spl2 | $spl4 | $spl22)),
  inference(backward_demodulation,[],[f3297,f90])).
fof(f3297,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),sK1))) | ($spl2 | $spl22)),
  inference(unit_resulting_resolution,[],[f88,f1805])).
fof(f1797,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl14),
  inference(cnf_transformation,[],[f1797_D])).
fof(f1797_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f3289,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK0,X32),s(fun(sK0,bool),sK1))),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X33),s(fun(bool,bool),X34)))))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X32),s(fun(sK0,bool),sK2))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X33))))) ) | $spl62),
  inference(cnf_transformation,[],[f3289_D])).
fof(f3289_D,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK0,X32),s(fun(sK0,bool),sK1))),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X33),s(fun(bool,bool),X34)))))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X32),s(fun(sK0,bool),sK2))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X33))))) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f3277,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK0,X32),s(fun(sK0,bool),sK1))),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X33),s(fun(bool,bool),X34)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X33)))) | s(bool,'T_0') = s(bool,X31) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X32),s(fun(sK0,bool),sK2)))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f3249])).
fof(f3249,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK0,X32),s(fun(sK0,bool),sK1))),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X33),s(fun(bool,bool),X34)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X33)))) | s(bool,'T_0') = s(bool,X31) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X32),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X31)) ) | $spl6),
  inference(superposition,[],[f632,f1270])).
fof(f1270,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK2))) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X0)) ) | $spl6),
  inference(condensation,[],[f1269])).
fof(f1269,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK2))) | s(bool,X2) = s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X2)) ) | $spl6),
  inference(duplicate_literal_removal,[],[f1000])).
fof(f1000,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK2))) | s(bool,X2) = s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) ) | $spl6),
  inference(superposition,[],[f654,f62])).
fof(f654,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) = s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK2))) | s(bool,X2) = s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X2)) ) | $spl6),
  inference(subsumption_resolution,[],[f648,f60])).
fof(f648,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s(bool,X1) = s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X1) | s(bool,X2) = s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X2)) ) | $spl6),
  inference(superposition,[],[f619,f118])).
fof(f118,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f62,f62])).
fof(f619,plain,(
  ( ! [X37,X38] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X37),s(fun(sK0,bool),sK1)))) | s(bool,X38) = s(bool,'const.bool.IN_2'(s(sK0,X37),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X38)) ) | $spl6),
  inference(subsumption_resolution,[],[f422,f60])).
fof(f422,plain,(
  ( ! [X37,X38] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X37),s(fun(sK0,bool),sK1)))) | s(bool,X38) = s(bool,'const.bool.IN_2'(s(sK0,X37),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X38)) ) | $spl6),
  inference(superposition,[],[f91,f118])).
fof(f632,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,'const.bool.IN_2'(s(bool,X30),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X31),s(fun(bool,bool),X32)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X31)))) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f77,f62])).
fof(f3288,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK0,X28),s(fun(sK0,bool),sK2))),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X29),s(fun(bool,bool),X30)))))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X28),s(fun(sK0,bool),sK1))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X29))))) ) | $spl60),
  inference(cnf_transformation,[],[f3288_D])).
fof(f3288_D,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK0,X28),s(fun(sK0,bool),sK2))),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X29),s(fun(bool,bool),X30)))))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X28),s(fun(sK0,bool),sK1))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X29))))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f3278,plain,(
  ( ! [X30,X28,X29,X27] : (~p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK0,X28),s(fun(sK0,bool),sK2))),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X29),s(fun(bool,bool),X30)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X29)))) | s(bool,'T_0') = s(bool,X27) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X28),s(fun(sK0,bool),sK1)))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f3248])).
fof(f3248,plain,(
  ( ! [X30,X28,X29,X27] : (~p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK0,X28),s(fun(sK0,bool),sK2))),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X29),s(fun(bool,bool),X30)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X29)))) | s(bool,'T_0') = s(bool,X27) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X28),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X27)) ) | $spl6),
  inference(superposition,[],[f632,f1250])).
fof(f1250,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK2))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X0)) ) | $spl6),
  inference(condensation,[],[f1249])).
fof(f1249,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK1))) | s(bool,X2) = s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) ) | $spl6),
  inference(duplicate_literal_removal,[],[f1026])).
fof(f1026,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK1))) | s(bool,X2) = s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl6),
  inference(superposition,[],[f654,f62])).
fof(f3148,plain,(
  ( ! [X4,X2,X3] : (s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,X4) | s(bool,X4) = s(bool,sK5(X2,X3,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X2,X3,bool,bool))))) ) | $spl56),
  inference(cnf_transformation,[],[f3148_D])).
fof(f3148_D,plain,(
  ( ! [X4,X2,X3] : (s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,X4) | s(bool,X4) = s(bool,sK5(X2,X3,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X2,X3,bool,bool))))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f3150,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | $spl59),
  inference(cnf_transformation,[],[f3150_D])).
fof(f3150_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl59),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl59])])).
fof(f3131,plain,(
  ( ! [X12,X13,X11] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X11) = s(fun(bool,bool),X12) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X11,X12,bool,bool)))) | s(bool,X13) = s(bool,sK5(X11,X12,bool,bool)) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f183,f118])).
fof(f183,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK5(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK5(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f84,f62])).
fof(f84,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK5(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK5(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f56])).
fof(f56,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK5(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK5(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5])],[f42])).
fof(f42,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f37])).
fof(f37,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f36])).
fof(f36,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.pred_set.IN_DISJOINT',aHL_EXT)).
fof(f3128,plain,(
  ( ! [X4,X2,X3] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X2,X3,bool,bool)))) | s(bool,X4) = s(bool,sK5(X2,X3,bool,bool)) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f183,f118])).
fof(f3022,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) ) | $spl50),
  inference(cnf_transformation,[],[f3022_D])).
fof(f3022_D,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f3019,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK5(X0,X1,bool,bool))) ) | $spl46),
  inference(cnf_transformation,[],[f3019_D])).
fof(f3019_D,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK5(X0,X1,bool,bool))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f3024,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | $spl53),
  inference(cnf_transformation,[],[f3024_D])).
fof(f3024_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl53),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl53])])).
fof(f3016,plain,(
  ( ! [X12,X13,X11] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X11) = s(fun(bool,bool),X12) | s(bool,'T_0') = s(bool,sK5(X11,X12,bool,bool)) | s(bool,X13) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f182,f118])).
fof(f182,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK5(X0,X1,X2,bool))) )),
  inference(superposition,[],[f84,f62])).
fof(f3018,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl44),
  inference(cnf_transformation,[],[f3018_D])).
fof(f3018_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f3025,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) ) | $spl54),
  inference(cnf_transformation,[],[f3025_D])).
fof(f3025_D,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f3015,plain,(
  ( ! [X10,X8,X9] : (s(bool,X8) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X9) = s(fun(bool,bool),X10) | s(bool,'T_0') = s(bool,sK5(X9,X10,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f182,f118])).
fof(f3014,plain,(
  ( ! [X6,X7,X5] : (s(bool,X5) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X6) = s(fun(bool,bool),X7) | s(bool,'T_0') = s(bool,sK5(X6,X7,bool,bool)) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f182,f118])).
fof(f3013,plain,(
  ( ! [X4,X2,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,sK5(X2,X3,bool,bool)) | s(bool,X4) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f182,f118])).
fof(f3021,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | $spl49),
  inference(cnf_transformation,[],[f3021_D])).
fof(f3021_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl49),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl49])])).
fof(f3012,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK5(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f182,f62])).
fof(f2594,plain,(
  ( ! [X47,X46] : (s(bool,'const.bool.IN_2'(s(sK0,X46),s(fun(sK0,bool),sK1))) = s(bool,'const.bool.IN_2'(s(sK0,X47),s(fun(sK0,bool),sK2))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X46),s(fun(sK0,bool),sK2))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X47),s(fun(sK0,bool),sK1)))) ) | $spl36),
  inference(cnf_transformation,[],[f2594_D])).
fof(f2594_D,plain,(
  ( ! [X47,X46] : (s(bool,'const.bool.IN_2'(s(sK0,X46),s(fun(sK0,bool),sK1))) = s(bool,'const.bool.IN_2'(s(sK0,X47),s(fun(sK0,bool),sK2))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X46),s(fun(sK0,bool),sK2))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X47),s(fun(sK0,bool),sK1)))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f1276,plain,(
  ( ! [X3] : (p(s(bool,X3))) ) | $spl8),
  inference(cnf_transformation,[],[f1276_D])).
fof(f1276_D,plain,(
  ( ! [X3] : (p(s(bool,X3))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f2614,plain,(
  ( ! [X52,X53,X51] : (p(s(bool,X52)) | s(bool,'const.bool.IN_2'(s(sK0,X51),s(fun(sK0,bool),sK1))) = s(bool,'const.bool.IN_2'(s(sK0,X53),s(fun(sK0,bool),sK2))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X53),s(fun(sK0,bool),sK1))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X51),s(fun(sK0,bool),sK2)))) ) | $spl6),
  inference(subsumption_resolution,[],[f2537,f64])).
fof(f2537,plain,(
  ( ! [X52,X53,X51] : (p(s(bool,X52)) | s(bool,'const.bool.IN_2'(s(sK0,X51),s(fun(sK0,bool),sK1))) = s(bool,'const.bool.IN_2'(s(sK0,X53),s(fun(sK0,bool),sK2))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X53),s(fun(sK0,bool),sK1))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X51),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X52)) ) | $spl6),
  inference(superposition,[],[f1822,f1270])).
fof(f1822,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK2))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK1)))) ) | $spl6),
  inference(condensation,[],[f1821])).
fof(f1821,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X17) = s(bool,'const.bool.IN_2'(s(sK0,X15),s(fun(sK0,bool),sK2))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X15),s(fun(sK0,bool),sK1)))) ) | $spl6),
  inference(subsumption_resolution,[],[f1752,f64])).
fof(f1752,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X17) = s(bool,'const.bool.IN_2'(s(sK0,X15),s(fun(sK0,bool),sK2))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X15),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X16)) ) | $spl6),
  inference(superposition,[],[f80,f1250])).
fof(f80,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f41])).
fof(f41,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f40])).
fof(f40,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f35])).
fof(f35,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X10,X11] : ((p(s(bool,X10)) => p(s(bool,X11))) => ((p(s(bool,X11)) => p(s(bool,X10))) => s(bool,X10) = s(bool,X11)))),
  file('thm.pred_set.IN_DISJOINT','thm.bool.IMP_ANTISYM_AX')).
fof(f2587,plain,(
  ( ! [X5] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X5),s(fun(sK0,bool),sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X5),s(fun(sK0,bool),sK2)))) ) | $spl34),
  inference(cnf_transformation,[],[f2587_D])).
fof(f2587_D,plain,(
  ( ! [X5] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X5),s(fun(sK0,bool),sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X5),s(fun(sK0,bool),sK2)))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f2603,plain,(
  ( ! [X109,X108] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X108),s(fun(sK0,bool),sK2))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X108),s(fun(sK0,bool),sK1)))) | s(bool,'T_0') = s(bool,X109)) ) | ($spl1 | $spl6)),
  inference(global_subsumption,[],[f2602,f59,f58,f57,f60,f62,f65,f64,f72,f83,f80,f84,f92,f111,f118,f150,f167,f210,f314,f366,f393,f402,f617,f618,f428,f719,f1275,f1291,f1796,f1800,f1789,f1803,f1782,f1806,f1808,f1735,f1737,f1303,f1238,f654,f1299,f1287,f1260,f1258,f1228,f1186,f1150,f1089,f619,f91,f1146,f1145,f1085,f1084,f670,f624,f1171,f1170,f1110,f1109,f676,f672,f630,f626,f74,f692,f646,f79,f78,f77,f73,f764,f76,f75,f365,f427,f1149,f1088,f1822,f1764,f1744,f1743,f1739,f1270,f1250,f678,f632,f183,f182,f110,f69,f68,f61,f2321,f2351,f2571,f2570,f2584,f2470,f2471,f2585,f2560,f2588,f2590,f2552,f2497,f2502,f2503,f2504])).
fof(f2471,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X68),s(fun(bool,bool),X69)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK0,X67),s(fun(sK0,bool),sK2))),s(fun(bool,bool),X68)))) | s(bool,X70) = s(bool,'const.bool.IN_2'(s(sK0,X67),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X70)) ) | $spl6),
  inference(superposition,[],[f77,f1270])).
fof(f2470,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X64),s(fun(bool,bool),X65)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK0,X63),s(fun(sK0,bool),sK2))),s(fun(bool,bool),X65)))) | s(bool,X66) = s(bool,'const.bool.IN_2'(s(sK0,X63),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X66)) ) | $spl6),
  inference(superposition,[],[f78,f1270])).
fof(f2584,plain,(
  ( ! [X61,X62,X60] : (s(bool,'const.bool.IN_2'(s(sK0,X60),s(fun(sK0,bool),sK2))) = s(bool,'const.bool.IN_2'(s(sK0,X61),s(fun(sK0,bool),sK2))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X61),s(fun(sK0,bool),sK1))) | s(bool,X62) = s(bool,'const.bool.IN_2'(s(sK0,X60),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X62)) ) | $spl6),
  inference(subsumption_resolution,[],[f2469,f61])).
fof(f2469,plain,(
  ( ! [X61,X62,X60] : (p(s(bool,'F_0')) | s(bool,'const.bool.IN_2'(s(sK0,X60),s(fun(sK0,bool),sK2))) = s(bool,'const.bool.IN_2'(s(sK0,X61),s(fun(sK0,bool),sK2))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X61),s(fun(sK0,bool),sK1))) | s(bool,X62) = s(bool,'const.bool.IN_2'(s(sK0,X60),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X62)) ) | $spl6),
  inference(superposition,[],[f1822,f1270])).
fof(f2351,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,'const.bool.IN_2'(s(bool,X60),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X61),s(fun(bool,bool),X62)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK0,X59),s(fun(sK0,bool),sK2))),s(fun(bool,bool),X62)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X59),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X60)) ) | $spl6),
  inference(superposition,[],[f78,f1250])).
fof(f2321,plain,(
  ( ! [X103,X101,X102,X100] : (~p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK0,X101),s(fun(sK0,bool),sK2))),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X102),s(fun(bool,bool),X103)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X100),s(fun(bool,bool),X103)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X101),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X100)) ) | $spl6),
  inference(superposition,[],[f78,f1250])).
fof(f68,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f48])).
fof(f678,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,'const.bool.IN_2'(s(bool,X30),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X31),s(fun(bool,bool),X32)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X32)))) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f78,f62])).
fof(f1739,plain,(
  ( ! [X88,X87,X85,X86] : (~p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK0,X86),s(fun(sK0,bool),sK2))),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X87),s(fun(bool,bool),X88)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X85),s(fun(bool,bool),X87)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X86),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X85)) ) | $spl6),
  inference(superposition,[],[f77,f1250])).
fof(f1743,plain,(
  ( ! [X103,X101,X102,X100] : (s(bool,f(s(X100,sK5(X101,X102,bool,X100)))) != s(bool,'const.bool.IN_2'(s(sK0,X103),s(fun(sK0,bool),sK2))) | s(fun(X100,bool),X101) = s(fun(X100,bool),X102) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X103),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,g(s(X100,sK5(X101,X102,bool,X100))))) ) | $spl6),
  inference(superposition,[],[f84,f1250])).
fof(f1744,plain,(
  ( ! [X107,X105,X106,X104] : (s(X106,f(s(bool,'const.bool.IN_2'(s(sK0,X107),s(fun(sK0,bool),sK2))))) != s(X106,g(s(bool,'const.bool.IN_2'(s(sK0,X107),s(fun(sK0,bool),sK2))))) | s(fun(bool,X106),X104) = s(fun(bool,X106),X105) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X107),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,sK5(X104,X105,X106,bool))) ) | $spl6),
  inference(superposition,[],[f84,f1250])).
fof(f1764,plain,(
  ( ! [X52,X50,X53,X51] : (~p(s(bool,'const.bool.IN_2'(s(bool,X51),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X52),s(fun(bool,bool),X53)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK0,X50),s(fun(sK0,bool),sK2))),s(fun(bool,bool),X52)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X50),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X51)) ) | $spl6),
  inference(superposition,[],[f77,f1250])).
fof(f1088,plain,(
  ( ! [X158,X156,X159,X157,X160] : (s(bool,f(s(X156,sK5(X157,X158,bool,X156)))) != s(bool,'const.bool.IN_2'(s(sK0,X159),s(fun(sK0,bool),sK2))) | s(fun(X156,bool),X157) = s(fun(X156,bool),X158) | s(bool,X160) = s(bool,'const.bool.IN_2'(s(sK0,X159),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,g(s(X156,sK5(X157,X158,bool,X156)))) | s(bool,'T_0') = s(bool,X160)) ) | $spl6),
  inference(superposition,[],[f84,f654])).
fof(f1149,plain,(
  ( ! [X158,X156,X159,X157,X160] : (s(bool,f(s(X156,sK5(X157,X158,bool,X156)))) != s(bool,'const.bool.IN_2'(s(sK0,X159),s(fun(sK0,bool),sK1))) | s(fun(X156,bool),X157) = s(fun(X156,bool),X158) | s(bool,X160) = s(bool,'const.bool.IN_2'(s(sK0,X159),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X160) | s(bool,'T_0') = s(bool,g(s(X156,sK5(X157,X158,bool,X156))))) ) | $spl6),
  inference(superposition,[],[f84,f654])).
fof(f427,plain,(
  ( ! [X54,X52,X55,X53] : (s(bool,'T_0') != s(bool,f(s(X52,sK5(X53,X54,bool,X52)))) | s(fun(X52,bool),X53) = s(fun(X52,bool),X54) | s(bool,X55) = s(bool,g(s(X52,sK5(X53,X54,bool,X52)))) | s(bool,'T_0') = s(bool,X55)) )),
  inference(superposition,[],[f84,f118])).
fof(f365,plain,(
  ( ! [X59,X57,X58,X56] : (s(bool,X59) != s(bool,f(s(X56,sK5(X57,X58,bool,X56)))) | s(fun(X56,bool),X57) = s(fun(X56,bool),X58) | s(bool,'T_0') = s(bool,X59) | s(bool,'T_0') = s(bool,g(s(X56,sK5(X57,X58,bool,X56))))) )),
  inference(superposition,[],[f84,f118])).
fof(f764,plain,(
  ( ! [X4,X2,X3,X1] : (s(fun(X1,bool),'const.pred_set.EMPTY_0') = s(fun(X1,bool),'const.pred_set.INTER_2'(s(fun(X1,bool),X2),s(fun(X1,bool),X3))) | s(bool,X4) = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X1,bool),X2),s(fun(X1,bool),X3))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f758,f60])).
fof(f758,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,'T_0')) | s(fun(X1,bool),'const.pred_set.EMPTY_0') = s(fun(X1,bool),'const.pred_set.INTER_2'(s(fun(X1,bool),X2),s(fun(X1,bool),X3))) | s(bool,X4) = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X1,bool),X2),s(fun(X1,bool),X3))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f75,f118])).
fof(f73,plain,(
  ( ! [X4,X2,X0,X1] : (s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) )),
  inference(cnf_transformation,[],[f52])).
fof(f646,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))) | s(bool,X4) = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f636,f60])).
fof(f636,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))) | s(bool,X4) = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f77,f118])).
fof(f692,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3)))) | s(bool,X4) = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f682,f60])).
fof(f682,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3)))) | s(bool,X4) = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f78,f118])).
fof(f626,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,'const.bool.IN_2'(s(bool,X9),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X10),s(fun(bool,bool),X11)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X8),s(fun(bool,bool),X10)))) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f77,f118])).
fof(f630,plain,(
  ( ! [X24,X23,X25,X22] : (~p(s(bool,'const.bool.IN_2'(s(bool,X22),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X23),s(fun(bool,bool),X24)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X23)))) | s(bool,X22) = s(bool,X25) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f77,f118])).
fof(f672,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,'const.bool.IN_2'(s(bool,X9),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X10),s(fun(bool,bool),X11)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X8),s(fun(bool,bool),X11)))) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f78,f118])).
fof(f676,plain,(
  ( ! [X24,X23,X25,X22] : (~p(s(bool,'const.bool.IN_2'(s(bool,X22),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X23),s(fun(bool,bool),X24)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X24)))) | s(bool,X22) = s(bool,X25) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f78,f118])).
fof(f1109,plain,(
  ( ! [X70,X72,X71,X69,X73] : (~p(s(bool,'const.bool.IN_2'(s(bool,X70),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X71),s(fun(bool,bool),X72)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK0,X69),s(fun(sK0,bool),sK2))),s(fun(bool,bool),X72)))) | s(bool,X73) = s(bool,'const.bool.IN_2'(s(sK0,X69),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X70) | s(bool,'T_0') = s(bool,X73)) ) | $spl6),
  inference(superposition,[],[f78,f654])).
fof(f1110,plain,(
  ( ! [X78,X76,X74,X77,X75] : (~p(s(bool,'const.bool.IN_2'(s(bool,X75),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X76),s(fun(bool,bool),X77)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK0,X74),s(fun(sK0,bool),sK2))),s(fun(bool,bool),X76)))) | s(bool,X78) = s(bool,'const.bool.IN_2'(s(sK0,X74),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X75) | s(bool,'T_0') = s(bool,X78)) ) | $spl6),
  inference(superposition,[],[f77,f654])).
fof(f1170,plain,(
  ( ! [X70,X74,X72,X71,X73] : (~p(s(bool,'const.bool.IN_2'(s(bool,X71),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X72),s(fun(bool,bool),X73)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK0,X70),s(fun(sK0,bool),sK1))),s(fun(bool,bool),X73)))) | s(bool,X74) = s(bool,'const.bool.IN_2'(s(sK0,X70),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X74) | s(bool,'T_0') = s(bool,X71)) ) | $spl6),
  inference(superposition,[],[f78,f654])).
fof(f1171,plain,(
  ( ! [X78,X76,X79,X77,X75] : (~p(s(bool,'const.bool.IN_2'(s(bool,X76),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X77),s(fun(bool,bool),X78)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK0,X75),s(fun(sK0,bool),sK1))),s(fun(bool,bool),X77)))) | s(bool,X79) = s(bool,'const.bool.IN_2'(s(sK0,X75),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X79) | s(bool,'T_0') = s(bool,X76)) ) | $spl6),
  inference(superposition,[],[f77,f654])).
fof(f624,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X1),s(fun(bool,bool),X2)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),X1)))) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f77,f118])).
fof(f670,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X1),s(fun(bool,bool),X2)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),X2)))) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f78,f118])).
fof(f1084,plain,(
  ( ! [X140,X138,X136,X139,X137] : (~p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK0,X137),s(fun(sK0,bool),sK2))),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X138),s(fun(bool,bool),X139)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X136),s(fun(bool,bool),X139)))) | s(bool,X140) = s(bool,'const.bool.IN_2'(s(sK0,X137),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X136) | s(bool,'T_0') = s(bool,X140)) ) | $spl6),
  inference(superposition,[],[f78,f654])).
fof(f1085,plain,(
  ( ! [X144,X142,X145,X143,X141] : (~p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK0,X142),s(fun(sK0,bool),sK2))),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X143),s(fun(bool,bool),X144)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X141),s(fun(bool,bool),X143)))) | s(bool,X145) = s(bool,'const.bool.IN_2'(s(sK0,X142),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X141) | s(bool,'T_0') = s(bool,X145)) ) | $spl6),
  inference(superposition,[],[f77,f654])).
fof(f1145,plain,(
  ( ! [X140,X138,X136,X139,X137] : (~p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK0,X137),s(fun(sK0,bool),sK1))),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X138),s(fun(bool,bool),X139)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X136),s(fun(bool,bool),X139)))) | s(bool,X140) = s(bool,'const.bool.IN_2'(s(sK0,X137),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X140) | s(bool,'T_0') = s(bool,X136)) ) | $spl6),
  inference(superposition,[],[f78,f654])).
fof(f1146,plain,(
  ( ! [X144,X142,X145,X143,X141] : (~p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK0,X142),s(fun(sK0,bool),sK1))),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X143),s(fun(bool,bool),X144)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X141),s(fun(bool,bool),X143)))) | s(bool,X145) = s(bool,'const.bool.IN_2'(s(sK0,X142),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X145) | s(bool,'T_0') = s(bool,X141)) ) | $spl6),
  inference(superposition,[],[f77,f654])).
fof(f1089,plain,(
  ( ! [X165,X163,X161,X164,X162] : (s(X163,f(s(bool,'const.bool.IN_2'(s(sK0,X164),s(fun(sK0,bool),sK2))))) != s(X163,g(s(bool,'const.bool.IN_2'(s(sK0,X164),s(fun(sK0,bool),sK2))))) | s(fun(bool,X163),X161) = s(fun(bool,X163),X162) | s(bool,X165) = s(bool,'const.bool.IN_2'(s(sK0,X164),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,sK5(X161,X162,X163,bool)) | s(bool,'T_0') = s(bool,X165)) ) | $spl6),
  inference(superposition,[],[f84,f654])).
fof(f1150,plain,(
  ( ! [X165,X163,X161,X164,X162] : (s(X163,f(s(bool,'const.bool.IN_2'(s(sK0,X164),s(fun(sK0,bool),sK1))))) != s(X163,g(s(bool,'const.bool.IN_2'(s(sK0,X164),s(fun(sK0,bool),sK1))))) | s(fun(bool,X163),X161) = s(fun(bool,X163),X162) | s(bool,X165) = s(bool,'const.bool.IN_2'(s(sK0,X164),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X165) | s(bool,'T_0') = s(bool,sK5(X161,X162,X163,bool))) ) | $spl6),
  inference(superposition,[],[f84,f654])).
fof(f1186,plain,(
  ( ! [X0,X1] : (s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK1))) != s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK2))) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X0)) ) | $spl6),
  inference(duplicate_literal_removal,[],[f1174])).
fof(f1174,plain,(
  ( ! [X0,X1] : (s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK1))) != s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK2))) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl6),
  inference(equality_factoring,[],[f654])).
fof(f1228,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK2)))) | s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK1)))) ) | $spl6),
  inference(condensation,[],[f1227])).
fof(f1227,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X7),s(fun(sK0,bool),sK2)))) | s(bool,'T_0') = s(bool,X6) | s(bool,X8) = s(bool,'const.bool.IN_2'(s(sK0,X7),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X8)) ) | $spl6),
  inference(duplicate_literal_removal,[],[f1054])).
fof(f1054,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X7),s(fun(sK0,bool),sK2)))) | s(bool,'T_0') = s(bool,X6) | s(bool,X8) = s(bool,'const.bool.IN_2'(s(sK0,X7),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,X8)) ) | $spl6),
  inference(superposition,[],[f65,f654])).
fof(f1258,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK1))) = s(bool,'const.bool.IN_2'(s(sK0,X2),s(fun(sK0,bool),sK2))) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,X2),s(fun(sK0,bool),sK1))) | s(bool,X3) = s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X3)) ) | $spl6),
  inference(condensation,[],[f1257])).
fof(f1257,plain,(
  ( ! [X35,X33,X36,X34,X32] : (s(bool,'const.bool.IN_2'(s(sK0,X33),s(fun(sK0,bool),sK1))) = s(bool,'const.bool.IN_2'(s(sK0,X34),s(fun(sK0,bool),sK2))) | s(bool,X35) = s(bool,'const.bool.IN_2'(s(sK0,X34),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X32) | s(bool,'T_0') = s(bool,X35) | s(bool,X36) = s(bool,'const.bool.IN_2'(s(sK0,X33),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X36)) ) | $spl6),
  inference(duplicate_literal_removal,[],[f1008])).
fof(f1008,plain,(
  ( ! [X35,X33,X36,X34,X32] : (s(bool,'const.bool.IN_2'(s(sK0,X33),s(fun(sK0,bool),sK1))) = s(bool,'const.bool.IN_2'(s(sK0,X34),s(fun(sK0,bool),sK2))) | s(bool,X35) = s(bool,'const.bool.IN_2'(s(sK0,X34),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X32) | s(bool,'T_0') = s(bool,X35) | s(bool,X36) = s(bool,'const.bool.IN_2'(s(sK0,X33),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X36) | s(bool,'T_0') = s(bool,X32)) ) | $spl6),
  inference(superposition,[],[f654,f654])).
fof(f1260,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK2))) = s(bool,'const.bool.IN_2'(s(sK0,X2),s(fun(sK0,bool),sK2))) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,X2),s(fun(sK0,bool),sK1))) | s(bool,X3) = s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X3)) ) | $spl6),
  inference(condensation,[],[f1259])).
fof(f1259,plain,(
  ( ! [X30,X28,X31,X29,X27] : (s(bool,'const.bool.IN_2'(s(sK0,X28),s(fun(sK0,bool),sK2))) = s(bool,'const.bool.IN_2'(s(sK0,X29),s(fun(sK0,bool),sK2))) | s(bool,X30) = s(bool,'const.bool.IN_2'(s(sK0,X29),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,X30) | s(bool,X31) = s(bool,'const.bool.IN_2'(s(sK0,X28),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X31)) ) | $spl6),
  inference(duplicate_literal_removal,[],[f1007])).
fof(f1007,plain,(
  ( ! [X30,X28,X31,X29,X27] : (s(bool,'const.bool.IN_2'(s(sK0,X28),s(fun(sK0,bool),sK2))) = s(bool,'const.bool.IN_2'(s(sK0,X29),s(fun(sK0,bool),sK2))) | s(bool,X30) = s(bool,'const.bool.IN_2'(s(sK0,X29),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,X30) | s(bool,X31) = s(bool,'const.bool.IN_2'(s(sK0,X28),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,X31)) ) | $spl6),
  inference(superposition,[],[f654,f654])).
fof(f1287,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK2))) | s(bool,X2) = s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X2)) ) | $spl6),
  inference(condensation,[],[f1286])).
fof(f1286,plain,(
  ( ! [X21,X19,X20,X18] : (p(s(bool,X19)) | p(s(bool,X20)) | s(bool,X20) = s(bool,'const.bool.IN_2'(s(sK0,X18),s(fun(sK0,bool),sK2))) | s(bool,X21) = s(bool,'const.bool.IN_2'(s(sK0,X18),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X21)) ) | $spl6),
  inference(subsumption_resolution,[],[f1096,f64])).
fof(f1096,plain,(
  ( ! [X21,X19,X20,X18] : (p(s(bool,X19)) | p(s(bool,X20)) | s(bool,X20) = s(bool,'const.bool.IN_2'(s(sK0,X18),s(fun(sK0,bool),sK2))) | s(bool,X21) = s(bool,'const.bool.IN_2'(s(sK0,X18),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') = s(bool,X21)) ) | $spl6),
  inference(superposition,[],[f80,f654])).
fof(f1299,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK1))) | s(bool,X2) = s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X2)) ) | $spl6),
  inference(condensation,[],[f1298])).
fof(f1298,plain,(
  ( ! [X21,X19,X22,X20] : (p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X21) = s(bool,'const.bool.IN_2'(s(sK0,X19),s(fun(sK0,bool),sK1))) | s(bool,X22) = s(bool,'const.bool.IN_2'(s(sK0,X19),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X22)) ) | $spl6),
  inference(subsumption_resolution,[],[f1157,f64])).
fof(f1157,plain,(
  ( ! [X21,X19,X22,X20] : (p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X21) = s(bool,'const.bool.IN_2'(s(sK0,X19),s(fun(sK0,bool),sK1))) | s(bool,X22) = s(bool,'const.bool.IN_2'(s(sK0,X19),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,X20)) ) | $spl6),
  inference(superposition,[],[f80,f654])).
fof(f1238,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK1))) = s(bool,'const.bool.IN_2'(s(sK0,X2),s(fun(sK0,bool),sK1))) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,X2),s(fun(sK0,bool),sK2))) | s(bool,X3) = s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X3)) ) | $spl6),
  inference(condensation,[],[f1237])).
fof(f1237,plain,(
  ( ! [X35,X33,X36,X34,X32] : (s(bool,'const.bool.IN_2'(s(sK0,X33),s(fun(sK0,bool),sK1))) = s(bool,'const.bool.IN_2'(s(sK0,X34),s(fun(sK0,bool),sK1))) | s(bool,X35) = s(bool,'const.bool.IN_2'(s(sK0,X34),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X35) | s(bool,'T_0') = s(bool,X32) | s(bool,X36) = s(bool,'const.bool.IN_2'(s(sK0,X33),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X36)) ) | $spl6),
  inference(duplicate_literal_removal,[],[f1034])).
fof(f1034,plain,(
  ( ! [X35,X33,X36,X34,X32] : (s(bool,'const.bool.IN_2'(s(sK0,X33),s(fun(sK0,bool),sK1))) = s(bool,'const.bool.IN_2'(s(sK0,X34),s(fun(sK0,bool),sK1))) | s(bool,X35) = s(bool,'const.bool.IN_2'(s(sK0,X34),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X35) | s(bool,'T_0') = s(bool,X32) | s(bool,X36) = s(bool,'const.bool.IN_2'(s(sK0,X33),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X36) | s(bool,'T_0') = s(bool,X32)) ) | $spl6),
  inference(superposition,[],[f654,f654])).
fof(f1303,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK2))) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X0)) ) | $spl6),
  inference(subsumption_resolution,[],[f1176,f118])).
fof(f1176,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK2))) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK2))) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X0)) ) | $spl6),
  inference(equality_factoring,[],[f654])).
fof(f719,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f617,f62])).
fof(f428,plain,(
  ( ! [X59,X57,X58,X56] : (s(X58,f(s(bool,'T_0'))) != s(X58,g(s(bool,'T_0'))) | s(fun(bool,X58),X56) = s(fun(bool,X58),X57) | s(bool,X59) = s(bool,sK5(X56,X57,X58,bool)) | s(bool,'T_0') = s(bool,X59)) )),
  inference(superposition,[],[f84,f118])).
fof(f618,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f411,f60])).
fof(f411,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f83,f118])).
fof(f617,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f407,f60])).
fof(f407,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f68,f118])).
fof(f402,plain,(
  ( ! [X31,X32] : (s(bool,X31) != s(bool,'const.bool.IN_2'(s(bool,X31),s(fun(bool,bool),'const.pred_set.EMPTY_0'))) | s(bool,X31) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(superposition,[],[f111,f118])).
fof(f393,plain,(
  ( ! [X4,X2,X3] : (s(bool,X2) != s(bool,X3) | p(s(bool,X3)) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f64,f118])).
fof(f366,plain,(
  ( ! [X61,X62,X60,X63] : (s(X62,f(s(bool,X63))) != s(X62,g(s(bool,X63))) | s(fun(bool,X62),X60) = s(fun(bool,X62),X61) | s(bool,'T_0') = s(bool,X63) | s(bool,'T_0') = s(bool,sK5(X60,X61,X62,bool))) )),
  inference(superposition,[],[f84,f118])).
fof(f314,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f118,f118])).
fof(f210,plain,(
  ( ! [X6,X7,X5] : (s(bool,X7) = s(bool,'const.bool.IN_2'(s(X5,X6),s(fun(X5,bool),'const.pred_set.EMPTY_0'))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f207,f61])).
fof(f207,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,'F_0')) | s(bool,X7) = s(bool,'const.bool.IN_2'(s(X5,X6),s(fun(X5,bool),'const.pred_set.EMPTY_0'))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f150,f110])).
fof(f167,plain,(
  ( ! [X1] : (s(bool,X1) = s(bool,'const.bool.IN_2'(s(bool,X1),s(fun(bool,bool),'const.pred_set.EMPTY_0'))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(superposition,[],[f110,f62])).
fof(f150,plain,(
  ( ! [X4,X3] : (p(s(bool,X4)) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f144,f61])).
fof(f144,plain,(
  ( ! [X4,X3] : (p(s(bool,'F_0')) | p(s(bool,X4)) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f80,f62])).
fof(f111,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f72,f64])).
fof(f83,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f41])).
fof(f2602,plain,(
  ( ! [X109,X107,X108] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X108),s(fun(sK0,bool),sK1)))) | s(bool,X109) = s(bool,'const.bool.IN_2'(s(sK0,X107),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X109) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X108),s(fun(sK0,bool),sK2)))) ) | ($spl1 | $spl6)),
  inference(subsumption_resolution,[],[f2510,f1296])).
fof(f1296,plain,(
  ( ! [X80,X81] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,X80),s(fun(sK0,bool),sK1))) | s(bool,X81) = s(bool,'const.bool.IN_2'(s(sK0,X80),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X81)) ) | ($spl1 | $spl6)),
  inference(subsumption_resolution,[],[f1132,f93])).
fof(f93,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | $spl1),
  inference(unit_resulting_resolution,[],[f86,f64])).
fof(f1132,plain,(
  ( ! [X80,X81] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,X80),s(fun(sK0,bool),sK1))) | s(bool,X81) = s(bool,'const.bool.IN_2'(s(sK0,X80),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X81) | s(bool,'T_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) ) | ($spl1 | $spl6)),
  inference(superposition,[],[f93,f654])).
fof(f2510,plain,(
  ( ! [X109,X107,X108] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X108),s(fun(sK0,bool),sK1)))) | s(bool,X109) = s(bool,'const.bool.IN_2'(s(sK0,X107),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X109) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X108),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,X107),s(fun(sK0,bool),sK1)))) ) | $spl6),
  inference(superposition,[],[f619,f1270])).
fof(f2601,plain,(
  ( ! [X103] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X103),s(fun(sK0,bool),sK1)))) ) | $spl42),
  inference(cnf_transformation,[],[f2601_D])).
fof(f2601_D,plain,(
  ( ! [X103] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X103),s(fun(sK0,bool),sK1)))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f2600,plain,(
  ( ! [X103,X104] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X104),s(fun(sK0,bool),sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X103),s(fun(sK0,bool),sK1))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X104),s(fun(sK0,bool),sK2)))) ) | ($spl1 | $spl6 | $spl22)),
  inference(subsumption_resolution,[],[f2508,f1812])).
fof(f1812,plain,(
  ( ! [X59] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,X59),s(fun(sK0,bool),sK2))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X59),s(fun(sK0,bool),sK1)))) ) | ($spl1 | $spl6)),
  inference(subsumption_resolution,[],[f1730,f93])).
fof(f1730,plain,(
  ( ! [X59] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,X59),s(fun(sK0,bool),sK2))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X59),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) ) | ($spl1 | $spl6)),
  inference(superposition,[],[f93,f1250])).
fof(f2508,plain,(
  ( ! [X103,X104] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X104),s(fun(sK0,bool),sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X103),s(fun(sK0,bool),sK1))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X104),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,X103),s(fun(sK0,bool),sK2)))) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f1805,f1270])).
fof(f2599,plain,(
  ( ! [X97,X96] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X96),s(fun(sK0,bool),sK2))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X96),s(fun(sK0,bool),sK1)))) | s(bool,'T_0') = s(bool,X97)) ) | ($spl1 | $spl6)),
  inference(global_subsumption,[],[f2598,f59,f58,f57,f60,f62,f65,f64,f72,f83,f80,f84,f92,f111,f118,f150,f167,f210,f314,f366,f393,f402,f617,f618,f428,f719,f1275,f1291,f1796,f1800,f1789,f1803,f1782,f1806,f1808,f1735,f1737,f1303,f1238,f654,f1299,f1287,f1260,f1258,f1228,f1186,f1150,f1089,f619,f91,f1146,f1145,f1085,f1084,f670,f624,f1171,f1170,f1110,f1109,f676,f672,f630,f626,f74,f692,f646,f79,f78,f77,f73,f764,f76,f75,f365,f427,f1149,f1088,f1822,f1764,f1744,f1743,f1739,f1270,f1250,f678,f632,f183,f182,f110,f69,f68,f61,f2321,f2351,f2571,f2570,f2584,f2470,f2471,f2585,f2560,f2588,f2590,f2552,f2497,f2502,f2503,f2504])).
fof(f2598,plain,(
  ( ! [X97,X95,X96] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X96),s(fun(sK0,bool),sK1)))) | s(bool,'T_0') = s(bool,X97) | s(bool,X97) = s(bool,'const.bool.IN_2'(s(sK0,X95),s(fun(sK0,bool),sK1))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X96),s(fun(sK0,bool),sK2)))) ) | ($spl1 | $spl6)),
  inference(subsumption_resolution,[],[f2505,f1284])).
fof(f1284,plain,(
  ( ! [X80,X81] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,X80),s(fun(sK0,bool),sK2))) | s(bool,X81) = s(bool,'const.bool.IN_2'(s(sK0,X80),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X81)) ) | ($spl1 | $spl6)),
  inference(subsumption_resolution,[],[f1071,f93])).
fof(f1071,plain,(
  ( ! [X80,X81] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,X80),s(fun(sK0,bool),sK2))) | s(bool,X81) = s(bool,'const.bool.IN_2'(s(sK0,X80),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X81)) ) | ($spl1 | $spl6)),
  inference(superposition,[],[f93,f654])).
fof(f2505,plain,(
  ( ! [X97,X95,X96] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X96),s(fun(sK0,bool),sK1)))) | s(bool,'T_0') = s(bool,X97) | s(bool,X97) = s(bool,'const.bool.IN_2'(s(sK0,X95),s(fun(sK0,bool),sK1))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X96),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,X95),s(fun(sK0,bool),sK2)))) ) | $spl6),
  inference(superposition,[],[f1228,f1270])).
fof(f1816,plain,(
  ( ! [X80] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X80),s(fun(sK0,bool),sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,X80),s(fun(sK0,bool),sK2)))) ) | $spl26),
  inference(cnf_transformation,[],[f1816_D])).
fof(f1816_D,plain,(
  ( ! [X80] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X80),s(fun(sK0,bool),sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,X80),s(fun(sK0,bool),sK2)))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f2504,plain,(
  ( ! [X94,X93] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X94),s(fun(sK0,bool),sK1)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X93),s(fun(sK0,bool),sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X94),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,X93),s(fun(sK0,bool),sK2)))) ) | $spl6),
  inference(superposition,[],[f91,f1270])).
fof(f1813,plain,(
  ( ! [X74,X72,X75,X73] : (p(s(bool,'const.bool.IN_2'(s(X72,X73),s(fun(X72,bool),X74)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X72,X73),s(fun(X72,bool),'const.pred_set.INTER_2'(s(fun(X72,bool),X74),s(fun(X72,bool),X75)))))) ) | $spl24),
  inference(cnf_transformation,[],[f1813_D])).
fof(f1813_D,plain,(
  ( ! [X74,X72,X75,X73] : (p(s(bool,'const.bool.IN_2'(s(X72,X73),s(fun(X72,bool),X74)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X72,X73),s(fun(X72,bool),'const.pred_set.INTER_2'(s(fun(X72,bool),X74),s(fun(X72,bool),X75)))))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f2503,plain,(
  ( ! [X92,X90,X88,X91,X89] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X92),s(fun(sK0,bool),sK1)))) | p(s(bool,'const.bool.IN_2'(s(X88,X89),s(fun(X88,bool),X90)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X92),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X88,X89),s(fun(X88,bool),'const.pred_set.INTER_2'(s(fun(X88,bool),X90),s(fun(X88,bool),X91)))))) ) | $spl6),
  inference(superposition,[],[f77,f1270])).
fof(f2597,plain,(
  ( ! [X85,X83,X86,X84] : (p(s(bool,'const.bool.IN_2'(s(X83,X84),s(fun(X83,bool),X86)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X83,X84),s(fun(X83,bool),'const.pred_set.INTER_2'(s(fun(X83,bool),X85),s(fun(X83,bool),X86)))))) ) | $spl40),
  inference(cnf_transformation,[],[f2597_D])).
fof(f2597_D,plain,(
  ( ! [X85,X83,X86,X84] : (p(s(bool,'const.bool.IN_2'(s(X83,X84),s(fun(X83,bool),X86)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X83,X84),s(fun(X83,bool),'const.pred_set.INTER_2'(s(fun(X83,bool),X85),s(fun(X83,bool),X86)))))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f2502,plain,(
  ( ! [X87,X85,X83,X86,X84] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X87),s(fun(sK0,bool),sK1)))) | p(s(bool,'const.bool.IN_2'(s(X83,X84),s(fun(X83,bool),X86)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X87),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X83,X84),s(fun(X83,bool),'const.pred_set.INTER_2'(s(fun(X83,bool),X85),s(fun(X83,bool),X86)))))) ) | $spl6),
  inference(superposition,[],[f78,f1270])).
fof(f2596,plain,(
  ( ! [X68,X69,X67] : (s(fun(X67,bool),'const.pred_set.EMPTY_0') = s(fun(X67,bool),'const.pred_set.INTER_2'(s(fun(X67,bool),X68),s(fun(X67,bool),X69))) | s(bool,'T_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X67,bool),X68),s(fun(X67,bool),X69)))) ) | $spl38),
  inference(cnf_transformation,[],[f2596_D])).
fof(f2596_D,plain,(
  ( ! [X68,X69,X67] : (s(fun(X67,bool),'const.pred_set.EMPTY_0') = s(fun(X67,bool),'const.pred_set.INTER_2'(s(fun(X67,bool),X68),s(fun(X67,bool),X69))) | s(bool,'T_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X67,bool),X68),s(fun(X67,bool),X69)))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f2497,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X70),s(fun(sK0,bool),sK1)))) | s(fun(X67,bool),'const.pred_set.EMPTY_0') = s(fun(X67,bool),'const.pred_set.INTER_2'(s(fun(X67,bool),X68),s(fun(X67,bool),X69))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X70),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X67,bool),X68),s(fun(X67,bool),X69)))) ) | $spl6),
  inference(superposition,[],[f75,f1270])).
fof(f2552,plain,(
  ( ! [X47,X45,X46] : (s(bool,'const.bool.IN_2'(s(sK0,X46),s(fun(sK0,bool),sK1))) = s(bool,'const.bool.IN_2'(s(sK0,X47),s(fun(sK0,bool),sK2))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X47),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X45) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X46),s(fun(sK0,bool),sK2)))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f2492])).
fof(f2492,plain,(
  ( ! [X47,X45,X46] : (s(bool,'const.bool.IN_2'(s(sK0,X46),s(fun(sK0,bool),sK1))) = s(bool,'const.bool.IN_2'(s(sK0,X47),s(fun(sK0,bool),sK2))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X47),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X45) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X46),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X45)) ) | $spl6),
  inference(superposition,[],[f1250,f1270])).
fof(f2590,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK2)))) ) | $spl6),
  inference(condensation,[],[f2589])).
fof(f2589,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X17),s(fun(sK0,bool),sK1)))) | s(bool,X16) = s(bool,X18) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X17),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X16)) ) | $spl6),
  inference(subsumption_resolution,[],[f2482,f150])).
fof(f2482,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X17),s(fun(sK0,bool),sK1)))) | ~p(s(bool,X18)) | s(bool,X16) = s(bool,X18) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X17),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X16)) ) | $spl6),
  inference(superposition,[],[f83,f1270])).
fof(f2588,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X7),s(fun(sK0,bool),sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X7),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X6)) ) | $spl6),
  inference(subsumption_resolution,[],[f2478,f62])).
fof(f2478,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X7),s(fun(sK0,bool),sK1)))) | s(bool,'F_0') != s(bool,X6) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X7),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X6)) ) | $spl6),
  inference(superposition,[],[f68,f1270])).
fof(f2560,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X5),s(fun(sK0,bool),sK1)))) | s(bool,'T_0') = s(bool,X4) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X5),s(fun(sK0,bool),sK2)))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f2477])).
fof(f2477,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X5),s(fun(sK0,bool),sK1)))) | s(bool,'T_0') = s(bool,X4) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X5),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X4)) ) | $spl6),
  inference(superposition,[],[f65,f1270])).
fof(f2586,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,X3),s(fun(sK0,bool),sK1))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X3),s(fun(sK0,bool),sK2)))) ) | $spl32),
  inference(cnf_transformation,[],[f2586_D])).
fof(f2586_D,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,X3),s(fun(sK0,bool),sK1))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X3),s(fun(sK0,bool),sK2)))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f2585,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,X3),s(fun(sK0,bool),sK1))) | p(s(bool,X2)) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X3),s(fun(sK0,bool),sK2)))) ) | $spl6),
  inference(subsumption_resolution,[],[f2476,f64])).
fof(f2476,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,X3),s(fun(sK0,bool),sK1))) | p(s(bool,X2)) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X3),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X2)) ) | $spl6),
  inference(superposition,[],[f64,f1270])).
fof(f2583,plain,(
  ( ! [X24,X25] : (s(bool,'const.bool.IN_2'(s(sK0,X24),s(fun(sK0,bool),sK1))) = s(bool,'const.bool.IN_2'(s(sK0,X25),s(fun(sK0,bool),sK1))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X24),s(fun(sK0,bool),sK2))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X25),s(fun(sK0,bool),sK2)))) ) | $spl30),
  inference(cnf_transformation,[],[f2583_D])).
fof(f2583_D,plain,(
  ( ! [X24,X25] : (s(bool,'const.bool.IN_2'(s(sK0,X24),s(fun(sK0,bool),sK1))) = s(bool,'const.bool.IN_2'(s(sK0,X25),s(fun(sK0,bool),sK1))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X24),s(fun(sK0,bool),sK2))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X25),s(fun(sK0,bool),sK2)))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f2570,plain,(
  ( ! [X24,X23,X25] : (s(bool,'const.bool.IN_2'(s(sK0,X24),s(fun(sK0,bool),sK1))) = s(bool,'const.bool.IN_2'(s(sK0,X25),s(fun(sK0,bool),sK1))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X25),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X23) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X24),s(fun(sK0,bool),sK2)))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f2427])).
fof(f2427,plain,(
  ( ! [X24,X23,X25] : (s(bool,'const.bool.IN_2'(s(sK0,X24),s(fun(sK0,bool),sK1))) = s(bool,'const.bool.IN_2'(s(sK0,X25),s(fun(sK0,bool),sK1))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X25),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X23) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X24),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X23)) ) | $spl6),
  inference(superposition,[],[f1270,f1270])).
fof(f2582,plain,(
  ( ! [X21,X22] : (s(bool,'const.bool.IN_2'(s(sK0,X21),s(fun(sK0,bool),sK2))) = s(bool,'const.bool.IN_2'(s(sK0,X22),s(fun(sK0,bool),sK1))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X21),s(fun(sK0,bool),sK1))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X22),s(fun(sK0,bool),sK2)))) ) | $spl28),
  inference(cnf_transformation,[],[f2582_D])).
fof(f2582_D,plain,(
  ( ! [X21,X22] : (s(bool,'const.bool.IN_2'(s(sK0,X21),s(fun(sK0,bool),sK2))) = s(bool,'const.bool.IN_2'(s(sK0,X22),s(fun(sK0,bool),sK1))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X21),s(fun(sK0,bool),sK1))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X22),s(fun(sK0,bool),sK2)))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f2571,plain,(
  ( ! [X21,X22,X20] : (s(bool,'const.bool.IN_2'(s(sK0,X21),s(fun(sK0,bool),sK2))) = s(bool,'const.bool.IN_2'(s(sK0,X22),s(fun(sK0,bool),sK1))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X22),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X20) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X21),s(fun(sK0,bool),sK1)))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f2426])).
fof(f2426,plain,(
  ( ! [X21,X22,X20] : (s(bool,'const.bool.IN_2'(s(sK0,X21),s(fun(sK0,bool),sK2))) = s(bool,'const.bool.IN_2'(s(sK0,X22),s(fun(sK0,bool),sK1))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X22),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X20) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X21),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X20)) ) | $spl6),
  inference(superposition,[],[f1270,f1250])).
fof(f1933,plain,(
  $false | ($spl6 | $spl14)),
  inference(subsumption_resolution,[],[f1932,f60])).
fof(f1932,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl14)),
  inference(forward_demodulation,[],[f1931,f1797])).
fof(f1931,plain,(
  ( ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X4),s(fun(sK0,bool),sK1))))) ) | ($spl6 | $spl14)),
  inference(subsumption_resolution,[],[f1878,f60])).
fof(f1878,plain,(
  ( ! [X4] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X4),s(fun(sK0,bool),sK1))))) ) | ($spl6 | $spl14)),
  inference(backward_demodulation,[],[f1797,f91])).
fof(f1930,plain,(
  $false | $spl14),
  inference(global_subsumption,[],[f59,f58,f57,f60,f62,f65,f64,f72,f83,f80,f84,f92,f111,f118,f150,f167,f210,f314,f366,f393,f402,f617,f618,f428,f719,f74,f73,f764,f365,f427,f183,f182,f110,f69,f68,f61,f678,f676,f672,f670,f632,f630,f626,f624,f692,f646,f79,f78,f77,f76,f75,f1797,f1923,f1925,f1929])).
fof(f1929,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | $spl14),
  inference(subsumption_resolution,[],[f1863,f1797])).
fof(f1863,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X1))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | $spl14),
  inference(backward_demodulation,[],[f1797,f74])).
fof(f1925,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) ) | $spl14),
  inference(subsumption_resolution,[],[f1854,f60])).
fof(f1854,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) ) | $spl14),
  inference(backward_demodulation,[],[f1797,f75])).
fof(f1923,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl14),
  inference(subsumption_resolution,[],[f1828,f1797])).
fof(f1828,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl14),
  inference(backward_demodulation,[],[f1797,f68])).
fof(f1926,plain,(
  $false | $spl14),
  inference(global_subsumption,[],[f59,f58,f57,f60,f62,f65,f64,f72,f83,f80,f84,f92,f111,f118,f150,f167,f210,f314,f366,f393,f402,f617,f618,f428,f719,f74,f73,f764,f365,f427,f183,f182,f110,f69,f68,f61,f678,f676,f672,f670,f632,f630,f626,f624,f692,f646,f79,f78,f77,f76,f75,f1797,f1923,f1925])).
fof(f1924,plain,(
  $false | $spl14),
  inference(global_subsumption,[],[f59,f58,f57,f60,f62,f65,f64,f72,f83,f80,f84,f92,f111,f118,f150,f167,f210,f314,f366,f393,f402,f617,f618,f428,f719,f74,f73,f764,f365,f427,f183,f182,f110,f69,f68,f61,f678,f676,f672,f670,f632,f630,f626,f624,f692,f646,f79,f78,f77,f76,f75,f1797,f1923])).
fof(f1922,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1827,f60])).
fof(f1827,plain,(
  ~p(s(bool,'T_0')) | $spl14),
  inference(backward_demodulation,[],[f1797,f61])).
fof(f1805,plain,(
  ( ! [X5] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X5),s(fun(sK0,bool),sK2)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X5),s(fun(sK0,bool),sK1)))) ) | $spl22),
  inference(cnf_transformation,[],[f1805_D])).
fof(f1805_D,plain,(
  ( ! [X5] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X5),s(fun(sK0,bool),sK2)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X5),s(fun(sK0,bool),sK1)))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f1818,plain,(
  ( ! [X83,X84] : (s(bool,'T_0') = s(bool,X84) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X83),s(fun(sK0,bool),sK2)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X83),s(fun(sK0,bool),sK1)))) ) | ($spl1 | $spl6)),
  inference(global_subsumption,[],[f1817,f59,f58,f57,f60,f62,f65,f64,f72,f83,f80,f84,f92,f111,f118,f150,f167,f210,f314,f366,f393,f402,f617,f618,f428,f719,f1275,f1291,f1303,f1238,f654,f74,f73,f764,f365,f427,f1149,f1270,f1250,f183,f182,f110,f69,f68,f61,f1299,f1287,f1260,f1258,f1228,f1186,f1150,f1089,f1088,f619,f91,f1171,f1170,f1146,f1145,f1110,f1109,f1085,f1084,f678,f676,f672,f670,f632,f630,f626,f624,f692,f646,f79,f78,f77,f76,f75,f1796,f1800,f1789,f1803,f1782,f1806,f1808,f1735,f1737])).
fof(f1817,plain,(
  ( ! [X83,X84,X82] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X83),s(fun(sK0,bool),sK2)))) | s(bool,X84) = s(bool,'const.bool.IN_2'(s(sK0,X82),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X84) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X83),s(fun(sK0,bool),sK1)))) ) | ($spl1 | $spl6)),
  inference(subsumption_resolution,[],[f1738,f1296])).
fof(f1738,plain,(
  ( ! [X83,X84,X82] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X83),s(fun(sK0,bool),sK2)))) | s(bool,X84) = s(bool,'const.bool.IN_2'(s(sK0,X82),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X84) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X83),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,X82),s(fun(sK0,bool),sK1)))) ) | $spl6),
  inference(superposition,[],[f619,f1250])).
fof(f1737,plain,(
  ( ! [X80,X81] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X81),s(fun(sK0,bool),sK2)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X80),s(fun(sK0,bool),sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X81),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,X80),s(fun(sK0,bool),sK2)))) ) | $spl6),
  inference(superposition,[],[f91,f1250])).
fof(f1815,plain,(
  ( ! [X78,X79] : (s(bool,'T_0') = s(bool,X79) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X78),s(fun(sK0,bool),sK2)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X78),s(fun(sK0,bool),sK1)))) ) | ($spl1 | $spl6)),
  inference(global_subsumption,[],[f1814,f59,f58,f57,f60,f62,f65,f64,f72,f83,f80,f84,f92,f111,f118,f150,f167,f210,f314,f366,f393,f402,f617,f618,f428,f719,f1275,f1291,f1303,f1238,f654,f74,f73,f764,f365,f427,f1149,f1270,f1250,f183,f182,f110,f69,f68,f61,f1299,f1287,f1260,f1258,f1228,f1186,f1150,f1089,f1088,f619,f91,f1171,f1170,f1146,f1145,f1110,f1109,f1085,f1084,f678,f676,f672,f670,f632,f630,f626,f624,f692,f646,f79,f78,f77,f76,f75,f1796,f1800,f1789,f1803,f1782,f1806,f1808,f1735])).
fof(f1814,plain,(
  ( ! [X78,X79,X77] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X78),s(fun(sK0,bool),sK2)))) | s(bool,'T_0') = s(bool,X79) | s(bool,X79) = s(bool,'const.bool.IN_2'(s(sK0,X77),s(fun(sK0,bool),sK1))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X78),s(fun(sK0,bool),sK1)))) ) | ($spl1 | $spl6)),
  inference(subsumption_resolution,[],[f1736,f1284])).
fof(f1736,plain,(
  ( ! [X78,X79,X77] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X78),s(fun(sK0,bool),sK2)))) | s(bool,'T_0') = s(bool,X79) | s(bool,X79) = s(bool,'const.bool.IN_2'(s(sK0,X77),s(fun(sK0,bool),sK1))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X78),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,X77),s(fun(sK0,bool),sK2)))) ) | $spl6),
  inference(superposition,[],[f1228,f1250])).
fof(f1735,plain,(
  ( ! [X76,X74,X72,X75,X73] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X76),s(fun(sK0,bool),sK2)))) | p(s(bool,'const.bool.IN_2'(s(X72,X73),s(fun(X72,bool),X74)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X76),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X72,X73),s(fun(X72,bool),'const.pred_set.INTER_2'(s(fun(X72,bool),X74),s(fun(X72,bool),X75)))))) ) | $spl6),
  inference(superposition,[],[f77,f1250])).
fof(f1808,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK2)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK1)))) ) | $spl6),
  inference(condensation,[],[f1807])).
fof(f1807,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X17),s(fun(sK0,bool),sK2)))) | s(bool,X16) = s(bool,X18) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X17),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X16)) ) | $spl6),
  inference(subsumption_resolution,[],[f1718,f150])).
fof(f1718,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X17),s(fun(sK0,bool),sK2)))) | ~p(s(bool,X18)) | s(bool,X16) = s(bool,X18) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X17),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X16)) ) | $spl6),
  inference(superposition,[],[f83,f1250])).
fof(f1806,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X7),s(fun(sK0,bool),sK2)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X7),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X6)) ) | $spl6),
  inference(subsumption_resolution,[],[f1714,f62])).
fof(f1714,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X7),s(fun(sK0,bool),sK2)))) | s(bool,'F_0') != s(bool,X6) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X7),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X6)) ) | $spl6),
  inference(superposition,[],[f68,f1250])).
fof(f1782,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X5),s(fun(sK0,bool),sK2)))) | s(bool,'T_0') = s(bool,X4) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X5),s(fun(sK0,bool),sK1)))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f1713])).
fof(f1713,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X5),s(fun(sK0,bool),sK2)))) | s(bool,'T_0') = s(bool,X4) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X5),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X4)) ) | $spl6),
  inference(superposition,[],[f65,f1250])).
fof(f1804,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,X3),s(fun(sK0,bool),sK2))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X3),s(fun(sK0,bool),sK1)))) ) | $spl20),
  inference(cnf_transformation,[],[f1804_D])).
fof(f1804_D,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,X3),s(fun(sK0,bool),sK2))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X3),s(fun(sK0,bool),sK1)))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f1803,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,X3),s(fun(sK0,bool),sK2))) | p(s(bool,X2)) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X3),s(fun(sK0,bool),sK1)))) ) | $spl6),
  inference(subsumption_resolution,[],[f1712,f64])).
fof(f1712,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,X3),s(fun(sK0,bool),sK2))) | p(s(bool,X2)) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X3),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X2)) ) | $spl6),
  inference(superposition,[],[f64,f1250])).
fof(f1801,plain,(
  ( ! [X21,X22] : (s(bool,'const.bool.IN_2'(s(sK0,X21),s(fun(sK0,bool),sK2))) = s(bool,'const.bool.IN_2'(s(sK0,X22),s(fun(sK0,bool),sK2))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X21),s(fun(sK0,bool),sK1))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X22),s(fun(sK0,bool),sK1)))) ) | $spl18),
  inference(cnf_transformation,[],[f1801_D])).
fof(f1801_D,plain,(
  ( ! [X21,X22] : (s(bool,'const.bool.IN_2'(s(sK0,X21),s(fun(sK0,bool),sK2))) = s(bool,'const.bool.IN_2'(s(sK0,X22),s(fun(sK0,bool),sK2))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X21),s(fun(sK0,bool),sK1))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X22),s(fun(sK0,bool),sK1)))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f1789,plain,(
  ( ! [X21,X22,X20] : (s(bool,'const.bool.IN_2'(s(sK0,X21),s(fun(sK0,bool),sK2))) = s(bool,'const.bool.IN_2'(s(sK0,X22),s(fun(sK0,bool),sK2))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X22),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X20) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X21),s(fun(sK0,bool),sK1)))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f1695])).
fof(f1695,plain,(
  ( ! [X21,X22,X20] : (s(bool,'const.bool.IN_2'(s(sK0,X21),s(fun(sK0,bool),sK2))) = s(bool,'const.bool.IN_2'(s(sK0,X22),s(fun(sK0,bool),sK2))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X22),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X20) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X21),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X20)) ) | $spl6),
  inference(superposition,[],[f1250,f1250])).
fof(f1798,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK2))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK1)))) ) | $spl16),
  inference(cnf_transformation,[],[f1798_D])).
fof(f1798_D,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK2))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK1)))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f1800,plain,(
  ( ! [X19,X16] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X19),s(fun(sK0,bool),sK2))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X19),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X16)) ) | $spl6),
  inference(forward_demodulation,[],[f1790,f110])).
fof(f1790,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,'const.bool.IN_2'(s(sK0,X19),s(fun(sK0,bool),sK2))) = s(bool,'const.bool.IN_2'(s(X17,X18),s(fun(X17,bool),'const.pred_set.EMPTY_0'))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X19),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X16)) ) | $spl6),
  inference(duplicate_literal_removal,[],[f1694])).
fof(f1694,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,'const.bool.IN_2'(s(sK0,X19),s(fun(sK0,bool),sK2))) = s(bool,'const.bool.IN_2'(s(X17,X18),s(fun(X17,bool),'const.pred_set.EMPTY_0'))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X19),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X16) | s(bool,'T_0') = s(bool,X16)) ) | $spl6),
  inference(superposition,[],[f1250,f210])).
fof(f1799,plain,(
  ( ! [X14,X15] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X15),s(fun(sK0,bool),sK2))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X15),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X14)) ) | ($spl1 | $spl6)),
  inference(forward_demodulation,[],[f1791,f100])).
fof(f100,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | $spl1),
  inference(unit_resulting_resolution,[],[f86,f69])).
fof(f1791,plain,(
  ( ! [X14,X15] : (s(bool,'const.bool.IN_2'(s(sK0,X15),s(fun(sK0,bool),sK2))) = s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X15),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X14)) ) | ($spl1 | $spl6)),
  inference(duplicate_literal_removal,[],[f1693])).
fof(f1693,plain,(
  ( ! [X14,X15] : (s(bool,'const.bool.IN_2'(s(sK0,X15),s(fun(sK0,bool),sK2))) = s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X15),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X14) | s(bool,'T_0') = s(bool,X14)) ) | ($spl1 | $spl6)),
  inference(superposition,[],[f1250,f209])).
fof(f209,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X4)) ) | $spl1),
  inference(subsumption_resolution,[],[f206,f61])).
fof(f206,plain,(
  ( ! [X4] : (p(s(bool,'F_0')) | s(bool,X4) = s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X4)) ) | $spl1),
  inference(superposition,[],[f150,f100])).
fof(f1796,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK2))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X0)) ) | $spl6),
  inference(duplicate_literal_removal,[],[f1688])).
fof(f1688,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK2))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl6),
  inference(superposition,[],[f1250,f62])).
fof(f1472,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f1308,f60])).
fof(f1308,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f92,f1276,f83])).
fof(f1471,plain,(
  $false | ($spl1 | $spl8)),
  inference(subsumption_resolution,[],[f1310,f60])).
fof(f1310,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl8)),
  inference(unit_resulting_resolution,[],[f93,f1276,f83])).
fof(f1470,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f1312,f60])).
fof(f1312,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f111,f1276,f83])).
fof(f1469,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f1314,f1276])).
fof(f1314,plain,(
  ~p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f92,f1276,f83])).
fof(f1468,plain,(
  $false | ($spl1 | $spl8)),
  inference(subsumption_resolution,[],[f1316,f1276])).
fof(f1316,plain,(
  ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) | ($spl1 | $spl8)),
  inference(unit_resulting_resolution,[],[f93,f1276,f83])).
fof(f1467,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f1318,f1276])).
fof(f1318,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f111,f1276,f83])).
fof(f1466,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f1322,f60])).
fof(f1322,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f92,f1276,f83])).
fof(f1465,plain,(
  $false | ($spl1 | $spl8)),
  inference(subsumption_resolution,[],[f1324,f60])).
fof(f1324,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl8)),
  inference(unit_resulting_resolution,[],[f93,f1276,f83])).
fof(f1464,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f1326,f60])).
fof(f1326,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f111,f1276,f83])).
fof(f1463,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f1328,f1276])).
fof(f1328,plain,(
  ~p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f92,f1276,f83])).
fof(f1462,plain,(
  $false | ($spl1 | $spl8)),
  inference(subsumption_resolution,[],[f1330,f1276])).
fof(f1330,plain,(
  ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) | ($spl1 | $spl8)),
  inference(unit_resulting_resolution,[],[f93,f1276,f83])).
fof(f1461,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f1332,f1276])).
fof(f1332,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f111,f1276,f83])).
fof(f1460,plain,(
  $false | ($spl6 | $spl8)),
  inference(subsumption_resolution,[],[f1338,f1276])).
fof(f1338,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK1))))) ) | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f1276,f91])).
fof(f1456,plain,(
  $false | ($spl6 | $spl8)),
  inference(subsumption_resolution,[],[f1455,f1276])).
fof(f1455,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl6 | $spl8)),
  inference(subsumption_resolution,[],[f1425,f1423])).
fof(f1423,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl6 | $spl8)),
  inference(global_subsumption,[],[f1422,f59,f58,f57,f60,f61,f62,f65,f64,f69,f68,f72,f74,f73,f76,f75,f79,f78,f77,f83,f80,f84,f92,f111,f110,f118,f150,f167,f182,f183,f210,f314,f365,f366,f393,f402,f617,f618,f427,f428,f624,f626,f630,f632,f646,f670,f672,f676,f678,f692,f719,f764,f1276,f1353])).
fof(f1353,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1276,f75])).
fof(f1422,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl6 | $spl8)),
  inference(subsumption_resolution,[],[f1414,f1250])).
fof(f1414,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK1))) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl6 | $spl8)),
  inference(backward_demodulation,[],[f1408,f1186])).
fof(f1408,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK2)))) ) | ($spl6 | $spl8)),
  inference(forward_demodulation,[],[f1342,f110])).
fof(f1342,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK2))) = s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')))) ) | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f111,f1276,f619])).
fof(f1425,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | ($spl6 | $spl8)),
  inference(backward_demodulation,[],[f1423,f68])).
fof(f1454,plain,(
  $false | ($spl6 | $spl8)),
  inference(subsumption_resolution,[],[f1424,f60])).
fof(f1424,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl8)),
  inference(backward_demodulation,[],[f1423,f61])).
fof(f1407,plain,(
  $false | ($spl6 | $spl8)),
  inference(subsumption_resolution,[],[f1343,f1276])).
fof(f1343,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK2))))) ) | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f1276,f91])).
fof(f1344,plain,(
  $false | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f1276,f1276,f91])).
fof(f1339,plain,(
  $false | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f1276,f1276,f91])).
fof(f1337,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f92,f1276,f65])).
fof(f1335,plain,(
  $false | ($spl1 | $spl8)),
  inference(unit_resulting_resolution,[],[f100,f1276,f68])).
fof(f1333,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1276,f111,f1276,f83])).
fof(f1331,plain,(
  $false | ($spl1 | $spl8)),
  inference(unit_resulting_resolution,[],[f1276,f93,f1276,f83])).
fof(f1329,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1276,f92,f1276,f83])).
fof(f1327,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f60,f111,f1276,f83])).
fof(f1325,plain,(
  $false | ($spl1 | $spl8)),
  inference(unit_resulting_resolution,[],[f60,f93,f1276,f83])).
fof(f1323,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f60,f92,f1276,f83])).
fof(f1319,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1276,f111,f1276,f83])).
fof(f1317,plain,(
  $false | ($spl1 | $spl8)),
  inference(unit_resulting_resolution,[],[f1276,f93,f1276,f83])).
fof(f1315,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1276,f92,f1276,f83])).
fof(f1313,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f60,f111,f1276,f83])).
fof(f1311,plain,(
  $false | ($spl1 | $spl8)),
  inference(unit_resulting_resolution,[],[f60,f93,f1276,f83])).
fof(f1309,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f60,f92,f1276,f83])).
fof(f1305,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f61,f1276])).
fof(f1292,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,X4),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X5) | s(bool,X5) = s(bool,'const.bool.IN_2'(s(sK0,X4),s(fun(sK0,bool),sK2)))) ) | $spl12),
  inference(cnf_transformation,[],[f1292_D])).
fof(f1292_D,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,X4),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X5) | s(bool,X5) = s(bool,'const.bool.IN_2'(s(sK0,X4),s(fun(sK0,bool),sK2)))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f1291,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,X4),s(fun(sK0,bool),sK1))) | p(s(bool,X3)) | s(bool,X5) = s(bool,'const.bool.IN_2'(s(sK0,X4),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X5)) ) | $spl6),
  inference(subsumption_resolution,[],[f1114,f64])).
fof(f1114,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,X4),s(fun(sK0,bool),sK1))) | p(s(bool,X3)) | s(bool,X5) = s(bool,'const.bool.IN_2'(s(sK0,X4),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,X3)) ) | $spl6),
  inference(superposition,[],[f64,f654])).
fof(f1277,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,X4),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X5) | s(bool,X5) = s(bool,'const.bool.IN_2'(s(sK0,X4),s(fun(sK0,bool),sK1)))) ) | $spl10),
  inference(cnf_transformation,[],[f1277_D])).
fof(f1277_D,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,X4),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X5) | s(bool,X5) = s(bool,'const.bool.IN_2'(s(sK0,X4),s(fun(sK0,bool),sK1)))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f1275,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,X4),s(fun(sK0,bool),sK2))) | p(s(bool,X3)) | s(bool,X5) = s(bool,'const.bool.IN_2'(s(sK0,X4),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X5)) ) | $spl6),
  inference(subsumption_resolution,[],[f1053,f64])).
fof(f1053,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,X4),s(fun(sK0,bool),sK2))) | p(s(bool,X3)) | s(bool,X5) = s(bool,'const.bool.IN_2'(s(sK0,X4),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,X5)) ) | $spl6),
  inference(superposition,[],[f64,f654])).
fof(f85,plain,(
  p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) | $spl0),
  inference(cnf_transformation,[],[f85_D])).
fof(f85_D,plain,(
  p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f91,plain,(
  ( ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X4),s(fun(sK0,bool),sK2)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X4),s(fun(sK0,bool),sK1))))) ) | $spl6),
  inference(cnf_transformation,[],[f91_D])).
fof(f91_D,plain,(
  ( ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X4),s(fun(sK0,bool),sK2)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X4),s(fun(sK0,bool),sK1))))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f57,plain,(
  ( ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X4),s(fun(sK0,bool),sK2)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X4),s(fun(sK0,bool),sK1)))) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) )),
  inference(cnf_transformation,[],[f45])).
fof(f45,plain,(
  (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X4),s(fun(sK0,bool),sK1)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X4),s(fun(sK0,bool),sK2)))))) & (~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) | (p(s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),sK1)))) & p(s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f44])).
fof(f44,plain,(
  ? [X0,X1,X2] : ((p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))))) & (~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))))),
  inference(rectify,[],[f43])).
fof(f43,plain,(
  ? [X0,X1,X2] : ((p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))) & (~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))))),
  inference(nnf_transformation,[],[f38])).
fof(f38,plain,(
  ? [X0,X1,X2] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <~> ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(ennf_transformation,[],[f18])).
fof(f18,plain,(
  ~! [X0,X1,X2] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ~? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(flattening,[],[f17])).
fof(f17,plain,(
  ~! [X0] : ! [X1,X2] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ~? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))) <=> ~? [X9] : (p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X7)))) & p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X8))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))) <=> ~? [X9] : (p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X7)))) & p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X8))))))),
  file('thm.pred_set.IN_DISJOINT','thm.pred_set.IN_DISJOINT')).
fof(f86,plain,(
  ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) | $spl1),
  inference(cnf_transformation,[],[f86_D])).
fof(f86_D,plain,(
  ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f90,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),sK1)))) | $spl4),
  inference(cnf_transformation,[],[f90_D])).
fof(f90_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),sK1)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f58,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),sK1)))) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),
  inference(cnf_transformation,[],[f45])).
fof(f88,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))) | $spl2),
  inference(cnf_transformation,[],[f88_D])).
fof(f88_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f59,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),
  inference(cnf_transformation,[],[f45])).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 129
% Passive clauses: 461
% Generated clauses: 5641
% Final active clauses: 62
% Final passive clauses: 123
% Input formulas: 15
% Initial clauses: 28
% 
% Duplicate literals: 331
% Trivial inequalities: 75
% Fw subsumption resolutions: 336
% Fw demodulations: 708
% Bw demodulations: 148
% Condensations: 314
% Global subsumptions: 8
% 
% Simple tautologies: 84
% Equational tautologies: 136
% Forward subsumptions: 3102
% Fw demodulations to eq. taut.: 2
% Bw demodulations to eq. taut.: 1
% 
% Unit resulting resolution: 586
% Factoring: 1
% Forward superposition: 1399
% Backward superposition: 1391
% Self superposition: 74
% Equality factoring: 60
% Equality resolution: 7
% 
% Split clauses: 41
% Split components: 87
% SAT solver clauses: 596
% SAT solver unit clauses: 32
% SAT solver binary clauses: 109
% 
% Sat splits: 41
% Sat splitting refutations: 43
% 
% Memory used [KB]: 3326
% Time elapsed: 0.195 s
% ------------------------------
% ------------------------------
% Success in time 1.51 s
