Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.integer
% (103681)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 181
% Passive clauses: 4552
% Generated clauses: 26735
% Final active clauses: 163
% Final passive clauses: 4195
% Input formulas: 69
% Initial clauses: 110
% 
% Duplicate literals: 808
% Trivial inequalities: 399
% Fw subsumption resolutions: 2806
% Fw demodulations: 6180
% Bw demodulations: 18
% 
% Simple tautologies: 409
% Equational tautologies: 59
% Forward subsumptions: 11253
% 
% Binary resolution: 608
% Factoring: 4
% Forward superposition: 7243
% Backward superposition: 8206
% Self superposition: 156
% Equality factoring: 111
% Equality resolution: 28
% 
% Split clauses: 218
% Split components: 438
% SAT solver clauses: 444
% SAT solver unit clauses: 6
% SAT solver binary clauses: 97
% 
% Sat splits: 218
% 
% TWLsolver clauses: 444
% TWLsolver calls for satisfiability: 51
% 
% Memory used [KB]: 15607
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 144 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.integer
% (103711)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 188
% Passive clauses: 2124
% Generated clauses: 20062
% Final active clauses: 113
% Final passive clauses: 1174
% Input formulas: 69
% Initial clauses: 110
% 
% Duplicate literals: 256
% Trivial inequalities: 192
% Fw subsumption resolutions: 1485
% Fw demodulations: 4462
% Bw demodulations: 739
% Condensations: 463
% Global subsumptions: 5
% 
% Simple tautologies: 184
% Equational tautologies: 25
% Forward subsumptions: 8756
% Bw demodulations to eq. taut.: 96
% 
% Binary resolution: 818
% Factoring: 4
% Forward superposition: 5711
% Backward superposition: 3503
% Self superposition: 76
% Equality factoring: 23
% Equality resolution: 16
% 
% Split clauses: 72
% Split components: 146
% SAT solver clauses: 2159
% SAT solver unit clauses: 129
% SAT solver binary clauses: 56
% SAT solver learnt clauses: 10
% 
% Sat splits: 72
% Sat splitting refutations: 42
% 
% TWLsolver clauses: 2139
% TWLsolver calls for satisfiability: 654
% 
% Memory used [KB]: 14072
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.integer
% (103730)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 218
% Passive clauses: 1322
% Generated clauses: 615417
% Final active clauses: 101
% Final passive clauses: 452
% Input formulas: 69
% Initial clauses: 110
% Discarded non-redundant clauses: 2221
% 
% Duplicate literals: 243
% Trivial inequalities: 163
% Fw subsumption resolutions: 1177
% Fw demodulations: 168411
% Bw demodulations: 345
% Condensations: 209
% Global subsumptions: 368
% 
% Simple tautologies: 127
% Equational tautologies: 103
% Forward subsumptions: 36931
% Fw demodulations to eq. taut.: 44589
% Bw demodulations to eq. taut.: 13
% 
% Unit resulting resolution: 435688
% Factoring: 7
% Forward superposition: 4142
% Backward superposition: 4076
% Self superposition: 102
% Equality factoring: 33
% Equality resolution: 14
% 
% Split clauses: 10
% Split components: 22
% SAT solver clauses: 2175
% SAT solver unit clauses: 339
% SAT solver binary clauses: 800
% 
% Sat splits: 10
% Sat splitting refutations: 463
% 
% Memory used [KB]: 316881
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.integer
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.integer
% SZS output start Proof for thm.integer
fof(f31271,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f136,f198,f135,f196,f134,f197,f195,f224,f249,f252,f254,f327,f396,f251,f574,f596,f623,f628,f631,f675,f790,f797,f799,f806,f758,f847,f870,f850,f900,f903,f659,f657,f996,f1040,f348,f1177,f1183,f1349,f1350,f655,f653,f1455,f1457,f1460,f1463,f1557,f1616,f1622,f1179,f1559,f1627,f1650,f1630,f1656,f1654,f1658,f1683,f1681,f1732,f1759,f1784,f1652,f1804,f1924,f2036,f1803,f2061,f2065,f2070,f2142,f2207,f2413,f3351,f3416,f3419,f3543,f3654,f3928,f3952,f4007,f4010,f4012,f4014,f4051,f4078,f4054,f4090,f4154,f4080,f4178,f4614,f4748,f3663,f3660,f3652,f3649,f1178,f4043,f4762,f4778,f4780,f4787,f4891,f4896,f4898,f5119,f5121,f5123,f5670,f5925,f6379,f3658,f3656,f6478,f3446,f7012,f7405,f7431,f7436,f7463,f7439,f7469,f7467,f7898,f4084,f4082,f8042,f7433,f8066,f8131,f7465,f8157,f8852,f8855,f8858,f8887,f8893,f8921,f8896,f8927,f8925,f8928,f8889,f8954,f8982,f8923,f9008,f10304,f10307,f10310,f10759,f10804,f10806,f10808,f10809,f10810,f6475,f3444,f3442,f10952,f10954,f10955,f11258,f11291,f11486,f11517,f11489,f11523,f11521,f11599,f11519,f11628,f11800,f11801,f11850,f11849,f11848,f11847,f11846,f11843,f11840,f11839,f11838,f11837,f11836,f11835,f11832,f11824,f11823,f11822,f11821,f11820,f11819,f11896,f11897,f11898,f11899,f11900,f11901,f11902,f11903,f11904,f11905,f13290,f13471,f14004,f15044,f15046,f15048,f15478,f15480,f15482,f15596,f15605,f15609,f15724,f15733,f15737,f16602,f10306,f16604,f10309,f16606,f10312,f16608,f15598,f16610,f15726,f16796,f16799,f16806,f16810,f16797,f16933,f16937,f16944,f16948,f17071,f17075,f17082,f17086,f18282,f18284,f18286,f18288,f18290,f18293,f18295,f18297,f18299,f18301,f18303,f18305,f18307,f18309,f18311,f18313,f18315,f18317,f18319,f18320,f18323,f18325,f18327,f18329,f18331,f18333,f18335,f18337,f18339,f18341,f18343,f18345,f18346,f18348,f18350,f18352,f18354,f18363,f18283,f18364,f18285,f18365,f18287,f18366,f18289,f18367,f18291,f18369,f18294,f18370,f18371,f18372,f18300,f18373,f18302,f18374,f18304,f18375,f18306,f18376,f18377,f18310,f18378,f18312,f18379,f18380,f18381,f18382,f18383,f18385,f18324,f18386,f18326,f18387,f18328,f18388,f18330,f18389,f18332,f18390,f18391,f18336,f18392,f18338,f18393,f18340,f18394,f18342,f18395,f18396,f18397,f18398,f18349,f18399,f18351,f18400,f18353,f18401,f18355,f18680,f18690,f19408,f19412,f19410,f19453,f19456,f19455,f20025,f20026,f18296,f20036,f20038,f20045,f20049,f16935,f20169,f20170,f18298,f20180,f20182,f20189,f20193,f17073,f20284,f20286,f20285,f21056,f21058,f21057,f21376,f21378,f21377,f21934,f22069,f22948,f22950,f22949,f23046,f23049,f23253,f23255,f23257,f23258,f23259,f23260,f23261,f23262,f23263,f23264,f23265,f23266,f23267,f23268,f23270,f23272,f23273,f23274,f23275,f23298,f23301,f23308,f23311,f23316,f23320,f23321,f23324,f23325,f23328,f23329,f23330,f23331,f23332,f23333,f23503,f869,f23504,f23515,f23509,f23518,f23513,f23742,f23743,f18318,f23776,f23777,f18344,f23804,f23805,f23048,f23856,f23859,f23858,f24061,f24064,f24063,f24410,f24411,f18321,f24475,f24478,f24477,f24731,f24734,f19367,f18880,f22847,f11317,f24736,f21965,f24735,f24738,f11286,f11284,f18981,f24740,f24817,f24818,f24820,f24989,f25282,f25287,f25288,f25292,f25298,f25302,f25300,f25304,f25321,f25323,f25329,f25331,f25332,f25333,f22229,f22226,f22199,f22196,f25358,f25368,f19835,f25375,f1181,f25376,f25387,f25393,f25396,f18504,f11318,f25429,f25430,f25441,f25389,f25445,f25395,f25446,f25541,f801,f25674,f2015,f25698,f2038,f25484,f25481,f25699,f25474,f25702,f755,f25461,f25458,f25361,f25706,f25708,f25710,f25712,f25714,f25717,f25719,f25726,f1558,f25730,f25707,f25731,f25330,f25732,f25335,f25735,f25728,f25743,f25745,f25715,f25747,f25394,f25748,f25700,f25716,f25750,f25360,f25752,f25704,f25734,f10294,f25809,f25968,f875,f25971,f25972,f18316,f25995,f25996,f26003,f26062,f26065,f26064,f26316,f26319,f26318,f26350,f26353,f26352,f26506,f26509,f26508,f26641,f26763,f26764,f18314,f26886,f26887,f18334,f26908,f26909,f18308,f26935,f26938,f26937,f26972,f26975,f26974,f26977,f27484,f27485,f4750,f18347,f27590,f27792,f27795,f27794,f19411,f28130,f28220,f28225,f28226,f28227,f28498,f28696,f28698,f598,f28494,f28702,f28703,f28700,f28712,f28715,f28721,f28724,f28701,f199,f28736,f28754,f28881,f28884,f28883,f28768,f25188,f25189,f28767,f28766,f25711,f28933,f28934,f28935,f22118,f22121,f800,f28903,f28951,f29023,f29241,f29244,f29246,f28902,f29300,f29303,f29468,f29587,f29590,f29591,f29594,f876,f874,f29774,f794,f792,f30039,f754,f756,f30068,f30094,f30095,f30145,f30146,f30173,f30174,f30198,f7404,f3448,f30200,f30201,f791,f803,f4048,f7006,f7791,f1456,f30205,f30209,f30370,f30432,f30436,f30438,f30440,f30434,f30618,f30730,f30759,f349,f347,f30898,f30986,f31092,f30760,f31094,f30260,f872,f200,f31095,f31213,f31269])).
fof(f31269,plain,(
  $false | ($spl1 | $spl46)),
  inference(subsumption_resolution,[],[f31249,f187])).
fof(f187,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X0))))) )),
  inference(cnf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0] : p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X0))))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X9] : p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X9),s('type.integer.int',X9))))),
  file('thm.integer.INT_ABS_LE|split|1','thm.integer.INT_LE_REFL')).
fof(f31249,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0)))))) | ($spl1 | $spl46)),
  inference(resolution,[],[f31217,f137])).
fof(f137,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f116])).
fof(f116,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int',X1)))) | (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0)))))) & (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int',X1))))))),
  inference(flattening,[],[f115])).
fof(f115,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int',X1)))) | (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0)))))) & ((~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int',X1))))))),
  inference(nnf_transformation,[],[f72])).
fof(f72,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int',X1)))) <=> (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0))))))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X6,X7] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X6))),s('type.integer.int',X7)))) <=> (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X6),s('type.integer.int',X7)))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X7))),s('type.integer.int',X6))))))),
  file('thm.integer.INT_ABS_LE|split|1','thm.integer.INT_ABS_LE|split|0')).
fof(f31217,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0)))))) | ($spl1 | $spl46)),
  inference(forward_demodulation,[],[f196,f872])).
fof(f31213,plain,(
  $false | ($spl44 | $spl46)),
  inference(subsumption_resolution,[],[f31212,f24590])).
fof(f24590,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X0))))) )),
  inference(factoring,[],[f169])).
fof(f169,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f111])).
fof(f111,plain,(
  ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(ennf_transformation,[],[f92])).
fof(f92,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) => ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(flattening,[],[f91])).
fof(f91,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) => ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X9,X10] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X9),s('type.integer.int',X10)))) => ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X10),s('type.integer.int',X9)))))),
  file('thm.integer.INT_ABS_LE|split|1','thm.integer.INT_LT_GT')).
fof(f31212,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK0),s('type.integer.int',sK0)))) | ($spl44 | $spl46)),
  inference(forward_demodulation,[],[f870,f872])).
fof(f31095,plain,(
  $false | ($spl7 | $spl20 | $spl45)),
  inference(subsumption_resolution,[],[f1348,f23606])).
fof(f23606,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int',sK1)))) | $spl45),
  inference(resolution,[],[f158,f869])).
fof(f158,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0,X1] : ((p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0))))) & (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))),
  inference(nnf_transformation,[],[f86])).
fof(f86,plain,(
  ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) <=> p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(flattening,[],[f85])).
fof(f85,plain,(
  ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) <=> p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X9,X10] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X9),s('type.integer.int',X10)))) <=> p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X10),s('type.integer.int',X9)))))),
  file('thm.integer.INT_ABS_LE|split|1','thm.integer.INT_NOT_LT')).
fof(f1348,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int',sK1)))) | ($spl7 | $spl20)),
  inference(subsumption_resolution,[],[f1324,f1020])).
fof(f1020,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int',sK0)))) | $spl7),
  inference(resolution,[],[f248,f161])).
fof(f161,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0))))) )),
  inference(cnf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0))))) & (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))),
  inference(nnf_transformation,[],[f88])).
fof(f88,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) <=> ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(flattening,[],[f87])).
fof(f87,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) <=> ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X9,X10] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X9),s('type.integer.int',X10)))) <=> ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X10),s('type.integer.int',X9)))))),
  file('thm.integer.INT_ABS_LE|split|1','thm.integer.int_le')).
fof(f248,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | $spl7),
  inference(cnf_transformation,[],[f248_D])).
fof(f248_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f1324,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int',sK0)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int',sK1)))) | $spl20),
  inference(resolution,[],[f1064,f139])).
fof(f139,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f116])).
fof(f1064,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int',sK1)))) | $spl20),
  inference(resolution,[],[f653,f160])).
fof(f160,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f120])).
fof(f200,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0)))) | $spl5),
  inference(cnf_transformation,[],[f200_D])).
fof(f200_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0)))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f872,plain,(
  s('type.integer.int',sK0) = s('type.integer.int',sK1) | $spl46),
  inference(cnf_transformation,[],[f872_D])).
fof(f872_D,plain,(
  s('type.integer.int',sK0) = s('type.integer.int',sK1) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f30260,plain,(
  s('type.integer.int',sK0) = s('type.integer.int',sK1) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0)))) | $spl45),
  inference(resolution,[],[f869,f162])).
fof(f162,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X0) = s('type.integer.int',X1) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | s('type.integer.int',X0) = s('type.integer.int',X1)) & ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & s('type.integer.int',X0) != s('type.integer.int',X1)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))),
  inference(flattening,[],[f121])).
fof(f121,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | s('type.integer.int',X0) = s('type.integer.int',X1))) & ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & s('type.integer.int',X0) != s('type.integer.int',X1)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))),
  inference(nnf_transformation,[],[f89])).
fof(f89,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) <=> (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | s('type.integer.int',X0) = s('type.integer.int',X1)))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X9,X10] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X9),s('type.integer.int',X10)))) <=> (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X9),s('type.integer.int',X10)))) | s('type.integer.int',X9) = s('type.integer.int',X10)))),
  file('thm.integer.INT_ABS_LE|split|1','thm.integer.INT_LE_LT')).
fof(f31094,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK0),s('type.integer.int',sK1)))) | $spl344),
  inference(cnf_transformation,[],[f31094_D])).
fof(f31094_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK0),s('type.integer.int',sK1)))) <=> ~$spl344),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl344])])).
fof(f30760,plain,(
  s('type.integer.int',sK0) = s('type.integer.int',sK1) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK0),s('type.integer.int',sK1)))) | $spl45),
  inference(resolution,[],[f23606,f162])).
fof(f31092,plain,(
  $false | ($spl3 | $spl6)),
  inference(subsumption_resolution,[],[f31070,f198])).
fof(f31070,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | $spl6),
  inference(resolution,[],[f249,f163])).
fof(f163,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f122])).
fof(f30986,plain,(
  $false | $spl338),
  inference(subsumption_resolution,[],[f30434,f24590])).
fof(f30898,plain,(
  $false | $spl327),
  inference(subsumption_resolution,[],[f28700,f187])).
fof(f347,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int',sK0)))) | $spl13),
  inference(cnf_transformation,[],[f347_D])).
fof(f347_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int',sK0)))) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f349,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int',sK1)))) | $spl14),
  inference(cnf_transformation,[],[f349_D])).
fof(f349_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int',sK1)))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f30759,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int',sK1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int',sK0)))) | $spl45),
  inference(resolution,[],[f23606,f139])).
fof(f30730,plain,(
  $false | ($spl38 | $spl42)),
  inference(subsumption_resolution,[],[f30312,f24590])).
fof(f30312,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK0),s('type.integer.int',sK0)))) | ($spl38 | $spl42)),
  inference(backward_demodulation,[],[f803,f799])).
fof(f30618,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int',sK1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int',sK0)))) | ($spl5 | $spl42)),
  inference(resolution,[],[f30566,f139])).
fof(f30566,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int',sK1)))) | ($spl5 | $spl42)),
  inference(resolution,[],[f30371,f200])).
fof(f30371,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',sK0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int',X0))))) ) | $spl42),
  inference(forward_demodulation,[],[f30314,f803])).
fof(f30314,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl42),
  inference(backward_demodulation,[],[f803,f3027])).
fof(f3027,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))))) )),
  inference(forward_demodulation,[],[f142,f145])).
fof(f145,plain,(
  ( ! [X0] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))) )),
  inference(cnf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0] : s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X9] : s(bool,'const.integer.int_le_2'(s('type.integer.int',X9),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9)))))),
  file('thm.integer.INT_ABS_LE|split|1','thm.integer.INT_NEG_GE|48|')).
fof(f142,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))))) )),
  inference(cnf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X9] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X9)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9)))))))),
  file('thm.integer.INT_ABS_LE|split|1','thm.integer.INT_LE_NEGTOTAL')).
fof(f30434,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK0),s('type.integer.int',sK0)))) | $spl338),
  inference(cnf_transformation,[],[f30434_D])).
fof(f30434_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK0),s('type.integer.int',sK0)))) <=> ~$spl338),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl338])])).
fof(f30440,plain,(
  s('type.integer.int',sK0) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | $spl342),
  inference(cnf_transformation,[],[f30440_D])).
fof(f30440_D,plain,(
  s('type.integer.int',sK0) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) <=> ~$spl342),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl342])])).
fof(f30438,plain,(
  s('type.integer.int',sK0) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK0),s('type.integer.int',sK0)))) | $spl42),
  inference(forward_demodulation,[],[f30437,f803])).
fof(f30437,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK0),s('type.integer.int',sK0)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl42),
  inference(forward_demodulation,[],[f20269,f803])).
fof(f20269,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(forward_demodulation,[],[f20222,f9243])).
fof(f9243,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(subsumption_resolution,[],[f9220,f2508])).
fof(f2508,plain,(
  ( ! [X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(resolution,[],[f190,f160])).
fof(f190,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))))) )),
  inference(cnf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0] : p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6] : p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X6))))))),
  file('thm.integer.INT_ABS_LE|split|1','thm.integer.INT_ABS_POS')).
fof(f9220,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(resolution,[],[f9151,f162])).
fof(f9151,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(resolution,[],[f2783,f187])).
fof(f2783,plain,(
  ( ! [X1] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X1))),s('type.integer.int',X1))))) )),
  inference(subsumption_resolution,[],[f2753,f187])).
fof(f2753,plain,(
  ( ! [X1] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X1))),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X1))))) )),
  inference(superposition,[],[f139,f149])).
fof(f149,plain,(
  ( ! [X0] : (s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int',X0)))) )),
  inference(cnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0] : s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int',X0)))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X9] : s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X9))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9))),s('type.integer.int',X9)))),
  file('thm.integer.INT_ABS_LE|split|1','thm.integer.INT_LE_NEGL')).
fof(f20222,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(superposition,[],[f2507,f151])).
fof(f151,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(cnf_transformation,[],[f33])).
fof(f33,axiom,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.integer.INT_ABS_LE|split|1','thm.integer.INT_NEG_|48|')).
fof(f2507,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))) )),
  inference(resolution,[],[f190,f162])).
fof(f30436,plain,(
  s('type.integer.int',sK0) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))))) | $spl340),
  inference(cnf_transformation,[],[f30436_D])).
fof(f30436_D,plain,(
  s('type.integer.int',sK0) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))))) <=> ~$spl340),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl340])])).
fof(f30432,plain,(
  s('type.integer.int',sK0) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK0),s('type.integer.int',sK0)))) | $spl42),
  inference(forward_demodulation,[],[f30431,f803])).
fof(f30431,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK0),s('type.integer.int',sK0)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl42),
  inference(forward_demodulation,[],[f20259,f803])).
fof(f20259,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(forward_demodulation,[],[f20218,f9243])).
fof(f20218,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(superposition,[],[f2507,f9243])).
fof(f30370,plain,(
  $false | ($spl38 | $spl42 | $spl262)),
  inference(subsumption_resolution,[],[f30312,f23518])).
fof(f30209,plain,(
  $false | ($spl5 | $spl44)),
  inference(subsumption_resolution,[],[f24268,f870])).
fof(f24268,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK0)))) | $spl5),
  inference(resolution,[],[f163,f200])).
fof(f30205,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl336),
  inference(cnf_transformation,[],[f30205_D])).
fof(f30205_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl336),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl336])])).
fof(f1456,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | $spl59),
  inference(cnf_transformation,[],[f1456_D])).
fof(f1456_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) <=> ~$spl59),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl59])])).
fof(f7791,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl86),
  inference(forward_demodulation,[],[f7790,f148])).
fof(f148,plain,(
  ( ! [X0] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))) )),
  inference(cnf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0] : s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X9] : s(bool,'const.integer.int_le_2'(s('type.integer.int',X9),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',X9),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9)))))),
  file('thm.integer.INT_ABS_LE|split|1','thm.integer.INT_LE_NEGR')).
fof(f7790,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl86),
  inference(forward_demodulation,[],[f7789,f140])).
fof(f140,plain,(
  ( ! [X0,X1] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))) )),
  inference(cnf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0,X1] : s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X9,X10] : s(bool,'const.integer.int_le_2'(s('type.integer.int',X10),s('type.integer.int',X9))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X10)))))),
  file('thm.integer.INT_ABS_LE|split|1','thm.integer.INT_LE_NEG')).
fof(f7789,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | $spl86),
  inference(subsumption_resolution,[],[f7757,f193])).
fof(f193,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f170])).
fof(f170,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f94])).
fof(f94,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f93])).
fof(f93,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X8] : ! [X9] : (s(X8,X9) = s(X8,X9) <=> p(s(bool,'T_0')))),
  file('thm.integer.INT_ABS_LE|split|1','thm.bool.REFL_CLAUSE')).
fof(f7757,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | $spl86),
  inference(superposition,[],[f139,f4020])).
fof(f4020,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl86),
  inference(resolution,[],[f3655,f186])).
fof(f186,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f132])).
fof(f132,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f101])).
fof(f101,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X11] : (s(bool,'T_0') = s(bool,X11) <=> p(s(bool,X11)))),
  file('thm.integer.INT_ABS_LE|split|1','thm.bool.EQ_CLAUSES|split|1')).
fof(f3655,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl86),
  inference(cnf_transformation,[],[f3655_D])).
fof(f3655_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f7006,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl86),
  inference(forward_demodulation,[],[f7005,f148])).
fof(f7005,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl86),
  inference(forward_demodulation,[],[f6977,f140])).
fof(f6977,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | $spl86),
  inference(resolution,[],[f3655,f139])).
fof(f4048,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl86),
  inference(forward_demodulation,[],[f4047,f148])).
fof(f4047,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl86),
  inference(forward_demodulation,[],[f4019,f140])).
fof(f4019,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | $spl86),
  inference(resolution,[],[f3655,f139])).
fof(f803,plain,(
  s('type.integer.int',sK0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl42),
  inference(cnf_transformation,[],[f803_D])).
fof(f803_D,plain,(
  s('type.integer.int',sK0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f791,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK0)))) | $spl35),
  inference(cnf_transformation,[],[f791_D])).
fof(f791_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK0)))) <=> ~$spl35),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl35])])).
fof(f30201,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK0)))) | s('type.integer.int',sK0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl38),
  inference(forward_demodulation,[],[f27629,f149])).
fof(f27629,plain,(
  s('type.integer.int',sK0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK0)))) | $spl38),
  inference(resolution,[],[f26323,f162])).
fof(f26323,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK0)))) | $spl38),
  inference(resolution,[],[f799,f169])).
fof(f30200,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))),s('type.integer.int',sK1)))) | $spl335),
  inference(cnf_transformation,[],[f30200_D])).
fof(f30200_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))),s('type.integer.int',sK1)))) <=> ~$spl335),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl335])])).
fof(f3448,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))))),s('type.integer.int',sK1)))) | $spl84),
  inference(cnf_transformation,[],[f3448_D])).
fof(f3448_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))))),s('type.integer.int',sK1)))) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f7404,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))))),s('type.integer.int',sK1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))),s('type.integer.int',sK1)))) | $spl82),
  inference(resolution,[],[f3445,f139])).
fof(f3445,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))))) | $spl82),
  inference(cnf_transformation,[],[f3445_D])).
fof(f3445_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))))) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f30198,plain,(
  $false | ($spl1 | $spl4 | $spl47)),
  inference(subsumption_resolution,[],[f569,f29797])).
fof(f29797,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int',sK1)))) | ($spl4 | $spl47)),
  inference(subsumption_resolution,[],[f29779,f871])).
fof(f871,plain,(
  s('type.integer.int',sK0) != s('type.integer.int',sK1) | $spl47),
  inference(cnf_transformation,[],[f871_D])).
fof(f871_D,plain,(
  s('type.integer.int',sK0) != s('type.integer.int',sK1) <=> ~$spl47),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl47])])).
fof(f29779,plain,(
  s('type.integer.int',sK0) = s('type.integer.int',sK1) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int',sK1)))) | $spl4),
  inference(resolution,[],[f848,f162])).
fof(f848,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK0),s('type.integer.int',sK1)))) | $spl4),
  inference(resolution,[],[f199,f160])).
fof(f569,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int',sK1)))) | $spl1),
  inference(resolution,[],[f351,f137])).
fof(f351,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int',sK1)))) | $spl1),
  inference(resolution,[],[f201,f161])).
fof(f201,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0)))))) | $spl1),
  inference(resolution,[],[f196,f163])).
fof(f30174,plain,(
  $false | ($spl8 | $spl44 | $spl270)),
  inference(subsumption_resolution,[],[f30151,f870])).
fof(f30151,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK0)))) | ($spl8 | $spl270)),
  inference(resolution,[],[f26639,f159])).
fof(f159,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0))))) )),
  inference(cnf_transformation,[],[f119])).
fof(f26639,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int',sK1)))) | ($spl8 | $spl270)),
  inference(forward_demodulation,[],[f26638,f680])).
fof(f680,plain,(
  s('type.integer.int',sK1) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))) | $spl8),
  inference(superposition,[],[f150,f251])).
fof(f150,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))) )),
  inference(cnf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0] : s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X9] : s('type.integer.int',X9) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9)))))),
  file('thm.integer.INT_ABS_LE|split|1','thm.integer.INT_NEGNEG')).
fof(f26638,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))) | $spl270),
  inference(forward_demodulation,[],[f24735,f148])).
fof(f30173,plain,(
  $false | ($spl4 | $spl8 | $spl47 | $spl270)),
  inference(subsumption_resolution,[],[f30172,f848])).
fof(f30172,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK0),s('type.integer.int',sK1)))) | ($spl8 | $spl47 | $spl270)),
  inference(subsumption_resolution,[],[f30150,f871])).
fof(f30150,plain,(
  s('type.integer.int',sK0) = s('type.integer.int',sK1) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK0),s('type.integer.int',sK1)))) | ($spl8 | $spl270)),
  inference(resolution,[],[f26639,f162])).
fof(f30146,plain,(
  $false | ($spl8 | $spl36 | $spl44)),
  inference(subsumption_resolution,[],[f30123,f870])).
fof(f30123,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK0)))) | ($spl8 | $spl36)),
  inference(resolution,[],[f26321,f159])).
fof(f26321,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int',sK1)))) | ($spl8 | $spl36)),
  inference(forward_demodulation,[],[f794,f680])).
fof(f30145,plain,(
  $false | ($spl4 | $spl8 | $spl36 | $spl47)),
  inference(subsumption_resolution,[],[f30144,f848])).
fof(f30144,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK0),s('type.integer.int',sK1)))) | ($spl8 | $spl36 | $spl47)),
  inference(subsumption_resolution,[],[f30122,f871])).
fof(f30122,plain,(
  s('type.integer.int',sK0) = s('type.integer.int',sK1) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK0),s('type.integer.int',sK1)))) | ($spl8 | $spl36)),
  inference(resolution,[],[f26321,f162])).
fof(f30095,plain,(
  $false | ($spl8 | $spl44 | $spl52)),
  inference(subsumption_resolution,[],[f30072,f870])).
fof(f30072,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK0)))) | ($spl8 | $spl52)),
  inference(resolution,[],[f25587,f159])).
fof(f25587,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int',sK1)))) | ($spl8 | $spl52)),
  inference(forward_demodulation,[],[f1178,f251])).
fof(f30094,plain,(
  $false | ($spl4 | $spl8 | $spl47 | $spl52)),
  inference(subsumption_resolution,[],[f30093,f848])).
fof(f30093,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK0),s('type.integer.int',sK1)))) | ($spl8 | $spl47 | $spl52)),
  inference(subsumption_resolution,[],[f30071,f871])).
fof(f30071,plain,(
  s('type.integer.int',sK0) = s('type.integer.int',sK1) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK0),s('type.integer.int',sK1)))) | ($spl8 | $spl52)),
  inference(resolution,[],[f25587,f162])).
fof(f30068,plain,(
  $false | ($spl4 | $spl8 | $spl59)),
  inference(subsumption_resolution,[],[f25272,f199])).
fof(f25272,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0)))) | ($spl8 | $spl59)),
  inference(forward_demodulation,[],[f1456,f251])).
fof(f756,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | $spl30),
  inference(cnf_transformation,[],[f756_D])).
fof(f756_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f754,plain,(
  s('type.integer.int',sK1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl28),
  inference(cnf_transformation,[],[f754_D])).
fof(f754_D,plain,(
  s('type.integer.int',sK1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f30039,plain,(
  s('type.integer.int',sK1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | ($spl8 | $spl38)),
  inference(forward_demodulation,[],[f30038,f680])).
fof(f30038,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ($spl8 | $spl38)),
  inference(subsumption_resolution,[],[f30037,f26323])).
fof(f30037,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK0)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl8),
  inference(forward_demodulation,[],[f29967,f147])).
fof(f147,plain,(
  ( ! [X0] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(cnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0] : s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X9] : s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X9))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.integer.INT_ABS_LE|split|1','thm.integer.INT_NEG_LT|48|')).
fof(f29967,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl8),
  inference(superposition,[],[f3355,f680])).
fof(f3355,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f143,f146])).
fof(f146,plain,(
  ( ! [X0] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))) )),
  inference(cnf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0] : s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X9] : s(bool,'const.integer.int_lt_2'(s('type.integer.int',X9),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9)))))),
  file('thm.integer.INT_ABS_LE|split|1','thm.integer.INT_NEG_GT|48|')).
fof(f143,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X9] : (s('type.integer.int',X9) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X9)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9)))))))),
  file('thm.integer.INT_ABS_LE|split|1','thm.integer.INT_LT_NEGTOTAL')).
fof(f792,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK0)))) | $spl34),
  inference(cnf_transformation,[],[f792_D])).
fof(f792_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK0)))) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f794,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))) | $spl36),
  inference(cnf_transformation,[],[f794_D])).
fof(f794_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f29774,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK0)))) | $spl8),
  inference(forward_demodulation,[],[f29773,f148])).
fof(f29773,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK0)))) | $spl8),
  inference(forward_demodulation,[],[f29772,f251])).
fof(f29772,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl8),
  inference(forward_demodulation,[],[f29710,f149])).
fof(f29710,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl8),
  inference(superposition,[],[f3027,f251])).
fof(f874,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK1)))) | $spl49),
  inference(cnf_transformation,[],[f874_D])).
fof(f874_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK1)))) <=> ~$spl49),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl49])])).
fof(f876,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int',sK0)))) | $spl50),
  inference(cnf_transformation,[],[f876_D])).
fof(f876_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int',sK0)))) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f29594,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int',sK0)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK1)))) | $spl4),
  inference(resolution,[],[f199,f139])).
fof(f29591,plain,(
  $false | ($spl8 | $spl46 | $spl73)),
  inference(subsumption_resolution,[],[f25271,f872])).
fof(f25271,plain,(
  s('type.integer.int',sK0) != s('type.integer.int',sK1) | ($spl8 | $spl73)),
  inference(forward_demodulation,[],[f1682,f251])).
fof(f1682,plain,(
  s('type.integer.int',sK1) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))) | $spl73),
  inference(cnf_transformation,[],[f1682_D])).
fof(f1682_D,plain,(
  s('type.integer.int',sK1) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))) <=> ~$spl73),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl73])])).
fof(f29590,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))) | $spl46),
  inference(forward_demodulation,[],[f29589,f872])).
fof(f29589,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int',sK1)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))) | $spl46),
  inference(forward_demodulation,[],[f29588,f149])).
fof(f29588,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | $spl46),
  inference(forward_demodulation,[],[f29526,f148])).
fof(f29526,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | $spl46),
  inference(superposition,[],[f3027,f872])).
fof(f29587,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK0)))) | $spl46),
  inference(forward_demodulation,[],[f29586,f872])).
fof(f29586,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK0)))) | $spl46),
  inference(forward_demodulation,[],[f29585,f148])).
fof(f29585,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl46),
  inference(forward_demodulation,[],[f29525,f149])).
fof(f29525,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl46),
  inference(superposition,[],[f3027,f872])).
fof(f29468,plain,(
  $false | ($spl1 | $spl8 | $spl72)),
  inference(subsumption_resolution,[],[f29448,f187])).
fof(f29448,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0)))))) | ($spl1 | $spl8 | $spl72)),
  inference(resolution,[],[f29297,f137])).
fof(f29297,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0)))))) | ($spl1 | $spl8 | $spl72)),
  inference(forward_demodulation,[],[f196,f28952])).
fof(f28952,plain,(
  s('type.integer.int',sK0) = s('type.integer.int',sK1) | ($spl8 | $spl72)),
  inference(forward_demodulation,[],[f1683,f251])).
fof(f29303,plain,(
  $false | ($spl8 | $spl49)),
  inference(subsumption_resolution,[],[f28885,f187])).
fof(f28885,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK1)))) | ($spl8 | $spl49)),
  inference(forward_demodulation,[],[f874,f680])).
fof(f29300,plain,(
  $false | ($spl8 | $spl44 | $spl72 | $spl262)),
  inference(subsumption_resolution,[],[f29299,f23518])).
fof(f29299,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK0),s('type.integer.int',sK0)))) | ($spl8 | $spl44 | $spl72)),
  inference(forward_demodulation,[],[f870,f28952])).
fof(f28902,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int',sK0)))) | $spl330),
  inference(cnf_transformation,[],[f28902_D])).
fof(f28902_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int',sK0)))) <=> ~$spl330),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl330])])).
fof(f29246,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int',sK0)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK0)))) | ($spl2 | $spl8)),
  inference(resolution,[],[f25277,f139])).
fof(f25277,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int',sK0)))) | ($spl2 | $spl8)),
  inference(forward_demodulation,[],[f197,f251])).
fof(f29244,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))) | ($spl8 | $spl72)),
  inference(forward_demodulation,[],[f29243,f28952])).
fof(f29243,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int',sK1)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))) | ($spl8 | $spl72)),
  inference(forward_demodulation,[],[f29242,f149])).
fof(f29242,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | ($spl8 | $spl72)),
  inference(forward_demodulation,[],[f29180,f148])).
fof(f29180,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | ($spl8 | $spl72)),
  inference(superposition,[],[f3027,f28952])).
fof(f29241,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK0)))) | ($spl8 | $spl72)),
  inference(forward_demodulation,[],[f29240,f28952])).
fof(f29240,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK0)))) | ($spl8 | $spl72)),
  inference(forward_demodulation,[],[f29239,f148])).
fof(f29239,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl8 | $spl72)),
  inference(forward_demodulation,[],[f29179,f149])).
fof(f29179,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl8 | $spl72)),
  inference(superposition,[],[f3027,f28952])).
fof(f29023,plain,(
  $false | ($spl1 | $spl8 | $spl28)),
  inference(trivial_inequality_removal,[],[f28959])).
fof(f28959,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))) | ($spl1 | $spl8 | $spl28)),
  inference(backward_demodulation,[],[f28950,f545])).
fof(f545,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))) | $spl1),
  inference(resolution,[],[f205,f164])).
fof(f164,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | s('type.integer.int',X0) != s('type.integer.int',X1)) )),
  inference(cnf_transformation,[],[f122])).
fof(f205,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0)))))) | $spl1),
  inference(resolution,[],[f196,f137])).
fof(f28950,plain,(
  s('type.integer.int',sK0) = s('type.integer.int',sK1) | ($spl8 | $spl28)),
  inference(forward_demodulation,[],[f2081,f251])).
fof(f2081,plain,(
  s('type.integer.int',sK1) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))) | $spl28),
  inference(backward_demodulation,[],[f754,f151])).
fof(f28951,plain,(
  $false | ($spl8 | $spl28 | $spl73)),
  inference(subsumption_resolution,[],[f28950,f25271])).
fof(f28903,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3)) ) | $spl332),
  inference(cnf_transformation,[],[f28903_D])).
fof(f28903_D,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3)) ) <=> ~$spl332),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl332])])).
fof(f800,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl41),
  inference(cnf_transformation,[],[f800_D])).
fof(f800_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl41),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl41])])).
fof(f22121,plain,(
  ( ! [X3] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,'T_0') != s(bool,X3)) ) | $spl96),
  inference(forward_demodulation,[],[f22120,f146])).
fof(f22120,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))))) ) | $spl96),
  inference(subsumption_resolution,[],[f22100,f183])).
fof(f183,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f100])).
fof(f100,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X11] : (s(bool,'T_0') = s(bool,X11) <=> p(s(bool,X11)))),
  file('thm.integer.INT_ABS_LE|split|1','thm.bool.EQ_CLAUSES|split|0')).
fof(f22100,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,X3)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))))) ) | $spl96),
  inference(superposition,[],[f22089,f168])).
fof(f168,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f110])).
fof(f110,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f109])).
fof(f109,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f90])).
fof(f90,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X12,X13] : ((p(s(bool,X12)) => p(s(bool,X13))) => ((p(s(bool,X13)) => p(s(bool,X12))) => s(bool,X12) = s(bool,X13)))),
  file('thm.integer.INT_ABS_LE|split|1','thm.bool.IMP_ANTISYM_AX')).
fof(f22089,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | $spl96),
  inference(forward_demodulation,[],[f8933,f9243])).
fof(f8933,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | $spl96),
  inference(resolution,[],[f8859,f183])).
fof(f8859,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | $spl96),
  inference(resolution,[],[f4081,f160])).
fof(f4081,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl96),
  inference(cnf_transformation,[],[f4081_D])).
fof(f4081_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f22118,plain,(
  ( ! [X2] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,'T_0') != s(bool,X2)) ) | $spl96),
  inference(forward_demodulation,[],[f22117,f146])).
fof(f22117,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))))) ) | $spl96),
  inference(subsumption_resolution,[],[f22099,f183])).
fof(f22099,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | ~p(s(bool,X2))) ) | $spl96),
  inference(superposition,[],[f22089,f168])).
fof(f28935,plain,(
  $false | $spl332),
  inference(trivial_inequality_removal,[],[f28924])).
fof(f28924,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl332),
  inference(superposition,[],[f28903,f2511])).
fof(f2511,plain,(
  ( ! [X4] : (s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X4)))))) )),
  inference(resolution,[],[f190,f186])).
fof(f28934,plain,(
  $false | $spl332),
  inference(trivial_inequality_removal,[],[f28929])).
fof(f28929,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl332),
  inference(superposition,[],[f28903,f9273])).
fof(f9273,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(superposition,[],[f2511,f9243])).
fof(f28933,plain,(
  $false | $spl332),
  inference(equality_resolution,[],[f28903])).
fof(f25711,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl300),
  inference(cnf_transformation,[],[f25711_D])).
fof(f25711_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl300),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl300])])).
fof(f28766,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0)))) | ($spl8 | $spl86)),
  inference(forward_demodulation,[],[f25270,f251])).
fof(f25270,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl8 | $spl86)),
  inference(forward_demodulation,[],[f4048,f251])).
fof(f28767,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0)))) | ($spl8 | $spl86)),
  inference(forward_demodulation,[],[f25199,f251])).
fof(f25199,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl8 | $spl86)),
  inference(forward_demodulation,[],[f7006,f251])).
fof(f25189,plain,(
  ( ! [X2] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0)))) | s(bool,'T_0') != s(bool,X2)) ) | ($spl8 | $spl29 | $spl41)),
  inference(forward_demodulation,[],[f7354,f251])).
fof(f7354,plain,(
  ( ! [X2] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | s(bool,'T_0') != s(bool,X2)) ) | ($spl29 | $spl41)),
  inference(forward_demodulation,[],[f7353,f148])).
fof(f7353,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl29 | $spl41)),
  inference(subsumption_resolution,[],[f7347,f183])).
fof(f7347,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | ~p(s(bool,X2)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl29 | $spl41)),
  inference(superposition,[],[f5951,f168])).
fof(f5951,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl29 | $spl41)),
  inference(resolution,[],[f5765,f183])).
fof(f5765,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl29 | $spl41)),
  inference(subsumption_resolution,[],[f5744,f753])).
fof(f753,plain,(
  s('type.integer.int',sK1) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl29),
  inference(cnf_transformation,[],[f753_D])).
fof(f753_D,plain,(
  s('type.integer.int',sK1) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f5744,plain,(
  s('type.integer.int',sK1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl41),
  inference(resolution,[],[f800,f162])).
fof(f25188,plain,(
  ( ! [X3] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0)))) | s(bool,'T_0') != s(bool,X3)) ) | ($spl8 | $spl29 | $spl41)),
  inference(forward_demodulation,[],[f7356,f251])).
fof(f7356,plain,(
  ( ! [X3] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | s(bool,'T_0') != s(bool,X3)) ) | ($spl29 | $spl41)),
  inference(forward_demodulation,[],[f7355,f148])).
fof(f7355,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl29 | $spl41)),
  inference(subsumption_resolution,[],[f7348,f183])).
fof(f7348,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,X3))) ) | ($spl29 | $spl41)),
  inference(superposition,[],[f5951,f168])).
fof(f28768,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0)))) | ($spl8 | $spl86)),
  inference(forward_demodulation,[],[f25175,f251])).
fof(f25175,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl8 | $spl86)),
  inference(forward_demodulation,[],[f7791,f251])).
fof(f28883,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int',sK0)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0)))) | ($spl7 | $spl8)),
  inference(forward_demodulation,[],[f3723,f251])).
fof(f3723,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))),s('type.integer.int',sK0)))) | $spl7),
  inference(forward_demodulation,[],[f3722,f140])).
fof(f3722,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))),s('type.integer.int',sK0)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | $spl7),
  inference(subsumption_resolution,[],[f3689,f193])).
fof(f3689,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))),s('type.integer.int',sK0)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | $spl7),
  inference(superposition,[],[f139,f1250])).
fof(f1250,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int',sK0))) | $spl7),
  inference(resolution,[],[f1020,f186])).
fof(f28884,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int',sK0)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0)))) | ($spl7 | $spl8)),
  inference(forward_demodulation,[],[f1271,f251])).
fof(f1271,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))),s('type.integer.int',sK0)))) | $spl7),
  inference(forward_demodulation,[],[f1249,f140])).
fof(f1249,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))),s('type.integer.int',sK0)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | $spl7),
  inference(resolution,[],[f1020,f139])).
fof(f28881,plain,(
  $false | ($spl1 | $spl8 | $spl90)),
  inference(subsumption_resolution,[],[f28880,f196])).
fof(f28880,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0)))))) | ($spl8 | $spl90)),
  inference(forward_demodulation,[],[f3951,f251])).
fof(f3951,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))))) | $spl90),
  inference(cnf_transformation,[],[f3951_D])).
fof(f3951_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))))) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f28754,plain,(
  $false | ($spl1 | $spl8 | $spl79)),
  inference(subsumption_resolution,[],[f28753,f203])).
fof(f203,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int',sK1)))) | $spl1),
  inference(resolution,[],[f196,f161])).
fof(f28753,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int',sK1)))) | ($spl8 | $spl79)),
  inference(forward_demodulation,[],[f3441,f251])).
fof(f3441,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))),s('type.integer.int',sK1)))) | $spl79),
  inference(cnf_transformation,[],[f3441_D])).
fof(f3441_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))),s('type.integer.int',sK1)))) <=> ~$spl79),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl79])])).
fof(f28736,plain,(
  $false | ($spl8 | $spl61 | $spl159)),
  inference(subsumption_resolution,[],[f28735,f26513])).
fof(f26513,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1)))))) | $spl61),
  inference(resolution,[],[f1558,f158])).
fof(f28735,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1)))))) | ($spl8 | $spl159)),
  inference(forward_demodulation,[],[f11673,f680])).
fof(f11673,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))))) | $spl159),
  inference(resolution,[],[f11628,f163])).
fof(f199,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0)))) | $spl4),
  inference(cnf_transformation,[],[f199_D])).
fof(f199_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f28701,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int',sK1)))) | $spl328),
  inference(cnf_transformation,[],[f28701_D])).
fof(f28701_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int',sK1)))) <=> ~$spl328),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl328])])).
fof(f28724,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int',sK1)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0)))) | $spl8),
  inference(forward_demodulation,[],[f28723,f680])).
fof(f28723,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0)))) | $spl8),
  inference(forward_demodulation,[],[f28722,f148])).
fof(f28722,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0)))) | $spl8),
  inference(forward_demodulation,[],[f26874,f145])).
fof(f26874,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))) | $spl8),
  inference(forward_demodulation,[],[f26873,f251])).
fof(f26873,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))) | $spl8),
  inference(forward_demodulation,[],[f26815,f148])).
fof(f26815,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))) | $spl8),
  inference(superposition,[],[f3027,f680])).
fof(f28721,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int',sK1)))) | $spl8),
  inference(forward_demodulation,[],[f28720,f251])).
fof(f28720,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int',sK1)))) | $spl8),
  inference(forward_demodulation,[],[f28719,f148])).
fof(f28719,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int',sK1)))) | $spl8),
  inference(forward_demodulation,[],[f26881,f145])).
fof(f26881,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int',sK1)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | $spl8),
  inference(forward_demodulation,[],[f26880,f680])).
fof(f26880,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | $spl8),
  inference(forward_demodulation,[],[f26816,f148])).
fof(f26816,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | $spl8),
  inference(superposition,[],[f3027,f251])).
fof(f28715,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int',sK1)))) | $spl8),
  inference(forward_demodulation,[],[f28714,f251])).
fof(f28714,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int',sK1)))) | $spl8),
  inference(forward_demodulation,[],[f28713,f148])).
fof(f28713,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int',sK1)))) | $spl8),
  inference(forward_demodulation,[],[f26896,f680])).
fof(f26896,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int',sK1)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl8),
  inference(forward_demodulation,[],[f26895,f251])).
fof(f26895,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int',sK1)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl8),
  inference(forward_demodulation,[],[f26840,f149])).
fof(f26840,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl8),
  inference(superposition,[],[f3027,f680])).
fof(f28712,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int',sK1)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0)))) | $spl8),
  inference(forward_demodulation,[],[f28711,f680])).
fof(f28711,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0)))) | $spl8),
  inference(forward_demodulation,[],[f28710,f148])).
fof(f28710,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0)))) | $spl8),
  inference(forward_demodulation,[],[f26899,f251])).
fof(f26899,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl8),
  inference(forward_demodulation,[],[f26898,f680])).
fof(f26898,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl8),
  inference(forward_demodulation,[],[f26841,f149])).
fof(f26841,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl8),
  inference(superposition,[],[f3027,f251])).
fof(f28700,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK1)))) | $spl327),
  inference(cnf_transformation,[],[f28700_D])).
fof(f28700_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK1)))) <=> ~$spl327),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl327])])).
fof(f28703,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0)))) | ($spl8 | $spl51)),
  inference(forward_demodulation,[],[f28317,f680])).
fof(f28317,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0)))) | $spl51),
  inference(resolution,[],[f139,f875])).
fof(f28702,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int',sK1)))) | $spl329),
  inference(cnf_transformation,[],[f28702_D])).
fof(f28702_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int',sK1)))) <=> ~$spl329),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl329])])).
fof(f28494,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int',sK1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK1)))) | ($spl8 | $spl61)),
  inference(forward_demodulation,[],[f28318,f251])).
fof(f28318,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int',sK1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK1)))) | $spl61),
  inference(resolution,[],[f139,f1558])).
fof(f598,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))))),s('type.integer.int',sK1)))) | $spl18),
  inference(cnf_transformation,[],[f598_D])).
fof(f598_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))))),s('type.integer.int',sK1)))) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f28698,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0)))))) | $spl325),
  inference(cnf_transformation,[],[f28698_D])).
fof(f28698_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0)))))) <=> ~$spl325),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl325])])).
fof(f28696,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))))),s('type.integer.int',sK1)))) | ($spl1 | $spl8)),
  inference(forward_demodulation,[],[f28344,f251])).
fof(f28344,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))))),s('type.integer.int',sK1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0)))))) | $spl1),
  inference(resolution,[],[f139,f351])).
fof(f28498,plain,(
  $false | ($spl5 | $spl8 | $spl61)),
  inference(subsumption_resolution,[],[f28497,f189])).
fof(f189,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.integer.INT_ABS_LE|split|1',aHL_TRUTH)).
fof(f28497,plain,(
  ~p(s(bool,'T_0')) | ($spl5 | $spl8 | $spl61)),
  inference(forward_demodulation,[],[f28496,f10265])).
fof(f10265,plain,(
  ( ! [X30] : (s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X30),s('type.integer.int',X30)))) ) | $spl5),
  inference(subsumption_resolution,[],[f10111,f200])).
fof(f10111,plain,(
  ( ! [X30] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X30),s('type.integer.int',X30)))) ) | $spl5),
  inference(superposition,[],[f187,f1892])).
fof(f1892,plain,(
  ( ! [X24] : (s(bool,X24) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) | s(bool,'T_0') = s(bool,X24)) ) | $spl5),
  inference(resolution,[],[f886,f186])).
fof(f886,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,X0) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0)))) ) | $spl5),
  inference(resolution,[],[f200,f165])).
fof(f165,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f110])).
fof(f28496,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK1)))) | ($spl5 | $spl8 | $spl61)),
  inference(subsumption_resolution,[],[f28495,f193])).
fof(f28495,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK1)))) | ($spl5 | $spl8 | $spl61)),
  inference(forward_demodulation,[],[f28494,f997])).
fof(f997,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))) | $spl5),
  inference(resolution,[],[f283,f186])).
fof(f283,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int',sK1)))) | $spl5),
  inference(resolution,[],[f255,f161])).
fof(f255,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK0)))) | $spl5),
  inference(resolution,[],[f200,f163])).
fof(f28227,plain,(
  $false | ($spl5 | $spl61 | $spl112)),
  inference(subsumption_resolution,[],[f28153,f26513])).
fof(f28153,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1)))))) | ($spl5 | $spl112)),
  inference(superposition,[],[f24024,f25846])).
fof(f25846,plain,(
  s('type.integer.int',sK1) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))))) | ($spl5 | $spl112)),
  inference(forward_demodulation,[],[f7465,f13286])).
fof(f13286,plain,(
  ( ! [X2] : (s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X2))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X2)))))) ) | $spl5),
  inference(subsumption_resolution,[],[f13244,f12794])).
fof(f12794,plain,(
  ( ! [X3] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X3))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X3))))))))) ) | $spl5),
  inference(resolution,[],[f11249,f160])).
fof(f11249,plain,(
  ( ! [X4] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X4))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X4))))))) ) | $spl5),
  inference(subsumption_resolution,[],[f11248,f193])).
fof(f11248,plain,(
  ( ! [X4] : (~p(s(bool,'T_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X4))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X4))))))) ) | $spl5),
  inference(forward_demodulation,[],[f11247,f10446])).
fof(f10446,plain,(
  ( ! [X7] : (s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X7)))))) ) | $spl5),
  inference(resolution,[],[f10381,f186])).
fof(f10381,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))))) ) | $spl5),
  inference(subsumption_resolution,[],[f10332,f193])).
fof(f10332,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))))) ) | $spl5),
  inference(superposition,[],[f137,f10265])).
fof(f11247,plain,(
  ( ! [X4] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X4),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X4)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X4))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X4))))))) ) | $spl5),
  inference(forward_demodulation,[],[f11212,f140])).
fof(f11212,plain,(
  ( ! [X4] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X4))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X4)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X4))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X4))))))) ) | $spl5),
  inference(resolution,[],[f10783,f139])).
fof(f10783,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))))) ) | $spl5),
  inference(superposition,[],[f10564,f150])).
fof(f10564,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))))))))) ) | $spl5),
  inference(superposition,[],[f10382,f140])).
fof(f10382,plain,(
  ( ! [X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X1))))),s('type.integer.int',X1))))) ) | $spl5),
  inference(subsumption_resolution,[],[f10333,f193])).
fof(f10333,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X1))))),s('type.integer.int',X1))))) ) | $spl5),
  inference(superposition,[],[f138,f10265])).
fof(f138,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f116])).
fof(f13244,plain,(
  ( ! [X2] : (s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X2))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X2))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X2))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X2))))))))) ) | $spl5),
  inference(resolution,[],[f12803,f162])).
fof(f12803,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))))))))) ) | $spl5),
  inference(superposition,[],[f11249,f150])).
fof(f24024,plain,(
  ( ! [X2] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X2))),s('type.integer.int',X2))))) ) | $spl5),
  inference(subsumption_resolution,[],[f23941,f193])).
fof(f23941,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X2))),s('type.integer.int',X2))))) ) | $spl5),
  inference(superposition,[],[f159,f10446])).
fof(f28226,plain,(
  $false | ($spl5 | $spl65 | $spl112)),
  inference(subsumption_resolution,[],[f28152,f1651])).
fof(f1651,plain,(
  s('type.integer.int',sK1) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))) | $spl65),
  inference(cnf_transformation,[],[f1651_D])).
fof(f1651_D,plain,(
  s('type.integer.int',sK1) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))) <=> ~$spl65),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl65])])).
fof(f28152,plain,(
  s('type.integer.int',sK1) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))) | ($spl5 | $spl112)),
  inference(superposition,[],[f21479,f25846])).
fof(f21479,plain,(
  ( ! [X10] : (s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X10))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X10)))))) ) | $spl5),
  inference(subsumption_resolution,[],[f21478,f10441])).
fof(f10441,plain,(
  ( ! [X2] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X2))),s('type.integer.int',X2))))) ) | $spl5),
  inference(resolution,[],[f10381,f160])).
fof(f21478,plain,(
  ( ! [X10] : (s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X10))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X10))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X10))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X10))))))) ) | $spl5),
  inference(subsumption_resolution,[],[f21423,f193])).
fof(f21423,plain,(
  ( ! [X10] : (~p(s(bool,'T_0')) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X10))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X10))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X10))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X10))))))) ) | $spl5),
  inference(superposition,[],[f162,f10277])).
fof(f10277,plain,(
  ( ! [X68] : (s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X68))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X68)))))) ) | $spl5),
  inference(subsumption_resolution,[],[f10135,f200])).
fof(f10135,plain,(
  ( ! [X68] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X68))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X68)))))) ) | $spl5),
  inference(superposition,[],[f9205,f1892])).
fof(f9205,plain,(
  ( ! [X5] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X5))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X5))))))) )),
  inference(forward_demodulation,[],[f9204,f144])).
fof(f144,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))) = s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int',X0)))) )),
  inference(cnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0] : s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))) = s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int',X0)))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X15] : s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X15))) = s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X15),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X15))),s('type.integer.int',X15)))),
  file('thm.integer.INT_ABS_LE|split|1','thm.integer.INT_ABS')).
fof(f9204,plain,(
  ( ! [X5] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X5),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X5))),s('type.integer.int',X5))))),s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X5),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X5))),s('type.integer.int',X5))))))) )),
  inference(subsumption_resolution,[],[f9203,f193])).
fof(f9203,plain,(
  ( ! [X5] : (~p(s(bool,'T_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X5),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X5))),s('type.integer.int',X5))))),s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X5),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X5))),s('type.integer.int',X5))))))) )),
  inference(forward_demodulation,[],[f9173,f2511])).
fof(f9173,plain,(
  ( ! [X5] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X5)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X5),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X5))),s('type.integer.int',X5))))),s('type.integer.int','const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X5),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X5))),s('type.integer.int',X5))))))) )),
  inference(superposition,[],[f2783,f144])).
fof(f28225,plain,(
  $false | ($spl5 | $spl65 | $spl112)),
  inference(subsumption_resolution,[],[f28141,f1651])).
fof(f28141,plain,(
  s('type.integer.int',sK1) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))) | ($spl5 | $spl112)),
  inference(superposition,[],[f21479,f25846])).
fof(f28220,plain,(
  $false | ($spl5 | $spl65 | $spl112)),
  inference(subsumption_resolution,[],[f28138,f1651])).
fof(f28138,plain,(
  s('type.integer.int',sK1) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))) | ($spl5 | $spl112)),
  inference(superposition,[],[f25846,f21479])).
fof(f28130,plain,(
  $false | ($spl5 | $spl61 | $spl116)),
  inference(subsumption_resolution,[],[f25823,f1558])).
fof(f25823,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int',sK1)))) | ($spl5 | $spl116)),
  inference(forward_demodulation,[],[f21936,f13286])).
fof(f21936,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))),s('type.integer.int',sK1)))) | ($spl5 | $spl116)),
  inference(forward_demodulation,[],[f21935,f21479])).
fof(f21935,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))))),s('type.integer.int',sK1)))) | ($spl5 | $spl116)),
  inference(forward_demodulation,[],[f7469,f21479])).
fof(f19411,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,X5)) ) | $spl244),
  inference(cnf_transformation,[],[f19411_D])).
fof(f19411_D,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,X5)) ) <=> ~$spl244),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl244])])).
fof(f27794,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))))) | $spl322),
  inference(cnf_transformation,[],[f27794_D])).
fof(f27794_D,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))))) <=> ~$spl322),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl322])])).
fof(f27795,plain,(
  ( ! [X6] : (s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))))) | s(bool,'T_0') = s(bool,X6)) ) | ($spl5 | $spl61)),
  inference(subsumption_resolution,[],[f27784,f184])).
fof(f184,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f131])).
fof(f27784,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))))) | s(bool,'T_0') = s(bool,X6)) ) | ($spl5 | $spl61)),
  inference(superposition,[],[f26513,f10081])).
fof(f10081,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | $spl5),
  inference(superposition,[],[f1892,f1892])).
fof(f27792,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1)))))) ) | ($spl5 | $spl61)),
  inference(subsumption_resolution,[],[f27783,f184])).
fof(f27783,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1)))))) ) | ($spl5 | $spl61)),
  inference(superposition,[],[f26513,f10081])).
fof(f27590,plain,(
  $false | ($spl5 | $spl60 | $spl65)),
  inference(subsumption_resolution,[],[f27589,f24024])).
fof(f27589,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int',sK1)))) | ($spl60 | $spl65)),
  inference(subsumption_resolution,[],[f27559,f1651])).
fof(f27559,plain,(
  s('type.integer.int',sK1) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int',sK1)))) | $spl60),
  inference(resolution,[],[f1559,f162])).
fof(f18347,plain,(
  ( ! [X290,X289] : (s('type.integer.int',X289) = s('type.integer.int',X290) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X289),s('type.integer.int',X290))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X289),s('type.integer.int',X290))))) ) | $spl232),
  inference(cnf_transformation,[],[f18347_D])).
fof(f18347_D,plain,(
  ( ! [X290,X289] : (s('type.integer.int',X289) = s('type.integer.int',X290) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X289),s('type.integer.int',X290))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X289),s('type.integer.int',X290))))) ) <=> ~$spl232),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl232])])).
fof(f4750,plain,(
  ( ! [X1] : (p(s(bool,X1))) ) | $spl102),
  inference(cnf_transformation,[],[f4750_D])).
fof(f4750_D,plain,(
  ( ! [X1] : (p(s(bool,X1))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f27485,plain,(
  ( ! [X33,X31,X32] : (p(s(bool,X33)) | s('type.integer.int',X31) = s('type.integer.int',X32) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X31),s('type.integer.int',X32)))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X31),s('type.integer.int',X32)))) ) | $spl5),
  inference(subsumption_resolution,[],[f27058,f183])).
fof(f27058,plain,(
  ( ! [X33,X31,X32] : (p(s(bool,X33)) | s('type.integer.int',X31) = s('type.integer.int',X32) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X31),s('type.integer.int',X32)))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X31),s('type.integer.int',X32))) | s(bool,'T_0') = s(bool,X33)) ) | $spl5),
  inference(superposition,[],[f162,f10081])).
fof(f27484,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X30)) | s('type.integer.int',X28) = s('type.integer.int',X29) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X28),s('type.integer.int',X29)))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X28),s('type.integer.int',X29)))) ) | $spl5),
  inference(subsumption_resolution,[],[f27057,f183])).
fof(f27057,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X30)) | s('type.integer.int',X28) = s('type.integer.int',X29) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X28),s('type.integer.int',X29)))) | s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X28),s('type.integer.int',X29)))) ) | $spl5),
  inference(superposition,[],[f162,f10081])).
fof(f26977,plain,(
  $false | ($spl5 | $spl61 | $spl84)),
  inference(subsumption_resolution,[],[f25858,f1558])).
fof(f25858,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int',sK1)))) | ($spl5 | $spl84)),
  inference(forward_demodulation,[],[f25857,f13286])).
fof(f25857,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))),s('type.integer.int',sK1)))) | ($spl5 | $spl84)),
  inference(forward_demodulation,[],[f3448,f21479])).
fof(f26974,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | $spl320),
  inference(cnf_transformation,[],[f26974_D])).
fof(f26974_D,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) <=> ~$spl320),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl320])])).
fof(f26975,plain,(
  ( ! [X6] : (s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X6)) ) | ($spl5 | $spl29 | $spl41)),
  inference(subsumption_resolution,[],[f26960,f184])).
fof(f26960,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X6)) ) | ($spl5 | $spl29 | $spl41)),
  inference(superposition,[],[f5146,f10081])).
fof(f5146,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | ($spl29 | $spl41)),
  inference(subsumption_resolution,[],[f5124,f753])).
fof(f5124,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | s('type.integer.int',sK1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl41),
  inference(resolution,[],[f800,f3355])).
fof(f26972,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) ) | ($spl5 | $spl29 | $spl41)),
  inference(subsumption_resolution,[],[f26959,f184])).
fof(f26959,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) ) | ($spl5 | $spl29 | $spl41)),
  inference(superposition,[],[f5146,f10081])).
fof(f26937,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))))) | $spl318),
  inference(cnf_transformation,[],[f26937_D])).
fof(f26937_D,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))))) <=> ~$spl318),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl318])])).
fof(f26938,plain,(
  ( ! [X6] : (s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))))) | s(bool,'T_0') = s(bool,X6)) ) | $spl5),
  inference(subsumption_resolution,[],[f26928,f184])).
fof(f26928,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))))) | s(bool,'T_0') = s(bool,X6)) ) | $spl5),
  inference(superposition,[],[f1687,f10081])).
fof(f1687,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1)))))) | $spl5),
  inference(resolution,[],[f1087,f161])).
fof(f1087,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int',sK0)))) | $spl5),
  inference(resolution,[],[f883,f163])).
fof(f883,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int',sK0)))) | $spl5),
  inference(resolution,[],[f200,f137])).
fof(f26935,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1)))))) ) | $spl5),
  inference(subsumption_resolution,[],[f26927,f184])).
fof(f26927,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1)))))) ) | $spl5),
  inference(superposition,[],[f1687,f10081])).
fof(f18308,plain,(
  ( ! [X135] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X135),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X135)))) ) | $spl196),
  inference(cnf_transformation,[],[f18308_D])).
fof(f18308_D,plain,(
  ( ! [X135] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X135),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X135)))) ) <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f26909,plain,(
  ( ! [X19,X18] : (p(s(bool,X19)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X18),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X18)))) ) | $spl5),
  inference(subsumption_resolution,[],[f26858,f183])).
fof(f26858,plain,(
  ( ! [X19,X18] : (p(s(bool,X19)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X18),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X18))) | s(bool,'T_0') = s(bool,X19)) ) | $spl5),
  inference(superposition,[],[f3027,f10081])).
fof(f26908,plain,(
  ( ! [X17,X16] : (p(s(bool,X17)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X16),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X16)))) ) | $spl5),
  inference(subsumption_resolution,[],[f26857,f183])).
fof(f26857,plain,(
  ( ! [X17,X16] : (p(s(bool,X17)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X16),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,'T_0') = s(bool,X17) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X16)))) ) | $spl5),
  inference(superposition,[],[f3027,f10081])).
fof(f18334,plain,(
  ( ! [X221] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X221)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X221),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl220),
  inference(cnf_transformation,[],[f18334_D])).
fof(f18334_D,plain,(
  ( ! [X221] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X221)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X221),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) <=> ~$spl220),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl220])])).
fof(f26887,plain,(
  ( ! [X17,X18] : (p(s(bool,X18)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X17)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X17),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl5),
  inference(subsumption_resolution,[],[f26832,f183])).
fof(f26832,plain,(
  ( ! [X17,X18] : (p(s(bool,X18)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X17)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X17),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X18)) ) | $spl5),
  inference(superposition,[],[f3027,f10081])).
fof(f26886,plain,(
  ( ! [X15,X16] : (p(s(bool,X16)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X15)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X15),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl5),
  inference(subsumption_resolution,[],[f26831,f183])).
fof(f26831,plain,(
  ( ! [X15,X16] : (p(s(bool,X16)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X15)))) | s(bool,'T_0') = s(bool,X16) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X15),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl5),
  inference(superposition,[],[f3027,f10081])).
fof(f18314,plain,(
  ( ! [X152,X153] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X153))),s('type.integer.int',X152)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X152))),s('type.integer.int',X153)))) ) | $spl202),
  inference(cnf_transformation,[],[f18314_D])).
fof(f18314_D,plain,(
  ( ! [X152,X153] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X153))),s('type.integer.int',X152)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X152))),s('type.integer.int',X153)))) ) <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f26764,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,X37)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X36))),s('type.integer.int',X35)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X35))),s('type.integer.int',X36)))) ) | $spl5),
  inference(subsumption_resolution,[],[f26695,f183])).
fof(f26695,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,X37)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X36))),s('type.integer.int',X35)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X35))),s('type.integer.int',X36))) | s(bool,'T_0') = s(bool,X37)) ) | $spl5),
  inference(superposition,[],[f138,f10081])).
fof(f26763,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,X34)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X33))),s('type.integer.int',X32)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X32))),s('type.integer.int',X33)))) ) | $spl5),
  inference(subsumption_resolution,[],[f26694,f183])).
fof(f26694,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,X34)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X33))),s('type.integer.int',X32)))) | s(bool,'T_0') = s(bool,X34) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X32))),s('type.integer.int',X33)))) ) | $spl5),
  inference(superposition,[],[f138,f10081])).
fof(f26641,plain,(
  $false | ($spl5 | $spl61 | $spl68)),
  inference(subsumption_resolution,[],[f25805,f1558])).
fof(f25805,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int',sK1)))) | ($spl5 | $spl68)),
  inference(forward_demodulation,[],[f1656,f21479])).
fof(f26508,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))))) | $spl316),
  inference(cnf_transformation,[],[f26508_D])).
fof(f26508_D,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))))) <=> ~$spl316),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl316])])).
fof(f26509,plain,(
  ( ! [X6] : (s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))))) | s(bool,'T_0') = s(bool,X6)) ) | $spl5),
  inference(subsumption_resolution,[],[f26499,f184])).
fof(f26499,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))))) | s(bool,'T_0') = s(bool,X6)) ) | $spl5),
  inference(superposition,[],[f1089,f10081])).
fof(f1089,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1)))))) | $spl5),
  inference(resolution,[],[f883,f161])).
fof(f26506,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1)))))) ) | $spl5),
  inference(subsumption_resolution,[],[f26498,f184])).
fof(f26498,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1)))))) ) | $spl5),
  inference(superposition,[],[f1089,f10081])).
fof(f26352,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl314),
  inference(cnf_transformation,[],[f26352_D])).
fof(f26352_D,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl314),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl314])])).
fof(f26353,plain,(
  ( ! [X6] : (s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X6)) ) | ($spl5 | $spl38)),
  inference(subsumption_resolution,[],[f26343,f184])).
fof(f26343,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X6)) ) | ($spl5 | $spl38)),
  inference(superposition,[],[f799,f10081])).
fof(f26350,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl5 | $spl38)),
  inference(subsumption_resolution,[],[f26342,f184])).
fof(f26342,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl5 | $spl38)),
  inference(superposition,[],[f799,f10081])).
fof(f26318,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int',sK1))) | $spl312),
  inference(cnf_transformation,[],[f26318_D])).
fof(f26318_D,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int',sK1))) <=> ~$spl312),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl312])])).
fof(f26319,plain,(
  ( ! [X6] : (s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X6)) ) | ($spl1 | $spl5)),
  inference(subsumption_resolution,[],[f26308,f184])).
fof(f26308,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X6)) ) | ($spl1 | $spl5)),
  inference(superposition,[],[f351,f10081])).
fof(f26316,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int',sK1)))) ) | ($spl1 | $spl5)),
  inference(subsumption_resolution,[],[f26307,f184])).
fof(f26307,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int',sK1)))) ) | ($spl1 | $spl5)),
  inference(superposition,[],[f351,f10081])).
fof(f26064,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int',sK1))) | $spl310),
  inference(cnf_transformation,[],[f26064_D])).
fof(f26064_D,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int',sK1))) <=> ~$spl310),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl310])])).
fof(f26065,plain,(
  ( ! [X6] : (s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X6)) ) | ($spl1 | $spl5)),
  inference(subsumption_resolution,[],[f26055,f184])).
fof(f26055,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X6)) ) | ($spl1 | $spl5)),
  inference(superposition,[],[f203,f10081])).
fof(f26062,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int',sK1)))) ) | ($spl1 | $spl5)),
  inference(subsumption_resolution,[],[f26054,f184])).
fof(f26054,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int',sK1)))) ) | ($spl1 | $spl5)),
  inference(superposition,[],[f203,f10081])).
fof(f26003,plain,(
  $false | ($spl1 | $spl5 | $spl158)),
  inference(subsumption_resolution,[],[f26002,f196])).
fof(f26002,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0)))))) | ($spl5 | $spl158)),
  inference(forward_demodulation,[],[f11627,f13286])).
fof(f11627,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))))) | $spl158),
  inference(cnf_transformation,[],[f11627_D])).
fof(f11627_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))))) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f25996,plain,(
  $false | ($spl8 | $spl147)),
  inference(subsumption_resolution,[],[f11285,f680])).
fof(f11285,plain,(
  s('type.integer.int',sK1) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))) | $spl147),
  inference(cnf_transformation,[],[f11285_D])).
fof(f11285_D,plain,(
  s('type.integer.int',sK1) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))) <=> ~$spl147),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl147])])).
fof(f25995,plain,(
  $false | ($spl65 | $spl146 | $spl152)),
  inference(subsumption_resolution,[],[f18512,f1651])).
fof(f18512,plain,(
  s('type.integer.int',sK1) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))) | ($spl146 | $spl152)),
  inference(backward_demodulation,[],[f11286,f11519])).
fof(f18316,plain,(
  ( ! [X158,X159] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X158))),s('type.integer.int',X159)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X158),s('type.integer.int',X159)))) ) | $spl204),
  inference(cnf_transformation,[],[f18316_D])).
fof(f18316_D,plain,(
  ( ! [X158,X159] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X158))),s('type.integer.int',X159)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X158),s('type.integer.int',X159)))) ) <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f25972,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,X34)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X32))),s('type.integer.int',X33)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X32),s('type.integer.int',X33)))) ) | $spl5),
  inference(subsumption_resolution,[],[f25911,f183])).
fof(f25911,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,X34)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X32))),s('type.integer.int',X33)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X32),s('type.integer.int',X33))) | s(bool,'T_0') = s(bool,X34)) ) | $spl5),
  inference(superposition,[],[f137,f10081])).
fof(f25971,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X31)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X29))),s('type.integer.int',X30)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X29),s('type.integer.int',X30)))) ) | $spl5),
  inference(subsumption_resolution,[],[f25910,f183])).
fof(f25910,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X31)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X29))),s('type.integer.int',X30)))) | s(bool,'T_0') = s(bool,X31) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X29),s('type.integer.int',X30)))) ) | $spl5),
  inference(superposition,[],[f137,f10081])).
fof(f875,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int',sK0)))) | $spl51),
  inference(cnf_transformation,[],[f875_D])).
fof(f875_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int',sK0)))) <=> ~$spl51),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl51])])).
fof(f25968,plain,(
  ( ! [X5] : (p(s(bool,X5)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int',sK0))))) ) | $spl5),
  inference(subsumption_resolution,[],[f25899,f183])).
fof(f25899,plain,(
  ( ! [X5] : (p(s(bool,X5)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int',sK0)))) | s(bool,'T_0') = s(bool,X5)) ) | $spl5),
  inference(superposition,[],[f137,f1892])).
fof(f25809,plain,(
  $false | ($spl1 | $spl8 | $spl80)),
  inference(subsumption_resolution,[],[f24761,f202])).
fof(f202,plain,(
  s('type.integer.int',sK1) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))) | $spl1),
  inference(resolution,[],[f196,f164])).
fof(f24761,plain,(
  s('type.integer.int',sK1) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))) | ($spl8 | $spl80)),
  inference(backward_demodulation,[],[f251,f3444])).
fof(f10294,plain,(
  ( ! [X3] : (p(s(bool,X3)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int',sK0))))) ) | $spl5),
  inference(subsumption_resolution,[],[f10217,f183])).
fof(f10217,plain,(
  ( ! [X3] : (p(s(bool,X3)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int',sK0)))) | s(bool,'T_0') = s(bool,X3)) ) | $spl5),
  inference(superposition,[],[f137,f1892])).
fof(f25734,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1)))))) | $spl308),
  inference(cnf_transformation,[],[f25734_D])).
fof(f25734_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1)))))) <=> ~$spl308),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl308])])).
fof(f25704,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl296),
  inference(cnf_transformation,[],[f25704_D])).
fof(f25704_D,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl296),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl296])])).
fof(f25752,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1)))))) | ($spl5 | $spl80)),
  inference(forward_demodulation,[],[f25751,f13286])).
fof(f25751,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1)))))) | ($spl5 | $spl80)),
  inference(forward_demodulation,[],[f20215,f21479])).
fof(f20215,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1)))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))))) | $spl80),
  inference(superposition,[],[f2507,f3444])).
fof(f25360,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl284),
  inference(cnf_transformation,[],[f25360_D])).
fof(f25360_D,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl284),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl284])])).
fof(f25750,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1)))))) | ($spl5 | $spl152)),
  inference(forward_demodulation,[],[f25749,f13286])).
fof(f25749,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1)))))) | ($spl5 | $spl152)),
  inference(forward_demodulation,[],[f20217,f21479])).
fof(f20217,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1)))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))))) | $spl152),
  inference(superposition,[],[f2507,f11519])).
fof(f25716,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl303),
  inference(cnf_transformation,[],[f25716_D])).
fof(f25716_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl303),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl303])])).
fof(f25700,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1)))))))) | $spl292),
  inference(cnf_transformation,[],[f25700_D])).
fof(f25700_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1)))))))) <=> ~$spl292),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl292])])).
fof(f25748,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1)))))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl5 | $spl80)),
  inference(forward_demodulation,[],[f20865,f13286])).
fof(f20865,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))))))) | $spl80),
  inference(superposition,[],[f2585,f3444])).
fof(f2585,plain,(
  ( ! [X5] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X5))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X5),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X5))))))) )),
  inference(superposition,[],[f137,f148])).
fof(f25394,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0)))))))) | $spl288),
  inference(cnf_transformation,[],[f25394_D])).
fof(f25394_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0)))))))) <=> ~$spl288),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl288])])).
fof(f25747,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0)))))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl5 | $spl152)),
  inference(forward_demodulation,[],[f20867,f13286])).
fof(f20867,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))))))) | $spl152),
  inference(superposition,[],[f2585,f11519])).
fof(f25715,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl302),
  inference(cnf_transformation,[],[f25715_D])).
fof(f25715_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl302),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl302])])).
fof(f25745,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int',sK0)))) | ($spl5 | $spl8 | $spl80)),
  inference(forward_demodulation,[],[f25744,f13286])).
fof(f25744,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int',sK0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl8 | $spl80)),
  inference(forward_demodulation,[],[f21972,f251])).
fof(f21972,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl80),
  inference(superposition,[],[f2619,f3444])).
fof(f2619,plain,(
  ( ! [X6] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X6))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X6)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X6),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f137,f148])).
fof(f25743,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int',sK0)))) | ($spl5 | $spl8 | $spl152)),
  inference(forward_demodulation,[],[f25742,f13286])).
fof(f25742,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int',sK0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl8 | $spl152)),
  inference(forward_demodulation,[],[f21974,f251])).
fof(f21974,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl152),
  inference(superposition,[],[f2619,f11519])).
fof(f25728,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int',sK0)))))) | $spl306),
  inference(cnf_transformation,[],[f25728_D])).
fof(f25728_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int',sK0)))))) <=> ~$spl306),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl306])])).
fof(f25735,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int',sK0)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl8 | $spl80 | $spl122)),
  inference(forward_demodulation,[],[f22321,f251])).
fof(f22321,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))))) | ($spl80 | $spl122)),
  inference(superposition,[],[f2214,f22208])).
fof(f22208,plain,(
  s('type.integer.int',sK1) = s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | ($spl80 | $spl122)),
  inference(forward_demodulation,[],[f22207,f3444])).
fof(f22207,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) = s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | $spl122),
  inference(forward_demodulation,[],[f22151,f150])).
fof(f22151,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) = s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | $spl122),
  inference(superposition,[],[f144,f22128])).
fof(f22128,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl122),
  inference(forward_demodulation,[],[f8960,f9243])).
fof(f8960,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl122),
  inference(resolution,[],[f8887,f186])).
fof(f2214,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))))) )),
  inference(superposition,[],[f138,f151])).
fof(f25335,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int',sK0))))) | $spl282),
  inference(cnf_transformation,[],[f25335_D])).
fof(f25335_D,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int',sK0))))) <=> ~$spl282),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl282])])).
fof(f25732,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int',sK0))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1)))))) | ($spl8 | $spl80 | $spl122)),
  inference(forward_demodulation,[],[f22323,f251])).
fof(f22323,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1)))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))))) | ($spl80 | $spl122)),
  inference(superposition,[],[f2507,f22208])).
fof(f25330,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int',sK0)))))))) | $spl280),
  inference(cnf_transformation,[],[f25330_D])).
fof(f25330_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int',sK0)))))))) <=> ~$spl280),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl280])])).
fof(f25731,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int',sK0)))))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl8 | $spl80 | $spl122)),
  inference(forward_demodulation,[],[f22329,f251])).
fof(f22329,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))))))) | ($spl80 | $spl122)),
  inference(superposition,[],[f2585,f22208])).
fof(f25707,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl298),
  inference(cnf_transformation,[],[f25707_D])).
fof(f25707_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl298),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl298])])).
fof(f25730,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int',sK0)))) | ($spl8 | $spl80 | $spl122)),
  inference(forward_demodulation,[],[f25729,f251])).
fof(f25729,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int',sK0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl8 | $spl80 | $spl122)),
  inference(forward_demodulation,[],[f22337,f251])).
fof(f22337,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl80 | $spl122)),
  inference(superposition,[],[f2619,f22208])).
fof(f1558,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int',sK1)))) | $spl61),
  inference(cnf_transformation,[],[f1558_D])).
fof(f1558_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int',sK1)))) <=> ~$spl61),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl61])])).
fof(f25726,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int',sK0)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int',sK1)))) | ($spl8 | $spl80 | $spl122)),
  inference(forward_demodulation,[],[f22348,f251])).
fof(f22348,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int',sK1)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))))) | ($spl80 | $spl122)),
  inference(superposition,[],[f2752,f22208])).
fof(f2752,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))))) )),
  inference(superposition,[],[f138,f149])).
fof(f25719,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int',sK0))))),s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int',sK0)))))) | $spl304),
  inference(cnf_transformation,[],[f25719_D])).
fof(f25719_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int',sK0))))),s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int',sK0)))))) <=> ~$spl304),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl304])])).
fof(f25717,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int',sK0))))),s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int',sK0)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl8 | $spl80 | $spl122)),
  inference(forward_demodulation,[],[f22738,f251])).
fof(f22738,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))))),s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))))) | ($spl80 | $spl122)),
  inference(superposition,[],[f2648,f22208])).
fof(f2648,plain,(
  ( ! [X2] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X2))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X2))),s('type.integer.int',X2))))) )),
  inference(forward_demodulation,[],[f2615,f150])).
fof(f2615,plain,(
  ( ! [X2] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X2))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X2))))))),s('type.integer.int',X2))))) )),
  inference(superposition,[],[f138,f148])).
fof(f25714,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int',sK0)))) | ($spl5 | $spl8 | $spl80)),
  inference(forward_demodulation,[],[f25713,f13286])).
fof(f25713,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int',sK0)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl8 | $spl80)),
  inference(forward_demodulation,[],[f22854,f251])).
fof(f22854,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl80),
  inference(superposition,[],[f2655,f3444])).
fof(f2655,plain,(
  ( ! [X7] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X7))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X7)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(subsumption_resolution,[],[f2654,f187])).
fof(f2654,plain,(
  ( ! [X7] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X7),s('type.integer.int',X7)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X7))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X7))))))) )),
  inference(forward_demodulation,[],[f2620,f150])).
fof(f2620,plain,(
  ( ! [X7] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X7))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X7)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X7))))),s('type.integer.int',X7))))) )),
  inference(superposition,[],[f139,f148])).
fof(f25712,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl301),
  inference(cnf_transformation,[],[f25712_D])).
fof(f25712_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl301),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl301])])).
fof(f25710,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int',sK0)))) | ($spl5 | $spl8 | $spl152)),
  inference(forward_demodulation,[],[f25709,f13286])).
fof(f25709,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int',sK0)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl8 | $spl152)),
  inference(forward_demodulation,[],[f22856,f251])).
fof(f22856,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl152),
  inference(superposition,[],[f2655,f11519])).
fof(f25708,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl299),
  inference(cnf_transformation,[],[f25708_D])).
fof(f25708_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl299),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl299])])).
fof(f25706,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int',sK0)))) | ($spl8 | $spl80 | $spl122)),
  inference(forward_demodulation,[],[f25705,f251])).
fof(f25705,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int',sK0)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl8 | $spl80 | $spl122)),
  inference(forward_demodulation,[],[f22863,f251])).
fof(f22863,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl80 | $spl122)),
  inference(superposition,[],[f2655,f22208])).
fof(f25361,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | ($spl5 | $spl80)),
  inference(forward_demodulation,[],[f20224,f13286])).
fof(f20224,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | $spl80),
  inference(superposition,[],[f2507,f3444])).
fof(f25458,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | ($spl5 | $spl41 | $spl80)),
  inference(forward_demodulation,[],[f11083,f13286])).
fof(f11083,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | ($spl41 | $spl80)),
  inference(forward_demodulation,[],[f11082,f3444])).
fof(f11082,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | ($spl41 | $spl80)),
  inference(subsumption_resolution,[],[f10934,f800])).
fof(f10934,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | $spl80),
  inference(superposition,[],[f3355,f3444])).
fof(f25461,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | ($spl5 | $spl80)),
  inference(forward_demodulation,[],[f11080,f13286])).
fof(f11080,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | $spl80),
  inference(subsumption_resolution,[],[f10933,f2508])).
fof(f10933,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | $spl80),
  inference(superposition,[],[f3355,f3444])).
fof(f755,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | $spl31),
  inference(cnf_transformation,[],[f755_D])).
fof(f755_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) <=> ~$spl31),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl31])])).
fof(f25702,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1)))))) | $spl294),
  inference(cnf_transformation,[],[f25702_D])).
fof(f25702_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1)))))) <=> ~$spl294),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl294])])).
fof(f25474,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | ($spl5 | $spl80)),
  inference(forward_demodulation,[],[f10918,f13286])).
fof(f10918,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))))) | $spl80),
  inference(superposition,[],[f2731,f3444])).
fof(f2731,plain,(
  ( ! [X7] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X7)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X7))),s('type.integer.int',X7))))) )),
  inference(superposition,[],[f163,f149])).
fof(f25699,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1)))))))) | $spl293),
  inference(cnf_transformation,[],[f25699_D])).
fof(f25699_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1)))))))) <=> ~$spl293),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl293])])).
fof(f25481,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | ($spl5 | $spl80)),
  inference(forward_demodulation,[],[f10908,f13286])).
fof(f10908,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))))))) | $spl80),
  inference(superposition,[],[f2588,f3444])).
fof(f2588,plain,(
  ( ! [X8] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X8)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X8),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X8))))))) )),
  inference(superposition,[],[f159,f148])).
fof(f25484,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1)))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | ($spl5 | $spl80)),
  inference(forward_demodulation,[],[f10907,f13286])).
fof(f10907,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))))))) | $spl80),
  inference(superposition,[],[f2587,f3444])).
fof(f2587,plain,(
  ( ! [X7] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X7)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X7))))) )),
  inference(superposition,[],[f158,f148])).
fof(f2038,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))))),s('type.integer.int',sK0)))) | $spl77),
  inference(cnf_transformation,[],[f2038_D])).
fof(f2038_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))))),s('type.integer.int',sK0)))) <=> ~$spl77),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl77])])).
fof(f25698,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1)))))) | $spl290),
  inference(cnf_transformation,[],[f25698_D])).
fof(f25698_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1)))))) <=> ~$spl290),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl290])])).
fof(f2015,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))))),s('type.integer.int',sK0)))) | $spl5),
  inference(resolution,[],[f1687,f139])).
fof(f25674,plain,(
  $false | ($spl1 | $spl5 | $spl8 | $spl64)),
  inference(subsumption_resolution,[],[f25673,f202])).
fof(f25673,plain,(
  s('type.integer.int',sK1) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))) | ($spl5 | $spl8 | $spl64)),
  inference(forward_demodulation,[],[f25634,f1652])).
fof(f25634,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))) | ($spl5 | $spl8)),
  inference(superposition,[],[f13286,f680])).
fof(f801,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl40),
  inference(cnf_transformation,[],[f801_D])).
fof(f801_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f25541,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | ($spl29 | $spl64)),
  inference(subsumption_resolution,[],[f25540,f753])).
fof(f25540,plain,(
  s('type.integer.int',sK1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | $spl64),
  inference(forward_demodulation,[],[f25539,f1652])).
fof(f25539,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl64),
  inference(forward_demodulation,[],[f6940,f1652])).
fof(f6940,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl64),
  inference(superposition,[],[f3355,f1652])).
fof(f25446,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0)))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | ($spl5 | $spl152)),
  inference(forward_demodulation,[],[f13408,f13286])).
fof(f13408,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))))))) | $spl152),
  inference(superposition,[],[f2587,f11519])).
fof(f25395,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0)))))))) | $spl289),
  inference(cnf_transformation,[],[f25395_D])).
fof(f25395_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0)))))))) <=> ~$spl289),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl289])])).
fof(f25445,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | ($spl5 | $spl152)),
  inference(forward_demodulation,[],[f13409,f13286])).
fof(f13409,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))))))) | $spl152),
  inference(superposition,[],[f2588,f11519])).
fof(f25389,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0)))))) | $spl286),
  inference(cnf_transformation,[],[f25389_D])).
fof(f25389_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0)))))) <=> ~$spl286),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl286])])).
fof(f25441,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | ($spl5 | $spl152)),
  inference(forward_demodulation,[],[f13419,f13286])).
fof(f13419,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))))) | $spl152),
  inference(superposition,[],[f2731,f11519])).
fof(f25430,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | ($spl5 | $spl152)),
  inference(forward_demodulation,[],[f13639,f13286])).
fof(f13639,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | $spl152),
  inference(subsumption_resolution,[],[f13434,f2508])).
fof(f13434,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | $spl152),
  inference(superposition,[],[f3355,f11519])).
fof(f25429,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | ($spl5 | $spl41 | $spl152)),
  inference(forward_demodulation,[],[f13642,f13286])).
fof(f13642,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | ($spl41 | $spl152)),
  inference(forward_demodulation,[],[f13641,f11519])).
fof(f13641,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | ($spl41 | $spl152)),
  inference(subsumption_resolution,[],[f13435,f800])).
fof(f13435,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | $spl152),
  inference(superposition,[],[f3355,f11519])).
fof(f11318,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))) | $spl149),
  inference(cnf_transformation,[],[f11318_D])).
fof(f11318_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))) <=> ~$spl149),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl149])])).
fof(f18504,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))) | $spl152),
  inference(forward_demodulation,[],[f18503,f148])).
fof(f18503,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))) | $spl152),
  inference(forward_demodulation,[],[f13318,f145])).
fof(f13318,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))) | $spl152),
  inference(superposition,[],[f2752,f11519])).
fof(f25396,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0)))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | ($spl5 | $spl152)),
  inference(forward_demodulation,[],[f19069,f13286])).
fof(f19069,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))))))) | $spl152),
  inference(superposition,[],[f2587,f11519])).
fof(f25393,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | ($spl5 | $spl152)),
  inference(forward_demodulation,[],[f19070,f13286])).
fof(f19070,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))))))) | $spl152),
  inference(superposition,[],[f2588,f11519])).
fof(f25387,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | ($spl5 | $spl152)),
  inference(forward_demodulation,[],[f19081,f13286])).
fof(f19081,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))))) | $spl152),
  inference(superposition,[],[f2731,f11519])).
fof(f25376,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | ($spl5 | $spl152)),
  inference(forward_demodulation,[],[f19346,f13286])).
fof(f19346,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | $spl152),
  inference(subsumption_resolution,[],[f19096,f2508])).
fof(f19096,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | $spl152),
  inference(superposition,[],[f3355,f11519])).
fof(f1181,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0)))))) | $spl54),
  inference(cnf_transformation,[],[f1181_D])).
fof(f1181_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0)))))) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f25375,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0)))))) | ($spl5 | $spl41 | $spl152)),
  inference(forward_demodulation,[],[f19349,f13286])).
fof(f19349,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0)))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | ($spl5 | $spl41 | $spl152)),
  inference(forward_demodulation,[],[f19348,f13286])).
fof(f19348,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | ($spl41 | $spl152)),
  inference(subsumption_resolution,[],[f19097,f800])).
fof(f19097,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | $spl152),
  inference(superposition,[],[f3355,f11519])).
fof(f19835,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int',sK1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK1)))) | ($spl5 | $spl152)),
  inference(forward_demodulation,[],[f19834,f13286])).
fof(f19834,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))),s('type.integer.int',sK1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK1)))) | ($spl5 | $spl152)),
  inference(subsumption_resolution,[],[f19778,f193])).
fof(f19778,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))),s('type.integer.int',sK1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK1)))) | ($spl5 | $spl152)),
  inference(superposition,[],[f139,f13327])).
fof(f13327,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | ($spl5 | $spl152)),
  inference(superposition,[],[f10543,f11519])).
fof(f10543,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X5))))),s('type.integer.int',X5)))) ) | $spl5),
  inference(resolution,[],[f10382,f186])).
fof(f25368,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | ($spl5 | $spl64 | $spl152)),
  inference(forward_demodulation,[],[f25367,f13286])).
fof(f25367,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | ($spl5 | $spl64 | $spl152)),
  inference(forward_demodulation,[],[f20257,f21479])).
fof(f20257,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))))) | ($spl64 | $spl152)),
  inference(forward_demodulation,[],[f20217,f1652])).
fof(f25358,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | ($spl5 | $spl152)),
  inference(forward_demodulation,[],[f20226,f13286])).
fof(f20226,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | $spl152),
  inference(superposition,[],[f2507,f11519])).
fof(f22196,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | p(s(bool,X0))) ) | $spl122),
  inference(forward_demodulation,[],[f22195,f147])).
fof(f22195,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl122),
  inference(subsumption_resolution,[],[f22136,f183])).
fof(f22136,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X0)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl122),
  inference(superposition,[],[f22128,f165])).
fof(f22199,plain,(
  ( ! [X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | p(s(bool,X1))) ) | $spl122),
  inference(forward_demodulation,[],[f22198,f147])).
fof(f22198,plain,(
  ( ! [X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X1))) ) | $spl122),
  inference(subsumption_resolution,[],[f22137,f183])).
fof(f22137,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,X1) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X1))) ) | $spl122),
  inference(superposition,[],[f22128,f165])).
fof(f22226,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | p(s(bool,X0))) ) | $spl122),
  inference(forward_demodulation,[],[f22225,f147])).
fof(f22225,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl122),
  inference(subsumption_resolution,[],[f22160,f183])).
fof(f22160,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X0)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl122),
  inference(superposition,[],[f165,f22128])).
fof(f22229,plain,(
  ( ! [X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | p(s(bool,X1))) ) | $spl122),
  inference(forward_demodulation,[],[f22228,f147])).
fof(f22228,plain,(
  ( ! [X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X1))) ) | $spl122),
  inference(subsumption_resolution,[],[f22161,f183])).
fof(f22161,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,X1) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X1))) ) | $spl122),
  inference(superposition,[],[f165,f22128])).
fof(f25333,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int',sK0))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | ($spl8 | $spl64 | $spl80 | $spl122)),
  inference(forward_demodulation,[],[f22466,f251])).
fof(f22466,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))))) | ($spl64 | $spl80 | $spl122)),
  inference(forward_demodulation,[],[f22323,f1652])).
fof(f25332,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int',sK0)))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | ($spl8 | $spl80 | $spl122)),
  inference(forward_demodulation,[],[f22331,f251])).
fof(f22331,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))))))) | ($spl80 | $spl122)),
  inference(superposition,[],[f2587,f22208])).
fof(f25331,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int',sK0)))))))) | $spl281),
  inference(cnf_transformation,[],[f25331_D])).
fof(f25331_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int',sK0)))))))) <=> ~$spl281),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl281])])).
fof(f25329,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int',sK0)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | ($spl8 | $spl80 | $spl122)),
  inference(forward_demodulation,[],[f22332,f251])).
fof(f22332,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))))))) | ($spl80 | $spl122)),
  inference(superposition,[],[f2588,f22208])).
fof(f25323,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int',sK0))))),s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int',sK0)))))) | $spl278),
  inference(cnf_transformation,[],[f25323_D])).
fof(f25323_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int',sK0))))),s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int',sK0)))))) <=> ~$spl278),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl278])])).
fof(f25321,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int',sK0))))),s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int',sK0)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | ($spl8 | $spl80 | $spl122)),
  inference(forward_demodulation,[],[f22346,f251])).
fof(f22346,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))))),s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))))) | ($spl80 | $spl122)),
  inference(superposition,[],[f2731,f22208])).
fof(f25304,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl276),
  inference(cnf_transformation,[],[f25304_D])).
fof(f25304_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl276),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl276])])).
fof(f25300,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int',sK0))) | $spl274),
  inference(cnf_transformation,[],[f25300_D])).
fof(f25300_D,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int',sK0))) <=> ~$spl274),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl274])])).
fof(f25302,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int',sK0))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | ($spl8 | $spl80 | $spl122)),
  inference(forward_demodulation,[],[f25301,f251])).
fof(f25301,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | ($spl8 | $spl80 | $spl122)),
  inference(forward_demodulation,[],[f22361,f251])).
fof(f22361,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | ($spl80 | $spl122)),
  inference(superposition,[],[f3355,f22208])).
fof(f25298,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int',sK0))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | ($spl5 | $spl8 | $spl41 | $spl80 | $spl122)),
  inference(forward_demodulation,[],[f22552,f251])).
fof(f22552,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | ($spl5 | $spl41 | $spl80 | $spl122)),
  inference(forward_demodulation,[],[f22551,f22208])).
fof(f22551,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | ($spl5 | $spl41 | $spl80 | $spl122)),
  inference(subsumption_resolution,[],[f22550,f200])).
fof(f22550,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | ($spl5 | $spl41 | $spl80 | $spl122)),
  inference(forward_demodulation,[],[f22362,f5132])).
fof(f5132,plain,(
  s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl5 | $spl41)),
  inference(resolution,[],[f800,f886])).
fof(f22362,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.bool.COND_3'(s(bool,'T_0'),s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | ($spl80 | $spl122)),
  inference(superposition,[],[f3355,f22208])).
fof(f25292,plain,(
  $false | ($spl1 | $spl5 | $spl22)),
  inference(subsumption_resolution,[],[f25291,f189])).
fof(f25291,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl5 | $spl22)),
  inference(forward_demodulation,[],[f25290,f10265])).
fof(f25290,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK1)))) | ($spl1 | $spl22)),
  inference(forward_demodulation,[],[f196,f655])).
fof(f25288,plain,(
  $false | ($spl1 | $spl22)),
  inference(subsumption_resolution,[],[f202,f655])).
fof(f25287,plain,(
  $false | ($spl1 | $spl22 | $spl261)),
  inference(subsumption_resolution,[],[f25286,f23515])).
fof(f25286,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK1)))) | ($spl1 | $spl22)),
  inference(forward_demodulation,[],[f203,f655])).
fof(f25282,plain,(
  $false | ($spl1 | $spl22 | $spl64)),
  inference(subsumption_resolution,[],[f25281,f655])).
fof(f25281,plain,(
  s('type.integer.int',sK1) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))) | ($spl1 | $spl64)),
  inference(forward_demodulation,[],[f545,f1652])).
fof(f24989,plain,(
  $false | ($spl5 | $spl8 | $spl138)),
  inference(subsumption_resolution,[],[f24988,f884])).
fof(f884,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) | $spl5),
  inference(resolution,[],[f200,f185])).
fof(f185,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f132])).
fof(f24988,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) | ($spl5 | $spl8 | $spl138)),
  inference(forward_demodulation,[],[f24987,f10265])).
fof(f24987,plain,(
  s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int',sK0))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) | ($spl8 | $spl138)),
  inference(forward_demodulation,[],[f10306,f251])).
fof(f24820,plain,(
  $false | ($spl3 | $spl5 | $spl8)),
  inference(subsumption_resolution,[],[f24819,f189])).
fof(f24819,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl5 | $spl8)),
  inference(forward_demodulation,[],[f24771,f10265])).
fof(f24771,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int',sK0)))) | ($spl3 | $spl8)),
  inference(backward_demodulation,[],[f251,f23995])).
fof(f23995,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | $spl3),
  inference(subsumption_resolution,[],[f23898,f193])).
fof(f23898,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | $spl3),
  inference(superposition,[],[f159,f1127])).
fof(f1127,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int',sK0))) | $spl3),
  inference(resolution,[],[f972,f186])).
fof(f972,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int',sK0)))) | $spl3),
  inference(resolution,[],[f198,f161])).
fof(f24818,plain,(
  $false | ($spl3 | $spl5 | $spl8)),
  inference(subsumption_resolution,[],[f24755,f10406])).
fof(f10406,plain,(
  ( ! [X7] : (s(bool,'T_0') != s(bool,'const.integer.int_lt_2'(s('type.integer.int',X7),s('type.integer.int',X7)))) ) | $spl5),
  inference(resolution,[],[f10385,f183])).
fof(f10385,plain,(
  ( ! [X8] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X8),s('type.integer.int',X8))))) ) | $spl5),
  inference(subsumption_resolution,[],[f10340,f193])).
fof(f10340,plain,(
  ( ! [X8] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X8),s('type.integer.int',X8))))) ) | $spl5),
  inference(superposition,[],[f159,f10265])).
fof(f24755,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK0),s('type.integer.int',sK0))) | ($spl3 | $spl8)),
  inference(backward_demodulation,[],[f251,f1127])).
fof(f24817,plain,(
  $false | ($spl3 | $spl5 | $spl8)),
  inference(subsumption_resolution,[],[f24753,f10265])).
fof(f24753,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int',sK0))) | ($spl3 | $spl8)),
  inference(backward_demodulation,[],[f251,f975])).
fof(f975,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | $spl3),
  inference(resolution,[],[f198,f185])).
fof(f24740,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))) | $spl272),
  inference(cnf_transformation,[],[f24740_D])).
fof(f24740_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))) <=> ~$spl272),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl272])])).
fof(f18981,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))) | $spl152),
  inference(superposition,[],[f2752,f11519])).
fof(f11284,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK1)))) | $spl144),
  inference(cnf_transformation,[],[f11284_D])).
fof(f11284_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK1)))) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f11286,plain,(
  s('type.integer.int',sK1) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))) | $spl146),
  inference(cnf_transformation,[],[f11286_D])).
fof(f11286_D,plain,(
  s('type.integer.int',sK1) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f24738,plain,(
  s('type.integer.int',sK1) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK1)))) | ($spl5 | $spl22)),
  inference(subsumption_resolution,[],[f19485,f193])).
fof(f19485,plain,(
  ~p(s(bool,'T_0')) | s('type.integer.int',sK1) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK1)))) | ($spl5 | $spl22)),
  inference(superposition,[],[f162,f11262])).
fof(f11262,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK1))) | ($spl5 | $spl22)),
  inference(resolution,[],[f11234,f186])).
fof(f11234,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK1)))) | ($spl5 | $spl22)),
  inference(superposition,[],[f10783,f655])).
fof(f24735,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl270),
  inference(cnf_transformation,[],[f24735_D])).
fof(f24735_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl270),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl270])])).
fof(f21965,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl22),
  inference(superposition,[],[f2619,f655])).
fof(f24736,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl271),
  inference(cnf_transformation,[],[f24736_D])).
fof(f24736_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl271),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl271])])).
fof(f11317,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))) | $spl148),
  inference(cnf_transformation,[],[f11317_D])).
fof(f11317_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f22847,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl22),
  inference(superposition,[],[f2655,f655])).
fof(f18880,plain,(
  s('type.integer.int',sK1) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))) | ($spl5 | $spl22)),
  inference(resolution,[],[f11259,f162])).
fof(f11259,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))) | ($spl5 | $spl22)),
  inference(resolution,[],[f11234,f160])).
fof(f19367,plain,(
  s('type.integer.int',sK1) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK1)))) | ($spl5 | $spl22)),
  inference(resolution,[],[f18881,f162])).
fof(f18881,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK1)))) | ($spl5 | $spl22)),
  inference(resolution,[],[f11259,f161])).
fof(f24734,plain,(
  s('type.integer.int',sK1) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK1)))) | ($spl5 | $spl22)),
  inference(subsumption_resolution,[],[f11440,f193])).
fof(f11440,plain,(
  ~p(s(bool,'T_0')) | s('type.integer.int',sK1) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK1)))) | ($spl5 | $spl22)),
  inference(superposition,[],[f162,f11262])).
fof(f24731,plain,(
  $false | ($spl5 | $spl22 | $spl149 | $spl152)),
  inference(subsumption_resolution,[],[f22896,f22990])).
fof(f22990,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK0)))) | ($spl22 | $spl149)),
  inference(forward_demodulation,[],[f22954,f149])).
fof(f22954,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK0)))) | ($spl22 | $spl149)),
  inference(resolution,[],[f22901,f3102])).
fof(f3102,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))))) )),
  inference(superposition,[],[f3027,f148])).
fof(f22901,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))) | ($spl22 | $spl149)),
  inference(forward_demodulation,[],[f22900,f148])).
fof(f22900,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl22 | $spl149)),
  inference(subsumption_resolution,[],[f22847,f11318])).
fof(f22896,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK0)))) | ($spl5 | $spl152)),
  inference(forward_demodulation,[],[f22895,f149])).
fof(f22895,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK0)))) | ($spl5 | $spl152)),
  inference(forward_demodulation,[],[f22894,f2216])).
fof(f2216,plain,(
  ( ! [X2] : (s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X2))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X2))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f140,f151])).
fof(f22894,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl5 | $spl152)),
  inference(subsumption_resolution,[],[f22893,f200])).
fof(f22893,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl152),
  inference(forward_demodulation,[],[f22845,f150])).
fof(f22845,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl152),
  inference(superposition,[],[f2655,f11519])).
fof(f24477,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK0))) | $spl268),
  inference(cnf_transformation,[],[f24477_D])).
fof(f24477_D,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK0))) <=> ~$spl268),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl268])])).
fof(f24478,plain,(
  ( ! [X6] : (s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK0))) | s(bool,'T_0') = s(bool,X6)) ) | ($spl5 | $spl22 | $spl149)),
  inference(subsumption_resolution,[],[f24467,f184])).
fof(f24467,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK0))) | s(bool,'T_0') = s(bool,X6)) ) | ($spl5 | $spl22 | $spl149)),
  inference(superposition,[],[f24282,f10081])).
fof(f24282,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK0)))) | ($spl22 | $spl149)),
  inference(resolution,[],[f163,f22955])).
fof(f22955,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK0)))) | ($spl22 | $spl149)),
  inference(resolution,[],[f22901,f2587])).
fof(f24475,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK0)))) ) | ($spl5 | $spl22 | $spl149)),
  inference(subsumption_resolution,[],[f24466,f184])).
fof(f24466,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK0)))) ) | ($spl5 | $spl22 | $spl149)),
  inference(superposition,[],[f24282,f10081])).
fof(f18321,plain,(
  ( ! [X180,X179] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X179),s('type.integer.int',X180)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X179),s('type.integer.int',X180)))) ) | $spl208),
  inference(cnf_transformation,[],[f18321_D])).
fof(f18321_D,plain,(
  ( ! [X180,X179] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X179),s('type.integer.int',X180)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X179),s('type.integer.int',X180)))) ) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f24411,plain,(
  ( ! [X35,X36,X34] : (p(s(bool,X36)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X34),s('type.integer.int',X35)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X34),s('type.integer.int',X35)))) ) | $spl5),
  inference(subsumption_resolution,[],[f24335,f183])).
fof(f24335,plain,(
  ( ! [X35,X36,X34] : (p(s(bool,X36)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X34),s('type.integer.int',X35)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X34),s('type.integer.int',X35))) | s(bool,'T_0') = s(bool,X36)) ) | $spl5),
  inference(superposition,[],[f163,f10081])).
fof(f24410,plain,(
  ( ! [X33,X31,X32] : (p(s(bool,X33)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X31),s('type.integer.int',X32)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X31),s('type.integer.int',X32)))) ) | $spl5),
  inference(subsumption_resolution,[],[f24334,f183])).
fof(f24334,plain,(
  ( ! [X33,X31,X32] : (p(s(bool,X33)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X31),s('type.integer.int',X32)))) | s(bool,'T_0') = s(bool,X33) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X31),s('type.integer.int',X32)))) ) | $spl5),
  inference(superposition,[],[f163,f10081])).
fof(f24063,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK0))) | $spl266),
  inference(cnf_transformation,[],[f24063_D])).
fof(f24063_D,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK0))) <=> ~$spl266),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl266])])).
fof(f24064,plain,(
  ( ! [X6] : (s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK0))) | s(bool,'T_0') = s(bool,X6)) ) | ($spl5 | $spl22 | $spl149)),
  inference(subsumption_resolution,[],[f24053,f184])).
fof(f24053,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK0))) | s(bool,'T_0') = s(bool,X6)) ) | ($spl5 | $spl22 | $spl149)),
  inference(superposition,[],[f22990,f10081])).
fof(f24061,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK0)))) ) | ($spl5 | $spl22 | $spl149)),
  inference(subsumption_resolution,[],[f24052,f184])).
fof(f24052,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK0)))) ) | ($spl5 | $spl22 | $spl149)),
  inference(superposition,[],[f22990,f10081])).
fof(f23858,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK0))) | $spl264),
  inference(cnf_transformation,[],[f23858_D])).
fof(f23858_D,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK0))) <=> ~$spl264),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl264])])).
fof(f23859,plain,(
  ( ! [X6] : (s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK0))) | s(bool,'T_0') = s(bool,X6)) ) | ($spl5 | $spl22 | $spl149)),
  inference(subsumption_resolution,[],[f23848,f184])).
fof(f23848,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK0))) | s(bool,'T_0') = s(bool,X6)) ) | ($spl5 | $spl22 | $spl149)),
  inference(superposition,[],[f22959,f10081])).
fof(f22959,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK0)))) | ($spl22 | $spl149)),
  inference(resolution,[],[f22901,f161])).
fof(f23856,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK0)))) ) | ($spl5 | $spl22 | $spl149)),
  inference(subsumption_resolution,[],[f23847,f184])).
fof(f23847,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK0)))) ) | ($spl5 | $spl22 | $spl149)),
  inference(superposition,[],[f22959,f10081])).
fof(f23048,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK0))) | $spl256),
  inference(cnf_transformation,[],[f23048_D])).
fof(f23048_D,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK0))) <=> ~$spl256),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl256])])).
fof(f23805,plain,(
  ( ! [X6] : (s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK0))) | s(bool,'T_0') = s(bool,X6)) ) | ($spl5 | $spl22 | $spl149)),
  inference(subsumption_resolution,[],[f23796,f184])).
fof(f23796,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK0))) | s(bool,'T_0') = s(bool,X6)) ) | ($spl5 | $spl22 | $spl149)),
  inference(superposition,[],[f22955,f10081])).
fof(f23804,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK0)))) ) | ($spl5 | $spl22 | $spl149)),
  inference(subsumption_resolution,[],[f23795,f184])).
fof(f23795,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK0)))) ) | ($spl5 | $spl22 | $spl149)),
  inference(superposition,[],[f22955,f10081])).
fof(f18344,plain,(
  ( ! [X277,X278] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X278),s('type.integer.int',X277)))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X277),s('type.integer.int',X278)))) ) | $spl230),
  inference(cnf_transformation,[],[f18344_D])).
fof(f18344_D,plain,(
  ( ! [X277,X278] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X278),s('type.integer.int',X277)))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X277),s('type.integer.int',X278)))) ) <=> ~$spl230),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl230])])).
fof(f23777,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,X34)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X33),s('type.integer.int',X32)))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X32),s('type.integer.int',X33)))) ) | $spl5),
  inference(subsumption_resolution,[],[f23704,f183])).
fof(f23704,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,X34)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X33),s('type.integer.int',X32)))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X32),s('type.integer.int',X33))) | s(bool,'T_0') = s(bool,X34)) ) | $spl5),
  inference(superposition,[],[f158,f10081])).
fof(f23776,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X31)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X30),s('type.integer.int',X29)))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X29),s('type.integer.int',X30)))) ) | $spl5),
  inference(subsumption_resolution,[],[f23703,f183])).
fof(f23703,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X31)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X30),s('type.integer.int',X29)))) | s(bool,'T_0') = s(bool,X31) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X29),s('type.integer.int',X30)))) ) | $spl5),
  inference(superposition,[],[f158,f10081])).
fof(f18318,plain,(
  ( ! [X165,X164] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X165),s('type.integer.int',X164)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X164),s('type.integer.int',X165)))) ) | $spl206),
  inference(cnf_transformation,[],[f18318_D])).
fof(f18318_D,plain,(
  ( ! [X165,X164] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X165),s('type.integer.int',X164)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X164),s('type.integer.int',X165)))) ) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f23743,plain,(
  ( ! [X35,X36,X34] : (p(s(bool,X36)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X35),s('type.integer.int',X34)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X34),s('type.integer.int',X35)))) ) | $spl5),
  inference(subsumption_resolution,[],[f23661,f183])).
fof(f23661,plain,(
  ( ! [X35,X36,X34] : (p(s(bool,X36)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X35),s('type.integer.int',X34)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X34),s('type.integer.int',X35))) | s(bool,'T_0') = s(bool,X36)) ) | $spl5),
  inference(superposition,[],[f158,f10081])).
fof(f23742,plain,(
  ( ! [X33,X31,X32] : (p(s(bool,X33)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X32),s('type.integer.int',X31)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X31),s('type.integer.int',X32)))) ) | $spl5),
  inference(subsumption_resolution,[],[f23660,f183])).
fof(f23660,plain,(
  ( ! [X33,X31,X32] : (p(s(bool,X33)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X32),s('type.integer.int',X31)))) | s(bool,'T_0') = s(bool,X33) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X31),s('type.integer.int',X32)))) ) | $spl5),
  inference(superposition,[],[f158,f10081])).
fof(f23513,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | $spl258),
  inference(cnf_transformation,[],[f23513_D])).
fof(f23513_D,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) <=> ~$spl258),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl258])])).
fof(f23518,plain,(
  ( ! [X8] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X8),s('type.integer.int',X8))))) ) | $spl262),
  inference(cnf_transformation,[],[f23518_D])).
fof(f23518_D,plain,(
  ( ! [X8] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X8),s('type.integer.int',X8))))) ) <=> ~$spl262),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl262])])).
fof(f23509,plain,(
  ( ! [X8] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X8),s('type.integer.int',X8)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))) ) | ($spl5 | $spl22 | $spl149)),
  inference(superposition,[],[f22901,f11923])).
fof(f11923,plain,(
  ( ! [X33,X32] : (s(bool,X32) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X33),s('type.integer.int',X33))) | s(bool,'T_0') = s(bool,X32)) ) | $spl5),
  inference(resolution,[],[f10258,f10385])).
fof(f10258,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | s(bool,'T_0') = s(bool,X3) | s(bool,X2) = s(bool,X3)) ) | $spl5),
  inference(subsumption_resolution,[],[f10082,f200])).
fof(f10082,plain,(
  ( ! [X2,X3] : (s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X3) | p(s(bool,X2)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))))) ) | $spl5),
  inference(superposition,[],[f1892,f165])).
fof(f23515,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK1)))) | $spl261),
  inference(cnf_transformation,[],[f23515_D])).
fof(f23515_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK1)))) <=> ~$spl261),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl261])])).
fof(f23504,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK1)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | ($spl0 | $spl22 | $spl149)),
  inference(superposition,[],[f22901,f4356])).
fof(f4356,plain,(
  ( ! [X25] : (s(bool,X25) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X25)) ) | ($spl0 | $spl22)),
  inference(resolution,[],[f1509,f186])).
fof(f1509,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,X0) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK1)))) ) | ($spl0 | $spl22)),
  inference(resolution,[],[f1353,f165])).
fof(f1353,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK1)))) | ($spl0 | $spl22)),
  inference(backward_demodulation,[],[f655,f629])).
fof(f629,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int',sK1)))) | $spl0),
  inference(resolution,[],[f195,f160])).
fof(f869,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK0)))) | $spl45),
  inference(cnf_transformation,[],[f869_D])).
fof(f869_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK0)))) <=> ~$spl45),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl45])])).
fof(f23503,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK0)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | ($spl5 | $spl22 | $spl149)),
  inference(superposition,[],[f22901,f4273])).
fof(f4273,plain,(
  ( ! [X25] : (s(bool,X25) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) | s(bool,'T_0') = s(bool,X25)) ) | $spl5),
  inference(resolution,[],[f933,f186])).
fof(f933,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,X0) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK0)))) ) | $spl5),
  inference(resolution,[],[f255,f165])).
fof(f23333,plain,(
  $false | ($spl5 | $spl22 | $spl244)),
  inference(subsumption_resolution,[],[f23133,f189])).
fof(f23133,plain,(
  ~p(s(bool,'T_0')) | ($spl5 | $spl22 | $spl244)),
  inference(backward_demodulation,[],[f19411,f11259])).
fof(f23332,plain,(
  $false | $spl244),
  inference(subsumption_resolution,[],[f23128,f189])).
fof(f23128,plain,(
  ~p(s(bool,'T_0')) | $spl244),
  inference(backward_demodulation,[],[f19411,f2508])).
fof(f23331,plain,(
  $false | ($spl5 | $spl244)),
  inference(subsumption_resolution,[],[f23127,f189])).
fof(f23127,plain,(
  ~p(s(bool,'T_0')) | ($spl5 | $spl244)),
  inference(backward_demodulation,[],[f19411,f13928])).
fof(f13928,plain,(
  ( ! [X160] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X160))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X160))))))))) ) | $spl5),
  inference(subsumption_resolution,[],[f13927,f193])).
fof(f13927,plain,(
  ( ! [X160] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X160))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X160))))))))) ) | $spl5),
  inference(forward_demodulation,[],[f13784,f2511])).
fof(f13784,plain,(
  ( ! [X160] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X160))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X160)))))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X160))))))))) ) | $spl5),
  inference(superposition,[],[f2757,f13286])).
fof(f2757,plain,(
  ( ! [X5] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X5),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X5)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X5))))) )),
  inference(superposition,[],[f159,f149])).
fof(f23330,plain,(
  $false | ($spl5 | $spl244)),
  inference(subsumption_resolution,[],[f23126,f189])).
fof(f23126,plain,(
  ~p(s(bool,'T_0')) | ($spl5 | $spl244)),
  inference(backward_demodulation,[],[f19411,f11176])).
fof(f11176,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))))))) ) | $spl5),
  inference(superposition,[],[f10740,f150])).
fof(f10740,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))))),s('type.integer.int',X0))))) ) | $spl5),
  inference(superposition,[],[f10540,f141])).
fof(f141,plain,(
  ( ! [X0,X1] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))) )),
  inference(cnf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0,X1] : s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X9,X10] : s(bool,'const.integer.int_lt_2'(s('type.integer.int',X10),s('type.integer.int',X9))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X10)))))),
  file('thm.integer.INT_ABS_LE|split|1','thm.integer.INT_LT_NEG')).
fof(f10540,plain,(
  ( ! [X2] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X2))))))))) ) | $spl5),
  inference(resolution,[],[f10382,f160])).
fof(f23329,plain,(
  $false | ($spl5 | $spl244)),
  inference(subsumption_resolution,[],[f23122,f189])).
fof(f23122,plain,(
  ~p(s(bool,'T_0')) | ($spl5 | $spl244)),
  inference(backward_demodulation,[],[f19411,f10540])).
fof(f23328,plain,(
  $false | $spl244),
  inference(subsumption_resolution,[],[f23327,f189])).
fof(f23327,plain,(
  ~p(s(bool,'T_0')) | $spl244),
  inference(forward_demodulation,[],[f23326,f19411])).
fof(f23326,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0))))) ) | $spl244),
  inference(subsumption_resolution,[],[f23121,f193])).
fof(f23121,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0))))) ) | $spl244),
  inference(backward_demodulation,[],[f19411,f169])).
fof(f23325,plain,(
  $false | ($spl5 | $spl244)),
  inference(subsumption_resolution,[],[f23110,f189])).
fof(f23110,plain,(
  ~p(s(bool,'T_0')) | ($spl5 | $spl244)),
  inference(backward_demodulation,[],[f19411,f200])).
fof(f23324,plain,(
  $false | $spl244),
  inference(subsumption_resolution,[],[f23323,f189])).
fof(f23323,plain,(
  ~p(s(bool,'T_0')) | $spl244),
  inference(forward_demodulation,[],[f23322,f19411])).
fof(f23322,plain,(
  ( ! [X9] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9))),s('type.integer.int',X9))))) ) | $spl244),
  inference(subsumption_resolution,[],[f23105,f193])).
fof(f23105,plain,(
  ( ! [X9] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9))),s('type.integer.int',X9))))) ) | $spl244),
  inference(backward_demodulation,[],[f19411,f2622])).
fof(f2622,plain,(
  ( ! [X9] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9))),s('type.integer.int',X9)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X9),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f159,f148])).
fof(f23321,plain,(
  $false | ($spl149 | $spl244)),
  inference(subsumption_resolution,[],[f23103,f189])).
fof(f23103,plain,(
  ~p(s(bool,'T_0')) | ($spl149 | $spl244)),
  inference(backward_demodulation,[],[f19411,f11318])).
fof(f23320,plain,(
  $false | ($spl22 | $spl149 | $spl244)),
  inference(subsumption_resolution,[],[f23100,f189])).
fof(f23100,plain,(
  ~p(s(bool,'T_0')) | ($spl22 | $spl149 | $spl244)),
  inference(backward_demodulation,[],[f19411,f22901])).
fof(f23316,plain,(
  $false | $spl244),
  inference(subsumption_resolution,[],[f23315,f189])).
fof(f23315,plain,(
  ~p(s(bool,'T_0')) | $spl244),
  inference(forward_demodulation,[],[f23314,f19411])).
fof(f23314,plain,(
  ( ! [X8] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X8))))) ) | $spl244),
  inference(subsumption_resolution,[],[f23090,f193])).
fof(f23090,plain,(
  ( ! [X8] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X8))))) ) | $spl244),
  inference(backward_demodulation,[],[f19411,f2588])).
fof(f23311,plain,(
  $false | $spl244),
  inference(subsumption_resolution,[],[f23310,f189])).
fof(f23310,plain,(
  ~p(s(bool,'T_0')) | $spl244),
  inference(forward_demodulation,[],[f23309,f19411])).
fof(f23309,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0))))) ) | $spl244),
  inference(subsumption_resolution,[],[f23081,f193])).
fof(f23081,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0))))) ) | $spl244),
  inference(backward_demodulation,[],[f19411,f160])).
fof(f23308,plain,(
  $false | $spl244),
  inference(subsumption_resolution,[],[f23307,f189])).
fof(f23307,plain,(
  ~p(s(bool,'T_0')) | $spl244),
  inference(forward_demodulation,[],[f23306,f19411])).
fof(f23306,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) ) | $spl244),
  inference(subsumption_resolution,[],[f23080,f193])).
fof(f23080,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) ) | $spl244),
  inference(backward_demodulation,[],[f19411,f159])).
fof(f23301,plain,(
  $false | $spl244),
  inference(subsumption_resolution,[],[f23300,f189])).
fof(f23300,plain,(
  ~p(s(bool,'T_0')) | $spl244),
  inference(forward_demodulation,[],[f23299,f19411])).
fof(f23299,plain,(
  ( ! [X5] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X5),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X5))))))) ) | $spl244),
  inference(subsumption_resolution,[],[f23065,f193])).
fof(f23065,plain,(
  ( ! [X5] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X5),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X5))))))) ) | $spl244),
  inference(backward_demodulation,[],[f19411,f2757])).
fof(f23298,plain,(
  $false | $spl244),
  inference(subsumption_resolution,[],[f23297,f189])).
fof(f23297,plain,(
  ~p(s(bool,'T_0')) | $spl244),
  inference(forward_demodulation,[],[f23296,f19411])).
fof(f23296,plain,(
  ( ! [X3] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl244),
  inference(subsumption_resolution,[],[f23058,f193])).
fof(f23058,plain,(
  ( ! [X3] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl244),
  inference(backward_demodulation,[],[f19411,f2727])).
fof(f2727,plain,(
  ( ! [X3] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X3))),s('type.integer.int',X3))))) )),
  inference(superposition,[],[f159,f149])).
fof(f23275,plain,(
  $false | ($spl3 | $spl244)),
  inference(trivial_inequality_removal,[],[f23230])).
fof(f23230,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl3 | $spl244)),
  inference(superposition,[],[f975,f19411])).
fof(f23274,plain,(
  $false | ($spl3 | $spl22 | $spl244)),
  inference(trivial_inequality_removal,[],[f23231])).
fof(f23231,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl3 | $spl22 | $spl244)),
  inference(superposition,[],[f1739,f19411])).
fof(f1739,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | ($spl3 | $spl22)),
  inference(resolution,[],[f1733,f185])).
fof(f1733,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | ($spl3 | $spl22)),
  inference(forward_demodulation,[],[f974,f655])).
fof(f974,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | $spl3),
  inference(resolution,[],[f198,f137])).
fof(f23273,plain,(
  $false | ($spl149 | $spl244)),
  inference(trivial_inequality_removal,[],[f23232])).
fof(f23232,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl149 | $spl244)),
  inference(superposition,[],[f18930,f19411])).
fof(f18930,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | $spl149),
  inference(resolution,[],[f11318,f183])).
fof(f23272,plain,(
  $false | ($spl29 | $spl41 | $spl244)),
  inference(trivial_inequality_removal,[],[f23233])).
fof(f23233,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl29 | $spl41 | $spl244)),
  inference(superposition,[],[f5951,f19411])).
fof(f23270,plain,(
  $false | ($spl5 | $spl244)),
  inference(trivial_inequality_removal,[],[f23235])).
fof(f23235,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl5 | $spl244)),
  inference(superposition,[],[f884,f19411])).
fof(f23268,plain,(
  $false | ($spl5 | $spl244)),
  inference(trivial_inequality_removal,[],[f23237])).
fof(f23237,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl5 | $spl244)),
  inference(superposition,[],[f10406,f19411])).
fof(f23267,plain,(
  $false | ($spl5 | $spl244)),
  inference(trivial_inequality_removal,[],[f23238])).
fof(f23238,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl5 | $spl244)),
  inference(superposition,[],[f10508,f19411])).
fof(f10508,plain,(
  ( ! [X7] : (s(bool,'T_0') != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X7))),s('type.integer.int',X7)))) ) | $spl5),
  inference(resolution,[],[f10441,f183])).
fof(f23266,plain,(
  $false | ($spl5 | $spl244)),
  inference(trivial_inequality_removal,[],[f23239])).
fof(f23239,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl5 | $spl244)),
  inference(superposition,[],[f10719,f19411])).
fof(f10719,plain,(
  ( ! [X7] : (s(bool,'T_0') != s(bool,'const.integer.int_lt_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X7)))))))) ) | $spl5),
  inference(resolution,[],[f10540,f183])).
fof(f23265,plain,(
  $false | ($spl5 | $spl244)),
  inference(trivial_inequality_removal,[],[f23240])).
fof(f23240,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl5 | $spl244)),
  inference(superposition,[],[f11384,f19411])).
fof(f11384,plain,(
  ( ! [X7] : (s(bool,'T_0') != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X7))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X7)))))) ) | $spl5),
  inference(resolution,[],[f11176,f183])).
fof(f23264,plain,(
  $false | ($spl5 | $spl152 | $spl244)),
  inference(trivial_inequality_removal,[],[f23241])).
fof(f23241,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl5 | $spl152 | $spl244)),
  inference(superposition,[],[f13331,f19411])).
fof(f13331,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | ($spl5 | $spl152)),
  inference(superposition,[],[f10719,f11519])).
fof(f23263,plain,(
  $false | ($spl96 | $spl244)),
  inference(trivial_inequality_removal,[],[f23242])).
fof(f23242,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl96 | $spl244)),
  inference(superposition,[],[f22089,f19411])).
fof(f23262,plain,(
  $false | ($spl7 | $spl244)),
  inference(trivial_inequality_removal,[],[f23243])).
fof(f23243,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl7 | $spl244)),
  inference(superposition,[],[f1022,f19411])).
fof(f1022,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | $spl7),
  inference(resolution,[],[f248,f185])).
fof(f23261,plain,(
  $false | ($spl22 | $spl24 | $spl244)),
  inference(trivial_inequality_removal,[],[f23244])).
fof(f23244,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl22 | $spl24 | $spl244)),
  inference(superposition,[],[f1831,f19411])).
fof(f1831,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | ($spl22 | $spl24)),
  inference(resolution,[],[f1659,f185])).
fof(f1659,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | ($spl22 | $spl24)),
  inference(resolution,[],[f1624,f160])).
fof(f1624,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int',sK1)))) | ($spl22 | $spl24)),
  inference(forward_demodulation,[],[f656,f655])).
fof(f656,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))))),s('type.integer.int',sK1)))) | $spl24),
  inference(cnf_transformation,[],[f656_D])).
fof(f656_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))))),s('type.integer.int',sK1)))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f23260,plain,(
  $false | ($spl5 | $spl22 | $spl244)),
  inference(trivial_inequality_removal,[],[f23245])).
fof(f23245,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl5 | $spl22 | $spl244)),
  inference(superposition,[],[f11296,f19411])).
fof(f11296,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | ($spl5 | $spl22)),
  inference(resolution,[],[f11259,f183])).
fof(f23259,plain,(
  $false | $spl244),
  inference(trivial_inequality_removal,[],[f23246])).
fof(f23246,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl244),
  inference(superposition,[],[f3032,f19411])).
fof(f3032,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X4))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(resolution,[],[f2508,f185])).
fof(f23258,plain,(
  $false | $spl244),
  inference(trivial_inequality_removal,[],[f23247])).
fof(f23247,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl244),
  inference(superposition,[],[f6576,f19411])).
fof(f6576,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(resolution,[],[f5396,f183])).
fof(f5396,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(forward_demodulation,[],[f5395,f147])).
fof(f5395,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(subsumption_resolution,[],[f5375,f163])).
fof(f5375,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(superposition,[],[f2622,f151])).
fof(f23257,plain,(
  $false | ($spl22 | $spl244)),
  inference(trivial_inequality_removal,[],[f23248])).
fof(f23248,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl22 | $spl244)),
  inference(superposition,[],[f4551,f19411])).
fof(f4551,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl22),
  inference(superposition,[],[f3032,f655])).
fof(f23255,plain,(
  $false | ($spl5 | $spl244)),
  inference(trivial_inequality_removal,[],[f23250])).
fof(f23250,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl5 | $spl244)),
  inference(superposition,[],[f931,f19411])).
fof(f931,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) | $spl5),
  inference(resolution,[],[f255,f185])).
fof(f23253,plain,(
  $false | ($spl0 | $spl22 | $spl244)),
  inference(trivial_inequality_removal,[],[f23252])).
fof(f23252,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl0 | $spl22 | $spl244)),
  inference(superposition,[],[f1507,f19411])).
fof(f1507,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK1))) | ($spl0 | $spl22)),
  inference(resolution,[],[f1353,f185])).
fof(f23049,plain,(
  ( ! [X6] : (s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK0))) | s(bool,'T_0') = s(bool,X6)) ) | ($spl5 | $spl22 | $spl149)),
  inference(subsumption_resolution,[],[f23031,f184])).
fof(f23031,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK0))) | s(bool,'T_0') = s(bool,X6)) ) | ($spl5 | $spl22 | $spl149)),
  inference(superposition,[],[f22955,f10081])).
fof(f23046,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK0)))) ) | ($spl5 | $spl22 | $spl149)),
  inference(subsumption_resolution,[],[f23030,f184])).
fof(f23030,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK0)))) ) | ($spl5 | $spl22 | $spl149)),
  inference(superposition,[],[f22955,f10081])).
fof(f22949,plain,(
  ( ! [X15] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X15),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X15))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X15)))))) ) | $spl254),
  inference(cnf_transformation,[],[f22949_D])).
fof(f22949_D,plain,(
  ( ! [X15] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X15),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X15))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X15)))))) ) <=> ~$spl254),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl254])])).
fof(f22950,plain,(
  ( ! [X17,X18] : (p(s(bool,X18)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X17),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X17))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X17)))))) ) | $spl5),
  inference(subsumption_resolution,[],[f22879,f183])).
fof(f22879,plain,(
  ( ! [X17,X18] : (p(s(bool,X18)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X17),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X17))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X17))))) | s(bool,'T_0') = s(bool,X18)) ) | $spl5),
  inference(superposition,[],[f2655,f10081])).
fof(f22948,plain,(
  ( ! [X15,X16] : (p(s(bool,X16)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X15),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X15))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X15)))))) ) | $spl5),
  inference(subsumption_resolution,[],[f22878,f183])).
fof(f22878,plain,(
  ( ! [X15,X16] : (p(s(bool,X16)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X15),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,'T_0') = s(bool,X16) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X15))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X15)))))) ) | $spl5),
  inference(superposition,[],[f2655,f10081])).
fof(f22069,plain,(
  $false | ($spl5 | $spl22 | $spl24 | $spl131)),
  inference(subsumption_resolution,[],[f22068,f189])).
fof(f22068,plain,(
  ~p(s(bool,'T_0')) | ($spl5 | $spl22 | $spl24 | $spl131)),
  inference(forward_demodulation,[],[f22067,f1662])).
fof(f1662,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int',sK1))) | ($spl22 | $spl24)),
  inference(resolution,[],[f1624,f186])).
fof(f22067,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int',sK1)))) | ($spl5 | $spl131)),
  inference(forward_demodulation,[],[f22066,f149])).
fof(f22066,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | ($spl5 | $spl131)),
  inference(forward_demodulation,[],[f22065,f2216])).
fof(f22065,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl5 | $spl131)),
  inference(forward_demodulation,[],[f22064,f9243])).
fof(f22064,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | ($spl5 | $spl131)),
  inference(forward_demodulation,[],[f8925,f21479])).
fof(f21934,plain,(
  $false | ($spl5 | $spl115)),
  inference(subsumption_resolution,[],[f21933,f189])).
fof(f21933,plain,(
  ~p(s(bool,'T_0')) | ($spl5 | $spl115)),
  inference(forward_demodulation,[],[f21932,f10446])).
fof(f21932,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))))) | ($spl5 | $spl115)),
  inference(forward_demodulation,[],[f7467,f21479])).
fof(f21377,plain,(
  ( ! [X18] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int',X18),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X18),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X18)))))) ) | $spl252),
  inference(cnf_transformation,[],[f21377_D])).
fof(f21377_D,plain,(
  ( ! [X18] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int',X18),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X18),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X18)))))) ) <=> ~$spl252),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl252])])).
fof(f21378,plain,(
  ( ! [X21,X20] : (p(s(bool,X21)) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int',X20),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X20),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X20)))))) ) | $spl5),
  inference(subsumption_resolution,[],[f21326,f183])).
fof(f21326,plain,(
  ( ! [X21,X20] : (p(s(bool,X21)) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int',X20),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X20),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X20))))) | s(bool,'T_0') = s(bool,X21)) ) | $spl5),
  inference(superposition,[],[f2605,f10081])).
fof(f2605,plain,(
  ( ! [X33] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X33),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X33)))))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int',X33),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f183,f148])).
fof(f21376,plain,(
  ( ! [X19,X18] : (p(s(bool,X19)) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int',X18),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X18),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X18)))))) ) | $spl5),
  inference(subsumption_resolution,[],[f21325,f183])).
fof(f21325,plain,(
  ( ! [X19,X18] : (p(s(bool,X19)) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int',X18),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X18),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X18)))))) ) | $spl5),
  inference(superposition,[],[f2605,f10081])).
fof(f21057,plain,(
  ( ! [X18] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int',X18),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X18))))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X18),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl250),
  inference(cnf_transformation,[],[f21057_D])).
fof(f21057_D,plain,(
  ( ! [X18] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int',X18),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X18))))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X18),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) <=> ~$spl250),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl250])])).
fof(f21058,plain,(
  ( ! [X21,X20] : (p(s(bool,X21)) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int',X20),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X20))))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X20),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl5),
  inference(subsumption_resolution,[],[f21004,f183])).
fof(f21004,plain,(
  ( ! [X21,X20] : (p(s(bool,X21)) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int',X20),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X20))))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X20),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X21)) ) | $spl5),
  inference(superposition,[],[f2604,f10081])).
fof(f2604,plain,(
  ( ! [X32] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X32),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int',X32),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X32)))))) )),
  inference(superposition,[],[f183,f148])).
fof(f21056,plain,(
  ( ! [X19,X18] : (p(s(bool,X19)) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int',X18),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X18))))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X18),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl5),
  inference(subsumption_resolution,[],[f21003,f183])).
fof(f21003,plain,(
  ( ! [X19,X18] : (p(s(bool,X19)) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int',X18),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X18))))) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X18),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl5),
  inference(superposition,[],[f2604,f10081])).
fof(f20285,plain,(
  ( ! [X15] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X15))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X15)))))) ) | $spl248),
  inference(cnf_transformation,[],[f20285_D])).
fof(f20285_D,plain,(
  ( ! [X15] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X15))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X15)))))) ) <=> ~$spl248),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl248])])).
fof(f20286,plain,(
  ( ! [X17,X18] : (p(s(bool,X18)) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X17))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X17)))))) ) | $spl5),
  inference(subsumption_resolution,[],[f20243,f183])).
fof(f20243,plain,(
  ( ! [X17,X18] : (p(s(bool,X18)) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X17))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X17))))) | s(bool,'T_0') = s(bool,X18)) ) | $spl5),
  inference(superposition,[],[f2507,f10081])).
fof(f20284,plain,(
  ( ! [X15,X16] : (p(s(bool,X16)) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X15))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X15)))))) ) | $spl5),
  inference(subsumption_resolution,[],[f20242,f183])).
fof(f20242,plain,(
  ( ! [X15,X16] : (p(s(bool,X16)) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X15))) | s(bool,'T_0') = s(bool,X16) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X15)))))) ) | $spl5),
  inference(superposition,[],[f2507,f10081])).
fof(f17073,plain,(
  s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | $spl170),
  inference(cnf_transformation,[],[f17073_D])).
fof(f17073_D,plain,(
  s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f20193,plain,(
  ( ! [X56] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | p(s(bool,X56))) ) | ($spl0 | $spl5 | $spl22 | $spl149)),
  inference(forward_demodulation,[],[f20192,f1891])).
fof(f1891,plain,(
  s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK1))) | ($spl0 | $spl5 | $spl22)),
  inference(resolution,[],[f886,f1353])).
fof(f20192,plain,(
  ( ! [X56] : (p(s(bool,X56)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK1))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))) ) | ($spl0 | $spl22 | $spl149)),
  inference(subsumption_resolution,[],[f20160,f183])).
fof(f20160,plain,(
  ( ! [X56] : (p(s(bool,X56)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK1))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | s(bool,'T_0') = s(bool,X56)) ) | ($spl0 | $spl22 | $spl149)),
  inference(superposition,[],[f18928,f4356])).
fof(f18928,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,X0) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))) ) | $spl149),
  inference(resolution,[],[f11318,f165])).
fof(f20189,plain,(
  ( ! [X55] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | p(s(bool,X55))) ) | ($spl5 | $spl149)),
  inference(forward_demodulation,[],[f20188,f1890])).
fof(f1890,plain,(
  s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) | $spl5),
  inference(resolution,[],[f886,f255])).
fof(f20188,plain,(
  ( ! [X55] : (p(s(bool,X55)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))) ) | ($spl5 | $spl149)),
  inference(subsumption_resolution,[],[f20154,f183])).
fof(f20154,plain,(
  ( ! [X55] : (p(s(bool,X55)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | s(bool,'T_0') = s(bool,X55)) ) | ($spl5 | $spl149)),
  inference(superposition,[],[f18928,f4273])).
fof(f20182,plain,(
  ( ! [X49] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | p(s(bool,X49))) ) | ($spl5 | $spl149)),
  inference(forward_demodulation,[],[f20181,f10408])).
fof(f10408,plain,(
  ( ! [X9] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X9),s('type.integer.int',X9)))) ) | $spl5),
  inference(resolution,[],[f10385,f886])).
fof(f20181,plain,(
  ( ! [X48,X49] : (p(s(bool,X49)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',X48),s('type.integer.int',X48))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))) ) | ($spl5 | $spl149)),
  inference(subsumption_resolution,[],[f20148,f183])).
fof(f20148,plain,(
  ( ! [X48,X49] : (p(s(bool,X49)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',X48),s('type.integer.int',X48))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | s(bool,'T_0') = s(bool,X49)) ) | ($spl5 | $spl149)),
  inference(superposition,[],[f18928,f11923])).
fof(f20180,plain,(
  ( ! [X47] : (p(s(bool,X47)) | s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))) ) | ($spl5 | $spl149)),
  inference(subsumption_resolution,[],[f20142,f183])).
fof(f20142,plain,(
  ( ! [X47] : (p(s(bool,X47)) | s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | s(bool,'T_0') = s(bool,X47)) ) | ($spl5 | $spl149)),
  inference(superposition,[],[f18928,f1892])).
fof(f18298,plain,(
  ( ! [X70] : (s(bool,X70) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | s(bool,'T_0') = s(bool,X70)) ) | $spl186),
  inference(cnf_transformation,[],[f18298_D])).
fof(f18298_D,plain,(
  ( ! [X70] : (s(bool,X70) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | s(bool,'T_0') = s(bool,X70)) ) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f20170,plain,(
  ( ! [X17,X18] : (p(s(bool,X18)) | s(bool,X17) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | s(bool,'T_0') = s(bool,X17)) ) | ($spl5 | $spl149)),
  inference(subsumption_resolution,[],[f20114,f183])).
fof(f20114,plain,(
  ( ! [X17,X18] : (p(s(bool,X18)) | s(bool,X17) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | s(bool,'T_0') = s(bool,X17) | s(bool,'T_0') = s(bool,X18)) ) | ($spl5 | $spl149)),
  inference(superposition,[],[f18928,f10081])).
fof(f20169,plain,(
  ( ! [X15,X16] : (p(s(bool,X16)) | s(bool,X15) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | s(bool,'T_0') = s(bool,X15)) ) | ($spl5 | $spl149)),
  inference(subsumption_resolution,[],[f20113,f183])).
fof(f20113,plain,(
  ( ! [X15,X16] : (p(s(bool,X16)) | s(bool,X15) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | s(bool,'T_0') = s(bool,X16) | s(bool,'T_0') = s(bool,X15)) ) | ($spl5 | $spl149)),
  inference(superposition,[],[f18928,f10081])).
fof(f16935,plain,(
  s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | $spl168),
  inference(cnf_transformation,[],[f16935_D])).
fof(f16935_D,plain,(
  s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f20049,plain,(
  ( ! [X56] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | p(s(bool,X56))) ) | ($spl0 | $spl5 | $spl22)),
  inference(forward_demodulation,[],[f20048,f1891])).
fof(f20048,plain,(
  ( ! [X56] : (p(s(bool,X56)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK1))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))) ) | ($spl0 | $spl5 | $spl22)),
  inference(subsumption_resolution,[],[f20016,f183])).
fof(f20016,plain,(
  ( ! [X56] : (p(s(bool,X56)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK1))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | s(bool,'T_0') = s(bool,X56)) ) | ($spl0 | $spl5 | $spl22)),
  inference(superposition,[],[f11294,f4356])).
fof(f11294,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,X0) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))) ) | ($spl5 | $spl22)),
  inference(resolution,[],[f11259,f165])).
fof(f20045,plain,(
  ( ! [X55] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | p(s(bool,X55))) ) | ($spl5 | $spl22)),
  inference(forward_demodulation,[],[f20044,f1890])).
fof(f20044,plain,(
  ( ! [X55] : (p(s(bool,X55)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))) ) | ($spl5 | $spl22)),
  inference(subsumption_resolution,[],[f20010,f183])).
fof(f20010,plain,(
  ( ! [X55] : (p(s(bool,X55)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | s(bool,'T_0') = s(bool,X55)) ) | ($spl5 | $spl22)),
  inference(superposition,[],[f11294,f4273])).
fof(f20038,plain,(
  ( ! [X49] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | p(s(bool,X49))) ) | ($spl5 | $spl22)),
  inference(forward_demodulation,[],[f20037,f10408])).
fof(f20037,plain,(
  ( ! [X48,X49] : (p(s(bool,X49)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',X48),s('type.integer.int',X48))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))) ) | ($spl5 | $spl22)),
  inference(subsumption_resolution,[],[f20004,f183])).
fof(f20004,plain,(
  ( ! [X48,X49] : (p(s(bool,X49)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',X48),s('type.integer.int',X48))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | s(bool,'T_0') = s(bool,X49)) ) | ($spl5 | $spl22)),
  inference(superposition,[],[f11294,f11923])).
fof(f20036,plain,(
  ( ! [X47] : (p(s(bool,X47)) | s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))) ) | ($spl5 | $spl22)),
  inference(subsumption_resolution,[],[f19998,f183])).
fof(f19998,plain,(
  ( ! [X47] : (p(s(bool,X47)) | s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | s(bool,'T_0') = s(bool,X47)) ) | ($spl5 | $spl22)),
  inference(superposition,[],[f11294,f1892])).
fof(f18296,plain,(
  ( ! [X68] : (s(bool,X68) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | s(bool,'T_0') = s(bool,X68)) ) | $spl184),
  inference(cnf_transformation,[],[f18296_D])).
fof(f18296_D,plain,(
  ( ! [X68] : (s(bool,X68) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | s(bool,'T_0') = s(bool,X68)) ) <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f20026,plain,(
  ( ! [X17,X18] : (p(s(bool,X18)) | s(bool,X17) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | s(bool,'T_0') = s(bool,X17)) ) | ($spl5 | $spl22)),
  inference(subsumption_resolution,[],[f19970,f183])).
fof(f19970,plain,(
  ( ! [X17,X18] : (p(s(bool,X18)) | s(bool,X17) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | s(bool,'T_0') = s(bool,X17) | s(bool,'T_0') = s(bool,X18)) ) | ($spl5 | $spl22)),
  inference(superposition,[],[f11294,f10081])).
fof(f20025,plain,(
  ( ! [X15,X16] : (p(s(bool,X16)) | s(bool,X15) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | s(bool,'T_0') = s(bool,X15)) ) | ($spl5 | $spl22)),
  inference(subsumption_resolution,[],[f19969,f183])).
fof(f19969,plain,(
  ( ! [X15,X16] : (p(s(bool,X16)) | s(bool,X15) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | s(bool,'T_0') = s(bool,X16) | s(bool,'T_0') = s(bool,X15)) ) | ($spl5 | $spl22)),
  inference(superposition,[],[f11294,f10081])).
fof(f19455,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK1))) | $spl246),
  inference(cnf_transformation,[],[f19455_D])).
fof(f19455_D,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK1))) <=> ~$spl246),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl246])])).
fof(f19456,plain,(
  ( ! [X6] : (s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X6)) ) | ($spl5 | $spl149)),
  inference(subsumption_resolution,[],[f19443,f184])).
fof(f19443,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X6)) ) | ($spl5 | $spl149)),
  inference(superposition,[],[f11348,f10081])).
fof(f11348,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK1)))) | $spl149),
  inference(resolution,[],[f11318,f161])).
fof(f19453,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK1)))) ) | ($spl5 | $spl149)),
  inference(subsumption_resolution,[],[f19442,f184])).
fof(f19442,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK1)))) ) | ($spl5 | $spl149)),
  inference(superposition,[],[f11348,f10081])).
fof(f19410,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK1))) | $spl242),
  inference(cnf_transformation,[],[f19410_D])).
fof(f19410_D,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK1))) <=> ~$spl242),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl242])])).
fof(f19412,plain,(
  ( ! [X6] : (s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X6)) ) | ($spl5 | $spl22)),
  inference(subsumption_resolution,[],[f19394,f184])).
fof(f19394,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X6)) ) | ($spl5 | $spl22)),
  inference(superposition,[],[f18881,f10081])).
fof(f19408,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK1)))) ) | ($spl5 | $spl22)),
  inference(subsumption_resolution,[],[f19393,f184])).
fof(f19393,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK1)))) ) | ($spl5 | $spl22)),
  inference(superposition,[],[f18881,f10081])).
fof(f18690,plain,(
  $false | ($spl9 | $spl146)),
  inference(subsumption_resolution,[],[f18529,f250])).
fof(f250,plain,(
  s('type.integer.int',sK0) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))) | $spl9),
  inference(cnf_transformation,[],[f250_D])).
fof(f250_D,plain,(
  s('type.integer.int',sK0) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f18529,plain,(
  s('type.integer.int',sK0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))) | $spl146),
  inference(superposition,[],[f150,f11286])).
fof(f18680,plain,(
  $false | ($spl0 | $spl22 | $spl144 | $spl146)),
  inference(subsumption_resolution,[],[f18511,f1507])).
fof(f18511,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK1))) | ($spl144 | $spl146)),
  inference(backward_demodulation,[],[f11286,f11324])).
fof(f11324,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK1))) | $spl144),
  inference(resolution,[],[f11284,f186])).
fof(f18355,plain,(
  ( ! [X361] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int',X361)))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X361),s('type.integer.int',sK1)))) ) | $spl240),
  inference(cnf_transformation,[],[f18355_D])).
fof(f18355_D,plain,(
  ( ! [X361] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int',X361)))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X361),s('type.integer.int',sK1)))) ) <=> ~$spl240),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl240])])).
fof(f18401,plain,(
  ( ! [X362,X361] : (p(s(bool,X362)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int',X361)))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X361),s('type.integer.int',sK1)))) ) | ($spl5 | $spl22)),
  inference(subsumption_resolution,[],[f17639,f183])).
fof(f17639,plain,(
  ( ! [X362,X361] : (p(s(bool,X362)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int',X361)))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X361),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X362)) ) | ($spl5 | $spl22)),
  inference(superposition,[],[f9886,f10081])).
fof(f9886,plain,(
  ( ! [X2] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X2),s('type.integer.int',sK1)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int',X2))))) ) | $spl22),
  inference(resolution,[],[f1359,f161])).
fof(f1359,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int',X0))))) ) | $spl22),
  inference(superposition,[],[f137,f655])).
fof(f18353,plain,(
  ( ! [X341] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X341)))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X341),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int',X341) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl238),
  inference(cnf_transformation,[],[f18353_D])).
fof(f18353_D,plain,(
  ( ! [X341] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X341)))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X341),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int',X341) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl238),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl238])])).
fof(f18400,plain,(
  ( ! [X341,X342] : (p(s(bool,X342)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X341)))) | s('type.integer.int',X341) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X341),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl5),
  inference(subsumption_resolution,[],[f17626,f183])).
fof(f17626,plain,(
  ( ! [X341,X342] : (p(s(bool,X342)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X341)))) | s('type.integer.int',X341) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X341),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X342)) ) | $spl5),
  inference(superposition,[],[f3355,f10081])).
fof(f18351,plain,(
  ( ! [X337] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X337))),s('type.integer.int',X337)))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X337),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl236),
  inference(cnf_transformation,[],[f18351_D])).
fof(f18351_D,plain,(
  ( ! [X337] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X337))),s('type.integer.int',X337)))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X337),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) <=> ~$spl236),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl236])])).
fof(f18399,plain,(
  ( ! [X337,X338] : (p(s(bool,X338)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X337))),s('type.integer.int',X337)))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X337),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl5),
  inference(subsumption_resolution,[],[f17624,f183])).
fof(f17624,plain,(
  ( ! [X337,X338] : (p(s(bool,X338)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X337))),s('type.integer.int',X337)))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X337),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X338)) ) | $spl5),
  inference(superposition,[],[f2726,f10081])).
fof(f2726,plain,(
  ( ! [X2] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X2))),s('type.integer.int',X2)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f158,f149])).
fof(f18349,plain,(
  ( ! [X305] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X305)))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X305),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X305)))))) ) | $spl234),
  inference(cnf_transformation,[],[f18349_D])).
fof(f18349_D,plain,(
  ( ! [X305] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X305)))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X305),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X305)))))) ) <=> ~$spl234),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl234])])).
fof(f18398,plain,(
  ( ! [X306,X305] : (p(s(bool,X306)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X305)))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X305),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X305)))))) ) | $spl5),
  inference(subsumption_resolution,[],[f17605,f183])).
fof(f17605,plain,(
  ( ! [X306,X305] : (p(s(bool,X306)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X305)))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X305),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X305))))) | s(bool,'T_0') = s(bool,X306)) ) | $spl5),
  inference(superposition,[],[f2756,f10081])).
fof(f2756,plain,(
  ( ! [X4] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X4),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X4)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X4))))) )),
  inference(superposition,[],[f158,f149])).
fof(f18397,plain,(
  ( ! [X290,X289,X291] : (p(s(bool,X291)) | s('type.integer.int',X289) = s('type.integer.int',X290) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X289),s('type.integer.int',X290)))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X289),s('type.integer.int',X290)))) ) | $spl5),
  inference(subsumption_resolution,[],[f17599,f183])).
fof(f17599,plain,(
  ( ! [X290,X289,X291] : (p(s(bool,X291)) | s('type.integer.int',X289) = s('type.integer.int',X290) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X289),s('type.integer.int',X290)))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X289),s('type.integer.int',X290))) | s(bool,'T_0') = s(bool,X291)) ) | $spl5),
  inference(superposition,[],[f162,f10081])).
fof(f18396,plain,(
  ( ! [X288,X287,X286] : (p(s(bool,X288)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X287),s('type.integer.int',X286)))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X286),s('type.integer.int',X287)))) ) | $spl5),
  inference(subsumption_resolution,[],[f17598,f183])).
fof(f17598,plain,(
  ( ! [X288,X287,X286] : (p(s(bool,X288)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X287),s('type.integer.int',X286)))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X286),s('type.integer.int',X287))) | s(bool,'T_0') = s(bool,X288)) ) | $spl5),
  inference(superposition,[],[f161,f10081])).
fof(f18395,plain,(
  ( ! [X277,X279,X278] : (p(s(bool,X279)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X278),s('type.integer.int',X277)))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X277),s('type.integer.int',X278)))) ) | $spl5),
  inference(subsumption_resolution,[],[f17595,f183])).
fof(f17595,plain,(
  ( ! [X277,X279,X278] : (p(s(bool,X279)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X278),s('type.integer.int',X277)))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X277),s('type.integer.int',X278))) | s(bool,'T_0') = s(bool,X279)) ) | $spl5),
  inference(superposition,[],[f158,f10081])).
fof(f18342,plain,(
  ( ! [X273] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X273),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X273))) | s('type.integer.int',X273) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl228),
  inference(cnf_transformation,[],[f18342_D])).
fof(f18342_D,plain,(
  ( ! [X273] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X273),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X273))) | s('type.integer.int',X273) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl228),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl228])])).
fof(f18394,plain,(
  ( ! [X273,X274] : (p(s(bool,X274)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X273),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X273) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X273)))) ) | $spl5),
  inference(subsumption_resolution,[],[f17593,f183])).
fof(f17593,plain,(
  ( ! [X273,X274] : (p(s(bool,X274)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X273),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X273) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X273))) | s(bool,'T_0') = s(bool,X274)) ) | $spl5),
  inference(superposition,[],[f3355,f10081])).
fof(f18340,plain,(
  ( ! [X271] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X271),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X271)))))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X271)))) ) | $spl226),
  inference(cnf_transformation,[],[f18340_D])).
fof(f18340_D,plain,(
  ( ! [X271] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X271),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X271)))))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X271)))) ) <=> ~$spl226),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl226])])).
fof(f18393,plain,(
  ( ! [X271,X272] : (p(s(bool,X272)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X271),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X271)))))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X271)))) ) | $spl5),
  inference(subsumption_resolution,[],[f17592,f183])).
fof(f17592,plain,(
  ( ! [X271,X272] : (p(s(bool,X272)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X271),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X271)))))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X271))) | s(bool,'T_0') = s(bool,X272)) ) | $spl5),
  inference(superposition,[],[f2587,f10081])).
fof(f18338,plain,(
  ( ! [X265] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X265),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X265))),s('type.integer.int',X265)))) ) | $spl224),
  inference(cnf_transformation,[],[f18338_D])).
fof(f18338_D,plain,(
  ( ! [X265] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X265),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X265))),s('type.integer.int',X265)))) ) <=> ~$spl224),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl224])])).
fof(f18392,plain,(
  ( ! [X265,X266] : (p(s(bool,X266)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X265),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X265))),s('type.integer.int',X265)))) ) | $spl5),
  inference(subsumption_resolution,[],[f17589,f183])).
fof(f17589,plain,(
  ( ! [X265,X266] : (p(s(bool,X266)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X265),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X265))),s('type.integer.int',X265))) | s(bool,'T_0') = s(bool,X266)) ) | $spl5),
  inference(superposition,[],[f2621,f10081])).
fof(f2621,plain,(
  ( ! [X8] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X8))),s('type.integer.int',X8)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X8),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f158,f148])).
fof(f18336,plain,(
  ( ! [X237] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int',X237)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X237),s('type.integer.int',sK1)))) ) | $spl222),
  inference(cnf_transformation,[],[f18336_D])).
fof(f18336_D,plain,(
  ( ! [X237] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int',X237)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X237),s('type.integer.int',sK1)))) ) <=> ~$spl222),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl222])])).
fof(f18391,plain,(
  ( ! [X237,X238] : (p(s(bool,X238)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int',X237)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X237),s('type.integer.int',sK1)))) ) | ($spl5 | $spl22)),
  inference(subsumption_resolution,[],[f17573,f183])).
fof(f17573,plain,(
  ( ! [X237,X238] : (p(s(bool,X238)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int',X237)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X237),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X238)) ) | ($spl5 | $spl22)),
  inference(superposition,[],[f15747,f10081])).
fof(f15747,plain,(
  ( ! [X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',sK1))))) ) | $spl22),
  inference(resolution,[],[f9884,f161])).
fof(f9884,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int',X0))))) ) | $spl22),
  inference(resolution,[],[f1359,f163])).
fof(f18390,plain,(
  ( ! [X222,X221] : (p(s(bool,X222)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X221)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X221),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl5),
  inference(subsumption_resolution,[],[f17562,f183])).
fof(f17562,plain,(
  ( ! [X222,X221] : (p(s(bool,X222)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X221)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X221),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X222)) ) | $spl5),
  inference(superposition,[],[f3027,f10081])).
fof(f18332,plain,(
  ( ! [X203] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X203)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X203),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X203)))))) ) | $spl218),
  inference(cnf_transformation,[],[f18332_D])).
fof(f18332_D,plain,(
  ( ! [X203] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X203)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X203),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X203)))))) ) <=> ~$spl218),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl218])])).
fof(f18389,plain,(
  ( ! [X204,X203] : (p(s(bool,X204)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X203)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X203),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X203)))))) ) | $spl5),
  inference(subsumption_resolution,[],[f17550,f183])).
fof(f17550,plain,(
  ( ! [X204,X203] : (p(s(bool,X204)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X203)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X203),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X203))))) | s(bool,'T_0') = s(bool,X204)) ) | $spl5),
  inference(superposition,[],[f2587,f10081])).
fof(f18330,plain,(
  ( ! [X201] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X201)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X201),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X201)))))) ) | $spl216),
  inference(cnf_transformation,[],[f18330_D])).
fof(f18330_D,plain,(
  ( ! [X201] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X201)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X201),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X201)))))) ) <=> ~$spl216),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl216])])).
fof(f18388,plain,(
  ( ! [X202,X201] : (p(s(bool,X202)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X201)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X201),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X201)))))) ) | $spl5),
  inference(subsumption_resolution,[],[f17549,f183])).
fof(f17549,plain,(
  ( ! [X202,X201] : (p(s(bool,X202)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X201)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X201),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X201))))) | s(bool,'T_0') = s(bool,X202)) ) | $spl5),
  inference(superposition,[],[f3102,f10081])).
fof(f18328,plain,(
  ( ! [X195] : (s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X195))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X195)))))))) ) | $spl214),
  inference(cnf_transformation,[],[f18328_D])).
fof(f18328_D,plain,(
  ( ! [X195] : (s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X195))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X195)))))))) ) <=> ~$spl214),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl214])])).
fof(f18387,plain,(
  ( ! [X196,X195] : (p(s(bool,X196)) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X195))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X195)))))))) ) | $spl5),
  inference(subsumption_resolution,[],[f17546,f183])).
fof(f17546,plain,(
  ( ! [X196,X195] : (p(s(bool,X196)) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X195))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X195))))))) | s(bool,'T_0') = s(bool,X196)) ) | $spl5),
  inference(superposition,[],[f12629,f10081])).
fof(f12629,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))))))) ) | $spl5),
  inference(superposition,[],[f12138,f150])).
fof(f12138,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))))))))))) ) | $spl5),
  inference(superposition,[],[f10445,f140])).
fof(f10445,plain,(
  ( ! [X6] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X6))))))),s('type.integer.int',X6))))) ) | $spl5),
  inference(resolution,[],[f10381,f138])).
fof(f18326,plain,(
  ( ! [X191] : (s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X191))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X191)))))) ) | $spl212),
  inference(cnf_transformation,[],[f18326_D])).
fof(f18326_D,plain,(
  ( ! [X191] : (s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X191))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X191)))))) ) <=> ~$spl212),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl212])])).
fof(f18386,plain,(
  ( ! [X191,X192] : (p(s(bool,X192)) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X191))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X191)))))) ) | $spl5),
  inference(subsumption_resolution,[],[f17544,f183])).
fof(f17544,plain,(
  ( ! [X191,X192] : (p(s(bool,X192)) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X191))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X191))))) | s(bool,'T_0') = s(bool,X192)) ) | $spl5),
  inference(superposition,[],[f9205,f10081])).
fof(f18324,plain,(
  ( ! [X185] : (s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X185),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X185)))))))))) ) | $spl210),
  inference(cnf_transformation,[],[f18324_D])).
fof(f18324_D,plain,(
  ( ! [X185] : (s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X185),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X185)))))))))) ) <=> ~$spl210),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl210])])).
fof(f18385,plain,(
  ( ! [X185,X186] : (p(s(bool,X186)) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X185),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X185)))))))))) ) | $spl5),
  inference(subsumption_resolution,[],[f17541,f183])).
fof(f17541,plain,(
  ( ! [X185,X186] : (p(s(bool,X186)) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X185),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X185))))))))) | s(bool,'T_0') = s(bool,X186)) ) | $spl5),
  inference(superposition,[],[f10589,f10081])).
fof(f10589,plain,(
  ( ! [X5] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X5),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X5))))))))))) ) | $spl5),
  inference(resolution,[],[f10444,f137])).
fof(f10444,plain,(
  ( ! [X5] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X5),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X5))))))))) ) | $spl5),
  inference(resolution,[],[f10381,f137])).
fof(f18383,plain,(
  ( ! [X180,X181,X179] : (p(s(bool,X181)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X179),s('type.integer.int',X180)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X179),s('type.integer.int',X180)))) ) | $spl5),
  inference(subsumption_resolution,[],[f17539,f183])).
fof(f17539,plain,(
  ( ! [X180,X181,X179] : (p(s(bool,X181)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X179),s('type.integer.int',X180)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X179),s('type.integer.int',X180))) | s(bool,'T_0') = s(bool,X181)) ) | $spl5),
  inference(superposition,[],[f163,f10081])).
fof(f18382,plain,(
  ( ! [X175,X173,X174] : (p(s(bool,X175)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X174),s('type.integer.int',X173)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X173),s('type.integer.int',X174)))) ) | $spl5),
  inference(subsumption_resolution,[],[f17537,f183])).
fof(f17537,plain,(
  ( ! [X175,X173,X174] : (p(s(bool,X175)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X174),s('type.integer.int',X173)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X173),s('type.integer.int',X174))) | s(bool,'T_0') = s(bool,X175)) ) | $spl5),
  inference(superposition,[],[f161,f10081])).
fof(f18381,plain,(
  ( ! [X165,X166,X164] : (p(s(bool,X166)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X165),s('type.integer.int',X164)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X164),s('type.integer.int',X165)))) ) | $spl5),
  inference(subsumption_resolution,[],[f17534,f183])).
fof(f17534,plain,(
  ( ! [X165,X166,X164] : (p(s(bool,X166)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X165),s('type.integer.int',X164)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X164),s('type.integer.int',X165))) | s(bool,'T_0') = s(bool,X166)) ) | $spl5),
  inference(superposition,[],[f158,f10081])).
fof(f18380,plain,(
  ( ! [X158,X159,X160] : (p(s(bool,X160)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X158))),s('type.integer.int',X159)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X158),s('type.integer.int',X159)))) ) | $spl5),
  inference(subsumption_resolution,[],[f17532,f183])).
fof(f17532,plain,(
  ( ! [X158,X159,X160] : (p(s(bool,X160)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X158))),s('type.integer.int',X159)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X158),s('type.integer.int',X159))) | s(bool,'T_0') = s(bool,X160)) ) | $spl5),
  inference(superposition,[],[f137,f10081])).
fof(f18379,plain,(
  ( ! [X154,X152,X153] : (p(s(bool,X154)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X153))),s('type.integer.int',X152)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X152))),s('type.integer.int',X153)))) ) | $spl5),
  inference(subsumption_resolution,[],[f17530,f183])).
fof(f17530,plain,(
  ( ! [X154,X152,X153] : (p(s(bool,X154)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X153))),s('type.integer.int',X152)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X152))),s('type.integer.int',X153))) | s(bool,'T_0') = s(bool,X154)) ) | $spl5),
  inference(superposition,[],[f138,f10081])).
fof(f18312,plain,(
  ( ! [X149,X150] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X150))),s('type.integer.int',X149)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X149))),s('type.integer.int',X150))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X149),s('type.integer.int',X150))))) ) | $spl200),
  inference(cnf_transformation,[],[f18312_D])).
fof(f18312_D,plain,(
  ( ! [X149,X150] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X150))),s('type.integer.int',X149)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X149))),s('type.integer.int',X150))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X149),s('type.integer.int',X150))))) ) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f18378,plain,(
  ( ! [X151,X149,X150] : (p(s(bool,X151)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X150))),s('type.integer.int',X149)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X149),s('type.integer.int',X150)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X149))),s('type.integer.int',X150)))) ) | $spl5),
  inference(subsumption_resolution,[],[f17529,f183])).
fof(f17529,plain,(
  ( ! [X151,X149,X150] : (p(s(bool,X151)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X150))),s('type.integer.int',X149)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X149),s('type.integer.int',X150)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X149))),s('type.integer.int',X150))) | s(bool,'T_0') = s(bool,X151)) ) | $spl5),
  inference(superposition,[],[f139,f10081])).
fof(f18310,plain,(
  ( ! [X139] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X139),s('type.integer.int',sK1)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int',X139)))) ) | $spl198),
  inference(cnf_transformation,[],[f18310_D])).
fof(f18310_D,plain,(
  ( ! [X139] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X139),s('type.integer.int',sK1)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int',X139)))) ) <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f18377,plain,(
  ( ! [X140,X139] : (p(s(bool,X140)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X139),s('type.integer.int',sK1)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int',X139)))) ) | ($spl5 | $spl22)),
  inference(subsumption_resolution,[],[f17525,f183])).
fof(f17525,plain,(
  ( ! [X140,X139] : (p(s(bool,X140)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X139),s('type.integer.int',sK1)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int',X139))) | s(bool,'T_0') = s(bool,X140)) ) | ($spl5 | $spl22)),
  inference(superposition,[],[f15747,f10081])).
fof(f18376,plain,(
  ( ! [X136,X135] : (p(s(bool,X136)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X135),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X135)))) ) | $spl5),
  inference(subsumption_resolution,[],[f17523,f183])).
fof(f17523,plain,(
  ( ! [X136,X135] : (p(s(bool,X136)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X135),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X135))) | s(bool,'T_0') = s(bool,X136)) ) | $spl5),
  inference(superposition,[],[f3027,f10081])).
fof(f18306,plain,(
  ( ! [X131] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X131),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X131)))))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X131)))) ) | $spl194),
  inference(cnf_transformation,[],[f18306_D])).
fof(f18306_D,plain,(
  ( ! [X131] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X131),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X131)))))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X131)))) ) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f18375,plain,(
  ( ! [X132,X131] : (p(s(bool,X132)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X131),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X131)))))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X131)))) ) | $spl5),
  inference(subsumption_resolution,[],[f17521,f183])).
fof(f17521,plain,(
  ( ! [X132,X131] : (p(s(bool,X132)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X131),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X131)))))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X131))) | s(bool,'T_0') = s(bool,X132)) ) | $spl5),
  inference(superposition,[],[f3102,f10081])).
fof(f18304,plain,(
  ( ! [X127] : (s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X127))))))),s('type.integer.int',X127)))) ) | $spl192),
  inference(cnf_transformation,[],[f18304_D])).
fof(f18304_D,plain,(
  ( ! [X127] : (s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X127))))))),s('type.integer.int',X127)))) ) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f18374,plain,(
  ( ! [X127,X128] : (p(s(bool,X128)) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X127))))))),s('type.integer.int',X127)))) ) | $spl5),
  inference(subsumption_resolution,[],[f17519,f183])).
fof(f17519,plain,(
  ( ! [X127,X128] : (p(s(bool,X128)) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X127))))))),s('type.integer.int',X127))) | s(bool,'T_0') = s(bool,X128)) ) | $spl5),
  inference(superposition,[],[f10445,f10081])).
fof(f18302,plain,(
  ( ! [X123] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X123),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X123))),s('type.integer.int',X123)))) ) | $spl190),
  inference(cnf_transformation,[],[f18302_D])).
fof(f18302_D,plain,(
  ( ! [X123] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X123),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X123))),s('type.integer.int',X123)))) ) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f18373,plain,(
  ( ! [X123,X124] : (p(s(bool,X124)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X123),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X123))),s('type.integer.int',X123)))) ) | $spl5),
  inference(subsumption_resolution,[],[f17517,f183])).
fof(f17517,plain,(
  ( ! [X123,X124] : (p(s(bool,X124)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X123),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X123))),s('type.integer.int',X123))) | s(bool,'T_0') = s(bool,X124)) ) | $spl5),
  inference(superposition,[],[f2726,f10081])).
fof(f18300,plain,(
  ( ! [X121] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X121),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X121))),s('type.integer.int',X121)))) ) | $spl188),
  inference(cnf_transformation,[],[f18300_D])).
fof(f18300_D,plain,(
  ( ! [X121] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X121),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X121))),s('type.integer.int',X121)))) ) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f18372,plain,(
  ( ! [X121,X122] : (p(s(bool,X122)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X121),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X121))),s('type.integer.int',X121)))) ) | $spl5),
  inference(subsumption_resolution,[],[f17516,f183])).
fof(f17516,plain,(
  ( ! [X121,X122] : (p(s(bool,X122)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X121),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X121))),s('type.integer.int',X121))) | s(bool,'T_0') = s(bool,X122)) ) | $spl5),
  inference(superposition,[],[f3118,f10081])).
fof(f3118,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f3027,f149])).
fof(f18371,plain,(
  ( ! [X70,X71] : (p(s(bool,X71)) | s(bool,X70) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | s(bool,'T_0') = s(bool,X70)) ) | ($spl5 | $spl149)),
  inference(subsumption_resolution,[],[f17502,f183])).
fof(f17502,plain,(
  ( ! [X70,X71] : (p(s(bool,X71)) | s(bool,X70) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | s(bool,'T_0') = s(bool,X70) | s(bool,'T_0') = s(bool,X71)) ) | ($spl5 | $spl149)),
  inference(superposition,[],[f11351,f10081])).
fof(f11351,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,X0) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))) ) | $spl149),
  inference(resolution,[],[f11318,f165])).
fof(f18370,plain,(
  ( ! [X68,X69] : (p(s(bool,X69)) | s(bool,X68) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | s(bool,'T_0') = s(bool,X68)) ) | ($spl5 | $spl22)),
  inference(subsumption_resolution,[],[f17501,f183])).
fof(f17501,plain,(
  ( ! [X68,X69] : (p(s(bool,X69)) | s(bool,X68) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | s(bool,'T_0') = s(bool,X68) | s(bool,'T_0') = s(bool,X69)) ) | ($spl5 | $spl22)),
  inference(superposition,[],[f11294,f10081])).
fof(f18294,plain,(
  ( ! [X67,X65] : (s(bool,X65) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X67))),s('type.integer.int',X67))) | s(bool,'T_0') = s(bool,X65)) ) | $spl182),
  inference(cnf_transformation,[],[f18294_D])).
fof(f18294_D,plain,(
  ( ! [X67,X65] : (s(bool,X65) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X67))),s('type.integer.int',X67))) | s(bool,'T_0') = s(bool,X65)) ) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f18369,plain,(
  ( ! [X66,X67,X65] : (p(s(bool,X66)) | s(bool,X65) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X67))),s('type.integer.int',X67))) | s(bool,'T_0') = s(bool,X65)) ) | $spl5),
  inference(subsumption_resolution,[],[f17500,f183])).
fof(f17500,plain,(
  ( ! [X66,X67,X65] : (p(s(bool,X66)) | s(bool,X65) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X67))),s('type.integer.int',X67))) | s(bool,'T_0') = s(bool,X65) | s(bool,'T_0') = s(bool,X66)) ) | $spl5),
  inference(superposition,[],[f10506,f10081])).
fof(f10506,plain,(
  ( ! [X4,X3] : (p(s(bool,X4)) | s(bool,X4) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X3))),s('type.integer.int',X3)))) ) | $spl5),
  inference(resolution,[],[f10441,f165])).
fof(f18291,plain,(
  ( ! [X60] : (s(bool,X60) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | s(bool,'T_0') = s(bool,X60)) ) | $spl180),
  inference(cnf_transformation,[],[f18291_D])).
fof(f18291_D,plain,(
  ( ! [X60] : (s(bool,X60) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | s(bool,'T_0') = s(bool,X60)) ) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f18367,plain,(
  ( ! [X61,X60] : (p(s(bool,X61)) | s(bool,X60) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | s(bool,'T_0') = s(bool,X60)) ) | ($spl5 | $spl22 | $spl24)),
  inference(subsumption_resolution,[],[f17498,f183])).
fof(f17498,plain,(
  ( ! [X61,X60] : (p(s(bool,X61)) | s(bool,X60) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | s(bool,'T_0') = s(bool,X60) | s(bool,'T_0') = s(bool,X61)) ) | ($spl5 | $spl22 | $spl24)),
  inference(superposition,[],[f5204,f10081])).
fof(f5204,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,X0) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) ) | ($spl22 | $spl24)),
  inference(resolution,[],[f1659,f165])).
fof(f18289,plain,(
  ( ! [X58] : (s(bool,X58) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X58)) ) | $spl178),
  inference(cnf_transformation,[],[f18289_D])).
fof(f18289_D,plain,(
  ( ! [X58] : (s(bool,X58) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X58)) ) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f18366,plain,(
  ( ! [X59,X58] : (p(s(bool,X59)) | s(bool,X58) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X58)) ) | ($spl5 | $spl41)),
  inference(subsumption_resolution,[],[f17497,f183])).
fof(f17497,plain,(
  ( ! [X59,X58] : (p(s(bool,X59)) | s(bool,X58) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X58) | s(bool,'T_0') = s(bool,X59)) ) | ($spl5 | $spl41)),
  inference(superposition,[],[f5128,f10081])).
fof(f5128,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,X0) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl41),
  inference(resolution,[],[f800,f165])).
fof(f18287,plain,(
  ( ! [X50] : (s(bool,X50) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | s(bool,'T_0') = s(bool,X50)) ) | $spl176),
  inference(cnf_transformation,[],[f18287_D])).
fof(f18287_D,plain,(
  ( ! [X50] : (s(bool,X50) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | s(bool,'T_0') = s(bool,X50)) ) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f18365,plain,(
  ( ! [X50,X51] : (p(s(bool,X51)) | s(bool,X50) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | s(bool,'T_0') = s(bool,X50)) ) | ($spl3 | $spl5 | $spl22)),
  inference(subsumption_resolution,[],[f17493,f183])).
fof(f17493,plain,(
  ( ! [X50,X51] : (p(s(bool,X51)) | s(bool,X50) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | s(bool,'T_0') = s(bool,X50) | s(bool,'T_0') = s(bool,X51)) ) | ($spl3 | $spl5 | $spl22)),
  inference(superposition,[],[f1741,f10081])).
fof(f1741,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,X0) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) ) | ($spl3 | $spl22)),
  inference(resolution,[],[f1733,f165])).
fof(f18285,plain,(
  ( ! [X48] : (s(bool,X48) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | s(bool,'T_0') = s(bool,X48)) ) | $spl174),
  inference(cnf_transformation,[],[f18285_D])).
fof(f18285_D,plain,(
  ( ! [X48] : (s(bool,X48) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | s(bool,'T_0') = s(bool,X48)) ) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f18364,plain,(
  ( ! [X48,X49] : (p(s(bool,X49)) | s(bool,X48) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | s(bool,'T_0') = s(bool,X48)) ) | ($spl5 | $spl7)),
  inference(subsumption_resolution,[],[f17492,f183])).
fof(f17492,plain,(
  ( ! [X48,X49] : (p(s(bool,X49)) | s(bool,X48) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | s(bool,'T_0') = s(bool,X48) | s(bool,'T_0') = s(bool,X49)) ) | ($spl5 | $spl7)),
  inference(superposition,[],[f1024,f10081])).
fof(f1024,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,X0) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) ) | $spl7),
  inference(resolution,[],[f248,f165])).
fof(f18283,plain,(
  ( ! [X46] : (s(bool,X46) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | s(bool,'T_0') = s(bool,X46)) ) | $spl172),
  inference(cnf_transformation,[],[f18283_D])).
fof(f18283_D,plain,(
  ( ! [X46] : (s(bool,X46) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | s(bool,'T_0') = s(bool,X46)) ) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f18363,plain,(
  ( ! [X47,X46] : (p(s(bool,X47)) | s(bool,X46) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | s(bool,'T_0') = s(bool,X46)) ) | ($spl3 | $spl5)),
  inference(subsumption_resolution,[],[f17491,f183])).
fof(f17491,plain,(
  ( ! [X47,X46] : (p(s(bool,X47)) | s(bool,X46) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | s(bool,'T_0') = s(bool,X46) | s(bool,'T_0') = s(bool,X47)) ) | ($spl3 | $spl5)),
  inference(superposition,[],[f977,f10081])).
fof(f977,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,X0) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) ) | $spl3),
  inference(resolution,[],[f198,f165])).
fof(f18354,plain,(
  ( ! [X362,X361] : (p(s(bool,X362)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int',X361)))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X361),s('type.integer.int',sK1)))) ) | ($spl5 | $spl22)),
  inference(subsumption_resolution,[],[f17464,f183])).
fof(f17464,plain,(
  ( ! [X362,X361] : (p(s(bool,X362)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int',X361)))) | s(bool,'T_0') = s(bool,X362) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X361),s('type.integer.int',sK1)))) ) | ($spl5 | $spl22)),
  inference(superposition,[],[f9886,f10081])).
fof(f18352,plain,(
  ( ! [X341,X342] : (p(s(bool,X342)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X341)))) | s('type.integer.int',X341) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X341),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl5),
  inference(subsumption_resolution,[],[f17451,f183])).
fof(f17451,plain,(
  ( ! [X341,X342] : (p(s(bool,X342)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X341)))) | s('type.integer.int',X341) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X342) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X341),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl5),
  inference(superposition,[],[f3355,f10081])).
fof(f18350,plain,(
  ( ! [X337,X338] : (p(s(bool,X338)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X337))),s('type.integer.int',X337)))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X337),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl5),
  inference(subsumption_resolution,[],[f17449,f183])).
fof(f17449,plain,(
  ( ! [X337,X338] : (p(s(bool,X338)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X337))),s('type.integer.int',X337)))) | s(bool,'T_0') = s(bool,X338) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X337),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl5),
  inference(superposition,[],[f2726,f10081])).
fof(f18348,plain,(
  ( ! [X306,X305] : (p(s(bool,X306)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X305)))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X305),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X305)))))) ) | $spl5),
  inference(subsumption_resolution,[],[f17430,f183])).
fof(f17430,plain,(
  ( ! [X306,X305] : (p(s(bool,X306)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X305)))) | s(bool,'T_0') = s(bool,X306) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X305),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X305)))))) ) | $spl5),
  inference(superposition,[],[f2756,f10081])).
fof(f18346,plain,(
  ( ! [X290,X289,X291] : (p(s(bool,X291)) | s('type.integer.int',X289) = s('type.integer.int',X290) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X289),s('type.integer.int',X290)))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X289),s('type.integer.int',X290)))) ) | $spl5),
  inference(subsumption_resolution,[],[f17424,f183])).
fof(f17424,plain,(
  ( ! [X290,X289,X291] : (p(s(bool,X291)) | s('type.integer.int',X289) = s('type.integer.int',X290) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X289),s('type.integer.int',X290)))) | s(bool,'T_0') = s(bool,X291) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X289),s('type.integer.int',X290)))) ) | $spl5),
  inference(superposition,[],[f162,f10081])).
fof(f18345,plain,(
  ( ! [X288,X287,X286] : (p(s(bool,X288)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X287),s('type.integer.int',X286)))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X286),s('type.integer.int',X287)))) ) | $spl5),
  inference(subsumption_resolution,[],[f17423,f183])).
fof(f17423,plain,(
  ( ! [X288,X287,X286] : (p(s(bool,X288)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X287),s('type.integer.int',X286)))) | s(bool,'T_0') = s(bool,X288) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X286),s('type.integer.int',X287)))) ) | $spl5),
  inference(superposition,[],[f161,f10081])).
fof(f18343,plain,(
  ( ! [X277,X279,X278] : (p(s(bool,X279)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X278),s('type.integer.int',X277)))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X277),s('type.integer.int',X278)))) ) | $spl5),
  inference(subsumption_resolution,[],[f17420,f183])).
fof(f17420,plain,(
  ( ! [X277,X279,X278] : (p(s(bool,X279)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X278),s('type.integer.int',X277)))) | s(bool,'T_0') = s(bool,X279) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X277),s('type.integer.int',X278)))) ) | $spl5),
  inference(superposition,[],[f158,f10081])).
fof(f18341,plain,(
  ( ! [X273,X274] : (p(s(bool,X274)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X273),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X273) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X273)))) ) | $spl5),
  inference(subsumption_resolution,[],[f17418,f183])).
fof(f17418,plain,(
  ( ! [X273,X274] : (p(s(bool,X274)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X273),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X273) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X274) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X273)))) ) | $spl5),
  inference(superposition,[],[f3355,f10081])).
fof(f18339,plain,(
  ( ! [X271,X272] : (p(s(bool,X272)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X271),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X271)))))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X271)))) ) | $spl5),
  inference(subsumption_resolution,[],[f17417,f183])).
fof(f17417,plain,(
  ( ! [X271,X272] : (p(s(bool,X272)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X271),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X271)))))) | s(bool,'T_0') = s(bool,X272) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X271)))) ) | $spl5),
  inference(superposition,[],[f2587,f10081])).
fof(f18337,plain,(
  ( ! [X265,X266] : (p(s(bool,X266)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X265),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X265))),s('type.integer.int',X265)))) ) | $spl5),
  inference(subsumption_resolution,[],[f17414,f183])).
fof(f17414,plain,(
  ( ! [X265,X266] : (p(s(bool,X266)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X265),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,'T_0') = s(bool,X266) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X265))),s('type.integer.int',X265)))) ) | $spl5),
  inference(superposition,[],[f2621,f10081])).
fof(f18335,plain,(
  ( ! [X237,X238] : (p(s(bool,X238)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int',X237)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X237),s('type.integer.int',sK1)))) ) | ($spl5 | $spl22)),
  inference(subsumption_resolution,[],[f17398,f183])).
fof(f17398,plain,(
  ( ! [X237,X238] : (p(s(bool,X238)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int',X237)))) | s(bool,'T_0') = s(bool,X238) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X237),s('type.integer.int',sK1)))) ) | ($spl5 | $spl22)),
  inference(superposition,[],[f15747,f10081])).
fof(f18333,plain,(
  ( ! [X222,X221] : (p(s(bool,X222)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X221)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X221),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl5),
  inference(subsumption_resolution,[],[f17387,f183])).
fof(f17387,plain,(
  ( ! [X222,X221] : (p(s(bool,X222)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X221)))) | s(bool,'T_0') = s(bool,X222) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X221),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl5),
  inference(superposition,[],[f3027,f10081])).
fof(f18331,plain,(
  ( ! [X204,X203] : (p(s(bool,X204)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X203)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X203),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X203)))))) ) | $spl5),
  inference(subsumption_resolution,[],[f17375,f183])).
fof(f17375,plain,(
  ( ! [X204,X203] : (p(s(bool,X204)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X203)))) | s(bool,'T_0') = s(bool,X204) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X203),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X203)))))) ) | $spl5),
  inference(superposition,[],[f2587,f10081])).
fof(f18329,plain,(
  ( ! [X202,X201] : (p(s(bool,X202)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X201)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X201),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X201)))))) ) | $spl5),
  inference(subsumption_resolution,[],[f17374,f183])).
fof(f17374,plain,(
  ( ! [X202,X201] : (p(s(bool,X202)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X201)))) | s(bool,'T_0') = s(bool,X202) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X201),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X201)))))) ) | $spl5),
  inference(superposition,[],[f3102,f10081])).
fof(f18327,plain,(
  ( ! [X196,X195] : (p(s(bool,X196)) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X195))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X195)))))))) ) | $spl5),
  inference(subsumption_resolution,[],[f17371,f183])).
fof(f17371,plain,(
  ( ! [X196,X195] : (p(s(bool,X196)) | s(bool,'T_0') = s(bool,X196) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X195))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X195)))))))) ) | $spl5),
  inference(superposition,[],[f12629,f10081])).
fof(f18325,plain,(
  ( ! [X191,X192] : (p(s(bool,X192)) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X191))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X191)))))) ) | $spl5),
  inference(subsumption_resolution,[],[f17369,f183])).
fof(f17369,plain,(
  ( ! [X191,X192] : (p(s(bool,X192)) | s(bool,'T_0') = s(bool,X192) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X191))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X191)))))) ) | $spl5),
  inference(superposition,[],[f9205,f10081])).
fof(f18323,plain,(
  ( ! [X185,X186] : (p(s(bool,X186)) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X185),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X185)))))))))) ) | $spl5),
  inference(subsumption_resolution,[],[f17366,f183])).
fof(f17366,plain,(
  ( ! [X185,X186] : (p(s(bool,X186)) | s(bool,'T_0') = s(bool,X186) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X185),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X185)))))))))) ) | $spl5),
  inference(superposition,[],[f10589,f10081])).
fof(f18320,plain,(
  ( ! [X180,X181,X179] : (p(s(bool,X181)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X179),s('type.integer.int',X180)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X179),s('type.integer.int',X180)))) ) | $spl5),
  inference(subsumption_resolution,[],[f17364,f183])).
fof(f17364,plain,(
  ( ! [X180,X181,X179] : (p(s(bool,X181)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X179),s('type.integer.int',X180)))) | s(bool,'T_0') = s(bool,X181) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X179),s('type.integer.int',X180)))) ) | $spl5),
  inference(superposition,[],[f163,f10081])).
fof(f18319,plain,(
  ( ! [X175,X173,X174] : (p(s(bool,X175)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X174),s('type.integer.int',X173)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X173),s('type.integer.int',X174)))) ) | $spl5),
  inference(subsumption_resolution,[],[f17362,f183])).
fof(f17362,plain,(
  ( ! [X175,X173,X174] : (p(s(bool,X175)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X174),s('type.integer.int',X173)))) | s(bool,'T_0') = s(bool,X175) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X173),s('type.integer.int',X174)))) ) | $spl5),
  inference(superposition,[],[f161,f10081])).
fof(f18317,plain,(
  ( ! [X165,X166,X164] : (p(s(bool,X166)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X165),s('type.integer.int',X164)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X164),s('type.integer.int',X165)))) ) | $spl5),
  inference(subsumption_resolution,[],[f17359,f183])).
fof(f17359,plain,(
  ( ! [X165,X166,X164] : (p(s(bool,X166)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X165),s('type.integer.int',X164)))) | s(bool,'T_0') = s(bool,X166) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X164),s('type.integer.int',X165)))) ) | $spl5),
  inference(superposition,[],[f158,f10081])).
fof(f18315,plain,(
  ( ! [X158,X159,X160] : (p(s(bool,X160)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X158))),s('type.integer.int',X159)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X158),s('type.integer.int',X159)))) ) | $spl5),
  inference(subsumption_resolution,[],[f17357,f183])).
fof(f17357,plain,(
  ( ! [X158,X159,X160] : (p(s(bool,X160)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X158))),s('type.integer.int',X159)))) | s(bool,'T_0') = s(bool,X160) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X158),s('type.integer.int',X159)))) ) | $spl5),
  inference(superposition,[],[f137,f10081])).
fof(f18313,plain,(
  ( ! [X154,X152,X153] : (p(s(bool,X154)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X153))),s('type.integer.int',X152)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X152))),s('type.integer.int',X153)))) ) | $spl5),
  inference(subsumption_resolution,[],[f17355,f183])).
fof(f17355,plain,(
  ( ! [X154,X152,X153] : (p(s(bool,X154)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X153))),s('type.integer.int',X152)))) | s(bool,'T_0') = s(bool,X154) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X152))),s('type.integer.int',X153)))) ) | $spl5),
  inference(superposition,[],[f138,f10081])).
fof(f18311,plain,(
  ( ! [X151,X149,X150] : (p(s(bool,X151)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X150))),s('type.integer.int',X149)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X149),s('type.integer.int',X150)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X149))),s('type.integer.int',X150)))) ) | $spl5),
  inference(subsumption_resolution,[],[f17354,f183])).
fof(f17354,plain,(
  ( ! [X151,X149,X150] : (p(s(bool,X151)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X150))),s('type.integer.int',X149)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X149),s('type.integer.int',X150)))) | s(bool,'T_0') = s(bool,X151) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X149))),s('type.integer.int',X150)))) ) | $spl5),
  inference(superposition,[],[f139,f10081])).
fof(f18309,plain,(
  ( ! [X140,X139] : (p(s(bool,X140)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X139),s('type.integer.int',sK1)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int',X139)))) ) | ($spl5 | $spl22)),
  inference(subsumption_resolution,[],[f17350,f183])).
fof(f17350,plain,(
  ( ! [X140,X139] : (p(s(bool,X140)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X139),s('type.integer.int',sK1)))) | s(bool,'T_0') = s(bool,X140) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int',X139)))) ) | ($spl5 | $spl22)),
  inference(superposition,[],[f15747,f10081])).
fof(f18307,plain,(
  ( ! [X136,X135] : (p(s(bool,X136)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X135),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X135)))) ) | $spl5),
  inference(subsumption_resolution,[],[f17348,f183])).
fof(f17348,plain,(
  ( ! [X136,X135] : (p(s(bool,X136)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X135),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,'T_0') = s(bool,X136) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X135)))) ) | $spl5),
  inference(superposition,[],[f3027,f10081])).
fof(f18305,plain,(
  ( ! [X132,X131] : (p(s(bool,X132)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X131),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X131)))))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X131)))) ) | $spl5),
  inference(subsumption_resolution,[],[f17346,f183])).
fof(f17346,plain,(
  ( ! [X132,X131] : (p(s(bool,X132)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X131),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X131)))))) | s(bool,'T_0') = s(bool,X132) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X131)))) ) | $spl5),
  inference(superposition,[],[f3102,f10081])).
fof(f18303,plain,(
  ( ! [X127,X128] : (p(s(bool,X128)) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X127))))))),s('type.integer.int',X127)))) ) | $spl5),
  inference(subsumption_resolution,[],[f17344,f183])).
fof(f17344,plain,(
  ( ! [X127,X128] : (p(s(bool,X128)) | s(bool,'T_0') = s(bool,X128) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X127))))))),s('type.integer.int',X127)))) ) | $spl5),
  inference(superposition,[],[f10445,f10081])).
fof(f18301,plain,(
  ( ! [X123,X124] : (p(s(bool,X124)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X123),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X123))),s('type.integer.int',X123)))) ) | $spl5),
  inference(subsumption_resolution,[],[f17342,f183])).
fof(f17342,plain,(
  ( ! [X123,X124] : (p(s(bool,X124)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X123),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,'T_0') = s(bool,X124) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X123))),s('type.integer.int',X123)))) ) | $spl5),
  inference(superposition,[],[f2726,f10081])).
fof(f18299,plain,(
  ( ! [X121,X122] : (p(s(bool,X122)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X121),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X121))),s('type.integer.int',X121)))) ) | $spl5),
  inference(subsumption_resolution,[],[f17341,f183])).
fof(f17341,plain,(
  ( ! [X121,X122] : (p(s(bool,X122)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X121),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,'T_0') = s(bool,X122) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X121))),s('type.integer.int',X121)))) ) | $spl5),
  inference(superposition,[],[f3118,f10081])).
fof(f18297,plain,(
  ( ! [X70,X71] : (p(s(bool,X71)) | s(bool,X70) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | s(bool,'T_0') = s(bool,X70)) ) | ($spl5 | $spl149)),
  inference(subsumption_resolution,[],[f17327,f183])).
fof(f17327,plain,(
  ( ! [X70,X71] : (p(s(bool,X71)) | s(bool,X70) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | s(bool,'T_0') = s(bool,X71) | s(bool,'T_0') = s(bool,X70)) ) | ($spl5 | $spl149)),
  inference(superposition,[],[f11351,f10081])).
fof(f18295,plain,(
  ( ! [X68,X69] : (p(s(bool,X69)) | s(bool,X68) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | s(bool,'T_0') = s(bool,X68)) ) | ($spl5 | $spl22)),
  inference(subsumption_resolution,[],[f17326,f183])).
fof(f17326,plain,(
  ( ! [X68,X69] : (p(s(bool,X69)) | s(bool,X68) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | s(bool,'T_0') = s(bool,X69) | s(bool,'T_0') = s(bool,X68)) ) | ($spl5 | $spl22)),
  inference(superposition,[],[f11294,f10081])).
fof(f18293,plain,(
  ( ! [X66,X67,X65] : (p(s(bool,X66)) | s(bool,X65) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X67))),s('type.integer.int',X67))) | s(bool,'T_0') = s(bool,X65)) ) | $spl5),
  inference(subsumption_resolution,[],[f17325,f183])).
fof(f17325,plain,(
  ( ! [X66,X67,X65] : (p(s(bool,X66)) | s(bool,X65) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X67))),s('type.integer.int',X67))) | s(bool,'T_0') = s(bool,X66) | s(bool,'T_0') = s(bool,X65)) ) | $spl5),
  inference(superposition,[],[f10506,f10081])).
fof(f18290,plain,(
  ( ! [X61,X60] : (p(s(bool,X61)) | s(bool,X60) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | s(bool,'T_0') = s(bool,X60)) ) | ($spl5 | $spl22 | $spl24)),
  inference(subsumption_resolution,[],[f17323,f183])).
fof(f17323,plain,(
  ( ! [X61,X60] : (p(s(bool,X61)) | s(bool,X60) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | s(bool,'T_0') = s(bool,X61) | s(bool,'T_0') = s(bool,X60)) ) | ($spl5 | $spl22 | $spl24)),
  inference(superposition,[],[f5204,f10081])).
fof(f18288,plain,(
  ( ! [X59,X58] : (p(s(bool,X59)) | s(bool,X58) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X58)) ) | ($spl5 | $spl41)),
  inference(subsumption_resolution,[],[f17322,f183])).
fof(f17322,plain,(
  ( ! [X59,X58] : (p(s(bool,X59)) | s(bool,X58) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X59) | s(bool,'T_0') = s(bool,X58)) ) | ($spl5 | $spl41)),
  inference(superposition,[],[f5128,f10081])).
fof(f18286,plain,(
  ( ! [X50,X51] : (p(s(bool,X51)) | s(bool,X50) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | s(bool,'T_0') = s(bool,X50)) ) | ($spl3 | $spl5 | $spl22)),
  inference(subsumption_resolution,[],[f17318,f183])).
fof(f17318,plain,(
  ( ! [X50,X51] : (p(s(bool,X51)) | s(bool,X50) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | s(bool,'T_0') = s(bool,X51) | s(bool,'T_0') = s(bool,X50)) ) | ($spl3 | $spl5 | $spl22)),
  inference(superposition,[],[f1741,f10081])).
fof(f18284,plain,(
  ( ! [X48,X49] : (p(s(bool,X49)) | s(bool,X48) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | s(bool,'T_0') = s(bool,X48)) ) | ($spl5 | $spl7)),
  inference(subsumption_resolution,[],[f17317,f183])).
fof(f17317,plain,(
  ( ! [X48,X49] : (p(s(bool,X49)) | s(bool,X48) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | s(bool,'T_0') = s(bool,X49) | s(bool,'T_0') = s(bool,X48)) ) | ($spl5 | $spl7)),
  inference(superposition,[],[f1024,f10081])).
fof(f18282,plain,(
  ( ! [X47,X46] : (p(s(bool,X47)) | s(bool,X46) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | s(bool,'T_0') = s(bool,X46)) ) | ($spl3 | $spl5)),
  inference(subsumption_resolution,[],[f17316,f183])).
fof(f17316,plain,(
  ( ! [X47,X46] : (p(s(bool,X47)) | s(bool,X46) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | s(bool,'T_0') = s(bool,X47) | s(bool,'T_0') = s(bool,X46)) ) | ($spl3 | $spl5)),
  inference(superposition,[],[f977,f10081])).
fof(f17086,plain,(
  ( ! [X44] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | p(s(bool,X44))) ) | ($spl0 | $spl5 | $spl22 | $spl149)),
  inference(forward_demodulation,[],[f17085,f1891])).
fof(f17085,plain,(
  ( ! [X44] : (p(s(bool,X44)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK1))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))) ) | ($spl0 | $spl22 | $spl149)),
  inference(subsumption_resolution,[],[f17053,f183])).
fof(f17053,plain,(
  ( ! [X44] : (p(s(bool,X44)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK1))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | s(bool,'T_0') = s(bool,X44)) ) | ($spl0 | $spl22 | $spl149)),
  inference(superposition,[],[f11351,f4356])).
fof(f17082,plain,(
  ( ! [X43] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | p(s(bool,X43))) ) | ($spl5 | $spl149)),
  inference(forward_demodulation,[],[f17081,f1890])).
fof(f17081,plain,(
  ( ! [X43] : (p(s(bool,X43)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))) ) | ($spl5 | $spl149)),
  inference(subsumption_resolution,[],[f17047,f183])).
fof(f17047,plain,(
  ( ! [X43] : (p(s(bool,X43)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | s(bool,'T_0') = s(bool,X43)) ) | ($spl5 | $spl149)),
  inference(superposition,[],[f11351,f4273])).
fof(f17075,plain,(
  ( ! [X37] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | p(s(bool,X37))) ) | ($spl5 | $spl149)),
  inference(forward_demodulation,[],[f17074,f10408])).
fof(f17074,plain,(
  ( ! [X37,X36] : (p(s(bool,X37)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',X36),s('type.integer.int',X36))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))) ) | ($spl5 | $spl149)),
  inference(subsumption_resolution,[],[f17041,f183])).
fof(f17041,plain,(
  ( ! [X37,X36] : (p(s(bool,X37)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',X36),s('type.integer.int',X36))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | s(bool,'T_0') = s(bool,X37)) ) | ($spl5 | $spl149)),
  inference(superposition,[],[f11351,f11923])).
fof(f17071,plain,(
  ( ! [X35] : (p(s(bool,X35)) | s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))) ) | ($spl5 | $spl149)),
  inference(subsumption_resolution,[],[f17035,f183])).
fof(f17035,plain,(
  ( ! [X35] : (p(s(bool,X35)) | s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | s(bool,'T_0') = s(bool,X35)) ) | ($spl5 | $spl149)),
  inference(superposition,[],[f11351,f1892])).
fof(f16948,plain,(
  ( ! [X44] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | p(s(bool,X44))) ) | ($spl0 | $spl5 | $spl22)),
  inference(forward_demodulation,[],[f16947,f1891])).
fof(f16947,plain,(
  ( ! [X44] : (p(s(bool,X44)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK1))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))) ) | ($spl0 | $spl5 | $spl22)),
  inference(subsumption_resolution,[],[f16915,f183])).
fof(f16915,plain,(
  ( ! [X44] : (p(s(bool,X44)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK1))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | s(bool,'T_0') = s(bool,X44)) ) | ($spl0 | $spl5 | $spl22)),
  inference(superposition,[],[f11294,f4356])).
fof(f16944,plain,(
  ( ! [X43] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | p(s(bool,X43))) ) | ($spl5 | $spl22)),
  inference(forward_demodulation,[],[f16943,f1890])).
fof(f16943,plain,(
  ( ! [X43] : (p(s(bool,X43)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))) ) | ($spl5 | $spl22)),
  inference(subsumption_resolution,[],[f16909,f183])).
fof(f16909,plain,(
  ( ! [X43] : (p(s(bool,X43)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | s(bool,'T_0') = s(bool,X43)) ) | ($spl5 | $spl22)),
  inference(superposition,[],[f11294,f4273])).
fof(f16937,plain,(
  ( ! [X37] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | p(s(bool,X37))) ) | ($spl5 | $spl22)),
  inference(forward_demodulation,[],[f16936,f10408])).
fof(f16936,plain,(
  ( ! [X37,X36] : (p(s(bool,X37)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',X36),s('type.integer.int',X36))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))) ) | ($spl5 | $spl22)),
  inference(subsumption_resolution,[],[f16903,f183])).
fof(f16903,plain,(
  ( ! [X37,X36] : (p(s(bool,X37)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',X36),s('type.integer.int',X36))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | s(bool,'T_0') = s(bool,X37)) ) | ($spl5 | $spl22)),
  inference(superposition,[],[f11294,f11923])).
fof(f16933,plain,(
  ( ! [X35] : (p(s(bool,X35)) | s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))) ) | ($spl5 | $spl22)),
  inference(subsumption_resolution,[],[f16897,f183])).
fof(f16897,plain,(
  ( ! [X35] : (p(s(bool,X35)) | s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | s(bool,'T_0') = s(bool,X35)) ) | ($spl5 | $spl22)),
  inference(superposition,[],[f11294,f1892])).
fof(f16797,plain,(
  ( ! [X69] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X69))),s('type.integer.int',X69)))) ) | $spl166),
  inference(cnf_transformation,[],[f16797_D])).
fof(f16797_D,plain,(
  ( ! [X69] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X69))),s('type.integer.int',X69)))) ) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f16810,plain,(
  ( ! [X95,X96] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X96))),s('type.integer.int',X96))) | p(s(bool,X95))) ) | ($spl0 | $spl5 | $spl22)),
  inference(forward_demodulation,[],[f16809,f1891])).
fof(f16809,plain,(
  ( ! [X95,X96] : (p(s(bool,X95)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK1))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X96))),s('type.integer.int',X96)))) ) | ($spl0 | $spl5 | $spl22)),
  inference(subsumption_resolution,[],[f16778,f183])).
fof(f16778,plain,(
  ( ! [X95,X96] : (p(s(bool,X95)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK1))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X96))),s('type.integer.int',X96))) | s(bool,'T_0') = s(bool,X95)) ) | ($spl0 | $spl5 | $spl22)),
  inference(superposition,[],[f10506,f4356])).
fof(f16806,plain,(
  ( ! [X88,X89] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X89))),s('type.integer.int',X89))) | p(s(bool,X88))) ) | $spl5),
  inference(forward_demodulation,[],[f16805,f1890])).
fof(f16805,plain,(
  ( ! [X88,X89] : (p(s(bool,X88)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X89))),s('type.integer.int',X89)))) ) | $spl5),
  inference(subsumption_resolution,[],[f16772,f183])).
fof(f16772,plain,(
  ( ! [X88,X89] : (p(s(bool,X88)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X89))),s('type.integer.int',X89))) | s(bool,'T_0') = s(bool,X88)) ) | $spl5),
  inference(superposition,[],[f10506,f4273])).
fof(f16799,plain,(
  ( ! [X76,X77] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X77))),s('type.integer.int',X77))) | p(s(bool,X76))) ) | $spl5),
  inference(forward_demodulation,[],[f16798,f10408])).
fof(f16798,plain,(
  ( ! [X76,X77,X75] : (p(s(bool,X76)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',X75),s('type.integer.int',X75))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X77))),s('type.integer.int',X77)))) ) | $spl5),
  inference(subsumption_resolution,[],[f16766,f183])).
fof(f16766,plain,(
  ( ! [X76,X77,X75] : (p(s(bool,X76)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',X75),s('type.integer.int',X75))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X77))),s('type.integer.int',X77))) | s(bool,'T_0') = s(bool,X76)) ) | $spl5),
  inference(superposition,[],[f10506,f11923])).
fof(f16796,plain,(
  ( ! [X68,X69] : (p(s(bool,X68)) | s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X69))),s('type.integer.int',X69)))) ) | $spl5),
  inference(subsumption_resolution,[],[f16760,f183])).
fof(f16760,plain,(
  ( ! [X68,X69] : (p(s(bool,X68)) | s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X69))),s('type.integer.int',X69))) | s(bool,'T_0') = s(bool,X68)) ) | $spl5),
  inference(superposition,[],[f10506,f1892])).
fof(f15726,plain,(
  s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | $spl164),
  inference(cnf_transformation,[],[f15726_D])).
fof(f15726_D,plain,(
  s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f16610,plain,(
  ( ! [X99] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | p(s(bool,X99))) ) | ($spl5 | $spl22 | $spl24)),
  inference(forward_demodulation,[],[f16609,f10408])).
fof(f16609,plain,(
  ( ! [X99,X98] : (p(s(bool,X99)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',X98),s('type.integer.int',X98))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) ) | ($spl5 | $spl22 | $spl24)),
  inference(subsumption_resolution,[],[f16381,f183])).
fof(f16381,plain,(
  ( ! [X99,X98] : (p(s(bool,X99)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',X98),s('type.integer.int',X98))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | s(bool,'T_0') = s(bool,X99)) ) | ($spl5 | $spl22 | $spl24)),
  inference(superposition,[],[f5204,f11923])).
fof(f15598,plain,(
  s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl162),
  inference(cnf_transformation,[],[f15598_D])).
fof(f15598_D,plain,(
  s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f16608,plain,(
  ( ! [X97] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,X97))) ) | ($spl5 | $spl41)),
  inference(forward_demodulation,[],[f16607,f10408])).
fof(f16607,plain,(
  ( ! [X97,X96] : (p(s(bool,X97)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',X96),s('type.integer.int',X96))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl5 | $spl41)),
  inference(subsumption_resolution,[],[f16380,f183])).
fof(f16380,plain,(
  ( ! [X97,X96] : (p(s(bool,X97)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',X96),s('type.integer.int',X96))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X97)) ) | ($spl5 | $spl41)),
  inference(superposition,[],[f5128,f11923])).
fof(f10312,plain,(
  s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | $spl142),
  inference(cnf_transformation,[],[f10312_D])).
fof(f10312_D,plain,(
  s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f16606,plain,(
  ( ! [X89] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | p(s(bool,X89))) ) | ($spl3 | $spl5 | $spl22)),
  inference(forward_demodulation,[],[f16605,f10408])).
fof(f16605,plain,(
  ( ! [X88,X89] : (p(s(bool,X89)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',X88),s('type.integer.int',X88))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) ) | ($spl3 | $spl5 | $spl22)),
  inference(subsumption_resolution,[],[f16376,f183])).
fof(f16376,plain,(
  ( ! [X88,X89] : (p(s(bool,X89)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',X88),s('type.integer.int',X88))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | s(bool,'T_0') = s(bool,X89)) ) | ($spl3 | $spl5 | $spl22)),
  inference(superposition,[],[f1741,f11923])).
fof(f10309,plain,(
  s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | $spl140),
  inference(cnf_transformation,[],[f10309_D])).
fof(f10309_D,plain,(
  s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f16604,plain,(
  ( ! [X87] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | p(s(bool,X87))) ) | ($spl5 | $spl7)),
  inference(forward_demodulation,[],[f16603,f10408])).
fof(f16603,plain,(
  ( ! [X87,X86] : (p(s(bool,X87)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',X86),s('type.integer.int',X86))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) ) | ($spl5 | $spl7)),
  inference(subsumption_resolution,[],[f16375,f183])).
fof(f16375,plain,(
  ( ! [X87,X86] : (p(s(bool,X87)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',X86),s('type.integer.int',X86))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | s(bool,'T_0') = s(bool,X87)) ) | ($spl5 | $spl7)),
  inference(superposition,[],[f1024,f11923])).
fof(f10306,plain,(
  s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | $spl138),
  inference(cnf_transformation,[],[f10306_D])).
fof(f10306_D,plain,(
  s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f16602,plain,(
  ( ! [X85] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | p(s(bool,X85))) ) | ($spl3 | $spl5)),
  inference(forward_demodulation,[],[f16601,f10408])).
fof(f16601,plain,(
  ( ! [X85,X84] : (p(s(bool,X85)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',X84),s('type.integer.int',X84))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) ) | ($spl3 | $spl5)),
  inference(subsumption_resolution,[],[f16374,f183])).
fof(f16374,plain,(
  ( ! [X85,X84] : (p(s(bool,X85)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',X84),s('type.integer.int',X84))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | s(bool,'T_0') = s(bool,X85)) ) | ($spl3 | $spl5)),
  inference(superposition,[],[f977,f11923])).
fof(f15737,plain,(
  ( ! [X40] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | p(s(bool,X40))) ) | ($spl0 | $spl5 | $spl22 | $spl24)),
  inference(forward_demodulation,[],[f15736,f1891])).
fof(f15736,plain,(
  ( ! [X40] : (p(s(bool,X40)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK1))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) ) | ($spl0 | $spl22 | $spl24)),
  inference(subsumption_resolution,[],[f15708,f183])).
fof(f15708,plain,(
  ( ! [X40] : (p(s(bool,X40)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK1))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | s(bool,'T_0') = s(bool,X40)) ) | ($spl0 | $spl22 | $spl24)),
  inference(superposition,[],[f5204,f4356])).
fof(f15733,plain,(
  ( ! [X39] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | p(s(bool,X39))) ) | ($spl5 | $spl22 | $spl24)),
  inference(forward_demodulation,[],[f15732,f1890])).
fof(f15732,plain,(
  ( ! [X39] : (p(s(bool,X39)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) ) | ($spl5 | $spl22 | $spl24)),
  inference(subsumption_resolution,[],[f15702,f183])).
fof(f15702,plain,(
  ( ! [X39] : (p(s(bool,X39)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | s(bool,'T_0') = s(bool,X39)) ) | ($spl5 | $spl22 | $spl24)),
  inference(superposition,[],[f5204,f4273])).
fof(f15724,plain,(
  ( ! [X33] : (p(s(bool,X33)) | s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) ) | ($spl5 | $spl22 | $spl24)),
  inference(subsumption_resolution,[],[f15691,f183])).
fof(f15691,plain,(
  ( ! [X33] : (p(s(bool,X33)) | s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | s(bool,'T_0') = s(bool,X33)) ) | ($spl5 | $spl22 | $spl24)),
  inference(superposition,[],[f5204,f1892])).
fof(f15609,plain,(
  ( ! [X40] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,X40))) ) | ($spl0 | $spl5 | $spl22 | $spl41)),
  inference(forward_demodulation,[],[f15608,f1891])).
fof(f15608,plain,(
  ( ! [X40] : (p(s(bool,X40)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK1))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl0 | $spl22 | $spl41)),
  inference(subsumption_resolution,[],[f15580,f183])).
fof(f15580,plain,(
  ( ! [X40] : (p(s(bool,X40)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK1))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X40)) ) | ($spl0 | $spl22 | $spl41)),
  inference(superposition,[],[f5128,f4356])).
fof(f15605,plain,(
  ( ! [X39] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,X39))) ) | ($spl5 | $spl41)),
  inference(forward_demodulation,[],[f15604,f1890])).
fof(f15604,plain,(
  ( ! [X39] : (p(s(bool,X39)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl5 | $spl41)),
  inference(subsumption_resolution,[],[f15574,f183])).
fof(f15574,plain,(
  ( ! [X39] : (p(s(bool,X39)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X39)) ) | ($spl5 | $spl41)),
  inference(superposition,[],[f5128,f4273])).
fof(f15596,plain,(
  ( ! [X33] : (p(s(bool,X33)) | s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl5 | $spl41)),
  inference(subsumption_resolution,[],[f15563,f183])).
fof(f15563,plain,(
  ( ! [X33] : (p(s(bool,X33)) | s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X33)) ) | ($spl5 | $spl41)),
  inference(superposition,[],[f5128,f1892])).
fof(f15482,plain,(
  ( ! [X41] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | p(s(bool,X41))) ) | ($spl0 | $spl3 | $spl5 | $spl22)),
  inference(forward_demodulation,[],[f15481,f1891])).
fof(f15481,plain,(
  ( ! [X41] : (p(s(bool,X41)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK1))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) ) | ($spl0 | $spl3 | $spl22)),
  inference(subsumption_resolution,[],[f15280,f183])).
fof(f15280,plain,(
  ( ! [X41] : (p(s(bool,X41)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK1))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | s(bool,'T_0') = s(bool,X41)) ) | ($spl0 | $spl3 | $spl22)),
  inference(superposition,[],[f1741,f4356])).
fof(f15480,plain,(
  ( ! [X40] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | p(s(bool,X40))) ) | ($spl0 | $spl5 | $spl7 | $spl22)),
  inference(forward_demodulation,[],[f15479,f1891])).
fof(f15479,plain,(
  ( ! [X40] : (p(s(bool,X40)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK1))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) ) | ($spl0 | $spl7 | $spl22)),
  inference(subsumption_resolution,[],[f15279,f183])).
fof(f15279,plain,(
  ( ! [X40] : (p(s(bool,X40)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK1))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | s(bool,'T_0') = s(bool,X40)) ) | ($spl0 | $spl7 | $spl22)),
  inference(superposition,[],[f1024,f4356])).
fof(f15478,plain,(
  ( ! [X39] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | p(s(bool,X39))) ) | ($spl0 | $spl3 | $spl5 | $spl22)),
  inference(forward_demodulation,[],[f15477,f1891])).
fof(f15477,plain,(
  ( ! [X39] : (p(s(bool,X39)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK1))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) ) | ($spl0 | $spl3 | $spl22)),
  inference(subsumption_resolution,[],[f15278,f183])).
fof(f15278,plain,(
  ( ! [X39] : (p(s(bool,X39)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK1))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | s(bool,'T_0') = s(bool,X39)) ) | ($spl0 | $spl3 | $spl22)),
  inference(superposition,[],[f977,f4356])).
fof(f15048,plain,(
  ( ! [X39] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | p(s(bool,X39))) ) | ($spl3 | $spl5 | $spl22)),
  inference(forward_demodulation,[],[f15047,f1890])).
fof(f15047,plain,(
  ( ! [X39] : (p(s(bool,X39)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) ) | ($spl3 | $spl5 | $spl22)),
  inference(subsumption_resolution,[],[f14853,f183])).
fof(f14853,plain,(
  ( ! [X39] : (p(s(bool,X39)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | s(bool,'T_0') = s(bool,X39)) ) | ($spl3 | $spl5 | $spl22)),
  inference(superposition,[],[f1741,f4273])).
fof(f15046,plain,(
  ( ! [X38] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | p(s(bool,X38))) ) | ($spl5 | $spl7)),
  inference(forward_demodulation,[],[f15045,f1890])).
fof(f15045,plain,(
  ( ! [X38] : (p(s(bool,X38)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) ) | ($spl5 | $spl7)),
  inference(subsumption_resolution,[],[f14852,f183])).
fof(f14852,plain,(
  ( ! [X38] : (p(s(bool,X38)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | s(bool,'T_0') = s(bool,X38)) ) | ($spl5 | $spl7)),
  inference(superposition,[],[f1024,f4273])).
fof(f15044,plain,(
  ( ! [X37] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | p(s(bool,X37))) ) | ($spl3 | $spl5)),
  inference(forward_demodulation,[],[f15043,f1890])).
fof(f15043,plain,(
  ( ! [X37] : (p(s(bool,X37)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) ) | ($spl3 | $spl5)),
  inference(subsumption_resolution,[],[f14851,f183])).
fof(f14851,plain,(
  ( ! [X37] : (p(s(bool,X37)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | s(bool,'T_0') = s(bool,X37)) ) | ($spl3 | $spl5)),
  inference(superposition,[],[f977,f4273])).
fof(f14004,plain,(
  $false | ($spl22 | $spl24 | $spl152 | $spl155)),
  inference(subsumption_resolution,[],[f14003,f189])).
fof(f14003,plain,(
  ~p(s(bool,'T_0')) | ($spl22 | $spl24 | $spl152 | $spl155)),
  inference(forward_demodulation,[],[f14002,f1662])).
fof(f14002,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int',sK1)))) | ($spl152 | $spl155)),
  inference(forward_demodulation,[],[f11521,f11519])).
fof(f13471,plain,(
  $false | ($spl0 | $spl5 | $spl22 | $spl150 | $spl152)),
  inference(subsumption_resolution,[],[f13470,f200])).
fof(f13470,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0)))) | ($spl0 | $spl5 | $spl22 | $spl150 | $spl152)),
  inference(forward_demodulation,[],[f13296,f1891])).
fof(f13296,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK1)))) | ($spl150 | $spl152)),
  inference(backward_demodulation,[],[f11519,f11517])).
fof(f13290,plain,(
  $false | ($spl5 | $spl22 | $spl159)),
  inference(subsumption_resolution,[],[f13256,f11628])).
fof(f13256,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))))) | ($spl5 | $spl22)),
  inference(superposition,[],[f12803,f655])).
fof(f11905,plain,(
  $false | $spl102),
  inference(subsumption_resolution,[],[f11845,f4750])).
fof(f11845,plain,(
  ( ! [X34] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X34),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X34))))))) ) | $spl102),
  inference(resolution,[],[f4750,f2588])).
fof(f11904,plain,(
  $false | $spl102),
  inference(subsumption_resolution,[],[f11842,f4750])).
fof(f11842,plain,(
  ( ! [X32] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X32),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl102),
  inference(resolution,[],[f4750,f2622])).
fof(f11903,plain,(
  $false | $spl102),
  inference(subsumption_resolution,[],[f11833,f4750])).
fof(f11833,plain,(
  ( ! [X29] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X29))),s('type.integer.int',X29))))) ) | $spl102),
  inference(resolution,[],[f4750,f2727])).
fof(f11902,plain,(
  $false | $spl102),
  inference(subsumption_resolution,[],[f11830,f4750])).
fof(f11830,plain,(
  ( ! [X27] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X27))))) ) | $spl102),
  inference(resolution,[],[f4750,f2757])).
fof(f11901,plain,(
  $false | $spl102),
  inference(subsumption_resolution,[],[f11829,f4750])).
fof(f11829,plain,(
  ( ! [X26,X25] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X25),s('type.integer.int',X26))))) ) | $spl102),
  inference(resolution,[],[f4750,f159])).
fof(f11900,plain,(
  $false | $spl102),
  inference(subsumption_resolution,[],[f11828,f4750])).
fof(f11828,plain,(
  ( ! [X24,X23] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X23),s('type.integer.int',X24))))) ) | $spl102),
  inference(resolution,[],[f4750,f160])).
fof(f11899,plain,(
  $false | $spl102),
  inference(subsumption_resolution,[],[f11826,f4750])).
fof(f11826,plain,(
  ( ! [X19,X20] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X19),s('type.integer.int',X20))))) ) | $spl102),
  inference(resolution,[],[f4750,f169])).
fof(f11898,plain,(
  $false | $spl102),
  inference(subsumption_resolution,[],[f11825,f4750])).
fof(f11825,plain,(
  ( ! [X17,X18] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X17),s('type.integer.int',X18))))) ) | $spl102),
  inference(resolution,[],[f4750,f169])).
fof(f11897,plain,(
  $false | $spl102),
  inference(subsumption_resolution,[],[f11815,f4750])).
fof(f11815,plain,(
  ( ! [X10,X9] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X9),s('type.integer.int',X10))))) ) | $spl102),
  inference(resolution,[],[f4750,f159])).
fof(f11896,plain,(
  $false | $spl102),
  inference(subsumption_resolution,[],[f11814,f4750])).
fof(f11814,plain,(
  ( ! [X8,X7] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X7),s('type.integer.int',X8))))) ) | $spl102),
  inference(resolution,[],[f4750,f160])).
fof(f11819,plain,(
  $false | ($spl102 | $spl159)),
  inference(resolution,[],[f4750,f11628])).
fof(f11820,plain,(
  $false | ($spl3 | $spl102)),
  inference(resolution,[],[f4750,f198])).
fof(f11821,plain,(
  $false | ($spl3 | $spl22 | $spl102)),
  inference(resolution,[],[f4750,f1733])).
fof(f11822,plain,(
  $false | ($spl102 | $spl149)),
  inference(resolution,[],[f4750,f11318])).
fof(f11823,plain,(
  $false | ($spl5 | $spl102)),
  inference(resolution,[],[f4750,f200])).
fof(f11824,plain,(
  $false | ($spl5 | $spl102)),
  inference(resolution,[],[f4750,f10385])).
fof(f11832,plain,(
  $false | ($spl5 | $spl102)),
  inference(resolution,[],[f4750,f10540])).
fof(f11835,plain,(
  $false | ($spl5 | $spl102)),
  inference(resolution,[],[f4750,f10441])).
fof(f11836,plain,(
  $false | ($spl5 | $spl102)),
  inference(resolution,[],[f4750,f10586])).
fof(f10586,plain,(
  ( ! [X2] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X2))))),s('type.integer.int',X2))))) ) | $spl5),
  inference(resolution,[],[f10444,f160])).
fof(f11837,plain,(
  $false | ($spl5 | $spl102)),
  inference(resolution,[],[f4750,f10740])).
fof(f11838,plain,(
  $false | $spl102),
  inference(resolution,[],[f4750,f9504])).
fof(f9504,plain,(
  ( ! [X4] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X4))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X4))))))))) )),
  inference(resolution,[],[f9205,f160])).
fof(f11839,plain,(
  $false | ($spl5 | $spl102)),
  inference(resolution,[],[f4750,f11176])).
fof(f11840,plain,(
  $false | $spl102),
  inference(resolution,[],[f4750,f2508])).
fof(f11843,plain,(
  $false | $spl102),
  inference(resolution,[],[f4750,f5396])).
fof(f11846,plain,(
  $false | ($spl7 | $spl102)),
  inference(resolution,[],[f4750,f248])).
fof(f11847,plain,(
  $false | ($spl5 | $spl22 | $spl102)),
  inference(resolution,[],[f4750,f11487])).
fof(f11487,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))))) | ($spl5 | $spl22)),
  inference(resolution,[],[f11289,f160])).
fof(f11289,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))),s('type.integer.int',sK1)))) | ($spl5 | $spl22)),
  inference(subsumption_resolution,[],[f11288,f193])).
fof(f11288,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))),s('type.integer.int',sK1)))) | ($spl5 | $spl22)),
  inference(forward_demodulation,[],[f11287,f997])).
fof(f11287,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int',sK1)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))),s('type.integer.int',sK1)))) | ($spl5 | $spl22)),
  inference(forward_demodulation,[],[f11261,f140])).
fof(f11261,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))),s('type.integer.int',sK1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))) | ($spl5 | $spl22)),
  inference(resolution,[],[f11234,f139])).
fof(f11848,plain,(
  $false | ($spl22 | $spl24 | $spl102)),
  inference(resolution,[],[f4750,f1659])).
fof(f11849,plain,(
  $false | ($spl5 | $spl22 | $spl102)),
  inference(resolution,[],[f4750,f11259])).
fof(f11850,plain,(
  $false | ($spl41 | $spl102)),
  inference(resolution,[],[f4750,f800])).
fof(f11801,plain,(
  ( ! [X45] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X45),s('type.integer.int',X45)))) ) | $spl160),
  inference(cnf_transformation,[],[f11801_D])).
fof(f11801_D,plain,(
  ( ! [X45] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X45),s('type.integer.int',X45)))) ) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f11800,plain,(
  ( ! [X45,X44] : (p(s(bool,X44)) | s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X45),s('type.integer.int',X45)))) ) | $spl5),
  inference(subsumption_resolution,[],[f11779,f183])).
fof(f11779,plain,(
  ( ! [X45,X44] : (p(s(bool,X44)) | s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X45),s('type.integer.int',X45))) | s(bool,'T_0') = s(bool,X44)) ) | $spl5),
  inference(superposition,[],[f10404,f1892])).
fof(f10404,plain,(
  ( ! [X4,X3] : (p(s(bool,X4)) | s(bool,X4) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X3),s('type.integer.int',X3)))) ) | $spl5),
  inference(resolution,[],[f10385,f165])).
fof(f11628,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))))) | $spl159),
  inference(cnf_transformation,[],[f11628_D])).
fof(f11628_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))))) <=> ~$spl159),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl159])])).
fof(f11519,plain,(
  s('type.integer.int',sK1) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | $spl152),
  inference(cnf_transformation,[],[f11519_D])).
fof(f11519_D,plain,(
  s('type.integer.int',sK1) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f11599,plain,(
  s('type.integer.int',sK1) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))))) | ($spl5 | $spl22)),
  inference(resolution,[],[f11487,f162])).
fof(f11521,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))))) | $spl155),
  inference(cnf_transformation,[],[f11521_D])).
fof(f11521_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))))) <=> ~$spl155),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl155])])).
fof(f11523,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))))),s('type.integer.int',sK1)))) | $spl156),
  inference(cnf_transformation,[],[f11523_D])).
fof(f11523_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))))),s('type.integer.int',sK1)))) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f11489,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))))),s('type.integer.int',sK1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))))) | ($spl5 | $spl22)),
  inference(resolution,[],[f11289,f139])).
fof(f11517,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))),s('type.integer.int',sK1)))) | $spl150),
  inference(cnf_transformation,[],[f11517_D])).
fof(f11517_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))),s('type.integer.int',sK1)))) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f11486,plain,(
  s('type.integer.int',sK1) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))))),s('type.integer.int',sK1)))) | ($spl5 | $spl22)),
  inference(resolution,[],[f11289,f162])).
fof(f11291,plain,(
  s('type.integer.int',sK1) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))) | ($spl5 | $spl22)),
  inference(resolution,[],[f11259,f162])).
fof(f11258,plain,(
  s('type.integer.int',sK1) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK1)))) | ($spl5 | $spl22)),
  inference(resolution,[],[f11234,f162])).
fof(f10955,plain,(
  $false | ($spl64 | $spl80 | $spl121)),
  inference(subsumption_resolution,[],[f10822,f1652])).
fof(f10822,plain,(
  s('type.integer.int',sK1) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))) | ($spl80 | $spl121)),
  inference(backward_demodulation,[],[f3444,f8629])).
fof(f8629,plain,(
  s('type.integer.int',sK1) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))))) | $spl121),
  inference(resolution,[],[f8157,f164])).
fof(f10954,plain,(
  $false | ($spl5 | $spl80 | $spl121)),
  inference(subsumption_resolution,[],[f10953,f189])).
fof(f10953,plain,(
  ~p(s(bool,'T_0')) | ($spl5 | $spl80 | $spl121)),
  inference(forward_demodulation,[],[f10820,f10446])).
fof(f10820,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1)))))) | ($spl80 | $spl121)),
  inference(backward_demodulation,[],[f3444,f8157])).
fof(f10952,plain,(
  $false | ($spl5 | $spl80 | $spl110)),
  inference(subsumption_resolution,[],[f10818,f10441])).
fof(f10818,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int',sK1)))) | ($spl80 | $spl110)),
  inference(backward_demodulation,[],[f3444,f7463])).
fof(f3442,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))),s('type.integer.int',sK1)))) | $spl78),
  inference(cnf_transformation,[],[f3442_D])).
fof(f3442_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))),s('type.integer.int',sK1)))) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f3444,plain,(
  s('type.integer.int',sK1) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | $spl80),
  inference(cnf_transformation,[],[f3444_D])).
fof(f3444_D,plain,(
  s('type.integer.int',sK1) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f6475,plain,(
  s('type.integer.int',sK1) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))),s('type.integer.int',sK1)))) | ($spl22 | $spl24)),
  inference(resolution,[],[f5094,f162])).
fof(f5094,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))),s('type.integer.int',sK1)))) | ($spl22 | $spl24)),
  inference(subsumption_resolution,[],[f5072,f187])).
fof(f5072,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))),s('type.integer.int',sK1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | ($spl22 | $spl24)),
  inference(resolution,[],[f4903,f139])).
fof(f4903,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int',sK1)))) | ($spl22 | $spl24)),
  inference(forward_demodulation,[],[f656,f655])).
fof(f10810,plain,(
  s('type.integer.int',sK1) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))),s('type.integer.int',sK1)))) | ($spl22 | $spl24)),
  inference(subsumption_resolution,[],[f7592,f193])).
fof(f7592,plain,(
  ~p(s(bool,'T_0')) | s('type.integer.int',sK1) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))),s('type.integer.int',sK1)))) | ($spl22 | $spl24)),
  inference(superposition,[],[f162,f6479])).
fof(f6479,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))),s('type.integer.int',sK1))) | ($spl22 | $spl24)),
  inference(resolution,[],[f5094,f186])).
fof(f10809,plain,(
  $false | ($spl5 | $spl78)),
  inference(subsumption_resolution,[],[f3956,f10564])).
fof(f3956,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))))) | $spl78),
  inference(resolution,[],[f3442,f160])).
fof(f10808,plain,(
  $false | ($spl5 | $spl78)),
  inference(subsumption_resolution,[],[f3957,f10564])).
fof(f3957,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))))) | $spl78),
  inference(resolution,[],[f3442,f159])).
fof(f10806,plain,(
  $false | ($spl5 | $spl78)),
  inference(subsumption_resolution,[],[f10805,f10564])).
fof(f10805,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))))) | $spl78),
  inference(subsumption_resolution,[],[f7699,f193])).
fof(f7699,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))))) | $spl78),
  inference(superposition,[],[f159,f3958])).
fof(f3958,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))),s('type.integer.int',sK1))) | $spl78),
  inference(resolution,[],[f3442,f186])).
fof(f10804,plain,(
  $false | ($spl5 | $spl78)),
  inference(subsumption_resolution,[],[f10803,f10564])).
fof(f10803,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))))) | $spl78),
  inference(subsumption_resolution,[],[f7700,f193])).
fof(f7700,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))))) | $spl78),
  inference(superposition,[],[f160,f3958])).
fof(f10759,plain,(
  $false | ($spl5 | $spl91)),
  inference(resolution,[],[f10564,f3952])).
fof(f10310,plain,(
  ( ! [X41] : (p(s(bool,X41)) | s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) ) | ($spl3 | $spl5 | $spl22)),
  inference(subsumption_resolution,[],[f10247,f183])).
fof(f10247,plain,(
  ( ! [X41] : (p(s(bool,X41)) | s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | s(bool,'T_0') = s(bool,X41)) ) | ($spl3 | $spl5 | $spl22)),
  inference(superposition,[],[f1741,f1892])).
fof(f10307,plain,(
  ( ! [X40] : (p(s(bool,X40)) | s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) ) | ($spl5 | $spl7)),
  inference(subsumption_resolution,[],[f10246,f183])).
fof(f10246,plain,(
  ( ! [X40] : (p(s(bool,X40)) | s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | s(bool,'T_0') = s(bool,X40)) ) | ($spl5 | $spl7)),
  inference(superposition,[],[f1024,f1892])).
fof(f10304,plain,(
  ( ! [X39] : (p(s(bool,X39)) | s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) ) | ($spl3 | $spl5)),
  inference(subsumption_resolution,[],[f10245,f183])).
fof(f10245,plain,(
  ( ! [X39] : (p(s(bool,X39)) | s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | s(bool,'T_0') = s(bool,X39)) ) | ($spl3 | $spl5)),
  inference(superposition,[],[f977,f1892])).
fof(f9008,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))))) | $spl137),
  inference(cnf_transformation,[],[f9008_D])).
fof(f9008_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))))) <=> ~$spl137),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl137])])).
fof(f8923,plain,(
  s('type.integer.int',sK1) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl128),
  inference(cnf_transformation,[],[f8923_D])).
fof(f8923_D,plain,(
  s('type.integer.int',sK1) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f8982,plain,(
  s('type.integer.int',sK1) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))))) | $spl98),
  inference(resolution,[],[f8894,f162])).
fof(f8894,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))))) | $spl98),
  inference(resolution,[],[f4084,f160])).
fof(f8954,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | $spl135),
  inference(cnf_transformation,[],[f8954_D])).
fof(f8954_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) <=> ~$spl135),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl135])])).
fof(f8889,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl124),
  inference(cnf_transformation,[],[f8889_D])).
fof(f8889_D,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f8928,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | $spl96),
  inference(resolution,[],[f8859,f162])).
fof(f8925,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))))) | $spl131),
  inference(cnf_transformation,[],[f8925_D])).
fof(f8925_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))))) <=> ~$spl131),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl131])])).
fof(f8927,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))),s('type.integer.int',sK1)))) | $spl132),
  inference(cnf_transformation,[],[f8927_D])).
fof(f8927_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))),s('type.integer.int',sK1)))) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f8896,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))),s('type.integer.int',sK1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))))) | $spl98),
  inference(resolution,[],[f4084,f139])).
fof(f8921,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.integer.int',sK1)))) | $spl126),
  inference(cnf_transformation,[],[f8921_D])).
fof(f8921_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.integer.int',sK1)))) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f8893,plain,(
  s('type.integer.int',sK1) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.integer.int',sK1)))) | $spl98),
  inference(resolution,[],[f4084,f162])).
fof(f8887,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl122),
  inference(cnf_transformation,[],[f8887_D])).
fof(f8887_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f8858,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl96),
  inference(resolution,[],[f4081,f162])).
fof(f8855,plain,(
  $false | ($spl3 | $spl22 | $spl97)),
  inference(subsumption_resolution,[],[f8854,f1733])).
fof(f8854,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | $spl97),
  inference(forward_demodulation,[],[f8853,f148])).
fof(f8853,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl97),
  inference(forward_demodulation,[],[f8823,f150])).
fof(f8823,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl97),
  inference(resolution,[],[f8498,f138])).
fof(f8498,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | $spl97),
  inference(resolution,[],[f7910,f161])).
fof(f7910,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl97),
  inference(resolution,[],[f4082,f163])).
fof(f8852,plain,(
  $false | ($spl3 | $spl22 | $spl97)),
  inference(subsumption_resolution,[],[f8851,f1733])).
fof(f8851,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | $spl97),
  inference(forward_demodulation,[],[f8850,f148])).
fof(f8850,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl97),
  inference(forward_demodulation,[],[f8822,f145])).
fof(f8822,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | $spl97),
  inference(resolution,[],[f8498,f137])).
fof(f8157,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))))))) | $spl121),
  inference(cnf_transformation,[],[f8157_D])).
fof(f8157_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))))))) <=> ~$spl121),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl121])])).
fof(f7465,plain,(
  s('type.integer.int',sK1) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))))) | $spl112),
  inference(cnf_transformation,[],[f7465_D])).
fof(f7465_D,plain,(
  s('type.integer.int',sK1) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))))) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f8131,plain,(
  s('type.integer.int',sK1) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))))))) | $spl84),
  inference(resolution,[],[f7437,f162])).
fof(f7437,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))))))) | $spl84),
  inference(resolution,[],[f3448,f160])).
fof(f8066,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | $spl119),
  inference(cnf_transformation,[],[f8066_D])).
fof(f8066_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) <=> ~$spl119),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl119])])).
fof(f7433,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | $spl108),
  inference(cnf_transformation,[],[f7433_D])).
fof(f7433_D,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f8042,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | $spl82),
  inference(resolution,[],[f7406,f162])).
fof(f7406,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | $spl82),
  inference(resolution,[],[f3445,f160])).
fof(f4082,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl97),
  inference(cnf_transformation,[],[f4082_D])).
fof(f4082_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) <=> ~$spl97),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl97])])).
fof(f4084,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.integer.int',sK1)))) | $spl98),
  inference(cnf_transformation,[],[f4084_D])).
fof(f4084_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.integer.int',sK1)))) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f7898,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.integer.int',sK1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl88),
  inference(subsumption_resolution,[],[f7861,f193])).
fof(f7861,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.integer.int',sK1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl88),
  inference(superposition,[],[f139,f4055])).
fof(f4055,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',sK1))) | $spl88),
  inference(resolution,[],[f3658,f186])).
fof(f7467,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))))))) | $spl115),
  inference(cnf_transformation,[],[f7467_D])).
fof(f7467_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))))))) <=> ~$spl115),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl115])])).
fof(f7469,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))))))),s('type.integer.int',sK1)))) | $spl116),
  inference(cnf_transformation,[],[f7469_D])).
fof(f7469_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))))))),s('type.integer.int',sK1)))) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f7439,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))))))),s('type.integer.int',sK1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))))))) | $spl84),
  inference(resolution,[],[f3448,f139])).
fof(f7463,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))))),s('type.integer.int',sK1)))) | $spl110),
  inference(cnf_transformation,[],[f7463_D])).
fof(f7463_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))))),s('type.integer.int',sK1)))) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f7436,plain,(
  s('type.integer.int',sK1) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))))),s('type.integer.int',sK1)))) | $spl84),
  inference(resolution,[],[f3448,f162])).
fof(f7431,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))))) | $spl106),
  inference(cnf_transformation,[],[f7431_D])).
fof(f7431_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))))) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f7405,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))))) | $spl82),
  inference(resolution,[],[f3445,f162])).
fof(f7012,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.integer.int',sK1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl88),
  inference(resolution,[],[f3658,f139])).
fof(f3446,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))))) | $spl83),
  inference(cnf_transformation,[],[f3446_D])).
fof(f3446_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))))) <=> ~$spl83),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl83])])).
fof(f6478,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))))),s('type.integer.int',sK1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))))) | ($spl22 | $spl24)),
  inference(resolution,[],[f5094,f139])).
fof(f3656,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl87),
  inference(cnf_transformation,[],[f3656_D])).
fof(f3656_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl87),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl87])])).
fof(f3658,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',sK1)))) | $spl88),
  inference(cnf_transformation,[],[f3658_D])).
fof(f3658_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',sK1)))) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f6379,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',sK1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl22),
  inference(subsumption_resolution,[],[f6345,f193])).
fof(f6345,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',sK1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl22),
  inference(superposition,[],[f139,f3556])).
fof(f3556,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | $spl22),
  inference(superposition,[],[f2511,f655])).
fof(f5925,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',sK1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl41),
  inference(resolution,[],[f5745,f139])).
fof(f5745,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | $spl41),
  inference(resolution,[],[f800,f161])).
fof(f5670,plain,(
  $false | ($spl3 | $spl22 | $spl104)),
  inference(subsumption_resolution,[],[f5669,f1735])).
fof(f1735,plain,(
  s('type.integer.int',sK1) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))) | ($spl3 | $spl22)),
  inference(resolution,[],[f1733,f164])).
fof(f5669,plain,(
  s('type.integer.int',sK1) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))) | $spl104),
  inference(forward_demodulation,[],[f5582,f150])).
fof(f5582,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | $spl104),
  inference(backward_demodulation,[],[f4762,f151])).
fof(f5123,plain,(
  $false | ($spl22 | $spl24 | $spl40)),
  inference(subsumption_resolution,[],[f5122,f4903])).
fof(f5122,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int',sK1)))) | $spl40),
  inference(forward_demodulation,[],[f5099,f149])).
fof(f5099,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | $spl40),
  inference(resolution,[],[f801,f159])).
fof(f5121,plain,(
  $false | ($spl22 | $spl24 | $spl40)),
  inference(subsumption_resolution,[],[f5120,f4903])).
fof(f5120,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int',sK1)))) | $spl40),
  inference(forward_demodulation,[],[f5098,f149])).
fof(f5098,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | $spl40),
  inference(resolution,[],[f801,f160])).
fof(f5119,plain,(
  $false | ($spl3 | $spl22 | $spl40)),
  inference(subsumption_resolution,[],[f5118,f1733])).
fof(f5118,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | $spl40),
  inference(forward_demodulation,[],[f5097,f148])).
fof(f5097,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl40),
  inference(resolution,[],[f801,f163])).
fof(f4898,plain,(
  $false | ($spl3 | $spl22 | $spl31)),
  inference(subsumption_resolution,[],[f4897,f1733])).
fof(f4897,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | $spl31),
  inference(forward_demodulation,[],[f4872,f148])).
fof(f4872,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl31),
  inference(resolution,[],[f755,f158])).
fof(f4896,plain,(
  $false | ($spl3 | $spl22 | $spl31)),
  inference(subsumption_resolution,[],[f4895,f1733])).
fof(f4895,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | $spl31),
  inference(forward_demodulation,[],[f4871,f148])).
fof(f4871,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl31),
  inference(resolution,[],[f755,f161])).
fof(f4891,plain,(
  $false | ($spl3 | $spl22 | $spl31)),
  inference(subsumption_resolution,[],[f4868,f1733])).
fof(f4868,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | $spl31),
  inference(resolution,[],[f755,f2587])).
fof(f4787,plain,(
  $false | ($spl22 | $spl25 | $spl64)),
  inference(subsumption_resolution,[],[f4786,f187])).
fof(f4786,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK1)))) | ($spl22 | $spl25 | $spl64)),
  inference(forward_demodulation,[],[f4785,f1652])).
fof(f4785,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int',sK1)))) | ($spl22 | $spl25)),
  inference(forward_demodulation,[],[f1147,f655])).
fof(f1147,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0)))))) | $spl25),
  inference(resolution,[],[f657,f138])).
fof(f4780,plain,(
  $false | ($spl22 | $spl25 | $spl64)),
  inference(subsumption_resolution,[],[f4779,f187])).
fof(f4779,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK1)))) | ($spl22 | $spl25 | $spl64)),
  inference(forward_demodulation,[],[f1564,f1652])).
fof(f1564,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int',sK1)))) | ($spl22 | $spl25)),
  inference(resolution,[],[f1352,f138])).
fof(f1352,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int',sK1)))) | ($spl22 | $spl25)),
  inference(backward_demodulation,[],[f655,f657])).
fof(f4778,plain,(
  $false | ($spl22 | $spl25 | $spl64)),
  inference(subsumption_resolution,[],[f4777,f187])).
fof(f4777,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK1)))) | ($spl22 | $spl25 | $spl64)),
  inference(forward_demodulation,[],[f2071,f1652])).
fof(f2071,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int',sK1)))) | ($spl22 | $spl25)),
  inference(forward_demodulation,[],[f1147,f655])).
fof(f4762,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl104),
  inference(cnf_transformation,[],[f4762_D])).
fof(f4762_D,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f4043,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl86),
  inference(forward_demodulation,[],[f4016,f147])).
fof(f4016,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl86),
  inference(resolution,[],[f3655,f162])).
fof(f1178,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int',sK1)))) | $spl52),
  inference(cnf_transformation,[],[f1178_D])).
fof(f1178_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int',sK1)))) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f3649,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int',sK1)))) | p(s(bool,X0))) ) | $spl22),
  inference(forward_demodulation,[],[f3648,f149])).
fof(f3648,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | p(s(bool,X0))) ) | $spl22),
  inference(subsumption_resolution,[],[f3614,f183])).
fof(f3614,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | p(s(bool,X0))) ) | $spl22),
  inference(superposition,[],[f3556,f165])).
fof(f3652,plain,(
  ( ! [X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int',sK1)))) | p(s(bool,X1))) ) | $spl22),
  inference(forward_demodulation,[],[f3651,f149])).
fof(f3651,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))))) ) | $spl22),
  inference(subsumption_resolution,[],[f3615,f183])).
fof(f3615,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,X1) | p(s(bool,X1)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))))) ) | $spl22),
  inference(superposition,[],[f3556,f165])).
fof(f3660,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int',sK1)))) | p(s(bool,X0))) ) | $spl22),
  inference(forward_demodulation,[],[f3659,f149])).
fof(f3659,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))))) ) | $spl22),
  inference(subsumption_resolution,[],[f3629,f183])).
fof(f3629,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X0)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))))) ) | $spl22),
  inference(superposition,[],[f165,f3556])).
fof(f3663,plain,(
  ( ! [X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int',sK1)))) | p(s(bool,X1))) ) | $spl22),
  inference(forward_demodulation,[],[f3662,f149])).
fof(f3662,plain,(
  ( ! [X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | p(s(bool,X1))) ) | $spl22),
  inference(subsumption_resolution,[],[f3630,f183])).
fof(f3630,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,X1) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | p(s(bool,X1))) ) | $spl22),
  inference(superposition,[],[f165,f3556])).
fof(f4748,plain,(
  $false | $spl83),
  inference(subsumption_resolution,[],[f4727,f187])).
fof(f4727,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))))) | $spl83),
  inference(resolution,[],[f3446,f137])).
fof(f4614,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))))),s('type.integer.int',sK1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))))) | ($spl22 | $spl24)),
  inference(subsumption_resolution,[],[f4580,f193])).
fof(f4580,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))))),s('type.integer.int',sK1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))))) | ($spl22 | $spl24)),
  inference(superposition,[],[f139,f3420])).
fof(f3420,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))),s('type.integer.int',sK1))) | ($spl22 | $spl24)),
  inference(resolution,[],[f1684,f186])).
fof(f1684,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))),s('type.integer.int',sK1)))) | ($spl22 | $spl24)),
  inference(subsumption_resolution,[],[f1661,f187])).
fof(f1661,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))),s('type.integer.int',sK1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | ($spl22 | $spl24)),
  inference(resolution,[],[f1624,f139])).
fof(f4178,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl101),
  inference(cnf_transformation,[],[f4178_D])).
fof(f4178_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) <=> ~$spl101),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl101])])).
fof(f4080,plain,(
  s('type.integer.int',sK1) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl94),
  inference(cnf_transformation,[],[f4080_D])).
fof(f4080_D,plain,(
  s('type.integer.int',sK1) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f4154,plain,(
  s('type.integer.int',sK1) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl88),
  inference(resolution,[],[f4052,f162])).
fof(f4052,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl88),
  inference(resolution,[],[f3658,f160])).
fof(f4090,plain,(
  $false | ($spl7 | $spl64 | $spl77)),
  inference(subsumption_resolution,[],[f4089,f189])).
fof(f4089,plain,(
  ~p(s(bool,'T_0')) | ($spl7 | $spl64 | $spl77)),
  inference(forward_demodulation,[],[f4088,f1250])).
fof(f4088,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int',sK0)))) | ($spl64 | $spl77)),
  inference(forward_demodulation,[],[f2038,f1652])).
fof(f4054,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.integer.int',sK1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl88),
  inference(resolution,[],[f3658,f139])).
fof(f4078,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',sK1)))) | $spl92),
  inference(cnf_transformation,[],[f4078_D])).
fof(f4078_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',sK1)))) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f4051,plain,(
  s('type.integer.int',sK1) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',sK1)))) | $spl88),
  inference(resolution,[],[f3658,f162])).
fof(f4014,plain,(
  $false | ($spl30 | $spl87)),
  inference(subsumption_resolution,[],[f4013,f1590])).
fof(f1590,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl30),
  inference(resolution,[],[f756,f169])).
fof(f4013,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl87),
  inference(forward_demodulation,[],[f3983,f146])).
fof(f3983,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | $spl87),
  inference(resolution,[],[f3656,f158])).
fof(f4012,plain,(
  $false | ($spl30 | $spl87)),
  inference(subsumption_resolution,[],[f4011,f1590])).
fof(f4011,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl87),
  inference(forward_demodulation,[],[f3982,f146])).
fof(f3982,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | $spl87),
  inference(resolution,[],[f3656,f161])).
fof(f4010,plain,(
  $false | ($spl30 | $spl87)),
  inference(subsumption_resolution,[],[f4009,f189])).
fof(f4009,plain,(
  ~p(s(bool,'T_0')) | ($spl30 | $spl87)),
  inference(forward_demodulation,[],[f4008,f2542])).
fof(f2542,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | $spl30),
  inference(resolution,[],[f756,f186])).
fof(f4008,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | $spl87),
  inference(forward_demodulation,[],[f3980,f147])).
fof(f3980,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl87),
  inference(resolution,[],[f3656,f163])).
fof(f4007,plain,(
  $false | ($spl3 | $spl22 | $spl87)),
  inference(subsumption_resolution,[],[f4006,f1733])).
fof(f4006,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | $spl87),
  inference(forward_demodulation,[],[f4005,f148])).
fof(f4005,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl87),
  inference(forward_demodulation,[],[f3978,f145])).
fof(f3978,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | $spl87),
  inference(resolution,[],[f3656,f3027])).
fof(f3952,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))))) | $spl91),
  inference(cnf_transformation,[],[f3952_D])).
fof(f3952_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))))) <=> ~$spl91),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl91])])).
fof(f3928,plain,(
  s('type.integer.int',sK1) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))))) | ($spl22 | $spl24)),
  inference(resolution,[],[f3417,f162])).
fof(f3417,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))))) | ($spl22 | $spl24)),
  inference(resolution,[],[f1684,f160])).
fof(f3654,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',sK1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl22),
  inference(subsumption_resolution,[],[f3621,f193])).
fof(f3621,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',sK1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl22),
  inference(superposition,[],[f139,f3556])).
fof(f3543,plain,(
  $false | ($spl17 | $spl22 | $spl24)),
  inference(subsumption_resolution,[],[f3542,f189])).
fof(f3542,plain,(
  ~p(s(bool,'T_0')) | ($spl17 | $spl22 | $spl24)),
  inference(forward_demodulation,[],[f3541,f1662])).
fof(f3541,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int',sK1)))) | ($spl17 | $spl22)),
  inference(forward_demodulation,[],[f596,f655])).
fof(f3419,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))))),s('type.integer.int',sK1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))))) | ($spl22 | $spl24)),
  inference(resolution,[],[f1684,f139])).
fof(f3416,plain,(
  s('type.integer.int',sK1) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))),s('type.integer.int',sK1)))) | ($spl22 | $spl24)),
  inference(resolution,[],[f1684,f162])).
fof(f3351,plain,(
  $false | ($spl22 | $spl24 | $spl64 | $spl67)),
  inference(subsumption_resolution,[],[f3350,f1624])).
fof(f3350,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int',sK1)))) | ($spl64 | $spl67)),
  inference(forward_demodulation,[],[f1654,f1652])).
fof(f2413,plain,(
  $false | ($spl0 | $spl22 | $spl62 | $spl64)),
  inference(subsumption_resolution,[],[f2348,f1353])).
fof(f2348,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK1)))) | ($spl62 | $spl64)),
  inference(backward_demodulation,[],[f1652,f1650])).
fof(f2207,plain,(
  $false | ($spl3 | $spl22 | $spl72)),
  inference(subsumption_resolution,[],[f1683,f1735])).
fof(f2142,plain,(
  $false | ($spl3 | $spl22 | $spl28)),
  inference(subsumption_resolution,[],[f2081,f1735])).
fof(f2070,plain,(
  $false | ($spl22 | $spl25 | $spl70)),
  inference(subsumption_resolution,[],[f2069,f1681])).
fof(f2069,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int',sK1)))) | ($spl22 | $spl25)),
  inference(forward_demodulation,[],[f1148,f655])).
fof(f1148,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))))),s('type.integer.int',sK1)))) | $spl25),
  inference(resolution,[],[f657,f163])).
fof(f2065,plain,(
  $false | ($spl22 | $spl25 | $spl70)),
  inference(subsumption_resolution,[],[f1565,f1681])).
fof(f1565,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int',sK1)))) | ($spl22 | $spl25)),
  inference(resolution,[],[f1352,f163])).
fof(f2061,plain,(
  $false | $spl75),
  inference(subsumption_resolution,[],[f2043,f187])).
fof(f2043,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1)))))) | $spl75),
  inference(resolution,[],[f1804,f137])).
fof(f1803,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1)))))) | $spl74),
  inference(cnf_transformation,[],[f1803_D])).
fof(f1803_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1)))))) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f2036,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))))),s('type.integer.int',sK0)))) | ($spl5 | $spl22)),
  inference(forward_demodulation,[],[f2015,f655])).
fof(f1924,plain,(
  $false | ($spl7 | $spl13)),
  inference(subsumption_resolution,[],[f347,f1020])).
fof(f1804,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1)))))) | $spl75),
  inference(cnf_transformation,[],[f1804_D])).
fof(f1804_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1)))))) <=> ~$spl75),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl75])])).
fof(f1652,plain,(
  s('type.integer.int',sK1) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))) | $spl64),
  inference(cnf_transformation,[],[f1652_D])).
fof(f1652_D,plain,(
  s('type.integer.int',sK1) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f1784,plain,(
  s('type.integer.int',sK1) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1)))))) | ($spl22 | $spl26)),
  inference(resolution,[],[f1628,f162])).
fof(f1628,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1)))))) | ($spl22 | $spl26)),
  inference(resolution,[],[f1623,f160])).
fof(f1623,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int',sK1)))) | ($spl22 | $spl26)),
  inference(forward_demodulation,[],[f659,f655])).
fof(f1759,plain,(
  $false | ($spl3 | $spl22 | $spl58)),
  inference(subsumption_resolution,[],[f1457,f1733])).
fof(f1732,plain,(
  $false | ($spl22 | $spl24 | $spl53)),
  inference(subsumption_resolution,[],[f1179,f1624])).
fof(f1681,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int',sK1)))) | $spl70),
  inference(cnf_transformation,[],[f1681_D])).
fof(f1681_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int',sK1)))) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f1683,plain,(
  s('type.integer.int',sK1) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))) | $spl72),
  inference(cnf_transformation,[],[f1683_D])).
fof(f1683_D,plain,(
  s('type.integer.int',sK1) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f1658,plain,(
  s('type.integer.int',sK1) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int',sK1)))) | ($spl22 | $spl24)),
  inference(resolution,[],[f1624,f162])).
fof(f1654,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1)))))) | $spl67),
  inference(cnf_transformation,[],[f1654_D])).
fof(f1654_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1)))))) <=> ~$spl67),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl67])])).
fof(f1656,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))))),s('type.integer.int',sK1)))) | $spl68),
  inference(cnf_transformation,[],[f1656_D])).
fof(f1656_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))))),s('type.integer.int',sK1)))) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f1630,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))))),s('type.integer.int',sK1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1)))))) | ($spl22 | $spl26)),
  inference(resolution,[],[f1623,f139])).
fof(f1650,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int',sK1)))) | $spl62),
  inference(cnf_transformation,[],[f1650_D])).
fof(f1650_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int',sK1)))) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f1627,plain,(
  s('type.integer.int',sK1) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int',sK1)))) | ($spl22 | $spl26)),
  inference(resolution,[],[f1623,f162])).
fof(f1559,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int',sK1)))) | $spl60),
  inference(cnf_transformation,[],[f1559_D])).
fof(f1559_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int',sK1)))) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f1179,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int',sK1)))) | $spl53),
  inference(cnf_transformation,[],[f1179_D])).
fof(f1179_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int',sK1)))) <=> ~$spl53),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl53])])).
fof(f1622,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int',sK1)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int',sK1)))) | ($spl0 | $spl22)),
  inference(forward_demodulation,[],[f1621,f655])).
fof(f1621,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int',sK1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))))),s('type.integer.int',sK1)))) | ($spl0 | $spl22)),
  inference(forward_demodulation,[],[f1620,f655])).
fof(f1620,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))))),s('type.integer.int',sK1)))) | $spl0),
  inference(subsumption_resolution,[],[f1284,f193])).
fof(f1284,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))))),s('type.integer.int',sK1)))) | $spl0),
  inference(superposition,[],[f139,f904])).
fof(f904,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))))) | $spl0),
  inference(resolution,[],[f195,f186])).
fof(f1616,plain,(
  $false | ($spl22 | $spl25 | $spl30)),
  inference(subsumption_resolution,[],[f1615,f1352])).
fof(f1615,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int',sK1)))) | $spl30),
  inference(forward_demodulation,[],[f1592,f149])).
fof(f1592,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | $spl30),
  inference(resolution,[],[f756,f163])).
fof(f1557,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int',sK1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int',sK1)))) | ($spl0 | $spl22)),
  inference(subsumption_resolution,[],[f1530,f193])).
fof(f1530,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int',sK1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int',sK1)))) | ($spl0 | $spl22)),
  inference(superposition,[],[f139,f1351])).
fof(f1351,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK1))) | ($spl0 | $spl22)),
  inference(backward_demodulation,[],[f655,f904])).
fof(f1463,plain,(
  s('type.integer.int',sK1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | $spl22),
  inference(forward_demodulation,[],[f1462,f655])).
fof(f1462,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl22),
  inference(forward_demodulation,[],[f1461,f655])).
fof(f1461,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl22),
  inference(forward_demodulation,[],[f1383,f146])).
fof(f1383,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0)))))))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl22),
  inference(superposition,[],[f143,f655])).
fof(f1460,plain,(
  s('type.integer.int',sK1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl22),
  inference(forward_demodulation,[],[f1459,f655])).
fof(f1459,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl22),
  inference(forward_demodulation,[],[f1458,f655])).
fof(f1458,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0)))))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl22),
  inference(forward_demodulation,[],[f1382,f146])).
fof(f1382,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0)))))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl22),
  inference(superposition,[],[f143,f655])).
fof(f1457,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | $spl58),
  inference(cnf_transformation,[],[f1457_D])).
fof(f1457_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f1455,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int',sK1)))) | $spl22),
  inference(forward_demodulation,[],[f1454,f148])).
fof(f1454,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int',sK1)))) | $spl22),
  inference(forward_demodulation,[],[f1453,f655])).
fof(f1453,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int',sK1)))) | $spl22),
  inference(forward_demodulation,[],[f1452,f145])).
fof(f1452,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int',sK1)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0)))))))) | $spl22),
  inference(forward_demodulation,[],[f1381,f149])).
fof(f1381,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0)))))))) | $spl22),
  inference(superposition,[],[f142,f655])).
fof(f653,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0)))))) | $spl20),
  inference(cnf_transformation,[],[f653_D])).
fof(f653_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0)))))) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f655,plain,(
  s('type.integer.int',sK1) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))) | $spl22),
  inference(cnf_transformation,[],[f655_D])).
fof(f655_D,plain,(
  s('type.integer.int',sK1) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f1350,plain,(
  s('type.integer.int',sK1) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0)))))) | $spl0),
  inference(subsumption_resolution,[],[f1289,f193])).
fof(f1289,plain,(
  ~p(s(bool,'T_0')) | s('type.integer.int',sK1) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0)))))) | $spl0),
  inference(superposition,[],[f162,f904])).
fof(f1349,plain,(
  $false | ($spl5 | $spl7 | $spl20)),
  inference(subsumption_resolution,[],[f1348,f283])).
fof(f1183,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl56),
  inference(cnf_transformation,[],[f1183_D])).
fof(f1183_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f1177,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int',sK1)))) | $spl25),
  inference(forward_demodulation,[],[f1176,f146])).
fof(f1176,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int',sK1)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0)))))))) | $spl25),
  inference(forward_demodulation,[],[f1175,f147])).
fof(f1175,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int',sK1)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0)))))))) | $spl25),
  inference(forward_demodulation,[],[f1174,f146])).
fof(f1174,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int',sK1)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0)))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0)))))))) | $spl25),
  inference(forward_demodulation,[],[f1163,f149])).
fof(f1163,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0)))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0)))))))) | $spl25),
  inference(superposition,[],[f657,f143])).
fof(f348,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int',sK1)))) | $spl15),
  inference(cnf_transformation,[],[f348_D])).
fof(f348_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int',sK1)))) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f1040,plain,(
  s('type.integer.int',sK1) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int',sK1)))) | $spl0),
  inference(resolution,[],[f629,f162])).
fof(f996,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int',sK1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int',sK0)))) | $spl5),
  inference(resolution,[],[f283,f139])).
fof(f657,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))))),s('type.integer.int',sK1)))) | $spl25),
  inference(cnf_transformation,[],[f657_D])).
fof(f657_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))))),s('type.integer.int',sK1)))) <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f659,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0)))))) | $spl26),
  inference(cnf_transformation,[],[f659_D])).
fof(f659_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0)))))) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f903,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))))),s('type.integer.int',sK1)))) | $spl0),
  inference(resolution,[],[f195,f139])).
fof(f900,plain,(
  s('type.integer.int',sK1) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0)))))) | $spl0),
  inference(resolution,[],[f195,f162])).
fof(f850,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int',sK0)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK1)))) | $spl4),
  inference(resolution,[],[f199,f139])).
fof(f870,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK0)))) | $spl44),
  inference(cnf_transformation,[],[f870_D])).
fof(f870_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK0)))) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f847,plain,(
  s('type.integer.int',sK0) = s('type.integer.int',sK1) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK0)))) | $spl4),
  inference(resolution,[],[f199,f162])).
fof(f758,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK0)))) | $spl32),
  inference(cnf_transformation,[],[f758_D])).
fof(f758_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK0)))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f806,plain,(
  s('type.integer.int',sK0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK0)))) | $spl8),
  inference(forward_demodulation,[],[f805,f251])).
fof(f805,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK0)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl8),
  inference(forward_demodulation,[],[f804,f147])).
fof(f804,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK0)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl8),
  inference(forward_demodulation,[],[f702,f146])).
fof(f702,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK0)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl8),
  inference(superposition,[],[f143,f251])).
fof(f799,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl38),
  inference(cnf_transformation,[],[f799_D])).
fof(f799_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f797,plain,(
  s('type.integer.int',sK0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl8),
  inference(forward_demodulation,[],[f796,f251])).
fof(f796,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl8),
  inference(forward_demodulation,[],[f795,f146])).
fof(f795,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl8),
  inference(forward_demodulation,[],[f701,f146])).
fof(f701,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl8),
  inference(superposition,[],[f143,f251])).
fof(f790,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK0)))) | $spl8),
  inference(forward_demodulation,[],[f789,f148])).
fof(f789,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK0)))) | $spl8),
  inference(forward_demodulation,[],[f788,f251])).
fof(f788,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK0)))) | $spl8),
  inference(forward_demodulation,[],[f787,f145])).
fof(f787,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK0))),s('type.integer.int',sK0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))))) | $spl8),
  inference(forward_demodulation,[],[f700,f149])).
fof(f700,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))))) | $spl8),
  inference(superposition,[],[f142,f251])).
fof(f675,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK0)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | s('type.integer.int',sK1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl8),
  inference(superposition,[],[f143,f251])).
fof(f631,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))))),s('type.integer.int',sK1)))) | $spl0),
  inference(resolution,[],[f195,f139])).
fof(f628,plain,(
  s('type.integer.int',sK1) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0)))))) | $spl0),
  inference(resolution,[],[f195,f162])).
fof(f623,plain,(
  $false | ($spl1 | $spl6)),
  inference(subsumption_resolution,[],[f599,f351])).
fof(f599,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int',sK1)))) | $spl6),
  inference(resolution,[],[f420,f138])).
fof(f420,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int',sK0)))) | $spl6),
  inference(resolution,[],[f249,f160])).
fof(f596,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0)))))) | $spl17),
  inference(cnf_transformation,[],[f596_D])).
fof(f596_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0)))))) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f574,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))))),s('type.integer.int',sK1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0)))))) | $spl1),
  inference(resolution,[],[f351,f139])).
fof(f251,plain,(
  s('type.integer.int',sK0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))) | $spl8),
  inference(cnf_transformation,[],[f251_D])).
fof(f251_D,plain,(
  s('type.integer.int',sK0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f396,plain,(
  s('type.integer.int',sK0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int',sK0)))) | $spl2),
  inference(resolution,[],[f225,f162])).
fof(f225,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int',sK0)))) | $spl2),
  inference(resolution,[],[f197,f160])).
fof(f327,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int',sK1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))),s('type.integer.int',sK0)))) | $spl5),
  inference(resolution,[],[f283,f139])).
fof(f254,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | $spl10),
  inference(cnf_transformation,[],[f254_D])).
fof(f254_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f252,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | $spl2),
  inference(forward_demodulation,[],[f227,f150])).
fof(f227,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))),s('type.integer.int',sK0)))) | $spl2),
  inference(resolution,[],[f197,f139])).
fof(f249,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | $spl6),
  inference(cnf_transformation,[],[f249_D])).
fof(f249_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f224,plain,(
  s('type.integer.int',sK0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | $spl2),
  inference(resolution,[],[f197,f162])).
fof(f195,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0)))))) | $spl0),
  inference(cnf_transformation,[],[f195_D])).
fof(f195_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0)))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f197,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | $spl2),
  inference(cnf_transformation,[],[f197_D])).
fof(f197_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f134,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))))))),
  inference(cnf_transformation,[],[f114])).
fof(f114,plain,(
  (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))))) & (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0)))))) | (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0)))) & ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f113])).
fof(f113,plain,(
  ? [X0,X1] : ((p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))))))) & (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))) | (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) & ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))))))),
  inference(flattening,[],[f112])).
fof(f112,plain,(
  ? [X0,X1] : ((p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))) | (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))))) & (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))) | (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) & ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))))))),
  inference(nnf_transformation,[],[f106])).
fof(f106,plain,(
  ? [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))) <~> (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))))))))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ~! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))) <=> (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X6)))))) <=> (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X7),s('type.integer.int',X6)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X6),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X7))))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X6)))))) <=> (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X7),s('type.integer.int',X6)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X6),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X7))))))))),
  file('thm.integer.INT_ABS_LE|split|1','thm.integer.INT_ABS_LE|split|1')).
fof(f196,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0)))))) | $spl1),
  inference(cnf_transformation,[],[f196_D])).
fof(f196_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0)))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f135,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int',sK0)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))))))),
  inference(cnf_transformation,[],[f114])).
fof(f198,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | $spl3),
  inference(cnf_transformation,[],[f198_D])).
fof(f198_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f136,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK1)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK0))))))),
  inference(cnf_transformation,[],[f114])).
% SZS output end Proof for thm.integer
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 520
% Passive clauses: 8417
% Generated clauses: 33431
% Final active clauses: 34
% Final passive clauses: 2497
% Input formulas: 69
% Initial clauses: 59
% 
% Selected by SInE selection: 32
% SInE iterations: 1
% 
% Duplicate literals: 882
% Trivial inequalities: 258
% Fw subsumption resolutions: 2049
% Fw demodulations: 7870
% Bw demodulations: 307
% 
% Simple tautologies: 1026
% Equational tautologies: 53
% Forward subsumptions: 12419
% Fw demodulations to eq. taut.: 5
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 3806
% Factoring: 8
% Forward superposition: 8432
% Backward superposition: 7018
% Self superposition: 50
% Equality factoring: 32
% Equality resolution: 8
% 
% Split clauses: 332
% Split components: 672
% SAT solver clauses: 495
% SAT solver unit clauses: 30
% SAT solver binary clauses: 50
% 
% Sat splits: 332
% Sat splitting refutations: 162
% 
% Memory used [KB]: 8699
% Time elapsed: 0.440 s
% ------------------------------
% ------------------------------
% Success in time 12.171 s
