Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.prelim
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.prelim
% SZS output start Proof for thm.prelim
fof(f474,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f473,f147,f182,f185,f124,f183])).
fof(f183,plain,(
  ~p(s(bool,'F_0')) | $spl1),
  inference(cnf_transformation,[],[f183_D])).
fof(f183_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f124,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f56])).
fof(f56,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.prelim.ordering_distinct|split|1',aHL_FALSITY)).
fof(f185,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))) ) | $spl4),
  inference(cnf_transformation,[],[f185_D])).
fof(f185_D,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f182,plain,(
  p(s(bool,'F_0')) | $spl0),
  inference(cnf_transformation,[],[f182_D])).
fof(f182_D,plain,(
  p(s(bool,'F_0')) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f147,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0] : ((s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))),
  inference(nnf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X6] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6))) <=> p(s(bool,'F_0')))),
  file('thm.prelim.ordering_distinct|split|1','thm.numeral.numeral_eq|split|3')).
fof(f473,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f463,f185])).
fof(f463,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(trivial_inequality_removal,[],[f462])).
fof(f462,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.|48|_0') | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(superposition,[],[f155,f380])).
fof(f380,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(forward_demodulation,[],[f379,f119])).
fof(f119,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.LESS_0')))),
  inference(cnf_transformation,[],[f51])).
fof(f51,axiom,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.LESS_0')))),
  file('thm.prelim.ordering_distinct|split|1','thm.prelim.ordering2num_thm|split|0')).
fof(f379,plain,(
  s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.LESS_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(forward_demodulation,[],[f123,f116])).
fof(f116,plain,(
  s('type.prelim.ordering','const.prelim.LESS_0') = s('type.prelim.ordering','const.prelim.GREATER_0')),
  inference(cnf_transformation,[],[f55])).
fof(f55,plain,(
  s('type.prelim.ordering','const.prelim.LESS_0') = s('type.prelim.ordering','const.prelim.GREATER_0')),
  inference(flattening,[],[f6])).
fof(f6,negated_conjecture,(
  ~~s('type.prelim.ordering','const.prelim.LESS_0') = s('type.prelim.ordering','const.prelim.GREATER_0')),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ~s('type.prelim.ordering','const.prelim.LESS_0') = s('type.prelim.ordering','const.prelim.GREATER_0')),
  file('thm.prelim.ordering_distinct|split|1','thm.prelim.ordering_distinct|split|1')).
fof(f123,plain,(
  s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(cnf_transformation,[],[f53])).
fof(f53,axiom,(
  s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  file('thm.prelim.ordering_distinct|split|1','thm.prelim.ordering2num_thm|split|2')).
fof(f155,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f106])).
fof(f106,plain,(
  ! [X0] : ((s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0)) & (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))),
  inference(nnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) <=> s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X6] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X6))) <=> s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X6))),
  file('thm.prelim.ordering_distinct|split|1','thm.numeral.numeral_distrib|split|17')).
% SZS output end Proof for thm.prelim
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 38
% Passive clauses: 90
% Generated clauses: 356
% Final active clauses: 30
% Final passive clauses: 44
% Input formulas: 53
% Initial clauses: 66
% 
% Duplicate literals: 1
% Trivial inequalities: 1
% Fw subsumption resolutions: 6
% Fw demodulations: 14
% Bw demodulations: 5
% 
% Simple tautologies: 2
% Equational tautologies: 2
% Forward subsumptions: 225
% 
% Binary resolution: 1
% Forward superposition: 131
% Backward superposition: 106
% Self superposition: 10
% Equality factoring: 2
% 
% Split clauses: 15
% Split components: 30
% SAT solver clauses: 20
% SAT solver unit clauses: 3
% SAT solver binary clauses: 16
% SAT solver learnt clauses: 1
% 
% Sat splits: 15
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 19
% TWLsolver calls for satisfiability: 6
% 
% Memory used [KB]: 767
% Time elapsed: 0.017 s
% ------------------------------
% ------------------------------
% Success in time 0.02 s
