Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.list
% (32068)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 132
% Passive clauses: 5088
% Generated clauses: 18851
% Final active clauses: 123
% Final passive clauses: 4755
% Input formulas: 37
% Initial clauses: 95
% 
% Duplicate literals: 805
% Trivial inequalities: 329
% Fw subsumption resolutions: 2977
% Fw demodulations: 1016
% Bw demodulations: 5
% 
% Simple tautologies: 461
% Equational tautologies: 83
% Forward subsumptions: 6655
% Bw demodulations to eq. taut.: 4
% 
% Binary resolution: 831
% Factoring: 4
% Forward superposition: 6192
% Backward superposition: 5931
% Self superposition: 239
% Equality factoring: 201
% Equality resolution: 27
% 
% Split clauses: 690
% Split components: 1414
% SAT solver clauses: 1638
% SAT solver unit clauses: 38
% SAT solver binary clauses: 435
% 
% Sat splits: 690
% 
% TWLsolver clauses: 1638
% TWLsolver calls for satisfiability: 253
% 
% Memory used [KB]: 10490
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 144 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.list
% (32114)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 170
% Passive clauses: 1254
% Generated clauses: 17370
% Final active clauses: 106
% Final passive clauses: 411
% Input formulas: 37
% Initial clauses: 95
% 
% Duplicate literals: 427
% Trivial inequalities: 393
% Fw subsumption resolutions: 2045
% Fw demodulations: 2595
% Bw demodulations: 681
% Condensations: 787
% Global subsumptions: 13
% 
% Simple tautologies: 326
% Equational tautologies: 31
% Forward subsumptions: 7965
% Bw demodulations to eq. taut.: 38
% 
% Binary resolution: 338
% Factoring: 1
% Forward superposition: 4910
% Backward superposition: 3347
% Self superposition: 118
% Equality factoring: 65
% Equality resolution: 28
% 
% Split clauses: 71
% Split components: 143
% SAT solver clauses: 1449
% SAT solver unit clauses: 95
% SAT solver binary clauses: 68
% SAT solver learnt clauses: 26
% 
% Sat splits: 71
% Sat splitting refutations: 19
% 
% TWLsolver clauses: 1397
% TWLsolver calls for satisfiability: 439
% 
% Memory used [KB]: 12025
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.list
% (32151)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 851
% Passive clauses: 8312
% Generated clauses: 248046
% Final active clauses: 218
% Final passive clauses: 1301
% Input formulas: 37
% Initial clauses: 95
% 
% Duplicate literals: 7242
% Trivial inequalities: 2011
% Fw subsumption resolutions: 25937
% Fw demodulations: 42616
% Bw demodulations: 4903
% Condensations: 13731
% Global subsumptions: 693
% 
% Simple tautologies: 2136
% Equational tautologies: 1251
% Forward subsumptions: 141180
% Fw demodulations to eq. taut.: 6
% Bw demodulations to eq. taut.: 170
% 
% Unit resulting resolution: 30605
% Factoring: 8
% Forward superposition: 66473
% Backward superposition: 57432
% Self superposition: 894
% Equality factoring: 383
% Equality resolution: 158
% 
% Split clauses: 687
% Split components: 1385
% SAT solver clauses: 11256
% SAT solver unit clauses: 19
% SAT solver binary clauses: 355
% 
% Sat splits: 687
% Sat splitting refutations: 875
% 
% Memory used [KB]: 137140
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.list
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.list
% SZS output start Proof for thm.list
fof(f1017,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f144,f240,f143,f242,f142,f244,f238,f432,f1016])).
fof(f1016,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f834,f590])).
fof(f590,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK2),s(sK0,sK4))),s(sK1,sK8(sK2,sK4,sK5,sK3,sK1,sK0))))) | $spl0),
  inference(resolution,[],[f238,f147])).
fof(f147,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X5),s(X0,X4))),s(X1,sK8(X5,X4,X3,X2,X1,X0))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X5),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X3))),s('type.list.list'(X1),X2))))) )),
  inference(cnf_transformation,[],[f105])).
fof(f105,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X5),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X3))),s('type.list.list'(X1),X2)))) | (s('type.list.list'(X1),X2) = s('type.list.list'(X1),'const.list.CONS_2'(s(X1,sK8(X5,X4,X3,X2,X1,X0)),s('type.list.list'(X1),sK9(X5,X4,X3,X2,X1,X0)))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X5),s(X0,X4))),s(X1,sK8(X5,X4,X3,X2,X1,X0))))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X5),s('type.list.list'(X0),X3),s('type.list.list'(X1),sK9(X5,X4,X3,X2,X1,X0))))))) & (! [X6,X7] : (s('type.list.list'(X1),X2) != s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X6),s('type.list.list'(X1),X7))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X5),s(X0,X4))),s(X1,X6)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X5),s('type.list.list'(X0),X3),s('type.list.list'(X1),X7))))) | p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X5),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X3))),s('type.list.list'(X1),X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK8,sK9])],[f104])).
fof(f104,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X5),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X3))),s('type.list.list'(X1),X2)))) | ? [X8,X9] : (s('type.list.list'(X1),X2) = s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X8),s('type.list.list'(X1),X9))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X5),s(X0,X4))),s(X1,X8)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X5),s('type.list.list'(X0),X3),s('type.list.list'(X1),X9)))))) & (! [X6,X7] : (s('type.list.list'(X1),X2) != s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X6),s('type.list.list'(X1),X7))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X5),s(X0,X4))),s(X1,X6)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X5),s('type.list.list'(X0),X3),s('type.list.list'(X1),X7))))) | p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X5),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X3))),s('type.list.list'(X1),X2))))))),
  inference(rectify,[],[f103])).
fof(f103,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X5),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X3))),s('type.list.list'(X1),X2)))) | ? [X6,X7] : (s('type.list.list'(X1),X2) = s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X6),s('type.list.list'(X1),X7))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X5),s(X0,X4))),s(X1,X6)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X5),s('type.list.list'(X0),X3),s('type.list.list'(X1),X7)))))) & (! [X6,X7] : (s('type.list.list'(X1),X2) != s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X6),s('type.list.list'(X1),X7))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X5),s(X0,X4))),s(X1,X6)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X5),s('type.list.list'(X0),X3),s('type.list.list'(X1),X7))))) | p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X5),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X3))),s('type.list.list'(X1),X2))))))),
  inference(nnf_transformation,[],[f42])).
fof(f42,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X5),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X3))),s('type.list.list'(X1),X2)))) <=> ? [X6,X7] : (s('type.list.list'(X1),X2) = s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X6),s('type.list.list'(X1),X7))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X5),s(X0,X4))),s(X1,X6)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X5),s('type.list.list'(X0),X3),s('type.list.list'(X1),X7))))))),
  inference(flattening,[],[f41])).
fof(f41,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X5),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X3))),s('type.list.list'(X1),X2)))) <=> ? [X6,X7] : (s('type.list.list'(X1),X2) = s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X6),s('type.list.list'(X1),X7))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X5),s(X0,X4))),s(X1,X6)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X5),s('type.list.list'(X0),X3),s('type.list.list'(X1),X7))))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X14,X15] : ! [X11,X16,X17,X18] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X14,fun(X15,bool)),X18),s('type.list.list'(X14),'const.list.CONS_2'(s(X14,X17),s('type.list.list'(X14),X16))),s('type.list.list'(X15),X11)))) <=> ? [X19,X20] : (s('type.list.list'(X15),X11) = s('type.list.list'(X15),'const.list.CONS_2'(s(X15,X19),s('type.list.list'(X15),X20))) & p(s(bool,happ(s(fun(X15,bool),happ(s(fun(X14,fun(X15,bool)),X18),s(X14,X17))),s(X15,X19)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X14,fun(X15,bool)),X18),s('type.list.list'(X14),X16),s('type.list.list'(X15),X20))))))),
  file('thm.list.EVERY2_THM|split|1','thm.list.LIST_REL_CONS|49|')).
fof(f834,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK2),s(sK0,sK4))),s(sK1,sK8(sK2,sK4,sK5,sK3,sK1,sK0))))) | $spl0),
  inference(subsumption_resolution,[],[f833,f591])).
fof(f591,plain,(
  p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK2),s('type.list.list'(sK0),sK5),s('type.list.list'(sK1),sK9(sK2,sK4,sK5,sK3,sK1,sK0))))) | $spl0),
  inference(resolution,[],[f238,f148])).
fof(f148,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X5),s('type.list.list'(X0),X3),s('type.list.list'(X1),sK9(X5,X4,X3,X2,X1,X0))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X5),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X3))),s('type.list.list'(X1),X2))))) )),
  inference(cnf_transformation,[],[f105])).
fof(f833,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK2),s(sK0,sK4))),s(sK1,sK8(sK2,sK4,sK5,sK3,sK1,sK0))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK2),s('type.list.list'(sK0),sK5),s('type.list.list'(sK1),sK9(sK2,sK4,sK5,sK3,sK1,sK0))))) | $spl0),
  inference(trivial_inequality_removal,[],[f797])).
fof(f797,plain,(
  s('type.list.list'(sK1),sK3) != s('type.list.list'(sK1),sK3) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK2),s(sK0,sK4))),s(sK1,sK8(sK2,sK4,sK5,sK3,sK1,sK0))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK2),s('type.list.list'(sK0),sK5),s('type.list.list'(sK1),sK9(sK2,sK4,sK5,sK3,sK1,sK0))))) | $spl0),
  inference(superposition,[],[f236,f589])).
fof(f589,plain,(
  s('type.list.list'(sK1),sK3) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK8(sK2,sK4,sK5,sK3,sK1,sK0)),s('type.list.list'(sK1),sK9(sK2,sK4,sK5,sK3,sK1,sK0)))) | $spl0),
  inference(resolution,[],[f238,f146])).
fof(f146,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.list.list'(X1),X2) = s('type.list.list'(X1),'const.list.CONS_2'(s(X1,sK8(X5,X4,X3,X2,X1,X0)),s('type.list.list'(X1),sK9(X5,X4,X3,X2,X1,X0)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X5),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X3))),s('type.list.list'(X1),X2))))) )),
  inference(cnf_transformation,[],[f105])).
fof(f236,plain,(
  ( ! [X6,X7] : (s('type.list.list'(sK1),sK3) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,X6),s('type.list.list'(sK1),X7))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK2),s(sK0,sK4))),s(sK1,X6)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK2),s('type.list.list'(sK0),sK5),s('type.list.list'(sK1),X7))))) )),
  inference(subsumption_resolution,[],[f145,f149])).
fof(f149,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X5),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X3))),s('type.list.list'(X1),X2)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X5),s('type.list.list'(X0),X3),s('type.list.list'(X1),X7)))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X5),s(X0,X4))),s(X1,X6)))) | s('type.list.list'(X1),X2) != s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X6),s('type.list.list'(X1),X7)))) )),
  inference(cnf_transformation,[],[f105])).
fof(f145,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK2),s('type.list.list'(sK0),sK5),s('type.list.list'(sK1),X7)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK2),s(sK0,sK4))),s(sK1,X6)))) | s('type.list.list'(sK1),sK3) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,X6),s('type.list.list'(sK1),X7))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK2),s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK4),s('type.list.list'(sK0),sK5))),s('type.list.list'(sK1),sK3))))) )),
  inference(cnf_transformation,[],[f102])).
fof(f102,plain,(
  (p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK2),s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK4),s('type.list.list'(sK0),sK5))),s('type.list.list'(sK1),sK3)))) | (s('type.list.list'(sK1),sK3) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6),s('type.list.list'(sK1),sK7))) & p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK2),s(sK0,sK4))),s(sK1,sK6)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK2),s('type.list.list'(sK0),sK5),s('type.list.list'(sK1),sK7)))))) & (~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK2),s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK4),s('type.list.list'(sK0),sK5))),s('type.list.list'(sK1),sK3)))) | ! [X6,X7] : (s('type.list.list'(sK1),sK3) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,X6),s('type.list.list'(sK1),X7))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK2),s(sK0,sK4))),s(sK1,X6)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK2),s('type.list.list'(sK0),sK5),s('type.list.list'(sK1),X7))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5,sK6,sK7])],[f101])).
fof(f101,plain,(
  ? [X0,X1,X2,X3,X4,X5] : ((p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X5))),s('type.list.list'(X1),X3)))) | ? [X8,X9] : (s('type.list.list'(X1),X3) = s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X8),s('type.list.list'(X1),X9))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X8)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X5),s('type.list.list'(X1),X9)))))) & (~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X5))),s('type.list.list'(X1),X3)))) | ! [X6,X7] : (s('type.list.list'(X1),X3) != s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X6),s('type.list.list'(X1),X7))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X6)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X5),s('type.list.list'(X1),X7)))))))),
  inference(rectify,[],[f100])).
fof(f100,plain,(
  ? [X0,X1,X2,X3,X4,X5] : ((p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X5))),s('type.list.list'(X1),X3)))) | ? [X6,X7] : (s('type.list.list'(X1),X3) = s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X6),s('type.list.list'(X1),X7))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X6)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X5),s('type.list.list'(X1),X7)))))) & (~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X5))),s('type.list.list'(X1),X3)))) | ! [X6,X7] : (s('type.list.list'(X1),X3) != s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X6),s('type.list.list'(X1),X7))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X6)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X5),s('type.list.list'(X1),X7)))))))),
  inference(nnf_transformation,[],[f84])).
fof(f84,plain,(
  ? [X0,X1,X2,X3,X4,X5] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X5))),s('type.list.list'(X1),X3)))) <~> ? [X6,X7] : (s('type.list.list'(X1),X3) = s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X6),s('type.list.list'(X1),X7))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X6)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X5),s('type.list.list'(X1),X7))))))),
  inference(ennf_transformation,[],[f40])).
fof(f40,plain,(
  ~! [X0,X1,X2,X3,X4,X5] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X5))),s('type.list.list'(X1),X3)))) <=> ? [X6,X7] : (s('type.list.list'(X1),X3) = s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X6),s('type.list.list'(X1),X7))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X6)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X5),s('type.list.list'(X1),X7))))))),
  inference(flattening,[],[f39])).
fof(f39,plain,(
  ~! [X0,X1] : ! [X2,X3,X4,X5] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X5))),s('type.list.list'(X1),X3)))) <=> ? [X6,X7] : (s('type.list.list'(X1),X3) = s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X6),s('type.list.list'(X1),X7))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X6)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X5),s('type.list.list'(X1),X7))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10,X11] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X6,fun(X7,bool)),X8),s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X10),s('type.list.list'(X6),X11))),s('type.list.list'(X7),X9)))) <=> ? [X12,X13] : (s('type.list.list'(X7),X9) = s('type.list.list'(X7),'const.list.CONS_2'(s(X7,X12),s('type.list.list'(X7),X13))) & p(s(bool,happ(s(fun(X7,bool),happ(s(fun(X6,fun(X7,bool)),X8),s(X6,X10))),s(X7,X12)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X6,fun(X7,bool)),X8),s('type.list.list'(X6),X11),s('type.list.list'(X7),X13))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10,X11] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X6,fun(X7,bool)),X8),s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X10),s('type.list.list'(X6),X11))),s('type.list.list'(X7),X9)))) <=> ? [X12,X13] : (s('type.list.list'(X7),X9) = s('type.list.list'(X7),'const.list.CONS_2'(s(X7,X12),s('type.list.list'(X7),X13))) & p(s(bool,happ(s(fun(X7,bool),happ(s(fun(X6,fun(X7,bool)),X8),s(X6,X10))),s(X7,X12)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X6,fun(X7,bool)),X8),s('type.list.list'(X6),X11),s('type.list.list'(X7),X13))))))),
  file('thm.list.EVERY2_THM|split|1','thm.list.EVERY2_THM|split|1')).
fof(f432,plain,(
  $false | ($spl2 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f431,f240])).
fof(f431,plain,(
  ~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK2),s('type.list.list'(sK0),sK5),s('type.list.list'(sK1),sK7)))) | ($spl4 | $spl6)),
  inference(subsumption_resolution,[],[f428,f242])).
fof(f428,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK2),s(sK0,sK4))),s(sK1,sK6)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK2),s('type.list.list'(sK0),sK5),s('type.list.list'(sK1),sK7)))) | $spl6),
  inference(trivial_inequality_removal,[],[f426])).
fof(f426,plain,(
  s('type.list.list'(sK1),sK3) != s('type.list.list'(sK1),sK3) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK2),s(sK0,sK4))),s(sK1,sK6)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK2),s('type.list.list'(sK0),sK5),s('type.list.list'(sK1),sK7)))) | $spl6),
  inference(superposition,[],[f236,f244])).
fof(f238,plain,(
  p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK2),s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK4),s('type.list.list'(sK0),sK5))),s('type.list.list'(sK1),sK3)))) | $spl0),
  inference(cnf_transformation,[],[f238_D])).
fof(f238_D,plain,(
  p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK2),s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK4),s('type.list.list'(sK0),sK5))),s('type.list.list'(sK1),sK3)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f244,plain,(
  s('type.list.list'(sK1),sK3) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6),s('type.list.list'(sK1),sK7))) | $spl6),
  inference(cnf_transformation,[],[f244_D])).
fof(f244_D,plain,(
  s('type.list.list'(sK1),sK3) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6),s('type.list.list'(sK1),sK7))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f142,plain,(
  s('type.list.list'(sK1),sK3) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6),s('type.list.list'(sK1),sK7))) | p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK2),s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK4),s('type.list.list'(sK0),sK5))),s('type.list.list'(sK1),sK3))))),
  inference(cnf_transformation,[],[f102])).
fof(f242,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK2),s(sK0,sK4))),s(sK1,sK6)))) | $spl4),
  inference(cnf_transformation,[],[f242_D])).
fof(f242_D,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK2),s(sK0,sK4))),s(sK1,sK6)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f143,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK2),s(sK0,sK4))),s(sK1,sK6)))) | p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK2),s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK4),s('type.list.list'(sK0),sK5))),s('type.list.list'(sK1),sK3))))),
  inference(cnf_transformation,[],[f102])).
fof(f240,plain,(
  p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK2),s('type.list.list'(sK0),sK5),s('type.list.list'(sK1),sK7)))) | $spl2),
  inference(cnf_transformation,[],[f240_D])).
fof(f240_D,plain,(
  p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK2),s('type.list.list'(sK0),sK5),s('type.list.list'(sK1),sK7)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f144,plain,(
  p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK2),s('type.list.list'(sK0),sK5),s('type.list.list'(sK1),sK7)))) | p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK2),s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK4),s('type.list.list'(sK0),sK5))),s('type.list.list'(sK1),sK3))))),
  inference(cnf_transformation,[],[f102])).
% SZS output end Proof for thm.list
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 62
% Passive clauses: 148
% Generated clauses: 782
% Final active clauses: 54
% Final passive clauses: 100
% Input formulas: 37
% Initial clauses: 92
% 
% Selected by SInE selection: 35
% SInE iterations: 1
% 
% Duplicate literals: 21
% Trivial inequalities: 26
% Fw subsumption resolutions: 78
% Fw demodulations: 143
% Bw demodulations: 5
% 
% Simple tautologies: 74
% Equational tautologies: 3
% Forward subsumptions: 350
% 
% Binary resolution: 91
% Forward superposition: 155
% Backward superposition: 253
% Equality resolution: 1
% 
% Split clauses: 3
% Split components: 6
% SAT solver clauses: 6
% SAT solver unit clauses: 1
% SAT solver binary clauses: 3
% 
% Sat splits: 3
% Sat splitting refutations: 2
% 
% Memory used [KB]: 1535
% Time elapsed: 0.021 s
% ------------------------------
% ------------------------------
% Success in time 11.75 s
