Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.arithmetic
% (7552)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 777
% Passive clauses: 8747
% Generated clauses: 28168
% Final active clauses: 709
% Final passive clauses: 7539
% Input formulas: 133
% Initial clauses: 369
% 
% Duplicate literals: 250
% Trivial inequalities: 91
% Fw subsumption resolutions: 1104
% Fw demodulations: 3224
% Bw demodulations: 63
% 
% Simple tautologies: 281
% Equational tautologies: 64
% Forward subsumptions: 14197
% Fw demodulations to eq. taut.: 6
% 
% Binary resolution: 10230
% Factoring: 1022
% Forward superposition: 5441
% Backward superposition: 6154
% Self superposition: 47
% Equality factoring: 2
% Equality resolution: 22
% 
% Split clauses: 260
% Split components: 521
% SAT solver clauses: 475
% SAT solver unit clauses: 12
% SAT solver binary clauses: 81
% 
% Sat splits: 260
% Sat splitting refutations: 7
% 
% TWLsolver clauses: 472
% TWLsolver calls for satisfiability: 319
% 
% Memory used [KB]: 15735
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.arithmetic
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.arithmetic
% SZS output start Proof for thm.arithmetic
fof(f16911,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f16830,f553,f937,f11968,f11969,f665,f997,f696,f1021,f4186,f4220,f601,f972,f4157,f4158,f554,f939,f563,f948,f754,f1055,f1196,f1198,f562,f947,f597,f970])).
fof(f970,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl52),
  inference(cnf_transformation,[],[f970_D])).
fof(f970_D,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f597,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f393])).
fof(f393,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f172])).
fof(f172,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X10] : (s(bool,'T_0') = s(bool,X10) <=> p(s(bool,X10)))),
  file('thm.arithmetic.MAX_|48||split|0','thm.bool.EQ_CLAUSES|split|0')).
fof(f947,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))) ) | $spl8),
  inference(cnf_transformation,[],[f947_D])).
fof(f947_D,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f562,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f140])).
fof(f140,plain,(
  ! [X0] : p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))),
  inference(rectify,[],[f120])).
fof(f120,axiom,(
  ! [X6] : p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X6))))),
  file('thm.arithmetic.MAX_|48||split|0','thm.arithmetic.ZERO_LESS_EQ')).
fof(f1198,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) ) | $spl454),
  inference(cnf_transformation,[],[f1198_D])).
fof(f1198_D,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) ) <=> ~$spl454),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl454])])).
fof(f1196,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) ) | ($spl8 | $spl52)),
  inference(resolution,[],[f970,f947])).
fof(f1055,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))))) ) | $spl188),
  inference(cnf_transformation,[],[f1055_D])).
fof(f1055_D,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))))) ) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f754,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f473])).
fof(f473,plain,(
  ! [X0,X1] : ((p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))))) & (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(nnf_transformation,[],[f262])).
fof(f262,plain,(
  ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(flattening,[],[f261])).
fof(f261,plain,(
  ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(rectify,[],[f106])).
fof(f106,axiom,(
  ! [X7,X6] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num',X6)))) <=> p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num',X7)))))),
  file('thm.arithmetic.MAX_|48||split|0','thm.arithmetic.NOT_LESS_EQUAL')).
fof(f948,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) ) | $spl10),
  inference(cnf_transformation,[],[f948_D])).
fof(f948_D,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f563,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f141])).
fof(f141,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f87])).
fof(f87,axiom,(
  ! [X7] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X7)))),
  file('thm.arithmetic.MAX_|48||split|0','thm.arithmetic.MULT_CLAUSES|split|0')).
fof(f939,plain,(
  p(s(bool,'T_0')) | $spl2),
  inference(cnf_transformation,[],[f939_D])).
fof(f939_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f554,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.arithmetic.MAX_|48||split|0',aHL_TRUTH)).
fof(f4158,plain,(
  ( ! [X7] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num','const.num.|48|_0'))))) ) | $spl1346),
  inference(cnf_transformation,[],[f4158_D])).
fof(f4158_D,plain,(
  ( ! [X7] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num','const.num.|48|_0'))))) ) <=> ~$spl1346),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1346])])).
fof(f4157,plain,(
  ( ! [X7] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num','const.num.|48|_0'))))) ) | ($spl2 | $spl10 | $spl188 | $spl454)),
  inference(forward_demodulation,[],[f4156,f948])).
fof(f4156,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X6))))))) ) | ($spl2 | $spl10 | $spl188 | $spl454)),
  inference(subsumption_resolution,[],[f4155,f939])).
fof(f4155,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X6))))))) ) | ($spl10 | $spl188 | $spl454)),
  inference(forward_demodulation,[],[f4112,f1198])).
fof(f4112,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X7)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X6))))))) ) | ($spl10 | $spl188)),
  inference(superposition,[],[f1055,f948])).
fof(f972,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) | $spl56),
  inference(cnf_transformation,[],[f972_D])).
fof(f972_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f601,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f395])).
fof(f395,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f175])).
fof(f175,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f174])).
fof(f174,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X10] : (s(bool,'F_0') = s(bool,X10) <=> ~p(s(bool,X10)))),
  file('thm.arithmetic.MAX_|48||split|0','thm.bool.EQ_CLAUSES|split|3')).
fof(f4220,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) ) | $spl1350),
  inference(cnf_transformation,[],[f4220_D])).
fof(f4220_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl1350),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1350])])).
fof(f4186,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) ) | ($spl56 | $spl1346)),
  inference(resolution,[],[f4158,f972])).
fof(f1021,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X1),s('type.num.num',X0)))) ) | $spl130),
  inference(cnf_transformation,[],[f1021_D])).
fof(f1021_D,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X1),s('type.num.num',X0)))) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f696,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X1),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f234])).
fof(f234,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X1),s('type.num.num',X0)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X7,X6] : s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X7),s('type.num.num',X6))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num',X6))),s('type.num.num',X6),s('type.num.num',X7)))),
  file('thm.arithmetic.MAX_|48||split|0','thm.arithmetic.MAX_DEF')).
fof(f997,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))) ) | $spl82),
  inference(cnf_transformation,[],[f997_D])).
fof(f997_D,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f665,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f213])).
fof(f213,plain,(
  ! [X0,X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f212])).
fof(f212,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f70])).
fof(f70,axiom,(
  ! [X13] : ! [X11,X12] : s(X13,X12) = s(X13,'const.bool.COND_3'(s(bool,'F_0'),s(X13,X11),s(X13,X12)))),
  file('thm.arithmetic.MAX_|48||split|0','thm.bool.COND_CLAUSES|split|1')).
fof(f11969,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) ) | $spl3050),
  inference(cnf_transformation,[],[f11969_D])).
fof(f11969_D,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl3050),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3050])])).
fof(f11968,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) ) | ($spl82 | $spl130 | $spl1350)),
  inference(forward_demodulation,[],[f11896,f997])).
fof(f11896,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) ) | ($spl130 | $spl1350)),
  inference(superposition,[],[f1021,f4220])).
fof(f937,plain,(
  s('type.num.num',sK13) != s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | $spl1),
  inference(cnf_transformation,[],[f937_D])).
fof(f937_D,plain,(
  s('type.num.num',sK13) != s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f553,plain,(
  s('type.num.num',sK13) != s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')))),
  inference(cnf_transformation,[],[f388])).
fof(f388,plain,(
  s('type.num.num',sK13) != s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13])],[f309])).
fof(f309,plain,(
  ? [X0] : s('type.num.num',X0) != s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(ennf_transformation,[],[f135])).
fof(f135,plain,(
  ~! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : s('type.num.num',X6) = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0')))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : s('type.num.num',X6) = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0')))),
  file('thm.arithmetic.MAX_|48||split|0','thm.arithmetic.MAX_|48||split|0')).
fof(f16830,plain,(
  $false | ($spl1 | $spl3050)),
  inference(trivial_inequality_removal,[],[f16631])).
fof(f16631,plain,(
  s('type.num.num',sK13) != s('type.num.num',sK13) | ($spl1 | $spl3050)),
  inference(superposition,[],[f937,f11969])).
% SZS output end Proof for thm.arithmetic
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 439
% Passive clauses: 2686
% Generated clauses: 16430
% Final active clauses: 433
% Final passive clauses: 2142
% Input formulas: 133
% Initial clauses: 369
% 
% Duplicate literals: 116
% Trivial inequalities: 31
% Fw subsumption resolutions: 693
% Fw demodulations: 2425
% Bw demodulations: 109
% Condensations: 134
% Global subsumptions: 22
% 
% Simple tautologies: 257
% Equational tautologies: 20
% Forward subsumptions: 7007
% Fw demodulations to eq. taut.: 9
% 
% Binary resolution: 3541
% Factoring: 328
% Forward superposition: 3782
% Backward superposition: 2125
% Self superposition: 42
% Equality factoring: 2
% Equality resolution: 15
% 
% Split clauses: 3
% Split components: 6
% SAT solver clauses: 3389
% SAT solver unit clauses: 377
% SAT solver binary clauses: 127
% SAT solver learnt clauses: 45
% SAT solver learnt literals: 10
% 
% Sat splits: 3
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 3300
% TWLsolver calls for satisfiability: 2814
% 
% Memory used [KB]: 15223
% Time elapsed: 0.574 s
% ------------------------------
% ------------------------------
% Success in time 1.2 s
