Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.extreal
% (31999)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 870
% Passive clauses: 3679
% Generated clauses: 20068
% Final active clauses: 850
% Final passive clauses: 2613
% Input formulas: 261
% Initial clauses: 517
% 
% Duplicate literals: 160
% Trivial inequalities: 46
% Fw subsumption resolutions: 1294
% Fw demodulations: 1197
% Bw demodulations: 19
% 
% Simple tautologies: 262
% Equational tautologies: 18
% Forward subsumptions: 13250
% 
% Binary resolution: 3805
% Factoring: 363
% Forward superposition: 6865
% Backward superposition: 5859
% Self superposition: 26
% Equality factoring: 2
% Equality resolution: 9
% 
% Split clauses: 258
% Split components: 518
% SAT solver clauses: 364
% SAT solver unit clauses: 21
% SAT solver binary clauses: 78
% 
% Sat splits: 258
% 
% TWLsolver clauses: 364
% TWLsolver calls for satisfiability: 285
% 
% Memory used [KB]: 9978
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.extreal
% (32011)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 557
% Passive clauses: 2026
% Generated clauses: 14997
% Final active clauses: 542
% Final passive clauses: 1437
% Input formulas: 261
% Initial clauses: 517
% 
% Duplicate literals: 95
% Trivial inequalities: 26
% Fw subsumption resolutions: 595
% Fw demodulations: 1299
% Bw demodulations: 46
% Condensations: 48
% Global subsumptions: 34
% 
% Simple tautologies: 256
% Equational tautologies: 16
% Forward subsumptions: 8123
% 
% Binary resolution: 1818
% Factoring: 93
% Forward superposition: 4553
% Backward superposition: 3767
% Self superposition: 45
% Equality factoring: 2
% Equality resolution: 15
% 
% Split clauses: 9
% Split components: 18
% SAT solver clauses: 2889
% SAT solver unit clauses: 703
% SAT solver binary clauses: 159
% SAT solver learnt clauses: 69
% SAT solver learnt literals: 11
% 
% Sat splits: 9
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 2752
% TWLsolver calls for satisfiability: 4886
% 
% Memory used [KB]: 14583
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.extreal
% (32031)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 921
% Passive clauses: 8251
% Generated clauses: 116341
% Final active clauses: 692
% Final passive clauses: 2036
% Input formulas: 261
% Initial clauses: 517
% Discarded non-redundant clauses: 4484
% 
% Duplicate literals: 10
% Trivial inequalities: 59
% Fw subsumption resolutions: 1322
% Fw demodulations: 40479
% Bw demodulations: 546
% Condensations: 35
% Global subsumptions: 1380
% 
% Simple tautologies: 106
% Equational tautologies: 87
% Forward subsumptions: 37480
% Fw demodulations to eq. taut.: 557
% Bw demodulations to eq. taut.: 14
% 
% Unit resulting resolution: 54822
% Factoring: 1
% Forward superposition: 10071
% Backward superposition: 6931
% Self superposition: 79
% Equality factoring: 2
% Equality resolution: 22
% 
% Split clauses: 39
% Split components: 78
% SAT solver clauses: 12525
% SAT solver unit clauses: 5315
% SAT solver binary clauses: 4609
% 
% Sat splits: 39
% Sat splitting refutations: 2813
% 
% Memory used [KB]: 177054
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 22 next slice time: 22
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.extreal
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.extreal
% SZS output start Proof for thm.extreal
fof(f6158,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f566,f835,f833,f565,f834,f832,f898,f908,f907,f906,f913,f972,f984,f988,f989,f990,f991,f992,f994,f1115,f995,f1083,f1113,f1123,f1151,f1153,f1155,f1156,f1157,f1152,f1158,f1148,f1154,f1178,f1180,f1182,f1183,f1184,f1179,f1185,f1175,f1181,f1187,f1188,f1189,f1190,f1191,f1316,f1192,f1194,f1285,f1315,f1322,f1325,f1328,f1329,f1330,f1324,f1331,f1294,f1327,f1492,f1445,f1546,f1547,f1548,f1549,f1550,f1551,f1571,f1573,f1575,f1576,f1577,f1572,f1578,f1568,f1574,f1598,f1600,f1602,f1603,f1604,f1599,f1605,f1595,f1601,f1626,f1628,f1630,f1631,f1632,f1627,f1633,f1622,f1629,f1654,f1656,f1658,f1659,f1660,f1655,f1661,f1650,f1657,f1682,f1684,f1686,f1687,f1688,f1683,f1689,f1678,f1685,f1709,f1711,f1713,f1714,f1715,f1710,f1716,f1707,f1712,f1736,f1738,f1740,f1741,f1742,f1737,f1743,f1734,f1739,f1763,f1765,f1767,f1768,f1769,f1764,f1770,f1761,f1766,f1791,f1793,f1795,f1796,f1797,f1792,f1798,f1788,f1794,f1160,f1801,f1802,f1803,f1804,f1805,f1935,f1806,f1808,f1899,f1934,f1944,f2029,f2035,f2052,f2056,f2060,f2031,f2102,f2121,f2123,f2125,f2126,f2147,f2149,f2151,f2152,f2172,f2174,f2176,f2177,f2198,f2200,f915,f2202,f2203,f2278,f2284,f2301,f2305,f2309,f2280,f2351,f2353,f2374,f2376,f2398,f2400,f2421,f2423,f2444,f2446,f917,f2522,f2528,f2545,f2549,f2553,f2524,f2577,f2578,f2579,f2580,f1114,f2581,f2583,f2708,f2584,f1117,f2672,f2707,f2812,f2811,f1119,f3003,f3002,f1118,f3170,f3147,f3146,f3144,f3142,f3135,f3134,f3133,f3238,f3239,f3240,f3241,f3242,f3243,f3248,f3249,f3400,f3377,f3376,f3374,f3372,f3365,f3364,f3363,f3468,f3469,f3470,f3471,f3472,f3473,f3478,f3479,f3519,f3520,f3521,f3522,f4266,f4267,f4271,f4283,f4284,f4288,f3006,f4426,f4427,f4849,f4850,f4867,f4868,f5099,f5100,f5643,f5644,f5661,f5662,f2815,f1116,f5924,f5925,f5926,f5927,f5928,f5929,f5930,f5931,f5932,f5933,f5934,f5935,f5936,f5938,f5939,f5941,f5944,f5945,f5946,f5947,f5948,f5949,f5950,f5951,f5952,f5953,f5954,f5955,f5956,f5957,f5958,f5959,f5960,f5961,f5962,f5963,f5964,f5965,f5966,f5967,f5968,f5969,f5970,f5971,f5972,f5973,f5974,f5975,f5976,f5977,f5979,f5980,f5982,f5983,f5984,f5985,f5986,f5987,f5989,f5990,f5992,f5993,f5994,f5995,f5996,f5997,f5998,f5999,f6000,f6001,f6002,f6003,f6004,f6005,f6006,f6007,f6008,f6009,f6010,f6011,f6012,f6013,f5940,f6015,f6018,f6019,f6020,f6021,f6022,f6023,f6024,f6025,f6026,f6027,f6028,f6029,f6030,f6031,f6032,f6033,f6034,f6035,f6036,f6037,f6038,f6039,f6040,f6041,f6042,f6043,f6044,f6045,f6046,f6047,f6048,f6049,f6050,f6051,f6052,f6053,f6054,f6055,f6056,f6057,f6058,f6059,f6060,f6061,f6062,f6063,f6064,f6065,f6066,f6067,f6068,f6069,f6070,f6071,f6072,f6073,f6074,f6075,f6076,f6077,f6078,f6079,f6080,f6081,f6082,f6083,f6084,f6085,f6086,f6087,f6088,f6089,f6090,f6091,f6092,f6093,f6094,f6095,f6096,f6097,f6098,f6099,f6100,f6101,f6102,f6103,f6104,f6105,f6106,f6107,f6108,f6109,f6110,f6111,f6112,f6113,f6114,f6116,f6117,f1326,f6118,f2447,f6119,f2445,f6120,f2424,f6121,f2422,f6122,f2401,f6123,f2399,f6124,f2377,f6125,f2375,f6126,f2354,f6127,f2352,f6128,f2201,f6129,f2199,f6130,f2175,f6131,f2173,f6132,f2150,f6133,f2148,f6134,f2124,f6135,f2122,f911,f909,f6156,f6157])).
fof(f6157,plain,(
  $false | ($spl4 | $spl10)),
  inference(trivial_inequality_removal,[],[f6146])).
fof(f6146,plain,(
  s('type.extreal.extreal',sK12) != s('type.extreal.extreal',sK12) | ($spl4 | $spl10)),
  inference(superposition,[],[f909,f915])).
fof(f6156,plain,(
  $false | ($spl4 | $spl12)),
  inference(trivial_inequality_removal,[],[f6148])).
fof(f6148,plain,(
  s('type.extreal.extreal',sK12) != s('type.extreal.extreal',sK12) | ($spl4 | $spl12)),
  inference(superposition,[],[f909,f917])).
fof(f909,plain,(
  ( ! [X0] : (s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))) ) | $spl4),
  inference(cnf_transformation,[],[f909_D])).
fof(f909_D,plain,(
  ( ! [X0] : (s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f911,plain,(
  s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | $spl6),
  inference(cnf_transformation,[],[f911_D])).
fof(f911_D,plain,(
  s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f2122,plain,(
  ( ! [X6,X7] : (sP2(X6,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X7)) ) | $spl88),
  inference(cnf_transformation,[],[f2122_D])).
fof(f2122_D,plain,(
  ( ! [X6,X7] : (sP2(X6,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X7)) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f6135,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X1) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f2107,f912])).
fof(f912,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',sK12) | $spl9),
  inference(cnf_transformation,[],[f912_D])).
fof(f912_D,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',sK12) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f2107,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X1) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | $spl0),
  inference(superposition,[],[f1199,f617])).
fof(f617,plain,(
  ( ! [X2,X0] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(X0)))) | s('type.extreal.extreal',X0) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) )),
  inference(cnf_transformation,[],[f454])).
fof(f454,plain,(
  ! [X0] : ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0) & ! [X2] : s('type.extreal.extreal',X0) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) | s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(X0)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK15])],[f453])).
fof(f453,plain,(
  ! [X0] : ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0) & ! [X2] : s('type.extreal.extreal',X0) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) | ? [X1] : s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0))),
  inference(rectify,[],[f379])).
fof(f379,plain,(
  ! [X0] : ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0) & ! [X1] : s('type.extreal.extreal',X0) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) | ? [X2] : s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0))),
  inference(flattening,[],[f378])).
fof(f378,plain,(
  ! [X0] : ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0) & ! [X1] : s('type.extreal.extreal',X0) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) | (? [X2] : s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)))),
  inference(ennf_transformation,[],[f285])).
fof(f285,plain,(
  ! [X0] : ((s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | ? [X1] : s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) => (? [X2] : s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)))),
  inference(rectify,[],[f55])).
fof(f55,axiom,(
  ! [X11] : ((s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X11) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X11) | ? [X6] : s('type.extreal.extreal',X11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X6)))) => (? [X6] : s('type.extreal.extreal',X11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X6))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X11) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X11)))),
  file('thm.extreal.le_infty|split|3','thm.extreal.extreal_46_prove_ltb')).
fof(f1199,plain,(
  ( ! [X33,X32] : (sP2(X32,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0')),X33)) ) | $spl0),
  inference(resolution,[],[f987,f584])).
fof(f584,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f442])).
fof(f442,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(rectify,[],[f441])).
fof(f441,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(flattening,[],[f440])).
fof(f440,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(nnf_transformation,[],[f417])).
fof(f417,plain,(
  ! [X2,X0,X1] : (sP2(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f987,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl0),
  inference(resolution,[],[f832,f806])).
fof(f806,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))) )),
  inference(cnf_transformation,[],[f563])).
fof(f563,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0))))) & (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))))),
  inference(nnf_transformation,[],[f363])).
fof(f363,plain,(
  ! [X0,X1] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) <=> ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))))),
  inference(flattening,[],[f362])).
fof(f362,plain,(
  ! [X0,X1] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) <=> ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X6,X12] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X12)))) <=> ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X12),s('type.extreal.extreal',X6)))))),
  file('thm.extreal.le_infty|split|3','thm.extreal.extreal_lt_def')).
fof(f2124,plain,(
  ( ! [X8,X9] : (sP2(X8,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X9)) ) | $spl90),
  inference(cnf_transformation,[],[f2124_D])).
fof(f2124_D,plain,(
  ( ! [X8,X9] : (sP2(X8,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X9)) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f6134,plain,(
  ( ! [X4,X5,X3] : (sP2(X3,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X4) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f2108,f835])).
fof(f2108,plain,(
  ( ! [X4,X5,X3] : (sP2(X3,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X4) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl0),
  inference(superposition,[],[f1199,f617])).
fof(f2148,plain,(
  ( ! [X6,X7] : (sP2('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X6,X7)) ) | $spl92),
  inference(cnf_transformation,[],[f2148_D])).
fof(f2148_D,plain,(
  ( ! [X6,X7] : (sP2('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X6,X7)) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f6133,plain,(
  ( ! [X2,X0,X1] : (sP2('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0,X1) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f2131,f912])).
fof(f2131,plain,(
  ( ! [X2,X0,X1] : (sP2('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0,X1) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | $spl0),
  inference(superposition,[],[f1200,f617])).
fof(f1200,plain,(
  ( ! [X35,X34] : (sP2('const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0')),X34,X35)) ) | $spl0),
  inference(resolution,[],[f987,f585])).
fof(f585,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f442])).
fof(f2150,plain,(
  ( ! [X8,X9] : (sP2('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X8,X9)) ) | $spl94),
  inference(cnf_transformation,[],[f2150_D])).
fof(f2150_D,plain,(
  ( ! [X8,X9] : (sP2('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X8,X9)) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f6132,plain,(
  ( ! [X4,X5,X3] : (sP2('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X3,X4) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f2132,f835])).
fof(f2132,plain,(
  ( ! [X4,X5,X3] : (sP2('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X3,X4) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl0),
  inference(superposition,[],[f1200,f617])).
fof(f2173,plain,(
  ( ! [X6,X7] : (sP1(X6,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X7)) ) | $spl96),
  inference(cnf_transformation,[],[f2173_D])).
fof(f2173_D,plain,(
  ( ! [X6,X7] : (sP1(X6,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X7)) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f6131,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X1) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f2157,f912])).
fof(f2157,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X1) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | $spl0),
  inference(superposition,[],[f1202,f617])).
fof(f1202,plain,(
  ( ! [X39,X38] : (sP1(X38,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0')),X39)) ) | $spl0),
  inference(resolution,[],[f987,f588])).
fof(f588,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f445])).
fof(f445,plain,(
  ! [X0,X1,X2] : ((~sP1(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1,X2)))),
  inference(rectify,[],[f444])).
fof(f444,plain,(
  ! [X1,X0,X2] : ((~sP1(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0,X2)))),
  inference(flattening,[],[f443])).
fof(f443,plain,(
  ! [X1,X0,X2] : ((~sP1(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0,X2)))),
  inference(nnf_transformation,[],[f416])).
fof(f416,plain,(
  ! [X1,X0,X2] : (sP1(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f2175,plain,(
  ( ! [X8,X9] : (sP1(X8,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X9)) ) | $spl98),
  inference(cnf_transformation,[],[f2175_D])).
fof(f2175_D,plain,(
  ( ! [X8,X9] : (sP1(X8,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X9)) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f6130,plain,(
  ( ! [X4,X5,X3] : (sP1(X3,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X4) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f2158,f835])).
fof(f2158,plain,(
  ( ! [X4,X5,X3] : (sP1(X3,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X4) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl0),
  inference(superposition,[],[f1202,f617])).
fof(f2199,plain,(
  ( ! [X6,X7] : (sP1('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X6,X7)) ) | $spl100),
  inference(cnf_transformation,[],[f2199_D])).
fof(f2199_D,plain,(
  ( ! [X6,X7] : (sP1('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X6,X7)) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f6129,plain,(
  ( ! [X2,X0,X1] : (sP1('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0,X1) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f2182,f912])).
fof(f2182,plain,(
  ( ! [X2,X0,X1] : (sP1('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0,X1) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | $spl0),
  inference(superposition,[],[f1203,f617])).
fof(f1203,plain,(
  ( ! [X41,X40] : (sP1('const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0')),X40,X41)) ) | $spl0),
  inference(resolution,[],[f987,f589])).
fof(f589,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f445])).
fof(f2201,plain,(
  ( ! [X8,X9] : (sP1('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X8,X9)) ) | $spl102),
  inference(cnf_transformation,[],[f2201_D])).
fof(f2201_D,plain,(
  ( ! [X8,X9] : (sP1('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X8,X9)) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f6128,plain,(
  ( ! [X4,X5,X3] : (sP1('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X3,X4) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f2183,f835])).
fof(f2183,plain,(
  ( ! [X4,X5,X3] : (sP1('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X3,X4) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl0),
  inference(superposition,[],[f1203,f617])).
fof(f2352,plain,(
  ( ! [X10,X11] : (sP4(X10,X11,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl106),
  inference(cnf_transformation,[],[f2352_D])).
fof(f2352_D,plain,(
  ( ! [X10,X11] : (sP4(X10,X11,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f6127,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f2337,f912])).
fof(f2337,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | $spl0),
  inference(superposition,[],[f1217,f617])).
fof(f1217,plain,(
  ( ! [X66,X67] : (sP4(X66,X67,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl0),
  inference(resolution,[],[f987,f629])).
fof(f629,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f462])).
fof(f462,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f461])).
fof(f461,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(flattening,[],[f460])).
fof(f460,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(nnf_transformation,[],[f420])).
fof(f420,plain,(
  ! [X2,X0,X1] : (sP4(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f2354,plain,(
  ( ! [X12,X13] : (sP4(X12,X13,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl108),
  inference(cnf_transformation,[],[f2354_D])).
fof(f2354_D,plain,(
  ( ! [X12,X13] : (sP4(X12,X13,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f6126,plain,(
  ( ! [X4,X5,X3] : (sP4(X3,X4,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f2338,f835])).
fof(f2338,plain,(
  ( ! [X4,X5,X3] : (sP4(X3,X4,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl0),
  inference(superposition,[],[f1217,f617])).
fof(f2375,plain,(
  ( ! [X10,X11] : (sP4('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X10,X11)) ) | $spl110),
  inference(cnf_transformation,[],[f2375_D])).
fof(f2375_D,plain,(
  ( ! [X10,X11] : (sP4('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X10,X11)) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f6125,plain,(
  ( ! [X2,X0,X1] : (sP4('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0,X1) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f2360,f912])).
fof(f2360,plain,(
  ( ! [X2,X0,X1] : (sP4('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0,X1) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | $spl0),
  inference(superposition,[],[f1218,f617])).
fof(f1218,plain,(
  ( ! [X68,X69] : (sP4('const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0')),X68,X69)) ) | $spl0),
  inference(resolution,[],[f987,f631])).
fof(f631,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f462])).
fof(f2377,plain,(
  ( ! [X12,X13] : (sP4('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X12,X13)) ) | $spl112),
  inference(cnf_transformation,[],[f2377_D])).
fof(f2377_D,plain,(
  ( ! [X12,X13] : (sP4('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X12,X13)) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f6124,plain,(
  ( ! [X4,X5,X3] : (sP4('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X3,X4) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f2361,f835])).
fof(f2361,plain,(
  ( ! [X4,X5,X3] : (sP4('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X3,X4) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl0),
  inference(superposition,[],[f1218,f617])).
fof(f2399,plain,(
  ( ! [X10,X11] : (sP6('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X10,X11)) ) | $spl114),
  inference(cnf_transformation,[],[f2399_D])).
fof(f2399_D,plain,(
  ( ! [X10,X11] : (sP6('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X10,X11)) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f6123,plain,(
  ( ! [X2,X0,X1] : (sP6('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0,X1) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f2383,f912])).
fof(f2383,plain,(
  ( ! [X2,X0,X1] : (sP6('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0,X1) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | $spl0),
  inference(superposition,[],[f1234,f617])).
fof(f1234,plain,(
  ( ! [X101,X100] : (sP6('const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0')),X100,X101)) ) | $spl0),
  inference(resolution,[],[f987,f648])).
fof(f648,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f470])).
fof(f470,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(rectify,[],[f469])).
fof(f469,plain,(
  ! [X2,X0,X1] : ((~sP6(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP6(X2,X0,X1)))),
  inference(flattening,[],[f468])).
fof(f468,plain,(
  ! [X2,X0,X1] : ((~sP6(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP6(X2,X0,X1)))),
  inference(nnf_transformation,[],[f423])).
fof(f423,plain,(
  ! [X2,X0,X1] : (sP6(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f2401,plain,(
  ( ! [X12,X13] : (sP6('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X12,X13)) ) | $spl116),
  inference(cnf_transformation,[],[f2401_D])).
fof(f2401_D,plain,(
  ( ! [X12,X13] : (sP6('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X12,X13)) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f6122,plain,(
  ( ! [X4,X5,X3] : (sP6('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X3,X4) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f2384,f835])).
fof(f2384,plain,(
  ( ! [X4,X5,X3] : (sP6('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X3,X4) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl0),
  inference(superposition,[],[f1234,f617])).
fof(f2422,plain,(
  ( ! [X10,X11] : (sP8(X10,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X11)) ) | $spl118),
  inference(cnf_transformation,[],[f2422_D])).
fof(f2422_D,plain,(
  ( ! [X10,X11] : (sP8(X10,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X11)) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f6121,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X1) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f2407,f912])).
fof(f2407,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X1) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | $spl0),
  inference(superposition,[],[f1247,f617])).
fof(f1247,plain,(
  ( ! [X127,X126] : (sP8(X126,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0')),X127)) ) | $spl0),
  inference(resolution,[],[f987,f664])).
fof(f664,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f477])).
fof(f477,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(flattening,[],[f476])).
fof(f476,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(nnf_transformation,[],[f426])).
fof(f426,plain,(
  ! [X0,X1,X2] : (sP8(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f2424,plain,(
  ( ! [X12,X13] : (sP8(X12,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X13)) ) | $spl120),
  inference(cnf_transformation,[],[f2424_D])).
fof(f2424_D,plain,(
  ( ! [X12,X13] : (sP8(X12,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X13)) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f6120,plain,(
  ( ! [X4,X5,X3] : (sP8(X3,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X4) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f2408,f835])).
fof(f2408,plain,(
  ( ! [X4,X5,X3] : (sP8(X3,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X4) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl0),
  inference(superposition,[],[f1247,f617])).
fof(f2445,plain,(
  ( ! [X10,X11] : (sP8('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X10,X11)) ) | $spl122),
  inference(cnf_transformation,[],[f2445_D])).
fof(f2445_D,plain,(
  ( ! [X10,X11] : (sP8('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X10,X11)) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f6119,plain,(
  ( ! [X2,X0,X1] : (sP8('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0,X1) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f2430,f912])).
fof(f2430,plain,(
  ( ! [X2,X0,X1] : (sP8('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0,X1) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | $spl0),
  inference(superposition,[],[f1248,f617])).
fof(f1248,plain,(
  ( ! [X128,X129] : (sP8('const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0')),X128,X129)) ) | $spl0),
  inference(resolution,[],[f987,f665])).
fof(f665,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f477])).
fof(f2447,plain,(
  ( ! [X12,X13] : (sP8('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X12,X13)) ) | $spl124),
  inference(cnf_transformation,[],[f2447_D])).
fof(f2447_D,plain,(
  ( ! [X12,X13] : (sP8('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X12,X13)) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f6118,plain,(
  ( ! [X4,X5,X3] : (sP8('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X3,X4) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f2431,f835])).
fof(f2431,plain,(
  ( ! [X4,X5,X3] : (sP8('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X3,X4) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl0),
  inference(superposition,[],[f1248,f617])).
fof(f1326,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl44),
  inference(cnf_transformation,[],[f1326_D])).
fof(f1326_D,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f6117,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f2686,f912])).
fof(f2686,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))) ) | $spl0),
  inference(superposition,[],[f1800,f617])).
fof(f1800,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK12)))) | $spl0),
  inference(resolution,[],[f1131,f807])).
fof(f807,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0))))) )),
  inference(cnf_transformation,[],[f563])).
fof(f1131,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl0),
  inference(subsumption_resolution,[],[f1130,f790])).
fof(f790,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f401])).
fof(f401,plain,(
  ! [X0] : ((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f348])).
fof(f348,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) => ~p(s(bool,X0)))),
  inference(flattening,[],[f347])).
fof(f347,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) => ~p(s(bool,X0)))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X21] : ((p(s(bool,X21)) => p(s(bool,'F_0'))) => ~p(s(bool,X21)))),
  file('thm.extreal.le_infty|split|3','thm.bool.IMP_F')).
fof(f1130,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl0),
  inference(forward_demodulation,[],[f1098,f818])).
fof(f818,plain,(
  s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))),
  inference(cnf_transformation,[],[f20])).
fof(f20,axiom,(
  s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))),
  file('thm.extreal.le_infty|split|3','thm.extreal.extreal_le_def|split|4')).
fof(f1098,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl0),
  inference(superposition,[],[f832,f809])).
fof(f809,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.NegInf_0'))))) )),
  inference(cnf_transformation,[],[f564])).
fof(f564,plain,(
  ! [X0] : ((~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0)) & (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.NegInf_0'))))))),
  inference(nnf_transformation,[],[f365])).
fof(f365,plain,(
  ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.NegInf_0')))) <=> s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X6] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal','const.extreal.NegInf_0')))) <=> s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X6))),
  file('thm.extreal.le_infty|split|3','thm.extreal.le_infty|split|2')).
fof(f6116,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl310),
  inference(cnf_transformation,[],[f6116_D])).
fof(f6116_D,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) <=> ~$spl310),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl310])])).
fof(f6114,plain,(
  ( ! [X1] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f2687,f835])).
fof(f2687,plain,(
  ( ! [X1] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | $spl0),
  inference(superposition,[],[f1800,f617])).
fof(f6113,plain,(
  ( ! [X0] : (sP11(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl308),
  inference(cnf_transformation,[],[f6113_D])).
fof(f6113_D,plain,(
  ( ! [X0] : (sP11(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) <=> ~$spl308),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl308])])).
fof(f6112,plain,(
  ( ! [X0,X1] : (sP11(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f2725,f912])).
fof(f2725,plain,(
  ( ! [X0,X1] : (sP11(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | $spl0),
  inference(superposition,[],[f2654,f617])).
fof(f2654,plain,(
  ( ! [X152] : (sP11(X152,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK12)))) ) | $spl0),
  inference(resolution,[],[f1800,f695])).
fof(f695,plain,(
  ( ! [X0,X1] : (sP11(X0,X1) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f489])).
fof(f489,plain,(
  ! [X0,X1] : ((~sP11(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP11(X0,X1)))),
  inference(flattening,[],[f488])).
fof(f488,plain,(
  ! [X0,X1] : ((~sP11(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP11(X0,X1)))),
  inference(nnf_transformation,[],[f431])).
fof(f431,plain,(
  ! [X0,X1] : (sP11(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f6111,plain,(
  ( ! [X2] : (sP11(X2,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl306),
  inference(cnf_transformation,[],[f6111_D])).
fof(f6111_D,plain,(
  ( ! [X2] : (sP11(X2,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) <=> ~$spl306),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl306])])).
fof(f6110,plain,(
  ( ! [X2,X3] : (sP11(X2,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f2726,f835])).
fof(f2726,plain,(
  ( ! [X2,X3] : (sP11(X2,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3)))) ) | $spl0),
  inference(superposition,[],[f2654,f617])).
fof(f6109,plain,(
  ( ! [X0] : (sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0)) ) | $spl304),
  inference(cnf_transformation,[],[f6109_D])).
fof(f6109_D,plain,(
  ( ! [X0] : (sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0)) ) <=> ~$spl304),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl304])])).
fof(f6108,plain,(
  ( ! [X0,X1] : (sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f2747,f912])).
fof(f2747,plain,(
  ( ! [X0,X1] : (sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | $spl0),
  inference(superposition,[],[f2655,f617])).
fof(f2655,plain,(
  ( ! [X153] : (sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK12)),X153)) ) | $spl0),
  inference(resolution,[],[f1800,f696])).
fof(f696,plain,(
  ( ! [X0,X1] : (sP11(X0,X1) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f489])).
fof(f6107,plain,(
  ( ! [X2] : (sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X2)) ) | $spl302),
  inference(cnf_transformation,[],[f6107_D])).
fof(f6107_D,plain,(
  ( ! [X2] : (sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X2)) ) <=> ~$spl302),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl302])])).
fof(f6106,plain,(
  ( ! [X2,X3] : (sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X2) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f2748,f835])).
fof(f2748,plain,(
  ( ! [X2,X3] : (sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X2) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3)))) ) | $spl0),
  inference(superposition,[],[f2655,f617])).
fof(f6105,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X1)) ) | $spl300),
  inference(cnf_transformation,[],[f6105_D])).
fof(f6105_D,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X1)) ) <=> ~$spl300),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl300])])).
fof(f6104,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X1) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f2830,f912])).
fof(f2830,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X1) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | $spl0),
  inference(superposition,[],[f1813,f617])).
fof(f1813,plain,(
  ( ! [X33,X32] : (sP2(X32,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.NegInf_0')),X33)) ) | $spl0),
  inference(resolution,[],[f1131,f584])).
fof(f6103,plain,(
  ( ! [X4,X3] : (sP2(X3,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X4)) ) | $spl298),
  inference(cnf_transformation,[],[f6103_D])).
fof(f6103_D,plain,(
  ( ! [X4,X3] : (sP2(X3,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X4)) ) <=> ~$spl298),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl298])])).
fof(f6102,plain,(
  ( ! [X4,X5,X3] : (sP2(X3,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X4) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f2831,f835])).
fof(f2831,plain,(
  ( ! [X4,X5,X3] : (sP2(X3,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X4) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl0),
  inference(superposition,[],[f1813,f617])).
fof(f6101,plain,(
  ( ! [X0,X1] : (sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X0,X1)) ) | $spl296),
  inference(cnf_transformation,[],[f6101_D])).
fof(f6101_D,plain,(
  ( ! [X0,X1] : (sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X0,X1)) ) <=> ~$spl296),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl296])])).
fof(f6100,plain,(
  ( ! [X2,X0,X1] : (sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X0,X1) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f2853,f912])).
fof(f2853,plain,(
  ( ! [X2,X0,X1] : (sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X0,X1) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | $spl0),
  inference(superposition,[],[f1814,f617])).
fof(f1814,plain,(
  ( ! [X35,X34] : (sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.NegInf_0')),X34,X35)) ) | $spl0),
  inference(resolution,[],[f1131,f585])).
fof(f6099,plain,(
  ( ! [X4,X3] : (sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X3,X4)) ) | $spl294),
  inference(cnf_transformation,[],[f6099_D])).
fof(f6099_D,plain,(
  ( ! [X4,X3] : (sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X3,X4)) ) <=> ~$spl294),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl294])])).
fof(f6098,plain,(
  ( ! [X4,X5,X3] : (sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X3,X4) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f2854,f835])).
fof(f2854,plain,(
  ( ! [X4,X5,X3] : (sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X3,X4) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl0),
  inference(superposition,[],[f1814,f617])).
fof(f6097,plain,(
  ( ! [X0,X1] : (sP1(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X1)) ) | $spl292),
  inference(cnf_transformation,[],[f6097_D])).
fof(f6097_D,plain,(
  ( ! [X0,X1] : (sP1(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X1)) ) <=> ~$spl292),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl292])])).
fof(f6096,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X1) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f2878,f912])).
fof(f2878,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X1) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | $spl0),
  inference(superposition,[],[f1816,f617])).
fof(f1816,plain,(
  ( ! [X39,X38] : (sP1(X38,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.NegInf_0')),X39)) ) | $spl0),
  inference(resolution,[],[f1131,f588])).
fof(f6095,plain,(
  ( ! [X4,X3] : (sP1(X3,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X4)) ) | $spl290),
  inference(cnf_transformation,[],[f6095_D])).
fof(f6095_D,plain,(
  ( ! [X4,X3] : (sP1(X3,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X4)) ) <=> ~$spl290),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl290])])).
fof(f6094,plain,(
  ( ! [X4,X5,X3] : (sP1(X3,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X4) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f2879,f835])).
fof(f2879,plain,(
  ( ! [X4,X5,X3] : (sP1(X3,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X4) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl0),
  inference(superposition,[],[f1816,f617])).
fof(f6093,plain,(
  ( ! [X0,X1] : (sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X0,X1)) ) | $spl288),
  inference(cnf_transformation,[],[f6093_D])).
fof(f6093_D,plain,(
  ( ! [X0,X1] : (sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X0,X1)) ) <=> ~$spl288),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl288])])).
fof(f6092,plain,(
  ( ! [X2,X0,X1] : (sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X0,X1) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f2902,f912])).
fof(f2902,plain,(
  ( ! [X2,X0,X1] : (sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X0,X1) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | $spl0),
  inference(superposition,[],[f1817,f617])).
fof(f1817,plain,(
  ( ! [X41,X40] : (sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.NegInf_0')),X40,X41)) ) | $spl0),
  inference(resolution,[],[f1131,f589])).
fof(f6091,plain,(
  ( ! [X4,X3] : (sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X3,X4)) ) | $spl286),
  inference(cnf_transformation,[],[f6091_D])).
fof(f6091_D,plain,(
  ( ! [X4,X3] : (sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X3,X4)) ) <=> ~$spl286),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl286])])).
fof(f6090,plain,(
  ( ! [X4,X5,X3] : (sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X3,X4) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f2903,f835])).
fof(f2903,plain,(
  ( ! [X4,X5,X3] : (sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X3,X4) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl0),
  inference(superposition,[],[f1817,f617])).
fof(f6089,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl284),
  inference(cnf_transformation,[],[f6089_D])).
fof(f6089_D,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) <=> ~$spl284),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl284])])).
fof(f6088,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f2928,f912])).
fof(f2928,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | $spl0),
  inference(superposition,[],[f1831,f617])).
fof(f1831,plain,(
  ( ! [X66,X67] : (sP4(X66,X67,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl0),
  inference(resolution,[],[f1131,f629])).
fof(f6087,plain,(
  ( ! [X4,X3] : (sP4(X3,X4,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl282),
  inference(cnf_transformation,[],[f6087_D])).
fof(f6087_D,plain,(
  ( ! [X4,X3] : (sP4(X3,X4,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) <=> ~$spl282),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl282])])).
fof(f6086,plain,(
  ( ! [X4,X5,X3] : (sP4(X3,X4,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f2929,f835])).
fof(f2929,plain,(
  ( ! [X4,X5,X3] : (sP4(X3,X4,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl0),
  inference(superposition,[],[f1831,f617])).
fof(f6085,plain,(
  ( ! [X0,X1] : (sP4('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X0,X1)) ) | $spl280),
  inference(cnf_transformation,[],[f6085_D])).
fof(f6085_D,plain,(
  ( ! [X0,X1] : (sP4('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X0,X1)) ) <=> ~$spl280),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl280])])).
fof(f6084,plain,(
  ( ! [X2,X0,X1] : (sP4('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X0,X1) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f3534,f912])).
fof(f3534,plain,(
  ( ! [X2,X0,X1] : (sP4('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X0,X1) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | $spl0),
  inference(superposition,[],[f1832,f617])).
fof(f1832,plain,(
  ( ! [X68,X69] : (sP4('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.NegInf_0')),X68,X69)) ) | $spl0),
  inference(resolution,[],[f1131,f631])).
fof(f6083,plain,(
  ( ! [X4,X3] : (sP4('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X3,X4)) ) | $spl278),
  inference(cnf_transformation,[],[f6083_D])).
fof(f6083_D,plain,(
  ( ! [X4,X3] : (sP4('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X3,X4)) ) <=> ~$spl278),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl278])])).
fof(f6082,plain,(
  ( ! [X4,X5,X3] : (sP4('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X3,X4) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f3535,f835])).
fof(f3535,plain,(
  ( ! [X4,X5,X3] : (sP4('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X3,X4) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl0),
  inference(superposition,[],[f1832,f617])).
fof(f6081,plain,(
  ( ! [X0,X1] : (sP6('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X0,X1)) ) | $spl276),
  inference(cnf_transformation,[],[f6081_D])).
fof(f6081_D,plain,(
  ( ! [X0,X1] : (sP6('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X0,X1)) ) <=> ~$spl276),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl276])])).
fof(f6080,plain,(
  ( ! [X2,X0,X1] : (sP6('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X0,X1) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f3558,f912])).
fof(f3558,plain,(
  ( ! [X2,X0,X1] : (sP6('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X0,X1) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | $spl0),
  inference(superposition,[],[f1848,f617])).
fof(f1848,plain,(
  ( ! [X101,X100] : (sP6('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.NegInf_0')),X100,X101)) ) | $spl0),
  inference(resolution,[],[f1131,f648])).
fof(f6079,plain,(
  ( ! [X4,X3] : (sP6('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X3,X4)) ) | $spl274),
  inference(cnf_transformation,[],[f6079_D])).
fof(f6079_D,plain,(
  ( ! [X4,X3] : (sP6('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X3,X4)) ) <=> ~$spl274),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl274])])).
fof(f6078,plain,(
  ( ! [X4,X5,X3] : (sP6('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X3,X4) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f3559,f835])).
fof(f3559,plain,(
  ( ! [X4,X5,X3] : (sP6('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X3,X4) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl0),
  inference(superposition,[],[f1848,f617])).
fof(f6077,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X1)) ) | $spl272),
  inference(cnf_transformation,[],[f6077_D])).
fof(f6077_D,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X1)) ) <=> ~$spl272),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl272])])).
fof(f6076,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X1) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f3583,f912])).
fof(f3583,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X1) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | $spl0),
  inference(superposition,[],[f1861,f617])).
fof(f1861,plain,(
  ( ! [X127,X126] : (sP8(X126,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.NegInf_0')),X127)) ) | $spl0),
  inference(resolution,[],[f1131,f664])).
fof(f6075,plain,(
  ( ! [X4,X3] : (sP8(X3,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X4)) ) | $spl270),
  inference(cnf_transformation,[],[f6075_D])).
fof(f6075_D,plain,(
  ( ! [X4,X3] : (sP8(X3,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X4)) ) <=> ~$spl270),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl270])])).
fof(f6074,plain,(
  ( ! [X4,X5,X3] : (sP8(X3,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X4) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f3584,f835])).
fof(f3584,plain,(
  ( ! [X4,X5,X3] : (sP8(X3,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X4) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl0),
  inference(superposition,[],[f1861,f617])).
fof(f6073,plain,(
  ( ! [X0,X1] : (sP8('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X0,X1)) ) | $spl268),
  inference(cnf_transformation,[],[f6073_D])).
fof(f6073_D,plain,(
  ( ! [X0,X1] : (sP8('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X0,X1)) ) <=> ~$spl268),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl268])])).
fof(f6072,plain,(
  ( ! [X2,X0,X1] : (sP8('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X0,X1) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f3607,f912])).
fof(f3607,plain,(
  ( ! [X2,X0,X1] : (sP8('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X0,X1) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | $spl0),
  inference(superposition,[],[f1862,f617])).
fof(f1862,plain,(
  ( ! [X128,X129] : (sP8('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.NegInf_0')),X128,X129)) ) | $spl0),
  inference(resolution,[],[f1131,f665])).
fof(f6071,plain,(
  ( ! [X4,X3] : (sP8('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X3,X4)) ) | $spl266),
  inference(cnf_transformation,[],[f6071_D])).
fof(f6071_D,plain,(
  ( ! [X4,X3] : (sP8('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X3,X4)) ) <=> ~$spl266),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl266])])).
fof(f6070,plain,(
  ( ! [X4,X5,X3] : (sP8('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X3,X4) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f3608,f835])).
fof(f3608,plain,(
  ( ! [X4,X5,X3] : (sP8('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X3,X4) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl0),
  inference(superposition,[],[f1862,f617])).
fof(f6069,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl264),
  inference(cnf_transformation,[],[f6069_D])).
fof(f6069_D,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) <=> ~$spl264),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl264])])).
fof(f6068,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f3641,f912])).
fof(f3641,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | $spl0),
  inference(superposition,[],[f2590,f617])).
fof(f2590,plain,(
  ( ! [X35,X34] : (sP2(X34,X35,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK12)))) ) | $spl0),
  inference(resolution,[],[f1800,f583])).
fof(f583,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f442])).
fof(f6067,plain,(
  ( ! [X4,X3] : (sP2(X3,X4,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl262),
  inference(cnf_transformation,[],[f6067_D])).
fof(f6067_D,plain,(
  ( ! [X4,X3] : (sP2(X3,X4,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) <=> ~$spl262),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl262])])).
fof(f6066,plain,(
  ( ! [X4,X5,X3] : (sP2(X3,X4,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f3642,f835])).
fof(f3642,plain,(
  ( ! [X4,X5,X3] : (sP2(X3,X4,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl0),
  inference(superposition,[],[f2590,f617])).
fof(f6065,plain,(
  ( ! [X0,X1] : (sP1(X0,X1,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl260),
  inference(cnf_transformation,[],[f6065_D])).
fof(f6065_D,plain,(
  ( ! [X0,X1] : (sP1(X0,X1,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) <=> ~$spl260),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl260])])).
fof(f6064,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f3665,f912])).
fof(f3665,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | $spl0),
  inference(superposition,[],[f2593,f617])).
fof(f2593,plain,(
  ( ! [X41,X40] : (sP1(X40,X41,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK12)))) ) | $spl0),
  inference(resolution,[],[f1800,f587])).
fof(f587,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f445])).
fof(f6063,plain,(
  ( ! [X4,X3] : (sP1(X3,X4,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl258),
  inference(cnf_transformation,[],[f6063_D])).
fof(f6063_D,plain,(
  ( ! [X4,X3] : (sP1(X3,X4,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) <=> ~$spl258),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl258])])).
fof(f6062,plain,(
  ( ! [X4,X5,X3] : (sP1(X3,X4,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f3666,f835])).
fof(f3666,plain,(
  ( ! [X4,X5,X3] : (sP1(X3,X4,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl0),
  inference(superposition,[],[f2593,f617])).
fof(f6061,plain,(
  ( ! [X0,X1] : (sP0(X0,X1,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl256),
  inference(cnf_transformation,[],[f6061_D])).
fof(f6061_D,plain,(
  ( ! [X0,X1] : (sP0(X0,X1,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) <=> ~$spl256),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl256])])).
fof(f6060,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f3689,f912])).
fof(f3689,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | $spl0),
  inference(superposition,[],[f2594,f617])).
fof(f2594,plain,(
  ( ! [X43,X42] : (sP0(X42,X43,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK12)))) ) | $spl0),
  inference(resolution,[],[f1800,f591])).
fof(f591,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f447])).
fof(f447,plain,(
  ! [X0,X1,X2] : ((~sP0(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP0(X0,X1,X2)))),
  inference(flattening,[],[f446])).
fof(f446,plain,(
  ! [X0,X1,X2] : ((~sP0(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP0(X0,X1,X2)))),
  inference(nnf_transformation,[],[f415])).
fof(f415,plain,(
  ! [X0,X1,X2] : (sP0(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f6059,plain,(
  ( ! [X4,X3] : (sP0(X3,X4,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl254),
  inference(cnf_transformation,[],[f6059_D])).
fof(f6059_D,plain,(
  ( ! [X4,X3] : (sP0(X3,X4,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) <=> ~$spl254),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl254])])).
fof(f6058,plain,(
  ( ! [X4,X5,X3] : (sP0(X3,X4,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f3690,f835])).
fof(f3690,plain,(
  ( ! [X4,X5,X3] : (sP0(X3,X4,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl0),
  inference(superposition,[],[f2594,f617])).
fof(f6057,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X1)) ) | $spl252),
  inference(cnf_transformation,[],[f6057_D])).
fof(f6057_D,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X1)) ) <=> ~$spl252),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl252])])).
fof(f6056,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X1) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f3714,f912])).
fof(f3714,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X1) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | $spl0),
  inference(superposition,[],[f2595,f617])).
fof(f2595,plain,(
  ( ! [X45,X44] : (sP0(X44,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK12)),X45)) ) | $spl0),
  inference(resolution,[],[f1800,f592])).
fof(f592,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f447])).
fof(f6055,plain,(
  ( ! [X4,X3] : (sP0(X3,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X4)) ) | $spl250),
  inference(cnf_transformation,[],[f6055_D])).
fof(f6055_D,plain,(
  ( ! [X4,X3] : (sP0(X3,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X4)) ) <=> ~$spl250),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl250])])).
fof(f6054,plain,(
  ( ! [X4,X5,X3] : (sP0(X3,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X4) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f3715,f835])).
fof(f3715,plain,(
  ( ! [X4,X5,X3] : (sP0(X3,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X4) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl0),
  inference(superposition,[],[f2595,f617])).
fof(f6053,plain,(
  ( ! [X0,X1] : (sP0('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0,X1)) ) | $spl248),
  inference(cnf_transformation,[],[f6053_D])).
fof(f6053_D,plain,(
  ( ! [X0,X1] : (sP0('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0,X1)) ) <=> ~$spl248),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl248])])).
fof(f6052,plain,(
  ( ! [X2,X0,X1] : (sP0('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0,X1) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f3739,f912])).
fof(f3739,plain,(
  ( ! [X2,X0,X1] : (sP0('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0,X1) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | $spl0),
  inference(superposition,[],[f2596,f617])).
fof(f2596,plain,(
  ( ! [X47,X46] : (sP0('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK12)),X46,X47)) ) | $spl0),
  inference(resolution,[],[f1800,f593])).
fof(f593,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f447])).
fof(f6051,plain,(
  ( ! [X4,X3] : (sP0('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X3,X4)) ) | $spl246),
  inference(cnf_transformation,[],[f6051_D])).
fof(f6051_D,plain,(
  ( ! [X4,X3] : (sP0('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X3,X4)) ) <=> ~$spl246),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl246])])).
fof(f6050,plain,(
  ( ! [X4,X5,X3] : (sP0('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X3,X4) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f3740,f835])).
fof(f3740,plain,(
  ( ! [X4,X5,X3] : (sP0('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X3,X4) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl0),
  inference(superposition,[],[f2596,f617])).
fof(f6049,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X1)) ) | $spl244),
  inference(cnf_transformation,[],[f6049_D])).
fof(f6049_D,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X1)) ) <=> ~$spl244),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl244])])).
fof(f6048,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X1) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f3765,f912])).
fof(f3765,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X1) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | $spl0),
  inference(superposition,[],[f2608,f617])).
fof(f2608,plain,(
  ( ! [X68,X69] : (sP4(X68,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK12)),X69)) ) | $spl0),
  inference(resolution,[],[f1800,f630])).
fof(f630,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f462])).
fof(f6047,plain,(
  ( ! [X4,X3] : (sP4(X3,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X4)) ) | $spl242),
  inference(cnf_transformation,[],[f6047_D])).
fof(f6047_D,plain,(
  ( ! [X4,X3] : (sP4(X3,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X4)) ) <=> ~$spl242),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl242])])).
fof(f6046,plain,(
  ( ! [X4,X5,X3] : (sP4(X3,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X4) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f3766,f835])).
fof(f3766,plain,(
  ( ! [X4,X5,X3] : (sP4(X3,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X4) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl0),
  inference(superposition,[],[f2608,f617])).
fof(f6045,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl240),
  inference(cnf_transformation,[],[f6045_D])).
fof(f6045_D,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) <=> ~$spl240),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl240])])).
fof(f6044,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f3789,f912])).
fof(f3789,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | $spl0),
  inference(superposition,[],[f2621,f617])).
fof(f2621,plain,(
  ( ! [X94,X95] : (sP6(X94,X95,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK12)))) ) | $spl0),
  inference(resolution,[],[f1800,f646])).
fof(f646,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f470])).
fof(f6043,plain,(
  ( ! [X4,X3] : (sP6(X3,X4,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl238),
  inference(cnf_transformation,[],[f6043_D])).
fof(f6043_D,plain,(
  ( ! [X4,X3] : (sP6(X3,X4,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) <=> ~$spl238),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl238])])).
fof(f6042,plain,(
  ( ! [X4,X5,X3] : (sP6(X3,X4,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f3790,f835])).
fof(f3790,plain,(
  ( ! [X4,X5,X3] : (sP6(X3,X4,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl0),
  inference(superposition,[],[f2621,f617])).
fof(f6041,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X1)) ) | $spl236),
  inference(cnf_transformation,[],[f6041_D])).
fof(f6041_D,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X1)) ) <=> ~$spl236),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl236])])).
fof(f6040,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X1) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f3813,f912])).
fof(f3813,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X1) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | $spl0),
  inference(superposition,[],[f2622,f617])).
fof(f2622,plain,(
  ( ! [X97,X96] : (sP6(X96,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK12)),X97)) ) | $spl0),
  inference(resolution,[],[f1800,f647])).
fof(f647,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f470])).
fof(f6039,plain,(
  ( ! [X4,X3] : (sP6(X3,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X4)) ) | $spl234),
  inference(cnf_transformation,[],[f6039_D])).
fof(f6039_D,plain,(
  ( ! [X4,X3] : (sP6(X3,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X4)) ) <=> ~$spl234),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl234])])).
fof(f6038,plain,(
  ( ! [X4,X5,X3] : (sP6(X3,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X4) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f3814,f835])).
fof(f3814,plain,(
  ( ! [X4,X5,X3] : (sP6(X3,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X4) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl0),
  inference(superposition,[],[f2622,f617])).
fof(f6037,plain,(
  ( ! [X0] : (~sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0) | ~p(s(bool,X0))) ) | $spl232),
  inference(cnf_transformation,[],[f6037_D])).
fof(f6037_D,plain,(
  ( ! [X0] : (~sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0) | ~p(s(bool,X0))) ) <=> ~$spl232),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl232])])).
fof(f6036,plain,(
  ( ! [X0,X1] : (~sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0) | ~p(s(bool,X0)) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f3831,f912])).
fof(f3831,plain,(
  ( ! [X0,X1] : (~sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0) | ~p(s(bool,X0)) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | $spl0),
  inference(superposition,[],[f1057,f617])).
fof(f1057,plain,(
  ( ! [X144] : (~sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)),X144) | ~p(s(bool,X144))) ) | $spl0),
  inference(resolution,[],[f832,f679])).
fof(f679,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | ~sP10(X0,X1)) )),
  inference(cnf_transformation,[],[f484])).
fof(f484,plain,(
  ! [X0,X1] : ((~sP10(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP10(X0,X1)))),
  inference(rectify,[],[f483])).
fof(f483,plain,(
  ! [X1,X0] : ((~sP10(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP10(X1,X0)))),
  inference(flattening,[],[f482])).
fof(f482,plain,(
  ! [X1,X0] : ((~sP10(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP10(X1,X0)))),
  inference(nnf_transformation,[],[f429])).
fof(f429,plain,(
  ! [X1,X0] : (sP10(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f6035,plain,(
  ( ! [X2] : (~sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X2) | ~p(s(bool,X2))) ) | $spl230),
  inference(cnf_transformation,[],[f6035_D])).
fof(f6035_D,plain,(
  ( ! [X2] : (~sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X2) | ~p(s(bool,X2))) ) <=> ~$spl230),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl230])])).
fof(f6034,plain,(
  ( ! [X2,X3] : (~sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X2) | ~p(s(bool,X2)) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f3832,f835])).
fof(f3832,plain,(
  ( ! [X2,X3] : (~sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X2) | ~p(s(bool,X2)) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3)))) ) | $spl0),
  inference(superposition,[],[f1057,f617])).
fof(f6033,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl228),
  inference(cnf_transformation,[],[f6033_D])).
fof(f6033_D,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) <=> ~$spl228),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl228])])).
fof(f6032,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f3856,f912])).
fof(f3856,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | $spl0),
  inference(superposition,[],[f2638,f617])).
fof(f2638,plain,(
  ( ! [X128,X129] : (sP8(X128,X129,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK12)))) ) | $spl0),
  inference(resolution,[],[f1800,f663])).
fof(f663,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f477])).
fof(f6031,plain,(
  ( ! [X4,X3] : (sP8(X3,X4,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl226),
  inference(cnf_transformation,[],[f6031_D])).
fof(f6031_D,plain,(
  ( ! [X4,X3] : (sP8(X3,X4,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) <=> ~$spl226),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl226])])).
fof(f6030,plain,(
  ( ! [X4,X5,X3] : (sP8(X3,X4,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f3857,f835])).
fof(f3857,plain,(
  ( ! [X4,X5,X3] : (sP8(X3,X4,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl0),
  inference(superposition,[],[f2638,f617])).
fof(f6029,plain,(
  ( ! [X0] : (~sP10(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | ~p(s(bool,X0))) ) | $spl224),
  inference(cnf_transformation,[],[f6029_D])).
fof(f6029_D,plain,(
  ( ! [X0] : (~sP10(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | ~p(s(bool,X0))) ) <=> ~$spl224),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl224])])).
fof(f6028,plain,(
  ( ! [X0,X1] : (~sP10(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | ~p(s(bool,X0)) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f3875,f912])).
fof(f3875,plain,(
  ( ! [X0,X1] : (~sP10(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | ~p(s(bool,X0)) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | $spl0),
  inference(superposition,[],[f1058,f617])).
fof(f1058,plain,(
  ( ! [X145] : (~sP10(X145,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12))) | ~p(s(bool,X145))) ) | $spl0),
  inference(resolution,[],[f832,f679])).
fof(f6027,plain,(
  ( ! [X2] : (~sP10(X2,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | ~p(s(bool,X2))) ) | $spl222),
  inference(cnf_transformation,[],[f6027_D])).
fof(f6027_D,plain,(
  ( ! [X2] : (~sP10(X2,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | ~p(s(bool,X2))) ) <=> ~$spl222),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl222])])).
fof(f6026,plain,(
  ( ! [X2,X3] : (~sP10(X2,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | ~p(s(bool,X2)) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f3876,f835])).
fof(f3876,plain,(
  ( ! [X2,X3] : (~sP10(X2,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | ~p(s(bool,X2)) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3)))) ) | $spl0),
  inference(superposition,[],[f1058,f617])).
fof(f6025,plain,(
  ( ! [X0] : (sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0) | p(s(bool,X0))) ) | $spl220),
  inference(cnf_transformation,[],[f6025_D])).
fof(f6025_D,plain,(
  ( ! [X0] : (sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0) | p(s(bool,X0))) ) <=> ~$spl220),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl220])])).
fof(f6024,plain,(
  ( ! [X0,X1] : (sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0) | p(s(bool,X0)) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f3896,f912])).
fof(f3896,plain,(
  ( ! [X0,X1] : (sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0) | p(s(bool,X0)) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | $spl0),
  inference(superposition,[],[f1059,f617])).
fof(f1059,plain,(
  ( ! [X146] : (sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)),X146) | p(s(bool,X146))) ) | $spl0),
  inference(resolution,[],[f832,f680])).
fof(f680,plain,(
  ( ! [X0,X1] : (sP10(X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f484])).
fof(f6023,plain,(
  ( ! [X2] : (sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X2) | p(s(bool,X2))) ) | $spl218),
  inference(cnf_transformation,[],[f6023_D])).
fof(f6023_D,plain,(
  ( ! [X2] : (sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X2) | p(s(bool,X2))) ) <=> ~$spl218),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl218])])).
fof(f6022,plain,(
  ( ! [X2,X3] : (sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X2) | p(s(bool,X2)) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f3897,f835])).
fof(f3897,plain,(
  ( ! [X2,X3] : (sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X2) | p(s(bool,X2)) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3)))) ) | $spl0),
  inference(superposition,[],[f1059,f617])).
fof(f6021,plain,(
  ( ! [X0] : (sP10(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | p(s(bool,X0))) ) | $spl216),
  inference(cnf_transformation,[],[f6021_D])).
fof(f6021_D,plain,(
  ( ! [X0] : (sP10(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | p(s(bool,X0))) ) <=> ~$spl216),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl216])])).
fof(f6020,plain,(
  ( ! [X0,X1] : (sP10(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | p(s(bool,X0)) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f3917,f912])).
fof(f3917,plain,(
  ( ! [X0,X1] : (sP10(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | p(s(bool,X0)) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | $spl0),
  inference(superposition,[],[f1060,f617])).
fof(f1060,plain,(
  ( ! [X147] : (sP10(X147,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12))) | p(s(bool,X147))) ) | $spl0),
  inference(resolution,[],[f832,f683])).
fof(f683,plain,(
  ( ! [X0,X1] : (sP10(X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f484])).
fof(f6019,plain,(
  ( ! [X2] : (sP10(X2,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | p(s(bool,X2))) ) | $spl214),
  inference(cnf_transformation,[],[f6019_D])).
fof(f6019_D,plain,(
  ( ! [X2] : (sP10(X2,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | p(s(bool,X2))) ) <=> ~$spl214),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl214])])).
fof(f6018,plain,(
  ( ! [X2,X3] : (sP10(X2,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | p(s(bool,X2)) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f3918,f835])).
fof(f3918,plain,(
  ( ! [X2,X3] : (sP10(X2,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | p(s(bool,X2)) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3)))) ) | $spl0),
  inference(superposition,[],[f1060,f617])).
fof(f6015,plain,(
  ( ! [X1] : (s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f6014,f835])).
fof(f6014,plain,(
  ( ! [X1] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | $spl0),
  inference(subsumption_resolution,[],[f3933,f818])).
fof(f3933,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | $spl0),
  inference(superposition,[],[f1078,f617])).
fof(f1078,plain,(
  s(bool,'F_0') != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12))) | $spl0),
  inference(resolution,[],[f832,f776])).
fof(f776,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f547])).
fof(f547,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f338])).
fof(f338,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f337])).
fof(f337,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X21] : (s(bool,'F_0') = s(bool,X21) <=> ~p(s(bool,X21)))),
  file('thm.extreal.le_infty|split|3','thm.bool.EQ_CLAUSES|split|2')).
fof(f5940,plain,(
  s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | $spl148),
  inference(cnf_transformation,[],[f5940_D])).
fof(f5940_D,plain,(
  s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f6013,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f3982,f835])).
fof(f3982,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | $spl0),
  inference(superposition,[],[f1080,f617])).
fof(f1080,plain,(
  s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12))) | $spl0),
  inference(resolution,[],[f832,f781])).
fof(f781,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f549])).
fof(f549,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f341])).
fof(f341,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X21] : (s(bool,'T_0') = s(bool,X21) <=> p(s(bool,X21)))),
  file('thm.extreal.le_infty|split|3','thm.bool.EQ_CLAUSES|split|0')).
fof(f6012,plain,(
  ( ! [X0] : (~sP10(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | p(s(bool,X0))) ) | $spl212),
  inference(cnf_transformation,[],[f6012_D])).
fof(f6012_D,plain,(
  ( ! [X0] : (~sP10(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | p(s(bool,X0))) ) <=> ~$spl212),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl212])])).
fof(f6011,plain,(
  ( ! [X0,X1] : (~sP10(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | p(s(bool,X0)) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f4445,f912])).
fof(f4445,plain,(
  ( ! [X0,X1] : (~sP10(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | p(s(bool,X0)) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | $spl0),
  inference(superposition,[],[f1256,f617])).
fof(f1256,plain,(
  ( ! [X144] : (~sP10(X144,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0'))) | p(s(bool,X144))) ) | $spl0),
  inference(resolution,[],[f987,f678])).
fof(f678,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | ~sP10(X0,X1)) )),
  inference(cnf_transformation,[],[f484])).
fof(f6010,plain,(
  ( ! [X2] : (~sP10(X2,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | p(s(bool,X2))) ) | $spl210),
  inference(cnf_transformation,[],[f6010_D])).
fof(f6010_D,plain,(
  ( ! [X2] : (~sP10(X2,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | p(s(bool,X2))) ) <=> ~$spl210),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl210])])).
fof(f6009,plain,(
  ( ! [X2,X3] : (~sP10(X2,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | p(s(bool,X2)) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f4446,f835])).
fof(f4446,plain,(
  ( ! [X2,X3] : (~sP10(X2,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | p(s(bool,X2)) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3)))) ) | $spl0),
  inference(superposition,[],[f1256,f617])).
fof(f6008,plain,(
  ( ! [X0] : (~sP10('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0) | p(s(bool,X0))) ) | $spl208),
  inference(cnf_transformation,[],[f6008_D])).
fof(f6008_D,plain,(
  ( ! [X0] : (~sP10('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0) | p(s(bool,X0))) ) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f6007,plain,(
  ( ! [X0,X1] : (~sP10('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0) | p(s(bool,X0)) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f4465,f912])).
fof(f4465,plain,(
  ( ! [X0,X1] : (~sP10('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0) | p(s(bool,X0)) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | $spl0),
  inference(superposition,[],[f1257,f617])).
fof(f1257,plain,(
  ( ! [X145] : (~sP10('const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0')),X145) | p(s(bool,X145))) ) | $spl0),
  inference(resolution,[],[f987,f678])).
fof(f6006,plain,(
  ( ! [X2] : (~sP10('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X2) | p(s(bool,X2))) ) | $spl206),
  inference(cnf_transformation,[],[f6006_D])).
fof(f6006_D,plain,(
  ( ! [X2] : (~sP10('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X2) | p(s(bool,X2))) ) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f6005,plain,(
  ( ! [X2,X3] : (~sP10('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X2) | p(s(bool,X2)) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f4466,f835])).
fof(f4466,plain,(
  ( ! [X2,X3] : (~sP10('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X2) | p(s(bool,X2)) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3)))) ) | $spl0),
  inference(superposition,[],[f1257,f617])).
fof(f6004,plain,(
  ( ! [X0] : (sP10(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | ~p(s(bool,X0))) ) | $spl204),
  inference(cnf_transformation,[],[f6004_D])).
fof(f6004_D,plain,(
  ( ! [X0] : (sP10(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | ~p(s(bool,X0))) ) <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f6003,plain,(
  ( ! [X0,X1] : (sP10(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | ~p(s(bool,X0)) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f4487,f912])).
fof(f4487,plain,(
  ( ! [X0,X1] : (sP10(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | ~p(s(bool,X0)) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | $spl0),
  inference(superposition,[],[f1258,f617])).
fof(f1258,plain,(
  ( ! [X146] : (sP10(X146,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0'))) | ~p(s(bool,X146))) ) | $spl0),
  inference(resolution,[],[f987,f680])).
fof(f6002,plain,(
  ( ! [X2] : (sP10(X2,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | ~p(s(bool,X2))) ) | $spl202),
  inference(cnf_transformation,[],[f6002_D])).
fof(f6002_D,plain,(
  ( ! [X2] : (sP10(X2,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | ~p(s(bool,X2))) ) <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f6001,plain,(
  ( ! [X2,X3] : (sP10(X2,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | ~p(s(bool,X2)) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f4488,f835])).
fof(f4488,plain,(
  ( ! [X2,X3] : (sP10(X2,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | ~p(s(bool,X2)) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3)))) ) | $spl0),
  inference(superposition,[],[f1258,f617])).
fof(f6000,plain,(
  ( ! [X0] : (sP10('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0) | ~p(s(bool,X0))) ) | $spl200),
  inference(cnf_transformation,[],[f6000_D])).
fof(f6000_D,plain,(
  ( ! [X0] : (sP10('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0) | ~p(s(bool,X0))) ) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f5999,plain,(
  ( ! [X0,X1] : (sP10('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0) | ~p(s(bool,X0)) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f4509,f912])).
fof(f4509,plain,(
  ( ! [X0,X1] : (sP10('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0) | ~p(s(bool,X0)) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | $spl0),
  inference(superposition,[],[f1259,f617])).
fof(f1259,plain,(
  ( ! [X147] : (sP10('const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0')),X147) | ~p(s(bool,X147))) ) | $spl0),
  inference(resolution,[],[f987,f683])).
fof(f5998,plain,(
  ( ! [X2] : (sP10('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X2) | ~p(s(bool,X2))) ) | $spl198),
  inference(cnf_transformation,[],[f5998_D])).
fof(f5998_D,plain,(
  ( ! [X2] : (sP10('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X2) | ~p(s(bool,X2))) ) <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f5997,plain,(
  ( ! [X2,X3] : (sP10('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X2) | ~p(s(bool,X2)) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f4510,f835])).
fof(f4510,plain,(
  ( ! [X2,X3] : (sP10('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X2) | ~p(s(bool,X2)) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3)))) ) | $spl0),
  inference(superposition,[],[f1259,f617])).
fof(f5996,plain,(
  ( ! [X0,X1] : (~sP5(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X1) | p(s(bool,X0))) ) | $spl196),
  inference(cnf_transformation,[],[f5996_D])).
fof(f5996_D,plain,(
  ( ! [X0,X1] : (~sP5(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X1) | p(s(bool,X0))) ) <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f5995,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X1) | p(s(bool,X0)) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f4537,f912])).
fof(f4537,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X1) | p(s(bool,X0)) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | $spl0),
  inference(superposition,[],[f1012,f617])).
fof(f1012,plain,(
  ( ! [X54,X55] : (~sP5(X54,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)),X55) | p(s(bool,X54))) ) | $spl0),
  inference(resolution,[],[f832,f622])).
fof(f622,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | ~sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f459])).
fof(f459,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP4(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP4(X0,X1,X2) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f458])).
fof(f458,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP4(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP4(X2,X0,X1) | sP5(X2,X0,X1)))),
  inference(flattening,[],[f457])).
fof(f457,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP4(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP4(X2,X0,X1)) | sP5(X2,X0,X1)))),
  inference(nnf_transformation,[],[f421])).
fof(f421,plain,(
  ! [X2,X0,X1] : (sP5(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP4(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f5994,plain,(
  ( ! [X4,X3] : (~sP5(X3,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X4) | p(s(bool,X3))) ) | $spl194),
  inference(cnf_transformation,[],[f5994_D])).
fof(f5994_D,plain,(
  ( ! [X4,X3] : (~sP5(X3,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X4) | p(s(bool,X3))) ) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f5993,plain,(
  ( ! [X4,X5,X3] : (~sP5(X3,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X4) | p(s(bool,X3)) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f4538,f835])).
fof(f4538,plain,(
  ( ! [X4,X5,X3] : (~sP5(X3,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X4) | p(s(bool,X3)) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl0),
  inference(superposition,[],[f1012,f617])).
fof(f5992,plain,(
  s(bool,'F_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | $spl192),
  inference(cnf_transformation,[],[f5992_D])).
fof(f5992_D,plain,(
  s(bool,'F_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f5990,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f4550,f912])).
fof(f4550,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))) ) | $spl0),
  inference(superposition,[],[f1280,f617])).
fof(f1280,plain,(
  s(bool,'F_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0'))) | $spl0),
  inference(resolution,[],[f987,f777])).
fof(f777,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f547])).
fof(f5989,plain,(
  s(bool,'F_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | $spl190),
  inference(cnf_transformation,[],[f5989_D])).
fof(f5989_D,plain,(
  s(bool,'F_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f5987,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f4551,f835])).
fof(f4551,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | $spl0),
  inference(superposition,[],[f1280,f617])).
fof(f5986,plain,(
  ( ! [X0,X1] : (sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0,X1) | p(s(bool,X1))) ) | $spl188),
  inference(cnf_transformation,[],[f5986_D])).
fof(f5986_D,plain,(
  ( ! [X0,X1] : (sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0,X1) | p(s(bool,X1))) ) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f5985,plain,(
  ( ! [X2,X0,X1] : (sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0,X1) | p(s(bool,X1)) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f5003,f912])).
fof(f5003,plain,(
  ( ! [X2,X0,X1] : (sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0,X1) | p(s(bool,X1)) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | $spl0),
  inference(superposition,[],[f1021,f617])).
fof(f1021,plain,(
  ( ! [X72,X73] : (sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)),X72,X73) | p(s(bool,X73))) ) | $spl0),
  inference(resolution,[],[f832,f633])).
fof(f633,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f464])).
fof(f464,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP5(X2,X0,X1)) & (~sP5(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f463])).
fof(f463,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP5(X2,X0,X1)) & (~sP5(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f422])).
fof(f422,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP5(X2,X0,X1))),
  inference(definition_folding,[],[f380,f421,f420])).
fof(f380,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f290])).
fof(f290,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f289])).
fof(f289,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f135])).
fof(f135,axiom,(
  ! [X35,X60,X61] : ((p(s(bool,X61)) <=> (p(s(bool,X60)) => p(s(bool,X35)))) <=> ((p(s(bool,X61)) | p(s(bool,X60))) & (p(s(bool,X61)) | ~p(s(bool,X35))) & (~p(s(bool,X60)) | p(s(bool,X35)) | ~p(s(bool,X61)))))),
  file('thm.extreal.le_infty|split|3','thm.sat.sat_27_prove_ltb')).
fof(f5984,plain,(
  ( ! [X4,X3] : (sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X3,X4) | p(s(bool,X4))) ) | $spl186),
  inference(cnf_transformation,[],[f5984_D])).
fof(f5984_D,plain,(
  ( ! [X4,X3] : (sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X3,X4) | p(s(bool,X4))) ) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f5983,plain,(
  ( ! [X4,X5,X3] : (sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X3,X4) | p(s(bool,X4)) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f5004,f835])).
fof(f5004,plain,(
  ( ! [X4,X5,X3] : (sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X3,X4) | p(s(bool,X4)) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl0),
  inference(superposition,[],[f1021,f617])).
fof(f5982,plain,(
  s(bool,'T_0') != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | $spl185),
  inference(cnf_transformation,[],[f5982_D])).
fof(f5982_D,plain,(
  s(bool,'T_0') != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) <=> ~$spl185),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl185])])).
fof(f5980,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f5017,f912])).
fof(f5017,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))) ) | $spl0),
  inference(superposition,[],[f1282,f617])).
fof(f1282,plain,(
  s(bool,'T_0') != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0'))) | $spl0),
  inference(resolution,[],[f987,f780])).
fof(f780,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f549])).
fof(f5979,plain,(
  s(bool,'T_0') != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | $spl183),
  inference(cnf_transformation,[],[f5979_D])).
fof(f5979_D,plain,(
  s(bool,'T_0') != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) <=> ~$spl183),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl183])])).
fof(f5977,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f5018,f835])).
fof(f5018,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | $spl0),
  inference(superposition,[],[f1282,f617])).
fof(f5976,plain,(
  ( ! [X0] : (~sP11(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | p(s(bool,X0))) ) | $spl180),
  inference(cnf_transformation,[],[f5976_D])).
fof(f5976_D,plain,(
  ( ! [X0] : (~sP11(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | p(s(bool,X0))) ) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f5975,plain,(
  ( ! [X0,X1] : (~sP11(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | p(s(bool,X0)) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f5121,f912])).
fof(f5121,plain,(
  ( ! [X0,X1] : (~sP11(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | p(s(bool,X0)) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | $spl0),
  inference(superposition,[],[f1318,f617])).
fof(f1318,plain,(
  ( ! [X152] : (~sP11(X152,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0'))) | p(s(bool,X152))) ) | $spl0),
  inference(subsumption_resolution,[],[f1264,f746])).
fof(f746,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f528])).
fof(f528,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f527])).
fof(f527,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f327])).
fof(f327,plain,(
  ! [X0] : ((p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f166])).
fof(f166,axiom,(
  ! [X21] : ((p(s(bool,'F_0')) | p(s(bool,X21))) <=> p(s(bool,X21)))),
  file('thm.extreal.le_infty|split|3','thm.bool.OR_CLAUSES|split|2')).
fof(f1264,plain,(
  ( ! [X152] : (p(s(bool,'F_0')) | p(s(bool,X152)) | ~sP11(X152,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl0),
  inference(resolution,[],[f987,f694])).
fof(f694,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP11(X0,X1)) )),
  inference(cnf_transformation,[],[f489])).
fof(f5974,plain,(
  ( ! [X2] : (~sP11(X2,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | p(s(bool,X2))) ) | $spl178),
  inference(cnf_transformation,[],[f5974_D])).
fof(f5974_D,plain,(
  ( ! [X2] : (~sP11(X2,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | p(s(bool,X2))) ) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f5973,plain,(
  ( ! [X2,X3] : (~sP11(X2,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | p(s(bool,X2)) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f5122,f835])).
fof(f5122,plain,(
  ( ! [X2,X3] : (~sP11(X2,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | p(s(bool,X2)) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3)))) ) | $spl0),
  inference(superposition,[],[f1318,f617])).
fof(f5972,plain,(
  ( ! [X0,X1] : (sP5(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X1) | ~p(s(bool,X0))) ) | $spl176),
  inference(cnf_transformation,[],[f5972_D])).
fof(f5972_D,plain,(
  ( ! [X0,X1] : (sP5(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X1) | ~p(s(bool,X0))) ) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f5971,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X1) | ~p(s(bool,X0)) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f5155,f912])).
fof(f5155,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X1) | ~p(s(bool,X0)) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | $spl0),
  inference(superposition,[],[f1022,f617])).
fof(f1022,plain,(
  ( ! [X74,X75] : (sP5(X74,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)),X75) | ~p(s(bool,X74))) ) | $spl0),
  inference(resolution,[],[f832,f634])).
fof(f634,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f464])).
fof(f5970,plain,(
  ( ! [X4,X3] : (sP5(X3,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X4) | ~p(s(bool,X3))) ) | $spl174),
  inference(cnf_transformation,[],[f5970_D])).
fof(f5970_D,plain,(
  ( ! [X4,X3] : (sP5(X3,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X4) | ~p(s(bool,X3))) ) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f5969,plain,(
  ( ! [X4,X5,X3] : (sP5(X3,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X4) | ~p(s(bool,X3)) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f5156,f835])).
fof(f5156,plain,(
  ( ! [X4,X5,X3] : (sP5(X3,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X4) | ~p(s(bool,X3)) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl0),
  inference(superposition,[],[f1022,f617])).
fof(f5968,plain,(
  ( ! [X0] : (~sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0) | p(s(bool,X0))) ) | $spl172),
  inference(cnf_transformation,[],[f5968_D])).
fof(f5968_D,plain,(
  ( ! [X0] : (~sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0) | p(s(bool,X0))) ) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f5967,plain,(
  ( ! [X0,X1] : (~sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0) | p(s(bool,X0)) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f5172,f912])).
fof(f5172,plain,(
  ( ! [X0,X1] : (~sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0) | p(s(bool,X0)) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | $spl0),
  inference(superposition,[],[f1319,f617])).
fof(f1319,plain,(
  ( ! [X153] : (~sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0')),X153) | p(s(bool,X153))) ) | $spl0),
  inference(subsumption_resolution,[],[f1265,f746])).
fof(f1265,plain,(
  ( ! [X153] : (p(s(bool,'F_0')) | p(s(bool,X153)) | ~sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0')),X153)) ) | $spl0),
  inference(resolution,[],[f987,f694])).
fof(f5966,plain,(
  ( ! [X2] : (~sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X2) | p(s(bool,X2))) ) | $spl170),
  inference(cnf_transformation,[],[f5966_D])).
fof(f5966_D,plain,(
  ( ! [X2] : (~sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X2) | p(s(bool,X2))) ) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f5965,plain,(
  ( ! [X2,X3] : (~sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X2) | p(s(bool,X2)) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f5173,f835])).
fof(f5173,plain,(
  ( ! [X2,X3] : (~sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X2) | p(s(bool,X2)) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3)))) ) | $spl0),
  inference(superposition,[],[f1319,f617])).
fof(f5964,plain,(
  ( ! [X0] : (~sP10(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | p(s(bool,X0))) ) | $spl168),
  inference(cnf_transformation,[],[f5964_D])).
fof(f5964_D,plain,(
  ( ! [X0] : (~sP10(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | p(s(bool,X0))) ) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f5963,plain,(
  ( ! [X0,X1] : (~sP10(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | p(s(bool,X0)) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f5198,f912])).
fof(f5198,plain,(
  ( ! [X0,X1] : (~sP10(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | p(s(bool,X0)) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | $spl0),
  inference(superposition,[],[f1870,f617])).
fof(f1870,plain,(
  ( ! [X144] : (~sP10(X144,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.NegInf_0'))) | p(s(bool,X144))) ) | $spl0),
  inference(resolution,[],[f1131,f678])).
fof(f5962,plain,(
  ( ! [X2] : (~sP10(X2,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | p(s(bool,X2))) ) | $spl166),
  inference(cnf_transformation,[],[f5962_D])).
fof(f5962_D,plain,(
  ( ! [X2] : (~sP10(X2,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | p(s(bool,X2))) ) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f5961,plain,(
  ( ! [X2,X3] : (~sP10(X2,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | p(s(bool,X2)) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f5199,f835])).
fof(f5199,plain,(
  ( ! [X2,X3] : (~sP10(X2,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | p(s(bool,X2)) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3)))) ) | $spl0),
  inference(superposition,[],[f1870,f617])).
fof(f5960,plain,(
  ( ! [X0] : (~sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X0) | p(s(bool,X0))) ) | $spl164),
  inference(cnf_transformation,[],[f5960_D])).
fof(f5960_D,plain,(
  ( ! [X0] : (~sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X0) | p(s(bool,X0))) ) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f5959,plain,(
  ( ! [X0,X1] : (~sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X0) | p(s(bool,X0)) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f5226,f912])).
fof(f5226,plain,(
  ( ! [X0,X1] : (~sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X0) | p(s(bool,X0)) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | $spl0),
  inference(superposition,[],[f1871,f617])).
fof(f1871,plain,(
  ( ! [X145] : (~sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.NegInf_0')),X145) | p(s(bool,X145))) ) | $spl0),
  inference(resolution,[],[f1131,f678])).
fof(f5958,plain,(
  ( ! [X2] : (~sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X2) | p(s(bool,X2))) ) | $spl162),
  inference(cnf_transformation,[],[f5958_D])).
fof(f5958_D,plain,(
  ( ! [X2] : (~sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X2) | p(s(bool,X2))) ) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f5957,plain,(
  ( ! [X2,X3] : (~sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X2) | p(s(bool,X2)) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f5227,f835])).
fof(f5227,plain,(
  ( ! [X2,X3] : (~sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X2) | p(s(bool,X2)) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3)))) ) | $spl0),
  inference(superposition,[],[f1871,f617])).
fof(f5956,plain,(
  ( ! [X0,X1] : (sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0,X1) | ~p(s(bool,X0))) ) | $spl160),
  inference(cnf_transformation,[],[f5956_D])).
fof(f5956_D,plain,(
  ( ! [X0,X1] : (sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0,X1) | ~p(s(bool,X0))) ) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f5955,plain,(
  ( ! [X2,X0,X1] : (sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0,X1) | ~p(s(bool,X0)) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f5261,f912])).
fof(f5261,plain,(
  ( ! [X2,X0,X1] : (sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0,X1) | ~p(s(bool,X0)) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | $spl0),
  inference(superposition,[],[f1023,f617])).
fof(f1023,plain,(
  ( ! [X76,X77] : (sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)),X76,X77) | ~p(s(bool,X76))) ) | $spl0),
  inference(resolution,[],[f832,f634])).
fof(f5954,plain,(
  ( ! [X4,X3] : (sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X3,X4) | ~p(s(bool,X3))) ) | $spl158),
  inference(cnf_transformation,[],[f5954_D])).
fof(f5954_D,plain,(
  ( ! [X4,X3] : (sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X3,X4) | ~p(s(bool,X3))) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f5953,plain,(
  ( ! [X4,X5,X3] : (sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X3,X4) | ~p(s(bool,X3)) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f5262,f835])).
fof(f5262,plain,(
  ( ! [X4,X5,X3] : (sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X3,X4) | ~p(s(bool,X3)) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl0),
  inference(superposition,[],[f1023,f617])).
fof(f5952,plain,(
  ( ! [X0] : (sP10(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | ~p(s(bool,X0))) ) | $spl156),
  inference(cnf_transformation,[],[f5952_D])).
fof(f5952_D,plain,(
  ( ! [X0] : (sP10(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | ~p(s(bool,X0))) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f5951,plain,(
  ( ! [X0,X1] : (sP10(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | ~p(s(bool,X0)) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f5282,f912])).
fof(f5282,plain,(
  ( ! [X0,X1] : (sP10(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | ~p(s(bool,X0)) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | $spl0),
  inference(superposition,[],[f1872,f617])).
fof(f1872,plain,(
  ( ! [X146] : (sP10(X146,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.NegInf_0'))) | ~p(s(bool,X146))) ) | $spl0),
  inference(resolution,[],[f1131,f680])).
fof(f5950,plain,(
  ( ! [X2] : (sP10(X2,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | ~p(s(bool,X2))) ) | $spl154),
  inference(cnf_transformation,[],[f5950_D])).
fof(f5950_D,plain,(
  ( ! [X2] : (sP10(X2,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | ~p(s(bool,X2))) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f5949,plain,(
  ( ! [X2,X3] : (sP10(X2,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | ~p(s(bool,X2)) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f5283,f835])).
fof(f5283,plain,(
  ( ! [X2,X3] : (sP10(X2,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | ~p(s(bool,X2)) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3)))) ) | $spl0),
  inference(superposition,[],[f1872,f617])).
fof(f5948,plain,(
  ( ! [X0] : (sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X0) | ~p(s(bool,X0))) ) | $spl152),
  inference(cnf_transformation,[],[f5948_D])).
fof(f5948_D,plain,(
  ( ! [X0] : (sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X0) | ~p(s(bool,X0))) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f5947,plain,(
  ( ! [X0,X1] : (sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X0) | ~p(s(bool,X0)) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f5312,f912])).
fof(f5312,plain,(
  ( ! [X0,X1] : (sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X0) | ~p(s(bool,X0)) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | $spl0),
  inference(superposition,[],[f1873,f617])).
fof(f1873,plain,(
  ( ! [X147] : (sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.NegInf_0')),X147) | ~p(s(bool,X147))) ) | $spl0),
  inference(resolution,[],[f1131,f683])).
fof(f5946,plain,(
  ( ! [X2] : (sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X2) | ~p(s(bool,X2))) ) | $spl150),
  inference(cnf_transformation,[],[f5946_D])).
fof(f5946_D,plain,(
  ( ! [X2] : (sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X2) | ~p(s(bool,X2))) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f5945,plain,(
  ( ! [X2,X3] : (sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X2) | ~p(s(bool,X2)) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f5313,f835])).
fof(f5313,plain,(
  ( ! [X2,X3] : (sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X2) | ~p(s(bool,X2)) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3)))) ) | $spl0),
  inference(superposition,[],[f1873,f617])).
fof(f5944,plain,(
  ( ! [X1] : (s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f5943,f3955])).
fof(f3955,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl0),
  inference(subsumption_resolution,[],[f3949,f832])).
fof(f3949,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)))) | $spl0),
  inference(superposition,[],[f1078,f781])).
fof(f5943,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,'F_0') | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | ($spl0 | $spl3)),
  inference(forward_demodulation,[],[f5942,f814])).
fof(f814,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)))) )),
  inference(cnf_transformation,[],[f369])).
fof(f369,plain,(
  ! [X0] : s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X11] : s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X11)))),
  file('thm.extreal.le_infty|split|3','thm.extreal.extreal_le_def|split|1')).
fof(f5942,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f5336,f835])).
fof(f5336,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | $spl0),
  inference(superposition,[],[f1894,f617])).
fof(f1894,plain,(
  s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.NegInf_0'))) | $spl0),
  inference(resolution,[],[f1131,f777])).
fof(f5941,plain,(
  s(bool,'T_0') != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | $spl149),
  inference(cnf_transformation,[],[f5941_D])).
fof(f5941_D,plain,(
  s(bool,'T_0') != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) <=> ~$spl149),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl149])])).
fof(f5939,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f5784,f912])).
fof(f5784,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))) ) | $spl0),
  inference(superposition,[],[f1896,f617])).
fof(f1896,plain,(
  s(bool,'T_0') != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.NegInf_0'))) | $spl0),
  inference(resolution,[],[f1131,f780])).
fof(f5938,plain,(
  ( ! [X1] : (s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f5937,f835])).
fof(f5937,plain,(
  ( ! [X1] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | $spl0),
  inference(subsumption_resolution,[],[f5785,f814])).
fof(f5785,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | $spl0),
  inference(superposition,[],[f1896,f617])).
fof(f5936,plain,(
  ( ! [X0,X1] : (~sP7(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | p(s(bool,X0))) ) | $spl146),
  inference(cnf_transformation,[],[f5936_D])).
fof(f5936_D,plain,(
  ( ! [X0,X1] : (~sP7(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | p(s(bool,X0))) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f5935,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | p(s(bool,X0)) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f5838,f912])).
fof(f5838,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | p(s(bool,X0)) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | $spl0),
  inference(superposition,[],[f1027,f617])).
fof(f1027,plain,(
  ( ! [X85,X84] : (~sP7(X84,X85,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12))) | p(s(bool,X84))) ) | $spl0),
  inference(resolution,[],[f832,f638])).
fof(f638,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X0)) | ~sP7(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f467])).
fof(f467,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP6(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP6(X0,X1,X2) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f466])).
fof(f466,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP6(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP6(X2,X0,X1) | sP7(X2,X0,X1)))),
  inference(flattening,[],[f465])).
fof(f465,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP6(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP6(X2,X0,X1)) | sP7(X2,X0,X1)))),
  inference(nnf_transformation,[],[f424])).
fof(f424,plain,(
  ! [X2,X0,X1] : (sP7(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP6(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f5934,plain,(
  ( ! [X4,X3] : (~sP7(X3,X4,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | p(s(bool,X3))) ) | $spl144),
  inference(cnf_transformation,[],[f5934_D])).
fof(f5934_D,plain,(
  ( ! [X4,X3] : (~sP7(X3,X4,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | p(s(bool,X3))) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f5933,plain,(
  ( ! [X4,X5,X3] : (~sP7(X3,X4,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | p(s(bool,X3)) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f5839,f835])).
fof(f5839,plain,(
  ( ! [X4,X5,X3] : (~sP7(X3,X4,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | p(s(bool,X3)) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl0),
  inference(superposition,[],[f1027,f617])).
fof(f5932,plain,(
  ( ! [X0] : (~sP11(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | p(s(bool,X0))) ) | $spl142),
  inference(cnf_transformation,[],[f5932_D])).
fof(f5932_D,plain,(
  ( ! [X0] : (~sP11(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | p(s(bool,X0))) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f5931,plain,(
  ( ! [X0,X1] : (~sP11(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | p(s(bool,X0)) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f5860,f912])).
fof(f5860,plain,(
  ( ! [X0,X1] : (~sP11(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | p(s(bool,X0)) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | $spl0),
  inference(superposition,[],[f1937,f617])).
fof(f1937,plain,(
  ( ! [X152] : (~sP11(X152,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.NegInf_0'))) | p(s(bool,X152))) ) | $spl0),
  inference(subsumption_resolution,[],[f1878,f746])).
fof(f1878,plain,(
  ( ! [X152] : (p(s(bool,'F_0')) | p(s(bool,X152)) | ~sP11(X152,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl0),
  inference(resolution,[],[f1131,f694])).
fof(f5930,plain,(
  ( ! [X2] : (~sP11(X2,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | p(s(bool,X2))) ) | $spl140),
  inference(cnf_transformation,[],[f5930_D])).
fof(f5930_D,plain,(
  ( ! [X2] : (~sP11(X2,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | p(s(bool,X2))) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f5929,plain,(
  ( ! [X2,X3] : (~sP11(X2,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | p(s(bool,X2)) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f5861,f835])).
fof(f5861,plain,(
  ( ! [X2,X3] : (~sP11(X2,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | p(s(bool,X2)) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3)))) ) | $spl0),
  inference(superposition,[],[f1937,f617])).
fof(f5928,plain,(
  ( ! [X0] : (~sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X0) | p(s(bool,X0))) ) | $spl138),
  inference(cnf_transformation,[],[f5928_D])).
fof(f5928_D,plain,(
  ( ! [X0] : (~sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X0) | p(s(bool,X0))) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f5927,plain,(
  ( ! [X0,X1] : (~sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X0) | p(s(bool,X0)) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f5896,f912])).
fof(f5896,plain,(
  ( ! [X0,X1] : (~sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X0) | p(s(bool,X0)) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | $spl0),
  inference(superposition,[],[f1938,f617])).
fof(f1938,plain,(
  ( ! [X153] : (~sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.NegInf_0')),X153) | p(s(bool,X153))) ) | $spl0),
  inference(subsumption_resolution,[],[f1879,f746])).
fof(f1879,plain,(
  ( ! [X153] : (p(s(bool,'F_0')) | p(s(bool,X153)) | ~sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.NegInf_0')),X153)) ) | $spl0),
  inference(resolution,[],[f1131,f694])).
fof(f5926,plain,(
  ( ! [X2] : (~sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X2) | p(s(bool,X2))) ) | $spl136),
  inference(cnf_transformation,[],[f5926_D])).
fof(f5926_D,plain,(
  ( ! [X2] : (~sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X2) | p(s(bool,X2))) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f5925,plain,(
  ( ! [X2,X3] : (~sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X2) | p(s(bool,X2)) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f5897,f835])).
fof(f5897,plain,(
  ( ! [X2,X3] : (~sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X2) | p(s(bool,X2)) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3)))) ) | $spl0),
  inference(superposition,[],[f1938,f617])).
fof(f5924,plain,(
  $false | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f5923,f3955])).
fof(f5923,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f1978,f1080])).
fof(f1978,plain,(
  s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12))) | $spl6),
  inference(superposition,[],[f811,f911])).
fof(f811,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))))) )),
  inference(cnf_transformation,[],[f366])).
fof(f366,plain,(
  ! [X0] : s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X22] : s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X22)))))),
  file('thm.extreal.le_infty|split|3','thm.extreal.extreal_le_def|split|6')).
fof(f1116,plain,(
  ( ! [X4,X3] : (p(s(bool,X3)) | s(bool,X3) != s(bool,X4) | ~p(s(bool,X4))) ) | $spl20),
  inference(cnf_transformation,[],[f1116_D])).
fof(f1116_D,plain,(
  ( ! [X4,X3] : (p(s(bool,X3)) | s(bool,X3) != s(bool,X4) | ~p(s(bool,X4))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f2815,plain,(
  ( ! [X29] : (s(bool,'F_0') != s(bool,X29)) ) | $spl132),
  inference(cnf_transformation,[],[f2815_D])).
fof(f2815_D,plain,(
  ( ! [X29] : (s(bool,'F_0') != s(bool,X29)) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f5662,plain,(
  ( ! [X85,X86,X84] : (s(bool,'F_0') != s(bool,X84) | p(s(bool,X85)) | ~p(s(bool,X86)) | s(bool,X85) != s(bool,X86)) ) | $spl0),
  inference(subsumption_resolution,[],[f5395,f572])).
fof(f572,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f436])).
fof(f436,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f435])).
fof(f435,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f375])).
fof(f375,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f374])).
fof(f374,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f264])).
fof(f264,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f184])).
fof(f184,axiom,(
  ! [X6,X64,X12,X65] : ((s(bool,X6) = s(bool,X64) & (p(s(bool,X64)) => s(bool,X12) = s(bool,X65))) => ((p(s(bool,X6)) => p(s(bool,X12))) <=> (p(s(bool,X64)) => p(s(bool,X65)))))),
  file('thm.extreal.le_infty|split|3','thm.bool.IMP_CONG')).
fof(f5395,plain,(
  ( ! [X85,X86,X84] : (s(bool,'F_0') != s(bool,X84) | p(s(bool,X85)) | ~p(s(bool,X84)) | ~p(s(bool,X86)) | s(bool,X85) != s(bool,X86)) ) | $spl0),
  inference(superposition,[],[f574,f1894])).
fof(f574,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f436])).
fof(f5661,plain,(
  ( ! [X83,X81,X82] : (s(bool,'F_0') != s(bool,X81) | p(s(bool,X82)) | ~p(s(bool,X83)) | s(bool,X82) != s(bool,X83)) ) | $spl0),
  inference(subsumption_resolution,[],[f5394,f824])).
fof(f824,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f571])).
fof(f571,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f436])).
fof(f5394,plain,(
  ( ! [X83,X81,X82] : (s(bool,'F_0') != s(bool,X81) | p(s(bool,X82)) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.NegInf_0')))) | ~p(s(bool,X83)) | s(bool,X82) != s(bool,X83)) ) | $spl0),
  inference(superposition,[],[f574,f1894])).
fof(f5644,plain,(
  ( ! [X26,X24,X25] : (s(bool,'F_0') != s(bool,X24) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X26)) ) | $spl0),
  inference(subsumption_resolution,[],[f5375,f824])).
fof(f5375,plain,(
  ( ! [X26,X24,X25] : (s(bool,'F_0') != s(bool,X24) | p(s(bool,X25)) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.NegInf_0')))) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X26)) ) | $spl0),
  inference(superposition,[],[f570,f1894])).
fof(f570,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f436])).
fof(f5643,plain,(
  ( ! [X23,X21,X22] : (s(bool,'F_0') != s(bool,X21) | p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X23)) ) | $spl0),
  inference(subsumption_resolution,[],[f5374,f572])).
fof(f5374,plain,(
  ( ! [X23,X21,X22] : (s(bool,'F_0') != s(bool,X21) | p(s(bool,X22)) | ~p(s(bool,X21)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X23)) ) | $spl0),
  inference(superposition,[],[f570,f1894])).
fof(f5100,plain,(
  $false | ($spl0 | $spl40)),
  inference(trivial_inequality_removal,[],[f5068])).
fof(f5068,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl0 | $spl40)),
  inference(superposition,[],[f1316,f1280])).
fof(f5099,plain,(
  $false | $spl40),
  inference(equality_resolution,[],[f1316])).
fof(f4868,plain,(
  ( ! [X85,X86,X84] : (s(bool,'F_0') != s(bool,X84) | p(s(bool,X85)) | ~p(s(bool,X86)) | s(bool,X85) != s(bool,X86)) ) | $spl0),
  inference(subsumption_resolution,[],[f4603,f572])).
fof(f4603,plain,(
  ( ! [X85,X86,X84] : (s(bool,'F_0') != s(bool,X84) | p(s(bool,X85)) | ~p(s(bool,X84)) | ~p(s(bool,X86)) | s(bool,X85) != s(bool,X86)) ) | $spl0),
  inference(superposition,[],[f574,f1280])).
fof(f4867,plain,(
  ( ! [X83,X81,X82] : (s(bool,'F_0') != s(bool,X81) | p(s(bool,X82)) | ~p(s(bool,X83)) | s(bool,X82) != s(bool,X83)) ) | $spl0),
  inference(subsumption_resolution,[],[f4602,f824])).
fof(f4602,plain,(
  ( ! [X83,X81,X82] : (s(bool,'F_0') != s(bool,X81) | p(s(bool,X82)) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0')))) | ~p(s(bool,X83)) | s(bool,X82) != s(bool,X83)) ) | $spl0),
  inference(superposition,[],[f574,f1280])).
fof(f4850,plain,(
  ( ! [X26,X24,X25] : (s(bool,'F_0') != s(bool,X24) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X26)) ) | $spl0),
  inference(subsumption_resolution,[],[f4583,f824])).
fof(f4583,plain,(
  ( ! [X26,X24,X25] : (s(bool,'F_0') != s(bool,X24) | p(s(bool,X25)) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0')))) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X26)) ) | $spl0),
  inference(superposition,[],[f570,f1280])).
fof(f4849,plain,(
  ( ! [X23,X21,X22] : (s(bool,'F_0') != s(bool,X21) | p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X23)) ) | $spl0),
  inference(subsumption_resolution,[],[f4582,f572])).
fof(f4582,plain,(
  ( ! [X23,X21,X22] : (s(bool,'F_0') != s(bool,X21) | p(s(bool,X22)) | ~p(s(bool,X21)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X23)) ) | $spl0),
  inference(superposition,[],[f570,f1280])).
fof(f4427,plain,(
  $false | ($spl0 | $spl18)),
  inference(trivial_inequality_removal,[],[f4395])).
fof(f4395,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl0 | $spl18)),
  inference(superposition,[],[f1115,f1080])).
fof(f4426,plain,(
  $false | $spl18),
  inference(equality_resolution,[],[f1115])).
fof(f3006,plain,(
  ( ! [X33] : (s(bool,'T_0') != s(bool,X33)) ) | $spl134),
  inference(cnf_transformation,[],[f3006_D])).
fof(f3006_D,plain,(
  ( ! [X33] : (s(bool,'T_0') != s(bool,X33)) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f4288,plain,(
  ( ! [X97,X98,X96] : (s(bool,'T_0') != s(bool,X98) | p(s(bool,X96)) | ~p(s(bool,X97)) | s(bool,X96) != s(bool,X97)) ) | $spl0),
  inference(forward_demodulation,[],[f4287,f1080])).
fof(f4287,plain,(
  ( ! [X97,X98,X96] : (p(s(bool,X96)) | ~p(s(bool,X97)) | s(bool,X96) != s(bool,X97) | s(bool,X98) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)))) ) | $spl0),
  inference(subsumption_resolution,[],[f4032,f824])).
fof(f4032,plain,(
  ( ! [X97,X98,X96] : (~p(s(bool,'T_0')) | p(s(bool,X96)) | ~p(s(bool,X97)) | s(bool,X96) != s(bool,X97) | s(bool,X98) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)))) ) | $spl0),
  inference(superposition,[],[f574,f1080])).
fof(f4284,plain,(
  ( ! [X85,X86,X84] : (s(bool,'T_0') != s(bool,X84) | p(s(bool,X85)) | ~p(s(bool,X86)) | s(bool,X85) != s(bool,X86)) ) | $spl0),
  inference(subsumption_resolution,[],[f4028,f780])).
fof(f4028,plain,(
  ( ! [X85,X86,X84] : (s(bool,'T_0') != s(bool,X84) | p(s(bool,X85)) | ~p(s(bool,X84)) | ~p(s(bool,X86)) | s(bool,X85) != s(bool,X86)) ) | $spl0),
  inference(superposition,[],[f574,f1080])).
fof(f4283,plain,(
  ( ! [X83,X81,X82] : (s(bool,'T_0') != s(bool,X81) | p(s(bool,X82)) | ~p(s(bool,X83)) | s(bool,X82) != s(bool,X83)) ) | $spl0),
  inference(subsumption_resolution,[],[f4027,f824])).
fof(f4027,plain,(
  ( ! [X83,X81,X82] : (s(bool,'T_0') != s(bool,X81) | p(s(bool,X82)) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)))) | ~p(s(bool,X83)) | s(bool,X82) != s(bool,X83)) ) | $spl0),
  inference(superposition,[],[f574,f1080])).
fof(f4271,plain,(
  ( ! [X37,X38,X36] : (s(bool,'T_0') != s(bool,X38) | p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,X37)) ) | $spl0),
  inference(forward_demodulation,[],[f4270,f1080])).
fof(f4270,plain,(
  ( ! [X37,X38,X36] : (p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,X37) | s(bool,X38) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)))) ) | $spl0),
  inference(subsumption_resolution,[],[f4012,f824])).
fof(f4012,plain,(
  ( ! [X37,X38,X36] : (~p(s(bool,'T_0')) | p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,X37) | s(bool,X38) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)))) ) | $spl0),
  inference(superposition,[],[f570,f1080])).
fof(f4267,plain,(
  ( ! [X26,X24,X25] : (s(bool,'T_0') != s(bool,X24) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X26)) ) | $spl0),
  inference(subsumption_resolution,[],[f4008,f824])).
fof(f4008,plain,(
  ( ! [X26,X24,X25] : (s(bool,'T_0') != s(bool,X24) | p(s(bool,X25)) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)))) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X26)) ) | $spl0),
  inference(superposition,[],[f570,f1080])).
fof(f4266,plain,(
  ( ! [X23,X21,X22] : (s(bool,'T_0') != s(bool,X21) | p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X23)) ) | $spl0),
  inference(subsumption_resolution,[],[f4007,f780])).
fof(f4007,plain,(
  ( ! [X23,X21,X22] : (s(bool,'T_0') != s(bool,X21) | p(s(bool,X22)) | ~p(s(bool,X21)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X23)) ) | $spl0),
  inference(superposition,[],[f570,f1080])).
fof(f3522,plain,(
  $false | $spl134),
  inference(trivial_inequality_removal,[],[f3508])).
fof(f3508,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl134),
  inference(superposition,[],[f3006,f814])).
fof(f3521,plain,(
  $false | $spl134),
  inference(trivial_inequality_removal,[],[f3509])).
fof(f3509,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl134),
  inference(superposition,[],[f3006,f813])).
fof(f813,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0')))) )),
  inference(cnf_transformation,[],[f368])).
fof(f368,plain,(
  ! [X0] : s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0')))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X36] : s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X36))),s('type.extreal.extreal','const.extreal.PosInf_0')))),
  file('thm.extreal.le_infty|split|3','thm.extreal.extreal_le_def|split|3')).
fof(f3520,plain,(
  $false | $spl134),
  inference(trivial_inequality_removal,[],[f3510])).
fof(f3510,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl134),
  inference(superposition,[],[f3006,f819])).
fof(f819,plain,(
  s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))),
  inference(cnf_transformation,[],[f21])).
fof(f21,axiom,(
  s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))),
  file('thm.extreal.le_infty|split|3','thm.extreal.extreal_le_def|split|2')).
fof(f3519,plain,(
  $false | $spl134),
  inference(equality_resolution,[],[f3006])).
fof(f3479,plain,(
  $false | $spl22),
  inference(subsumption_resolution,[],[f3399,f3001])).
fof(f3001,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl22),
  inference(equality_resolution,[],[f1117])).
fof(f3399,plain,(
  ( ! [X350,X349] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X349),s('type.extreal.extreal',X350))))) ) | $spl22),
  inference(resolution,[],[f3001,f807])).
fof(f3478,plain,(
  $false | $spl22),
  inference(subsumption_resolution,[],[f3373,f3001])).
fof(f3373,plain,(
  ( ! [X280,X279] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X279),s('type.extreal.extreal',X280))))) ) | $spl22),
  inference(resolution,[],[f3001,f807])).
fof(f3473,plain,(
  $false | $spl22),
  inference(subsumption_resolution,[],[f3362,f3001])).
fof(f3362,plain,(
  p(s(bool,'F_0')) | $spl22),
  inference(resolution,[],[f3001,f802])).
fof(f802,plain,(
  p(s(bool,'F_0')) | p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f562])).
fof(f562,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f361])).
fof(f361,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f45])).
fof(f45,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.extreal.le_infty|split|3','thm.bool.NOT_CLAUSES|split|1')).
fof(f3472,plain,(
  $false | $spl22),
  inference(subsumption_resolution,[],[f3361,f3001])).
fof(f3361,plain,(
  p(s(bool,'F_0')) | $spl22),
  inference(resolution,[],[f3001,f800])).
fof(f800,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f561])).
fof(f561,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f360])).
fof(f360,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f44])).
fof(f44,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.extreal.le_infty|split|3','thm.bool.NOT_CLAUSES|split|2')).
fof(f3471,plain,(
  $false | $spl22),
  inference(subsumption_resolution,[],[f3355,f3001])).
fof(f3355,plain,(
  ( ! [X267] : (p(s(bool,X267))) ) | $spl22),
  inference(resolution,[],[f3001,f743])).
fof(f743,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f526])).
fof(f526,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))))),
  inference(flattening,[],[f525])).
fof(f525,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f396])).
fof(f396,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f326])).
fof(f326,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f120])).
fof(f120,axiom,(
  ! [X21] : ((p(s(bool,X21)) => p(s(bool,X21))) <=> p(s(bool,'T_0')))),
  file('thm.extreal.le_infty|split|3','thm.bool.IMP_CLAUSES|split|3')).
fof(f3470,plain,(
  $false | $spl22),
  inference(subsumption_resolution,[],[f3354,f3001])).
fof(f3354,plain,(
  ( ! [X266] : (p(s(bool,X266))) ) | $spl22),
  inference(resolution,[],[f3001,f740])).
fof(f740,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f524])).
fof(f524,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,'T_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f523])).
fof(f523,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (~p(s(bool,'T_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f395])).
fof(f395,plain,(
  ! [X0] : ((~p(s(bool,'T_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(ennf_transformation,[],[f325])).
fof(f325,plain,(
  ! [X0] : ((p(s(bool,'T_0')) => p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f117])).
fof(f117,axiom,(
  ! [X21] : ((p(s(bool,'T_0')) => p(s(bool,X21))) <=> p(s(bool,X21)))),
  file('thm.extreal.le_infty|split|3','thm.bool.IMP_CLAUSES|split|0')).
fof(f3469,plain,(
  $false | $spl22),
  inference(subsumption_resolution,[],[f3352,f3001])).
fof(f3352,plain,(
  p(s(bool,'F_0')) | $spl22),
  inference(resolution,[],[f3001,f737])).
fof(f737,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f522])).
fof(f522,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f521])).
fof(f521,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f394])).
fof(f394,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f324])).
fof(f324,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f119])).
fof(f119,axiom,(
  ! [X21] : ((p(s(bool,'F_0')) => p(s(bool,X21))) <=> p(s(bool,'T_0')))),
  file('thm.extreal.le_infty|split|3','thm.bool.IMP_CLAUSES|split|2')).
fof(f3468,plain,(
  $false | $spl22),
  inference(subsumption_resolution,[],[f3351,f3001])).
fof(f3351,plain,(
  ( ! [X264] : (p(s(bool,X264))) ) | $spl22),
  inference(resolution,[],[f3001,f734])).
fof(f734,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f520])).
fof(f520,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,'T_0'))))),
  inference(flattening,[],[f519])).
fof(f519,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f393])).
fof(f393,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f323])).
fof(f323,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f118])).
fof(f118,axiom,(
  ! [X21] : ((p(s(bool,X21)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  file('thm.extreal.le_infty|split|3','thm.bool.IMP_CLAUSES|split|1')).
fof(f3363,plain,(
  $false | $spl22),
  inference(resolution,[],[f3001,f804])).
fof(f804,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f33])).
fof(f33,axiom,(
  p(s(bool,'T_0'))),
  file('thm.extreal.le_infty|split|3','thm.bool.TRUTH')).
fof(f3364,plain,(
  $false | $spl22),
  inference(resolution,[],[f3001,f805])).
fof(f805,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.extreal.le_infty|split|3',aHL_TRUTH)).
fof(f3365,plain,(
  $false | $spl22),
  inference(resolution,[],[f3001,f828])).
fof(f828,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f727])).
fof(f727,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f515])).
fof(f515,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f318])).
fof(f318,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f317])).
fof(f317,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X7] : ! [X6] : (s(X7,X6) = s(X7,X6) <=> p(s(bool,'T_0')))),
  file('thm.extreal.le_infty|split|3','thm.bool.REFL_CLAUSE')).
fof(f3372,plain,(
  $false | $spl22),
  inference(resolution,[],[f3001,f816])).
fof(f816,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0))))) )),
  inference(cnf_transformation,[],[f371])).
fof(f371,plain,(
  ! [X0] : p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0))))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X6] : p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X6))))),
  file('thm.extreal.le_infty|split|3','thm.extreal.le_refl')).
fof(f3374,plain,(
  $false | $spl22),
  inference(resolution,[],[f3001,f817])).
fof(f817,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.PosInf_0'))))) )),
  inference(cnf_transformation,[],[f372])).
fof(f372,plain,(
  ! [X0] : p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.PosInf_0'))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X6] : p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal','const.extreal.PosInf_0'))))),
  file('thm.extreal.le_infty|split|3','thm.extreal.le_infty|split|1')).
fof(f3376,plain,(
  $false | $spl22),
  inference(resolution,[],[f3001,f815])).
fof(f815,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0))))) )),
  inference(cnf_transformation,[],[f370])).
fof(f370,plain,(
  ! [X0] : p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0))))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X6] : p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X6))))),
  file('thm.extreal.le_infty|split|3','thm.extreal.le_infty|split|0')).
fof(f3377,plain,(
  $false | ($spl0 | $spl22)),
  inference(resolution,[],[f3001,f832])).
fof(f3400,plain,(
  $false | ($spl0 | $spl22)),
  inference(resolution,[],[f3001,f1800])).
fof(f3249,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f3169,f1118])).
fof(f3169,plain,(
  ( ! [X350,X349] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X349),s('type.extreal.extreal',X350))))) ) | $spl24),
  inference(resolution,[],[f1118,f807])).
fof(f3248,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f3143,f1118])).
fof(f3143,plain,(
  ( ! [X280,X279] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X279),s('type.extreal.extreal',X280))))) ) | $spl24),
  inference(resolution,[],[f1118,f807])).
fof(f3243,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f3132,f1118])).
fof(f3132,plain,(
  p(s(bool,'F_0')) | $spl24),
  inference(resolution,[],[f1118,f802])).
fof(f3242,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f3131,f1118])).
fof(f3131,plain,(
  p(s(bool,'F_0')) | $spl24),
  inference(resolution,[],[f1118,f800])).
fof(f3241,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f3125,f1118])).
fof(f3125,plain,(
  ( ! [X267] : (p(s(bool,X267))) ) | $spl24),
  inference(resolution,[],[f1118,f743])).
fof(f3240,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f3124,f1118])).
fof(f3124,plain,(
  ( ! [X266] : (p(s(bool,X266))) ) | $spl24),
  inference(resolution,[],[f1118,f740])).
fof(f3239,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f3122,f1118])).
fof(f3122,plain,(
  p(s(bool,'F_0')) | $spl24),
  inference(resolution,[],[f1118,f737])).
fof(f3238,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f3121,f1118])).
fof(f3121,plain,(
  ( ! [X264] : (p(s(bool,X264))) ) | $spl24),
  inference(resolution,[],[f1118,f734])).
fof(f3133,plain,(
  $false | $spl24),
  inference(resolution,[],[f1118,f804])).
fof(f3134,plain,(
  $false | $spl24),
  inference(resolution,[],[f1118,f805])).
fof(f3135,plain,(
  $false | $spl24),
  inference(resolution,[],[f1118,f828])).
fof(f3142,plain,(
  $false | $spl24),
  inference(resolution,[],[f1118,f816])).
fof(f3144,plain,(
  $false | $spl24),
  inference(resolution,[],[f1118,f817])).
fof(f3146,plain,(
  $false | $spl24),
  inference(resolution,[],[f1118,f815])).
fof(f3147,plain,(
  $false | ($spl0 | $spl24)),
  inference(resolution,[],[f1118,f832])).
fof(f3170,plain,(
  $false | ($spl0 | $spl24)),
  inference(resolution,[],[f1118,f1800])).
fof(f1118,plain,(
  ( ! [X13] : (~p(s(bool,X13))) ) | $spl24),
  inference(cnf_transformation,[],[f1118_D])).
fof(f1118_D,plain,(
  ( ! [X13] : (~p(s(bool,X13))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f3002,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') != s(bool,X35) | ~p(s(bool,X34))) ) | $spl22),
  inference(duplicate_literal_removal,[],[f2958])).
fof(f2958,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') != s(bool,X35) | ~p(s(bool,X34)) | ~p(s(bool,X34))) ) | $spl22),
  inference(superposition,[],[f1117,f783])).
fof(f783,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f550])).
fof(f550,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f342])).
fof(f342,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X21] : (s(bool,'T_0') = s(bool,X21) <=> p(s(bool,X21)))),
  file('thm.extreal.le_infty|split|3','thm.bool.EQ_CLAUSES|split|1')).
fof(f3003,plain,(
  ( ! [X33,X32] : (s(bool,'T_0') != s(bool,X33) | ~p(s(bool,X32))) ) | $spl22),
  inference(duplicate_literal_removal,[],[f2957])).
fof(f2957,plain,(
  ( ! [X33,X32] : (s(bool,'T_0') != s(bool,X33) | ~p(s(bool,X32)) | ~p(s(bool,X32))) ) | $spl22),
  inference(superposition,[],[f1117,f781])).
fof(f1119,plain,(
  ( ! [X12] : (p(s(bool,X12))) ) | $spl26),
  inference(cnf_transformation,[],[f1119_D])).
fof(f1119_D,plain,(
  ( ! [X12] : (p(s(bool,X12))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f2811,plain,(
  ( ! [X30,X31] : (s(bool,'F_0') != s(bool,X31) | p(s(bool,X30))) ) | $spl16),
  inference(duplicate_literal_removal,[],[f2765])).
fof(f2765,plain,(
  ( ! [X30,X31] : (s(bool,'F_0') != s(bool,X31) | p(s(bool,X30)) | p(s(bool,X30))) ) | $spl16),
  inference(superposition,[],[f1114,f779])).
fof(f779,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f548])).
fof(f548,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f340])).
fof(f340,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f339])).
fof(f339,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X21] : (s(bool,'F_0') = s(bool,X21) <=> ~p(s(bool,X21)))),
  file('thm.extreal.le_infty|split|3','thm.bool.EQ_CLAUSES|split|3')).
fof(f2812,plain,(
  ( ! [X28,X29] : (s(bool,'F_0') != s(bool,X29) | p(s(bool,X28))) ) | $spl16),
  inference(duplicate_literal_removal,[],[f2764])).
fof(f2764,plain,(
  ( ! [X28,X29] : (s(bool,'F_0') != s(bool,X29) | p(s(bool,X28)) | p(s(bool,X28))) ) | $spl16),
  inference(superposition,[],[f1114,f777])).
fof(f2707,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK12)))) ) | $spl128),
  inference(cnf_transformation,[],[f2707_D])).
fof(f2707_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK12)))) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f2672,plain,(
  ( ! [X163,X162] : (p(s(bool,X162)) | p(s(bool,X163)) | s(bool,X163) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK12)))) ) | $spl0),
  inference(resolution,[],[f1800,f824])).
fof(f1117,plain,(
  ( ! [X8,X7] : (s(bool,X7) != s(bool,X8) | ~p(s(bool,X7))) ) | $spl22),
  inference(cnf_transformation,[],[f1117_D])).
fof(f1117_D,plain,(
  ( ! [X8,X7] : (s(bool,X7) != s(bool,X8) | ~p(s(bool,X7))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f2584,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK12))) | s(bool,X22) != s(bool,X23)) ) | $spl0),
  inference(resolution,[],[f1800,f574])).
fof(f2708,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK12)))) ) | $spl130),
  inference(cnf_transformation,[],[f2708_D])).
fof(f2708_D,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK12)))) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f2583,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,X19) | s(bool,X20) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK12)))) ) | $spl0),
  inference(resolution,[],[f1800,f574])).
fof(f2581,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK12)))) ) | $spl0),
  inference(resolution,[],[f1800,f573])).
fof(f573,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f436])).
fof(f1114,plain,(
  ( ! [X2,X0] : (s(bool,X0) != s(bool,X2) | p(s(bool,X0))) ) | $spl16),
  inference(cnf_transformation,[],[f1114_D])).
fof(f1114_D,plain,(
  ( ! [X2,X0] : (s(bool,X0) != s(bool,X2) | p(s(bool,X0))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f2580,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK12)))) ) | $spl0),
  inference(resolution,[],[f1800,f572])).
fof(f2579,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK12))) | s(bool,X7) != s(bool,X8)) ) | $spl0),
  inference(resolution,[],[f1800,f570])).
fof(f2578,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X4) | s(bool,X5) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK12)))) ) | $spl0),
  inference(resolution,[],[f1800,f570])).
fof(f2577,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK12)))) ) | $spl0),
  inference(resolution,[],[f1800,f568])).
fof(f568,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f436])).
fof(f2524,plain,(
  s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))))) | $spl126),
  inference(cnf_transformation,[],[f2524_D])).
fof(f2524_D,plain,(
  s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))))) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f2553,plain,(
  ( ! [X13] : (s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))))))) ) | ($spl9 | $spl12)),
  inference(forward_demodulation,[],[f2552,f917])).
fof(f2552,plain,(
  ( ! [X13] : (s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))) ) | ($spl9 | $spl12)),
  inference(forward_demodulation,[],[f2551,f917])).
fof(f2551,plain,(
  ( ! [X13] : (s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))) ) | ($spl9 | $spl12)),
  inference(subsumption_resolution,[],[f2550,f612])).
fof(f612,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))) )),
  inference(cnf_transformation,[],[f280])).
fof(f280,plain,(
  ! [X0] : s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))),
  inference(flattening,[],[f279])).
fof(f279,plain,(
  ! [X0] : ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))),
  inference(rectify,[],[f61])).
fof(f61,axiom,(
  ! [X6] : ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X6)))),
  file('thm.extreal.le_infty|split|3','thm.extreal.extreal_not_infty|split|1')).
fof(f2550,plain,(
  ( ! [X13] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))) ) | ($spl9 | $spl12)),
  inference(subsumption_resolution,[],[f2476,f912])).
fof(f2476,plain,(
  ( ! [X13] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))) ) | $spl12),
  inference(superposition,[],[f617,f917])).
fof(f2549,plain,(
  ( ! [X12] : (s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X12))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))))))) ) | ($spl3 | $spl12)),
  inference(forward_demodulation,[],[f2548,f917])).
fof(f2548,plain,(
  ( ! [X12] : (s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X12))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))) ) | ($spl3 | $spl12)),
  inference(forward_demodulation,[],[f2547,f917])).
fof(f2547,plain,(
  ( ! [X12] : (s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X12))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))) ) | ($spl3 | $spl12)),
  inference(subsumption_resolution,[],[f2546,f611])).
fof(f611,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))) )),
  inference(cnf_transformation,[],[f278])).
fof(f278,plain,(
  ! [X0] : s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))),
  inference(flattening,[],[f277])).
fof(f277,plain,(
  ! [X0] : ~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))),
  inference(rectify,[],[f156])).
fof(f156,axiom,(
  ! [X11] : ~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X11)))),
  file('thm.extreal.le_infty|split|3','thm.extreal.extreal_distinct|split|1')).
fof(f2546,plain,(
  ( ! [X12] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X12))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))) ) | ($spl3 | $spl12)),
  inference(subsumption_resolution,[],[f2475,f835])).
fof(f2475,plain,(
  ( ! [X12] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X12))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))) ) | $spl12),
  inference(superposition,[],[f617,f917])).
fof(f2545,plain,(
  ( ! [X11] : (s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X11)))) ) | $spl12),
  inference(forward_demodulation,[],[f2544,f917])).
fof(f2544,plain,(
  ( ! [X11] : (s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X11))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))))))) ) | $spl12),
  inference(subsumption_resolution,[],[f2543,f611])).
fof(f2543,plain,(
  ( ! [X11] : (s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X11))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))))))) ) | $spl12),
  inference(subsumption_resolution,[],[f2474,f612])).
fof(f2474,plain,(
  ( ! [X11] : (s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X11))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))))))) ) | $spl12),
  inference(superposition,[],[f617,f917])).
fof(f2528,plain,(
  ( ! [X3] : (s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))))))) ) | ($spl9 | $spl12)),
  inference(forward_demodulation,[],[f2527,f917])).
fof(f2527,plain,(
  ( ! [X3] : (s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))) ) | ($spl9 | $spl12)),
  inference(forward_demodulation,[],[f2526,f917])).
fof(f2526,plain,(
  ( ! [X3] : (s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))) ) | ($spl9 | $spl12)),
  inference(subsumption_resolution,[],[f2525,f612])).
fof(f2525,plain,(
  ( ! [X3] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))) ) | ($spl9 | $spl12)),
  inference(subsumption_resolution,[],[f2453,f912])).
fof(f2453,plain,(
  ( ! [X3] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))) ) | $spl12),
  inference(superposition,[],[f917,f617])).
fof(f2522,plain,(
  ( ! [X2] : (s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))))))) ) | ($spl3 | $spl12)),
  inference(forward_demodulation,[],[f2521,f917])).
fof(f2521,plain,(
  ( ! [X2] : (s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))) ) | ($spl3 | $spl12)),
  inference(forward_demodulation,[],[f2520,f917])).
fof(f2520,plain,(
  ( ! [X2] : (s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))) ) | ($spl3 | $spl12)),
  inference(subsumption_resolution,[],[f2519,f611])).
fof(f2519,plain,(
  ( ! [X2] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))) ) | ($spl3 | $spl12)),
  inference(subsumption_resolution,[],[f2452,f835])).
fof(f2452,plain,(
  ( ! [X2] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))) ) | $spl12),
  inference(superposition,[],[f917,f617])).
fof(f917,plain,(
  s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) | $spl12),
  inference(cnf_transformation,[],[f917_D])).
fof(f917_D,plain,(
  s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f2446,plain,(
  ( ! [X12,X13] : (sP8('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X12,X13) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f2435,f835])).
fof(f2435,plain,(
  ( ! [X12,X13] : (sP8('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X12,X13) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1248,f619])).
fof(f619,plain,(
  ( ! [X0] : (s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(X0)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0)) )),
  inference(cnf_transformation,[],[f456])).
fof(f456,plain,(
  ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK17])],[f287])).
fof(f287,plain,(
  ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | ? [X1] : s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1))))),
  inference(rectify,[],[f66])).
fof(f66,axiom,(
  ! [X44] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X44) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X44) | ? [X35] : s('type.extreal.extreal',X44) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X35))))),
  file('thm.extreal.le_infty|split|3','thm.extreal.extreal_nchotomy')).
fof(f2444,plain,(
  ( ! [X10,X11] : (sP8('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X10,X11) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f2434,f912])).
fof(f2434,plain,(
  ( ! [X10,X11] : (sP8('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X10,X11) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1248,f619])).
fof(f2423,plain,(
  ( ! [X12,X13] : (sP8(X12,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X13) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f2412,f835])).
fof(f2412,plain,(
  ( ! [X12,X13] : (sP8(X12,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X13) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1247,f619])).
fof(f2421,plain,(
  ( ! [X10,X11] : (sP8(X10,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X11) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f2411,f912])).
fof(f2411,plain,(
  ( ! [X10,X11] : (sP8(X10,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X11) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1247,f619])).
fof(f2400,plain,(
  ( ! [X12,X13] : (sP6('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X12,X13) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f2388,f835])).
fof(f2388,plain,(
  ( ! [X12,X13] : (sP6('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X12,X13) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1234,f619])).
fof(f2398,plain,(
  ( ! [X10,X11] : (sP6('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X10,X11) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f2387,f912])).
fof(f2387,plain,(
  ( ! [X10,X11] : (sP6('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X10,X11) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1234,f619])).
fof(f2376,plain,(
  ( ! [X12,X13] : (sP4('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X12,X13) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f2365,f835])).
fof(f2365,plain,(
  ( ! [X12,X13] : (sP4('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X12,X13) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1218,f619])).
fof(f2374,plain,(
  ( ! [X10,X11] : (sP4('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X10,X11) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f2364,f912])).
fof(f2364,plain,(
  ( ! [X10,X11] : (sP4('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X10,X11) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1218,f619])).
fof(f2353,plain,(
  ( ! [X12,X13] : (sP4(X12,X13,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f2342,f835])).
fof(f2342,plain,(
  ( ! [X12,X13] : (sP4(X12,X13,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1217,f619])).
fof(f2351,plain,(
  ( ! [X10,X11] : (sP4(X10,X11,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f2341,f912])).
fof(f2341,plain,(
  ( ! [X10,X11] : (sP4(X10,X11,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1217,f619])).
fof(f2280,plain,(
  s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK16(sK12))))))) | $spl104),
  inference(cnf_transformation,[],[f2280_D])).
fof(f2280_D,plain,(
  s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK16(sK12))))))) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f2309,plain,(
  ( ! [X13] : (s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))))))) ) | ($spl9 | $spl10)),
  inference(forward_demodulation,[],[f2308,f915])).
fof(f2308,plain,(
  ( ! [X13] : (s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK16(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12))))) ) | ($spl9 | $spl10)),
  inference(forward_demodulation,[],[f2307,f915])).
fof(f2307,plain,(
  ( ! [X13] : (s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK16(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12))))) ) | ($spl9 | $spl10)),
  inference(subsumption_resolution,[],[f2306,f612])).
fof(f2306,plain,(
  ( ! [X13] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK16(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12))))) ) | ($spl9 | $spl10)),
  inference(subsumption_resolution,[],[f2232,f912])).
fof(f2232,plain,(
  ( ! [X13] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK16(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12))))) ) | $spl10),
  inference(superposition,[],[f617,f915])).
fof(f2305,plain,(
  ( ! [X12] : (s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X12))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))))))) ) | ($spl3 | $spl10)),
  inference(forward_demodulation,[],[f2304,f915])).
fof(f2304,plain,(
  ( ! [X12] : (s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK16(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X12))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12))))) ) | ($spl3 | $spl10)),
  inference(forward_demodulation,[],[f2303,f915])).
fof(f2303,plain,(
  ( ! [X12] : (s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK16(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X12))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12))))) ) | ($spl3 | $spl10)),
  inference(subsumption_resolution,[],[f2302,f611])).
fof(f2302,plain,(
  ( ! [X12] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK16(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X12))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12))))) ) | ($spl3 | $spl10)),
  inference(subsumption_resolution,[],[f2231,f835])).
fof(f2231,plain,(
  ( ! [X12] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK16(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X12))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12))))) ) | $spl10),
  inference(superposition,[],[f617,f915])).
fof(f2301,plain,(
  ( ! [X11] : (s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK16(sK12))))))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X11)))) ) | $spl10),
  inference(forward_demodulation,[],[f2300,f915])).
fof(f2300,plain,(
  ( ! [X11] : (s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X11))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))))))) ) | $spl10),
  inference(subsumption_resolution,[],[f2299,f611])).
fof(f2299,plain,(
  ( ! [X11] : (s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X11))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))))))) ) | $spl10),
  inference(subsumption_resolution,[],[f2230,f612])).
fof(f2230,plain,(
  ( ! [X11] : (s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X11))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))))))) ) | $spl10),
  inference(superposition,[],[f617,f915])).
fof(f2284,plain,(
  ( ! [X3] : (s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))))))) ) | ($spl9 | $spl10)),
  inference(forward_demodulation,[],[f2283,f915])).
fof(f2283,plain,(
  ( ! [X3] : (s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK16(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12))))) ) | ($spl9 | $spl10)),
  inference(forward_demodulation,[],[f2282,f915])).
fof(f2282,plain,(
  ( ! [X3] : (s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK16(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12))))) ) | ($spl9 | $spl10)),
  inference(subsumption_resolution,[],[f2281,f612])).
fof(f2281,plain,(
  ( ! [X3] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK16(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12))))) ) | ($spl9 | $spl10)),
  inference(subsumption_resolution,[],[f2209,f912])).
fof(f2209,plain,(
  ( ! [X3] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK16(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12))))) ) | $spl10),
  inference(superposition,[],[f915,f617])).
fof(f2278,plain,(
  ( ! [X2] : (s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))))))) ) | ($spl3 | $spl10)),
  inference(forward_demodulation,[],[f2277,f915])).
fof(f2277,plain,(
  ( ! [X2] : (s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK16(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12))))) ) | ($spl3 | $spl10)),
  inference(forward_demodulation,[],[f2276,f915])).
fof(f2276,plain,(
  ( ! [X2] : (s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK16(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12))))) ) | ($spl3 | $spl10)),
  inference(subsumption_resolution,[],[f2275,f611])).
fof(f2275,plain,(
  ( ! [X2] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK16(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12))))) ) | ($spl3 | $spl10)),
  inference(subsumption_resolution,[],[f2208,f835])).
fof(f2208,plain,(
  ( ! [X2] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK16(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12))))) ) | $spl10),
  inference(superposition,[],[f915,f617])).
fof(f2203,plain,(
  ( ! [X12,X13] : (sP1('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X12,X13) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f2187,f835])).
fof(f2187,plain,(
  ( ! [X12,X13] : (sP1('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X12,X13) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1203,f619])).
fof(f2202,plain,(
  ( ! [X10,X11] : (sP1('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X10,X11) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f2186,f912])).
fof(f2186,plain,(
  ( ! [X10,X11] : (sP1('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X10,X11) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1203,f619])).
fof(f915,plain,(
  s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))) | $spl10),
  inference(cnf_transformation,[],[f915_D])).
fof(f915_D,plain,(
  s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f2200,plain,(
  ( ! [X8,X9] : (sP1('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X8,X9) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12))))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f2185,f835])).
fof(f2185,plain,(
  ( ! [X8,X9] : (sP1('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X8,X9) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1203,f618])).
fof(f618,plain,(
  ( ! [X0] : (s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(X0)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0)) )),
  inference(cnf_transformation,[],[f455])).
fof(f455,plain,(
  ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK16])],[f286])).
fof(f286,plain,(
  ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | ? [X1] : s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1))))),
  inference(rectify,[],[f62])).
fof(f62,axiom,(
  ! [X6] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X6) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X6) | ? [X35] : s('type.extreal.extreal',X6) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X35))))),
  file('thm.extreal.le_infty|split|3','thm.extreal.extreal_cases')).
fof(f2198,plain,(
  ( ! [X6,X7] : (sP1('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X6,X7) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12))))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f2184,f912])).
fof(f2184,plain,(
  ( ! [X6,X7] : (sP1('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X6,X7) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1203,f618])).
fof(f2177,plain,(
  ( ! [X12,X13] : (sP1(X12,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X13) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f2162,f835])).
fof(f2162,plain,(
  ( ! [X12,X13] : (sP1(X12,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X13) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1202,f619])).
fof(f2176,plain,(
  ( ! [X10,X11] : (sP1(X10,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X11) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f2161,f912])).
fof(f2161,plain,(
  ( ! [X10,X11] : (sP1(X10,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X11) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1202,f619])).
fof(f2174,plain,(
  ( ! [X8,X9] : (sP1(X8,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X9) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12))))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f2160,f835])).
fof(f2160,plain,(
  ( ! [X8,X9] : (sP1(X8,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X9) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1202,f618])).
fof(f2172,plain,(
  ( ! [X6,X7] : (sP1(X6,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X7) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12))))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f2159,f912])).
fof(f2159,plain,(
  ( ! [X6,X7] : (sP1(X6,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X7) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1202,f618])).
fof(f2152,plain,(
  ( ! [X12,X13] : (sP2('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X12,X13) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f2136,f835])).
fof(f2136,plain,(
  ( ! [X12,X13] : (sP2('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X12,X13) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1200,f619])).
fof(f2151,plain,(
  ( ! [X10,X11] : (sP2('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X10,X11) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f2135,f912])).
fof(f2135,plain,(
  ( ! [X10,X11] : (sP2('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X10,X11) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1200,f619])).
fof(f2149,plain,(
  ( ! [X8,X9] : (sP2('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X8,X9) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12))))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f2134,f835])).
fof(f2134,plain,(
  ( ! [X8,X9] : (sP2('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X8,X9) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1200,f618])).
fof(f2147,plain,(
  ( ! [X6,X7] : (sP2('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X6,X7) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12))))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f2133,f912])).
fof(f2133,plain,(
  ( ! [X6,X7] : (sP2('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X6,X7) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1200,f618])).
fof(f2126,plain,(
  ( ! [X12,X13] : (sP2(X12,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X13) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f2112,f835])).
fof(f2112,plain,(
  ( ! [X12,X13] : (sP2(X12,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X13) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1199,f619])).
fof(f2125,plain,(
  ( ! [X10,X11] : (sP2(X10,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X11) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f2111,f912])).
fof(f2111,plain,(
  ( ! [X10,X11] : (sP2(X10,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X11) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1199,f619])).
fof(f2123,plain,(
  ( ! [X8,X9] : (sP2(X8,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X9) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12))))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f2110,f835])).
fof(f2110,plain,(
  ( ! [X8,X9] : (sP2(X8,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X9) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1199,f618])).
fof(f2121,plain,(
  ( ! [X6,X7] : (sP2(X6,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X7) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12))))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f2109,f912])).
fof(f2109,plain,(
  ( ! [X6,X7] : (sP2(X6,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X7) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1199,f618])).
fof(f2102,plain,(
  $false | ($spl4 | $spl6)),
  inference(trivial_inequality_removal,[],[f2092])).
fof(f2092,plain,(
  s('type.extreal.extreal',sK12) != s('type.extreal.extreal',sK12) | ($spl4 | $spl6)),
  inference(superposition,[],[f909,f911])).
fof(f2031,plain,(
  s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK15(sK12))))))) | $spl86),
  inference(cnf_transformation,[],[f2031_D])).
fof(f2031_D,plain,(
  s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK15(sK12))))))) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f2060,plain,(
  ( ! [X13] : (s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))))))) ) | ($spl6 | $spl9)),
  inference(forward_demodulation,[],[f2059,f911])).
fof(f2059,plain,(
  ( ! [X13] : (s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK15(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12))))) ) | ($spl6 | $spl9)),
  inference(forward_demodulation,[],[f2058,f911])).
fof(f2058,plain,(
  ( ! [X13] : (s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK15(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12))))) ) | ($spl6 | $spl9)),
  inference(subsumption_resolution,[],[f2057,f612])).
fof(f2057,plain,(
  ( ! [X13] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK15(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12))))) ) | ($spl6 | $spl9)),
  inference(subsumption_resolution,[],[f1983,f912])).
fof(f1983,plain,(
  ( ! [X13] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK15(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12))))) ) | $spl6),
  inference(superposition,[],[f617,f911])).
fof(f2056,plain,(
  ( ! [X12] : (s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X12))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))))))) ) | ($spl3 | $spl6)),
  inference(forward_demodulation,[],[f2055,f911])).
fof(f2055,plain,(
  ( ! [X12] : (s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK15(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X12))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12))))) ) | ($spl3 | $spl6)),
  inference(forward_demodulation,[],[f2054,f911])).
fof(f2054,plain,(
  ( ! [X12] : (s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK15(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X12))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12))))) ) | ($spl3 | $spl6)),
  inference(subsumption_resolution,[],[f2053,f611])).
fof(f2053,plain,(
  ( ! [X12] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK15(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X12))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12))))) ) | ($spl3 | $spl6)),
  inference(subsumption_resolution,[],[f1982,f835])).
fof(f1982,plain,(
  ( ! [X12] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK15(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X12))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12))))) ) | $spl6),
  inference(superposition,[],[f617,f911])).
fof(f2052,plain,(
  ( ! [X11] : (s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK15(sK12))))))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X11)))) ) | $spl6),
  inference(forward_demodulation,[],[f2051,f911])).
fof(f2051,plain,(
  ( ! [X11] : (s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X11))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))))))) ) | $spl6),
  inference(subsumption_resolution,[],[f2050,f611])).
fof(f2050,plain,(
  ( ! [X11] : (s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X11))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))))))) ) | $spl6),
  inference(subsumption_resolution,[],[f1981,f612])).
fof(f1981,plain,(
  ( ! [X11] : (s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X11))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))))))) ) | $spl6),
  inference(superposition,[],[f617,f911])).
fof(f2035,plain,(
  ( ! [X3] : (s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))))))) ) | ($spl6 | $spl9)),
  inference(forward_demodulation,[],[f2034,f911])).
fof(f2034,plain,(
  ( ! [X3] : (s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK15(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12))))) ) | ($spl6 | $spl9)),
  inference(forward_demodulation,[],[f2033,f911])).
fof(f2033,plain,(
  ( ! [X3] : (s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK15(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12))))) ) | ($spl6 | $spl9)),
  inference(subsumption_resolution,[],[f2032,f612])).
fof(f2032,plain,(
  ( ! [X3] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK15(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12))))) ) | ($spl6 | $spl9)),
  inference(subsumption_resolution,[],[f1960,f912])).
fof(f1960,plain,(
  ( ! [X3] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK15(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12))))) ) | $spl6),
  inference(superposition,[],[f911,f617])).
fof(f2029,plain,(
  ( ! [X2] : (s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))))))) ) | ($spl3 | $spl6)),
  inference(forward_demodulation,[],[f2028,f911])).
fof(f2028,plain,(
  ( ! [X2] : (s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK15(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12))))) ) | ($spl3 | $spl6)),
  inference(forward_demodulation,[],[f2027,f911])).
fof(f2027,plain,(
  ( ! [X2] : (s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK15(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12))))) ) | ($spl3 | $spl6)),
  inference(subsumption_resolution,[],[f2026,f611])).
fof(f2026,plain,(
  ( ! [X2] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK15(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12))))) ) | ($spl3 | $spl6)),
  inference(subsumption_resolution,[],[f1959,f835])).
fof(f1959,plain,(
  ( ! [X2] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15('const.extreal.Normal_1'(s('type.realax.real',sK15(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12))))) ) | $spl6),
  inference(superposition,[],[f911,f617])).
fof(f1944,plain,(
  ( ! [X1] : (s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1943,f835])).
fof(f1943,plain,(
  ( ! [X1] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | $spl0),
  inference(subsumption_resolution,[],[f1903,f815])).
fof(f1903,plain,(
  ( ! [X1] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | $spl0),
  inference(superposition,[],[f1131,f617])).
fof(f1934,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl82),
  inference(cnf_transformation,[],[f1934_D])).
fof(f1934_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f1899,plain,(
  ( ! [X169,X168] : (p(s(bool,X168)) | ~p(s(bool,X169)) | s(bool,X169) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl0),
  inference(resolution,[],[f1131,f824])).
fof(f1808,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s(bool,X21) != s(bool,X23)) ) | $spl0),
  inference(resolution,[],[f1131,f574])).
fof(f1806,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl0),
  inference(resolution,[],[f1131,f573])).
fof(f1935,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl84),
  inference(cnf_transformation,[],[f1935_D])).
fof(f1935_D,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f1805,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s(bool,X12) != s(bool,X13)) ) | $spl0),
  inference(resolution,[],[f1131,f572])).
fof(f1804,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl0),
  inference(resolution,[],[f1131,f572])).
fof(f1803,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s(bool,X6) != s(bool,X8)) ) | $spl0),
  inference(resolution,[],[f1131,f570])).
fof(f1802,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s(bool,X3) != s(bool,X4)) ) | $spl0),
  inference(resolution,[],[f1131,f568])).
fof(f1801,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl0),
  inference(resolution,[],[f1131,f568])).
fof(f1160,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl33),
  inference(cnf_transformation,[],[f1160_D])).
fof(f1160_D,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.NegInf_0')))) <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f1794,plain,(
  ( ! [X4,X3] : (sP8(X3,X4,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl80),
  inference(cnf_transformation,[],[f1794_D])).
fof(f1794_D,plain,(
  ( ! [X4,X3] : (sP8(X3,X4,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f1788,plain,(
  ( ! [X14,X15] : (sP8(X14,X15,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.NegInf_0'))))) ) | $spl0),
  inference(superposition,[],[f1049,f809])).
fof(f1049,plain,(
  ( ! [X128,X129] : (sP8(X128,X129,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)))) ) | $spl0),
  inference(resolution,[],[f832,f663])).
fof(f1798,plain,(
  ( ! [X12,X13] : (sP8(X12,X13,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1787,f835])).
fof(f1787,plain,(
  ( ! [X12,X13] : (sP8(X12,X13,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1049,f619])).
fof(f1792,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl78),
  inference(cnf_transformation,[],[f1792_D])).
fof(f1792_D,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f1797,plain,(
  ( ! [X10,X11] : (sP8(X10,X11,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f1786,f912])).
fof(f1786,plain,(
  ( ! [X10,X11] : (sP8(X10,X11,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1049,f619])).
fof(f1796,plain,(
  ( ! [X8,X9] : (sP8(X8,X9,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12))))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1785,f835])).
fof(f1785,plain,(
  ( ! [X8,X9] : (sP8(X8,X9,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1049,f618])).
fof(f1795,plain,(
  ( ! [X6,X7] : (sP8(X6,X7,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12))))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f1784,f912])).
fof(f1784,plain,(
  ( ! [X6,X7] : (sP8(X6,X7,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1049,f618])).
fof(f1793,plain,(
  ( ! [X4,X5,X3] : (sP8(X3,X4,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1783,f835])).
fof(f1783,plain,(
  ( ! [X4,X5,X3] : (sP8(X3,X4,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl0),
  inference(superposition,[],[f1049,f617])).
fof(f1791,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f1782,f912])).
fof(f1782,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | $spl0),
  inference(superposition,[],[f1049,f617])).
fof(f1766,plain,(
  ( ! [X4,X3] : (sP6(X3,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X4)) ) | $spl76),
  inference(cnf_transformation,[],[f1766_D])).
fof(f1766_D,plain,(
  ( ! [X4,X3] : (sP6(X3,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X4)) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f1761,plain,(
  ( ! [X14,X15] : (sP6(X14,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X15) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.NegInf_0'))))) ) | $spl0),
  inference(superposition,[],[f1033,f809])).
fof(f1033,plain,(
  ( ! [X97,X96] : (sP6(X96,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)),X97)) ) | $spl0),
  inference(resolution,[],[f832,f647])).
fof(f1770,plain,(
  ( ! [X12,X13] : (sP6(X12,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X13) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1760,f835])).
fof(f1760,plain,(
  ( ! [X12,X13] : (sP6(X12,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X13) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1033,f619])).
fof(f1764,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X1)) ) | $spl74),
  inference(cnf_transformation,[],[f1764_D])).
fof(f1764_D,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X1)) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f1769,plain,(
  ( ! [X10,X11] : (sP6(X10,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X11) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f1759,f912])).
fof(f1759,plain,(
  ( ! [X10,X11] : (sP6(X10,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X11) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1033,f619])).
fof(f1768,plain,(
  ( ! [X8,X9] : (sP6(X8,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X9) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12))))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1758,f835])).
fof(f1758,plain,(
  ( ! [X8,X9] : (sP6(X8,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X9) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1033,f618])).
fof(f1767,plain,(
  ( ! [X6,X7] : (sP6(X6,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X7) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12))))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f1757,f912])).
fof(f1757,plain,(
  ( ! [X6,X7] : (sP6(X6,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X7) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1033,f618])).
fof(f1765,plain,(
  ( ! [X4,X5,X3] : (sP6(X3,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X4) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1756,f835])).
fof(f1756,plain,(
  ( ! [X4,X5,X3] : (sP6(X3,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X4) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl0),
  inference(superposition,[],[f1033,f617])).
fof(f1763,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X1) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f1755,f912])).
fof(f1755,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X1) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | $spl0),
  inference(superposition,[],[f1033,f617])).
fof(f1739,plain,(
  ( ! [X4,X3] : (sP6(X3,X4,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl72),
  inference(cnf_transformation,[],[f1739_D])).
fof(f1739_D,plain,(
  ( ! [X4,X3] : (sP6(X3,X4,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f1734,plain,(
  ( ! [X14,X15] : (sP6(X14,X15,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.NegInf_0'))))) ) | $spl0),
  inference(superposition,[],[f1032,f809])).
fof(f1032,plain,(
  ( ! [X94,X95] : (sP6(X94,X95,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)))) ) | $spl0),
  inference(resolution,[],[f832,f646])).
fof(f1743,plain,(
  ( ! [X12,X13] : (sP6(X12,X13,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1733,f835])).
fof(f1733,plain,(
  ( ! [X12,X13] : (sP6(X12,X13,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1032,f619])).
fof(f1737,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl70),
  inference(cnf_transformation,[],[f1737_D])).
fof(f1737_D,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f1742,plain,(
  ( ! [X10,X11] : (sP6(X10,X11,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f1732,f912])).
fof(f1732,plain,(
  ( ! [X10,X11] : (sP6(X10,X11,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1032,f619])).
fof(f1741,plain,(
  ( ! [X8,X9] : (sP6(X8,X9,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12))))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1731,f835])).
fof(f1731,plain,(
  ( ! [X8,X9] : (sP6(X8,X9,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1032,f618])).
fof(f1740,plain,(
  ( ! [X6,X7] : (sP6(X6,X7,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12))))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f1730,f912])).
fof(f1730,plain,(
  ( ! [X6,X7] : (sP6(X6,X7,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1032,f618])).
fof(f1738,plain,(
  ( ! [X4,X5,X3] : (sP6(X3,X4,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1729,f835])).
fof(f1729,plain,(
  ( ! [X4,X5,X3] : (sP6(X3,X4,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl0),
  inference(superposition,[],[f1032,f617])).
fof(f1736,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f1728,f912])).
fof(f1728,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | $spl0),
  inference(superposition,[],[f1032,f617])).
fof(f1712,plain,(
  ( ! [X4,X3] : (sP4(X3,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X4)) ) | $spl68),
  inference(cnf_transformation,[],[f1712_D])).
fof(f1712_D,plain,(
  ( ! [X4,X3] : (sP4(X3,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X4)) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f1707,plain,(
  ( ! [X14,X15] : (sP4(X14,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X15) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.NegInf_0'))))) ) | $spl0),
  inference(superposition,[],[f1019,f809])).
fof(f1019,plain,(
  ( ! [X68,X69] : (sP4(X68,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)),X69)) ) | $spl0),
  inference(resolution,[],[f832,f630])).
fof(f1716,plain,(
  ( ! [X12,X13] : (sP4(X12,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X13) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1706,f835])).
fof(f1706,plain,(
  ( ! [X12,X13] : (sP4(X12,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X13) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1019,f619])).
fof(f1710,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X1)) ) | $spl66),
  inference(cnf_transformation,[],[f1710_D])).
fof(f1710_D,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X1)) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f1715,plain,(
  ( ! [X10,X11] : (sP4(X10,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X11) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f1705,f912])).
fof(f1705,plain,(
  ( ! [X10,X11] : (sP4(X10,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X11) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1019,f619])).
fof(f1714,plain,(
  ( ! [X8,X9] : (sP4(X8,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X9) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12))))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1704,f835])).
fof(f1704,plain,(
  ( ! [X8,X9] : (sP4(X8,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X9) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1019,f618])).
fof(f1713,plain,(
  ( ! [X6,X7] : (sP4(X6,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X7) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12))))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f1703,f912])).
fof(f1703,plain,(
  ( ! [X6,X7] : (sP4(X6,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X7) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1019,f618])).
fof(f1711,plain,(
  ( ! [X4,X5,X3] : (sP4(X3,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X4) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1702,f835])).
fof(f1702,plain,(
  ( ! [X4,X5,X3] : (sP4(X3,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X4) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl0),
  inference(superposition,[],[f1019,f617])).
fof(f1709,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X1) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f1701,f912])).
fof(f1701,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X1) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | $spl0),
  inference(superposition,[],[f1019,f617])).
fof(f1685,plain,(
  ( ! [X4,X3] : (sP0('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X3,X4)) ) | $spl64),
  inference(cnf_transformation,[],[f1685_D])).
fof(f1685_D,plain,(
  ( ! [X4,X3] : (sP0('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X3,X4)) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f1678,plain,(
  ( ! [X14,X15] : (sP0('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X14,X15) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.NegInf_0'))))) ) | $spl0),
  inference(superposition,[],[f1007,f809])).
fof(f1007,plain,(
  ( ! [X47,X46] : (sP0('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)),X46,X47)) ) | $spl0),
  inference(resolution,[],[f832,f593])).
fof(f1689,plain,(
  ( ! [X12,X13] : (sP0('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X12,X13) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1677,f835])).
fof(f1677,plain,(
  ( ! [X12,X13] : (sP0('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X12,X13) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1007,f619])).
fof(f1683,plain,(
  ( ! [X0,X1] : (sP0('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0,X1)) ) | $spl62),
  inference(cnf_transformation,[],[f1683_D])).
fof(f1683_D,plain,(
  ( ! [X0,X1] : (sP0('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0,X1)) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f1688,plain,(
  ( ! [X10,X11] : (sP0('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X10,X11) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f1676,f912])).
fof(f1676,plain,(
  ( ! [X10,X11] : (sP0('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X10,X11) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1007,f619])).
fof(f1687,plain,(
  ( ! [X8,X9] : (sP0('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X8,X9) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12))))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1675,f835])).
fof(f1675,plain,(
  ( ! [X8,X9] : (sP0('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X8,X9) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1007,f618])).
fof(f1686,plain,(
  ( ! [X6,X7] : (sP0('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X6,X7) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12))))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f1674,f912])).
fof(f1674,plain,(
  ( ! [X6,X7] : (sP0('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X6,X7) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1007,f618])).
fof(f1684,plain,(
  ( ! [X4,X5,X3] : (sP0('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X3,X4) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1673,f835])).
fof(f1673,plain,(
  ( ! [X4,X5,X3] : (sP0('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X3,X4) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl0),
  inference(superposition,[],[f1007,f617])).
fof(f1682,plain,(
  ( ! [X2,X0,X1] : (sP0('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0,X1) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f1672,f912])).
fof(f1672,plain,(
  ( ! [X2,X0,X1] : (sP0('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0,X1) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | $spl0),
  inference(superposition,[],[f1007,f617])).
fof(f1657,plain,(
  ( ! [X4,X3] : (sP0(X3,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X4)) ) | $spl60),
  inference(cnf_transformation,[],[f1657_D])).
fof(f1657_D,plain,(
  ( ! [X4,X3] : (sP0(X3,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X4)) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f1650,plain,(
  ( ! [X14,X15] : (sP0(X14,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X15) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.NegInf_0'))))) ) | $spl0),
  inference(superposition,[],[f1006,f809])).
fof(f1006,plain,(
  ( ! [X45,X44] : (sP0(X44,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)),X45)) ) | $spl0),
  inference(resolution,[],[f832,f592])).
fof(f1661,plain,(
  ( ! [X12,X13] : (sP0(X12,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X13) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1649,f835])).
fof(f1649,plain,(
  ( ! [X12,X13] : (sP0(X12,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X13) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1006,f619])).
fof(f1655,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X1)) ) | $spl58),
  inference(cnf_transformation,[],[f1655_D])).
fof(f1655_D,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X1)) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f1660,plain,(
  ( ! [X10,X11] : (sP0(X10,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X11) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f1648,f912])).
fof(f1648,plain,(
  ( ! [X10,X11] : (sP0(X10,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X11) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1006,f619])).
fof(f1659,plain,(
  ( ! [X8,X9] : (sP0(X8,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X9) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12))))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1647,f835])).
fof(f1647,plain,(
  ( ! [X8,X9] : (sP0(X8,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X9) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1006,f618])).
fof(f1658,plain,(
  ( ! [X6,X7] : (sP0(X6,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X7) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12))))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f1646,f912])).
fof(f1646,plain,(
  ( ! [X6,X7] : (sP0(X6,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X7) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1006,f618])).
fof(f1656,plain,(
  ( ! [X4,X5,X3] : (sP0(X3,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X4) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1645,f835])).
fof(f1645,plain,(
  ( ! [X4,X5,X3] : (sP0(X3,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X4) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl0),
  inference(superposition,[],[f1006,f617])).
fof(f1654,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X1) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f1644,f912])).
fof(f1644,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X1) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | $spl0),
  inference(superposition,[],[f1006,f617])).
fof(f1629,plain,(
  ( ! [X4,X3] : (sP0(X3,X4,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl56),
  inference(cnf_transformation,[],[f1629_D])).
fof(f1629_D,plain,(
  ( ! [X4,X3] : (sP0(X3,X4,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f1622,plain,(
  ( ! [X14,X15] : (sP0(X14,X15,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.NegInf_0'))))) ) | $spl0),
  inference(superposition,[],[f1005,f809])).
fof(f1005,plain,(
  ( ! [X43,X42] : (sP0(X42,X43,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)))) ) | $spl0),
  inference(resolution,[],[f832,f591])).
fof(f1633,plain,(
  ( ! [X12,X13] : (sP0(X12,X13,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1621,f835])).
fof(f1621,plain,(
  ( ! [X12,X13] : (sP0(X12,X13,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1005,f619])).
fof(f1627,plain,(
  ( ! [X0,X1] : (sP0(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl54),
  inference(cnf_transformation,[],[f1627_D])).
fof(f1627_D,plain,(
  ( ! [X0,X1] : (sP0(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f1632,plain,(
  ( ! [X10,X11] : (sP0(X10,X11,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f1620,f912])).
fof(f1620,plain,(
  ( ! [X10,X11] : (sP0(X10,X11,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1005,f619])).
fof(f1631,plain,(
  ( ! [X8,X9] : (sP0(X8,X9,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12))))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1619,f835])).
fof(f1619,plain,(
  ( ! [X8,X9] : (sP0(X8,X9,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1005,f618])).
fof(f1630,plain,(
  ( ! [X6,X7] : (sP0(X6,X7,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12))))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f1618,f912])).
fof(f1618,plain,(
  ( ! [X6,X7] : (sP0(X6,X7,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1005,f618])).
fof(f1628,plain,(
  ( ! [X4,X5,X3] : (sP0(X3,X4,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1617,f835])).
fof(f1617,plain,(
  ( ! [X4,X5,X3] : (sP0(X3,X4,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl0),
  inference(superposition,[],[f1005,f617])).
fof(f1626,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f1616,f912])).
fof(f1616,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | $spl0),
  inference(superposition,[],[f1005,f617])).
fof(f1601,plain,(
  ( ! [X4,X3] : (sP1(X3,X4,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl52),
  inference(cnf_transformation,[],[f1601_D])).
fof(f1601_D,plain,(
  ( ! [X4,X3] : (sP1(X3,X4,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f1595,plain,(
  ( ! [X14,X15] : (sP1(X14,X15,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.NegInf_0'))))) ) | $spl0),
  inference(superposition,[],[f1004,f809])).
fof(f1004,plain,(
  ( ! [X41,X40] : (sP1(X40,X41,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)))) ) | $spl0),
  inference(resolution,[],[f832,f587])).
fof(f1605,plain,(
  ( ! [X12,X13] : (sP1(X12,X13,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1594,f835])).
fof(f1594,plain,(
  ( ! [X12,X13] : (sP1(X12,X13,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1004,f619])).
fof(f1599,plain,(
  ( ! [X0,X1] : (sP1(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl50),
  inference(cnf_transformation,[],[f1599_D])).
fof(f1599_D,plain,(
  ( ! [X0,X1] : (sP1(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f1604,plain,(
  ( ! [X10,X11] : (sP1(X10,X11,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f1593,f912])).
fof(f1593,plain,(
  ( ! [X10,X11] : (sP1(X10,X11,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1004,f619])).
fof(f1603,plain,(
  ( ! [X8,X9] : (sP1(X8,X9,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12))))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1592,f835])).
fof(f1592,plain,(
  ( ! [X8,X9] : (sP1(X8,X9,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1004,f618])).
fof(f1602,plain,(
  ( ! [X6,X7] : (sP1(X6,X7,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12))))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f1591,f912])).
fof(f1591,plain,(
  ( ! [X6,X7] : (sP1(X6,X7,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1004,f618])).
fof(f1600,plain,(
  ( ! [X4,X5,X3] : (sP1(X3,X4,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1590,f835])).
fof(f1590,plain,(
  ( ! [X4,X5,X3] : (sP1(X3,X4,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl0),
  inference(superposition,[],[f1004,f617])).
fof(f1598,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f1589,f912])).
fof(f1589,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | $spl0),
  inference(superposition,[],[f1004,f617])).
fof(f1574,plain,(
  ( ! [X4,X3] : (sP2(X3,X4,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl48),
  inference(cnf_transformation,[],[f1574_D])).
fof(f1574_D,plain,(
  ( ! [X4,X3] : (sP2(X3,X4,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f1568,plain,(
  ( ! [X14,X15] : (sP2(X14,X15,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.NegInf_0'))))) ) | $spl0),
  inference(superposition,[],[f1001,f809])).
fof(f1001,plain,(
  ( ! [X35,X34] : (sP2(X34,X35,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)))) ) | $spl0),
  inference(resolution,[],[f832,f583])).
fof(f1578,plain,(
  ( ! [X12,X13] : (sP2(X12,X13,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1567,f835])).
fof(f1567,plain,(
  ( ! [X12,X13] : (sP2(X12,X13,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1001,f619])).
fof(f1572,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl46),
  inference(cnf_transformation,[],[f1572_D])).
fof(f1572_D,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f1577,plain,(
  ( ! [X10,X11] : (sP2(X10,X11,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f1566,f912])).
fof(f1566,plain,(
  ( ! [X10,X11] : (sP2(X10,X11,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1001,f619])).
fof(f1576,plain,(
  ( ! [X8,X9] : (sP2(X8,X9,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12))))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1565,f835])).
fof(f1565,plain,(
  ( ! [X8,X9] : (sP2(X8,X9,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1001,f618])).
fof(f1575,plain,(
  ( ! [X6,X7] : (sP2(X6,X7,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12))))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f1564,f912])).
fof(f1564,plain,(
  ( ! [X6,X7] : (sP2(X6,X7,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1001,f618])).
fof(f1573,plain,(
  ( ! [X4,X5,X3] : (sP2(X3,X4,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1563,f835])).
fof(f1563,plain,(
  ( ! [X4,X5,X3] : (sP2(X3,X4,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl0),
  inference(superposition,[],[f1001,f617])).
fof(f1571,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f1562,f912])).
fof(f1562,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | $spl0),
  inference(superposition,[],[f1001,f617])).
fof(f1551,plain,(
  $false | $spl26),
  inference(subsumption_resolution,[],[f1491,f1119])).
fof(f1491,plain,(
  ( ! [X389,X388] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X388),s('type.extreal.extreal',X389))))) ) | $spl26),
  inference(resolution,[],[f1119,f806])).
fof(f1550,plain,(
  $false | $spl26),
  inference(subsumption_resolution,[],[f1448,f1119])).
fof(f1448,plain,(
  ( ! [X271,X270] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X270),s('type.extreal.extreal',X271))))) ) | $spl26),
  inference(resolution,[],[f1119,f806])).
fof(f1549,plain,(
  $false | $spl26),
  inference(subsumption_resolution,[],[f1447,f1119])).
fof(f1447,plain,(
  ~p(s(bool,'T_0')) | $spl26),
  inference(resolution,[],[f1119,f803])).
fof(f803,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f562])).
fof(f1548,plain,(
  $false | $spl26),
  inference(subsumption_resolution,[],[f1446,f1119])).
fof(f1446,plain,(
  ~p(s(bool,'T_0')) | $spl26),
  inference(resolution,[],[f1119,f801])).
fof(f801,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f561])).
fof(f1547,plain,(
  $false | $spl26),
  inference(subsumption_resolution,[],[f1443,f1119])).
fof(f1443,plain,(
  ( ! [X268] : (~p(s(bool,X268))) ) | $spl26),
  inference(resolution,[],[f1119,f790])).
fof(f1546,plain,(
  $false | $spl26),
  inference(subsumption_resolution,[],[f1437,f1119])).
fof(f1437,plain,(
  ( ! [X262] : (~p(s(bool,X262))) ) | $spl26),
  inference(resolution,[],[f1119,f732])).
fof(f732,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f518])).
fof(f518,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f517])).
fof(f517,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f392])).
fof(f392,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f322])).
fof(f322,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f321])).
fof(f321,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f121])).
fof(f121,axiom,(
  ! [X21] : ((p(s(bool,X21)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X21)))),
  file('thm.extreal.le_infty|split|3','thm.bool.IMP_CLAUSES|split|4')).
fof(f1445,plain,(
  $false | $spl26),
  inference(resolution,[],[f1119,f799])).
fof(f799,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f359])).
fof(f359,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.extreal.le_infty|split|3',aHL_FALSITY)).
fof(f1492,plain,(
  $false | ($spl0 | $spl26)),
  inference(resolution,[],[f1119,f987])).
fof(f1327,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl45),
  inference(cnf_transformation,[],[f1327_D])).
fof(f1327_D,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) <=> ~$spl45),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl45])])).
fof(f1294,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl0),
  inference(superposition,[],[f987,f809])).
fof(f1331,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1293,f835])).
fof(f1293,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | $spl0),
  inference(superposition,[],[f987,f619])).
fof(f1324,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl43),
  inference(cnf_transformation,[],[f1324_D])).
fof(f1324_D,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) <=> ~$spl43),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl43])])).
fof(f1330,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f1292,f912])).
fof(f1292,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | $spl0),
  inference(superposition,[],[f987,f619])).
fof(f1329,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1291,f835])).
fof(f1291,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | $spl0),
  inference(superposition,[],[f987,f618])).
fof(f1328,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f1290,f912])).
fof(f1290,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | $spl0),
  inference(superposition,[],[f987,f618])).
fof(f1325,plain,(
  ( ! [X1] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1289,f835])).
fof(f1289,plain,(
  ( ! [X1] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | $spl0),
  inference(superposition,[],[f987,f617])).
fof(f1322,plain,(
  ( ! [X0] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f1288,f912])).
fof(f1288,plain,(
  ( ! [X0] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))) ) | $spl0),
  inference(superposition,[],[f987,f617])).
fof(f1315,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl38),
  inference(cnf_transformation,[],[f1315_D])).
fof(f1315_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f1285,plain,(
  ( ! [X169,X168] : (p(s(bool,X168)) | ~p(s(bool,X169)) | s(bool,X169) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl0),
  inference(resolution,[],[f987,f824])).
fof(f1194,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s(bool,X21) != s(bool,X23)) ) | $spl0),
  inference(resolution,[],[f987,f574])).
fof(f1192,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl0),
  inference(resolution,[],[f987,f573])).
fof(f1316,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl40),
  inference(cnf_transformation,[],[f1316_D])).
fof(f1316_D,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f1191,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s(bool,X12) != s(bool,X13)) ) | $spl0),
  inference(resolution,[],[f987,f572])).
fof(f1190,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl0),
  inference(resolution,[],[f987,f572])).
fof(f1189,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s(bool,X6) != s(bool,X8)) ) | $spl0),
  inference(resolution,[],[f987,f570])).
fof(f1188,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s(bool,X3) != s(bool,X4)) ) | $spl0),
  inference(resolution,[],[f987,f568])).
fof(f1187,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl0),
  inference(resolution,[],[f987,f568])).
fof(f1181,plain,(
  ( ! [X2] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X2)) ) | $spl36),
  inference(cnf_transformation,[],[f1181_D])).
fof(f1181_D,plain,(
  ( ! [X2] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X2)) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f1175,plain,(
  ( ! [X8] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X8) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.NegInf_0'))))) ) | $spl0),
  inference(superposition,[],[f1066,f809])).
fof(f1066,plain,(
  ( ! [X153] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)),X153)) ) | $spl0),
  inference(resolution,[],[f832,f696])).
fof(f1185,plain,(
  ( ! [X7] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X7) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1174,f835])).
fof(f1174,plain,(
  ( ! [X7] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X7) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1066,f619])).
fof(f1179,plain,(
  ( ! [X0] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0)) ) | $spl34),
  inference(cnf_transformation,[],[f1179_D])).
fof(f1179_D,plain,(
  ( ! [X0] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0)) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f1184,plain,(
  ( ! [X6] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X6) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f1173,f912])).
fof(f1173,plain,(
  ( ! [X6] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X6) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1066,f619])).
fof(f1183,plain,(
  ( ! [X5] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X5) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12))))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1172,f835])).
fof(f1172,plain,(
  ( ! [X5] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X5) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1066,f618])).
fof(f1182,plain,(
  ( ! [X4] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X4) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12))))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f1171,f912])).
fof(f1171,plain,(
  ( ! [X4] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X4) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1066,f618])).
fof(f1180,plain,(
  ( ! [X2,X3] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X2) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1170,f835])).
fof(f1170,plain,(
  ( ! [X2,X3] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X2) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3)))) ) | $spl0),
  inference(superposition,[],[f1066,f617])).
fof(f1178,plain,(
  ( ! [X0,X1] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f1169,f912])).
fof(f1169,plain,(
  ( ! [X0,X1] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | $spl0),
  inference(superposition,[],[f1066,f617])).
fof(f1154,plain,(
  ( ! [X2] : (sP11(X2,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl30),
  inference(cnf_transformation,[],[f1154_D])).
fof(f1154_D,plain,(
  ( ! [X2] : (sP11(X2,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f1148,plain,(
  ( ! [X8] : (sP11(X8,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.NegInf_0'))))) ) | $spl0),
  inference(superposition,[],[f1065,f809])).
fof(f1065,plain,(
  ( ! [X152] : (sP11(X152,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)))) ) | $spl0),
  inference(resolution,[],[f832,f695])).
fof(f1158,plain,(
  ( ! [X7] : (sP11(X7,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1147,f835])).
fof(f1147,plain,(
  ( ! [X7] : (sP11(X7,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1065,f619])).
fof(f1152,plain,(
  ( ! [X0] : (sP11(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl28),
  inference(cnf_transformation,[],[f1152_D])).
fof(f1152_D,plain,(
  ( ! [X0] : (sP11(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f1157,plain,(
  ( ! [X6] : (sP11(X6,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12))))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f1146,f912])).
fof(f1146,plain,(
  ( ! [X6] : (sP11(X6,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1065,f619])).
fof(f1156,plain,(
  ( ! [X5] : (sP11(X5,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12))))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1145,f835])).
fof(f1145,plain,(
  ( ! [X5] : (sP11(X5,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1065,f618])).
fof(f1155,plain,(
  ( ! [X4] : (sP11(X4,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12))))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f1144,f912])).
fof(f1144,plain,(
  ( ! [X4] : (sP11(X4,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) ) | $spl0),
  inference(superposition,[],[f1065,f618])).
fof(f1153,plain,(
  ( ! [X2,X3] : (sP11(X2,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1143,f835])).
fof(f1143,plain,(
  ( ! [X2,X3] : (sP11(X2,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3)))) ) | $spl0),
  inference(superposition,[],[f1065,f617])).
fof(f1151,plain,(
  ( ! [X0,X1] : (sP11(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f1142,f912])).
fof(f1142,plain,(
  ( ! [X0,X1] : (sP11(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | $spl0),
  inference(superposition,[],[f1065,f617])).
fof(f1123,plain,(
  ( ! [X1] : (s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1122,f799])).
fof(f1122,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | ($spl0 | $spl3)),
  inference(forward_demodulation,[],[f1121,f818])).
fof(f1121,plain,(
  ( ! [X1] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1093,f835])).
fof(f1093,plain,(
  ( ! [X1] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | $spl0),
  inference(superposition,[],[f832,f617])).
fof(f1113,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)))) ) | $spl14),
  inference(cnf_transformation,[],[f1113_D])).
fof(f1113_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f1083,plain,(
  ( ! [X163,X162] : (p(s(bool,X162)) | p(s(bool,X163)) | s(bool,X163) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)))) ) | $spl0),
  inference(resolution,[],[f832,f824])).
fof(f995,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12))) | s(bool,X22) != s(bool,X23)) ) | $spl0),
  inference(resolution,[],[f832,f574])).
fof(f1115,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)))) ) | $spl18),
  inference(cnf_transformation,[],[f1115_D])).
fof(f1115_D,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f994,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,X19) | s(bool,X20) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)))) ) | $spl0),
  inference(resolution,[],[f832,f574])).
fof(f992,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)))) ) | $spl0),
  inference(resolution,[],[f832,f573])).
fof(f991,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)))) ) | $spl0),
  inference(resolution,[],[f832,f572])).
fof(f990,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12))) | s(bool,X7) != s(bool,X8)) ) | $spl0),
  inference(resolution,[],[f832,f570])).
fof(f989,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X4) | s(bool,X5) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)))) ) | $spl0),
  inference(resolution,[],[f832,f570])).
fof(f988,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)))) ) | $spl0),
  inference(resolution,[],[f832,f568])).
fof(f984,plain,(
  ( ! [X1] : (s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | ($spl3 | $spl9)),
  inference(subsumption_resolution,[],[f983,f835])).
fof(f983,plain,(
  ( ! [X1] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | $spl9),
  inference(trivial_inequality_removal,[],[f974])).
fof(f974,plain,(
  ( ! [X1] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.NegInf_0') | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | $spl9),
  inference(superposition,[],[f912,f617])).
fof(f972,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f971,f799])).
fof(f971,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl8)),
  inference(forward_demodulation,[],[f970,f818])).
fof(f970,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | ($spl0 | $spl8)),
  inference(forward_demodulation,[],[f832,f913])).
fof(f913,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | $spl8),
  inference(cnf_transformation,[],[f913_D])).
fof(f913_D,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f906,plain,(
  s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | $spl3),
  inference(trivial_inequality_removal,[],[f903])).
fof(f903,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.PosInf_0') | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | $spl3),
  inference(superposition,[],[f835,f619])).
fof(f907,plain,(
  s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | $spl3),
  inference(trivial_inequality_removal,[],[f901])).
fof(f901,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.PosInf_0') | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK16(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | $spl3),
  inference(superposition,[],[f835,f618])).
fof(f908,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))) ) | $spl3),
  inference(trivial_inequality_removal,[],[f899])).
fof(f899,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.PosInf_0') | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK15(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))) ) | $spl3),
  inference(superposition,[],[f835,f617])).
fof(f898,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f897,f816])).
fof(f897,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f833,f834])).
fof(f832,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)))) | $spl0),
  inference(cnf_transformation,[],[f832_D])).
fof(f832_D,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f834,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | $spl2),
  inference(cnf_transformation,[],[f834_D])).
fof(f834_D,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f565,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12))))),
  inference(cnf_transformation,[],[f434])).
fof(f434,plain,(
  (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) & (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',sK12))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12])],[f433])).
fof(f433,plain,(
  ? [X0] : ((p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X0)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)) & (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X0)))) | s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0)))),
  inference(nnf_transformation,[],[f373])).
fof(f373,plain,(
  ? [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X0)))) <~> s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0))),
  inference(ennf_transformation,[],[f263])).
fof(f263,plain,(
  ~! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X0)))) <=> s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X6)))) <=> s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X6))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X6)))) <=> s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X6))),
  file('thm.extreal.le_infty|split|3','thm.extreal.le_infty|split|3')).
fof(f833,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)))) | $spl1),
  inference(cnf_transformation,[],[f833_D])).
fof(f833_D,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f835,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',sK12) | $spl3),
  inference(cnf_transformation,[],[f835_D])).
fof(f835_D,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',sK12) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f566,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',sK12) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12))))),
  inference(cnf_transformation,[],[f434])).
% SZS output end Proof for thm.extreal
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 270
% Passive clauses: 987
% Generated clauses: 5205
% Final active clauses: 254
% Final passive clauses: 825
% Input formulas: 261
% Initial clauses: 259
% 
% Selected by SInE selection: 86
% SInE iterations: 1
% 
% Duplicate literals: 81
% Trivial inequalities: 35
% Fw subsumption resolutions: 960
% Fw demodulations: 275
% Bw demodulations: 4
% 
% Simple tautologies: 195
% Equational tautologies: 2
% Forward subsumptions: 2267
% 
% Binary resolution: 1173
% Forward superposition: 1530
% Backward superposition: 1007
% Equality resolution: 5
% 
% Split clauses: 294
% Split components: 725
% SAT solver clauses: 349
% SAT solver unit clauses: 41
% SAT solver binary clauses: 15
% 
% Sat splits: 294
% Sat splitting refutations: 54
% 
% Memory used [KB]: 3454
% Time elapsed: 0.170 s
% ------------------------------
% ------------------------------
% Success in time 12.94 s
