Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f14366,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f14325,f52,f55,f53,f65,f14025,f30,f42,f227,f229,f29,f44,f40,f12361,f12329,f331,f335,f230,f28,f45,f39])).
fof(f39,plain,(
  p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK0,bool),bool),sK1)))))) | $spl0),
  inference(cnf_transformation,[],[f39_D])).
fof(f39_D,plain,(
  p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK0,bool),bool),sK1)))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f45,plain,(
  ( ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X4),s(fun(fun(sK0,bool),bool),sK1)))) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X4))))) ) | $spl6),
  inference(cnf_transformation,[],[f45_D])).
fof(f45_D,plain,(
  ( ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X4),s(fun(fun(sK0,bool),bool),sK1)))) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X4))))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f28,plain,(
  ( ! [X4] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X4)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X4),s(fun(fun(sK0,bool),bool),sK1)))) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK0,bool),bool),sK1))))))) )),
  inference(cnf_transformation,[],[f23])).
fof(f23,plain,(
  (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK0,bool),bool),sK1)))))) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X4),s(fun(fun(sK0,bool),bool),sK1)))) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X4)))))) & (~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK0,bool),bool),sK1)))))) | (p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK3),s(fun(fun(sK0,bool),bool),sK1)))) & ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f22])).
fof(f22,plain,(
  ? [X0,X1,X2] : ((p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X2),s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X1)))))) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X4),s(fun(fun(X0,bool),bool),X1)))) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X4)))))) & (~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X2),s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X1)))))) | ? [X3] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),X1)))) & ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))))))),
  inference(rectify,[],[f21])).
fof(f21,plain,(
  ? [X0,X1,X2] : ((p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X2),s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X1)))))) | ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),X1)))) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))))) & (~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X2),s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X1)))))) | ? [X3] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),X1)))) & ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))))))),
  inference(nnf_transformation,[],[f18])).
fof(f18,plain,(
  ? [X0,X1,X2] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X2),s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X1)))))) <~> ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),X1)))) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))))),
  inference(ennf_transformation,[],[f10])).
fof(f10,plain,(
  ~! [X0,X1,X2] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X2),s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X1)))))) <=> ! [X3] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),X1)))) => p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))))),
  inference(flattening,[],[f9])).
fof(f9,plain,(
  ~! [X0] : ! [X1,X2] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X2),s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X1)))))) <=> ! [X3] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),X1)))) => p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X6,bool),X8),s(fun(X6,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X6,bool),bool),X7)))))) <=> ! [X9] : (p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),X9),s(fun(fun(X6,bool),bool),X7)))) => p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X9))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X6,bool),X8),s(fun(X6,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X6,bool),bool),X7)))))) <=> ! [X9] : (p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),X9),s(fun(fun(X6,bool),bool),X7)))) => p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X9))))))),
  file('thm.pred_set.DISJOINT_BIGUNION|split|1','thm.pred_set.DISJOINT_BIGUNION|split|1')).
fof(f230,plain,(
  ( ! [X11] : (s(bool,'T_0') = s(bool,X11)) ) | $spl14),
  inference(cnf_transformation,[],[f230_D])).
fof(f230_D,plain,(
  ( ! [X11] : (s(bool,'T_0') = s(bool,X11)) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f335,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))) | $spl16),
  inference(cnf_transformation,[],[f335_D])).
fof(f335_D,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f331,plain,(
  ( ! [X9] : (s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))) | s(bool,'T_0') = s(bool,X9)) ) | $spl3),
  inference(duplicate_literal_removal,[],[f289])).
fof(f289,plain,(
  ( ! [X9] : (s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) ) | $spl3),
  inference(superposition,[],[f232,f33])).
fof(f33,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.pred_set.DISJOINT_BIGUNION|split|1',aHL_BOOL_CASES)).
fof(f232,plain,(
  ( ! [X9] : (s(bool,X9) = s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))) | s(bool,'T_0') = s(bool,X9)) ) | $spl3),
  inference(subsumption_resolution,[],[f151,f31])).
fof(f31,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pred_set.DISJOINT_BIGUNION|split|1',aHL_TRUTH)).
fof(f151,plain,(
  ( ! [X9] : (~p(s(bool,'T_0')) | s(bool,X9) = s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))) | s(bool,'T_0') = s(bool,X9)) ) | $spl3),
  inference(superposition,[],[f42,f46])).
fof(f46,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f33,f33])).
fof(f12329,plain,(
  $false | ($spl4 | $spl12 | $spl16)),
  inference(subsumption_resolution,[],[f12328,f32])).
fof(f32,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f11])).
fof(f11,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.pred_set.DISJOINT_BIGUNION|split|1',aHL_FALSITY)).
fof(f12328,plain,(
  p(s(bool,'F_0')) | ($spl4 | $spl12 | $spl16)),
  inference(forward_demodulation,[],[f12327,f335])).
fof(f12327,plain,(
  p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))) | ($spl4 | $spl12)),
  inference(forward_demodulation,[],[f12326,f34])).
fof(f34,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))) )),
  inference(cnf_transformation,[],[f13])).
fof(f13,plain,(
  ! [X0,X1,X2] : s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))),
  inference(flattening,[],[f12])).
fof(f12,plain,(
  ! [X0] : ! [X1,X2] : s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6] : ! [X7,X8] : s(bool,'const.pred_set.DISJOINT_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8))) = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X7)))),
  file('thm.pred_set.DISJOINT_BIGUNION|split|1','thm.pred_set.DISJOINT_SYM')).
fof(f12326,plain,(
  p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),sK2)))) | ($spl4 | $spl12)),
  inference(subsumption_resolution,[],[f12280,f31])).
fof(f12280,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),sK2)))) | ($spl4 | $spl12)),
  inference(superposition,[],[f4298,f433])).
fof(f433,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK0,bool),bool),sK1))),s(fun(sK0,bool),sK2))) | $spl12),
  inference(subsumption_resolution,[],[f424,f31])).
fof(f424,plain,(
  ~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK0,bool),bool),sK1))),s(fun(sK0,bool),sK2))) | $spl12),
  inference(superposition,[],[f263,f229])).
fof(f263,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X19,bool),X21),s(fun(X19,bool),X20)))) | s(bool,'T_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X19,bool),X20),s(fun(X19,bool),X21)))) )),
  inference(superposition,[],[f48,f34])).
fof(f48,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f32,f33])).
fof(f4298,plain,(
  ( ! [X0] : (~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK0,bool),bool),sK1))),s(fun(sK0,bool),X0)))) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),X0))))) ) | $spl4),
  inference(subsumption_resolution,[],[f4271,f31])).
fof(f4271,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),X0)))) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK0,bool),bool),sK1))),s(fun(sK0,bool),X0))))) ) | $spl4),
  inference(superposition,[],[f35,f79])).
fof(f79,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK3),s(fun(fun(sK0,bool),bool),sK1))) | $spl4),
  inference(subsumption_resolution,[],[f78,f32])).
fof(f78,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK3),s(fun(fun(sK0,bool),bool),sK1))) | $spl4),
  inference(superposition,[],[f44,f33])).
fof(f35,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X4),s(fun(fun(X0,bool),bool),X1)))) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X4),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X1))),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f26])).
fof(f26,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X1))),s(fun(X0,bool),X2)))) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X4),s(fun(fun(X0,bool),bool),X1)))) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X4),s(fun(X0,bool),X2)))))) & ((p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),sK4(X2,X1,X0)),s(fun(fun(X0,bool),bool),X1)))) & ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),sK4(X2,X1,X0)),s(fun(X0,bool),X2))))) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X1))),s(fun(X0,bool),X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f25])).
fof(f25,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X1))),s(fun(X0,bool),X2)))) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X4),s(fun(fun(X0,bool),bool),X1)))) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X4),s(fun(X0,bool),X2)))))) & (? [X3] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),X1)))) & ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X1))),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f24])).
fof(f24,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X1))),s(fun(X0,bool),X2)))) | ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),X1)))) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X3),s(fun(X0,bool),X2)))))) & (? [X3] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),X1)))) & ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X1))),s(fun(X0,bool),X2))))))),
  inference(nnf_transformation,[],[f19])).
fof(f19,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X1))),s(fun(X0,bool),X2)))) <=> ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),X1)))) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X3),s(fun(X0,bool),X2))))))),
  inference(ennf_transformation,[],[f15])).
fof(f15,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X1))),s(fun(X0,bool),X2)))) <=> ! [X3] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),X1)))) => p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X3),s(fun(X0,bool),X2))))))),
  inference(flattening,[],[f14])).
fof(f14,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X1))),s(fun(X0,bool),X2)))) <=> ! [X3] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),X1)))) => p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X3),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : ! [X7,X8] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X6,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X6,bool),bool),X7))),s(fun(X6,bool),X8)))) <=> ! [X9] : (p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),X9),s(fun(fun(X6,bool),bool),X7)))) => p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X6,bool),X9),s(fun(X6,bool),X8))))))),
  file('thm.pred_set.DISJOINT_BIGUNION|split|1','thm.pred_set.pred_set_282_prove_ltb')).
fof(f12361,plain,(
  $false | ($spl1 | $spl12)),
  inference(subsumption_resolution,[],[f12360,f31])).
fof(f12360,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl12)),
  inference(forward_demodulation,[],[f40,f229])).
fof(f40,plain,(
  ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK0,bool),bool),sK1)))))) | $spl1),
  inference(cnf_transformation,[],[f40_D])).
fof(f40_D,plain,(
  ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK0,bool),bool),sK1)))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f44,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK3),s(fun(fun(sK0,bool),bool),sK1)))) | $spl4),
  inference(cnf_transformation,[],[f44_D])).
fof(f44_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK3),s(fun(fun(sK0,bool),bool),sK1)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f29,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK3),s(fun(fun(sK0,bool),bool),sK1)))) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK0,bool),bool),sK1))))))),
  inference(cnf_transformation,[],[f23])).
fof(f229,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK0,bool),bool),sK1))))) | $spl12),
  inference(cnf_transformation,[],[f229_D])).
fof(f229_D,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK0,bool),bool),sK1))))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f227,plain,(
  ( ! [X11] : (s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK0,bool),bool),sK1)))))) ) | $spl0),
  inference(subsumption_resolution,[],[f132,f48])).
fof(f132,plain,(
  ( ! [X11] : (p(s(bool,X11)) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK0,bool),bool),sK1)))))) ) | $spl0),
  inference(superposition,[],[f39,f46])).
fof(f42,plain,(
  ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))) | $spl3),
  inference(cnf_transformation,[],[f42_D])).
fof(f42_D,plain,(
  ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f30,plain,(
  ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK0,bool),bool),sK1))))))),
  inference(cnf_transformation,[],[f23])).
fof(f14025,plain,(
  $false | ($spl1 | $spl6)),
  inference(subsumption_resolution,[],[f14024,f40])).
fof(f14024,plain,(
  p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK0,bool),bool),sK1)))))) | ($spl1 | $spl6)),
  inference(forward_demodulation,[],[f13976,f34])).
fof(f13976,plain,(
  p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK0,bool),bool),sK1))),s(fun(sK0,bool),sK2)))) | ($spl1 | $spl6)),
  inference(resolution,[],[f13896,f544])).
fof(f544,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X2),s(fun(X0,bool),sK4(X2,X1,X0))))) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X1))),s(fun(X0,bool),X2))))) )),
  inference(forward_demodulation,[],[f37,f34])).
fof(f37,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X1))),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),sK4(X2,X1,X0)),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f26])).
fof(f13896,plain,(
  p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK4(sK2,sK1,sK0))))) | ($spl1 | $spl6)),
  inference(resolution,[],[f45,f12607])).
fof(f12607,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK4(sK2,sK1,sK0)),s(fun(fun(sK0,bool),bool),sK1)))) | $spl1),
  inference(resolution,[],[f40,f2470])).
fof(f2470,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X6,bool),X8),s(fun(X6,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X6,bool),bool),X7)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),sK4(X8,X7,X6)),s(fun(fun(X6,bool),bool),X7))))) )),
  inference(superposition,[],[f36,f34])).
fof(f36,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X1))),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),sK4(X2,X1,X0)),s(fun(fun(X0,bool),bool),X1))))) )),
  inference(cnf_transformation,[],[f26])).
fof(f65,plain,(
  $false | $spl8),
  inference(resolution,[],[f62,f31])).
fof(f62,plain,(
  ( ! [X2] : (~p(s(bool,X2))) ) | $spl8),
  inference(superposition,[],[f32,f53])).
fof(f53,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl8),
  inference(cnf_transformation,[],[f53_D])).
fof(f53_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f55,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl11),
  inference(cnf_transformation,[],[f55_D])).
fof(f55_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f52,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f33])).
fof(f14325,plain,(
  $false | ($spl11 | $spl14)),
  inference(trivial_inequality_removal,[],[f14313])).
fof(f14313,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl11 | $spl14)),
  inference(superposition,[],[f55,f230])).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 83
% Passive clauses: 2379
% Generated clauses: 14079
% Final active clauses: 20
% Final passive clauses: 1952
% Input formulas: 7
% Initial clauses: 11
% 
% Duplicate literals: 1016
% Trivial inequalities: 171
% Fw subsumption resolutions: 1664
% Fw demodulations: 946
% Bw demodulations: 42
% 
% Simple tautologies: 254
% Equational tautologies: 95
% Forward subsumptions: 5710
% Bw demodulations to eq. taut.: 2
% 
% Binary resolution: 106
% Forward superposition: 5693
% Backward superposition: 3541
% Self superposition: 328
% Equality factoring: 261
% Equality resolution: 8
% 
% Split clauses: 809
% Split components: 1624
% SAT solver clauses: 1919
% SAT solver unit clauses: 267
% SAT solver binary clauses: 524
% SAT solver learnt clauses: 6
% SAT solver learnt literals: 7
% 
% Sat splits: 809
% Sat splitting refutations: 8
% 
% TWLsolver clauses: 1891
% TWLsolver calls for satisfiability: 84
% 
% Memory used [KB]: 5628
% Time elapsed: 0.469 s
% ------------------------------
% ------------------------------
% Success in time 0.49 s
